2020-06-06 00:29:25.303833: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_8_bb1/premsel_enigma_01_2020_T30_loop02_epoch_243
2020-06-06 00:29:25.368778: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-06-06 00:29:25.491305: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-06-06 00:29:25.869930: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-06-06 00:29:27.231719: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 1927918 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_8_bb1/premsel_enigma_01_2020_T30_loop02_epoch_243' loaded (query=128; context=512; weigths=real; len_mult=0.000000).
# Preprocessing time       : 2.603 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(dt_k1_nat_1, axiom, ![X1, X2]:((v7_ordinal1(X1)&m1_subset_1(X2,k5_numbers))=>m2_subset_1(k1_nat_1(X1,X2),k1_numbers,k5_numbers)), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', dt_k1_nat_1)).
fof(commutativity_k1_nat_1, axiom, ![X1, X2]:((v7_ordinal1(X1)&m1_subset_1(X2,k5_numbers))=>k1_nat_1(X1,X2)=k1_nat_1(X2,X1)), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', commutativity_k1_nat_1)).
fof(redefinition_k1_nat_1, axiom, ![X1, X2]:((v7_ordinal1(X1)&m1_subset_1(X2,k5_numbers))=>k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', redefinition_k1_nat_1)).
fof(dt_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)=>m1_subset_1(X3,X1))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', dt_m2_subset_1)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', cc1_subset_1)).
fof(fc6_ordinal1, axiom, (~(v1_xboole_0(k4_ordinal1))&v3_ordinal1(k4_ordinal1)), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', fc6_ordinal1)).
fof(commutativity_k2_xcmplx_0, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', commutativity_k2_xcmplx_0)).
fof(cc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xcmplx_0(X1)), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', cc3_xreal_0)).
fof(reflexivity_r1_xxreal_0, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>r1_xxreal_0(X1,X1)), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', reflexivity_r1_xxreal_0)).
fof(rc2_xreal_0, axiom, ?[X1]:(((v1_xcmplx_0(X1)&v1_xxreal_0(X1))&v2_xxreal_0(X1))&v1_xreal_0(X1)), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', rc2_xreal_0)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', redefinition_k5_numbers)).
fof(cc4_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xxreal_0(X1)), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', cc4_xreal_0)).
fof(cc2_xreal_0, axiom, ![X1]:(v7_ordinal1(X1)=>v1_xreal_0(X1)), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', cc2_xreal_0)).
fof(t8_xreal_1, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>![X3]:(v1_xreal_0(X3)=>![X4]:(v1_xreal_0(X4)=>~(((~(r1_xxreal_0(X2,X1))&r1_xxreal_0(X3,X4))&r1_xxreal_0(k2_xcmplx_0(X2,X4),k2_xcmplx_0(X1,X3)))))))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', t8_xreal_1)).
fof(cc1_xreal_0, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>v1_xreal_0(X1)), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', cc1_xreal_0)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', dt_k5_numbers)).
fof(connectedness_r1_xxreal_0, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>(r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', connectedness_r1_xxreal_0)).
fof(rc4_ordinal1, axiom, ?[X1]:v7_ordinal1(X1), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', rc4_ordinal1)).
fof(fraenkel_a_1_1_pre_circ, axiom, ![X1, X2]:(v7_ordinal1(X2)=>(r2_hidden(X1,a_1_1_pre_circ(X2))<=>?[X3]:((m1_subset_1(X3,k5_numbers)&X1=X3)&~(r1_xxreal_0(X3,X2))))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', fraenkel_a_1_1_pre_circ)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', t2_subset)).
fof(d8_xxreal_2, axiom, ![X1]:((v2_membered(X1)&v2_xxreal_2(X1))=>![X2]:(v1_xxreal_0(X2)=>(X2=k1_xxreal_2(X1)<=>(r2_hidden(X2,X1)&![X3]:(v1_xxreal_0(X3)=>(r2_hidden(X3,X1)=>r1_xxreal_0(X3,X2))))))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', d8_xxreal_2)).
fof(spc6_arithm, axiom, ![X1, X2, X3]:(((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))&v1_xcmplx_0(X3))=>k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', spc6_arithm)).
fof(dt_k1_xxreal_2, axiom, ![X1]:(v2_membered(X1)=>v1_xxreal_0(k1_xxreal_2(X1))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', dt_k1_xxreal_2)).
fof(t23_pre_circ, conjecture, ![X1]:(v7_ordinal1(X1)=>~(v1_finset_1(a_1_1_pre_circ(X1)))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', t23_pre_circ)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', t1_subset)).
fof(fc10_xxreal_2, axiom, ![X1]:((v6_membered(X1)&v2_xxreal_2(X1))=>(v1_xxreal_0(k1_xxreal_2(X1))&v7_ordinal1(k1_xxreal_2(X1)))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', fc10_xxreal_2)).
fof(fc5_xreal_0, axiom, ![X1, X2]:((v1_xreal_0(X1)&v1_xreal_0(X2))=>v1_xreal_0(k2_xcmplx_0(X1,X2))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', fc5_xreal_0)).
fof(t7_boole, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', t7_boole)).
fof(cc13_membered, axiom, ![X1]:(v2_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xxreal_0(X2))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', cc13_membered)).
fof(cc14_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xreal_0(X2))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', cc14_membered)).
fof(t2_xxreal_0, axiom, ![X1]:(v1_xxreal_0(X1)=>![X2]:(v1_xxreal_0(X2)=>![X3]:(v1_xxreal_0(X3)=>((r1_xxreal_0(X1,X2)&r1_xxreal_0(X2,X3))=>r1_xxreal_0(X1,X3))))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', t2_xxreal_0)).
fof(spc0_numerals, axiom, ((m2_subset_1(np__0,k1_numbers,k5_numbers)&m1_subset_1(np__0,k5_numbers))&m1_subset_1(np__0,k1_numbers)), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', spc0_numerals)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', fc3_membered)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', spc1_numerals)).
fof(rqLessOrEqual__r1_xxreal_0__r1_r0, axiom, ~(r1_xxreal_0(np__1,np__0)), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', rqLessOrEqual__r1_xxreal_0__r1_r0)).
fof(rqRealAdd__k2_xcmplx_0__r0_r1_r1, axiom, k2_xcmplx_0(np__0,np__1)=np__1, file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', rqRealAdd__k2_xcmplx_0__r0_r1_r1)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', redefinition_m2_subset_1)).
fof(cc8_ordinal1, axiom, ![X1]:(m1_subset_1(X1,k4_ordinal1)=>v7_ordinal1(X1)), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', cc8_ordinal1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', d3_tarski)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', t3_subset)).
fof(cc20_membered, axiom, ![X1]:(v2_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v2_membered(X2))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', cc20_membered)).
fof(cc2_xxreal_2, axiom, ![X1]:(((v2_membered(X1)&~(v1_xboole_0(X1)))&v1_finset_1(X1))=>(((v2_membered(X1)&~(v1_xboole_0(X1)))&v1_xxreal_2(X1))&v2_xxreal_2(X1))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', cc2_xxreal_2)).
fof(cc24_membered, axiom, ![X1]:(v6_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v6_membered(X2))), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', cc24_membered)).
fof(fc6_membered, axiom, v6_membered(k4_ordinal1), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', fc6_membered)).
fof(cc4_membered, axiom, ![X1]:(v3_membered(X1)=>v2_membered(X1)), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', cc4_membered)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', t6_boole)).
fof(spc0_boole, axiom, v1_xboole_0(np__0), file('/local1/mptp/loop2/hardpred/preds__160/t23_pre_circ', spc0_boole)).
fof(c_0_47, plain, ![X30, X31]:(~v7_ordinal1(X30)|~m1_subset_1(X31,k5_numbers)|m2_subset_1(k1_nat_1(X30,X31),k1_numbers,k5_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_nat_1])])).
fof(c_0_48, plain, ![X25, X26]:(~v7_ordinal1(X25)|~m1_subset_1(X26,k5_numbers)|k1_nat_1(X25,X26)=k1_nat_1(X26,X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k1_nat_1])])).
fof(c_0_49, plain, ![X16, X17]:(~v7_ordinal1(X16)|~m1_subset_1(X17,k5_numbers)|k1_nat_1(X16,X17)=k2_xcmplx_0(X16,X17)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_nat_1])])).
fof(c_0_50, plain, ![X118, X119, X120]:(v1_xboole_0(X118)|v1_xboole_0(X119)|~m1_subset_1(X119,k1_zfmisc_1(X118))|(~m2_subset_1(X120,X118,X119)|m1_subset_1(X120,X118))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_subset_1])])])])).
fof(c_0_51, plain, ![X52, X53]:(~v1_xboole_0(X52)|(~m1_subset_1(X53,k1_zfmisc_1(X52))|v1_xboole_0(X53))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
cnf(c_0_52, plain, (m2_subset_1(k1_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_53, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_54, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_55, plain, (~v1_xboole_0(k4_ordinal1)&v3_ordinal1(k4_ordinal1)), inference(fof_simplification,[status(thm)],[fc6_ordinal1])).
fof(c_0_56, plain, ![X39, X40]:(~v1_xcmplx_0(X39)|~v1_xcmplx_0(X40)|k2_xcmplx_0(X39,X40)=k2_xcmplx_0(X40,X39)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k2_xcmplx_0])])).
fof(c_0_57, plain, ![X72]:(~v1_xreal_0(X72)|v1_xcmplx_0(X72)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_xreal_0])])).
fof(c_0_58, plain, ![X54, X55]:(~v1_xxreal_0(X54)|~v1_xxreal_0(X55)|r1_xxreal_0(X54,X54)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[reflexivity_r1_xxreal_0])])).
fof(c_0_59, plain, (((v1_xcmplx_0(esk22_0)&v1_xxreal_0(esk22_0))&v2_xxreal_0(esk22_0))&v1_xreal_0(esk22_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc2_xreal_0])])).
cnf(c_0_60, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m2_subset_1(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_61, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_62, plain, (m2_subset_1(k1_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_52])).
cnf(c_0_63, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_53])).
cnf(c_0_64, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_54])).
cnf(c_0_65, plain, (~v1_xboole_0(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_66, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
cnf(c_0_67, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_68, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
fof(c_0_69, plain, ![X80]:(~v1_xreal_0(X80)|v1_xxreal_0(X80)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_xreal_0])])).
fof(c_0_70, plain, ![X64]:(~v7_ordinal1(X64)|v1_xreal_0(X64)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_xreal_0])])).
fof(c_0_71, plain, ![X21, X22, X23, X24]:(~v1_xreal_0(X21)|(~v1_xreal_0(X22)|(~v1_xreal_0(X23)|(~v1_xreal_0(X24)|(r1_xxreal_0(X22,X21)|~r1_xxreal_0(X23,X24)|~r1_xxreal_0(k2_xcmplx_0(X22,X24),k2_xcmplx_0(X21,X23))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_xreal_1])])])])).
cnf(c_0_72, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_73, plain, (v1_xxreal_0(esk22_0)), inference(split_conjunct,[status(thm)],[c_0_59])).
fof(c_0_74, plain, ![X63]:(~m1_subset_1(X63,k1_numbers)|v1_xreal_0(X63)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_xreal_0])])).
cnf(c_0_75, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_60])).
cnf(c_0_76, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_61])).
cnf(c_0_77, plain, (m2_subset_1(k1_nat_1(X2,X1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63])])).
cnf(c_0_78, plain, (k1_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_64])])).
cnf(c_0_79, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
cnf(c_0_80, plain, (~v1_xboole_0(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_81, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_66])).
cnf(c_0_82, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_67])).
cnf(c_0_83, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_68])).
fof(c_0_84, plain, ![X56, X57]:(~v1_xxreal_0(X56)|~v1_xxreal_0(X57)|(r1_xxreal_0(X56,X57)|r1_xxreal_0(X57,X56))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[connectedness_r1_xxreal_0])])).
cnf(c_0_85, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_86, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_87, plain, (r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(k2_xcmplx_0(X2,X4),k2_xcmplx_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_88, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_72])).
cnf(c_0_89, plain, (v1_xxreal_0(esk22_0)), inference(evalgc,[status(thm)],[c_0_73])).
cnf(c_0_90, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_91, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~m2_subset_1(X2,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(csr,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_92, plain, (m2_subset_1(k2_xcmplx_0(X2,X1),k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_78])])).
cnf(c_0_93, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_79])).
cnf(c_0_94, plain, (~v1_xboole_0(k5_numbers)), inference(rw,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_95, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83])])).
cnf(c_0_96, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_97, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_98, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_99, plain, (r1_xxreal_0(X2,X1)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(k2_xcmplx_0(X2,X4),k2_xcmplx_0(X1,X3))), inference(evalgc,[status(thm)],[c_0_87])).
cnf(c_0_100, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89])])).
cnf(c_0_101, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_102, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_93])]), c_0_94])])).
cnf(c_0_103, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_83])])).
cnf(c_0_104, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_96])).
cnf(c_0_105, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98])])).
fof(c_0_106, plain, v7_ordinal1(esk12_0), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc4_ordinal1])])).
fof(c_0_107, plain, ![X158, X159, X161]:((((m1_subset_1(esk15_2(X158,X159),k5_numbers)|~r2_hidden(X158,a_1_1_pre_circ(X159))|~v7_ordinal1(X159))&(X158=esk15_2(X158,X159)|~r2_hidden(X158,a_1_1_pre_circ(X159))|~v7_ordinal1(X159)))&(~r1_xxreal_0(esk15_2(X158,X159),X159)|~r2_hidden(X158,a_1_1_pre_circ(X159))|~v7_ordinal1(X159)))&(~m1_subset_1(X161,k5_numbers)|X158!=X161|r1_xxreal_0(X161,X159)|r2_hidden(X158,a_1_1_pre_circ(X159))|~v7_ordinal1(X159))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_1_1_pre_circ])])])])])])).
fof(c_0_108, plain, ![X76, X77]:(~m1_subset_1(X76,X77)|(v1_xboole_0(X77)|r2_hidden(X76,X77))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
fof(c_0_109, plain, ![X12, X13, X14]:(((r2_hidden(X13,X12)|X13!=k1_xxreal_2(X12)|~v1_xxreal_0(X13)|(~v2_membered(X12)|~v2_xxreal_2(X12)))&(~v1_xxreal_0(X14)|(~r2_hidden(X14,X12)|r1_xxreal_0(X14,X13))|X13!=k1_xxreal_2(X12)|~v1_xxreal_0(X13)|(~v2_membered(X12)|~v2_xxreal_2(X12))))&((v1_xxreal_0(esk3_2(X12,X13))|~r2_hidden(X13,X12)|X13=k1_xxreal_2(X12)|~v1_xxreal_0(X13)|(~v2_membered(X12)|~v2_xxreal_2(X12)))&((r2_hidden(esk3_2(X12,X13),X12)|~r2_hidden(X13,X12)|X13=k1_xxreal_2(X12)|~v1_xxreal_0(X13)|(~v2_membered(X12)|~v2_xxreal_2(X12)))&(~r1_xxreal_0(esk3_2(X12,X13),X13)|~r2_hidden(X13,X12)|X13=k1_xxreal_2(X12)|~v1_xxreal_0(X13)|(~v2_membered(X12)|~v2_xxreal_2(X12)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d8_xxreal_2])])])])])).
fof(c_0_110, plain, ![X115, X116, X117]:(~v1_xcmplx_0(X115)|~v1_xcmplx_0(X116)|~v1_xcmplx_0(X117)|k2_xcmplx_0(k2_xcmplx_0(X115,X116),X117)=k2_xcmplx_0(X115,k2_xcmplx_0(X116,X117))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc6_arithm])])).
cnf(c_0_111, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X2)|~v1_xxreal_0(k2_xcmplx_0(X1,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100])])).
cnf(c_0_112, plain, (v1_xxreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_101])])).
cnf(c_0_113, plain, (m1_subset_1(k2_xcmplx_0(X2,X1),k1_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_103])])).
cnf(c_0_114, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105])])).
cnf(c_0_115, plain, (v7_ordinal1(esk12_0)), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_116, plain, (X1=esk15_2(X1,X2)|~r2_hidden(X1,a_1_1_pre_circ(X2))|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_117, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_108])).
cnf(c_0_118, plain, (r2_hidden(X1,X2)|X1!=k1_xxreal_2(X2)|~v1_xxreal_0(X1)|~v2_membered(X2)|~v2_xxreal_2(X2)), inference(split_conjunct,[status(thm)],[c_0_109])).
fof(c_0_119, plain, ![X47]:(~v2_membered(X47)|v1_xxreal_0(k1_xxreal_2(X47))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_xxreal_2])])).
fof(c_0_120, negated_conjecture, ~(![X1]:(v7_ordinal1(X1)=>~(v1_finset_1(a_1_1_pre_circ(X1))))), inference(assume_negation,[status(cth)],[t23_pre_circ])).
cnf(c_0_121, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),k2_xcmplx_0(X2,X4))|~r1_xxreal_0(X4,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_103])])).
cnf(c_0_122, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_123, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X2)|~m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112])])).
cnf(c_0_124, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(csr,[status(thm)],[c_0_113, c_0_98])).
cnf(c_0_125, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_105])])).
cnf(c_0_126, plain, (v7_ordinal1(esk12_0)), inference(evalgc,[status(thm)],[c_0_115])).
cnf(c_0_127, plain, (~r1_xxreal_0(esk15_2(X1,X2),X2)|~r2_hidden(X1,a_1_1_pre_circ(X2))|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_128, plain, (esk15_2(X1,X2)=X1|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_1_pre_circ(X2))), inference(evalgc,[status(thm)],[c_0_116])).
cnf(c_0_129, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_117])).
fof(c_0_130, plain, ![X58, X59]:(~r2_hidden(X58,X59)|m1_subset_1(X58,X59)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_131, plain, (r2_hidden(X1,X2)|X1!=k1_xxreal_2(X2)|~v2_membered(X2)|~v2_xxreal_2(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_118])).
cnf(c_0_132, plain, (v1_xxreal_0(k1_xxreal_2(X1))|~v2_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_119])).
fof(c_0_133, negated_conjecture, (v7_ordinal1(esk1_0)&v1_finset_1(a_1_1_pre_circ(esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_120])])])])).
cnf(c_0_134, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),k2_xcmplx_0(X4,X2))|~r1_xxreal_0(X4,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_103])])).
cnf(c_0_135, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_122])).
cnf(c_0_136, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_124])])).
cnf(c_0_137, plain, (r1_xxreal_0(X1,esk12_0)|r1_xxreal_0(esk12_0,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126])])).
cnf(c_0_138, plain, (~v7_ordinal1(X2)|~r2_hidden(X1,a_1_1_pre_circ(X2))|~r1_xxreal_0(esk15_2(X1,X2),X2)), inference(evalgc,[status(thm)],[c_0_127])).
cnf(c_0_139, plain, (esk15_2(X1,X2)=X1|v1_xboole_0(a_1_1_pre_circ(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,a_1_1_pre_circ(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129])])).
cnf(c_0_140, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_130])).
cnf(c_0_141, plain, (r2_hidden(k1_xxreal_2(X1),X1)|~v1_xxreal_0(k1_xxreal_2(X1))|~v2_xxreal_2(X1)|~v2_membered(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_131])])).
cnf(c_0_142, plain, (v1_xxreal_0(k1_xxreal_2(X1))|~v2_membered(X1)), inference(evalgc,[status(thm)],[c_0_132])).
cnf(c_0_143, negated_conjecture, (v7_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_133])).
fof(c_0_144, plain, ![X139]:((v1_xxreal_0(k1_xxreal_2(X139))|(~v6_membered(X139)|~v2_xxreal_2(X139)))&(v7_ordinal1(k1_xxreal_2(X139))|(~v6_membered(X139)|~v2_xxreal_2(X139)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc10_xxreal_2])])])).
cnf(c_0_145, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X4,X5))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),k2_xcmplx_0(X4,k2_xcmplx_0(X5,X2)))|~r1_xxreal_0(k2_xcmplx_0(X4,X5),X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135])])).
cnf(c_0_146, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(X2,X2)|~v7_ordinal1(X2)), inference(csr,[status(thm)],[c_0_136, c_0_98])).
cnf(c_0_147, plain, (r1_xxreal_0(esk12_0,esk12_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_126])])).
fof(c_0_148, plain, ![X99, X100]:(~v1_xreal_0(X99)|~v1_xreal_0(X100)|v1_xreal_0(k2_xcmplx_0(X99,X100))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc5_xreal_0])])).
cnf(c_0_149, plain, (v1_xboole_0(a_1_1_pre_circ(X2))|~r1_xxreal_0(X1,X2)|~r2_hidden(X1,a_1_1_pre_circ(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,a_1_1_pre_circ(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139])])).
cnf(c_0_150, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_140])).
cnf(c_0_151, plain, (r2_hidden(k1_xxreal_2(X1),X1)|~v2_xxreal_2(X1)|~v2_membered(X1)), inference(csr,[status(thm)],[c_0_141, c_0_142])).
cnf(c_0_152, negated_conjecture, (v7_ordinal1(esk1_0)), inference(evalgc,[status(thm)],[c_0_143])).
cnf(c_0_153, plain, (v7_ordinal1(k1_xxreal_2(X1))|~v6_membered(X1)|~v2_xxreal_2(X1)), inference(split_conjunct,[status(thm)],[c_0_144])).
fof(c_0_154, plain, ![X74, X75]:(~r2_hidden(X74,X75)|~v1_xboole_0(X75)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
cnf(c_0_155, plain, (r1_xxreal_0(X1,X3)|~v1_xxreal_0(X1)|~r2_hidden(X1,X2)|X3!=k1_xxreal_2(X2)|~v1_xxreal_0(X3)|~v2_membered(X2)|~v2_xxreal_2(X2)), inference(split_conjunct,[status(thm)],[c_0_109])).
fof(c_0_156, plain, ![X101, X102]:(~v2_membered(X101)|(~m1_subset_1(X102,X101)|v1_xxreal_0(X102))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc13_membered])])])).
cnf(c_0_157, plain, (r1_xxreal_0(X1,X3)|r2_hidden(X2,a_1_1_pre_circ(X3))|~m1_subset_1(X1,k5_numbers)|X2!=X1|~v7_ordinal1(X3)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_158, plain, (r1_xxreal_0(X1,X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4)|~v1_xreal_0(k2_xcmplx_0(X4,X3))|~v1_xreal_0(X5)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X5,X1),k2_xcmplx_0(X4,k2_xcmplx_0(X3,X2)))|~r1_xxreal_0(k2_xcmplx_0(X4,X3),X5)), inference(csr,[status(thm)],[c_0_145, c_0_83])).
cnf(c_0_159, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_147]), c_0_126])])])).
cnf(c_0_160, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_148])).
cnf(c_0_161, plain, (v1_xboole_0(a_1_1_pre_circ(X1))|~m1_subset_1(X2,a_1_1_pre_circ(X1))|~r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_149, c_0_129])).
cnf(c_0_162, plain, (m1_subset_1(k1_xxreal_2(X1),X1)|~v2_xxreal_2(X1)|~v2_membered(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_151])])).
cnf(c_0_163, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_152])])).
cnf(c_0_164, plain, (v7_ordinal1(k1_xxreal_2(X1))|~v2_xxreal_2(X1)|~v6_membered(X1)), inference(evalgc,[status(thm)],[c_0_153])).
cnf(c_0_165, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_154])).
cnf(c_0_166, plain, (r1_xxreal_0(X1,X3)|X3!=k1_xxreal_2(X2)|~v2_membered(X2)|~v2_xxreal_2(X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_155])).
cnf(c_0_167, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_156])).
cnf(c_0_168, plain, (r1_xxreal_0(X1,X3)|r2_hidden(X2,a_1_1_pre_circ(X3))|X1!=X2|~v7_ordinal1(X3)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_157])).
cnf(c_0_169, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~v1_xreal_0(k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2)))|~m1_subset_1(k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2)),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159])])).
cnf(c_0_170, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_160])).
fof(c_0_171, plain, ![X70, X71]:(~v3_membered(X70)|(~m1_subset_1(X71,X70)|v1_xreal_0(X71))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc14_membered])])])).
fof(c_0_172, plain, ![X18, X19, X20]:(~v1_xxreal_0(X18)|(~v1_xxreal_0(X19)|(~v1_xxreal_0(X20)|(~r1_xxreal_0(X18,X19)|~r1_xxreal_0(X19,X20)|r1_xxreal_0(X18,X20))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_xxreal_0])])])).
cnf(c_0_173, plain, (v1_xboole_0(a_1_1_pre_circ(X1))|~r1_xxreal_0(k1_xxreal_2(a_1_1_pre_circ(X1)),X1)|~v7_ordinal1(X1)|~v2_xxreal_2(a_1_1_pre_circ(X1))|~v2_membered(a_1_1_pre_circ(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_162])])).
cnf(c_0_174, negated_conjecture, (r1_xxreal_0(esk1_0,k1_xxreal_2(X1))|r1_xxreal_0(k1_xxreal_2(X1),esk1_0)|~v6_membered(X1)|~v2_xxreal_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_164])])).
cnf(c_0_175, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_165])).
cnf(c_0_176, plain, (r1_xxreal_0(X1,X2)|v1_xboole_0(X3)|X2!=k1_xxreal_2(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v2_xxreal_2(X3)|~v2_membered(X3)|~m1_subset_1(X1,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_129])])).
cnf(c_0_177, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_167])).
cnf(c_0_178, plain, (r1_xxreal_0(X1,X2)|r2_hidden(X1,a_1_1_pre_circ(X2))|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_168])])).
cnf(c_0_179, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2)),k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_169, c_0_170]), c_0_83])).
cnf(c_0_180, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_171])).
cnf(c_0_181, plain, (m1_subset_1(np__0,k1_numbers)), inference(split_conjunct,[status(thm)],[spc0_numerals])).
cnf(c_0_182, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
cnf(c_0_183, plain, (m1_subset_1(np__1,k1_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
fof(c_0_184, plain, ~r1_xxreal_0(np__1,np__0), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r1_r0])).
cnf(c_0_185, plain, (r1_xxreal_0(X1,X3)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_172])).
cnf(c_0_186, negated_conjecture, (v1_xboole_0(a_1_1_pre_circ(esk1_0))|r1_xxreal_0(esk1_0,k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))|~v6_membered(a_1_1_pre_circ(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_174]), c_0_152])])])).
cnf(c_0_187, plain, (~v1_xboole_0(X1)|~v2_xxreal_2(X1)|~v2_membered(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_151])])).
cnf(c_0_188, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,X3)|X3!=k1_xxreal_2(X1)|~m1_subset_1(X2,X1)|~v1_xxreal_0(X3)|~v2_xxreal_2(X1)|~v2_membered(X1)), inference(csr,[status(thm)],[c_0_176, c_0_177])).
cnf(c_0_189, plain, (m1_subset_1(X1,a_1_1_pre_circ(X2))|r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_178])])).
cnf(c_0_190, plain, (r1_xxreal_0(k2_xcmplx_0(X2,X1),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(k2_xcmplx_0(X3,k2_xcmplx_0(X2,X1)),k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_103])])).
cnf(c_0_191, plain, (k2_xcmplx_0(np__0,np__1)=np__1), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r0_r1_r1])).
cnf(c_0_192, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_180])).
cnf(c_0_193, plain, (m1_subset_1(np__0,k1_numbers)), inference(evalgc,[status(thm)],[c_0_181])).
cnf(c_0_194, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_182])).
cnf(c_0_195, plain, (m1_subset_1(np__1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_183])).
cnf(c_0_196, plain, (~r1_xxreal_0(np__1,np__0)), inference(split_conjunct,[status(thm)],[c_0_184])).
cnf(c_0_197, plain, (r1_xxreal_0(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_185])).
cnf(c_0_198, negated_conjecture, (r1_xxreal_0(esk1_0,k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~v6_membered(a_1_1_pre_circ(esk1_0))|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))), inference(csr,[status(thm)],[c_0_186, c_0_187])).
cnf(c_0_199, plain, (v1_xboole_0(a_1_1_pre_circ(X1))|r1_xxreal_0(X2,X3)|r1_xxreal_0(X2,X1)|X3!=k1_xxreal_2(a_1_1_pre_circ(X1))|~v1_xxreal_0(X3)|~v2_xxreal_2(a_1_1_pre_circ(X1))|~v2_membered(a_1_1_pre_circ(X1))|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_189])])).
cnf(c_0_200, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2)),k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X3,X2),X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_190, c_0_170]), c_0_170]), c_0_83])).
cnf(c_0_201, plain, (k2_xcmplx_0(np__0,np__1)=np__1), inference(evalgc,[status(thm)],[c_0_191])).
cnf(c_0_202, plain, (v1_xreal_0(np__0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_193]), c_0_194])])])).
cnf(c_0_203, plain, (v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_195]), c_0_194])])])).
cnf(c_0_204, plain, (~r1_xxreal_0(np__1,np__0)), inference(evalgc,[status(thm)],[c_0_196])).
cnf(c_0_205, negated_conjecture, (r1_xxreal_0(X1,k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~r1_xxreal_0(X1,esk1_0)|~v1_xxreal_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~v1_xxreal_0(esk1_0)|~v1_xxreal_0(X1)|~v6_membered(a_1_1_pre_circ(esk1_0))|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_198])])).
cnf(c_0_206, plain, (v1_xboole_0(a_1_1_pre_circ(X1))|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,k1_xxreal_2(a_1_1_pre_circ(X1)))|~m1_subset_1(X2,k5_numbers)|~v1_xxreal_0(k1_xxreal_2(a_1_1_pre_circ(X1)))|~v2_xxreal_2(a_1_1_pre_circ(X1))|~v2_membered(a_1_1_pre_circ(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_199])])).
fof(c_0_207, plain, ![X27, X28, X29]:((~m2_subset_1(X29,X27,X28)|m1_subset_1(X29,X28)|(v1_xboole_0(X27)|v1_xboole_0(X28)|~m1_subset_1(X28,k1_zfmisc_1(X27))))&(~m1_subset_1(X29,X28)|m2_subset_1(X29,X27,X28)|(v1_xboole_0(X27)|v1_xboole_0(X28)|~m1_subset_1(X28,k1_zfmisc_1(X27))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])).
cnf(c_0_208, plain, (~v1_xreal_0(X1)|~m1_subset_1(k2_xcmplx_0(X1,np__1),k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X1,np__1),X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_201]), c_0_202]), c_0_203])]), c_0_204])])).
cnf(c_0_209, negated_conjecture, (r1_xxreal_0(X1,k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~v6_membered(a_1_1_pre_circ(esk1_0))|~r1_xxreal_0(X1,esk1_0)|~v1_xxreal_0(esk1_0)|~v1_xxreal_0(X1)|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))), inference(csr,[status(thm)],[c_0_205, c_0_142])).
cnf(c_0_210, plain, (r1_xxreal_0(X1,k1_xxreal_2(a_1_1_pre_circ(X2)))|r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k5_numbers)|~v2_xxreal_2(a_1_1_pre_circ(X2))|~v2_membered(a_1_1_pre_circ(X2))|~v7_ordinal1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_206, c_0_142]), c_0_187])).
fof(c_0_211, plain, ![X32]:(~m1_subset_1(X32,k4_ordinal1)|v7_ordinal1(X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_ordinal1])])).
cnf(c_0_212, plain, (m1_subset_1(X1,X3)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_207])).
cnf(c_0_213, negated_conjecture, (~v1_xreal_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~m1_subset_1(k2_xcmplx_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)),np__1),k5_numbers)|~v6_membered(a_1_1_pre_circ(esk1_0))|~r1_xxreal_0(k2_xcmplx_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)),np__1),esk1_0)|~v1_xxreal_0(esk1_0)|~v1_xxreal_0(k2_xcmplx_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)),np__1))|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_209])])).
cnf(c_0_214, plain, (r1_xxreal_0(k2_xcmplx_0(k1_xxreal_2(a_1_1_pre_circ(X1)),np__1),X1)|~v1_xreal_0(k1_xxreal_2(a_1_1_pre_circ(X1)))|~m1_subset_1(k2_xcmplx_0(k1_xxreal_2(a_1_1_pre_circ(X1)),np__1),k5_numbers)|~v2_xxreal_2(a_1_1_pre_circ(X1))|~v2_membered(a_1_1_pre_circ(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_210])])).
cnf(c_0_215, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_211])).
cnf(c_0_216, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_212])).
cnf(c_0_217, negated_conjecture, (~v6_membered(a_1_1_pre_circ(esk1_0))|~v1_xreal_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~m1_subset_1(k2_xcmplx_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)),np__1),k5_numbers)|~v1_xxreal_0(k2_xcmplx_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)),np__1))|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_214]), c_0_152])])])).
cnf(c_0_218, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_215])).
cnf(c_0_219, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X1)|~m2_subset_1(X2,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(csr,[status(thm)],[c_0_216, c_0_76])).
fof(c_0_220, plain, ![X6, X7, X8, X9, X10]:((~r1_tarski(X6,X7)|(~r2_hidden(X8,X6)|r2_hidden(X8,X7)))&((r2_hidden(esk2_2(X9,X10),X9)|r1_tarski(X9,X10))&(~r2_hidden(esk2_2(X9,X10),X10)|r1_tarski(X9,X10)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_221, negated_conjecture, (~v6_membered(a_1_1_pre_circ(esk1_0))|~v1_xreal_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~m1_subset_1(k2_xcmplx_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)),np__1),k5_numbers)|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))|~v7_ordinal1(k2_xcmplx_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)),np__1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_105])])).
cnf(c_0_222, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_218, c_0_81])).
cnf(c_0_223, plain, (m1_subset_1(k1_nat_1(X1,X2),k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_62]), c_0_93])]), c_0_94])])).
cnf(c_0_224, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_225, plain, (m1_subset_1(esk15_2(X1,X2),k5_numbers)|~r2_hidden(X1,a_1_1_pre_circ(X2))|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_107])).
cnf(c_0_226, plain, (r2_hidden(esk2_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_220])).
cnf(c_0_227, negated_conjecture, (~v6_membered(a_1_1_pre_circ(esk1_0))|~v1_xreal_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~m1_subset_1(k2_xcmplx_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)),np__1),k5_numbers)|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))), inference(csr,[status(thm)],[c_0_221, c_0_222])).
cnf(c_0_228, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_64])])).
cnf(c_0_229, plain, (m1_subset_1(np__1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_224])).
cnf(c_0_230, plain, (r1_tarski(X1,X2)|~r2_hidden(esk2_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_220])).
cnf(c_0_231, plain, (m1_subset_1(esk15_2(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_1_pre_circ(X2))), inference(evalgc,[status(thm)],[c_0_225])).
cnf(c_0_232, plain, (r1_tarski(X1,X2)|r2_hidden(esk2_2(X1,X2),X1)), inference(evalgc,[status(thm)],[c_0_226])).
cnf(c_0_233, negated_conjecture, (~v6_membered(a_1_1_pre_circ(esk1_0))|~v1_xreal_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))|~v7_ordinal1(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_228]), c_0_229])])])).
fof(c_0_234, plain, ![X49, X50]:((~m1_subset_1(X49,k1_zfmisc_1(X50))|r1_tarski(X49,X50))&(~r1_tarski(X49,X50)|m1_subset_1(X49,k1_zfmisc_1(X50)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_235, plain, (r1_tarski(X1,X2)|~r2_hidden(esk2_2(X1,X2),X2)), inference(evalgc,[status(thm)],[c_0_230])).
cnf(c_0_236, plain, (m1_subset_1(esk15_2(esk2_2(a_1_1_pre_circ(X1),X2),X1),k5_numbers)|r1_tarski(a_1_1_pre_circ(X1),X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_232])])).
cnf(c_0_237, plain, (esk15_2(esk2_2(a_1_1_pre_circ(X1),X2),X1)=esk2_2(a_1_1_pre_circ(X1),X2)|r1_tarski(a_1_1_pre_circ(X1),X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_232])])).
fof(c_0_238, plain, ![X141, X142]:(~v2_membered(X141)|(~m1_subset_1(X142,k1_zfmisc_1(X141))|v2_membered(X142))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc20_membered])])])).
cnf(c_0_239, negated_conjecture, (~v6_membered(a_1_1_pre_circ(esk1_0))|~v1_xreal_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))), inference(csr,[status(thm)],[c_0_233, c_0_164])).
fof(c_0_240, plain, ![X190]:((((v2_membered(X190)|(~v2_membered(X190)|v1_xboole_0(X190)|~v1_finset_1(X190)))&(~v1_xboole_0(X190)|(~v2_membered(X190)|v1_xboole_0(X190)|~v1_finset_1(X190))))&(v1_xxreal_2(X190)|(~v2_membered(X190)|v1_xboole_0(X190)|~v1_finset_1(X190))))&(v2_xxreal_2(X190)|(~v2_membered(X190)|v1_xboole_0(X190)|~v1_finset_1(X190)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_xxreal_2])])])])).
fof(c_0_241, plain, ![X176, X177]:(~v6_membered(X176)|(~m1_subset_1(X177,k1_zfmisc_1(X176))|v6_membered(X177))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc24_membered])])])).
cnf(c_0_242, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_234])).
cnf(c_0_243, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(esk2_2(X1,X2),X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_129])])).
cnf(c_0_244, plain, (m1_subset_1(esk2_2(a_1_1_pre_circ(X1),X2),k5_numbers)|r1_tarski(a_1_1_pre_circ(X1),X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_237])])).
cnf(c_0_245, plain, (v6_membered(k4_ordinal1)), inference(split_conjunct,[status(thm)],[fc6_membered])).
cnf(c_0_246, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_238])).
fof(c_0_247, plain, ![X68]:(~v3_membered(X68)|v2_membered(X68)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_membered])])).
fof(c_0_248, plain, ![X48]:(~v1_xboole_0(X48)|X48=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_249, negated_conjecture, (~v6_membered(a_1_1_pre_circ(esk1_0))|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))|~v7_ordinal1(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_98])])).
cnf(c_0_250, plain, (v2_xxreal_2(X1)|v1_xboole_0(X1)|~v2_membered(X1)|~v1_finset_1(X1)), inference(split_conjunct,[status(thm)],[c_0_240])).
cnf(c_0_251, negated_conjecture, (v1_finset_1(a_1_1_pre_circ(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_133])).
cnf(c_0_252, plain, (v6_membered(X2)|~v6_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_241])).
cnf(c_0_253, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_242])).
cnf(c_0_254, plain, (r1_tarski(a_1_1_pre_circ(X1),k5_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_244]), c_0_94])])).
cnf(c_0_255, plain, (v6_membered(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_245])).
cnf(c_0_256, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_246])).
cnf(c_0_257, plain, (v2_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_247])).
cnf(c_0_258, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_248])).
cnf(c_0_259, plain, (v1_xboole_0(np__0)), inference(split_conjunct,[status(thm)],[spc0_boole])).
cnf(c_0_260, negated_conjecture, (~v6_membered(a_1_1_pre_circ(esk1_0))|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))), inference(csr,[status(thm)],[c_0_249, c_0_164])).
cnf(c_0_261, plain, (v2_xxreal_2(X1)|v1_xboole_0(X1)|~v1_finset_1(X1)|~v2_membered(X1)), inference(evalgc,[status(thm)],[c_0_250])).
cnf(c_0_262, negated_conjecture, (v1_finset_1(a_1_1_pre_circ(esk1_0))), inference(evalgc,[status(thm)],[c_0_251])).
cnf(c_0_263, plain, (v6_membered(X2)|~v6_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_252])).
cnf(c_0_264, plain, (m1_subset_1(a_1_1_pre_circ(X1),k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_254])])).
cnf(c_0_265, plain, (v6_membered(k5_numbers)), inference(rw,[status(thm)],[c_0_255, c_0_81])).
cnf(c_0_266, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_93])])).
cnf(c_0_267, plain, (v2_membered(X1)|~v3_membered(X1)), inference(evalgc,[status(thm)],[c_0_257])).
cnf(c_0_268, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_258])).
cnf(c_0_269, plain, (v1_xboole_0(np__0)), inference(evalgc,[status(thm)],[c_0_259])).
cnf(c_0_270, negated_conjecture, (v1_xboole_0(a_1_1_pre_circ(esk1_0))|~v6_membered(a_1_1_pre_circ(esk1_0))|~v1_xxreal_0(esk1_0)|~v2_membered(a_1_1_pre_circ(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_261]), c_0_262])])])).
cnf(c_0_271, plain, (v6_membered(a_1_1_pre_circ(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_264]), c_0_265])])])).
cnf(c_0_272, plain, (v2_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_267]), c_0_194])])])).
cnf(c_0_273, plain, (k1_xboole_0=np__0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_269])])).
cnf(c_0_274, negated_conjecture, (v1_xboole_0(a_1_1_pre_circ(esk1_0))|~v1_xxreal_0(esk1_0)|~v2_membered(a_1_1_pre_circ(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_271]), c_0_152])])])).
cnf(c_0_275, plain, (v2_membered(a_1_1_pre_circ(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_264]), c_0_272])])])).
cnf(c_0_276, plain, (X1=np__0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_268, c_0_273])])).
cnf(c_0_277, negated_conjecture, (v1_xboole_0(a_1_1_pre_circ(esk1_0))|~v1_xxreal_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_274, c_0_275]), c_0_152])])])).
cnf(c_0_278, plain, (r1_xxreal_0(X2,X1)|~v1_xboole_0(a_1_1_pre_circ(X1))|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_178])])).
cnf(c_0_279, negated_conjecture, (a_1_1_pre_circ(esk1_0)=np__0|~v1_xxreal_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_276, c_0_277])])).
cnf(c_0_280, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|~m1_subset_1(X1,k5_numbers)|~v1_xxreal_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_279]), c_0_269]), c_0_152])])])).
cnf(c_0_281, negated_conjecture, (~v1_xreal_0(esk1_0)|~m1_subset_1(k2_xcmplx_0(esk1_0,np__1),k5_numbers)|~v1_xxreal_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_280])])).
cnf(c_0_282, negated_conjecture, (~v1_xreal_0(esk1_0)|~m1_subset_1(k2_xcmplx_0(esk1_0,np__1),k5_numbers)), inference(csr,[status(thm)],[c_0_281, c_0_97])).
cnf(c_0_283, negated_conjecture, (~v1_xreal_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_228]), c_0_229]), c_0_152])])])).
cnf(c_0_284, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_98]), c_0_152])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 285
# Proof object clause steps            : 198
# Proof object formula steps           : 87
# Proof object conjectures             : 30
# Proof object clause conjectures      : 27
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 54
# Proof object initial formulas used   : 47
# Proof object generating inferences   : 68
# Proof object simplifying inferences  : 69
# Training examples: 123 positive, 768 negative
# Training: Positive examples begin
cnf(c_0_81, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_152, negated_conjecture, (v7_ordinal1(esk1_0))).# trainpos
cnf(c_0_126, plain, (v7_ordinal1(esk12_0))).# trainpos
cnf(c_0_262, negated_conjecture, (v1_finset_1(a_1_1_pre_circ(esk1_0)))).# trainpos
cnf(c_0_80, plain, (~v1_xboole_0(k4_ordinal1))).# trainpos
cnf(c_0_89, plain, (v1_xxreal_0(esk22_0))).# trainpos
cnf(c_0_269, plain, (v1_xboole_0(np__0))).# trainpos
cnf(c_0_255, plain, (v6_membered(k4_ordinal1))).# trainpos
cnf(c_0_194, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_98, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_204, plain, (~r1_xxreal_0(np__1,np__0))).# trainpos
cnf(c_0_267, plain, (v2_membered(X1)|~v3_membered(X1))).# trainpos
cnf(c_0_268, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_193, plain, (m1_subset_1(np__0,k1_numbers))).# trainpos
cnf(c_0_201, plain, (k2_xcmplx_0(np__0,np__1)=np__1)).# trainpos
cnf(c_0_229, plain, (m1_subset_1(np__1,k5_numbers))).# trainpos
cnf(c_0_195, plain, (m1_subset_1(np__1,k1_numbers))).# trainpos
cnf(c_0_93, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_97, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_218, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).# trainpos
cnf(c_0_175, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_83, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_138, plain, (~v7_ordinal1(X2)|~r2_hidden(X1,a_1_1_pre_circ(X2))|~r1_xxreal_0(esk15_2(X1,X2),X2))).# trainpos
cnf(c_0_101, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_261, plain, (v2_xxreal_2(X1)|v1_xboole_0(X1)|~v1_finset_1(X1)|~v2_membered(X1))).# trainpos
cnf(c_0_142, plain, (v1_xxreal_0(k1_xxreal_2(X1))|~v2_membered(X1))).# trainpos
cnf(c_0_164, plain, (v7_ordinal1(k1_xxreal_2(X1))|~v2_xxreal_2(X1)|~v6_membered(X1))).# trainpos
cnf(c_0_150, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_253, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_88, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_177, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_192, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_128, plain, (esk15_2(X1,X2)=X1|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_1_pre_circ(X2)))).# trainpos
cnf(c_0_129, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_104, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_232, plain, (r1_tarski(X1,X2)|r2_hidden(esk2_2(X1,X2),X1))).# trainpos
cnf(c_0_235, plain, (r1_tarski(X1,X2)|~r2_hidden(esk2_2(X1,X2),X2))).# trainpos
cnf(c_0_256, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_76, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_263, plain, (v6_membered(X2)|~v6_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_170, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_62, plain, (m2_subset_1(k1_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_197, plain, (r1_xxreal_0(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_82, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_131, plain, (r2_hidden(X1,X2)|X1!=k1_xxreal_2(X2)|~v2_membered(X2)|~v2_xxreal_2(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_63, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_64, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_231, plain, (m1_subset_1(esk15_2(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_1_pre_circ(X2)))).# trainpos
cnf(c_0_168, plain, (r1_xxreal_0(X1,X3)|r2_hidden(X2,a_1_1_pre_circ(X3))|X1!=X2|~v7_ordinal1(X3)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_135, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_216, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_75, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_166, plain, (r1_xxreal_0(X1,X3)|X3!=k1_xxreal_2(X2)|~v2_membered(X2)|~v2_xxreal_2(X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_99, plain, (r1_xxreal_0(X2,X1)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(k2_xcmplx_0(X2,X4),k2_xcmplx_0(X1,X3)))).# trainpos
cnf(c_0_237, plain, (esk15_2(esk2_2(a_1_1_pre_circ(X1),X2),X1)=esk2_2(a_1_1_pre_circ(X1),X2)|r1_tarski(a_1_1_pre_circ(X1),X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_243, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(esk2_2(X1,X2),X2))).# trainpos
cnf(c_0_139, plain, (esk15_2(X1,X2)=X1|v1_xboole_0(a_1_1_pre_circ(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,a_1_1_pre_circ(X2)))).# trainpos
cnf(c_0_273, plain, (k1_xboole_0=np__0)).# trainpos
cnf(c_0_203, plain, (v1_xreal_0(np__1))).# trainpos
cnf(c_0_105, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_178, plain, (r1_xxreal_0(X1,X2)|r2_hidden(X1,a_1_1_pre_circ(X2))|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_114, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_141, plain, (r2_hidden(k1_xxreal_2(X1),X1)|~v1_xxreal_0(k1_xxreal_2(X1))|~v2_xxreal_2(X1)|~v2_membered(X1))).# trainpos
cnf(c_0_149, plain, (v1_xboole_0(a_1_1_pre_circ(X2))|~r1_xxreal_0(X1,X2)|~r2_hidden(X1,a_1_1_pre_circ(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,a_1_1_pre_circ(X2)))).# trainpos
cnf(c_0_276, plain, (X1=np__0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_100, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_95, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_78, plain, (k1_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_77, plain, (m2_subset_1(k1_nat_1(X2,X1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_223, plain, (m1_subset_1(k1_nat_1(X1,X2),k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_176, plain, (r1_xxreal_0(X1,X2)|v1_xboole_0(X3)|X2!=k1_xxreal_2(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v2_xxreal_2(X3)|~v2_membered(X3)|~m1_subset_1(X1,X3))).# trainpos
cnf(c_0_202, plain, (v1_xreal_0(np__0))).# trainpos
cnf(c_0_125, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_112, plain, (v1_xxreal_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_228, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_163, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_162, plain, (m1_subset_1(k1_xxreal_2(X1),X1)|~v2_xxreal_2(X1)|~v2_membered(X1))).# trainpos
cnf(c_0_189, plain, (m1_subset_1(X1,a_1_1_pre_circ(X2))|r1_xxreal_0(X1,X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_137, plain, (r1_xxreal_0(X1,esk12_0)|r1_xxreal_0(esk12_0,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_173, plain, (v1_xboole_0(a_1_1_pre_circ(X1))|~r1_xxreal_0(k1_xxreal_2(a_1_1_pre_circ(X1)),X1)|~v7_ordinal1(X1)|~v2_xxreal_2(a_1_1_pre_circ(X1))|~v2_membered(a_1_1_pre_circ(X1)))).# trainpos
cnf(c_0_174, negated_conjecture, (r1_xxreal_0(esk1_0,k1_xxreal_2(X1))|r1_xxreal_0(k1_xxreal_2(X1),esk1_0)|~v6_membered(X1)|~v2_xxreal_2(X1))).# trainpos
cnf(c_0_147, plain, (r1_xxreal_0(esk12_0,esk12_0))).# trainpos
cnf(c_0_199, plain, (v1_xboole_0(a_1_1_pre_circ(X1))|r1_xxreal_0(X2,X3)|r1_xxreal_0(X2,X1)|X3!=k1_xxreal_2(a_1_1_pre_circ(X1))|~v1_xxreal_0(X3)|~v2_xxreal_2(a_1_1_pre_circ(X1))|~v2_membered(a_1_1_pre_circ(X1))|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_187, plain, (~v1_xboole_0(X1)|~v2_xxreal_2(X1)|~v2_membered(X1))).# trainpos
cnf(c_0_92, plain, (m2_subset_1(k2_xcmplx_0(X2,X1),k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_103, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_186, negated_conjecture, (v1_xboole_0(a_1_1_pre_circ(esk1_0))|r1_xxreal_0(esk1_0,k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))|~v6_membered(a_1_1_pre_circ(esk1_0)))).# trainpos
cnf(c_0_206, plain, (v1_xboole_0(a_1_1_pre_circ(X1))|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,k1_xxreal_2(a_1_1_pre_circ(X1)))|~m1_subset_1(X2,k5_numbers)|~v1_xxreal_0(k1_xxreal_2(a_1_1_pre_circ(X1)))|~v2_xxreal_2(a_1_1_pre_circ(X1))|~v2_membered(a_1_1_pre_circ(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_102, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_121, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),k2_xcmplx_0(X2,X4))|~r1_xxreal_0(X4,X3))).# trainpos
cnf(c_0_111, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X2)|~v1_xxreal_0(k2_xcmplx_0(X1,X2)))).# trainpos
cnf(c_0_278, plain, (r1_xxreal_0(X2,X1)|~v1_xboole_0(a_1_1_pre_circ(X1))|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_134, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),k2_xcmplx_0(X4,X2))|~r1_xxreal_0(X4,X3))).# trainpos
cnf(c_0_113, plain, (m1_subset_1(k2_xcmplx_0(X2,X1),k1_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_123, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X2)|~m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers))).# trainpos
cnf(c_0_136, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_205, negated_conjecture, (r1_xxreal_0(X1,k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~r1_xxreal_0(X1,esk1_0)|~v1_xxreal_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~v1_xxreal_0(esk1_0)|~v1_xxreal_0(X1)|~v6_membered(a_1_1_pre_circ(esk1_0))|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0)))).# trainpos
cnf(c_0_159, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_266, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers))).# trainpos
cnf(c_0_145, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X4,X5))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),k2_xcmplx_0(X4,k2_xcmplx_0(X5,X2)))|~r1_xxreal_0(k2_xcmplx_0(X4,X5),X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X4))).# trainpos
cnf(c_0_272, plain, (v2_membered(k5_numbers))).# trainpos
cnf(c_0_169, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~v1_xreal_0(k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2)))|~m1_subset_1(k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2)),k5_numbers))).# trainpos
cnf(c_0_190, plain, (r1_xxreal_0(k2_xcmplx_0(X2,X1),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(k2_xcmplx_0(X3,k2_xcmplx_0(X2,X1)),k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_208, plain, (~v1_xreal_0(X1)|~m1_subset_1(k2_xcmplx_0(X1,np__1),k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X1,np__1),X1))).# trainpos
cnf(c_0_213, negated_conjecture, (~v1_xreal_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~m1_subset_1(k2_xcmplx_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)),np__1),k5_numbers)|~v6_membered(a_1_1_pre_circ(esk1_0))|~r1_xxreal_0(k2_xcmplx_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)),np__1),esk1_0)|~v1_xxreal_0(esk1_0)|~v1_xxreal_0(k2_xcmplx_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)),np__1))|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0)))).# trainpos
cnf(c_0_214, plain, (r1_xxreal_0(k2_xcmplx_0(k1_xxreal_2(a_1_1_pre_circ(X1)),np__1),X1)|~v1_xreal_0(k1_xxreal_2(a_1_1_pre_circ(X1)))|~m1_subset_1(k2_xcmplx_0(k1_xxreal_2(a_1_1_pre_circ(X1)),np__1),k5_numbers)|~v2_xxreal_2(a_1_1_pre_circ(X1))|~v2_membered(a_1_1_pre_circ(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_217, negated_conjecture, (~v6_membered(a_1_1_pre_circ(esk1_0))|~v1_xreal_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~m1_subset_1(k2_xcmplx_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)),np__1),k5_numbers)|~v1_xxreal_0(k2_xcmplx_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)),np__1))|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0)))).# trainpos
cnf(c_0_221, negated_conjecture, (~v6_membered(a_1_1_pre_circ(esk1_0))|~v1_xreal_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~m1_subset_1(k2_xcmplx_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)),np__1),k5_numbers)|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))|~v7_ordinal1(k2_xcmplx_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)),np__1)))).# trainpos
cnf(c_0_233, negated_conjecture, (~v6_membered(a_1_1_pre_circ(esk1_0))|~v1_xreal_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))|~v7_ordinal1(k1_xxreal_2(a_1_1_pre_circ(esk1_0))))).# trainpos
cnf(c_0_249, negated_conjecture, (~v6_membered(a_1_1_pre_circ(esk1_0))|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))|~v7_ordinal1(k1_xxreal_2(a_1_1_pre_circ(esk1_0))))).# trainpos
cnf(c_0_236, plain, (m1_subset_1(esk15_2(esk2_2(a_1_1_pre_circ(X1),X2),X1),k5_numbers)|r1_tarski(a_1_1_pre_circ(X1),X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_244, plain, (m1_subset_1(esk2_2(a_1_1_pre_circ(X1),X2),k5_numbers)|r1_tarski(a_1_1_pre_circ(X1),X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_270, negated_conjecture, (v1_xboole_0(a_1_1_pre_circ(esk1_0))|~v6_membered(a_1_1_pre_circ(esk1_0))|~v1_xxreal_0(esk1_0)|~v2_membered(a_1_1_pre_circ(esk1_0)))).# trainpos
cnf(c_0_254, plain, (r1_tarski(a_1_1_pre_circ(X1),k5_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_264, plain, (m1_subset_1(a_1_1_pre_circ(X1),k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_275, plain, (v2_membered(a_1_1_pre_circ(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_271, plain, (v6_membered(a_1_1_pre_circ(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_274, negated_conjecture, (v1_xboole_0(a_1_1_pre_circ(esk1_0))|~v1_xxreal_0(esk1_0)|~v2_membered(a_1_1_pre_circ(esk1_0)))).# trainpos
cnf(c_0_277, negated_conjecture, (v1_xboole_0(a_1_1_pre_circ(esk1_0))|~v1_xxreal_0(esk1_0))).# trainpos
cnf(c_0_279, negated_conjecture, (a_1_1_pre_circ(esk1_0)=np__0|~v1_xxreal_0(esk1_0))).# trainpos
cnf(c_0_280, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|~m1_subset_1(X1,k5_numbers)|~v1_xxreal_0(esk1_0))).# trainpos
cnf(c_0_281, negated_conjecture, (~v1_xreal_0(esk1_0)|~m1_subset_1(k2_xcmplx_0(esk1_0,np__1),k5_numbers)|~v1_xxreal_0(esk1_0))).# trainpos
cnf(c_0_283, negated_conjecture, (~v1_xreal_0(esk1_0))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_52, plain, (o_0_0_xboole_0=k1_xboole_0)).#trainneg
cnf(i_0_101, plain, (v7_ordinal1(esk6_0))).#trainneg
cnf(i_0_109, plain, (v1_finset_1(esk7_0))).#trainneg
cnf(i_0_132, plain, (v1_finset_1(esk10_0))).#trainneg
cnf(i_0_142, plain, (v1_finset_1(esk11_0))).#trainneg
cnf(i_0_181, plain, (v1_finset_1(esk16_0))).#trainneg
cnf(i_0_194, plain, (v1_finset_1(esk17_0))).#trainneg
cnf(i_0_207, plain, (v1_finset_1(esk20_0))).#trainneg
cnf(i_0_213, plain, (v1_finset_1(esk21_0))).#trainneg
cnf(i_0_62, plain, (v1_finset_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_235, plain, (v1_finset_1(esk24_0))).#trainneg
cnf(i_0_67, plain, (~v1_finset_1(k4_ordinal1))).#trainneg
cnf(i_0_239, plain, (v1_finset_1(esk25_0))).#trainneg
cnf(i_0_17, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_197, plain, (~v1_finset_1(esk18_0))).#trainneg
cnf(i_0_43, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_70, plain, (v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_85, plain, (v1_xxreal_0(esk4_0))).#trainneg
cnf(i_0_58, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_87, plain, (v1_xboole_0(esk4_0))).#trainneg
cnf(i_0_151, plain, (v1_xboole_0(esk13_0))).#trainneg
cnf(i_0_84, plain, (v1_xreal_0(esk4_0))).#trainneg
cnf(i_0_221, plain, (v1_xreal_0(esk22_0))).#trainneg
cnf(i_0_26, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_193, plain, (v5_finset_1(esk17_0))).#trainneg
cnf(i_0_86, plain, (v1_xcmplx_0(esk4_0))).#trainneg
cnf(i_0_224, plain, (v1_xcmplx_0(esk22_0))).#trainneg
cnf(i_0_122, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_102, plain, (~v1_xboole_0(esk6_0))).#trainneg
cnf(i_0_148, plain, (v1_finset_1(X1)|~v1_finset_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_47, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_46, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_110, plain, (~v1_xboole_0(esk7_0))).#trainneg
cnf(i_0_133, plain, (~v1_xboole_0(esk10_0))).#trainneg
cnf(i_0_88, plain, (v3_ordinal1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_222, plain, (v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_114, plain, (v1_relat_1(esk7_0))).#trainneg
cnf(i_0_136, plain, (v1_relat_1(esk10_0))).#trainneg
cnf(i_0_185, plain, (v1_relat_1(esk16_0))).#trainneg
cnf(i_0_199, plain, (v1_relat_1(esk18_0))).#trainneg
cnf(i_0_209, plain, (v1_relat_1(esk20_0))).#trainneg
cnf(i_0_215, plain, (v1_relat_1(esk21_0))).#trainneg
cnf(i_0_260, plain, (v1_relat_1(esk28_0))).#trainneg
cnf(i_0_267, plain, (v1_relat_1(esk29_0))).#trainneg
cnf(i_0_270, plain, (v1_relat_1(esk30_0))).#trainneg
cnf(i_0_204, plain, (v1_card_1(k4_ordinal1))).#trainneg
cnf(i_0_166, plain, (v5_card_3(k4_ordinal1))).#trainneg
cnf(i_0_111, plain, (v1_funct_1(esk7_0))).#trainneg
cnf(i_0_134, plain, (v1_funct_1(esk10_0))).#trainneg
cnf(i_0_143, plain, (~v1_xboole_0(esk11_0))).#trainneg
cnf(i_0_182, plain, (~v1_xboole_0(esk16_0))).#trainneg
cnf(i_0_183, plain, (v1_funct_1(esk16_0))).#trainneg
cnf(i_0_195, plain, (~v1_xboole_0(esk17_0))).#trainneg
cnf(i_0_198, plain, (v1_funct_1(esk18_0))).#trainneg
cnf(i_0_203, plain, (~v1_xboole_0(esk19_0))).#trainneg
cnf(i_0_210, plain, (~v1_xboole_0(esk20_0))).#trainneg
cnf(i_0_82, plain, (v1_card_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_153, plain, (v4_card_3(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_208, plain, (v1_funct_1(esk20_0))).#trainneg
cnf(i_0_214, plain, (v1_funct_1(esk21_0))).#trainneg
cnf(i_0_259, plain, (v1_funct_1(esk28_0))).#trainneg
cnf(i_0_266, plain, (v1_funct_1(esk29_0))).#trainneg
cnf(i_0_269, plain, (v1_funct_1(esk30_0))).#trainneg
cnf(i_0_113, plain, (v2_relat_1(esk7_0))).#trainneg
cnf(i_0_108, plain, (v1_finseq_1(esk7_0))).#trainneg
cnf(i_0_131, plain, (v1_finseq_1(esk10_0))).#trainneg
cnf(i_0_180, plain, (v1_finseq_1(esk16_0))).#trainneg
cnf(i_0_107, plain, (v2_finseq_1(esk7_0))).#trainneg
cnf(i_0_130, plain, (v2_finseq_1(esk10_0))).#trainneg
cnf(i_0_179, plain, (v2_finseq_1(esk16_0))).#trainneg
cnf(i_0_129, plain, (v6_trees_3(esk10_0))).#trainneg
cnf(i_0_264, plain, (v6_trees_3(esk29_0))).#trainneg
cnf(i_0_236, plain, (~v1_xboole_0(esk24_0))).#trainneg
cnf(i_0_240, plain, (~v1_xboole_0(esk25_0))).#trainneg
cnf(i_0_238, plain, (v1_trees_3(esk25_0))).#trainneg
cnf(i_0_261, plain, (~v1_xboole_0(esk28_0))).#trainneg
cnf(i_0_141, plain, (v1_trees_1(esk11_0))).#trainneg
cnf(i_0_29, plain, (r1_xxreal_0(np__0,np__0))).#trainneg
cnf(i_0_265, plain, (~v1_xboole_0(esk29_0))).#trainneg
cnf(i_0_119, plain, (~v4_xxreal_2(k1_numbers))).#trainneg
cnf(i_0_126, plain, (v6_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_202, plain, (v1_trees_1(esk19_0))).#trainneg
cnf(i_0_178, plain, (v4_trees_3(esk16_0))).#trainneg
cnf(i_0_177, plain, (v5_trees_3(esk16_0))).#trainneg
cnf(i_0_234, plain, (v3_trees_3(esk24_0))).#trainneg
cnf(i_0_212, plain, (v3_trees_2(esk21_0))).#trainneg
cnf(i_0_268, plain, (v3_trees_2(esk30_0))).#trainneg
cnf(i_0_237, plain, (v2_trees_3(esk25_0))).#trainneg
cnf(i_0_258, plain, (v2_finset_1(esk28_0))).#trainneg
cnf(i_0_120, plain, (~v3_xxreal_2(k1_numbers))).#trainneg
cnf(i_0_117, plain, (v1_xboole_0(esk8_1(X1)))).#trainneg
cnf(i_0_228, plain, (v1_relat_1(esk23_1(X1)))).#trainneg
cnf(i_0_225, plain, (v1_funct_1(esk23_1(X1)))).#trainneg
cnf(i_0_227, plain, (v2_relat_1(esk23_1(X1)))).#trainneg
cnf(i_0_54, plain, (k2_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_96, plain, (~v1_finset_1(X1)|~v5_card_3(X1))).#trainneg
cnf(i_0_53, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_125, plain, (r1_xxreal_0(np__0,np__1))).#trainneg
cnf(i_0_80, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_68, plain, (r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_16, plain, (m1_subset_1(np__0,k5_numbers))).#trainneg
cnf(i_0_127, plain, (v3_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_244, plain, (v5_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_61, plain, (k2_xcmplx_0(np__1,np__0)=np__1)).#trainneg
cnf(i_0_72, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_112, plain, (v4_relat_1(esk7_0,k5_numbers))).#trainneg
cnf(i_0_135, plain, (v4_relat_1(esk10_0,k5_numbers))).#trainneg
cnf(i_0_184, plain, (v4_relat_1(esk16_0,k5_numbers))).#trainneg
cnf(i_0_33, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_123, plain, (m1_subset_1(esk9_1(X1),X1))).#trainneg
cnf(i_0_226, plain, (v4_relat_1(esk23_1(X1),X1))).#trainneg
cnf(i_0_250, plain, (~v1_subset_1(esk26_1(X1),X1))).#trainneg
cnf(i_0_63, plain, (~r2_hidden(X2,X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_272, plain, (v1_xboole_0(X1)|v1_finset_1(esk31_1(X1)))).#trainneg
cnf(i_0_81, plain, (v3_ordinal1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_192, plain, (v1_xboole_0(X1)|~v3_card_1(X1,k1_xboole_0))).#trainneg
cnf(i_0_100, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_76, plain, (v1_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_91, plain, (v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_89, plain, (v3_membered(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_124, plain, (v1_card_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_128, plain, (v4_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_94, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_95, plain, (v4_card_3(X1)|~v5_card_3(X1))).#trainneg
cnf(i_0_97, plain, (v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_190, plain, (v3_xxreal_2(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_243, plain, (v1_trees_3(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_257, plain, (v1_trees_3(X1)|~v2_trees_3(X1))).#trainneg
cnf(i_0_150, plain, (v1_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_149, plain, (v2_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_241, plain, (v3_trees_3(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_242, plain, (v2_trees_3(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_263, plain, (v3_finseq_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_65, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_152, plain, (v1_finset_1(k1_zfmisc_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_252, plain, (v5_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_74, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_254, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk27_1(X1)))).#trainneg
cnf(i_0_273, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk31_1(X1)))).#trainneg
cnf(i_0_205, plain, (v5_xxreal_2(X1)|~v1_finset_1(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_159, plain, (v1_finset_1(esk14_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_118, plain, (m1_subset_1(esk8_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_251, plain, (m1_subset_1(esk26_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_35, plain, (v5_finset_1(k1_zfmisc_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_162, plain, (v1_relat_1(esk14_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_160, plain, (v1_funct_1(esk14_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_157, plain, (v1_finseq_1(esk14_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_156, plain, (v2_finseq_1(esk14_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_271, plain, (v6_membered(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_161, plain, (v4_relat_1(esk14_1(X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_158, plain, (v3_card_1(esk14_1(X1),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_189, plain, (v6_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_83, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_186, plain, (v1_finset_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_31, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_98, plain, (~r2_hidden(X3,X2)|~r2_hidden(X1,X2)|~r2_hidden(X1,esk5_2(X3,X2)))).#trainneg
cnf(i_0_262, plain, (v3_ordinal1(X1)|~v1_ordinal1(X1)|~v2_ordinal1(X1))).#trainneg
cnf(i_0_232, plain, (v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_140, plain, (v1_xxreal_0(k1_xxreal_2(X1))|~v2_xxreal_2(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_147, plain, (v1_xxreal_0(k1_xxreal_2(X1))|~v2_xxreal_2(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_246, plain, (v1_xboole_0(X1)|v1_xxreal_2(X1)|~v1_finset_1(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_216, plain, (v6_trees_3(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_146, plain, (v1_xreal_0(k1_xxreal_2(X1))|~v2_xxreal_2(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_218, plain, (v4_trees_3(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_255, plain, (v1_xboole_0(X1)|m1_subset_1(esk27_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_274, plain, (v1_xboole_0(X1)|m1_subset_1(esk31_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_42, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_30, plain, (v1_finset_1(X2)|~v5_finset_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_78, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_121, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_201, plain, (v1_relat_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_36, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_200, plain, (v1_funct_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_217, plain, (v5_trees_3(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_104, plain, (v1_xboole_0(X1)|v1_xxreal_0(k1_xxreal_2(X1))|~v3_membered(X1)|~v4_xxreal_2(X1))).#trainneg
cnf(i_0_165, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_249, plain, (v1_finseq_1(X2)|~v3_finseq_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_155, plain, (v1_int_1(X2)|~v5_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_103, plain, (v1_xboole_0(X1)|v1_xreal_0(k1_xxreal_2(X1))|~v3_membered(X1)|~v4_xxreal_2(X1))).#trainneg
cnf(i_0_229, plain, (v1_rat_1(X2)|~v4_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_99, plain, (r2_hidden(esk5_2(X1,X2),X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_77, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_154, plain, (v1_membered(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_79, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_256, plain, (v1_trees_3(X2)|~v1_trees_3(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_253, plain, (v3_trees_3(X2)|~v3_trees_3(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_138, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~v1_trees_3(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_5, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1))).#trainneg
cnf(i_0_6, plain, (X2=k1_xxreal_2(X1)|~v2_membered(X1)|~v2_xxreal_2(X1)|~v1_xxreal_0(X2)|~r2_hidden(X2,X1)|~r1_xxreal_0(esk3_2(X1,X2),X2))).#trainneg
cnf(i_0_137, plain, (v1_xboole_0(X2)|v1_trees_1(X1)|~v1_trees_3(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_196, plain, (v1_xboole_0(X1)|v3_trees_2(X2)|~v3_trees_3(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_275, plain, (v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_115, plain, (v2_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_28, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_34, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_90, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_55, plain, (m2_subset_1(k2_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_21, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_8, plain, (X2=k1_xxreal_2(X1)|v1_xxreal_0(esk3_2(X1,X2))|~v2_membered(X1)|~v2_xxreal_2(X1)|~v1_xxreal_0(X2)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_7, plain, (X2=k1_xxreal_2(X1)|r2_hidden(esk3_2(X1,X2),X1)|~v2_membered(X1)|~v2_xxreal_2(X1)|~v1_xxreal_0(X2)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_276, plain, (v1_finset_1(esk12_0))).#trainneg
cnf(i_0_286, plain, (esk13_0=k1_xboole_0)).#trainneg
cnf(i_0_280, plain, (~v1_xboole_0(esk18_0))).#trainneg
cnf(i_0_278, negated_conjecture, (v1_finset_1(esk1_0))).#trainneg
cnf(i_0_277, plain, (v1_finset_1(esk6_0))).#trainneg
cnf(i_0_551, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_511, plain, (v3_membered(k5_numbers))).#trainneg
cnf(i_0_544, plain, (~v1_trees_3(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_479, plain, (m1_subset_1(esk2_2(X1,X2),X1)|r1_tarski(X1,X2))).#trainneg
cnf(i_0_478, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_415, plain, (~v1_finset_1(k1_numbers))).#trainneg
cnf(i_0_482, plain, (esk15_2(esk5_2(X1,a_1_1_pre_circ(X2)),X2)=esk5_2(X1,a_1_1_pre_circ(X2))|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_1_pre_circ(X2)))).#trainneg
cnf(i_0_326, plain, (v7_ordinal1(np__1))).#trainneg
cnf(i_0_328, plain, (v7_ordinal1(np__0))).#trainneg
cnf(i_0_301, plain, (~v6_membered(k1_numbers))).#trainneg
cnf(i_0_337, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_646, plain, (r1_tarski(a_1_1_pre_circ(X1),X2)|~r1_xxreal_0(esk2_2(a_1_1_pre_circ(X1),X2),X1)|~r2_hidden(esk2_2(a_1_1_pre_circ(X1),X2),a_1_1_pre_circ(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_694, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_690, plain, (v1_xxreal_0(np__1))).#trainneg
cnf(i_0_627, plain, (r1_xxreal_0(esk2_2(X1,X2),X3)|r1_tarski(X1,X2)|X3!=k1_xxreal_2(X1)|~v1_xxreal_0(X3)|~v1_xxreal_0(esk2_2(X1,X2))|~v2_xxreal_2(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_608, plain, (m2_subset_1(k2_nat_1(X2,X1),k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_697, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_698, plain, (v1_finset_1(X1)|~v5_finset_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_286, plain, (esk13_0=np__0)).#trainneg
cnf(i_0_300, plain, (v4_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_52, plain, (o_0_0_xboole_0=np__0)).#trainneg
cnf(i_0_305, plain, (v1_finset_1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_309, plain, (X1=esk8_1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_595, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X2,X1)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_596, plain, (k2_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_597, plain, (k2_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_688, plain, (v1_finset_1(np__1))).#trainneg
cnf(i_0_617, plain, (m1_subset_1(k2_nat_1(X1,X2),k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_723, plain, (r1_tarski(X1,a_1_1_pre_circ(X2))|r1_xxreal_0(esk2_2(X1,a_1_1_pre_circ(X2)),X2)|~m1_subset_1(esk2_2(X1,a_1_1_pre_circ(X2)),k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_771, plain, (esk4_0=np__0)).#trainneg
cnf(i_0_192, plain, (v1_xboole_0(X1)|~v3_card_1(X1,np__0))).#trainneg
cnf(i_0_499, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_549, plain, (~v1_trees_3(k5_numbers))).#trainneg
cnf(i_0_772, plain, (v3_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_787, plain, (k2_nat_1(X2,X1)=k2_xcmplx_0(X2,X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_782, plain, (k1_nat_1(X2,X1)=k2_xcmplx_0(X2,X1)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_783, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_689, plain, (v1_finset_1(np__0))).#trainneg
cnf(i_0_766, plain, (v7_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_384, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_321, plain, (v7_ordinal1(esk26_1(X1))|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_386, plain, (v3_ordinal1(np__0))).#trainneg
cnf(i_0_304, plain, (~m1_subset_1(esk18_0,k5_numbers))).#trainneg
cnf(i_0_730, plain, (r1_xxreal_0(X1,esk22_0)|r1_xxreal_0(esk22_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_754, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_322, plain, (v7_ordinal1(esk9_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_323, plain, (v7_ordinal1(esk8_1(X1))|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_800, plain, (m1_subset_1(k1_nat_1(X1,X2),k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_760, plain, (m1_subset_1(k2_nat_1(X1,X2),k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_859, plain, (r1_xxreal_0(X1,X2)|r1_tarski(X3,a_1_1_pre_circ(X2))|~r1_xxreal_0(X1,esk2_2(X3,a_1_1_pre_circ(X2)))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk2_2(X3,a_1_1_pre_circ(X2)))|~v1_xxreal_0(X1)|~m1_subset_1(esk2_2(X3,a_1_1_pre_circ(X2)),k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_876, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_757, plain, (r1_tarski(a_1_1_pre_circ(X1),X2)|~v7_ordinal1(X1)|k1_xxreal_2(a_1_1_pre_circ(X1))!=X1|~v1_xxreal_0(esk2_2(a_1_1_pre_circ(X1),X2))|~v1_xxreal_0(X1)|~v2_xxreal_2(a_1_1_pre_circ(X1))|~v2_membered(a_1_1_pre_circ(X1)))).#trainneg
cnf(i_0_767, plain, (v7_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_86, plain, (v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_85, plain, (v1_xxreal_0(np__0))).#trainneg
cnf(i_0_912, plain, (r1_xxreal_0(esk22_0,esk12_0)|r1_xxreal_0(esk12_0,esk22_0))).#trainneg
cnf(i_0_306, plain, (~m1_subset_1(k5_numbers,k5_numbers))).#trainneg
cnf(i_0_866, plain, (v3_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_875, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_329, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_824, plain, (m1_subset_1(k2_xcmplx_0(X2,X1),k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_854, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_450, plain, (v1_int_1(np__1)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_354, plain, (~v5_finset_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_330, plain, (v1_finset_1(k1_xxreal_2(X1))|~v6_membered(X1)|~v2_xxreal_2(X1))).#trainneg
cnf(i_0_913, plain, (r1_xxreal_0(esk22_0,esk6_0)|r1_xxreal_0(esk6_0,esk22_0))).#trainneg
cnf(i_0_686, plain, (~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_867, plain, (v3_membered(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_333, plain, (v7_ordinal1(esk27_1(X1))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1023, plain, (r1_tarski(a_1_1_pre_circ(X1),X2)|k1_xxreal_2(a_1_1_pre_circ(X1))!=X1|~v2_xxreal_2(a_1_1_pre_circ(X1))|~v2_membered(a_1_1_pre_circ(X1))|~v7_ordinal1(X1)|~v7_ordinal1(esk2_2(a_1_1_pre_circ(X1),X2)))).#trainneg
cnf(i_0_1000, plain, (r1_xxreal_0(esk2_2(X1,a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(X3)))),X3)|r1_tarski(X2,a_1_1_pre_circ(X3))|r1_tarski(X1,a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(X3))))|~m1_subset_1(esk2_2(X2,a_1_1_pre_circ(X3)),k5_numbers)|~v1_xxreal_0(esk2_2(X2,a_1_1_pre_circ(X3)))|~v1_xxreal_0(esk2_2(X1,a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(X3)))))|~v7_ordinal1(X3)|~m1_subset_1(esk2_2(X1,a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(X3)))),k5_numbers)|~v7_ordinal1(esk2_2(X2,a_1_1_pre_circ(X3))))).#trainneg
cnf(i_0_1024, plain, (r1_tarski(a_1_1_pre_circ(X1),X2)|k1_xxreal_2(a_1_1_pre_circ(X1))!=X1|~v2_xxreal_2(a_1_1_pre_circ(X1))|~v2_membered(a_1_1_pre_circ(X1))|~v7_ordinal1(X1)|~m1_subset_1(esk2_2(a_1_1_pre_circ(X1),X2),k1_numbers))).#trainneg
cnf(i_0_1017, negated_conjecture, (r1_xxreal_0(esk1_0,esk1_0))).#trainneg
cnf(i_0_412, plain, (v1_finset_1(esk26_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_451, plain, (v1_int_1(np__1)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_335, plain, (v7_ordinal1(esk31_1(X1))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1021, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|r1_xxreal_0(X1,esk1_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1043, plain, (v1_finset_1(k1_xxreal_2(X1))|~v5_finset_1(X1)|~v2_xxreal_2(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_345, plain, (v1_finset_1(esk27_1(X1))|v1_xboole_0(X1)|~v5_finset_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_758, plain, (v1_xboole_0(a_1_1_pre_circ(X1))|~r1_xxreal_0(esk9_1(a_1_1_pre_circ(X1)),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_751, plain, (r1_xxreal_0(k1_xxreal_2(X1),X2)|X2!=k1_xxreal_2(X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(k1_xxreal_2(X1))|~v2_xxreal_2(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_878, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_727, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X1,X3)|X2!=k1_xxreal_2(a_1_1_pre_circ(X3))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v2_xxreal_2(a_1_1_pre_circ(X3))|~v2_membered(a_1_1_pre_circ(X3))|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1001, plain, (r1_xxreal_0(esk2_2(X1,X2),X3)|r1_tarski(X4,a_1_1_pre_circ(X3))|r1_tarski(X1,X2)|~m1_subset_1(esk2_2(X4,a_1_1_pre_circ(X3)),k5_numbers)|~v1_xxreal_0(esk2_2(X4,a_1_1_pre_circ(X3)))|~v1_xxreal_0(esk2_2(X1,X2))|~v7_ordinal1(X3)|esk2_2(X4,a_1_1_pre_circ(X3))!=k1_xxreal_2(X1)|~v2_xxreal_2(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_414, plain, (v1_finset_1(esk8_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_452, plain, (v1_int_1(np__0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_693, plain, (~m1_subset_1(k1_numbers,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_346, plain, (v1_finset_1(esk26_1(X1))|~v5_finset_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_914, negated_conjecture, (r1_xxreal_0(esk22_0,esk1_0)|r1_xxreal_0(esk1_0,esk22_0))).#trainneg
cnf(i_0_1048, plain, (v1_relat_1(k1_xxreal_2(X1))|~v4_funct_1(X1)|~v2_xxreal_2(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_1049, plain, (v1_funct_1(k1_xxreal_2(X1))|~v4_funct_1(X1)|~v2_xxreal_2(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_347, plain, (v1_finset_1(esk9_1(X1))|~v5_finset_1(X1))).#trainneg
cnf(i_0_348, plain, (v1_finset_1(esk8_1(X1))|~v5_finset_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_762, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_746, plain, (esk15_2(k1_xxreal_2(a_1_1_pre_circ(X1)),X1)=k1_xxreal_2(a_1_1_pre_circ(X1))|~v7_ordinal1(X1)|~v2_xxreal_2(a_1_1_pre_circ(X1))|~v2_membered(a_1_1_pre_circ(X1)))).#trainneg
cnf(i_0_874, plain, (r1_xxreal_0(X1,k1_xxreal_2(X2))|r1_xxreal_0(k1_xxreal_2(X2),X1)|~v7_ordinal1(X1)|~v6_membered(X2)|~v2_xxreal_2(X2))).#trainneg
cnf(i_0_862, plain, (v1_xboole_0(esk14_1(np__0)))).#trainneg
cnf(i_0_453, plain, (v1_int_1(np__0)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_355, plain, (v1_xxreal_0(esk31_1(X1))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_915, plain, (r1_xxreal_0(esk22_0,np__1)|r1_xxreal_0(np__1,esk22_0))).#trainneg
cnf(i_0_870, plain, (r1_xxreal_0(X1,esk6_0)|r1_xxreal_0(esk6_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_356, plain, (v1_xxreal_0(esk27_1(X1))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_357, plain, (v1_xxreal_0(esk26_1(X1))|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1022, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|r1_xxreal_0(X1,esk1_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_798, plain, (m2_subset_1(k2_xcmplx_0(X2,X1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1081, plain, (v7_ordinal1(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_413, plain, (v1_finset_1(esk9_1(k1_zfmisc_1(X1)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_471, plain, (v1_rat_1(np__1)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_358, plain, (v1_xxreal_0(esk9_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_916, plain, (r1_xxreal_0(esk22_0,np__0)|r1_xxreal_0(np__0,esk22_0))).#trainneg
cnf(i_0_873, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_359, plain, (v1_xxreal_0(esk8_1(X1))|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_365, plain, (v1_xxreal_0(k5_numbers)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1178, plain, (v1_xboole_0(a_1_1_pre_circ(X1))|~v2_xxreal_2(a_1_1_pre_circ(X1))|~v2_membered(a_1_1_pre_circ(X1))|~v7_ordinal1(X1)|k1_xxreal_2(a_1_1_pre_circ(X1))!=X1|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1201, plain, (~r1_xxreal_0(k1_xxreal_2(a_1_1_pre_circ(X1)),X1)|~r2_hidden(k1_xxreal_2(a_1_1_pre_circ(X1)),a_1_1_pre_circ(X1))|~v7_ordinal1(X1)|~v2_xxreal_2(a_1_1_pre_circ(X1))|~v2_membered(a_1_1_pre_circ(X1)))).#trainneg
cnf(i_0_1153, plain, (r1_tarski(a_1_1_pre_circ(X1),a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(X1))))|r1_tarski(X2,a_1_1_pre_circ(X1))|~v7_ordinal1(X1)|~m1_subset_1(esk2_2(a_1_1_pre_circ(X1),a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(X1)))),k5_numbers)|~m1_subset_1(esk2_2(X2,a_1_1_pre_circ(X1)),k5_numbers)|~v1_xxreal_0(esk2_2(a_1_1_pre_circ(X1),a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(X1)))))|~v1_xxreal_0(esk2_2(X2,a_1_1_pre_circ(X1))))).#trainneg
cnf(i_0_1151, plain, (r1_xxreal_0(X1,X2)|r1_tarski(X3,a_1_1_pre_circ(esk2_2(X4,a_1_1_pre_circ(X2))))|r1_tarski(X4,a_1_1_pre_circ(X2))|~r1_xxreal_0(X1,esk2_2(X3,a_1_1_pre_circ(esk2_2(X4,a_1_1_pre_circ(X2)))))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk2_2(X3,a_1_1_pre_circ(esk2_2(X4,a_1_1_pre_circ(X2)))))|~v1_xxreal_0(X1)|~m1_subset_1(esk2_2(X3,a_1_1_pre_circ(esk2_2(X4,a_1_1_pre_circ(X2)))),k5_numbers)|~m1_subset_1(esk2_2(X4,a_1_1_pre_circ(X2)),k5_numbers)|~v1_xxreal_0(esk2_2(X4,a_1_1_pre_circ(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1209, plain, (esk14_1(np__0)=np__0)).#trainneg
cnf(i_0_366, plain, (v1_xreal_0(esk31_1(X1))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1222, plain, (r1_xxreal_0(esk6_0,esk6_0))).#trainneg
cnf(i_0_367, plain, (v1_xreal_0(esk27_1(X1))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_368, plain, (v1_xreal_0(esk26_1(X1))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1150, plain, (r1_xxreal_0(esk2_2(X1,a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(X4)))))),X4)|r1_tarski(X3,a_1_1_pre_circ(X4))|r1_tarski(X1,a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(X4))))))|r1_tarski(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(X4))))|~m1_subset_1(esk2_2(X3,a_1_1_pre_circ(X4)),k5_numbers)|~v1_xxreal_0(esk2_2(X3,a_1_1_pre_circ(X4)))|~v1_xxreal_0(esk2_2(X1,a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(X4)))))))|~v7_ordinal1(X4)|~m1_subset_1(esk2_2(X1,a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(X4)))))),k5_numbers)|~m1_subset_1(esk2_2(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(X4)))),k5_numbers)|~v1_xxreal_0(esk2_2(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(X4)))))|~v7_ordinal1(esk2_2(X3,a_1_1_pre_circ(X4))))).#trainneg
cnf(i_0_1011, plain, (r1_xxreal_0(k1_nat_1(X1,X2),X3)|r1_xxreal_0(X3,k1_nat_1(X1,X2))|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_877, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1194, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1082, plain, (v7_ordinal1(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_776, plain, (esk8_1(X1)=esk8_1(X2))).#trainneg
cnf(i_0_1327, plain, (v1_relat_1(np__0))).#trainneg
cnf(i_0_1439, plain, (v1_finset_1(esk8_1(X2))|~v1_finset_1(X1))).#trainneg
cnf(i_0_369, plain, (v1_xreal_0(esk9_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_1331, plain, (v4_relat_1(np__0,k5_numbers))).#trainneg
cnf(i_0_1328, plain, (v1_funct_1(np__0))).#trainneg
cnf(i_0_1329, plain, (v1_finseq_1(np__0))).#trainneg
cnf(i_0_370, plain, (v1_xreal_0(esk8_1(X1))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_376, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1309, plain, (v1_xboole_0(a_1_1_pre_circ(X1))|k1_xxreal_2(a_1_1_pre_circ(X1))!=X1|~v2_membered(a_1_1_pre_circ(X1))|~v7_ordinal1(X1)|~v1_finset_1(a_1_1_pre_circ(X1)))).#trainneg
cnf(i_0_1312, plain, (r1_tarski(a_1_1_pre_circ(X1),a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(X1))))|r1_tarski(X2,a_1_1_pre_circ(X1))|~m1_subset_1(esk2_2(a_1_1_pre_circ(X1),a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(X1)))),k5_numbers)|~m1_subset_1(esk2_2(X2,a_1_1_pre_circ(X1)),k5_numbers)|~v1_xxreal_0(esk2_2(X2,a_1_1_pre_circ(X1)))|~v7_ordinal1(X1)|~v7_ordinal1(esk2_2(a_1_1_pre_circ(X1),a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(X1))))))).#trainneg
cnf(i_0_1330, plain, (v2_finseq_1(np__0))).#trainneg
cnf(i_0_1332, plain, (v3_card_1(np__0,np__0))).#trainneg
cnf(i_0_472, plain, (v1_rat_1(np__1)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_377, plain, (v3_ordinal1(esk31_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_921, plain, (r1_xxreal_0(np__1,esk12_0)|r1_xxreal_0(esk12_0,np__1))).#trainneg
cnf(i_0_1078, plain, (r1_xxreal_0(esk8_1(X1),X2)|r1_xxreal_0(X2,esk8_1(X1))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_378, plain, (v3_ordinal1(esk27_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_379, plain, (v3_ordinal1(esk26_1(X1))|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1392, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X2)|~v7_ordinal1(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_1343, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1388, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X2)|~v1_xxreal_0(k2_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_394, plain, (v1_xcmplx_0(np__1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_473, plain, (v1_rat_1(np__0)|~v4_membered(k1_numbers))).#trainneg
cnf(i_0_750, plain, (~r2_hidden(X1,k1_xxreal_2(X1))|~v2_xxreal_2(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_380, plain, (v3_ordinal1(esk9_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_922, plain, (r1_xxreal_0(np__1,esk6_0)|r1_xxreal_0(esk6_0,np__1))).#trainneg
cnf(i_0_395, plain, (v1_xcmplx_0(np__1)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_460, plain, (r1_xxreal_0(X1,esk22_0)|r1_xxreal_0(esk22_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_381, plain, (v3_ordinal1(esk8_1(X1))|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_388, plain, (v1_xcmplx_0(esk31_1(X1))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1566, plain, (v1_xxreal_0(k2_xcmplx_0(X1,X2))|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1496, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(X4,X2))|~r1_xxreal_0(X4,X3))).#trainneg
cnf(i_0_1564, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X2)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1585, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_1453, plain, (v1_finset_1(esk8_1(X1)))).#trainneg
cnf(i_0_474, plain, (v1_rat_1(np__0)|~v4_membered(k5_numbers))).#trainneg
cnf(i_0_389, plain, (v1_xcmplx_0(esk27_1(X1))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1594, plain, (r1_xxreal_0(esk22_0,esk22_0))).#trainneg
cnf(i_0_405, plain, (v1_relat_1(np__1)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_753, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_390, plain, (v1_xcmplx_0(esk26_1(X1))|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_391, plain, (v1_xcmplx_0(esk9_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_1563, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X2)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1524, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X2)|~v7_ordinal1(k2_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_1529, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X2)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers))).#trainneg
cnf(i_0_1534, plain, (v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_406, plain, (v1_relat_1(np__1)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_392, plain, (v1_xcmplx_0(esk8_1(X1))|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_923, negated_conjecture, (r1_xxreal_0(np__1,esk1_0)|r1_xxreal_0(esk1_0,np__1))).#trainneg
cnf(i_0_918, plain, (r1_xxreal_0(esk22_0,X1)|r1_xxreal_0(X1,esk22_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_927, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_398, plain, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_399, plain, (v1_relat_1(esk31_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1376, plain, (r1_xxreal_0(np__1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(np__1,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X2,np__0))).#trainneg
cnf(i_0_1644, plain, (r1_xxreal_0(X1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_422, plain, (v1_funct_1(np__1)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_1755, negated_conjecture, (~v2_xxreal_2(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))))|~v2_membered(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))))|~v7_ordinal1(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~v6_membered(a_1_1_pre_circ(esk1_0))|~r1_xxreal_0(k1_xxreal_2(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))),esk1_0)|~v1_xxreal_0(esk1_0)|~v1_xxreal_0(k1_xxreal_2(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))))|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0)))).#trainneg
cnf(i_0_400, plain, (v1_relat_1(esk27_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1013, negated_conjecture, (r1_xxreal_0(esk1_0,esk12_0)|r1_xxreal_0(esk12_0,esk1_0))).#trainneg
cnf(i_0_423, plain, (v1_funct_1(np__1)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_1074, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_401, plain, (v1_relat_1(esk26_1(X1))|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_402, plain, (v1_relat_1(esk9_1(X1))|~v4_funct_1(X1))).#trainneg
cnf(i_0_1706, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X2)|~m1_subset_1(k2_xcmplx_0(X2,X1),k5_numbers))).#trainneg
cnf(i_0_1754, negated_conjecture, (r1_tarski(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))),X1)|~v7_ordinal1(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~v6_membered(a_1_1_pre_circ(esk1_0))|~r1_xxreal_0(esk2_2(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))),X1),esk1_0)|~v1_xxreal_0(esk1_0)|~v1_xxreal_0(esk2_2(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))),X1))|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0)))).#trainneg
cnf(i_0_571, plain, (v3_trees_2(np__1)|~v3_trees_3(k5_numbers))).#trainneg
cnf(i_0_1751, negated_conjecture, (~r2_hidden(X1,a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))))|~v7_ordinal1(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~v6_membered(a_1_1_pre_circ(esk1_0))|~r1_xxreal_0(esk15_2(X1,k1_xxreal_2(a_1_1_pre_circ(esk1_0))),esk1_0)|~v1_xxreal_0(esk1_0)|~v1_xxreal_0(esk15_2(X1,k1_xxreal_2(a_1_1_pre_circ(esk1_0))))|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0)))).#trainneg
cnf(i_0_403, plain, (v1_relat_1(esk8_1(X1))|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1016, negated_conjecture, (r1_xxreal_0(esk1_0,esk6_0)|r1_xxreal_0(esk6_0,esk1_0))).#trainneg
cnf(i_0_439, plain, (v1_finseq_1(np__1)|~v3_finseq_1(k1_numbers))).#trainneg
cnf(i_0_440, plain, (v1_finseq_1(np__1)|~v3_finseq_1(k5_numbers))).#trainneg
cnf(i_0_1140, plain, (r1_xxreal_0(esk12_0,X1)|r1_xxreal_0(X1,esk12_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1227, plain, (r1_xxreal_0(esk6_0,X1)|r1_xxreal_0(X1,esk6_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_409, plain, (v1_relat_1(k5_numbers)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1417, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1416, plain, (v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v2_membered(k1_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1559, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k2_xcmplx_0(X2,X1),k1_numbers)|~r1_xxreal_0(X2,X2))).#trainneg
cnf(i_0_929, plain, (v1_finset_1(esk9_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_573, plain, (v3_trees_2(np__0)|~v3_trees_3(k5_numbers))).#trainneg
cnf(i_0_411, plain, (v1_finset_1(esk27_1(X1))|v1_xboole_0(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_1276, plain, (r1_xxreal_0(np__0,X1)|r1_xxreal_0(X1,np__0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_416, plain, (v1_funct_1(esk31_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_417, plain, (v1_funct_1(esk27_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1366, plain, (r1_tarski(a_1_1_pre_circ(X1),a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(X1))))))|r1_tarski(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(X1))))|r1_tarski(X3,a_1_1_pre_circ(X1))|~v7_ordinal1(X1)|~m1_subset_1(esk2_2(a_1_1_pre_circ(X1),a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(X1)))))),k5_numbers)|~m1_subset_1(esk2_2(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(X1)))),k5_numbers)|~m1_subset_1(esk2_2(X3,a_1_1_pre_circ(X1)),k5_numbers)|~v1_xxreal_0(esk2_2(a_1_1_pre_circ(X1),a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(X1)))))))|~v1_xxreal_0(esk2_2(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(X1)))))|~v1_xxreal_0(esk2_2(X3,a_1_1_pre_circ(X1))))).#trainneg
cnf(i_0_1364, plain, (r1_xxreal_0(X1,X2)|r1_tarski(X3,a_1_1_pre_circ(esk2_2(X4,a_1_1_pre_circ(esk2_2(X5,a_1_1_pre_circ(X2))))))|r1_tarski(X4,a_1_1_pre_circ(esk2_2(X5,a_1_1_pre_circ(X2))))|r1_tarski(X5,a_1_1_pre_circ(X2))|~r1_xxreal_0(X1,esk2_2(X3,a_1_1_pre_circ(esk2_2(X4,a_1_1_pre_circ(esk2_2(X5,a_1_1_pre_circ(X2)))))))|~v1_xxreal_0(X2)|~v1_xxreal_0(esk2_2(X3,a_1_1_pre_circ(esk2_2(X4,a_1_1_pre_circ(esk2_2(X5,a_1_1_pre_circ(X2)))))))|~v1_xxreal_0(X1)|~m1_subset_1(esk2_2(X3,a_1_1_pre_circ(esk2_2(X4,a_1_1_pre_circ(esk2_2(X5,a_1_1_pre_circ(X2)))))),k5_numbers)|~m1_subset_1(esk2_2(X4,a_1_1_pre_circ(esk2_2(X5,a_1_1_pre_circ(X2)))),k5_numbers)|~m1_subset_1(esk2_2(X5,a_1_1_pre_circ(X2)),k5_numbers)|~v1_xxreal_0(esk2_2(X4,a_1_1_pre_circ(esk2_2(X5,a_1_1_pre_circ(X2)))))|~v1_xxreal_0(esk2_2(X5,a_1_1_pre_circ(X2)))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_517, plain, (v1_membered(k5_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1437, plain, (m1_subset_1(esk8_1(X2),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1771, negated_conjecture, (~v6_membered(a_1_1_pre_circ(esk1_0))|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))))|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))))|~v2_membered(a_1_1_pre_circ(esk1_0))|k1_xxreal_2(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))))!=esk1_0)).#trainneg
cnf(i_0_418, plain, (v1_funct_1(esk26_1(X1))|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1019, negated_conjecture, (r1_xxreal_0(esk1_0,np__0)|r1_xxreal_0(np__0,esk1_0))).#trainneg
cnf(i_0_1900, plain, (v1_membered(k5_numbers))).#trainneg
cnf(i_0_523, plain, (v1_relat_1(k5_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1917, plain, (v2_membered(esk8_1(X1))|~v2_membered(X2))).#trainneg
cnf(i_0_1919, plain, (v6_membered(esk8_1(X1))|~v6_membered(X2))).#trainneg
cnf(i_0_419, plain, (v1_funct_1(esk9_1(X1))|~v4_funct_1(X1))).#trainneg
cnf(i_0_426, plain, (v1_funct_1(k5_numbers)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1828, plain, (v1_xxreal_0(k2_xcmplx_0(X1,X2))|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1897, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~v7_ordinal1(k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))))).#trainneg
cnf(i_0_1899, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~v1_xxreal_0(k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))))).#trainneg
cnf(i_0_529, plain, (v1_trees_3(k5_numbers)|~v1_trees_3(k1_numbers))).#trainneg
cnf(i_0_427, plain, (v1_xboole_0(esk23_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1135, plain, (r1_xxreal_0(esk12_0,esk6_0)|r1_xxreal_0(esk6_0,esk12_0))).#trainneg
cnf(i_0_535, plain, (v3_trees_3(k5_numbers)|~v3_trees_3(k1_numbers))).#trainneg
cnf(i_0_1920, plain, (v3_membered(esk8_1(X1))|~v3_membered(X2))).#trainneg
cnf(i_0_1921, plain, (v1_membered(esk8_1(X1))|~v1_membered(X2))).#trainneg
cnf(i_0_443, plain, (v1_finseq_1(k5_numbers)|~v3_finseq_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_436, plain, (v1_finseq_1(esk9_1(X1))|~v3_finseq_1(X1))).#trainneg
cnf(i_0_1264, plain, (v2_xxreal_0(np__0)|r1_xxreal_0(np__0,esk22_0))).#trainneg
cnf(i_0_1806, negated_conjecture, (~v6_membered(a_1_1_pre_circ(esk1_0))|~r1_xxreal_0(esk5_2(X1,a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))),esk1_0)|~v1_xxreal_0(esk5_2(X1,a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))))|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))|~r2_hidden(esk5_2(X1,a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))),a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))))|~r2_hidden(X1,a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))))|~v7_ordinal1(k1_xxreal_2(a_1_1_pre_circ(esk1_0))))).#trainneg
cnf(i_0_433, plain, (v1_finseq_1(esk31_1(X1))|v1_xboole_0(X1)|~v3_finseq_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_454, plain, (v1_int_1(k5_numbers)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_475, plain, (v1_rat_1(k5_numbers)|~v4_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1922, plain, (v1_relat_1(esk8_1(X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_1923, plain, (v1_trees_3(esk8_1(X1))|~v1_trees_3(X2))).#trainneg
cnf(i_0_434, plain, (v1_finseq_1(esk27_1(X1))|v1_xboole_0(X1)|~v3_finseq_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_435, plain, (v1_finseq_1(esk26_1(X1))|~v3_finseq_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1802, negated_conjecture, (r1_tarski(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))),a_1_1_pre_circ(esk1_0))|~v6_membered(a_1_1_pre_circ(esk1_0))|~v1_xxreal_0(esk2_2(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))),a_1_1_pre_circ(esk1_0)))|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))|~m1_subset_1(esk2_2(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))),a_1_1_pre_circ(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_1375, plain, (r1_xxreal_0(np__0,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(np__1,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X2,np__1))).#trainneg
cnf(i_0_1546, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,k2_xcmplx_0(X1,X2))|~v7_ordinal1(X3)|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2054, negated_conjecture, (r1_tarski(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))),a_1_1_pre_circ(esk1_0))|~v6_membered(a_1_1_pre_circ(esk1_0))|~m1_subset_1(esk2_2(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))),a_1_1_pre_circ(esk1_0)),k5_numbers)|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))|~m1_subset_1(esk2_2(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))),a_1_1_pre_circ(esk1_0)),k1_numbers))).#trainneg
cnf(i_0_2053, negated_conjecture, (r1_tarski(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))),a_1_1_pre_circ(esk1_0))|~v6_membered(a_1_1_pre_circ(esk1_0))|~m1_subset_1(esk2_2(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))),a_1_1_pre_circ(esk1_0)),k5_numbers)|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))|~v7_ordinal1(esk2_2(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))),a_1_1_pre_circ(esk1_0))))).#trainneg
cnf(i_0_1935, plain, (v2_membered(esk8_1(X1)))).#trainneg
cnf(i_0_1931, negated_conjecture, (v1_xboole_0(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))))|k1_xxreal_2(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))))!=esk1_0|~v6_membered(a_1_1_pre_circ(esk1_0))|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))))|~v2_membered(a_1_1_pre_circ(esk1_0))|~v1_finset_1(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))))).#trainneg
cnf(i_0_444, plain, (v1_int_1(esk31_1(X1))|v1_xboole_0(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_574, plain, (v3_trees_2(k5_numbers)|~v3_trees_3(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2042, plain, (v1_relat_1(esk8_1(X1))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1924, plain, (v3_trees_3(esk8_1(X1))|~v3_trees_3(X2))).#trainneg
cnf(i_0_445, plain, (v1_int_1(esk27_1(X1))|v1_xboole_0(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_446, plain, (v1_int_1(esk26_1(X1))|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1969, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2001, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~m1_subset_1(k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2)),k1_numbers))).#trainneg
cnf(i_0_1971, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~m1_subset_1(X3,k5_numbers)|~v7_ordinal1(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_1387, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X3,X3)|~v1_xxreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_699, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1937, plain, (v6_membered(esk8_1(X1)))).#trainneg
cnf(i_0_447, plain, (v1_int_1(esk9_1(X1))|~v5_membered(X1))).#trainneg
cnf(i_0_1138, plain, (r1_xxreal_0(esk12_0,np__0)|r1_xxreal_0(np__0,esk12_0))).#trainneg
cnf(i_0_1280, plain, (v1_xxreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1742, plain, (v1_xcmplx_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2022, plain, (v3_membered(esk8_1(X1))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2025, plain, (v1_membered(esk8_1(X1))|~v3_membered(X2))).#trainneg
cnf(i_0_462, plain, (r1_xxreal_0(X1,k1_xxreal_2(X2))|r1_xxreal_0(k1_xxreal_2(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(X2))).#trainneg
cnf(i_0_463, plain, (r1_xxreal_0(X1,k1_xxreal_2(X2))|r1_xxreal_0(k1_xxreal_2(X2),X1)|~v1_xxreal_0(X1)|~v6_membered(X2)|~v2_xxreal_2(X2))).#trainneg
cnf(i_0_1363, plain, (r1_xxreal_0(esk2_2(X1,a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(esk2_2(X4,a_1_1_pre_circ(X5)))))))),X5)|r1_tarski(X4,a_1_1_pre_circ(X5))|r1_tarski(X1,a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(esk2_2(X4,a_1_1_pre_circ(X5))))))))|r1_tarski(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(esk2_2(X4,a_1_1_pre_circ(X5))))))|r1_tarski(X3,a_1_1_pre_circ(esk2_2(X4,a_1_1_pre_circ(X5))))|~m1_subset_1(esk2_2(X4,a_1_1_pre_circ(X5)),k5_numbers)|~v1_xxreal_0(esk2_2(X4,a_1_1_pre_circ(X5)))|~v1_xxreal_0(esk2_2(X1,a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(esk2_2(X4,a_1_1_pre_circ(X5)))))))))|~v7_ordinal1(X5)|~m1_subset_1(esk2_2(X1,a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(esk2_2(X4,a_1_1_pre_circ(X5)))))))),k5_numbers)|~m1_subset_1(esk2_2(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(esk2_2(X4,a_1_1_pre_circ(X5)))))),k5_numbers)|~m1_subset_1(esk2_2(X3,a_1_1_pre_circ(esk2_2(X4,a_1_1_pre_circ(X5)))),k5_numbers)|~v1_xxreal_0(esk2_2(X2,a_1_1_pre_circ(esk2_2(X3,a_1_1_pre_circ(esk2_2(X4,a_1_1_pre_circ(X5)))))))|~v1_xxreal_0(esk2_2(X3,a_1_1_pre_circ(esk2_2(X4,a_1_1_pre_circ(X5)))))|~v7_ordinal1(esk2_2(X4,a_1_1_pre_circ(X5))))).#trainneg
cnf(i_0_1545, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,k2_xcmplx_0(X1,X2))|~v1_xxreal_0(X3)|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2180, plain, (r1_xxreal_0(k1_xxreal_2(X1),k1_xxreal_2(X2))|r1_xxreal_0(k1_xxreal_2(X2),k1_xxreal_2(X1))|~v2_membered(X1)|~v2_membered(X2))).#trainneg
cnf(i_0_2188, plain, (r1_xxreal_0(k1_xxreal_2(X1),X2)|r1_xxreal_0(X2,k1_xxreal_2(X1))|~v1_xxreal_0(X2)|~v2_xxreal_2(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2182, plain, (r1_xxreal_0(k1_xxreal_2(X1),k1_xxreal_2(X2))|r1_xxreal_0(k1_xxreal_2(X2),k1_xxreal_2(X1))|~v2_membered(X1)|~v3_membered(X2)|~v2_xxreal_2(X2))).#trainneg
cnf(i_0_2097, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X3))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X1,X3),X3)|~v7_ordinal1(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_700, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2019, plain, (v3_membered(esk8_1(X1)))).#trainneg
cnf(i_0_1035, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k1_xxreal_2(k1_zfmisc_1(X1)))|~v2_xxreal_2(k1_zfmisc_1(X1))|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_464, plain, (r1_xxreal_0(X1,k1_xxreal_2(X2))|r1_xxreal_0(k1_xxreal_2(X2),X1)|~v1_xxreal_0(X1)|~v3_membered(X2)|~v2_xxreal_2(X2))).#trainneg
cnf(i_0_2029, plain, (v1_int_1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2030, plain, (v1_rat_1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2044, plain, (v1_relat_1(esk8_1(X1))|~v4_funct_1(X2))).#trainneg
cnf(i_0_2048, plain, (v1_relat_1(esk8_1(X1))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_465, plain, (v1_rat_1(esk31_1(X1))|v1_xboole_0(X1)|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_466, plain, (v1_rat_1(esk27_1(X1))|v1_xboole_0(X1)|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2183, plain, (r1_xxreal_0(k1_xxreal_2(X1),X2)|r1_xxreal_0(X2,k1_xxreal_2(X1))|~v2_membered(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2124, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X3,k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X1,X3),X3)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2098, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X3))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X1,X3),X3)|~v7_ordinal1(X3)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2184, plain, (r1_xxreal_0(k1_xxreal_2(X1),X2)|r1_xxreal_0(X2,k1_xxreal_2(X1))|~v2_membered(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2108, plain, (r1_xxreal_0(k2_xcmplx_0(X2,X1),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(k2_xcmplx_0(X3,k2_xcmplx_0(X2,X1)),k1_numbers)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_701, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2024, plain, (v1_membered(esk8_1(X1)))).#trainneg
cnf(i_0_315, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk26_1(X1)))).#trainneg
cnf(i_0_467, plain, (v1_rat_1(esk26_1(X1))|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2050, plain, (v1_trees_3(esk8_1(X1))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2082, plain, (v1_relat_1(esk8_1(X1))|~v6_membered(X2))).#trainneg
cnf(i_0_468, plain, (v1_rat_1(esk9_1(X1))|~v4_membered(X1))).#trainneg
cnf(i_0_484, plain, (m1_subset_1(esk5_2(X1,X2),X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_2233, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~v7_ordinal1(X3)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_2268, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~v7_ordinal1(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_2304, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X2),X3)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2306, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X2),X3)|~m1_subset_1(X3,k5_numbers)|~v7_ordinal1(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_2278, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3))).#trainneg
cnf(i_0_2295, plain, (r1_xxreal_0(k2_xcmplx_0(X2,X1),X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(k2_xcmplx_0(X3,k2_xcmplx_0(X2,X1)),k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X3,X2),X3))).#trainneg
cnf(i_0_2323, plain, (v3_ordinal1(X1)|~v7_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2032, plain, (v1_relat_1(esk8_1(X1)))).#trainneg
cnf(i_0_2249, plain, (~v1_xboole_0(X1)|~v2_xxreal_2(k1_zfmisc_1(X1))|~v2_membered(k1_zfmisc_1(X1))|~v2_xxreal_2(k1_xxreal_2(k1_zfmisc_1(X1)))|~v2_membered(k1_xxreal_2(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_488, plain, (v2_membered(esk31_1(X1))|v1_xboole_0(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_580, plain, (v1_funct_1(k5_numbers)|~v1_funct_1(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2092, plain, (v3_trees_3(esk8_1(X1))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_489, plain, (v2_membered(esk27_1(X1))|v1_xboole_0(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_490, plain, (v2_membered(esk26_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_2271, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~v1_xreal_0(X2)|~m1_subset_1(k1_xxreal_2(a_1_1_pre_circ(esk1_0)),k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v6_membered(a_1_1_pre_circ(esk1_0))|~r1_xxreal_0(k2_xcmplx_0(X1,k1_xxreal_2(a_1_1_pre_circ(esk1_0))),esk1_0)|~v1_xxreal_0(esk1_0)|~v1_xxreal_0(k2_xcmplx_0(X1,k1_xxreal_2(a_1_1_pre_circ(esk1_0))))|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0)))).#trainneg
cnf(i_0_1970, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers))).#trainneg
cnf(i_0_1377, plain, (r1_xxreal_0(np__0,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(np__0,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X2,np__0))).#trainneg
cnf(i_0_1977, plain, (r1_xxreal_0(k2_xcmplx_0(X2,X1),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~v7_ordinal1(k2_xcmplx_0(X3,k2_xcmplx_0(X2,X1)))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1881, plain, (r1_xxreal_0(X1,X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4)|~v1_xreal_0(k2_xcmplx_0(X4,X3))|~v1_xreal_0(X5)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X5,X1),k2_xcmplx_0(X4,k2_xcmplx_0(X2,X3)))|~r1_xxreal_0(k2_xcmplx_0(X4,X3),X5)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2412, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(X3,k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X2,X3),X3)|~v7_ordinal1(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_2441, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2445, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~v7_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2049, plain, (v1_trees_3(esk8_1(X1)))).#trainneg
cnf(i_0_2330, plain, (~v1_xboole_0(X1)|~v2_xxreal_2(esk26_1(X1))|~v2_membered(esk26_1(X1)))).#trainneg
cnf(i_0_491, plain, (v2_membered(esk9_1(k1_zfmisc_1(X1)))|~v2_membered(X1))).#trainneg
cnf(i_0_1483, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_2156, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1925, plain, (v1_funct_1(esk8_1(X1))|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_496, plain, (v1_xboole_0(esk26_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_497, plain, (v1_xboole_0(esk9_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2402, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X2,X3),X3)|~v7_ordinal1(X3)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_2443, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~m1_subset_1(X3,k5_numbers)|~v7_ordinal1(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_2442, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers))).#trainneg
cnf(i_0_702, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2091, plain, (v3_trees_3(esk8_1(X1)))).#trainneg
cnf(i_0_570, plain, (v3_trees_2(np__1)|v1_xboole_0(k1_numbers)|~v3_trees_3(k1_numbers))).#trainneg
cnf(i_0_500, plain, (v6_membered(esk31_1(X1))|v1_xboole_0(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_1225, plain, (r1_xxreal_0(esk6_0,np__0)|r1_xxreal_0(np__0,esk6_0))).#trainneg
cnf(i_0_2159, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_2006, plain, (X1=esk23_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_501, plain, (v6_membered(esk27_1(X1))|v1_xboole_0(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_502, plain, (v6_membered(esk26_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_2491, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X3,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X3,X2),X3)|~v7_ordinal1(k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3))|~v1_xreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_2535, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X1,X3),X3)|~v7_ordinal1(X3)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_2608, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X3,k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X1,X3),X3)|~v7_ordinal1(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_572, plain, (v3_trees_2(np__0)|v1_xboole_0(k1_numbers)|~v3_trees_3(k1_numbers))).#trainneg
cnf(i_0_1926, plain, (~v1_xboole_0(X1)|~r2_hidden(X3,esk8_1(X2)))).#trainneg
cnf(i_0_503, plain, (v6_membered(esk9_1(k1_zfmisc_1(X1)))|~v6_membered(X1))).#trainneg
cnf(i_0_1481, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2571, plain, (v1_funct_1(esk8_1(X1))|~v1_relat_1(X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2572, plain, (X1=esk26_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_506, plain, (v3_membered(esk31_1(X1))|v1_xboole_0(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_507, plain, (v3_membered(esk27_1(X1))|v1_xboole_0(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_2618, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~m1_subset_1(X3,k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X1,X3),X3))).#trainneg
cnf(i_0_2659, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X3,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X3,X2),X3)|~m1_subset_1(k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3),k5_numbers))).#trainneg
cnf(i_0_703, plain, (v1_relat_1(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_508, plain, (v3_membered(esk26_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_2154, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2157, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_509, plain, (v3_membered(esk9_1(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_512, plain, (v1_membered(esk31_1(X1))|v1_xboole_0(X1)|~v1_membered(X1))).#trainneg
cnf(i_0_704, plain, (v1_funct_1(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2556, plain, (v1_funct_1(esk8_1(X1)))).#trainneg
cnf(i_0_297, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk9_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_513, plain, (v1_membered(esk27_1(X1))|v1_xboole_0(X1)|~v1_membered(X1))).#trainneg
cnf(i_0_1169, plain, (v1_finset_1(k1_xxreal_2(esk17_0))|~v2_xxreal_2(esk17_0)|~v2_membered(esk17_0))).#trainneg
cnf(i_0_1933, negated_conjecture, (v2_xxreal_0(np__0)|r1_xxreal_0(np__0,esk1_0)|~v2_xxreal_0(esk1_0)|~v1_xreal_0(esk1_0))).#trainneg
cnf(i_0_514, plain, (v1_membered(esk26_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_515, plain, (v1_membered(esk9_1(k1_zfmisc_1(X1)))|~v1_membered(X1))).#trainneg
cnf(i_0_1961, plain, (r1_xxreal_0(k2_xcmplx_0(X2,X1),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(k2_xcmplx_0(X3,k2_xcmplx_0(X2,X1)),k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2107, plain, (r1_xxreal_0(k2_xcmplx_0(X2,X1),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(k2_xcmplx_0(X3,k2_xcmplx_0(X2,X1)),k1_numbers)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_705, plain, (v1_finseq_1(X1)|~v3_finseq_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_520, plain, (v1_relat_1(esk26_1(X1))|~v1_relat_1(X1))).#trainneg
cnf(i_0_518, plain, (v1_relat_1(esk31_1(X1))|v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1516, plain, (r1_xxreal_0(X1,esk8_1(X2))|r1_xxreal_0(esk8_1(X2),X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2583, plain, (X1=esk9_1(k1_zfmisc_1(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_519, plain, (v1_relat_1(esk27_1(X1))|v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_521, plain, (v1_relat_1(esk9_1(k1_zfmisc_1(X1)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_1976, plain, (r1_xxreal_0(k2_xcmplx_0(X2,X1),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3)|~v7_ordinal1(k2_xcmplx_0(X3,k2_xcmplx_0(X2,X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_706, plain, (v1_int_1(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_526, plain, (v1_trees_3(esk26_1(X1))|~v1_trees_3(X1))).#trainneg
cnf(i_0_524, plain, (v1_trees_3(esk31_1(X1))|v1_xboole_0(X1)|~v1_trees_3(X1))).#trainneg
cnf(i_0_2153, plain, (v2_xxreal_0(np__0)|r1_xxreal_0(np__0,esk12_0)|~v2_xxreal_0(esk12_0)|~v1_xreal_0(esk12_0))).#trainneg
cnf(i_0_2631, plain, (v2_xxreal_0(np__0)|r1_xxreal_0(np__0,esk6_0)|~v2_xxreal_0(esk6_0)|~v1_xreal_0(esk6_0))).#trainneg
cnf(i_0_525, plain, (v1_trees_3(esk27_1(X1))|v1_xboole_0(X1)|~v1_trees_3(X1))).#trainneg
cnf(i_0_527, plain, (v1_trees_3(esk9_1(k1_zfmisc_1(X1)))|~v1_trees_3(X1))).#trainneg
cnf(i_0_2778, plain, (~v1_xreal_0(X1)|~m1_subset_1(k2_xcmplx_0(np__1,X1),k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(np__1,X1),X1))).#trainneg
cnf(i_0_2729, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X3,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X3,X2),X3)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2803, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X2),X3)|~v7_ordinal1(k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3))|~v1_xreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_707, plain, (v1_rat_1(X1)|~v4_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_532, plain, (v3_trees_3(esk26_1(X1))|~v3_trees_3(X1))).#trainneg
cnf(i_0_1233, negated_conjecture, (r1_xxreal_0(esk9_1(k5_numbers),esk1_0)|r1_xxreal_0(esk1_0,esk9_1(k5_numbers)))).#trainneg
cnf(i_0_2685, plain, (~v1_xboole_0(X1)|~v2_xxreal_2(esk8_1(X2)))).#trainneg
cnf(i_0_530, plain, (v3_trees_3(esk31_1(X1))|v1_xboole_0(X1)|~v3_trees_3(X1))).#trainneg
cnf(i_0_1774, negated_conjecture, (v2_xxreal_0(esk1_0)|r1_xxreal_0(esk1_0,esk12_0)|~v2_xxreal_0(esk12_0)|~v1_xreal_0(esk1_0)|~v1_xreal_0(esk12_0))).#trainneg
cnf(i_0_1813, negated_conjecture, (v2_xxreal_0(esk1_0)|r1_xxreal_0(esk1_0,esk6_0)|~v2_xxreal_0(esk6_0)|~v1_xreal_0(esk1_0)|~v1_xreal_0(esk6_0))).#trainneg
cnf(i_0_2017, plain, (v2_xxreal_0(esk6_0)|r1_xxreal_0(esk6_0,esk12_0)|~v2_xxreal_0(esk12_0)|~v1_xreal_0(esk6_0)|~v1_xreal_0(esk12_0))).#trainneg
cnf(i_0_932, plain, (r1_xxreal_0(X1,esk9_1(X2))|r1_xxreal_0(esk9_1(X2),X1)|~v7_ordinal1(X1)|~v6_membered(X2))).#trainneg
cnf(i_0_1262, plain, (r1_xxreal_0(X1,esk9_1(X2))|r1_xxreal_0(esk9_1(X2),X1)|~v7_ordinal1(X1)|~v2_membered(X2))).#trainneg
cnf(i_0_531, plain, (v3_trees_3(esk27_1(X1))|v1_xboole_0(X1)|~v3_trees_3(X1))).#trainneg
cnf(i_0_533, plain, (v3_trees_3(esk9_1(k1_zfmisc_1(X1)))|~v3_trees_3(X1))).#trainneg
cnf(i_0_2772, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X2),X3)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_2763, plain, (r1_xxreal_0(k2_xcmplx_0(X2,X1),X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(k2_xcmplx_0(X3,k2_xcmplx_0(X2,X1)),k1_numbers)|~r1_xxreal_0(k2_xcmplx_0(X3,X2),X3))).#trainneg
cnf(i_0_2798, plain, (r1_xxreal_0(k2_xcmplx_0(X2,X1),X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X2),X3)|~v7_ordinal1(k2_xcmplx_0(X3,k2_xcmplx_0(X2,X1))))).#trainneg
cnf(i_0_2727, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X3,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X3,X2),X3)|~m1_subset_1(X3,k5_numbers)|~v7_ordinal1(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_2147, plain, (v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1475, plain, (v1_xxreal_0(esk9_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_2809, negated_conjecture, (~v6_membered(a_1_1_pre_circ(esk1_0))|~v1_xreal_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~m1_subset_1(k2_xcmplx_0(np__1,k1_xxreal_2(a_1_1_pre_circ(esk1_0))),k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(np__1,k1_xxreal_2(a_1_1_pre_circ(esk1_0))),esk1_0)|~v1_xxreal_0(k2_xcmplx_0(np__1,k1_xxreal_2(a_1_1_pre_circ(esk1_0))))|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0)))).#trainneg
cnf(i_0_536, plain, (v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2552, plain, (r1_xxreal_0(np__1,k5_numbers)|r1_xxreal_0(k5_numbers,np__1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_2553, plain, (r1_xxreal_0(esk22_0,k5_numbers)|r1_xxreal_0(k5_numbers,esk22_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_2554, plain, (r1_xxreal_0(np__0,k5_numbers)|r1_xxreal_0(k5_numbers,np__0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_2813, plain, (v1_int_1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2846, plain, (v1_rat_1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_543, plain, (v1_xboole_0(X1)|~v1_trees_3(X1)|~v1_xboole_0(esk9_1(X1)))).#trainneg
cnf(i_0_556, plain, (v1_trees_1(esk9_1(X1))|v1_xboole_0(X1)|~v1_trees_3(X1))).#trainneg
cnf(i_0_2826, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X2,X3),X3)|~v7_ordinal1(X3)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_2901, plain, (~v1_xreal_0(X1)|~m1_subset_1(k2_xcmplx_0(X1,np__1),k1_numbers)|~r1_xxreal_0(k2_xcmplx_0(X1,np__1),X1))).#trainneg
cnf(i_0_2917, plain, (~v1_xreal_0(X1)|~r1_xxreal_0(k2_xcmplx_0(X1,np__1),X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__1)))).#trainneg
cnf(i_0_2825, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X2,X3),X3)|~v7_ordinal1(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_2624, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2642, plain, (v3_membered(esk26_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_1244, negated_conjecture, (r1_xxreal_0(k1_xxreal_2(k5_numbers),esk1_0)|r1_xxreal_0(esk1_0,k1_xxreal_2(k5_numbers))|~v2_xxreal_2(k5_numbers)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_567, plain, (v3_trees_2(esk9_1(X1))|v1_xboole_0(X1)|~v3_trees_3(X1))).#trainneg
cnf(i_0_603, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_2848, negated_conjecture, (v2_xxreal_0(esk1_0)|r1_xxreal_0(esk1_0,esk9_1(k5_numbers))|~v2_xxreal_0(esk9_1(k5_numbers))|~v1_xreal_0(esk1_0)|~v1_xreal_0(esk9_1(k5_numbers)))).#trainneg
cnf(i_0_1308, negated_conjecture, (v2_xxreal_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~v2_xxreal_0(esk1_0)|~v1_xreal_0(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~v1_xreal_0(esk1_0)|~v6_membered(a_1_1_pre_circ(esk1_0))|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0)))).#trainneg
cnf(i_0_2012, plain, (r1_xxreal_0(esk23_1(X1),X2)|r1_xxreal_0(X2,esk23_1(X1))|~v7_ordinal1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_575, plain, (v1_funct_1(esk31_1(X1))|v1_xboole_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_576, plain, (v1_funct_1(esk27_1(X1))|v1_xboole_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2872, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X1)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~m1_subset_1(X3,k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X2,X3),X3))).#trainneg
cnf(i_0_2444, plain, (v3_ordinal1(X1)|~v7_ordinal1(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_908, plain, (v1_finset_1(esk26_1(X1))|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1772, negated_conjecture, (r1_xxreal_0(esk1_0,k1_xxreal_2(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))))|~v6_membered(a_1_1_pre_circ(esk1_0))|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))))|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))))|~v2_membered(a_1_1_pre_circ(esk1_0))|~v6_membered(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))))).#trainneg
cnf(i_0_577, plain, (v1_funct_1(esk26_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_652, plain, (r2_hidden(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_1752, negated_conjecture, (v1_xboole_0(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0))))|~v7_ordinal1(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))|~v6_membered(a_1_1_pre_circ(esk1_0))|~r1_xxreal_0(esk9_1(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))),esk1_0)|~v1_xxreal_0(esk1_0)|~v1_xxreal_0(esk9_1(a_1_1_pre_circ(k1_xxreal_2(a_1_1_pre_circ(esk1_0)))))|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0)))).#trainneg
cnf(i_0_2262, plain, (r1_xxreal_0(X1,k1_xxreal_2(X2))|r1_xxreal_0(k1_xxreal_2(X2),X1)|~v7_ordinal1(X1)|~v3_membered(X2))).#trainneg
cnf(i_0_2579, plain, (r1_xxreal_0(esk26_1(X1),X2)|r1_xxreal_0(X2,esk26_1(X1))|~v7_ordinal1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_578, plain, (v1_funct_1(esk9_1(k1_zfmisc_1(X1)))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_585, plain, (r1_xxreal_0(X1,np__1)|~r1_xxreal_0(X1,np__0)|~v1_xxreal_0(np__1)|~v1_xxreal_0(np__0)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_3035, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2959, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X2,X2))).#trainneg
cnf(i_0_3068, plain, (r1_xxreal_0(k1_xxreal_2(k5_numbers),X1)|r1_xxreal_0(X1,k1_xxreal_2(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3007, plain, (~v1_xreal_0(X1)|~r1_xxreal_0(k2_xcmplx_0(np__1,X1),X1)|~v7_ordinal1(k2_xcmplx_0(np__1,X1)))).#trainneg
cnf(i_0_3033, plain, (m1_subset_1(esk2_2(k5_numbers,X1),k1_numbers)|r1_tarski(k5_numbers,X1))).#trainneg
cnf(i_0_3053, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~v7_ordinal1(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_587, plain, (k2_xcmplx_0(X1,esk22_0)=k2_xcmplx_0(esk22_0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2970, plain, (v2_xxreal_0(k5_numbers)|r1_xxreal_0(k5_numbers,np__1)|~v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_3136, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2858, plain, (r1_xxreal_0(esk9_1(X1),X2)|r1_xxreal_0(X2,esk9_1(X1))|~v7_ordinal1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3123, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3189, plain, (r1_tarski(k5_numbers,k1_numbers))).#trainneg
cnf(i_0_3122, plain, (v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3134, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3133, plain, (r1_xxreal_0(X1,k1_xxreal_2(X2))|r1_xxreal_0(k1_xxreal_2(X2),X1)|~v2_membered(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3056, plain, (r1_tarski(X1,k1_numbers)|~r2_hidden(esk2_2(X1,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_2957, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3))).#trainneg
cnf(i_0_3018, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X3,X2),X3)|~v7_ordinal1(X3)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_3191, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k5_numbers))).#trainneg
cnf(i_0_3204, plain, (k2_xcmplx_0(esk22_0,np__1)=k2_xcmplx_0(np__1,esk22_0))).#trainneg
cnf(i_0_598, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk31_1(X2)))).#trainneg
cnf(i_0_2972, plain, (v2_xxreal_0(k5_numbers)|r1_xxreal_0(k5_numbers,esk22_0)|~v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_3348, plain, (~r1_xxreal_0(k2_xcmplx_0(np__1,esk22_0),esk22_0)|~v7_ordinal1(k2_xcmplx_0(np__1,esk22_0)))).#trainneg
cnf(i_0_3346, plain, (~m1_subset_1(k2_xcmplx_0(np__1,esk22_0),k1_numbers)|~r1_xxreal_0(k2_xcmplx_0(np__1,esk22_0),esk22_0))).#trainneg
cnf(i_0_2867, plain, (r1_xxreal_0(esk9_1(X1),X2)|r1_xxreal_0(X2,esk9_1(X1))|~v7_ordinal1(X2)|~v3_membered(X1))).#trainneg
cnf(i_0_599, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk27_1(X2)))).#trainneg
cnf(i_0_600, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk26_1(X2)))).#trainneg
cnf(i_0_3225, plain, (v1_xxreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3315, plain, (k2_xcmplx_0(k2_xcmplx_0(np__1,esk22_0),X1)=k2_xcmplx_0(esk22_0,k2_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3272, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X3,k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(X2,X3),X3)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_3351, negated_conjecture, (~v6_membered(a_1_1_pre_circ(esk1_0))|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))|~m1_subset_1(k1_xxreal_2(a_1_1_pre_circ(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_3285, plain, (v1_xreal_0(k2_xcmplx_0(np__1,esk22_0)))).#trainneg
cnf(i_0_3205, plain, (k2_xcmplx_0(esk22_0,np__0)=k2_xcmplx_0(np__0,esk22_0))).#trainneg
cnf(i_0_3108, plain, (~v2_xxreal_2(a_1_1_pre_circ(np__1))|~v2_membered(a_1_1_pre_circ(np__1))|~r1_xxreal_0(k1_xxreal_2(a_1_1_pre_circ(np__1)),np__0)|~v1_xxreal_0(k1_xxreal_2(a_1_1_pre_circ(np__1))))).#trainneg
cnf(i_0_601, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk9_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_2974, plain, (v2_xxreal_0(k5_numbers)|r1_xxreal_0(k5_numbers,np__0)|~v2_xxreal_0(np__0)|~v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_3347, plain, (~m1_subset_1(k2_xcmplx_0(np__1,esk22_0),k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(np__1,esk22_0),esk22_0))).#trainneg
cnf(i_0_3084, plain, (~v7_ordinal1(k2_xcmplx_0(np__1,np__1))|~r1_xxreal_0(k2_xcmplx_0(np__1,np__1),np__0)|~v1_xxreal_0(k2_xcmplx_0(np__1,np__1)))).#trainneg
cnf(i_0_3070, plain, (r1_xxreal_0(k1_xxreal_2(X1),X2)|r1_xxreal_0(X2,k1_xxreal_2(X1))|~v7_ordinal1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_604, plain, (m1_subset_1(esk15_2(esk5_2(X1,a_1_1_pre_circ(X2)),X2),k5_numbers)|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_1_pre_circ(X2)))).#trainneg
cnf(i_0_3528, plain, (k2_xcmplx_0(k2_xcmplx_0(np__0,esk22_0),X1)=k2_xcmplx_0(esk22_0,k2_xcmplx_0(np__0,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3602, plain, (v7_ordinal1(esk15_2(esk2_2(a_1_1_pre_circ(X1),X2),X1))|r1_tarski(a_1_1_pre_circ(X1),X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3627, plain, (v7_ordinal1(esk2_2(a_1_1_pre_circ(X1),X2))|r1_tarski(a_1_1_pre_circ(X1),X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3630, plain, (v1_xreal_0(esk2_2(a_1_1_pre_circ(X1),X2))|r1_tarski(a_1_1_pre_circ(X1),X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3752, negated_conjecture, (v1_xboole_0(a_1_1_pre_circ(esk1_0))|~v1_xxreal_0(esk1_0)|~v2_membered(a_1_1_pre_circ(esk1_0))|~m1_subset_1(a_1_1_pre_circ(esk1_0),k5_numbers))).#trainneg
cnf(i_0_3629, plain, (v1_xxreal_0(esk2_2(a_1_1_pre_circ(X1),X2))|r1_tarski(a_1_1_pre_circ(X1),X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3649, plain, (r1_tarski(a_1_1_pre_circ(X1),a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(X1))))|r1_tarski(X2,a_1_1_pre_circ(X1))|~m1_subset_1(esk2_2(X2,a_1_1_pre_circ(X1)),k5_numbers)|~v1_xxreal_0(esk2_2(X2,a_1_1_pre_circ(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3739, plain, (v1_finset_1(esk15_2(esk2_2(a_1_1_pre_circ(X1),X2),X1))|r1_tarski(a_1_1_pre_circ(X1),X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3316, plain, (v1_xxreal_0(k2_xcmplx_0(np__1,esk22_0)))).#trainneg
cnf(i_0_3559, plain, (~v2_xxreal_2(a_1_1_pre_circ(np__1))|~v2_membered(a_1_1_pre_circ(np__1))|k1_xxreal_2(a_1_1_pre_circ(np__1))!=np__0)).#trainneg
cnf(i_0_606, plain, (m1_subset_1(esk15_2(X1,X2),k5_numbers)|v1_xboole_0(a_1_1_pre_circ(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,a_1_1_pre_circ(X2)))).#trainneg
cnf(i_0_3565, plain, (v2_xxreal_0(k5_numbers)|r1_xxreal_0(k5_numbers,np__0)|r1_xxreal_0(np__0,esk22_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_1261, plain, (r1_xxreal_0(X1,esk9_1(X2))|r1_xxreal_0(esk9_1(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(X2))).#trainneg
cnf(i_0_2945, plain, (r1_xxreal_0(X1,esk9_1(X2))|r1_xxreal_0(esk9_1(X2),X1)|~v1_xxreal_0(X1)|~v3_membered(X2))).#trainneg
cnf(i_0_3621, negated_conjecture, (r1_xxreal_0(esk15_2(esk2_2(a_1_1_pre_circ(X1),X2),X1),esk1_0)|r1_xxreal_0(esk1_0,esk15_2(esk2_2(a_1_1_pre_circ(X1),X2),X1))|r1_tarski(a_1_1_pre_circ(X1),X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3620, plain, (r1_xxreal_0(esk15_2(esk2_2(a_1_1_pre_circ(X1),X2),X1),X3)|r1_xxreal_0(X3,esk15_2(esk2_2(a_1_1_pre_circ(X1),X2),X1))|r1_tarski(a_1_1_pre_circ(X1),X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3261, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X3))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X1,X3),X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3856, plain, (r1_xxreal_0(esk15_2(esk2_2(a_1_1_pre_circ(X1),X2),X1),esk15_2(esk2_2(a_1_1_pre_circ(X3),X4),X3))|r1_xxreal_0(esk15_2(esk2_2(a_1_1_pre_circ(X3),X4),X3),esk15_2(esk2_2(a_1_1_pre_circ(X1),X2),X1))|r1_tarski(a_1_1_pre_circ(X3),X4)|r1_tarski(a_1_1_pre_circ(X1),X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3777, plain, (r1_tarski(a_1_1_pre_circ(X1),X2)|v1_finset_1(esk2_2(a_1_1_pre_circ(X1),X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1441, plain, (v7_ordinal1(esk8_1(X2))|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3498, plain, (v1_xreal_0(k2_xcmplx_0(np__0,esk22_0)))).#trainneg
cnf(i_0_3529, plain, (v1_xxreal_0(k2_xcmplx_0(np__0,esk22_0)))).#trainneg
cnf(i_0_3786, plain, (v1_xboole_0(a_1_1_pre_circ(np__1))|k1_xxreal_2(a_1_1_pre_circ(np__1))!=np__0|~v2_membered(a_1_1_pre_circ(np__1))|~v1_finset_1(a_1_1_pre_circ(np__1)))).#trainneg
cnf(i_0_611, plain, (v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3312, plain, (r1_xxreal_0(np__0,esk22_0)|~r1_xxreal_0(np__1,k2_xcmplx_0(np__1,esk22_0)))).#trainneg
cnf(i_0_3813, plain, (v1_xboole_0(a_1_1_pre_circ(X1))|m1_subset_1(X2,k5_numbers)|~m1_subset_1(X2,a_1_1_pre_circ(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_612, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X4)=k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_613, plain, (k2_xcmplx_0(np__1,k2_xcmplx_0(np__0,X1))=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_3909, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,a_1_1_pre_circ(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3902, plain, (v3_membered(a_1_1_pre_circ(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3925, plain, (r1_xxreal_0(esk15_2(esk2_2(a_1_1_pre_circ(X1),X2),X1),esk15_2(esk2_2(a_1_1_pre_circ(X3),X4),X3))|r1_xxreal_0(esk15_2(esk2_2(a_1_1_pre_circ(X3),X4),X3),esk15_2(esk2_2(a_1_1_pre_circ(X1),X2),X1))|r1_tarski(a_1_1_pre_circ(X3),X4)|r1_tarski(a_1_1_pre_circ(X1),X2)|~v7_ordinal1(X3)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3876, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X2)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X3))).#trainneg
cnf(i_0_3034, plain, (m1_subset_1(k1_xxreal_2(k5_numbers),k1_numbers)|~v2_xxreal_2(k5_numbers))).#trainneg
cnf(i_0_614, plain, (k2_xcmplx_0(np__0,k2_xcmplx_0(np__1,X1))=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_3525, plain, (r1_xxreal_0(np__0,esk22_0)|~r1_xxreal_0(np__1,k2_xcmplx_0(np__0,esk22_0)))).#trainneg
cnf(i_0_3081, plain, (~m1_subset_1(k2_xcmplx_0(np__1,np__1),k5_numbers)|~r1_xxreal_0(k2_xcmplx_0(np__1,np__1),np__0)|~v1_xxreal_0(k2_xcmplx_0(np__1,np__1)))).#trainneg
cnf(i_0_3082, plain, (~m1_subset_1(k2_xcmplx_0(np__1,np__1),k1_numbers)|~r1_xxreal_0(k2_xcmplx_0(np__1,np__1),np__0)|~v1_xxreal_0(k2_xcmplx_0(np__1,np__1)))).#trainneg
cnf(i_0_3846, plain, (r1_xxreal_0(esk9_1(X1),X2)|r1_xxreal_0(X2,esk9_1(X1))|~v1_xxreal_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_615, plain, (k2_xcmplx_0(np__0,k2_xcmplx_0(np__0,X1))=k2_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_625, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_3774, plain, (m1_subset_1(a_1_1_pre_circ(X1),k1_zfmisc_1(a_1_1_pre_circ(esk2_2(X2,a_1_1_pre_circ(X1)))))|r1_tarski(X2,a_1_1_pre_circ(X1))|~m1_subset_1(esk2_2(X2,a_1_1_pre_circ(X1)),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3850, negated_conjecture, (r1_xxreal_0(esk1_0,esk15_2(esk2_2(a_1_1_pre_circ(esk1_0),X1),esk1_0))|r1_tarski(a_1_1_pre_circ(esk1_0),X1)|~r2_hidden(esk2_2(a_1_1_pre_circ(esk1_0),X1),a_1_1_pre_circ(esk1_0)))).#trainneg
cnf(i_0_3616, plain, (r1_xxreal_0(X1,esk15_2(esk2_2(a_1_1_pre_circ(X2),X3),X2))|r1_xxreal_0(esk15_2(esk2_2(a_1_1_pre_circ(X2),X3),X2),X1)|r1_tarski(a_1_1_pre_circ(X2),X3)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3910, plain, (r2_hidden(X1,k5_numbers)|~r2_hidden(X1,a_1_1_pre_circ(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4220, plain, (v1_xreal_0(k1_xxreal_2(k5_numbers))|~v2_xxreal_2(k5_numbers))).#trainneg
cnf(i_0_4218, plain, (v1_finset_1(k1_xxreal_2(k5_numbers))|~v5_finset_1(k1_numbers)|~v2_xxreal_2(k5_numbers))).#trainneg
cnf(i_0_626, plain, (r1_xxreal_0(esk5_2(X1,X2),X3)|X3!=k1_xxreal_2(X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(esk5_2(X1,X2))|~v2_xxreal_2(X2)|~v2_membered(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_3526, plain, (r1_xxreal_0(np__1,esk22_0)|~r1_xxreal_0(np__1,k2_xcmplx_0(np__0,esk22_0)))).#trainneg
cnf(i_0_630, plain, (r1_xxreal_0(X1,esk3_2(X2,X3))|r1_xxreal_0(esk3_2(X2,X3),X1)|X3=k1_xxreal_2(X2)|~v1_xxreal_0(X1)|~v1_xxreal_0(X3)|~v2_xxreal_2(X2)|~v2_membered(X2)|~r2_hidden(X3,X2))).#trainneg
cnf(i_0_632, plain, (esk15_2(esk3_2(a_1_1_pre_circ(X1),X2),X1)=esk3_2(a_1_1_pre_circ(X1),X2)|X2=k1_xxreal_2(a_1_1_pre_circ(X1))|~v7_ordinal1(X1)|~v1_xxreal_0(X2)|~v2_xxreal_2(a_1_1_pre_circ(X1))|~v2_membered(a_1_1_pre_circ(X1))|~r2_hidden(X2,a_1_1_pre_circ(X1)))).#trainneg
cnf(i_0_3617, plain, (r1_xxreal_0(X1,k1_xxreal_2(esk15_2(esk2_2(a_1_1_pre_circ(X2),X3),X2)))|r1_xxreal_0(k1_xxreal_2(esk15_2(esk2_2(a_1_1_pre_circ(X2),X3),X2)),X1)|r1_tarski(a_1_1_pre_circ(X2),X3)|~v1_xxreal_0(X1)|~v2_xxreal_2(esk15_2(esk2_2(a_1_1_pre_circ(X2),X3),X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3740, plain, (r1_xxreal_0(esk15_2(esk2_2(a_1_1_pre_circ(X1),X2),X1),k1_xxreal_2(k5_numbers))|r1_xxreal_0(k1_xxreal_2(k5_numbers),esk15_2(esk2_2(a_1_1_pre_circ(X1),X2),X1))|r1_tarski(a_1_1_pre_circ(X1),X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3903, plain, (v1_membered(a_1_1_pre_circ(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3789, plain, (v7_ordinal1(esk15_2(X1,X2))|v1_xboole_0(a_1_1_pre_circ(X2))|~m1_subset_1(X1,a_1_1_pre_circ(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4541, plain, (v1_xboole_0(a_1_1_pre_circ(X1))|v7_ordinal1(X2)|~m1_subset_1(X2,a_1_1_pre_circ(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4471, negated_conjecture, (r1_xxreal_0(esk1_0,esk2_2(a_1_1_pre_circ(esk1_0),X1))|r1_tarski(a_1_1_pre_circ(esk1_0),X1))).#trainneg
cnf(i_0_868, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_4502, plain, (v1_xxreal_0(k1_xxreal_2(k5_numbers))|~v2_xxreal_2(k5_numbers))).#trainneg
cnf(i_0_635, plain, (m1_subset_1(esk3_2(X1,X2),X1)|X2=k1_xxreal_2(X1)|~v1_xxreal_0(X2)|~v2_xxreal_2(X1)|~v2_membered(X1)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_3527, plain, (r1_xxreal_0(np__0,esk22_0)|~r1_xxreal_0(np__0,k2_xcmplx_0(np__0,esk22_0)))).#trainneg
cnf(i_0_1040, plain, (v1_xboole_0(X1)|~v1_trees_3(X1)|~v1_xboole_0(k1_xxreal_2(X1))|~v2_xxreal_2(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_2953, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|r1_xxreal_0(k2_xcmplx_0(X2,X3),X1)|~v7_ordinal1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_633, plain, (m1_subset_1(esk15_2(esk3_2(a_1_1_pre_circ(X1),X2),X1),k5_numbers)|X2=k1_xxreal_2(a_1_1_pre_circ(X1))|~v7_ordinal1(X1)|~v1_xxreal_0(X2)|~v2_xxreal_2(a_1_1_pre_circ(X1))|~v2_membered(a_1_1_pre_circ(X1))|~r2_hidden(X2,a_1_1_pre_circ(X1)))).#trainneg
cnf(i_0_637, plain, (r1_xxreal_0(esk3_2(X1,X2),X3)|X2=k1_xxreal_2(X1)|X3!=k1_xxreal_2(X1)|~v1_xxreal_0(X3)|~v1_xxreal_0(esk3_2(X1,X2))|~v2_xxreal_2(X1)|~v2_membered(X1)|~v1_xxreal_0(X2)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_4515, plain, (X2=k1_xxreal_2(a_1_1_pre_circ(X1))|~r1_xxreal_0(esk3_2(a_1_1_pre_circ(X1),X2),X1)|~r2_hidden(esk3_2(a_1_1_pre_circ(X1),X2),a_1_1_pre_circ(X1))|~v7_ordinal1(X1)|~v1_xxreal_0(X2)|~v2_xxreal_2(a_1_1_pre_circ(X1))|~r2_hidden(X2,a_1_1_pre_circ(X1)))).#trainneg
cnf(i_0_4501, plain, (r2_hidden(X1,k5_numbers)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4737, negated_conjecture, (k1_xxreal_2(X1)=esk1_0|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(X1)|~v2_membered(X1)|~r2_hidden(esk1_0,X1)|~m1_subset_1(esk3_2(X1,esk1_0),k5_numbers))).#trainneg
cnf(i_0_4765, negated_conjecture, (~v2_xxreal_2(a_1_1_pre_circ(esk1_0))|~v2_membered(a_1_1_pre_circ(esk1_0))|~m1_subset_1(k1_xxreal_2(a_1_1_pre_circ(esk1_0)),k5_numbers)|~v1_xxreal_0(esk1_0))).#trainneg
cnf(i_0_4738, negated_conjecture, (r1_tarski(a_1_1_pre_circ(esk1_0),X1)|~m1_subset_1(esk2_2(a_1_1_pre_circ(esk1_0),X1),k5_numbers)|~v1_xxreal_0(esk1_0))).#trainneg
cnf(i_0_4735, negated_conjecture, (~r2_hidden(X1,a_1_1_pre_circ(esk1_0))|~m1_subset_1(esk15_2(X1,esk1_0),k5_numbers)|~v1_xxreal_0(esk1_0))).#trainneg
cnf(i_0_4772, plain, (X1=k1_xxreal_2(a_1_1_pre_circ(X2))|~r1_xxreal_0(esk3_2(a_1_1_pre_circ(X2),X1),X2)|~v1_xxreal_0(X1)|~v2_xxreal_2(a_1_1_pre_circ(X2))|~r2_hidden(X1,a_1_1_pre_circ(X2))|~v7_ordinal1(X2)|~v2_membered(a_1_1_pre_circ(X2)))).#trainneg
cnf(i_0_4740, negated_conjecture, (~v1_xreal_0(esk1_0)|~m1_subset_1(k2_xcmplx_0(np__1,esk1_0),k5_numbers)|~v1_xxreal_0(esk1_0))).#trainneg
cnf(i_0_2241, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_2011, plain, (esk23_1(X1)=esk8_1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4586, negated_conjecture, (v2_membered(np__0)|~v1_xxreal_0(esk1_0))).#trainneg
cnf(i_0_640, plain, (r1_xxreal_0(X1,np__0)|~v1_xreal_0(X2)|~v1_xreal_0(np__0)|~v1_xreal_0(X1)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__0)|~r1_xxreal_0(np__0,X2))).#trainneg
cnf(i_0_4677, plain, (~v1_xboole_0(k1_xxreal_2(esk25_0))|~v2_xxreal_2(esk25_0)|~v2_membered(esk25_0))).#trainneg
cnf(i_0_4624, negated_conjecture, (~v6_membered(np__0)|~v1_xxreal_0(esk1_0)|~v2_xxreal_2(np__0)|~v2_membered(np__0))).#trainneg
cnf(i_0_2187, plain, (r1_xxreal_0(k1_xxreal_2(X1),X2)|r1_xxreal_0(X2,k1_xxreal_2(X1))|~v1_xxreal_0(X2)|~v2_xxreal_2(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_638, plain, (r1_xxreal_0(X1,np__1)|~v1_xreal_0(X2)|~v1_xreal_0(np__0)|~v1_xreal_0(X1)|~v1_xreal_0(np__1)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__1)|~r1_xxreal_0(np__0,X2))).#trainneg
cnf(i_0_639, plain, (r1_xxreal_0(X1,np__0)|~v1_xreal_0(X2)|~v1_xreal_0(np__1)|~v1_xreal_0(X1)|~v1_xreal_0(np__0)|~r1_xxreal_0(k2_xcmplx_0(X1,X2),np__1)|~r1_xxreal_0(np__1,X2))).#trainneg
cnf(i_0_4828, plain, (r1_xxreal_0(X1,X2)|r2_hidden(X3,k5_numbers)|~r1_xxreal_0(X1,X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~m1_subset_1(X3,k5_numbers)|~v7_ordinal1(X2))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 161
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 275
# Removed in clause preprocessing      : 19
# Initial clauses in saturation        : 256
# Processed clauses                    : 1395
# ...of these trivial                  : 29
# ...subsumed                          : 475
# ...remaining for further processing  : 891
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 66
# Backward-rewritten                   : 36
# Generated clauses                    : 4728
# ...of the previous two non-trivial   : 4420
# Contextual simplify-reflections      : 123
# Paramodulations                      : 4724
# Factorizations                       : 0
# Equation resolutions                 : 4
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 788
#    Positive orientable unit clauses  : 138
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 34
#    Non-unit-clauses                  : 615
# Current number of unprocessed clauses: 3092
# ...number of literals in the above   : 17591
# Current number of archived formulas  : 0
# Current number of archived clauses   : 993
# Proof object given clauses           : 123
# Proof search given clauses           : 891
# Clause-clause subsumption calls (NU) : 158829
# Rec. Clause-clause subsumption calls : 48588
# Non-unit clause-clause subsumptions  : 594
# Unit Clause-clause subsumption calls : 2361
# Rewrite failures with RHS unbound    : 16
# BW rewrite match attempts            : 32
# BW rewrite match successes           : 21
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 103675

# -------------------------------------------------
# User time                : 8.126 s
# System time              : 0.257 s
# Total time               : 8.382 s
# Maximum resident set size: 295032 pages
8.12user 0.33system 0:08.55elapsed 98%CPU (0avgtext+0avgdata 295756maxresident)k
0inputs+328outputs (0major+91461minor)pagefaults 0swaps
