2020-06-08 18:34:17.236349: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_49
2020-06-08 18:34:17.317885: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-06-08 18:34:17.482292: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-06-08 18:34:18.047241: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-06-08 18:34:20.764215: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 3527878 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_49' loaded (query=128; context=512; weigths=real; len_mult=0.000000).
# Preprocessing time       : 4.924 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t21_asympt_0, axiom, ![X1]:((v1_funct_1(X1)&(v1_funct_2(X1,k5_numbers,k1_numbers)&(v2_asympt_0(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))=>![X2]:((v1_funct_1(X2)&(v1_funct_2(X2,k5_numbers,k1_numbers)&(v2_asympt_0(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,k5_numbers,k1_numbers)&(v2_asympt_0(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))=>((r2_hidden(X1,k7_asympt_0(X2))&r2_hidden(X2,k7_asympt_0(X3)))=>r2_hidden(X1,k7_asympt_0(X3)))))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t26_asympt_0', t21_asympt_0)).
fof(t18_asympt_0, axiom, ![X1, X2]:((v1_funct_1(X2)&(v1_funct_2(X2,k5_numbers,k1_numbers)&(v2_asympt_0(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))=>(r2_hidden(X1,k7_asympt_0(X2))=>(v1_funct_1(X1)&(v1_funct_2(X1,k5_numbers,k1_numbers)&(v2_asympt_0(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t26_asympt_0', t18_asympt_0)).
fof(t19_asympt_0, axiom, ![X1]:((v1_funct_1(X1)&(v1_funct_2(X1,k5_numbers,k1_numbers)&(v2_asympt_0(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))=>![X2]:((v1_funct_1(X2)&(v1_funct_2(X2,k5_numbers,k1_numbers)&(v2_asympt_0(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))=>(r2_hidden(X1,k7_asympt_0(X2))<=>r2_hidden(X2,k6_asympt_0(X1))))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t26_asympt_0', t19_asympt_0)).
fof(t6_asympt_0, axiom, ![X1, X2]:((v1_funct_1(X2)&(v1_funct_2(X2,k5_numbers,k1_numbers)&(v2_asympt_0(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))=>(r2_hidden(X1,k6_asympt_0(X2))=>(v1_funct_1(X1)&(v1_funct_2(X1,k5_numbers,k1_numbers)&(v2_asympt_0(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t26_asympt_0', t6_asympt_0)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t26_asympt_0', d3_tarski)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t26_asympt_0', d10_xboole_0)).
fof(t10_asympt_0, axiom, ![X1]:((v1_funct_1(X1)&(v1_funct_2(X1,k5_numbers,k1_numbers)&(v2_asympt_0(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))=>r2_hidden(X1,k6_asympt_0(X1))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t26_asympt_0', t10_asympt_0)).
fof(t9_asympt_0, axiom, ![X1]:((v1_funct_1(X1)&(v1_funct_2(X1,k5_numbers,k1_numbers)&(v2_asympt_0(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))=>![X2]:((v1_funct_1(X2)&(v1_funct_2(X2,k5_numbers,k1_numbers)&(v2_asympt_0(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))=>k6_asympt_0(k1_asympt_0(X1,X2))=k6_asympt_0(k5_asympt_0(X1,X2)))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t26_asympt_0', t9_asympt_0)).
fof(dt_k1_asympt_0, axiom, ![X1, X2]:(((v1_funct_1(X1)&(v1_funct_2(X1,k5_numbers,k1_numbers)&(v2_asympt_0(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))&(v1_funct_1(X2)&(v1_funct_2(X2,k5_numbers,k1_numbers)&(v2_asympt_0(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))=>(v1_funct_1(k1_asympt_0(X1,X2))&(v1_funct_2(k1_asympt_0(X1,X2),k5_numbers,k1_numbers)&(v2_asympt_0(k1_asympt_0(X1,X2))&m1_subset_1(k1_asympt_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t26_asympt_0', dt_k1_asympt_0)).
fof(dt_k5_asympt_0, axiom, ![X1, X2]:(((v1_funct_1(X1)&(v1_funct_2(X1,k5_numbers,k1_numbers)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(v1_funct_1(X2)&(v1_funct_2(X2,k5_numbers,k1_numbers)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))=>(v1_funct_1(k5_asympt_0(X1,X2))&(v1_funct_2(k5_asympt_0(X1,X2),k5_numbers,k1_numbers)&m1_subset_1(k5_asympt_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t26_asympt_0', dt_k5_asympt_0)).
fof(fc1_asympt_0, axiom, ![X1, X2]:(((v1_funct_1(X1)&(v1_funct_2(X1,k5_numbers,k1_numbers)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(v1_funct_1(X2)&(v1_funct_2(X2,k5_numbers,k1_numbers)&(v2_asympt_0(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))=>(v1_funct_1(k5_asympt_0(X1,X2))&(v1_funct_2(k5_asympt_0(X1,X2),k5_numbers,k1_numbers)&v2_asympt_0(k5_asympt_0(X1,X2))))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t26_asympt_0', fc1_asympt_0)).
fof(l23_asympt_0, axiom, ![X1]:((v1_funct_1(X1)&(v1_funct_2(X1,k5_numbers,k1_numbers)&(v2_asympt_0(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))=>![X2]:((v1_funct_1(X2)&(v1_funct_2(X2,k5_numbers,k1_numbers)&(v2_asympt_0(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))=>((r2_hidden(X1,k6_asympt_0(X2))&r2_hidden(X2,k6_asympt_0(X1)))<=>k6_asympt_0(X1)=k6_asympt_0(X2)))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t26_asympt_0', l23_asympt_0)).
fof(t26_asympt_0, conjecture, ![X1]:((v1_funct_1(X1)&(v1_funct_2(X1,k5_numbers,k1_numbers)&(v2_asympt_0(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))=>![X2]:((v1_funct_1(X2)&(v1_funct_2(X2,k5_numbers,k1_numbers)&(v2_asympt_0(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))=>k7_asympt_0(k1_asympt_0(X1,X2))=k7_asympt_0(k5_asympt_0(X1,X2)))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t26_asympt_0', t26_asympt_0)).
fof(c_0_13, plain, ![X147, X148, X149]:(~v1_funct_1(X147)|(~v1_funct_2(X147,k5_numbers,k1_numbers)|(~v2_asympt_0(X147)|~m1_subset_1(X147,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))|(~v1_funct_1(X148)|(~v1_funct_2(X148,k5_numbers,k1_numbers)|(~v2_asympt_0(X148)|~m1_subset_1(X148,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))|(~v1_funct_1(X149)|(~v1_funct_2(X149,k5_numbers,k1_numbers)|(~v2_asympt_0(X149)|~m1_subset_1(X149,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))|(~r2_hidden(X147,k7_asympt_0(X148))|~r2_hidden(X148,k7_asympt_0(X149))|r2_hidden(X147,k7_asympt_0(X149)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t21_asympt_0])])])).
fof(c_0_14, plain, ![X237, X238]:((v1_funct_1(X237)|~r2_hidden(X237,k7_asympt_0(X238))|(~v1_funct_1(X238)|(~v1_funct_2(X238,k5_numbers,k1_numbers)|(~v2_asympt_0(X238)|~m1_subset_1(X238,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))&((v1_funct_2(X237,k5_numbers,k1_numbers)|~r2_hidden(X237,k7_asympt_0(X238))|(~v1_funct_1(X238)|(~v1_funct_2(X238,k5_numbers,k1_numbers)|(~v2_asympt_0(X238)|~m1_subset_1(X238,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))&((v2_asympt_0(X237)|~r2_hidden(X237,k7_asympt_0(X238))|(~v1_funct_1(X238)|(~v1_funct_2(X238,k5_numbers,k1_numbers)|(~v2_asympt_0(X238)|~m1_subset_1(X238,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))&(m1_subset_1(X237,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r2_hidden(X237,k7_asympt_0(X238))|(~v1_funct_1(X238)|(~v1_funct_2(X238,k5_numbers,k1_numbers)|(~v2_asympt_0(X238)|~m1_subset_1(X238,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t18_asympt_0])])])).
fof(c_0_15, plain, ![X90, X91]:((~r2_hidden(X90,k7_asympt_0(X91))|r2_hidden(X91,k6_asympt_0(X90))|(~v1_funct_1(X91)|(~v1_funct_2(X91,k5_numbers,k1_numbers)|(~v2_asympt_0(X91)|~m1_subset_1(X91,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))|(~v1_funct_1(X90)|(~v1_funct_2(X90,k5_numbers,k1_numbers)|(~v2_asympt_0(X90)|~m1_subset_1(X90,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))&(~r2_hidden(X91,k6_asympt_0(X90))|r2_hidden(X90,k7_asympt_0(X91))|(~v1_funct_1(X91)|(~v1_funct_2(X91,k5_numbers,k1_numbers)|(~v2_asympt_0(X91)|~m1_subset_1(X91,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))|(~v1_funct_1(X90)|(~v1_funct_2(X90,k5_numbers,k1_numbers)|(~v2_asympt_0(X90)|~m1_subset_1(X90,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t19_asympt_0])])])])).
fof(c_0_16, plain, ![X235, X236]:((v1_funct_1(X235)|~r2_hidden(X235,k6_asympt_0(X236))|(~v1_funct_1(X236)|(~v1_funct_2(X236,k5_numbers,k1_numbers)|(~v2_asympt_0(X236)|~m1_subset_1(X236,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))&((v1_funct_2(X235,k5_numbers,k1_numbers)|~r2_hidden(X235,k6_asympt_0(X236))|(~v1_funct_1(X236)|(~v1_funct_2(X236,k5_numbers,k1_numbers)|(~v2_asympt_0(X236)|~m1_subset_1(X236,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))&((v2_asympt_0(X235)|~r2_hidden(X235,k6_asympt_0(X236))|(~v1_funct_1(X236)|(~v1_funct_2(X236,k5_numbers,k1_numbers)|(~v2_asympt_0(X236)|~m1_subset_1(X236,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))&(m1_subset_1(X235,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r2_hidden(X235,k6_asympt_0(X236))|(~v1_funct_1(X236)|(~v1_funct_2(X236,k5_numbers,k1_numbers)|(~v2_asympt_0(X236)|~m1_subset_1(X236,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_asympt_0])])])).
cnf(c_0_17, plain, (r2_hidden(X1,k7_asympt_0(X3))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v2_asympt_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v2_asympt_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v2_asympt_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r2_hidden(X1,k7_asympt_0(X2))|~r2_hidden(X2,k7_asympt_0(X3))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, plain, (v1_funct_1(X1)|~r2_hidden(X1,k7_asympt_0(X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v2_asympt_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (v1_funct_2(X1,k5_numbers,k1_numbers)|~r2_hidden(X1,k7_asympt_0(X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v2_asympt_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_20, plain, (v2_asympt_0(X1)|~r2_hidden(X1,k7_asympt_0(X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v2_asympt_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_21, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r2_hidden(X1,k7_asympt_0(X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v2_asympt_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_22, plain, (r2_hidden(X2,k7_asympt_0(X1))|~r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v2_asympt_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v2_asympt_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (v1_funct_1(X1)|~r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v2_asympt_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_24, plain, (v1_funct_2(X1,k5_numbers,k1_numbers)|~r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v2_asympt_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_25, plain, (v2_asympt_0(X1)|~r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v2_asympt_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_26, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v2_asympt_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_27, plain, (r2_hidden(X1,k7_asympt_0(X3))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X3)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~r2_hidden(X2,k7_asympt_0(X3))|~r2_hidden(X1,k7_asympt_0(X2))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_17])).
cnf(c_0_28, plain, (v1_funct_1(X1)|~v1_funct_1(X2)|~v2_asympt_0(X2)|~r2_hidden(X1,k7_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_18])).
cnf(c_0_29, plain, (v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v2_asympt_0(X2)|~r2_hidden(X1,k7_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_19])).
cnf(c_0_30, plain, (v2_asympt_0(X1)|~v1_funct_1(X2)|~v2_asympt_0(X2)|~r2_hidden(X1,k7_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_20])).
cnf(c_0_31, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v2_asympt_0(X2)|~r2_hidden(X1,k7_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_21])).
cnf(c_0_32, plain, (r2_hidden(X2,k7_asympt_0(X1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_22])).
cnf(c_0_33, plain, (v1_funct_1(X1)|~v1_funct_1(X2)|~v2_asympt_0(X2)|~r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_23])).
cnf(c_0_34, plain, (v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v2_asympt_0(X2)|~r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_24])).
cnf(c_0_35, plain, (v2_asympt_0(X1)|~v1_funct_1(X2)|~v2_asympt_0(X2)|~r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_25])).
cnf(c_0_36, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v2_asympt_0(X2)|~r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_26])).
cnf(c_0_37, plain, (r2_hidden(X1,k7_asympt_0(X2))|~r2_hidden(X3,k7_asympt_0(X2))|~r2_hidden(X1,k7_asympt_0(X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_27, c_0_28]), c_0_28]), c_0_29]), c_0_30]), c_0_31])).
cnf(c_0_38, plain, (r2_hidden(X1,k7_asympt_0(X2))|~r2_hidden(X2,k6_asympt_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_32, c_0_33]), c_0_34]), c_0_35]), c_0_36])).
fof(c_0_39, plain, ![X50, X51, X52, X53, X54]:((~r1_tarski(X50,X51)|(~r2_hidden(X52,X50)|r2_hidden(X52,X51)))&((r2_hidden(esk9_2(X53,X54),X53)|r1_tarski(X53,X54))&(~r2_hidden(esk9_2(X53,X54),X54)|r1_tarski(X53,X54)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_40, plain, (r2_hidden(X1,k7_asympt_0(X2))|~r2_hidden(X1,k7_asympt_0(X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~r2_hidden(X2,k6_asympt_0(X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38])])).
cnf(c_0_41, plain, (r2_hidden(esk9_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_42, plain, ![X35, X36]:(((r1_tarski(X35,X36)|X35!=X36)&(r1_tarski(X36,X35)|X35!=X36))&(~r1_tarski(X35,X36)|~r1_tarski(X36,X35)|X35=X36)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_43, plain, (r1_tarski(X1,X2)|~r2_hidden(esk9_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_44, plain, (r2_hidden(X1,k7_asympt_0(X2))|~r2_hidden(X1,k7_asympt_0(X3))|~r2_hidden(X2,k6_asympt_0(X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_40, c_0_33]), c_0_29]), c_0_34]), c_0_30]), c_0_35]), c_0_31]), c_0_36])).
cnf(c_0_45, plain, (r1_tarski(X1,X2)|r2_hidden(esk9_2(X1,X2),X1)), inference(evalgc,[status(thm)],[c_0_41])).
cnf(c_0_46, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_47, plain, (r1_tarski(X1,X2)|~r2_hidden(esk9_2(X1,X2),X2)), inference(evalgc,[status(thm)],[c_0_43])).
cnf(c_0_48, plain, (r2_hidden(esk9_2(k7_asympt_0(X1),X2),k7_asympt_0(X3))|r1_tarski(k7_asympt_0(X1),X2)|~r2_hidden(X3,k6_asympt_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45])])).
fof(c_0_49, plain, ![X106]:(~v1_funct_1(X106)|(~v1_funct_2(X106,k5_numbers,k1_numbers)|(~v2_asympt_0(X106)|~m1_subset_1(X106,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))|r2_hidden(X106,k6_asympt_0(X106))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_asympt_0])])).
fof(c_0_50, plain, ![X95, X96]:(~v1_funct_1(X95)|(~v1_funct_2(X95,k5_numbers,k1_numbers)|(~v2_asympt_0(X95)|~m1_subset_1(X95,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))|(~v1_funct_1(X96)|(~v1_funct_2(X96,k5_numbers,k1_numbers)|(~v2_asympt_0(X96)|~m1_subset_1(X96,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))|k6_asympt_0(k1_asympt_0(X95,X96))=k6_asympt_0(k5_asympt_0(X95,X96)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_asympt_0])])])).
cnf(c_0_51, plain, (X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_46])).
cnf(c_0_52, plain, (r1_tarski(k7_asympt_0(X1),k7_asympt_0(X2))|~r2_hidden(X2,k6_asympt_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_48])])).
cnf(c_0_53, plain, (r2_hidden(X1,k6_asympt_0(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v2_asympt_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_54, plain, (k6_asympt_0(k1_asympt_0(X1,X2))=k6_asympt_0(k5_asympt_0(X1,X2))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v2_asympt_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v2_asympt_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_50])).
fof(c_0_55, plain, ![X56, X57]:((v1_funct_1(k1_asympt_0(X56,X57))|(~v1_funct_1(X56)|(~v1_funct_2(X56,k5_numbers,k1_numbers)|(~v2_asympt_0(X56)|~m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))|(~v1_funct_1(X57)|(~v1_funct_2(X57,k5_numbers,k1_numbers)|(~v2_asympt_0(X57)|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))))&((v1_funct_2(k1_asympt_0(X56,X57),k5_numbers,k1_numbers)|(~v1_funct_1(X56)|(~v1_funct_2(X56,k5_numbers,k1_numbers)|(~v2_asympt_0(X56)|~m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))|(~v1_funct_1(X57)|(~v1_funct_2(X57,k5_numbers,k1_numbers)|(~v2_asympt_0(X57)|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))))&((v2_asympt_0(k1_asympt_0(X56,X57))|(~v1_funct_1(X56)|(~v1_funct_2(X56,k5_numbers,k1_numbers)|(~v2_asympt_0(X56)|~m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))|(~v1_funct_1(X57)|(~v1_funct_2(X57,k5_numbers,k1_numbers)|(~v2_asympt_0(X57)|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))))&(m1_subset_1(k1_asympt_0(X56,X57),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|(~v1_funct_1(X56)|(~v1_funct_2(X56,k5_numbers,k1_numbers)|(~v2_asympt_0(X56)|~m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))|(~v1_funct_1(X57)|(~v1_funct_2(X57,k5_numbers,k1_numbers)|(~v2_asympt_0(X57)|~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_asympt_0])])])).
cnf(c_0_56, plain, (k7_asympt_0(X1)=k7_asympt_0(X2)|~r1_tarski(k7_asympt_0(X1),k7_asympt_0(X2))|~r2_hidden(X1,k6_asympt_0(X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52])])).
cnf(c_0_57, plain, (r2_hidden(X1,k6_asympt_0(X1))|~v1_funct_1(X1)|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_53])).
cnf(c_0_58, plain, (k6_asympt_0(k1_asympt_0(X1,X2))=k6_asympt_0(k5_asympt_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_54])).
cnf(c_0_59, plain, (v1_funct_1(k1_asympt_0(X1,X2))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v2_asympt_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v2_asympt_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_60, plain, (v1_funct_2(k1_asympt_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v2_asympt_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v2_asympt_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_61, plain, (v2_asympt_0(k1_asympt_0(X1,X2))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v2_asympt_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v2_asympt_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_62, plain, (m1_subset_1(k1_asympt_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v2_asympt_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v2_asympt_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_63, plain, (k7_asympt_0(X1)=k7_asympt_0(X2)|~r2_hidden(X1,k6_asympt_0(X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~r2_hidden(X2,k6_asympt_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_56, c_0_52])])).
cnf(c_0_64, plain, (r2_hidden(k1_asympt_0(X1,X2),k6_asympt_0(k5_asympt_0(X1,X2)))|~m1_subset_1(k1_asympt_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(k1_asympt_0(X1,X2))|~v1_funct_2(k1_asympt_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_asympt_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58])])).
cnf(c_0_65, plain, (v1_funct_1(k1_asympt_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_59])).
cnf(c_0_66, plain, (v1_funct_2(k1_asympt_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_60])).
cnf(c_0_67, plain, (v2_asympt_0(k1_asympt_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_61])).
cnf(c_0_68, plain, (m1_subset_1(k1_asympt_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_62])).
fof(c_0_69, plain, ![X65, X66]:((v1_funct_1(k5_asympt_0(X65,X66))|(~v1_funct_1(X65)|(~v1_funct_2(X65,k5_numbers,k1_numbers)|~m1_subset_1(X65,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|(~v1_funct_1(X66)|(~v1_funct_2(X66,k5_numbers,k1_numbers)|~m1_subset_1(X66,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))&((v1_funct_2(k5_asympt_0(X65,X66),k5_numbers,k1_numbers)|(~v1_funct_1(X65)|(~v1_funct_2(X65,k5_numbers,k1_numbers)|~m1_subset_1(X65,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|(~v1_funct_1(X66)|(~v1_funct_2(X66,k5_numbers,k1_numbers)|~m1_subset_1(X66,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))&(m1_subset_1(k5_asympt_0(X65,X66),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|(~v1_funct_1(X65)|(~v1_funct_2(X65,k5_numbers,k1_numbers)|~m1_subset_1(X65,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|(~v1_funct_1(X66)|(~v1_funct_2(X66,k5_numbers,k1_numbers)|~m1_subset_1(X66,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_asympt_0])])])).
fof(c_0_70, plain, ![X101, X102]:((v1_funct_1(k5_asympt_0(X101,X102))|(~v1_funct_1(X101)|(~v1_funct_2(X101,k5_numbers,k1_numbers)|~m1_subset_1(X101,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|(~v1_funct_1(X102)|(~v1_funct_2(X102,k5_numbers,k1_numbers)|(~v2_asympt_0(X102)|~m1_subset_1(X102,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))))&((v1_funct_2(k5_asympt_0(X101,X102),k5_numbers,k1_numbers)|(~v1_funct_1(X101)|(~v1_funct_2(X101,k5_numbers,k1_numbers)|~m1_subset_1(X101,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|(~v1_funct_1(X102)|(~v1_funct_2(X102,k5_numbers,k1_numbers)|(~v2_asympt_0(X102)|~m1_subset_1(X102,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))))&(v2_asympt_0(k5_asympt_0(X101,X102))|(~v1_funct_1(X101)|(~v1_funct_2(X101,k5_numbers,k1_numbers)|~m1_subset_1(X101,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|(~v1_funct_1(X102)|(~v1_funct_2(X102,k5_numbers,k1_numbers)|(~v2_asympt_0(X102)|~m1_subset_1(X102,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_asympt_0])])])).
cnf(c_0_71, plain, (k7_asympt_0(X1)=k7_asympt_0(X2)|~r2_hidden(X1,k6_asympt_0(X2))|~r2_hidden(X2,k6_asympt_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_63, c_0_33]), c_0_34]), c_0_35]), c_0_36])).
cnf(c_0_72, plain, (r2_hidden(k1_asympt_0(X1,X2),k6_asympt_0(k5_asympt_0(X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_64, c_0_65]), c_0_66]), c_0_67]), c_0_68])).
cnf(c_0_73, plain, (v1_funct_1(k5_asympt_0(X1,X2))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_74, plain, (v1_funct_2(k5_asympt_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_75, plain, (v2_asympt_0(k5_asympt_0(X1,X2))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v2_asympt_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_76, plain, (m1_subset_1(k5_asympt_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_69])).
fof(c_0_77, plain, ![X192, X193]:((~r2_hidden(X192,k6_asympt_0(X193))|~r2_hidden(X193,k6_asympt_0(X192))|k6_asympt_0(X192)=k6_asympt_0(X193)|(~v1_funct_1(X193)|(~v1_funct_2(X193,k5_numbers,k1_numbers)|(~v2_asympt_0(X193)|~m1_subset_1(X193,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))|(~v1_funct_1(X192)|(~v1_funct_2(X192,k5_numbers,k1_numbers)|(~v2_asympt_0(X192)|~m1_subset_1(X192,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))&((r2_hidden(X192,k6_asympt_0(X193))|k6_asympt_0(X192)!=k6_asympt_0(X193)|(~v1_funct_1(X193)|(~v1_funct_2(X193,k5_numbers,k1_numbers)|(~v2_asympt_0(X193)|~m1_subset_1(X193,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))|(~v1_funct_1(X192)|(~v1_funct_2(X192,k5_numbers,k1_numbers)|(~v2_asympt_0(X192)|~m1_subset_1(X192,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))&(r2_hidden(X193,k6_asympt_0(X192))|k6_asympt_0(X192)!=k6_asympt_0(X193)|(~v1_funct_1(X193)|(~v1_funct_2(X193,k5_numbers,k1_numbers)|(~v2_asympt_0(X193)|~m1_subset_1(X193,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))|(~v1_funct_1(X192)|(~v1_funct_2(X192,k5_numbers,k1_numbers)|(~v2_asympt_0(X192)|~m1_subset_1(X192,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[l23_asympt_0])])])])).
fof(c_0_78, negated_conjecture, ~(![X1]:((v1_funct_1(X1)&(v1_funct_2(X1,k5_numbers,k1_numbers)&(v2_asympt_0(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))=>![X2]:((v1_funct_1(X2)&(v1_funct_2(X2,k5_numbers,k1_numbers)&(v2_asympt_0(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))=>k7_asympt_0(k1_asympt_0(X1,X2))=k7_asympt_0(k5_asympt_0(X1,X2))))), inference(assume_negation,[status(cth)],[t26_asympt_0])).
cnf(c_0_79, plain, (k7_asympt_0(k1_asympt_0(X1,X2))=k7_asympt_0(k5_asympt_0(X1,X2))|~r2_hidden(k5_asympt_0(X1,X2),k6_asympt_0(k1_asympt_0(X1,X2)))|~m1_subset_1(k5_asympt_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(k5_asympt_0(X1,X2))|~v1_funct_2(k5_asympt_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k5_asympt_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72])])).
cnf(c_0_80, plain, (v1_funct_1(k5_asympt_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_73])).
cnf(c_0_81, plain, (v1_funct_2(k5_asympt_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_74])).
cnf(c_0_82, plain, (v2_asympt_0(k5_asympt_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_75])).
cnf(c_0_83, plain, (m1_subset_1(k5_asympt_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_84, plain, (r2_hidden(X1,k6_asympt_0(X2))|k6_asympt_0(X2)!=k6_asympt_0(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v2_asympt_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v2_asympt_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_77])).
fof(c_0_85, negated_conjecture, ((v1_funct_1(esk1_0)&(v1_funct_2(esk1_0,k5_numbers,k1_numbers)&(v2_asympt_0(esk1_0)&m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))&((v1_funct_1(esk2_0)&(v1_funct_2(esk2_0,k5_numbers,k1_numbers)&(v2_asympt_0(esk2_0)&m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))&k7_asympt_0(k1_asympt_0(esk1_0,esk2_0))!=k7_asympt_0(k5_asympt_0(esk1_0,esk2_0)))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_78])])])).
cnf(c_0_86, plain, (k7_asympt_0(k1_asympt_0(X1,X2))=k7_asympt_0(k5_asympt_0(X1,X2))|~r2_hidden(k5_asympt_0(X1,X2),k6_asympt_0(k1_asympt_0(X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_79, c_0_80]), c_0_81]), c_0_82]), c_0_83])).
cnf(c_0_87, plain, (r2_hidden(X1,k6_asympt_0(X2))|k6_asympt_0(X1)!=k6_asympt_0(X2)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_84])).
cnf(c_0_88, negated_conjecture, (k7_asympt_0(k1_asympt_0(esk1_0,esk2_0))!=k7_asympt_0(k5_asympt_0(esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_89, plain, (k7_asympt_0(k1_asympt_0(X1,X2))=k7_asympt_0(k5_asympt_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|k6_asympt_0(k1_asympt_0(X1,X2))!=k6_asympt_0(k5_asympt_0(X1,X2))|~m1_subset_1(k1_asympt_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k5_asympt_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(k1_asympt_0(X1,X2))|~v2_asympt_0(k5_asympt_0(X1,X2))|~v1_funct_2(k1_asympt_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_2(k5_asympt_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_asympt_0(X1,X2))|~v1_funct_1(k5_asympt_0(X1,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87])])).
cnf(c_0_90, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_91, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_92, negated_conjecture, (v2_asympt_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_93, negated_conjecture, (v2_asympt_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_94, negated_conjecture, (v1_funct_2(esk2_0,k5_numbers,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_95, negated_conjecture, (v1_funct_2(esk1_0,k5_numbers,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_96, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_97, negated_conjecture, (v1_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_98, negated_conjecture, (k7_asympt_0(k1_asympt_0(esk1_0,esk2_0))!=k7_asympt_0(k5_asympt_0(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[c_0_88])).
cnf(c_0_99, plain, (k7_asympt_0(k1_asympt_0(X1,X2))=k7_asympt_0(k5_asympt_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_89, c_0_80]), c_0_65]), c_0_81]), c_0_66]), c_0_82]), c_0_67]), c_0_83]), c_0_68]), c_0_58])).
cnf(c_0_100, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_101, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_91])).
cnf(c_0_102, negated_conjecture, (v2_asympt_0(esk2_0)), inference(evalgc,[status(thm)],[c_0_92])).
cnf(c_0_103, negated_conjecture, (v2_asympt_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_93])).
cnf(c_0_104, negated_conjecture, (v1_funct_2(esk2_0,k5_numbers,k1_numbers)), inference(evalgc,[status(thm)],[c_0_94])).
cnf(c_0_105, negated_conjecture, (v1_funct_2(esk1_0,k5_numbers,k1_numbers)), inference(evalgc,[status(thm)],[c_0_95])).
cnf(c_0_106, negated_conjecture, (v1_funct_1(esk2_0)), inference(evalgc,[status(thm)],[c_0_96])).
cnf(c_0_107, negated_conjecture, (v1_funct_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_97])).
cnf(c_0_108, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106]), c_0_107])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 109
# Proof object clause steps            : 82
# Proof object formula steps           : 27
# Proof object conjectures             : 22
# Proof object clause conjectures      : 19
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 33
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 9
# Proof object simplifying inferences  : 46
# Training examples: 41 positive, 353 negative
# Training: Positive examples begin
cnf(c_0_107, negated_conjecture, (v1_funct_1(esk1_0))).# trainpos
cnf(c_0_106, negated_conjecture, (v1_funct_1(esk2_0))).# trainpos
cnf(c_0_103, negated_conjecture, (v2_asympt_0(esk1_0))).# trainpos
cnf(c_0_102, negated_conjecture, (v2_asympt_0(esk2_0))).# trainpos
cnf(c_0_98, negated_conjecture, (k7_asympt_0(k1_asympt_0(esk1_0,esk2_0))!=k7_asympt_0(k5_asympt_0(esk1_0,esk2_0)))).# trainpos
cnf(c_0_105, negated_conjecture, (v1_funct_2(esk1_0,k5_numbers,k1_numbers))).# trainpos
cnf(c_0_104, negated_conjecture, (v1_funct_2(esk2_0,k5_numbers,k1_numbers))).# trainpos
cnf(c_0_101, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_100, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_51, plain, (X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_45, plain, (r1_tarski(X1,X2)|r2_hidden(esk9_2(X1,X2),X1))).# trainpos
cnf(c_0_47, plain, (r1_tarski(X1,X2)|~r2_hidden(esk9_2(X1,X2),X2))).# trainpos
cnf(c_0_28, plain, (v1_funct_1(X1)|~v1_funct_1(X2)|~v2_asympt_0(X2)|~r2_hidden(X1,k7_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_30, plain, (v2_asympt_0(X1)|~v1_funct_1(X2)|~v2_asympt_0(X2)|~r2_hidden(X1,k7_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_57, plain, (r2_hidden(X1,k6_asympt_0(X1))|~v1_funct_1(X1)|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_33, plain, (v1_funct_1(X1)|~v1_funct_1(X2)|~v2_asympt_0(X2)|~r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_29, plain, (v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v2_asympt_0(X2)|~r2_hidden(X1,k7_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_35, plain, (v2_asympt_0(X1)|~v1_funct_1(X2)|~v2_asympt_0(X2)|~r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_34, plain, (v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v2_asympt_0(X2)|~r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_80, plain, (v1_funct_1(k5_asympt_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_31, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v2_asympt_0(X2)|~r2_hidden(X1,k7_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_36, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v2_asympt_0(X2)|~r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_81, plain, (v1_funct_2(k5_asympt_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_83, plain, (m1_subset_1(k5_asympt_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_82, plain, (v2_asympt_0(k5_asympt_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_65, plain, (v1_funct_1(k1_asympt_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_67, plain, (v2_asympt_0(k1_asympt_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_87, plain, (r2_hidden(X1,k6_asympt_0(X2))|k6_asympt_0(X1)!=k6_asympt_0(X2)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_66, plain, (v1_funct_2(k1_asympt_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_68, plain, (m1_subset_1(k1_asympt_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_27, plain, (r2_hidden(X1,k7_asympt_0(X3))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X3)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~r2_hidden(X2,k7_asympt_0(X3))|~r2_hidden(X1,k7_asympt_0(X2))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_32, plain, (r2_hidden(X2,k7_asympt_0(X1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_58, plain, (k6_asympt_0(k1_asympt_0(X1,X2))=k6_asympt_0(k5_asympt_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_40, plain, (r2_hidden(X1,k7_asympt_0(X2))|~r2_hidden(X1,k7_asympt_0(X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~r2_hidden(X2,k6_asympt_0(X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).# trainpos
cnf(c_0_48, plain, (r2_hidden(esk9_2(k7_asympt_0(X1),X2),k7_asympt_0(X3))|r1_tarski(k7_asympt_0(X1),X2)|~r2_hidden(X3,k6_asympt_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).# trainpos
cnf(c_0_64, plain, (r2_hidden(k1_asympt_0(X1,X2),k6_asympt_0(k5_asympt_0(X1,X2)))|~m1_subset_1(k1_asympt_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(k1_asympt_0(X1,X2))|~v1_funct_2(k1_asympt_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_asympt_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1))).# trainpos
cnf(c_0_52, plain, (r1_tarski(k7_asympt_0(X1),k7_asympt_0(X2))|~r2_hidden(X2,k6_asympt_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).# trainpos
cnf(c_0_56, plain, (k7_asympt_0(X1)=k7_asympt_0(X2)|~r1_tarski(k7_asympt_0(X1),k7_asympt_0(X2))|~r2_hidden(X1,k6_asympt_0(X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).# trainpos
cnf(c_0_63, plain, (k7_asympt_0(X1)=k7_asympt_0(X2)|~r2_hidden(X1,k6_asympt_0(X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~r2_hidden(X2,k6_asympt_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).# trainpos
cnf(c_0_79, plain, (k7_asympt_0(k1_asympt_0(X1,X2))=k7_asympt_0(k5_asympt_0(X1,X2))|~r2_hidden(k5_asympt_0(X1,X2),k6_asympt_0(k1_asympt_0(X1,X2)))|~m1_subset_1(k5_asympt_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(k5_asympt_0(X1,X2))|~v1_funct_2(k5_asympt_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k5_asympt_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1))).# trainpos
cnf(c_0_89, plain, (k7_asympt_0(k1_asympt_0(X1,X2))=k7_asympt_0(k5_asympt_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|k6_asympt_0(k1_asympt_0(X1,X2))!=k6_asympt_0(k5_asympt_0(X1,X2))|~m1_subset_1(k1_asympt_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k5_asympt_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(k1_asympt_0(X1,X2))|~v2_asympt_0(k5_asympt_0(X1,X2))|~v1_funct_2(k1_asympt_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_2(k5_asympt_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_asympt_0(X1,X2))|~v1_funct_1(k5_asympt_0(X1,X2)))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_37, plain, (k4_ordinal1=k5_numbers)).#trainneg
cnf(i_0_185, plain, (k1_xboole_0=k6_numbers)).#trainneg
cnf(i_0_208, plain, (o_0_0_xboole_0=k1_xboole_0)).#trainneg
cnf(i_0_23, plain, (v3_membered(k1_numbers))).#trainneg
cnf(i_0_99, plain, (v1_xboole_0(np__0))).#trainneg
cnf(i_0_50, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_96, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_131, plain, (v6_membered(k4_ordinal1))).#trainneg
cnf(i_0_46, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_51, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_80, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_162, plain, (~r1_xxreal_0(np__1,np__0))).#trainneg
cnf(i_0_56, plain, (m1_subset_1(np__0,k5_numbers))).#trainneg
cnf(i_0_160, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_129, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_55, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_94, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_93, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_163, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_38, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_57, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_95, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_157, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_132, plain, (v1_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_149, plain, (v2_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_27, plain, (r1_tarski(X1,X2)|X1!=X2)).#trainneg
cnf(i_0_28, plain, (r1_tarski(X1,X2)|X1!=X2)).#trainneg
cnf(i_0_128, plain, (k5_xcmplx_0(k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_79, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_217, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11, plain, (X1=k5_asympt_0(X2,X3)|k2_square_1(k1_seq_1(X2,esk3_3(X2,X3,X1)),k1_seq_1(X3,esk3_3(X2,X3,X1)))!=k1_seq_1(X1,esk3_3(X2,X3,X1))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_150, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_90, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_142, plain, (~v1_xboole_0(X1)|~m1_funct_2(X1,X2,X3))).#trainneg
cnf(i_0_205, plain, (m1_subset_1(k2_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_171, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_141, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_119, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_125, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_164, plain, (k2_real_1(k2_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_159, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_170, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_48, plain, (k4_xxreal_0(X1,X1)=X1|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_176, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_215, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_49, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_210, plain, (v1_xxreal_0(k4_xxreal_0(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_75, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_199, plain, (v2_asympt_0(X1)|~v1_funct_1(X1)|~v4_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_116, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_102, plain, (v1_funct_1(k3_series_1(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_148, plain, (~v1_funct_1(X2)|~v1_funct_1(X1)|~v4_asympt_0(X2)|~v4_asympt_0(X1)|~r2_hidden(X1,k7_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_limfunc1(k52_valued_1(k5_numbers,k1_numbers,k1_numbers,X2,X1)))).#trainneg
cnf(i_0_118, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_195, plain, (v1_series_1(X1)|~v1_funct_1(X1)|~v2_series_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_169, plain, (v1_xreal_0(k1_seq_2(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_161, plain, (m1_subset_1(k2_seq_2(X1),k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_198, plain, (v5_asympt_0(X1)|~v1_funct_1(X1)|~v4_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_120, plain, (v1_xboole_0(X4)|m2_funct_2(X1,X3,X4,X2)|~m1_subset_1(X1,X2)|~m1_funct_2(X2,X3,X4))).#trainneg
cnf(i_0_134, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_47, plain, (k4_xxreal_0(X1,X2)=k4_xxreal_0(X2,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_74, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_17, plain, (r1_xxreal_0(X1,k4_xxreal_0(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_177, plain, (k2_square_1(X1,X1)=X1|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_172, plain, (k7_nat_1(X1,X1)=X1|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_41, plain, (r2_hidden(X3,X2)|~r2_hidden(X3,X1)|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_216, plain, (r1_tarski(X1,X3)|~r1_tarski(X2,X3)|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_127, plain, (k7_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_180, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_117, plain, (k8_real_1(X1,X2)=k8_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_146, plain, (m1_subset_1(k7_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_158, plain, (m1_subset_1(k8_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_103, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_36, plain, (a_1_1_asympt_0(X1)=k7_asympt_0(X1)|~v1_funct_1(X1)|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_101, plain, (v1_funct_2(k3_series_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_100, plain, (m1_subset_1(k3_series_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_182, plain, (~v1_funct_1(X2)|~v1_funct_1(X1)|~v4_asympt_0(X2)|~v4_asympt_0(X1)|~r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_limfunc1(k52_valued_1(k5_numbers,k1_numbers,k1_numbers,X1,X2)))).#trainneg
cnf(i_0_32, plain, (k4_xxreal_0(X1,X2)=X2|k4_xxreal_0(X1,X2)=X1|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_138, plain, (k2_square_1(X1,X2)=k4_xxreal_0(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_151, plain, (v1_series_1(X1)|~v1_funct_1(X1)|~v2_comseq_2(k3_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_123, plain, (v2_series_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_series_1(k56_valued_1(k5_numbers,k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_97, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_165, plain, (v4_asympt_0(X1)|~v1_funct_1(X1)|~v2_asympt_0(X1)|~v5_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_133, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_179, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_143, plain, (k2_square_1(X1,X2)=k2_square_1(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_135, plain, (k7_nat_1(X1,X2)=k4_xxreal_0(X1,X2)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_168, plain, (k7_nat_1(X1,X2)=k7_nat_1(X2,X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_24, plain, (X1=X2|r2_hidden(esk6_2(X1,X2),X2)|r2_hidden(esk6_2(X1,X2),X1))).#trainneg
cnf(i_0_25, plain, (X1=X2|~r2_hidden(esk6_2(X1,X2),X2)|~r2_hidden(esk6_2(X1,X2),X1))).#trainneg
cnf(i_0_122, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_121, plain, (v1_xboole_0(X3)|m1_subset_1(X1,X4)|~m1_funct_2(X4,X2,X3)|~m2_funct_2(X1,X2,X3,X4))).#trainneg
cnf(i_0_104, plain, (k2_series_1(X1)=k3_series_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_113, plain, (k1_seq_2(X1)=k2_seq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_98, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_82, plain, (k2_seq_2(k3_series_1(X1))=k4_series_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_16, plain, (r1_xxreal_0(X1,X3)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X3)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_18, plain, (v2_asympt_0(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r1_xxreal_0(k6_numbers,k1_seq_1(X1,esk5_2(X1,X2))))).#trainneg
cnf(i_0_78, plain, (r2_hidden(X1,k7_asympt_0(X1))|~v1_funct_1(X1)|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_62, plain, (v2_limfunc1(X1)|m1_subset_1(esk11_1(X1),k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_63, plain, (~v1_funct_1(X2)|~v2_limfunc1(X2)|~m1_subset_1(X3,k1_numbers)|~r1_xxreal_0(X3,k1_seq_1(X2,X1))|~r1_xxreal_0(esk10_2(X2,X3),X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_136, plain, (k1_valued_1(X1,X2)=k1_valued_1(X2,X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1))).#trainneg
cnf(i_0_58, plain, (a_1_0_asympt_0(X1)=k6_asympt_0(X1)|~v1_funct_1(X1)|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_152, plain, (v2_comseq_2(k3_series_1(X1))|~v1_funct_1(X1)|~v1_series_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_137, plain, (k18_valued_1(X1,X2)=k18_valued_1(X2,X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1))).#trainneg
cnf(i_0_144, plain, (v1_funct_1(k1_valued_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1))).#trainneg
cnf(i_0_145, plain, (v1_relat_1(k1_valued_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1))).#trainneg
cnf(i_0_68, plain, (v1_limfunc1(X1)|m1_subset_1(esk14_1(X1),k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_203, plain, (v1_xboole_0(k1_seq_2(X1))|~v1_funct_1(X1)|~v1_fdiff_1(X1,k6_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_213, plain, (v1_limfunc1(X1)|k1_seq_1(X1,esk20_1(X1))!=esk20_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_214, plain, (v1_limfunc1(X1)|m2_subset_1(esk20_1(X1),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_212, plain, (v2_series_1(X1)|m2_subset_1(esk19_1(X1),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_22, plain, (m2_subset_1(esk4_1(X1),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_211, plain, (v2_series_1(X1)|k8_nat_1(k1_numbers,X1,esk19_1(X1))!=k6_numbers|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_124, plain, (v1_series_1(k56_valued_1(k5_numbers,k1_numbers,X1))|~v1_funct_1(X1)|~v2_series_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_14, plain, (r1_xxreal_0(k3_xcmplx_0(X1,X3),k3_xcmplx_0(X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(k6_numbers,X3))).#trainneg
cnf(i_0_15, plain, (r1_xxreal_0(X3,X2)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k3_xcmplx_0(X3,X1),k3_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_64, plain, (m2_subset_1(esk10_2(X1,X2),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v2_limfunc1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_209, plain, (k18_complex1(k2_seq_2(X1))=k2_seq_2(k56_valued_1(k5_numbers,k1_numbers,X1))|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_69, plain, (~v1_funct_1(X2)|~v1_limfunc1(X2)|~m1_subset_1(X3,k1_numbers)|~r1_xxreal_0(k1_seq_1(X2,X1),X3)|~r1_xxreal_0(esk13_2(X2,X3),X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_10, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(X2,X4))|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_83, plain, (k8_nat_1(X2,X1,X3)=k1_funct_1(X1,X3)|~v1_funct_1(X1)|~v7_ordinal1(X3)|~v1_funct_2(X1,k5_numbers,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_70, plain, (m2_subset_1(esk13_2(X1,X2),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_limfunc1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_19, plain, (v2_asympt_0(X2)|r1_xxreal_0(X1,esk5_2(X2,X1))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_60, plain, (v2_limfunc1(X2)|r1_xxreal_0(X1,esk12_2(X2,X1))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_66, plain, (v1_limfunc1(X2)|r1_xxreal_0(X1,esk15_2(X2,X1))|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_20, plain, (v2_asympt_0(X1)|m2_subset_1(esk5_2(X1,X2),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_61, plain, (v2_limfunc1(X1)|m2_subset_1(esk12_2(X1,X2),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_67, plain, (v1_limfunc1(X1)|m2_subset_1(esk15_2(X1,X2),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_59, plain, (v2_limfunc1(X1)|r1_xxreal_0(esk11_1(X1),k1_seq_1(X1,esk12_2(X1,X2)))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_65, plain, (v1_limfunc1(X1)|r1_xxreal_0(k1_seq_1(X1,esk15_2(X1,X2)),esk14_1(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_207, plain, (r1_xxreal_0(X1,k2_seq_2(X2))|m2_subset_1(esk18_2(X1,X2),k1_numbers,k5_numbers)|~v1_funct_1(X2)|~v1_xreal_0(X1)|~v2_comseq_2(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_130, plain, (k5_asympt_0(X1,X2)=k5_asympt_0(X2,X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_184, plain, (r2_funct_2(X2,X3,X1,X1)|~v1_funct_1(X4)|~v1_funct_1(X1)|~v1_funct_2(X4,X2,X3)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_109, plain, (k2_seq_2(X1)=k2_seq_2(X2)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_comseq_2(X2)|~m2_valued_0(X1,k1_numbers,X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_173, plain, (v1_partfun1(k1_valued_1(X1,X2),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v4_relat_1(X2,X3)|~v4_relat_1(X1,X3)|~v1_partfun1(X2,X3)|~v1_partfun1(X1,X3))).#trainneg
cnf(i_0_204, plain, (v7_valued_0(X2)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v7_valued_0(X1)|~m2_valued_0(X2,k1_numbers,X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_21, plain, (r1_xxreal_0(k6_numbers,k1_seq_1(X2,X1))|~v1_funct_1(X2)|~v2_asympt_0(X2)|~r1_xxreal_0(esk4_1(X2),X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_206, plain, (r1_xxreal_0(X1,k2_seq_2(X2))|~v1_funct_1(X2)|~v1_xreal_0(X1)|~v2_comseq_2(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~r1_xxreal_0(X1,k1_seq_1(X2,esk18_2(X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_87, plain, (v4_asympt_0(k5_asympt_0(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v4_asympt_0(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_114, plain, (r2_funct_2(X3,X4,X1,X2)|X1!=X2|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,X3,X4)|~v1_funct_2(X1,X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_115, plain, (X3=X4|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_2(X4,X1,X2)|~v1_funct_2(X3,X1,X2)|~r2_funct_2(X1,X2,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_112, plain, (v1_funct_1(k1_valued_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_comseq_2(X2)|~v2_comseq_2(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_110, plain, (v2_comseq_2(k1_valued_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_comseq_2(X2)|~v2_comseq_2(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_153, plain, (k1_asympt_0(X1,X2)=k1_asympt_0(X2,X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_105, plain, (k1_asympt_0(X1,X2)=k1_valued_1(X1,X2)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_218, plain, (r2_funct_2(X2,X3,X4,X1)|~v1_funct_1(X4)|~v1_funct_1(X1)|~v1_funct_2(X4,X2,X3)|~v1_funct_2(X1,X2,X3)|~r2_funct_2(X2,X3,X1,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_107, plain, (m2_subset_1(esk17_2(X1,X2),k1_numbers,k5_numbers)|r2_funct_2(k5_numbers,k1_numbers,X1,k56_valued_1(k5_numbers,k1_numbers,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_106, plain, (r2_funct_2(k5_numbers,k1_numbers,X1,k56_valued_1(k5_numbers,k1_numbers,X2))|k1_seq_1(X1,esk17_2(X1,X2))!=k18_complex1(k1_seq_1(X2,esk17_2(X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_183, plain, (v2_comseq_2(k3_valued_1(k5_numbers,k1_numbers,k1_numbers,X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_comseq_2(X2)|~v2_comseq_2(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_111, plain, (v1_funct_2(k1_valued_1(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_comseq_2(X2)|~v2_comseq_2(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_77, plain, (r2_hidden(X2,k6_asympt_0(X1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~r2_hidden(X1,k7_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_139, plain, (r1_tarski(k6_asympt_0(X1),k6_asympt_0(X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_178, plain, (k9_real_1(k2_seq_2(X1),k2_seq_2(X2))=k2_seq_2(k47_valued_1(k5_numbers,k1_numbers,k1_numbers,X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_comseq_2(X2)|~v2_comseq_2(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_147, plain, (r2_hidden(X1,k7_asympt_0(X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v4_asympt_0(X2)|~v4_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_limfunc1(k52_valued_1(k5_numbers,k1_numbers,k1_numbers,X1,X2)))).#trainneg
cnf(i_0_181, plain, (r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v4_asympt_0(X2)|~v4_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_limfunc1(k52_valued_1(k5_numbers,k1_numbers,k1_numbers,X2,X1)))).#trainneg
cnf(i_0_156, plain, (k6_asympt_0(X1)=k6_asympt_0(X2)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~r2_hidden(X2,k6_asympt_0(X1))|~r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_108, plain, (k1_seq_1(X1,X3)=k18_complex1(k1_seq_1(X2,X3))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r2_funct_2(k5_numbers,k1_numbers,X1,k56_valued_1(k5_numbers,k1_numbers,X2)))).#trainneg
cnf(i_0_72, plain, (m2_subset_1(esk16_3(X1,X2,X3),k1_numbers,k5_numbers)|r2_funct_2(k5_numbers,k1_numbers,X2,k26_valued_1(k5_numbers,k1_numbers,X3,X1))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_xreal_0(X1)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_73, plain, (k1_seq_1(X1,X4)=k4_real_1(X3,k1_seq_1(X2,X4))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r2_funct_2(k5_numbers,k1_numbers,X1,k26_valued_1(k5_numbers,k1_numbers,X2,X3)))).#trainneg
cnf(i_0_186, plain, (k3_valued_1(k5_numbers,k1_numbers,k1_numbers,k1_valued_0(k1_numbers,X2,X1),k1_valued_0(k1_numbers,X3,X1))=k1_valued_0(k1_numbers,k3_valued_1(k5_numbers,k1_numbers,k1_numbers,X2,X3),X1)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_12, plain, (X3=k5_asympt_0(X1,X2)|m2_subset_1(esk3_3(X1,X2,X3),k1_numbers,k5_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_71, plain, (r2_funct_2(k5_numbers,k1_numbers,X1,k26_valued_1(k5_numbers,k1_numbers,X3,X2))|k1_seq_1(X1,esk16_3(X2,X1,X3))!=k4_real_1(X2,k1_seq_1(X3,esk16_3(X2,X1,X3)))|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_13, plain, (k1_seq_1(X1,X4)=k2_square_1(k1_seq_1(X2,X4),k1_seq_1(X3,X4))|X1!=k5_asympt_0(X2,X3)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_140, plain, (r2_hidden(X1,k6_asympt_0(X3))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_asympt_0(X3)|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~r2_hidden(X2,k6_asympt_0(X3))|~r2_hidden(X1,k6_asympt_0(X2))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_31, plain, (k1_seq_1(X1,X4)=k7_real_1(k1_seq_1(X2,X4),k1_seq_1(X3,X4))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r2_funct_2(k5_numbers,k1_numbers,X1,k3_valued_1(k5_numbers,k1_numbers,k1_numbers,X2,X3)))).#trainneg
cnf(i_0_30, plain, (m2_subset_1(esk7_3(X1,X2,X3),k1_numbers,k5_numbers)|r2_funct_2(k5_numbers,k1_numbers,X1,k3_valued_1(k5_numbers,k1_numbers,k1_numbers,X2,X3))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_29, plain, (r2_funct_2(k5_numbers,k1_numbers,X1,k3_valued_1(k5_numbers,k1_numbers,k1_numbers,X2,X3))|k7_real_1(k1_seq_1(X2,esk7_3(X1,X2,X3)),k1_seq_1(X3,esk7_3(X1,X2,X3)))!=k1_seq_1(X1,esk7_3(X1,X2,X3))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_34, plain, (m2_subset_1(esk8_3(X1,X2,X3),k1_numbers,k5_numbers)|r2_funct_2(k5_numbers,k1_numbers,X1,k20_valued_1(k5_numbers,k1_numbers,k1_numbers,X2,X3))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_33, plain, (r2_funct_2(k5_numbers,k1_numbers,X1,k20_valued_1(k5_numbers,k1_numbers,k1_numbers,X2,X3))|k8_real_1(k1_seq_1(X2,esk8_3(X1,X2,X3)),k1_seq_1(X3,esk8_3(X1,X2,X3)))!=k1_seq_1(X1,esk8_3(X1,X2,X3))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_35, plain, (k1_seq_1(X1,X4)=k8_real_1(k1_seq_1(X2,X4),k1_seq_1(X3,X4))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r2_funct_2(k5_numbers,k1_numbers,X1,k20_valued_1(k5_numbers,k1_numbers,k1_numbers,X2,X3)))).#trainneg
cnf(i_0_239, plain, (v1_xreal_0(np__1))).#trainneg
cnf(i_0_255, negated_conjecture, (v1_relat_1(esk2_0))).#trainneg
cnf(i_0_256, negated_conjecture, (v1_relat_1(esk1_0))).#trainneg
cnf(i_0_219, plain, (np__0=k6_numbers)).#trainneg
cnf(i_0_222, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_465, negated_conjecture, (k7_asympt_0(k1_valued_1(esk1_0,esk2_0))!=k7_asympt_0(k5_asympt_0(esk1_0,esk2_0)))).#trainneg
cnf(i_0_380, plain, (k2_seq_2(k56_valued_1(k5_numbers,k1_numbers,k3_series_1(X1)))=k18_complex1(k4_series_1(X1))|~v2_comseq_2(k3_series_1(X1))|~m1_subset_1(k3_series_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k3_series_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_470, plain, (r1_tarski(X1,k6_asympt_0(X2))|k6_asympt_0(esk9_2(X1,k6_asympt_0(X2)))!=k6_asympt_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk9_2(X1,k6_asympt_0(X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(esk9_2(X1,k6_asympt_0(X2)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(esk9_2(X1,k6_asympt_0(X2)),k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(esk9_2(X1,k6_asympt_0(X2))))).#trainneg
cnf(i_0_309, plain, (k4_xxreal_0(X1,X2)=X1|k4_xxreal_0(X1,X2)=X2|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_282, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_503, plain, (r2_hidden(X1,k6_asympt_0(esk6_2(X2,k7_asympt_0(X1))))|X2=k7_asympt_0(X1)|r2_hidden(esk6_2(X2,k7_asympt_0(X1)),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_401, plain, (m1_subset_1(esk9_2(k7_asympt_0(X1),X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|r1_tarski(k7_asympt_0(X1),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_384, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X4,X3))|~r1_xxreal_0(X2,X4)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_234, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_298, negated_conjecture, (v3_valued_0(esk2_0))).#trainneg
cnf(i_0_299, negated_conjecture, (v3_valued_0(esk1_0))).#trainneg
cnf(i_0_114, plain, (r2_funct_2(X1,X2,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_223, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|r2_hidden(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_259, negated_conjecture, (v4_relat_1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_247, plain, (v1_xcmplx_0(np__1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_260, negated_conjecture, (v4_relat_1(esk1_0,k5_numbers))).#trainneg
cnf(i_0_270, negated_conjecture, (v5_relat_1(esk2_0,k1_numbers))).#trainneg
cnf(i_0_224, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|r2_hidden(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_230, plain, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|r2_hidden(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_389, plain, (v1_funct_2(esk9_2(k7_asympt_0(X1),X2),k5_numbers,k1_numbers)|r1_tarski(k7_asympt_0(X1),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_391, plain, (v1_funct_2(esk6_2(X1,k7_asympt_0(X2)),k5_numbers,k1_numbers)|X1=k7_asympt_0(X2)|r2_hidden(esk6_2(X1,k7_asympt_0(X2)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_99, plain, (v1_xboole_0(k6_numbers))).#trainneg
cnf(i_0_576, plain, (k4_xxreal_0(X1,X2)=X2|k4_xxreal_0(X1,X2)=X1|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_620, plain, (v1_xboole_0(X1)|m2_subset_1(X2,X1,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_55, plain, (m1_subset_1(k6_numbers,k1_numbers))).#trainneg
cnf(i_0_630, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_56, plain, (m1_subset_1(k6_numbers,k5_numbers))).#trainneg
cnf(i_0_226, plain, (r2_hidden(np__1,k1_numbers))).#trainneg
cnf(i_0_162, plain, (~r1_xxreal_0(np__1,k6_numbers))).#trainneg
cnf(i_0_231, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_57, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_243, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_271, negated_conjecture, (v5_relat_1(esk1_0,k1_numbers))).#trainneg
cnf(i_0_615, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_235, plain, (k4_xxreal_0(X1,X1)=X1|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_236, negated_conjecture, (v1_xreal_0(esk2_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_568, plain, (r2_hidden(X1,k7_asympt_0(X2))|~r2_hidden(X2,k6_asympt_0(X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~r2_hidden(X3,k6_asympt_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_471, plain, (X1=k6_asympt_0(X2)|~r2_hidden(esk6_2(X1,k6_asympt_0(X2)),X1)|k6_asympt_0(esk6_2(X1,k6_asympt_0(X2)))!=k6_asympt_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk6_2(X1,k6_asympt_0(X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(esk6_2(X1,k6_asympt_0(X2)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(esk6_2(X1,k6_asympt_0(X2)),k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(esk6_2(X1,k6_asympt_0(X2))))).#trainneg
cnf(i_0_501, plain, (r2_hidden(X1,k6_asympt_0(esk9_2(k7_asympt_0(X1),X2)))|r1_tarski(k7_asympt_0(X1),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_415, plain, (r1_xxreal_0(X1,esk15_2(X2,X3))|v1_limfunc1(X2)|~v1_xxreal_0(esk15_2(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_422, plain, (m1_subset_1(esk15_2(X1,X2),k5_numbers)|v1_limfunc1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_567, plain, (r2_hidden(esk6_2(X1,k7_asympt_0(X2)),k7_asympt_0(X3))|X1=k7_asympt_0(X2)|r2_hidden(esk6_2(X1,k7_asympt_0(X2)),X1)|~r2_hidden(X3,k6_asympt_0(X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_335, plain, (m1_subset_1(X1,k7_asympt_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_358, plain, (m1_subset_1(X1,k6_asympt_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_227, plain, (r2_hidden(np__1,k5_numbers))).#trainneg
cnf(i_0_237, negated_conjecture, (v1_xreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_582, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k5_numbers))).#trainneg
cnf(i_0_251, plain, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_244, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_618, plain, (v3_valued_0(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_616, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_238, plain, (v1_xreal_0(k2_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_245, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_414, plain, (v1_limfunc1(X2)|~v2_limfunc1(X1)|~m2_subset_1(esk15_2(X2,esk10_2(X1,X3)),k1_numbers,k5_numbers)|~r1_xxreal_0(X3,k1_seq_1(X1,esk15_2(X2,esk10_2(X1,X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m2_subset_1(esk10_2(X1,X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_545, plain, (k1_seq_1(k5_asympt_0(X1,X2),X3)=k2_square_1(k1_seq_1(X4,X3),k1_seq_1(X5,X3))|k5_asympt_0(X1,X2)!=k5_asympt_0(X4,X5)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X5,k5_numbers,k1_numbers)|~v1_funct_2(X4,k5_numbers,k1_numbers)|~v1_funct_2(k5_asympt_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v1_funct_1(k5_asympt_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_399, plain, (v1_funct_2(esk6_2(X1,k6_asympt_0(X2)),k5_numbers,k1_numbers)|X1=k6_asympt_0(X2)|r2_hidden(esk6_2(X1,k6_asympt_0(X2)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_624, plain, (k7_real_1(k1_seq_1(X1,X3),k1_seq_1(X2,X3))=k1_seq_1(k3_valued_1(k5_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_valued_1(k5_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(k3_valued_1(k5_numbers,k1_numbers,k1_numbers,X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_1(k3_valued_1(k5_numbers,k1_numbers,k1_numbers,X1,X2)))).#trainneg
cnf(i_0_527, plain, (k6_asympt_0(k1_valued_1(X1,X2))=k6_asympt_0(k5_asympt_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_671, plain, (v1_xreal_0(k6_numbers))).#trainneg
cnf(i_0_751, plain, (r1_tarski(k5_numbers,k1_numbers))).#trainneg
cnf(i_0_669, plain, (r2_hidden(k6_numbers,k1_numbers))).#trainneg
cnf(i_0_246, plain, (v1_xcmplx_0(k2_real_1(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_672, plain, (v1_xcmplx_0(k6_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_677, plain, (v1_xcmplx_0(k6_numbers)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_667, plain, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|m2_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_750, plain, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_252, plain, (X1=X2|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_253, plain, (m1_subset_1(esk9_2(X1,X2),X1)|r1_tarski(X1,X2))).#trainneg
cnf(i_0_759, plain, (k2_square_1(k1_seq_1(X1,X3),k1_seq_1(X2,X3))=k1_seq_1(k5_asympt_0(X1,X2),X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_651, plain, (k4_xxreal_0(X1,X2)=X1|k4_xxreal_0(X1,X2)=X2|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_664, plain, (m2_subset_1(k7_real_1(X1,X2),k1_numbers,k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_738, plain, (r2_hidden(k1_valued_1(X1,X2),k6_asympt_0(k5_asympt_0(X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_466, plain, (k1_asympt_0(X2,X1)=k1_valued_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_610, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_266, plain, (v1_xreal_0(k2_seq_2(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_674, plain, (r2_hidden(k6_numbers,k5_numbers))).#trainneg
cnf(i_0_752, plain, (v1_xcmplx_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_753, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_758, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_619, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_681, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_263, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k6_numbers))).#trainneg
cnf(i_0_517, plain, (r1_tarski(X1,k6_asympt_0(X2))|~r1_tarski(X1,k6_asympt_0(X3))|~r2_hidden(X3,k6_asympt_0(X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_403, plain, (m1_subset_1(esk6_2(X1,k7_asympt_0(X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|X1=k7_asympt_0(X2)|r2_hidden(esk6_2(X1,k7_asympt_0(X2)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_516, plain, (r2_hidden(X1,k6_asympt_0(X2))|~r2_hidden(X1,k6_asympt_0(X3))|~r2_hidden(X3,k6_asympt_0(X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_760, plain, (k1_seq_1(k5_asympt_0(X1,X2),X3)=k2_square_1(k1_seq_1(X4,X3),k1_seq_1(X5,X3))|k5_asympt_0(X1,X2)!=k5_asympt_0(X5,X4)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X5,k5_numbers,k1_numbers)|~v1_funct_2(X4,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_869, negated_conjecture, (k7_asympt_0(k1_valued_1(esk2_0,esk1_0))!=k7_asympt_0(k5_asympt_0(esk1_0,esk2_0)))).#trainneg
cnf(i_0_803, plain, (k5_numbers=k1_numbers|~r1_tarski(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_816, plain, (k5_numbers=k1_numbers|~m1_subset_1(k1_numbers,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_808, plain, (v1_xcmplx_0(k6_numbers))).#trainneg
cnf(i_0_817, plain, (X1=k1_numbers|~m1_subset_1(k1_numbers,k1_zfmisc_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_264, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k6_numbers))).#trainneg
cnf(i_0_804, plain, (r2_hidden(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_659, plain, (m2_subset_1(np__1,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_660, plain, (m2_subset_1(np__1,k5_numbers,k5_numbers))).#trainneg
cnf(i_0_884, plain, (v3_valued_0(k2_zfmisc_1(X1,k1_numbers)))).#trainneg
cnf(i_0_887, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_267, plain, (v1_xcmplx_0(k2_seq_2(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_272, plain, (r1_xxreal_0(X1,k4_xxreal_0(X2,X1))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_832, plain, (v1_xboole_0(X1)|m2_subset_1(esk9_2(X1,X2),X1,X1)|r1_tarski(X1,X2))).#trainneg
cnf(i_0_856, plain, (r1_tarski(X1,k7_asympt_0(X2))|~r1_tarski(X1,k7_asympt_0(X3))|~r2_hidden(X2,k6_asympt_0(X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_875, plain, (v1_funct_2(k1_valued_1(X2,X1),k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_931, plain, (k2_square_1(k1_seq_1(X2,X3),k1_seq_1(X1,X3))=k1_seq_1(k5_asympt_0(X1,X2),X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_873, plain, (v2_asympt_0(k1_valued_1(X2,X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_872, plain, (v1_funct_1(k1_valued_1(X2,X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_811, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_303, plain, (v1_relat_1(k3_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_670, plain, (m2_subset_1(k6_numbers,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_687, plain, (k4_xxreal_0(X1,X1)=X1|~v1_xxreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_675, plain, (m2_subset_1(k6_numbers,k5_numbers,k5_numbers))).#trainneg
cnf(i_0_633, plain, (k1_zfmisc_1(k1_numbers)=k6_numbers|r2_hidden(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_611, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_609, plain, (v1_xboole_0(k1_zfmisc_1(X1))|r2_hidden(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_275, plain, (k2_square_1(X1,X1)=X1|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_278, plain, (k7_nat_1(X1,X1)=X1|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_871, plain, (k1_valued_1(X2,X1)=k1_valued_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_685, plain, (k4_xxreal_0(X1,X1)=X1|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_405, plain, (v2_series_1(X1)|k1_funct_1(X1,esk19_1(X1))!=k6_numbers|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v7_ordinal1(esk19_1(X1)))).#trainneg
cnf(i_0_515, plain, (k6_asympt_0(X1)=k6_asympt_0(X2)|~r1_tarski(k6_asympt_0(X1),k6_asympt_0(X2))|~r2_hidden(X2,k6_asympt_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_383, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X3)|~r1_xxreal_0(X2,k6_numbers)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(k6_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_505, plain, (r1_tarski(X1,k7_asympt_0(X2))|~r2_hidden(X2,k6_asympt_0(esk9_2(X1,k7_asympt_0(X2))))|~m1_subset_1(esk9_2(X1,k7_asympt_0(X2)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(esk9_2(X1,k7_asympt_0(X2)))|~v1_funct_2(esk9_2(X1,k7_asympt_0(X2)),k5_numbers,k1_numbers)|~v1_funct_1(esk9_2(X1,k7_asympt_0(X2))))).#trainneg
cnf(i_0_733, plain, (r2_hidden(X1,k7_asympt_0(k1_asympt_0(X2,X3)))|~r2_hidden(k5_asympt_0(X2,X3),k6_asympt_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X3)|~v2_asympt_0(X2)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_957, plain, (r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_960, plain, (r1_xxreal_0(k6_numbers,k6_numbers))).#trainneg
cnf(i_0_280, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1049, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_961, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_284, plain, (k8_real_1(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_285, plain, (k8_real_1(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_982, plain, (r1_tarski(X1,k7_asympt_0(k1_asympt_0(X2,X3)))|~r1_tarski(X1,k7_asympt_0(k5_asympt_0(X2,X3)))|~m1_subset_1(k5_asympt_0(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(k5_asympt_0(X2,X3))|~v1_funct_2(k5_asympt_0(X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(k5_asympt_0(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X3)|~v2_asympt_0(X2)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_981, plain, (r1_tarski(X1,k7_asympt_0(k1_valued_1(X2,X3)))|~r1_tarski(X1,k7_asympt_0(k5_asympt_0(X2,X3)))|~m1_subset_1(k5_asympt_0(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(k5_asympt_0(X2,X3))|~v1_funct_2(k5_asympt_0(X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(k5_asympt_0(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X3)|~v2_asympt_0(X2)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_812, plain, (k1_numbers=X1|~r1_tarski(k1_numbers,X1)|~m1_subset_1(X1,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_306, plain, (v3_valued_0(k3_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_810, plain, (k1_zfmisc_1(k1_numbers)=k6_numbers|m2_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_287, plain, (v1_xreal_0(k7_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_962, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_629, negated_conjecture, (k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))=k6_numbers|r2_hidden(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_632, negated_conjecture, (k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))=k6_numbers|r2_hidden(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1011, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_577, negated_conjecture, (r1_tarski(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|~r1_tarski(X1,esk2_0))).#trainneg
cnf(i_0_288, plain, (v1_xcmplx_0(k7_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_290, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_974, plain, (r1_tarski(X1,k7_asympt_0(X2))|~r1_tarski(X1,k7_asympt_0(k1_valued_1(X3,X4)))|~r2_hidden(X2,k6_asympt_0(k5_asympt_0(X3,X4)))|~m1_subset_1(k1_valued_1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(k1_valued_1(X3,X4))|~v1_funct_2(k1_valued_1(X3,X4),k5_numbers,k1_numbers)|~v1_funct_1(k1_valued_1(X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X4)|~v2_asympt_0(X3)|~v1_funct_2(X4,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3))).#trainneg
cnf(i_0_1052, plain, (v2_series_1(X1)|k1_seq_1(X1,esk19_1(X1))!=k6_numbers|~v7_ordinal1(esk19_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1079, plain, (r2_hidden(X1,k7_asympt_0(k1_valued_1(X2,X3)))|~r2_hidden(k5_asympt_0(X2,X3),k6_asympt_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v2_asympt_0(X3)|~v2_asympt_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_975, plain, (r1_tarski(X1,k7_asympt_0(X2))|~r1_tarski(X1,k7_asympt_0(k1_asympt_0(X3,X4)))|~r2_hidden(X2,k6_asympt_0(k5_asympt_0(X3,X4)))|~m1_subset_1(k1_asympt_0(X3,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(k1_asympt_0(X3,X4))|~v1_funct_2(k1_asympt_0(X3,X4),k5_numbers,k1_numbers)|~v1_funct_1(k1_asympt_0(X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X4)|~v2_asympt_0(X3)|~v1_funct_2(X4,k5_numbers,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3))).#trainneg
cnf(i_0_940, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(k6_numbers,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_318, plain, (X1=k5_asympt_0(X2,X3)|k2_square_1(k1_seq_1(X3,esk3_3(X2,X3,X1)),k1_seq_1(X2,esk3_3(X2,X3,X1)))!=k1_seq_1(X1,esk3_3(X2,X3,X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X2,esk3_3(X2,X3,X1)),k1_numbers)|~m1_subset_1(k1_seq_1(X3,esk3_3(X2,X3,X1)),k1_numbers))).#trainneg
cnf(i_0_1015, plain, (k1_seq_1(k5_asympt_0(X1,X1),X2)=k1_seq_1(X1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers))).#trainneg
cnf(i_0_1127, plain, (k7_asympt_0(k1_valued_1(X1,X2))=k7_asympt_0(k5_asympt_0(X1,X2))|~r2_hidden(k5_asympt_0(X1,X2),k6_asympt_0(k1_valued_1(X1,X2)))|~m1_subset_1(k5_asympt_0(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(k5_asympt_0(X1,X2))|~v1_funct_2(k5_asympt_0(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k5_asympt_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1159, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r1_tarski(X1,esk2_0))).#trainneg
cnf(i_0_1157, negated_conjecture, (k2_zfmisc_1(k5_numbers,k1_numbers)=k1_numbers|~m1_subset_1(k2_zfmisc_1(k5_numbers,k1_numbers),k1_zfmisc_1(k5_numbers))|~r1_tarski(k1_numbers,esk2_0))).#trainneg
cnf(i_0_292, plain, (v1_xreal_0(k8_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_656, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|m2_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_657, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|m2_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1170, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_293, plain, (v1_xcmplx_0(k8_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_295, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1166, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1121, plain, (k7_asympt_0(X1)=k7_asympt_0(k1_asympt_0(X2,X3))|~r2_hidden(X1,k6_asympt_0(k5_asympt_0(X2,X3)))|~r2_hidden(k1_asympt_0(X2,X3),k6_asympt_0(X1))|~m1_subset_1(k1_asympt_0(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(k1_asympt_0(X2,X3))|~v1_funct_2(k1_asympt_0(X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(k1_asympt_0(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X3)|~v2_asympt_0(X2)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1120, plain, (k7_asympt_0(X1)=k7_asympt_0(k1_valued_1(X2,X3))|~r2_hidden(X1,k6_asympt_0(k5_asympt_0(X2,X3)))|~r2_hidden(k1_valued_1(X2,X3),k6_asympt_0(X1))|~m1_subset_1(k1_valued_1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(k1_valued_1(X2,X3))|~v1_funct_2(k1_valued_1(X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(k1_valued_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X3)|~v2_asympt_0(X2)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1102, plain, (k8_real_1(X2,X1)=k3_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1110, plain, (k7_asympt_0(X1)=k7_asympt_0(k1_asympt_0(X2,X3))|~r2_hidden(X1,k6_asympt_0(k1_asympt_0(X2,X3)))|~m1_subset_1(k1_asympt_0(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(k1_asympt_0(X2,X3))|~v1_funct_2(k1_asympt_0(X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(k1_asympt_0(X2,X3))|~r1_tarski(k7_asympt_0(X1),k7_asympt_0(k5_asympt_0(X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X3)|~v2_asympt_0(X2)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1131, plain, (k7_asympt_0(X1)=k7_asympt_0(k1_valued_1(X2,X3))|~r2_hidden(X1,k6_asympt_0(k1_valued_1(X2,X3)))|~m1_subset_1(k1_valued_1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(k1_valued_1(X2,X3))|~v1_funct_2(k1_valued_1(X2,X3),k5_numbers,k1_numbers)|~v1_funct_1(k1_valued_1(X2,X3))|~r1_tarski(k7_asympt_0(X1),k7_asympt_0(k5_asympt_0(X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X3)|~v2_asympt_0(X2)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1122, plain, (k7_asympt_0(esk9_2(k6_asympt_0(X1),X2))=k7_asympt_0(X1)|r1_tarski(k6_asympt_0(X1),X2)|~r2_hidden(X1,k6_asympt_0(esk9_2(k6_asympt_0(X1),X2)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1130, plain, (k7_asympt_0(X1)=k7_asympt_0(X2)|~r2_hidden(X2,k6_asympt_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|k6_asympt_0(X1)!=k6_asympt_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1113, plain, (k7_asympt_0(k1_asympt_0(X1,X2))=X3|~r1_tarski(k7_asympt_0(k1_asympt_0(X1,X2)),X3)|~r1_tarski(X3,k7_asympt_0(k5_asympt_0(X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_818, plain, (X1=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_622, negated_conjecture, (m1_subset_1(k1_seq_1(esk2_0,X1),k1_numbers))).#trainneg
cnf(i_0_623, negated_conjecture, (m1_subset_1(k1_seq_1(esk1_0,X1),k1_numbers))).#trainneg
cnf(i_0_852, plain, (k7_asympt_0(X1)=k7_asympt_0(X2)|~m1_subset_1(k7_asympt_0(X2),k1_zfmisc_1(k7_asympt_0(X1)))|~r2_hidden(X2,k6_asympt_0(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_296, plain, (m1_subset_1(k8_real_1(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1368, negated_conjecture, (k2_zfmisc_1(k5_numbers,k1_numbers)=esk2_0|~m1_subset_1(k2_zfmisc_1(k5_numbers,k1_numbers),k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_1369, negated_conjecture, (k2_zfmisc_1(k5_numbers,k1_numbers)=esk1_0|~m1_subset_1(k2_zfmisc_1(k5_numbers,k1_numbers),k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1288, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_578, negated_conjecture, (r1_tarski(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|~r1_tarski(X1,esk1_0))).#trainneg
cnf(i_0_300, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|r2_hidden(k3_series_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1323, plain, (k7_asympt_0(k1_valued_1(X1,X2))=k7_asympt_0(k5_asympt_0(X1,X2))|~r2_hidden(k5_asympt_0(X1,X2),k6_asympt_0(k5_asympt_0(X1,X2)))|~m1_subset_1(k1_valued_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v2_asympt_0(X2)|~v2_asympt_0(X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 129
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 218
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 211
# Processed clauses                    : 483
# ...of these trivial                  : 3
# ...subsumed                          : 86
# ...remaining for further processing  : 394
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 11
# Backward-rewritten                   : 9
# Generated clauses                    : 1227
# ...of the previous two non-trivial   : 1128
# Contextual simplify-reflections      : 110
# Paramodulations                      : 1220
# Factorizations                       : 2
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 371
#    Positive orientable unit clauses  : 55
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 310
# Current number of unprocessed clauses: 790
# ...number of literals in the above   : 7241
# Current number of archived formulas  : 0
# Current number of archived clauses   : 414
# Proof object given clauses           : 41
# Proof search given clauses           : 394
# Clause-clause subsumption calls (NU) : 88636
# Rec. Clause-clause subsumption calls : 4041
# Non-unit clause-clause subsumptions  : 203
# Unit Clause-clause subsumption calls : 2230
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 27
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 68136

# -------------------------------------------------
# User time                : 13.164 s
# System time              : 0.540 s
# Total time               : 13.704 s
# Maximum resident set size: 358284 pages
13.16user 0.64system 0:13.89elapsed 99%CPU (0avgtext+0avgdata 358996maxresident)k
0inputs+256outputs (0major+150126minor)pagefaults 0swaps
