# SZS status Theorem
# protokoll_atpstr_my_9f72062a403172caf26fbe878a32becf1403c276
# SinE strategy is GSinE(CountFormulas, hypos, 2.0,, 3, 20,1.0)
# Preprocessing time       : 0.029 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t8_conmetr, conjecture, ![X1]:((~(v2_struct_0(X1))&(v3_analmetr(X1)&l1_analmetr(X1)))=>(v11_aff_2(k3_analmetr(X1))=>v8_conmetr(X1))), file('/home/mptp/big1/gnn_pred_train_big_knn512_all.probs/preds__-1/t8_conmetr', t8_conmetr)).
fof(d8_conmetr, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_analmetr(X1)&l1_analmetr(X1)))=>(v8_conmetr(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>![X6]:(m1_subset_1(X6,u1_struct_0(X1))=>![X7]:(m1_subset_1(X7,u1_struct_0(X1))=>((r2_analoaf(X1,X2,X3,X4,X5)&(r2_analoaf(X1,X2,X3,X6,X7)&(r2_analoaf(X1,X2,X4,X3,X5)&r2_analoaf(X1,X2,X6,X3,X7))))=>(r5_analmetr(X1,X2,X3,X4)|(r5_analmetr(X1,X2,X3,X6)|r2_analoaf(X1,X4,X6,X5,X7)))))))))))), file('/home/mptp/big1/gnn_pred_train_big_knn512_all.probs/preds__-1/t8_conmetr', d8_conmetr)).
fof(t36_analmetr, axiom, ![X1]:((~(v2_struct_0(X1))&l1_analmetr(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>![X6]:(m1_subset_1(X6,u1_struct_0(k3_analmetr(X1)))=>![X7]:(m1_subset_1(X7,u1_struct_0(k3_analmetr(X1)))=>![X8]:(m1_subset_1(X8,u1_struct_0(k3_analmetr(X1)))=>![X9]:(m1_subset_1(X9,u1_struct_0(k3_analmetr(X1)))=>((X2=X6&(X3=X7&(X4=X8&X5=X9)))=>(r2_analoaf(X1,X2,X3,X4,X5)<=>r2_analoaf(k3_analmetr(X1),X6,X7,X8,X9)))))))))))), file('/home/mptp/big1/gnn_pred_train_big_knn512_all.probs/preds__-1/t8_conmetr', t36_analmetr)).
fof(t35_analmetr, axiom, ![X1]:((~(v2_struct_0(X1))&l1_analmetr(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))<=>m1_subset_1(X2,u1_struct_0(k3_analmetr(X1))))), file('/home/mptp/big1/gnn_pred_train_big_knn512_all.probs/preds__-1/t8_conmetr', t35_analmetr)).
fof(dt_k3_analmetr, axiom, ![X1]:((~(v2_struct_0(X1))&l1_analmetr(X1))=>(v1_analoaf(k3_analmetr(X1))&l1_analoaf(k3_analmetr(X1)))), file('/home/mptp/big1/gnn_pred_train_big_knn512_all.probs/preds__-1/t8_conmetr', dt_k3_analmetr)).
fof(fc4_analmetr, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_analmetr(X1)&l1_analmetr(X1)))=>(~(v7_struct_0(k3_analmetr(X1)))&(v1_analoaf(k3_analmetr(X1))&(v1_diraf(k3_analmetr(X1))&v2_diraf(k3_analmetr(X1)))))), file('/home/mptp/big1/gnn_pred_train_big_knn512_all.probs/preds__-1/t8_conmetr', fc4_analmetr)).
fof(t4_translac, axiom, ![X1]:((~(v7_struct_0(X1))&(v1_diraf(X1)&(v2_diraf(X1)&l1_analoaf(X1))))=>(v11_aff_2(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>![X6]:(m1_subset_1(X6,u1_struct_0(X1))=>![X7]:(m1_subset_1(X7,u1_struct_0(X1))=>((r2_analoaf(X1,X2,X3,X4,X6)&(r2_analoaf(X1,X2,X3,X5,X7)&(r2_analoaf(X1,X2,X4,X3,X6)&r2_analoaf(X1,X2,X5,X3,X7))))=>(r1_aff_1(X1,X2,X3,X4)|(r1_aff_1(X1,X2,X3,X5)|r2_analoaf(X1,X4,X5,X6,X7)))))))))))), file('/home/mptp/big1/gnn_pred_train_big_knn512_all.probs/preds__-1/t8_conmetr', t4_translac)).
fof(t40_analmetr, axiom, ![X1]:((~(v2_struct_0(X1))&(v2_analmetr(X1)&l1_analmetr(X1)))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(k3_analmetr(X1)))=>![X6]:(m1_subset_1(X6,u1_struct_0(k3_analmetr(X1)))=>![X7]:(m1_subset_1(X7,u1_struct_0(k3_analmetr(X1)))=>((X2=X5&(X3=X6&X4=X7))=>(r5_analmetr(X1,X2,X3,X4)<=>r1_aff_1(k3_analmetr(X1),X5,X6,X7)))))))))), file('/home/mptp/big1/gnn_pred_train_big_knn512_all.probs/preds__-1/t8_conmetr', t40_analmetr)).
fof(t37_analmetr, axiom, ![X1]:((~(v2_struct_0(X1))&(v3_analmetr(X1)&l1_analmetr(X1)))=>(~(v2_struct_0(X1))&(v2_analmetr(X1)&l1_analmetr(X1)))), file('/home/mptp/big1/gnn_pred_train_big_knn512_all.probs/preds__-1/t8_conmetr', t37_analmetr)).
fof(c_0_9, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(v3_analmetr(X1)&l1_analmetr(X1)))=>(v11_aff_2(k3_analmetr(X1))=>v8_conmetr(X1)))), inference(assume_negation,[status(cth)],[t8_conmetr])).
fof(c_0_10, plain, ![X14, X15, X16, X17, X18, X19, X20]:((~v8_conmetr(X14)|(~m1_subset_1(X15,u1_struct_0(X14))|(~m1_subset_1(X16,u1_struct_0(X14))|(~m1_subset_1(X17,u1_struct_0(X14))|(~m1_subset_1(X18,u1_struct_0(X14))|(~m1_subset_1(X19,u1_struct_0(X14))|(~m1_subset_1(X20,u1_struct_0(X14))|(~r2_analoaf(X14,X15,X16,X17,X18)|(~r2_analoaf(X14,X15,X16,X19,X20)|(~r2_analoaf(X14,X15,X17,X16,X18)|~r2_analoaf(X14,X15,X19,X16,X20)))|(r5_analmetr(X14,X15,X16,X17)|(r5_analmetr(X14,X15,X16,X19)|r2_analoaf(X14,X17,X19,X18,X20))))))))))|(v2_struct_0(X14)|(~v3_analmetr(X14)|~l1_analmetr(X14))))&((m1_subset_1(esk2_1(X14),u1_struct_0(X14))|v8_conmetr(X14)|(v2_struct_0(X14)|(~v3_analmetr(X14)|~l1_analmetr(X14))))&((m1_subset_1(esk3_1(X14),u1_struct_0(X14))|v8_conmetr(X14)|(v2_struct_0(X14)|(~v3_analmetr(X14)|~l1_analmetr(X14))))&((m1_subset_1(esk4_1(X14),u1_struct_0(X14))|v8_conmetr(X14)|(v2_struct_0(X14)|(~v3_analmetr(X14)|~l1_analmetr(X14))))&((m1_subset_1(esk5_1(X14),u1_struct_0(X14))|v8_conmetr(X14)|(v2_struct_0(X14)|(~v3_analmetr(X14)|~l1_analmetr(X14))))&((m1_subset_1(esk6_1(X14),u1_struct_0(X14))|v8_conmetr(X14)|(v2_struct_0(X14)|(~v3_analmetr(X14)|~l1_analmetr(X14))))&((m1_subset_1(esk7_1(X14),u1_struct_0(X14))|v8_conmetr(X14)|(v2_struct_0(X14)|(~v3_analmetr(X14)|~l1_analmetr(X14))))&(((r2_analoaf(X14,esk2_1(X14),esk3_1(X14),esk4_1(X14),esk5_1(X14))|v8_conmetr(X14)|(v2_struct_0(X14)|(~v3_analmetr(X14)|~l1_analmetr(X14))))&((r2_analoaf(X14,esk2_1(X14),esk3_1(X14),esk6_1(X14),esk7_1(X14))|v8_conmetr(X14)|(v2_struct_0(X14)|(~v3_analmetr(X14)|~l1_analmetr(X14))))&((r2_analoaf(X14,esk2_1(X14),esk4_1(X14),esk3_1(X14),esk5_1(X14))|v8_conmetr(X14)|(v2_struct_0(X14)|(~v3_analmetr(X14)|~l1_analmetr(X14))))&(r2_analoaf(X14,esk2_1(X14),esk6_1(X14),esk3_1(X14),esk7_1(X14))|v8_conmetr(X14)|(v2_struct_0(X14)|(~v3_analmetr(X14)|~l1_analmetr(X14)))))))&((~r5_analmetr(X14,esk2_1(X14),esk3_1(X14),esk4_1(X14))|v8_conmetr(X14)|(v2_struct_0(X14)|(~v3_analmetr(X14)|~l1_analmetr(X14))))&((~r5_analmetr(X14,esk2_1(X14),esk3_1(X14),esk6_1(X14))|v8_conmetr(X14)|(v2_struct_0(X14)|(~v3_analmetr(X14)|~l1_analmetr(X14))))&(~r2_analoaf(X14,esk4_1(X14),esk6_1(X14),esk5_1(X14),esk7_1(X14))|v8_conmetr(X14)|(v2_struct_0(X14)|(~v3_analmetr(X14)|~l1_analmetr(X14)))))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d8_conmetr])])])])])])).
fof(c_0_11, negated_conjecture, ((~v2_struct_0(esk1_0)&(v3_analmetr(esk1_0)&l1_analmetr(esk1_0)))&(v11_aff_2(k3_analmetr(esk1_0))&~v8_conmetr(esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])).
fof(c_0_12, plain, ![X89, X90, X91, X92, X93, X94, X95, X96, X97]:((~r2_analoaf(X89,X90,X91,X92,X93)|r2_analoaf(k3_analmetr(X89),X94,X95,X96,X97)|(X90!=X94|(X91!=X95|(X92!=X96|X93!=X97)))|~m1_subset_1(X97,u1_struct_0(k3_analmetr(X89)))|~m1_subset_1(X96,u1_struct_0(k3_analmetr(X89)))|~m1_subset_1(X95,u1_struct_0(k3_analmetr(X89)))|~m1_subset_1(X94,u1_struct_0(k3_analmetr(X89)))|~m1_subset_1(X93,u1_struct_0(X89))|~m1_subset_1(X92,u1_struct_0(X89))|~m1_subset_1(X91,u1_struct_0(X89))|~m1_subset_1(X90,u1_struct_0(X89))|(v2_struct_0(X89)|~l1_analmetr(X89)))&(~r2_analoaf(k3_analmetr(X89),X94,X95,X96,X97)|r2_analoaf(X89,X90,X91,X92,X93)|(X90!=X94|(X91!=X95|(X92!=X96|X93!=X97)))|~m1_subset_1(X97,u1_struct_0(k3_analmetr(X89)))|~m1_subset_1(X96,u1_struct_0(k3_analmetr(X89)))|~m1_subset_1(X95,u1_struct_0(k3_analmetr(X89)))|~m1_subset_1(X94,u1_struct_0(k3_analmetr(X89)))|~m1_subset_1(X93,u1_struct_0(X89))|~m1_subset_1(X92,u1_struct_0(X89))|~m1_subset_1(X91,u1_struct_0(X89))|~m1_subset_1(X90,u1_struct_0(X89))|(v2_struct_0(X89)|~l1_analmetr(X89)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t36_analmetr])])])])])).
fof(c_0_13, plain, ![X98, X99]:((~m1_subset_1(X99,u1_struct_0(X98))|m1_subset_1(X99,u1_struct_0(k3_analmetr(X98)))|(v2_struct_0(X98)|~l1_analmetr(X98)))&(~m1_subset_1(X99,u1_struct_0(k3_analmetr(X98)))|m1_subset_1(X99,u1_struct_0(X98))|(v2_struct_0(X98)|~l1_analmetr(X98)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t35_analmetr])])])])])).
cnf(c_0_14, plain, (m1_subset_1(esk7_1(X1),u1_struct_0(X1))|v8_conmetr(X1)|v2_struct_0(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_15, negated_conjecture, (l1_analmetr(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_16, negated_conjecture, (v3_analmetr(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_17, negated_conjecture, (~v8_conmetr(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_18, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_19, plain, (r2_analoaf(k3_analmetr(X1),X6,X7,X8,X9)|v2_struct_0(X1)|~r2_analoaf(X1,X2,X3,X4,X5)|X2!=X6|X3!=X7|X4!=X8|X5!=X9|~m1_subset_1(X9,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X8,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X7,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X6,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_20, plain, (m1_subset_1(X1,u1_struct_0(k3_analmetr(X2)))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_analmetr(X2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_21, plain, (v2_struct_0(X1)|v8_conmetr(X1)|m1_subset_1(esk7_1(X1),u1_struct_0(X1))|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(evalgc,[status(thm)],[c_0_14])).
cnf(c_0_22, negated_conjecture, (l1_analmetr(esk1_0)), inference(evalgc,[status(thm)],[c_0_15])).
cnf(c_0_23, negated_conjecture, (v3_analmetr(esk1_0)), inference(evalgc,[status(thm)],[c_0_16])).
cnf(c_0_24, negated_conjecture, (~v8_conmetr(esk1_0)), inference(evalgc,[status(thm)],[c_0_17])).
cnf(c_0_25, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_18])).
cnf(c_0_26, plain, (m1_subset_1(esk3_1(X1),u1_struct_0(X1))|v8_conmetr(X1)|v2_struct_0(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_27, plain, (m1_subset_1(esk6_1(X1),u1_struct_0(X1))|v8_conmetr(X1)|v2_struct_0(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_28, plain, (m1_subset_1(esk2_1(X1),u1_struct_0(X1))|v8_conmetr(X1)|v2_struct_0(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_29, plain, (v2_struct_0(X1)|r2_analoaf(k3_analmetr(X1),X6,X7,X8,X9)|X5!=X9|X4!=X8|X3!=X7|X2!=X6|~l1_analmetr(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r2_analoaf(X1,X2,X3,X4,X5)|~m1_subset_1(X9,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X8,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X7,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X6,u1_struct_0(k3_analmetr(X1)))), inference(evalgc,[status(thm)],[c_0_19])).
cnf(c_0_30, plain, (v2_struct_0(X2)|m1_subset_1(X1,u1_struct_0(k3_analmetr(X2)))|~l1_analmetr(X2)|~m1_subset_1(X1,u1_struct_0(X2))), inference(evalgc,[status(thm)],[c_0_20])).
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk7_1(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_21, c_0_22]), c_0_23])]), c_0_24]), c_0_25])])).
cnf(c_0_32, plain, (r2_analoaf(X1,esk2_1(X1),esk6_1(X1),esk3_1(X1),esk7_1(X1))|v8_conmetr(X1)|v2_struct_0(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_33, plain, (v2_struct_0(X1)|v8_conmetr(X1)|m1_subset_1(esk3_1(X1),u1_struct_0(X1))|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(evalgc,[status(thm)],[c_0_26])).
cnf(c_0_34, plain, (v2_struct_0(X1)|v8_conmetr(X1)|m1_subset_1(esk6_1(X1),u1_struct_0(X1))|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(evalgc,[status(thm)],[c_0_27])).
cnf(c_0_35, plain, (v2_struct_0(X1)|v8_conmetr(X1)|m1_subset_1(esk2_1(X1),u1_struct_0(X1))|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(evalgc,[status(thm)],[c_0_28])).
fof(c_0_36, plain, ![X79]:((v1_analoaf(k3_analmetr(X79))|(v2_struct_0(X79)|~l1_analmetr(X79)))&(l1_analoaf(k3_analmetr(X79))|(v2_struct_0(X79)|~l1_analmetr(X79)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_analmetr])])])])).
fof(c_0_37, plain, ![X78]:((~v7_struct_0(k3_analmetr(X78))|(v2_struct_0(X78)|(~v3_analmetr(X78)|~l1_analmetr(X78))))&((v1_analoaf(k3_analmetr(X78))|(v2_struct_0(X78)|(~v3_analmetr(X78)|~l1_analmetr(X78))))&((v1_diraf(k3_analmetr(X78))|(v2_struct_0(X78)|(~v3_analmetr(X78)|~l1_analmetr(X78))))&(v2_diraf(k3_analmetr(X78))|(v2_struct_0(X78)|(~v3_analmetr(X78)|~l1_analmetr(X78))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc4_analmetr])])])])).
cnf(c_0_38, plain, (m1_subset_1(esk5_1(X1),u1_struct_0(X1))|v8_conmetr(X1)|v2_struct_0(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
fof(c_0_39, plain, ![X27, X28, X29, X30, X31, X32, X33]:((~v11_aff_2(X27)|(~m1_subset_1(X28,u1_struct_0(X27))|(~m1_subset_1(X29,u1_struct_0(X27))|(~m1_subset_1(X30,u1_struct_0(X27))|(~m1_subset_1(X31,u1_struct_0(X27))|(~m1_subset_1(X32,u1_struct_0(X27))|(~m1_subset_1(X33,u1_struct_0(X27))|(~r2_analoaf(X27,X28,X29,X30,X32)|(~r2_analoaf(X27,X28,X29,X31,X33)|(~r2_analoaf(X27,X28,X30,X29,X32)|~r2_analoaf(X27,X28,X31,X29,X33)))|(r1_aff_1(X27,X28,X29,X30)|(r1_aff_1(X27,X28,X29,X31)|r2_analoaf(X27,X30,X31,X32,X33))))))))))|(v7_struct_0(X27)|(~v1_diraf(X27)|(~v2_diraf(X27)|~l1_analoaf(X27)))))&((m1_subset_1(esk8_1(X27),u1_struct_0(X27))|v11_aff_2(X27)|(v7_struct_0(X27)|(~v1_diraf(X27)|(~v2_diraf(X27)|~l1_analoaf(X27)))))&((m1_subset_1(esk9_1(X27),u1_struct_0(X27))|v11_aff_2(X27)|(v7_struct_0(X27)|(~v1_diraf(X27)|(~v2_diraf(X27)|~l1_analoaf(X27)))))&((m1_subset_1(esk10_1(X27),u1_struct_0(X27))|v11_aff_2(X27)|(v7_struct_0(X27)|(~v1_diraf(X27)|(~v2_diraf(X27)|~l1_analoaf(X27)))))&((m1_subset_1(esk11_1(X27),u1_struct_0(X27))|v11_aff_2(X27)|(v7_struct_0(X27)|(~v1_diraf(X27)|(~v2_diraf(X27)|~l1_analoaf(X27)))))&((m1_subset_1(esk12_1(X27),u1_struct_0(X27))|v11_aff_2(X27)|(v7_struct_0(X27)|(~v1_diraf(X27)|(~v2_diraf(X27)|~l1_analoaf(X27)))))&((m1_subset_1(esk13_1(X27),u1_struct_0(X27))|v11_aff_2(X27)|(v7_struct_0(X27)|(~v1_diraf(X27)|(~v2_diraf(X27)|~l1_analoaf(X27)))))&(((r2_analoaf(X27,esk8_1(X27),esk9_1(X27),esk10_1(X27),esk12_1(X27))|v11_aff_2(X27)|(v7_struct_0(X27)|(~v1_diraf(X27)|(~v2_diraf(X27)|~l1_analoaf(X27)))))&((r2_analoaf(X27,esk8_1(X27),esk9_1(X27),esk11_1(X27),esk13_1(X27))|v11_aff_2(X27)|(v7_struct_0(X27)|(~v1_diraf(X27)|(~v2_diraf(X27)|~l1_analoaf(X27)))))&((r2_analoaf(X27,esk8_1(X27),esk10_1(X27),esk9_1(X27),esk12_1(X27))|v11_aff_2(X27)|(v7_struct_0(X27)|(~v1_diraf(X27)|(~v2_diraf(X27)|~l1_analoaf(X27)))))&(r2_analoaf(X27,esk8_1(X27),esk11_1(X27),esk9_1(X27),esk13_1(X27))|v11_aff_2(X27)|(v7_struct_0(X27)|(~v1_diraf(X27)|(~v2_diraf(X27)|~l1_analoaf(X27))))))))&((~r1_aff_1(X27,esk8_1(X27),esk9_1(X27),esk10_1(X27))|v11_aff_2(X27)|(v7_struct_0(X27)|(~v1_diraf(X27)|(~v2_diraf(X27)|~l1_analoaf(X27)))))&((~r1_aff_1(X27,esk8_1(X27),esk9_1(X27),esk11_1(X27))|v11_aff_2(X27)|(v7_struct_0(X27)|(~v1_diraf(X27)|(~v2_diraf(X27)|~l1_analoaf(X27)))))&(~r2_analoaf(X27,esk10_1(X27),esk11_1(X27),esk12_1(X27),esk13_1(X27))|v11_aff_2(X27)|(v7_struct_0(X27)|(~v1_diraf(X27)|(~v2_diraf(X27)|~l1_analoaf(X27))))))))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_translac])])])])])])).
cnf(c_0_40, plain, (r2_analoaf(k3_analmetr(X1),X2,X3,X4,X5)|v2_struct_0(X1)|~r2_analoaf(X1,X2,X3,X4,X5)|~m1_subset_1(X5,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X4,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_analmetr(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_29])])])])])).
cnf(c_0_41, negated_conjecture, (m1_subset_1(esk7_1(esk1_0),u1_struct_0(k3_analmetr(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_31]), c_0_22])]), c_0_25])])).
cnf(c_0_42, plain, (v2_struct_0(X1)|v8_conmetr(X1)|r2_analoaf(X1,esk2_1(X1),esk6_1(X1),esk3_1(X1),esk7_1(X1))|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(evalgc,[status(thm)],[c_0_32])).
cnf(c_0_43, negated_conjecture, (m1_subset_1(esk3_1(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_22]), c_0_23])]), c_0_24]), c_0_25])])).
cnf(c_0_44, negated_conjecture, (m1_subset_1(esk6_1(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_22]), c_0_23])]), c_0_24]), c_0_25])])).
cnf(c_0_45, negated_conjecture, (m1_subset_1(esk2_1(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_22]), c_0_23])]), c_0_24]), c_0_25])])).
cnf(c_0_46, plain, (l1_analoaf(k3_analmetr(X1))|v2_struct_0(X1)|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
cnf(c_0_47, plain, (v2_diraf(k3_analmetr(X1))|v2_struct_0(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_48, plain, (v1_diraf(k3_analmetr(X1))|v2_struct_0(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_49, plain, (r2_analoaf(X1,esk2_1(X1),esk3_1(X1),esk6_1(X1),esk7_1(X1))|v8_conmetr(X1)|v2_struct_0(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_50, plain, (v2_struct_0(X1)|v8_conmetr(X1)|m1_subset_1(esk5_1(X1),u1_struct_0(X1))|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(evalgc,[status(thm)],[c_0_38])).
cnf(c_0_51, plain, (m1_subset_1(esk4_1(X1),u1_struct_0(X1))|v8_conmetr(X1)|v2_struct_0(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_52, plain, (r1_aff_1(X1,X2,X3,X4)|r1_aff_1(X1,X2,X3,X5)|r2_analoaf(X1,X4,X5,X6,X7)|v7_struct_0(X1)|~v11_aff_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X7,u1_struct_0(X1))|~r2_analoaf(X1,X2,X3,X4,X6)|~r2_analoaf(X1,X2,X3,X5,X7)|~r2_analoaf(X1,X2,X4,X3,X6)|~r2_analoaf(X1,X2,X5,X3,X7)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_53, negated_conjecture, (r2_analoaf(k3_analmetr(esk1_0),X1,X2,X3,esk7_1(esk1_0))|~r2_analoaf(esk1_0,X1,X2,X3,esk7_1(esk1_0))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_41]), c_0_31]), c_0_22])]), c_0_25])])).
cnf(c_0_54, negated_conjecture, (r2_analoaf(esk1_0,esk2_1(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_22]), c_0_23])]), c_0_24]), c_0_25])])).
cnf(c_0_55, negated_conjecture, (m1_subset_1(esk3_1(esk1_0),u1_struct_0(k3_analmetr(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_43]), c_0_22])]), c_0_25])])).
cnf(c_0_56, negated_conjecture, (m1_subset_1(esk6_1(esk1_0),u1_struct_0(k3_analmetr(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_44]), c_0_22])]), c_0_25])])).
cnf(c_0_57, negated_conjecture, (m1_subset_1(esk2_1(esk1_0),u1_struct_0(k3_analmetr(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_45]), c_0_22])]), c_0_25])])).
cnf(c_0_58, plain, (v2_struct_0(X1)|l1_analoaf(k3_analmetr(X1))|~l1_analmetr(X1)), inference(evalgc,[status(thm)],[c_0_46])).
cnf(c_0_59, plain, (v2_struct_0(X1)|v2_diraf(k3_analmetr(X1))|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(evalgc,[status(thm)],[c_0_47])).
cnf(c_0_60, plain, (v2_struct_0(X1)|v1_diraf(k3_analmetr(X1))|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(evalgc,[status(thm)],[c_0_48])).
cnf(c_0_61, negated_conjecture, (v11_aff_2(k3_analmetr(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_62, plain, (v2_struct_0(X1)|v8_conmetr(X1)|r2_analoaf(X1,esk2_1(X1),esk3_1(X1),esk6_1(X1),esk7_1(X1))|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(evalgc,[status(thm)],[c_0_49])).
cnf(c_0_63, negated_conjecture, (m1_subset_1(esk5_1(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_22]), c_0_23])]), c_0_24]), c_0_25])])).
cnf(c_0_64, plain, (r2_analoaf(X1,esk2_1(X1),esk4_1(X1),esk3_1(X1),esk5_1(X1))|v8_conmetr(X1)|v2_struct_0(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_65, plain, (v2_struct_0(X1)|v8_conmetr(X1)|m1_subset_1(esk4_1(X1),u1_struct_0(X1))|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(evalgc,[status(thm)],[c_0_51])).
cnf(c_0_66, plain, (r2_analoaf(X1,esk2_1(X1),esk3_1(X1),esk4_1(X1),esk5_1(X1))|v8_conmetr(X1)|v2_struct_0(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_67, plain, (v7_struct_0(X1)|r1_aff_1(X1,X2,X3,X5)|r1_aff_1(X1,X2,X3,X4)|r2_analoaf(X1,X4,X5,X6,X7)|~v11_aff_2(X1)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r2_analoaf(X1,X2,X5,X3,X7)|~r2_analoaf(X1,X2,X4,X3,X6)|~r2_analoaf(X1,X2,X3,X5,X7)|~r2_analoaf(X1,X2,X3,X4,X6)), inference(evalgc,[status(thm)],[c_0_52])).
cnf(c_0_68, negated_conjecture, (r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54]), c_0_55]), c_0_56]), c_0_57]), c_0_43]), c_0_44]), c_0_45])])])).
cnf(c_0_69, negated_conjecture, (l1_analoaf(k3_analmetr(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_22]), c_0_25])])).
cnf(c_0_70, negated_conjecture, (v2_diraf(k3_analmetr(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_22]), c_0_23])]), c_0_25])])).
cnf(c_0_71, negated_conjecture, (v1_diraf(k3_analmetr(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_22]), c_0_23])]), c_0_25])])).
cnf(c_0_72, negated_conjecture, (v11_aff_2(k3_analmetr(esk1_0))), inference(evalgc,[status(thm)],[c_0_61])).
cnf(c_0_73, negated_conjecture, (r2_analoaf(esk1_0,esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_22]), c_0_23])]), c_0_24]), c_0_25])])).
cnf(c_0_74, negated_conjecture, (m1_subset_1(esk5_1(esk1_0),u1_struct_0(k3_analmetr(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_63]), c_0_22])]), c_0_25])])).
cnf(c_0_75, plain, (v2_struct_0(X1)|v8_conmetr(X1)|r2_analoaf(X1,esk2_1(X1),esk4_1(X1),esk3_1(X1),esk5_1(X1))|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(evalgc,[status(thm)],[c_0_64])).
cnf(c_0_76, negated_conjecture, (m1_subset_1(esk4_1(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_22]), c_0_23])]), c_0_24]), c_0_25])])).
cnf(c_0_77, plain, (v2_struct_0(X1)|v8_conmetr(X1)|r2_analoaf(X1,esk2_1(X1),esk3_1(X1),esk4_1(X1),esk5_1(X1))|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(evalgc,[status(thm)],[c_0_66])).
cnf(c_0_78, plain, (r2_analoaf(X1,X6,X7,X8,X9)|v2_struct_0(X1)|~r2_analoaf(k3_analmetr(X1),X2,X3,X4,X5)|X6!=X2|X7!=X3|X8!=X4|X9!=X5|~m1_subset_1(X5,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X4,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X9,u1_struct_0(X1))|~m1_subset_1(X8,u1_struct_0(X1))|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_12])).
cnf(c_0_79, negated_conjecture, (r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0))|r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),X1)|v7_struct_0(k3_analmetr(esk1_0))|r2_analoaf(k3_analmetr(esk1_0),X1,esk6_1(esk1_0),X2,esk7_1(esk1_0))|~r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),X1,esk3_1(esk1_0),X2)|~r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0))|~r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),X1,X2)|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68]), c_0_69]), c_0_70]), c_0_71]), c_0_41]), c_0_56]), c_0_55]), c_0_57]), c_0_72])])])).
cnf(c_0_80, negated_conjecture, (r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_73]), c_0_56]), c_0_55]), c_0_57]), c_0_44]), c_0_43]), c_0_45])])])).
cnf(c_0_81, negated_conjecture, (r2_analoaf(k3_analmetr(esk1_0),X1,X2,X3,esk5_1(esk1_0))|~r2_analoaf(esk1_0,X1,X2,X3,esk5_1(esk1_0))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_74]), c_0_63]), c_0_22])]), c_0_25])])).
cnf(c_0_82, negated_conjecture, (r2_analoaf(esk1_0,esk2_1(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_22]), c_0_23])]), c_0_24]), c_0_25])])).
cnf(c_0_83, negated_conjecture, (m1_subset_1(esk4_1(esk1_0),u1_struct_0(k3_analmetr(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_76]), c_0_22])]), c_0_25])])).
cnf(c_0_84, negated_conjecture, (r2_analoaf(esk1_0,esk2_1(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0),esk5_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_22]), c_0_23])]), c_0_24]), c_0_25])])).
fof(c_0_85, plain, ![X82, X83, X84, X85, X86, X87, X88]:((~r5_analmetr(X82,X83,X84,X85)|r1_aff_1(k3_analmetr(X82),X86,X87,X88)|(X83!=X86|(X84!=X87|X85!=X88))|~m1_subset_1(X88,u1_struct_0(k3_analmetr(X82)))|~m1_subset_1(X87,u1_struct_0(k3_analmetr(X82)))|~m1_subset_1(X86,u1_struct_0(k3_analmetr(X82)))|~m1_subset_1(X85,u1_struct_0(X82))|~m1_subset_1(X84,u1_struct_0(X82))|~m1_subset_1(X83,u1_struct_0(X82))|(v2_struct_0(X82)|(~v2_analmetr(X82)|~l1_analmetr(X82))))&(~r1_aff_1(k3_analmetr(X82),X86,X87,X88)|r5_analmetr(X82,X83,X84,X85)|(X83!=X86|(X84!=X87|X85!=X88))|~m1_subset_1(X88,u1_struct_0(k3_analmetr(X82)))|~m1_subset_1(X87,u1_struct_0(k3_analmetr(X82)))|~m1_subset_1(X86,u1_struct_0(k3_analmetr(X82)))|~m1_subset_1(X85,u1_struct_0(X82))|~m1_subset_1(X84,u1_struct_0(X82))|~m1_subset_1(X83,u1_struct_0(X82))|(v2_struct_0(X82)|(~v2_analmetr(X82)|~l1_analmetr(X82))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t40_analmetr])])])])])).
cnf(c_0_86, plain, (v2_struct_0(X1)|r2_analoaf(X1,X6,X7,X8,X9)|X5!=X9|X4!=X8|X3!=X7|X2!=X6|~l1_analmetr(X1)|~m1_subset_1(X9,u1_struct_0(X1))|~m1_subset_1(X8,u1_struct_0(X1))|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X4,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(X1)))|~r2_analoaf(k3_analmetr(X1),X2,X3,X4,X5)), inference(evalgc,[status(thm)],[c_0_78])).
cnf(c_0_87, negated_conjecture, (r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0))|r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),X1)|v7_struct_0(k3_analmetr(esk1_0))|r2_analoaf(k3_analmetr(esk1_0),X1,esk6_1(esk1_0),X2,esk7_1(esk1_0))|~r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),X1,esk3_1(esk1_0),X2)|~r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),X1,X2)|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_80])])).
cnf(c_0_88, negated_conjecture, (r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_55]), c_0_83]), c_0_57]), c_0_43]), c_0_76]), c_0_45])])])).
cnf(c_0_89, negated_conjecture, (r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0),esk5_1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_84]), c_0_83]), c_0_55]), c_0_57]), c_0_76]), c_0_43]), c_0_45])])])).
fof(c_0_90, plain, ![X107]:((~v2_struct_0(X107)|(v2_struct_0(X107)|(~v3_analmetr(X107)|~l1_analmetr(X107))))&((v2_analmetr(X107)|(v2_struct_0(X107)|(~v3_analmetr(X107)|~l1_analmetr(X107))))&(l1_analmetr(X107)|(v2_struct_0(X107)|(~v3_analmetr(X107)|~l1_analmetr(X107)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t37_analmetr])])])])).
cnf(c_0_91, plain, (r5_analmetr(X1,X5,X6,X7)|v2_struct_0(X1)|~r1_aff_1(k3_analmetr(X1),X2,X3,X4)|X5!=X2|X6!=X3|X7!=X4|~m1_subset_1(X4,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~v2_analmetr(X1)|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_92, plain, (v8_conmetr(X1)|v2_struct_0(X1)|~r2_analoaf(X1,esk4_1(X1),esk6_1(X1),esk5_1(X1),esk7_1(X1))|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_93, plain, (r2_analoaf(X1,X2,X3,X4,X5)|v2_struct_0(X1)|~r2_analoaf(k3_analmetr(X1),X2,X3,X4,X5)|~m1_subset_1(X5,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X4,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_analmetr(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_86])])])])])).
cnf(c_0_94, negated_conjecture, (r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0))|r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))|r2_analoaf(k3_analmetr(esk1_0),esk4_1(esk1_0),esk6_1(esk1_0),esk5_1(esk1_0),esk7_1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_89]), c_0_74]), c_0_83])])])).
cnf(c_0_95, plain, (v2_analmetr(X1)|v2_struct_0(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_96, plain, (v2_struct_0(X1)|r5_analmetr(X1,X5,X6,X7)|X4!=X7|X3!=X6|X2!=X5|~l1_analmetr(X1)|~v2_analmetr(X1)|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(X1)))|~r1_aff_1(k3_analmetr(X1),X2,X3,X4)), inference(evalgc,[status(thm)],[c_0_91])).
cnf(c_0_97, plain, (v2_struct_0(X1)|v8_conmetr(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)|~r2_analoaf(X1,esk4_1(X1),esk6_1(X1),esk5_1(X1),esk7_1(X1))), inference(evalgc,[status(thm)],[c_0_92])).
cnf(c_0_98, negated_conjecture, (r2_analoaf(esk1_0,esk4_1(esk1_0),esk6_1(esk1_0),esk5_1(esk1_0),esk7_1(esk1_0))|r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0))|r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_41]), c_0_74]), c_0_56]), c_0_83]), c_0_31]), c_0_63]), c_0_44]), c_0_76]), c_0_22])]), c_0_25])])).
cnf(c_0_99, plain, (v2_struct_0(X1)|v2_analmetr(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(evalgc,[status(thm)],[c_0_95])).
cnf(c_0_100, plain, (r5_analmetr(X1,X2,X3,X4)|v2_struct_0(X1)|~v2_analmetr(X1)|~r1_aff_1(k3_analmetr(X1),X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_analmetr(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_96])])])])).
cnf(c_0_101, negated_conjecture, (r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0))|r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_22]), c_0_23])]), c_0_24]), c_0_25])])).
cnf(c_0_102, negated_conjecture, (v2_analmetr(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_22]), c_0_23])]), c_0_25])])).
cnf(c_0_103, plain, (v8_conmetr(X1)|v2_struct_0(X1)|~r5_analmetr(X1,esk2_1(X1),esk3_1(X1),esk6_1(X1))|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_104, negated_conjecture, (r5_analmetr(esk1_0,esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0))|r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101]), c_0_102]), c_0_56]), c_0_55]), c_0_57]), c_0_44]), c_0_43]), c_0_45]), c_0_22])]), c_0_25])])).
cnf(c_0_105, plain, (v8_conmetr(X1)|v2_struct_0(X1)|~r5_analmetr(X1,esk2_1(X1),esk3_1(X1),esk4_1(X1))|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_106, plain, (v2_struct_0(X1)|v8_conmetr(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)|~r5_analmetr(X1,esk2_1(X1),esk3_1(X1),esk6_1(X1))), inference(evalgc,[status(thm)],[c_0_103])).
cnf(c_0_107, negated_conjecture, (r5_analmetr(esk1_0,esk2_1(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))|r5_analmetr(esk1_0,esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_104]), c_0_102]), c_0_83]), c_0_55]), c_0_57]), c_0_76]), c_0_43]), c_0_45]), c_0_22])]), c_0_25])])).
cnf(c_0_108, plain, (v2_struct_0(X1)|~v7_struct_0(k3_analmetr(X1))|~v3_analmetr(X1)|~l1_analmetr(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_109, plain, (v2_struct_0(X1)|v8_conmetr(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)|~r5_analmetr(X1,esk2_1(X1),esk3_1(X1),esk4_1(X1))), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_110, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|r5_analmetr(esk1_0,esk2_1(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_22]), c_0_23])]), c_0_24]), c_0_25])])).
cnf(c_0_111, plain, (v2_struct_0(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)|~v7_struct_0(k3_analmetr(X1))), inference(evalgc,[status(thm)],[c_0_108])).
cnf(c_0_112, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110]), c_0_22]), c_0_23])]), c_0_24]), c_0_25])])).
cnf(c_0_113, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_22]), c_0_23])]), c_0_25]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 114
# Proof object clause steps            : 95
# Proof object formula steps           : 19
# Proof object conjectures             : 49
# Proof object clause conjectures      : 46
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 28
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 35
# Proof object simplifying inferences  : 180
# Training examples: 65 positive, 629 negative
# Training: Positive examples begin
cnf(c_0_25, negated_conjecture, (~v2_struct_0(esk1_0))).# trainpos
cnf(c_0_24, negated_conjecture, (~v8_conmetr(esk1_0))).# trainpos
cnf(c_0_23, negated_conjecture, (v3_analmetr(esk1_0))).# trainpos
cnf(c_0_22, negated_conjecture, (l1_analmetr(esk1_0))).# trainpos
cnf(c_0_72, negated_conjecture, (v11_aff_2(k3_analmetr(esk1_0)))).# trainpos
cnf(c_0_58, plain, (v2_struct_0(X1)|l1_analoaf(k3_analmetr(X1))|~l1_analmetr(X1))).# trainpos
cnf(c_0_99, plain, (v2_struct_0(X1)|v2_analmetr(X1)|~v3_analmetr(X1)|~l1_analmetr(X1))).# trainpos
cnf(c_0_111, plain, (v2_struct_0(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)|~v7_struct_0(k3_analmetr(X1)))).# trainpos
cnf(c_0_60, plain, (v2_struct_0(X1)|v1_diraf(k3_analmetr(X1))|~v3_analmetr(X1)|~l1_analmetr(X1))).# trainpos
cnf(c_0_59, plain, (v2_struct_0(X1)|v2_diraf(k3_analmetr(X1))|~v3_analmetr(X1)|~l1_analmetr(X1))).# trainpos
cnf(c_0_30, plain, (v2_struct_0(X2)|m1_subset_1(X1,u1_struct_0(k3_analmetr(X2)))|~l1_analmetr(X2)|~m1_subset_1(X1,u1_struct_0(X2)))).# trainpos
cnf(c_0_109, plain, (v2_struct_0(X1)|v8_conmetr(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)|~r5_analmetr(X1,esk2_1(X1),esk3_1(X1),esk4_1(X1)))).# trainpos
cnf(c_0_106, plain, (v2_struct_0(X1)|v8_conmetr(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)|~r5_analmetr(X1,esk2_1(X1),esk3_1(X1),esk6_1(X1)))).# trainpos
cnf(c_0_35, plain, (v2_struct_0(X1)|v8_conmetr(X1)|m1_subset_1(esk2_1(X1),u1_struct_0(X1))|~v3_analmetr(X1)|~l1_analmetr(X1))).# trainpos
cnf(c_0_33, plain, (v2_struct_0(X1)|v8_conmetr(X1)|m1_subset_1(esk3_1(X1),u1_struct_0(X1))|~v3_analmetr(X1)|~l1_analmetr(X1))).# trainpos
cnf(c_0_65, plain, (v2_struct_0(X1)|v8_conmetr(X1)|m1_subset_1(esk4_1(X1),u1_struct_0(X1))|~v3_analmetr(X1)|~l1_analmetr(X1))).# trainpos
cnf(c_0_50, plain, (v2_struct_0(X1)|v8_conmetr(X1)|m1_subset_1(esk5_1(X1),u1_struct_0(X1))|~v3_analmetr(X1)|~l1_analmetr(X1))).# trainpos
cnf(c_0_34, plain, (v2_struct_0(X1)|v8_conmetr(X1)|m1_subset_1(esk6_1(X1),u1_struct_0(X1))|~v3_analmetr(X1)|~l1_analmetr(X1))).# trainpos
cnf(c_0_21, plain, (v2_struct_0(X1)|v8_conmetr(X1)|m1_subset_1(esk7_1(X1),u1_struct_0(X1))|~v3_analmetr(X1)|~l1_analmetr(X1))).# trainpos
cnf(c_0_97, plain, (v2_struct_0(X1)|v8_conmetr(X1)|~v3_analmetr(X1)|~l1_analmetr(X1)|~r2_analoaf(X1,esk4_1(X1),esk6_1(X1),esk5_1(X1),esk7_1(X1)))).# trainpos
cnf(c_0_77, plain, (v2_struct_0(X1)|v8_conmetr(X1)|r2_analoaf(X1,esk2_1(X1),esk3_1(X1),esk4_1(X1),esk5_1(X1))|~v3_analmetr(X1)|~l1_analmetr(X1))).# trainpos
cnf(c_0_62, plain, (v2_struct_0(X1)|v8_conmetr(X1)|r2_analoaf(X1,esk2_1(X1),esk3_1(X1),esk6_1(X1),esk7_1(X1))|~v3_analmetr(X1)|~l1_analmetr(X1))).# trainpos
cnf(c_0_75, plain, (v2_struct_0(X1)|v8_conmetr(X1)|r2_analoaf(X1,esk2_1(X1),esk4_1(X1),esk3_1(X1),esk5_1(X1))|~v3_analmetr(X1)|~l1_analmetr(X1))).# trainpos
cnf(c_0_42, plain, (v2_struct_0(X1)|v8_conmetr(X1)|r2_analoaf(X1,esk2_1(X1),esk6_1(X1),esk3_1(X1),esk7_1(X1))|~v3_analmetr(X1)|~l1_analmetr(X1))).# trainpos
cnf(c_0_96, plain, (v2_struct_0(X1)|r5_analmetr(X1,X5,X6,X7)|X4!=X7|X3!=X6|X2!=X5|~l1_analmetr(X1)|~v2_analmetr(X1)|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(X1)))|~r1_aff_1(k3_analmetr(X1),X2,X3,X4))).# trainpos
cnf(c_0_86, plain, (v2_struct_0(X1)|r2_analoaf(X1,X6,X7,X8,X9)|X5!=X9|X4!=X8|X3!=X7|X2!=X6|~l1_analmetr(X1)|~m1_subset_1(X9,u1_struct_0(X1))|~m1_subset_1(X8,u1_struct_0(X1))|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X4,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(X1)))|~r2_analoaf(k3_analmetr(X1),X2,X3,X4,X5))).# trainpos
cnf(c_0_29, plain, (v2_struct_0(X1)|r2_analoaf(k3_analmetr(X1),X6,X7,X8,X9)|X5!=X9|X4!=X8|X3!=X7|X2!=X6|~l1_analmetr(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r2_analoaf(X1,X2,X3,X4,X5)|~m1_subset_1(X9,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X8,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X7,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X6,u1_struct_0(k3_analmetr(X1))))).# trainpos
cnf(c_0_67, plain, (v7_struct_0(X1)|r1_aff_1(X1,X2,X3,X5)|r1_aff_1(X1,X2,X3,X4)|r2_analoaf(X1,X4,X5,X6,X7)|~v11_aff_2(X1)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r2_analoaf(X1,X2,X5,X3,X7)|~r2_analoaf(X1,X2,X4,X3,X6)|~r2_analoaf(X1,X2,X3,X5,X7)|~r2_analoaf(X1,X2,X3,X4,X6))).# trainpos
cnf(c_0_102, negated_conjecture, (v2_analmetr(esk1_0))).# trainpos
cnf(c_0_69, negated_conjecture, (l1_analoaf(k3_analmetr(esk1_0)))).# trainpos
cnf(c_0_71, negated_conjecture, (v1_diraf(k3_analmetr(esk1_0)))).# trainpos
cnf(c_0_70, negated_conjecture, (v2_diraf(k3_analmetr(esk1_0)))).# trainpos
cnf(c_0_45, negated_conjecture, (m1_subset_1(esk2_1(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_43, negated_conjecture, (m1_subset_1(esk3_1(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_76, negated_conjecture, (m1_subset_1(esk4_1(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_63, negated_conjecture, (m1_subset_1(esk5_1(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_44, negated_conjecture, (m1_subset_1(esk6_1(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_31, negated_conjecture, (m1_subset_1(esk7_1(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_57, negated_conjecture, (m1_subset_1(esk2_1(esk1_0),u1_struct_0(k3_analmetr(esk1_0))))).# trainpos
cnf(c_0_55, negated_conjecture, (m1_subset_1(esk3_1(esk1_0),u1_struct_0(k3_analmetr(esk1_0))))).# trainpos
cnf(c_0_83, negated_conjecture, (m1_subset_1(esk4_1(esk1_0),u1_struct_0(k3_analmetr(esk1_0))))).# trainpos
cnf(c_0_74, negated_conjecture, (m1_subset_1(esk5_1(esk1_0),u1_struct_0(k3_analmetr(esk1_0))))).# trainpos
cnf(c_0_56, negated_conjecture, (m1_subset_1(esk6_1(esk1_0),u1_struct_0(k3_analmetr(esk1_0))))).# trainpos
cnf(c_0_41, negated_conjecture, (m1_subset_1(esk7_1(esk1_0),u1_struct_0(k3_analmetr(esk1_0))))).# trainpos
cnf(c_0_84, negated_conjecture, (r2_analoaf(esk1_0,esk2_1(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0),esk5_1(esk1_0)))).# trainpos
cnf(c_0_73, negated_conjecture, (r2_analoaf(esk1_0,esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0)))).# trainpos
cnf(c_0_82, negated_conjecture, (r2_analoaf(esk1_0,esk2_1(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0)))).# trainpos
cnf(c_0_54, negated_conjecture, (r2_analoaf(esk1_0,esk2_1(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0)))).# trainpos
cnf(c_0_100, plain, (r5_analmetr(X1,X2,X3,X4)|v2_struct_0(X1)|~v2_analmetr(X1)|~r1_aff_1(k3_analmetr(X1),X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_analmetr(X1))).# trainpos
cnf(c_0_93, plain, (r2_analoaf(X1,X2,X3,X4,X5)|v2_struct_0(X1)|~r2_analoaf(k3_analmetr(X1),X2,X3,X4,X5)|~m1_subset_1(X5,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X4,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_analmetr(X1))).# trainpos
cnf(c_0_40, plain, (r2_analoaf(k3_analmetr(X1),X2,X3,X4,X5)|v2_struct_0(X1)|~r2_analoaf(X1,X2,X3,X4,X5)|~m1_subset_1(X5,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X4,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_analmetr(X1))).# trainpos
cnf(c_0_53, negated_conjecture, (r2_analoaf(k3_analmetr(esk1_0),X1,X2,X3,esk7_1(esk1_0))|~r2_analoaf(esk1_0,X1,X2,X3,esk7_1(esk1_0))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_68, negated_conjecture, (r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0)))).# trainpos
cnf(c_0_80, negated_conjecture, (r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0)))).# trainpos
cnf(c_0_81, negated_conjecture, (r2_analoaf(k3_analmetr(esk1_0),X1,X2,X3,esk5_1(esk1_0))|~r2_analoaf(esk1_0,X1,X2,X3,esk5_1(esk1_0))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_88, negated_conjecture, (r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0)))).# trainpos
cnf(c_0_89, negated_conjecture, (r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0),esk5_1(esk1_0)))).# trainpos
cnf(c_0_79, negated_conjecture, (r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0))|r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),X1)|v7_struct_0(k3_analmetr(esk1_0))|r2_analoaf(k3_analmetr(esk1_0),X1,esk6_1(esk1_0),X2,esk7_1(esk1_0))|~r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),X1,esk3_1(esk1_0),X2)|~r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0))|~r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),X1,X2)|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).# trainpos
cnf(c_0_94, negated_conjecture, (r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0))|r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))|r2_analoaf(k3_analmetr(esk1_0),esk4_1(esk1_0),esk6_1(esk1_0),esk5_1(esk1_0),esk7_1(esk1_0)))).# trainpos
cnf(c_0_98, negated_conjecture, (r2_analoaf(esk1_0,esk4_1(esk1_0),esk6_1(esk1_0),esk5_1(esk1_0),esk7_1(esk1_0))|r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0))|r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).# trainpos
cnf(c_0_101, negated_conjecture, (r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0))|r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).# trainpos
cnf(c_0_104, negated_conjecture, (r5_analmetr(esk1_0,esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0))|r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).# trainpos
cnf(c_0_107, negated_conjecture, (r5_analmetr(esk1_0,esk2_1(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))|r5_analmetr(esk1_0,esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0)))).# trainpos
cnf(c_0_110, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|r5_analmetr(esk1_0,esk2_1(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0)))).# trainpos
cnf(c_0_112, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0)))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_73, plain, (l1_analoaf(X1)|~l1_analmetr(X1))).#trainneg
cnf(i_0_56, plain, (v2_struct_0(X1)|~l1_analmetr(X1)|~v2_struct_0(k3_analmetr(X1)))).#trainneg
cnf(i_0_54, plain, (v2_struct_0(X1)|v1_analoaf(k3_analmetr(X1))|~l1_analmetr(X1))).#trainneg
cnf(i_0_59, plain, (v2_struct_0(X1)|~l1_analmetr(X1)|~v2_analmetr(X1)|~v7_struct_0(k3_analmetr(X1)))).#trainneg
cnf(i_0_57, plain, (v2_struct_0(X1)|v1_diraf(k3_analmetr(X1))|~l1_analmetr(X1)|~v2_analmetr(X1))).#trainneg
cnf(i_0_64, plain, (v2_struct_0(X2)|m1_subset_1(X1,u1_struct_0(X2))|~l1_analmetr(X2)|~m1_subset_1(X1,u1_struct_0(k3_analmetr(X2))))).#trainneg
cnf(i_0_68, plain, (v2_struct_0(X2)|m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_analmetr(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_analmetr(X2)))))).#trainneg
cnf(i_0_22, plain, (v11_aff_2(X1)|v7_struct_0(X1)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)|~r1_aff_1(X1,esk8_1(X1),esk9_1(X1),esk10_1(X1)))).#trainneg
cnf(i_0_21, plain, (v11_aff_2(X1)|v7_struct_0(X1)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)|~r1_aff_1(X1,esk8_1(X1),esk9_1(X1),esk11_1(X1)))).#trainneg
cnf(i_0_32, plain, (v11_aff_2(X1)|v7_struct_0(X1)|m1_subset_1(esk8_1(X1),u1_struct_0(X1))|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1))).#trainneg
cnf(i_0_31, plain, (v11_aff_2(X1)|v7_struct_0(X1)|m1_subset_1(esk9_1(X1),u1_struct_0(X1))|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1))).#trainneg
cnf(i_0_71, plain, (v2_struct_0(X1)|v3_analmetr(X1)|v7_struct_0(k3_analmetr(X1))|~l1_analmetr(X1)|~v2_analmetr(X1)|~v1_diraf(k3_analmetr(X1))|~v2_diraf(k3_analmetr(X1))|~l1_analoaf(k3_analmetr(X1)))).#trainneg
cnf(i_0_69, plain, (v2_struct_0(X2)|m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_analmetr(X2))))|~l1_analmetr(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_30, plain, (v11_aff_2(X1)|v7_struct_0(X1)|m1_subset_1(esk10_1(X1),u1_struct_0(X1))|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1))).#trainneg
cnf(i_0_29, plain, (v11_aff_2(X1)|v7_struct_0(X1)|m1_subset_1(esk11_1(X1),u1_struct_0(X1))|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1))).#trainneg
cnf(i_0_28, plain, (v11_aff_2(X1)|v7_struct_0(X1)|m1_subset_1(esk12_1(X1),u1_struct_0(X1))|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1))).#trainneg
cnf(i_0_27, plain, (v11_aff_2(X1)|v7_struct_0(X1)|m1_subset_1(esk13_1(X1),u1_struct_0(X1))|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1))).#trainneg
cnf(i_0_26, plain, (v11_aff_2(X1)|v7_struct_0(X1)|r2_analoaf(X1,esk8_1(X1),esk9_1(X1),esk10_1(X1),esk12_1(X1))|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1))).#trainneg
cnf(i_0_25, plain, (v11_aff_2(X1)|v7_struct_0(X1)|r2_analoaf(X1,esk8_1(X1),esk9_1(X1),esk11_1(X1),esk13_1(X1))|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1))).#trainneg
cnf(i_0_24, plain, (v11_aff_2(X1)|v7_struct_0(X1)|r2_analoaf(X1,esk8_1(X1),esk10_1(X1),esk9_1(X1),esk12_1(X1))|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1))).#trainneg
cnf(i_0_20, plain, (v11_aff_2(X1)|v7_struct_0(X1)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)|~r2_analoaf(X1,esk10_1(X1),esk11_1(X1),esk12_1(X1),esk13_1(X1)))).#trainneg
cnf(i_0_42, plain, (v7_struct_0(X1)|v1_funct_1(esk14_3(X1,X2,X3))|~v11_aff_2(X1)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_38, plain, (v7_struct_0(X1)|v7_transgeo(esk14_3(X1,X2,X3),X1)|~v11_aff_2(X1)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_48, plain, (X2=X3|v7_struct_0(X1)|v1_funct_1(esk15_3(X1,X2,X3))|~v11_aff_2(X1)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_66, plain, (v2_struct_0(X2)|v4_analmetr(X3,X2)|X1!=X3|~l1_analmetr(X2)|~v2_analmetr(X2)|~v1_aff_1(X1,k3_analmetr(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_analmetr(X2)))))).#trainneg
cnf(i_0_67, plain, (v2_struct_0(X2)|v1_aff_1(X3,k3_analmetr(X2))|X1!=X3|~l1_analmetr(X2)|~v2_analmetr(X2)|~v4_analmetr(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k3_analmetr(X2)))))).#trainneg
cnf(i_0_44, plain, (X2=X3|v7_struct_0(X1)|v7_transgeo(esk15_3(X1,X2,X3),X1)|~v11_aff_2(X1)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_23, plain, (v11_aff_2(X1)|v7_struct_0(X1)|r2_analoaf(X1,esk8_1(X1),esk11_1(X1),esk9_1(X1),esk13_1(X1))|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1))).#trainneg
cnf(i_0_41, plain, (v7_struct_0(X1)|v1_funct_2(esk14_3(X1,X2,X3),u1_struct_0(X1),u1_struct_0(X1))|~v11_aff_2(X1)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_40, plain, (v7_struct_0(X1)|v3_funct_2(esk14_3(X1,X2,X3),u1_struct_0(X1),u1_struct_0(X1))|~v11_aff_2(X1)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_47, plain, (X2=X3|v7_struct_0(X1)|v1_funct_2(esk15_3(X1,X2,X3),u1_struct_0(X1),u1_struct_0(X1))|~v11_aff_2(X1)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_46, plain, (X2=X3|v7_struct_0(X1)|v3_funct_2(esk15_3(X1,X2,X3),u1_struct_0(X1),u1_struct_0(X1))|~v11_aff_2(X1)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_39, plain, (v7_struct_0(X1)|m1_subset_1(esk14_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~v11_aff_2(X1)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_37, plain, (k3_funct_2(u1_struct_0(X1),u1_struct_0(X1),esk14_3(X1,X2,X3),X2)=X3|v7_struct_0(X1)|~v11_aff_2(X1)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_45, plain, (X2=X3|v7_struct_0(X1)|m1_subset_1(esk15_3(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~v11_aff_2(X1)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_43, plain, (X2=X3|k3_funct_2(u1_struct_0(X1),u1_struct_0(X1),esk15_3(X1,X2,X3),X2)=X3|v7_struct_0(X1)|~v11_aff_2(X1)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_61, plain, (v2_struct_0(X1)|r1_aff_1(k3_analmetr(X1),X5,X6,X7)|X4!=X7|X3!=X6|X2!=X5|~l1_analmetr(X1)|~v2_analmetr(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r5_analmetr(X1,X2,X3,X4)|~m1_subset_1(X7,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X6,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X5,u1_struct_0(k3_analmetr(X1))))).#trainneg
cnf(i_0_19, plain, (v2_struct_0(X1)|r5_analmetr(X1,X2,X3,X6)|r5_analmetr(X1,X2,X3,X4)|r2_analoaf(X1,X4,X6,X5,X7)|~v3_analmetr(X1)|~l1_analmetr(X1)|~v8_conmetr(X1)|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r2_analoaf(X1,X2,X6,X3,X7)|~r2_analoaf(X1,X2,X4,X3,X5)|~r2_analoaf(X1,X2,X3,X6,X7)|~r2_analoaf(X1,X2,X3,X4,X5))).#trainneg
cnf(i_0_34, plain, (X9=X10|X2=X3|v7_struct_0(X1)|r1_aff_1(X1,X2,X3,X7)|r1_aff_1(X1,X2,X3,X5)|~v11_aff_2(X1)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)|~m1_subset_1(X10,u1_struct_0(X1))|~m1_subset_1(X9,u1_struct_0(X1))|~m1_subset_1(X8,u1_struct_0(X1))|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r1_aff_1(X1,X2,X3,X4)|~r2_analoaf(X1,X7,X8,X4,X10)|~r2_analoaf(X1,X7,X4,X8,X10)|~r2_analoaf(X1,X5,X6,X4,X9)|~r2_analoaf(X1,X5,X4,X6,X9)|~r2_analoaf(X1,X2,X7,X3,X8)|~r2_analoaf(X1,X2,X5,X3,X6)|~r2_analoaf(X1,X2,X3,X7,X8)|~r2_analoaf(X1,X2,X3,X5,X6))).#trainneg
cnf(i_0_35, plain, (X4=X10|X2=X3|v7_struct_0(X1)|r1_aff_1(X1,X2,X3,X7)|r1_aff_1(X1,X2,X3,X5)|~v11_aff_2(X1)|~v1_diraf(X1)|~v2_diraf(X1)|~l1_analoaf(X1)|~m1_subset_1(X10,u1_struct_0(X1))|~m1_subset_1(X9,u1_struct_0(X1))|~m1_subset_1(X8,u1_struct_0(X1))|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~r1_aff_1(X1,X2,X3,X4)|~r2_analoaf(X1,X7,X10,X8,X9)|~r2_analoaf(X1,X7,X8,X10,X9)|~r2_analoaf(X1,X5,X6,X4,X9)|~r2_analoaf(X1,X5,X4,X6,X9)|~r2_analoaf(X1,X2,X7,X3,X8)|~r2_analoaf(X1,X2,X5,X3,X6)|~r2_analoaf(X1,X2,X3,X7,X8)|~r2_analoaf(X1,X2,X3,X5,X6))).#trainneg
cnf(i_0_79, negated_conjecture, (l1_analoaf(esk1_0))).#trainneg
cnf(i_0_80, negated_conjecture, (v1_analoaf(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_218, negated_conjecture, (m1_subset_1(esk2_1(esk1_0),u1_struct_0(k3_analmetr(k3_analmetr(esk1_0))))|v2_struct_0(k3_analmetr(esk1_0))|~l1_analmetr(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_111, negated_conjecture, (v7_struct_0(esk1_0)|m1_subset_1(esk8_1(esk1_0),u1_struct_0(esk1_0))|v11_aff_2(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0))).#trainneg
cnf(i_0_232, negated_conjecture, (m1_subset_1(esk3_1(esk1_0),u1_struct_0(k3_analmetr(k3_analmetr(esk1_0))))|v2_struct_0(k3_analmetr(esk1_0))|~l1_analmetr(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_246, negated_conjecture, (m1_subset_1(esk4_1(esk1_0),u1_struct_0(k3_analmetr(k3_analmetr(esk1_0))))|v2_struct_0(k3_analmetr(esk1_0))|~l1_analmetr(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_260, negated_conjecture, (m1_subset_1(esk5_1(esk1_0),u1_struct_0(k3_analmetr(k3_analmetr(esk1_0))))|v2_struct_0(k3_analmetr(esk1_0))|~l1_analmetr(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_274, negated_conjecture, (m1_subset_1(esk6_1(esk1_0),u1_struct_0(k3_analmetr(k3_analmetr(esk1_0))))|v2_struct_0(k3_analmetr(esk1_0))|~l1_analmetr(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_288, negated_conjecture, (m1_subset_1(esk7_1(esk1_0),u1_struct_0(k3_analmetr(k3_analmetr(esk1_0))))|v2_struct_0(k3_analmetr(esk1_0))|~l1_analmetr(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_133, negated_conjecture, (v1_funct_1(esk14_3(esk1_0,X1,esk2_1(esk1_0)))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_146, negated_conjecture, (v1_funct_1(esk14_3(esk1_0,X1,esk3_1(esk1_0)))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_159, negated_conjecture, (v1_funct_1(esk14_3(esk1_0,X1,esk4_1(esk1_0)))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_172, negated_conjecture, (v1_funct_1(esk14_3(esk1_0,X1,esk5_1(esk1_0)))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_185, negated_conjecture, (v1_funct_1(esk14_3(esk1_0,X1,esk6_1(esk1_0)))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_112, negated_conjecture, (v7_struct_0(esk1_0)|m1_subset_1(esk9_1(esk1_0),u1_struct_0(esk1_0))|v11_aff_2(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0))).#trainneg
cnf(i_0_113, negated_conjecture, (v7_struct_0(esk1_0)|m1_subset_1(esk10_1(esk1_0),u1_struct_0(esk1_0))|v11_aff_2(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0))).#trainneg
cnf(i_0_114, negated_conjecture, (v7_struct_0(esk1_0)|m1_subset_1(esk11_1(esk1_0),u1_struct_0(esk1_0))|v11_aff_2(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0))).#trainneg
cnf(i_0_115, negated_conjecture, (v7_struct_0(esk1_0)|m1_subset_1(esk12_1(esk1_0),u1_struct_0(esk1_0))|v11_aff_2(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0))).#trainneg
cnf(i_0_116, negated_conjecture, (v7_struct_0(esk1_0)|m1_subset_1(esk13_1(esk1_0),u1_struct_0(esk1_0))|v11_aff_2(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0))).#trainneg
cnf(i_0_211, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),X1,esk2_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_311, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk2_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_312, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk2_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_313, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk2_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_314, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk2_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_225, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),X1,esk3_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_239, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),X1,esk4_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_253, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),X1,esk5_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_267, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),X1,esk6_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_281, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),X1,esk7_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_315, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk2_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_316, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk2_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_317, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk3_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_318, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_319, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk3_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_66, plain, (v4_analmetr(X1,X2)|v2_struct_0(X2)|~v2_analmetr(X2)|~v1_aff_1(X1,k3_analmetr(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_analmetr(X2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_analmetr(X2))).#trainneg
cnf(i_0_320, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_198, negated_conjecture, (v1_funct_1(esk14_3(esk1_0,X1,esk7_1(esk1_0)))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_134, negated_conjecture, (v7_transgeo(esk14_3(esk1_0,X1,esk2_1(esk1_0)),esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_147, negated_conjecture, (v7_transgeo(esk14_3(esk1_0,X1,esk3_1(esk1_0)),esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_160, negated_conjecture, (v7_transgeo(esk14_3(esk1_0,X1,esk4_1(esk1_0)),esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_173, negated_conjecture, (v7_transgeo(esk14_3(esk1_0,X1,esk5_1(esk1_0)),esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_186, negated_conjecture, (v7_transgeo(esk14_3(esk1_0,X1,esk6_1(esk1_0)),esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_199, negated_conjecture, (v7_transgeo(esk14_3(esk1_0,X1,esk7_1(esk1_0)),esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_321, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk3_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_322, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_323, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk4_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_324, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk4_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_325, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk4_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_326, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk4_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_327, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_328, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_329, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk5_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_330, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk5_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_331, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk5_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_332, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk5_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_333, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_334, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk5_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_335, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk6_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_336, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk6_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_337, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk6_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_338, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk6_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_339, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_340, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_341, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk7_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_342, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_343, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk7_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_344, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk7_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_345, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_346, negated_conjecture, (v1_funct_1(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk7_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_67, plain, (v1_aff_1(X1,k3_analmetr(X2))|v2_struct_0(X2)|~v4_analmetr(X1,X2)|~v2_analmetr(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k3_analmetr(X2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_analmetr(X2))).#trainneg
cnf(i_0_212, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),X1,esk2_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_226, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),X1,esk3_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_240, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),X1,esk4_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_254, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),X1,esk5_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_268, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),X1,esk6_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_282, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),X1,esk7_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_141, negated_conjecture, (X1=esk2_1(esk1_0)|v1_funct_1(esk15_3(esk1_0,X1,esk2_1(esk1_0)))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_154, negated_conjecture, (X1=esk3_1(esk1_0)|v1_funct_1(esk15_3(esk1_0,X1,esk3_1(esk1_0)))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_167, negated_conjecture, (X1=esk4_1(esk1_0)|v1_funct_1(esk15_3(esk1_0,X1,esk4_1(esk1_0)))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_180, negated_conjecture, (X1=esk5_1(esk1_0)|v1_funct_1(esk15_3(esk1_0,X1,esk5_1(esk1_0)))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_193, negated_conjecture, (X1=esk6_1(esk1_0)|v1_funct_1(esk15_3(esk1_0,X1,esk6_1(esk1_0)))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_347, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk2_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_348, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk2_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_349, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk2_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_350, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk2_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_351, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk2_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_352, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk2_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_353, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk3_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_354, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_355, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk3_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_356, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_357, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk3_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_358, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_359, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk4_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_360, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk4_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_361, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk4_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_362, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk4_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_363, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_364, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_365, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk5_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_366, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk5_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_367, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk5_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_368, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk5_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_369, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_117, negated_conjecture, (v7_struct_0(esk1_0)|r2_analoaf(esk1_0,esk8_1(esk1_0),esk9_1(esk1_0),esk10_1(esk1_0),esk12_1(esk1_0))|v11_aff_2(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0))).#trainneg
cnf(i_0_206, negated_conjecture, (X1=esk7_1(esk1_0)|v1_funct_1(esk15_3(esk1_0,X1,esk7_1(esk1_0)))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_219, negated_conjecture, (X1=esk2_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),X1,esk2_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_233, negated_conjecture, (X1=esk3_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),X1,esk3_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_247, negated_conjecture, (X1=esk4_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),X1,esk4_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_261, negated_conjecture, (X1=esk5_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),X1,esk5_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_275, negated_conjecture, (X1=esk6_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),X1,esk6_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_289, negated_conjecture, (X1=esk7_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),X1,esk7_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_135, negated_conjecture, (v1_funct_2(esk14_3(esk1_0,X1,esk2_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_136, negated_conjecture, (v3_funct_2(esk14_3(esk1_0,X1,esk2_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_148, negated_conjecture, (v1_funct_2(esk14_3(esk1_0,X1,esk3_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_370, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk5_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_371, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk6_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_372, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk6_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_373, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk6_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_374, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk6_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_375, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_376, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_377, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk7_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_378, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_379, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk7_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_380, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk7_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_381, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_382, negated_conjecture, (v7_transgeo(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk7_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_118, negated_conjecture, (v7_struct_0(esk1_0)|r2_analoaf(esk1_0,esk8_1(esk1_0),esk9_1(esk1_0),esk11_1(esk1_0),esk13_1(esk1_0))|v11_aff_2(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0))).#trainneg
cnf(i_0_119, negated_conjecture, (v7_struct_0(esk1_0)|r2_analoaf(esk1_0,esk8_1(esk1_0),esk10_1(esk1_0),esk9_1(esk1_0),esk12_1(esk1_0))|v11_aff_2(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0))).#trainneg
cnf(i_0_120, negated_conjecture, (v7_struct_0(esk1_0)|r2_analoaf(esk1_0,esk8_1(esk1_0),esk11_1(esk1_0),esk9_1(esk1_0),esk13_1(esk1_0))|v11_aff_2(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0))).#trainneg
cnf(i_0_383, negated_conjecture, (esk7_1(esk1_0)=esk2_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk2_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_384, negated_conjecture, (esk6_1(esk1_0)=esk2_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk2_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_385, negated_conjecture, (esk5_1(esk1_0)=esk2_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk2_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_386, negated_conjecture, (esk4_1(esk1_0)=esk2_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk2_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_387, negated_conjecture, (esk3_1(esk1_0)=esk2_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk2_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_388, negated_conjecture, (esk7_1(esk1_0)=esk3_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk3_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_389, negated_conjecture, (esk6_1(esk1_0)=esk3_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_61, plain, (r1_aff_1(k3_analmetr(X1),X2,X3,X4)|v2_struct_0(X1)|~v2_analmetr(X1)|~r5_analmetr(X1,X2,X3,X4)|~m1_subset_1(X4,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_analmetr(X1))).#trainneg
cnf(i_0_390, negated_conjecture, (esk5_1(esk1_0)=esk3_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk3_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_149, negated_conjecture, (v3_funct_2(esk14_3(esk1_0,X1,esk3_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_161, negated_conjecture, (v1_funct_2(esk14_3(esk1_0,X1,esk4_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_162, negated_conjecture, (v3_funct_2(esk14_3(esk1_0,X1,esk4_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_174, negated_conjecture, (v1_funct_2(esk14_3(esk1_0,X1,esk5_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_175, negated_conjecture, (v3_funct_2(esk14_3(esk1_0,X1,esk5_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_187, negated_conjecture, (v1_funct_2(esk14_3(esk1_0,X1,esk6_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_188, negated_conjecture, (v3_funct_2(esk14_3(esk1_0,X1,esk6_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_200, negated_conjecture, (v1_funct_2(esk14_3(esk1_0,X1,esk7_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_201, negated_conjecture, (v3_funct_2(esk14_3(esk1_0,X1,esk7_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_142, negated_conjecture, (X1=esk2_1(esk1_0)|v7_transgeo(esk15_3(esk1_0,X1,esk2_1(esk1_0)),esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_391, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_392, negated_conjecture, (esk3_1(esk1_0)=esk2_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_393, negated_conjecture, (esk7_1(esk1_0)=esk4_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk4_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_394, negated_conjecture, (esk6_1(esk1_0)=esk4_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk4_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_395, negated_conjecture, (esk5_1(esk1_0)=esk4_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk4_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_396, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_397, negated_conjecture, (esk4_1(esk1_0)=esk2_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_398, negated_conjecture, (esk7_1(esk1_0)=esk5_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk5_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_399, negated_conjecture, (esk6_1(esk1_0)=esk5_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk5_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_400, negated_conjecture, (esk5_1(esk1_0)=esk4_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk5_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_401, negated_conjecture, (esk5_1(esk1_0)=esk3_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_402, negated_conjecture, (esk5_1(esk1_0)=esk2_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk5_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_403, negated_conjecture, (esk7_1(esk1_0)=esk6_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk6_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_404, negated_conjecture, (esk6_1(esk1_0)=esk5_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk6_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_405, negated_conjecture, (esk6_1(esk1_0)=esk4_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk6_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_406, negated_conjecture, (esk6_1(esk1_0)=esk3_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_407, negated_conjecture, (esk6_1(esk1_0)=esk2_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_408, negated_conjecture, (esk7_1(esk1_0)=esk6_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_409, negated_conjecture, (esk7_1(esk1_0)=esk5_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk7_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_410, negated_conjecture, (esk7_1(esk1_0)=esk4_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk7_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_411, negated_conjecture, (esk7_1(esk1_0)=esk3_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_412, negated_conjecture, (esk7_1(esk1_0)=esk2_1(esk1_0)|v1_funct_1(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk7_1(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_220, negated_conjecture, (X1=esk2_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),X1,esk2_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_155, negated_conjecture, (X1=esk3_1(esk1_0)|v7_transgeo(esk15_3(esk1_0,X1,esk3_1(esk1_0)),esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_168, negated_conjecture, (X1=esk4_1(esk1_0)|v7_transgeo(esk15_3(esk1_0,X1,esk4_1(esk1_0)),esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_181, negated_conjecture, (X1=esk5_1(esk1_0)|v7_transgeo(esk15_3(esk1_0,X1,esk5_1(esk1_0)),esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_194, negated_conjecture, (X1=esk6_1(esk1_0)|v7_transgeo(esk15_3(esk1_0,X1,esk6_1(esk1_0)),esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_207, negated_conjecture, (X1=esk7_1(esk1_0)|v7_transgeo(esk15_3(esk1_0,X1,esk7_1(esk1_0)),esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_234, negated_conjecture, (X1=esk3_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),X1,esk3_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_248, negated_conjecture, (X1=esk4_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),X1,esk4_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_262, negated_conjecture, (X1=esk5_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),X1,esk5_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_276, negated_conjecture, (X1=esk6_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),X1,esk6_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_290, negated_conjecture, (X1=esk7_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),X1,esk7_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_137, negated_conjecture, (v7_struct_0(esk1_0)|m1_subset_1(esk14_3(esk1_0,X1,esk2_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_419, negated_conjecture, (esk7_1(esk1_0)=esk2_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk2_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_420, negated_conjecture, (esk6_1(esk1_0)=esk2_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk2_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_421, negated_conjecture, (esk5_1(esk1_0)=esk2_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk2_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_422, negated_conjecture, (esk4_1(esk1_0)=esk2_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk2_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_423, negated_conjecture, (esk3_1(esk1_0)=esk2_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk2_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_424, negated_conjecture, (esk7_1(esk1_0)=esk3_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk3_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_425, negated_conjecture, (esk6_1(esk1_0)=esk3_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_426, negated_conjecture, (esk5_1(esk1_0)=esk3_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk3_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_427, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_428, negated_conjecture, (esk3_1(esk1_0)=esk2_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_429, negated_conjecture, (esk7_1(esk1_0)=esk4_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk4_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_430, negated_conjecture, (esk6_1(esk1_0)=esk4_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk4_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_431, negated_conjecture, (esk5_1(esk1_0)=esk4_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk4_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_213, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),X1,esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_214, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),X1,esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_227, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),X1,esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_228, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),X1,esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_241, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),X1,esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_242, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),X1,esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_255, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),X1,esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_256, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),X1,esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_269, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),X1,esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_270, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),X1,esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_283, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),X1,esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_150, negated_conjecture, (v7_struct_0(esk1_0)|m1_subset_1(esk14_3(esk1_0,X1,esk3_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_163, negated_conjecture, (v7_struct_0(esk1_0)|m1_subset_1(esk14_3(esk1_0,X1,esk4_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_176, negated_conjecture, (v7_struct_0(esk1_0)|m1_subset_1(esk14_3(esk1_0,X1,esk5_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_189, negated_conjecture, (v7_struct_0(esk1_0)|m1_subset_1(esk14_3(esk1_0,X1,esk6_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_202, negated_conjecture, (v7_struct_0(esk1_0)|m1_subset_1(esk14_3(esk1_0,X1,esk7_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_284, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),X1,esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_138, negated_conjecture, (k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),esk14_3(esk1_0,X1,esk2_1(esk1_0)),X1)=esk2_1(esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_151, negated_conjecture, (k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),esk14_3(esk1_0,X1,esk3_1(esk1_0)),X1)=esk3_1(esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_164, negated_conjecture, (k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),esk14_3(esk1_0,X1,esk4_1(esk1_0)),X1)=esk4_1(esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_177, negated_conjecture, (k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),esk14_3(esk1_0,X1,esk5_1(esk1_0)),X1)=esk5_1(esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_432, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_433, negated_conjecture, (esk4_1(esk1_0)=esk2_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_434, negated_conjecture, (esk7_1(esk1_0)=esk5_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk5_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_435, negated_conjecture, (esk6_1(esk1_0)=esk5_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk5_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_436, negated_conjecture, (esk5_1(esk1_0)=esk4_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk5_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_437, negated_conjecture, (esk5_1(esk1_0)=esk3_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_438, negated_conjecture, (esk5_1(esk1_0)=esk2_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk5_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_439, negated_conjecture, (esk7_1(esk1_0)=esk6_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk6_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_440, negated_conjecture, (esk6_1(esk1_0)=esk5_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk6_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_441, negated_conjecture, (esk6_1(esk1_0)=esk4_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk6_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_442, negated_conjecture, (esk6_1(esk1_0)=esk3_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_443, negated_conjecture, (esk6_1(esk1_0)=esk2_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_444, negated_conjecture, (esk7_1(esk1_0)=esk6_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_445, negated_conjecture, (esk7_1(esk1_0)=esk5_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk7_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_446, negated_conjecture, (esk7_1(esk1_0)=esk4_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk7_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_447, negated_conjecture, (esk7_1(esk1_0)=esk3_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_448, negated_conjecture, (esk7_1(esk1_0)=esk2_1(esk1_0)|v7_transgeo(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk7_1(esk1_0)),k3_analmetr(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_216, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),X1,esk2_1(esk1_0)),X1)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_230, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),X1,esk3_1(esk1_0)),X1)=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_244, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),X1,esk4_1(esk1_0)),X1)=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_258, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),X1,esk5_1(esk1_0)),X1)=esk5_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_272, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),X1,esk6_1(esk1_0)),X1)=esk6_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_286, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),X1,esk7_1(esk1_0)),X1)=esk7_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_139, negated_conjecture, (k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),esk15_3(esk1_0,X1,esk2_1(esk1_0)),X1)=esk2_1(esk1_0)|X1=esk2_1(esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_190, negated_conjecture, (k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),esk14_3(esk1_0,X1,esk6_1(esk1_0)),X1)=esk6_1(esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_203, negated_conjecture, (k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),esk14_3(esk1_0,X1,esk7_1(esk1_0)),X1)=esk7_1(esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_413, negated_conjecture, (r1_aff_1(k3_analmetr(esk1_0),X1,X2,esk7_1(esk1_0))|~r5_analmetr(esk1_0,X1,X2,esk7_1(esk1_0))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_414, negated_conjecture, (r1_aff_1(k3_analmetr(esk1_0),X1,X2,esk6_1(esk1_0))|~r5_analmetr(esk1_0,X1,X2,esk6_1(esk1_0))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_415, negated_conjecture, (r1_aff_1(k3_analmetr(esk1_0),X1,X2,esk5_1(esk1_0))|~r5_analmetr(esk1_0,X1,X2,esk5_1(esk1_0))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_416, negated_conjecture, (r1_aff_1(k3_analmetr(esk1_0),X1,X2,esk4_1(esk1_0))|~r5_analmetr(esk1_0,X1,X2,esk4_1(esk1_0))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_417, negated_conjecture, (r1_aff_1(k3_analmetr(esk1_0),X1,X2,esk3_1(esk1_0))|~r5_analmetr(esk1_0,X1,X2,esk3_1(esk1_0))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_418, negated_conjecture, (r1_aff_1(k3_analmetr(esk1_0),X1,X2,esk2_1(esk1_0))|~r5_analmetr(esk1_0,X1,X2,esk2_1(esk1_0))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_143, negated_conjecture, (X1=esk2_1(esk1_0)|v1_funct_2(esk15_3(esk1_0,X1,esk2_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_144, negated_conjecture, (X1=esk2_1(esk1_0)|v3_funct_2(esk15_3(esk1_0,X1,esk2_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_156, negated_conjecture, (X1=esk3_1(esk1_0)|v1_funct_2(esk15_3(esk1_0,X1,esk3_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_449, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_450, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_451, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_452, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_453, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_454, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_455, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_456, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_457, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_458, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_459, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_460, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_461, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_215, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),X1,esk2_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_221, negated_conjecture, (X1=esk2_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),X1,esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_222, negated_conjecture, (X1=esk2_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),X1,esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_229, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),X1,esk3_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_235, negated_conjecture, (X1=esk3_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),X1,esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_236, negated_conjecture, (X1=esk3_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),X1,esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_243, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),X1,esk4_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_249, negated_conjecture, (X1=esk4_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),X1,esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_250, negated_conjecture, (X1=esk4_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),X1,esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_257, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),X1,esk5_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_145, negated_conjecture, (X1=esk2_1(esk1_0)|v7_struct_0(esk1_0)|m1_subset_1(esk15_3(esk1_0,X1,esk2_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_157, negated_conjecture, (X1=esk3_1(esk1_0)|v3_funct_2(esk15_3(esk1_0,X1,esk3_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_169, negated_conjecture, (X1=esk4_1(esk1_0)|v1_funct_2(esk15_3(esk1_0,X1,esk4_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_170, negated_conjecture, (X1=esk4_1(esk1_0)|v3_funct_2(esk15_3(esk1_0,X1,esk4_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_182, negated_conjecture, (X1=esk5_1(esk1_0)|v1_funct_2(esk15_3(esk1_0,X1,esk5_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_183, negated_conjecture, (X1=esk5_1(esk1_0)|v3_funct_2(esk15_3(esk1_0,X1,esk5_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_195, negated_conjecture, (X1=esk6_1(esk1_0)|v1_funct_2(esk15_3(esk1_0,X1,esk6_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_196, negated_conjecture, (X1=esk6_1(esk1_0)|v3_funct_2(esk15_3(esk1_0,X1,esk6_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_208, negated_conjecture, (X1=esk7_1(esk1_0)|v1_funct_2(esk15_3(esk1_0,X1,esk7_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_209, negated_conjecture, (X1=esk7_1(esk1_0)|v3_funct_2(esk15_3(esk1_0,X1,esk7_1(esk1_0)),u1_struct_0(esk1_0),u1_struct_0(esk1_0))|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_271, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),X1,esk6_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_285, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),X1,esk7_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_462, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_463, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_464, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_465, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_466, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_467, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_468, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_469, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_470, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_471, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_472, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_473, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_474, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_263, negated_conjecture, (X1=esk5_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),X1,esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_264, negated_conjecture, (X1=esk5_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),X1,esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_277, negated_conjecture, (X1=esk6_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),X1,esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_278, negated_conjecture, (X1=esk6_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),X1,esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_291, negated_conjecture, (X1=esk7_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),X1,esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_292, negated_conjecture, (X1=esk7_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),X1,esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_475, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_476, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_477, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_478, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_152, negated_conjecture, (k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),esk15_3(esk1_0,X1,esk3_1(esk1_0)),X1)=esk3_1(esk1_0)|X1=esk3_1(esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_479, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_158, negated_conjecture, (X1=esk3_1(esk1_0)|v7_struct_0(esk1_0)|m1_subset_1(esk15_3(esk1_0,X1,esk3_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_171, negated_conjecture, (X1=esk4_1(esk1_0)|v7_struct_0(esk1_0)|m1_subset_1(esk15_3(esk1_0,X1,esk4_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_184, negated_conjecture, (X1=esk5_1(esk1_0)|v7_struct_0(esk1_0)|m1_subset_1(esk15_3(esk1_0,X1,esk5_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_197, negated_conjecture, (X1=esk6_1(esk1_0)|v7_struct_0(esk1_0)|m1_subset_1(esk15_3(esk1_0,X1,esk6_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_210, negated_conjecture, (X1=esk7_1(esk1_0)|v7_struct_0(esk1_0)|m1_subset_1(esk15_3(esk1_0,X1,esk7_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_165, negated_conjecture, (k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),esk15_3(esk1_0,X1,esk4_1(esk1_0)),X1)=esk4_1(esk1_0)|X1=esk4_1(esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_178, negated_conjecture, (k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),esk15_3(esk1_0,X1,esk5_1(esk1_0)),X1)=esk5_1(esk1_0)|X1=esk5_1(esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_191, negated_conjecture, (k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),esk15_3(esk1_0,X1,esk6_1(esk1_0)),X1)=esk6_1(esk1_0)|X1=esk6_1(esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_204, negated_conjecture, (k3_funct_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),esk15_3(esk1_0,X1,esk7_1(esk1_0)),X1)=esk7_1(esk1_0)|X1=esk7_1(esk1_0)|v7_struct_0(esk1_0)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_480, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_481, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_482, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_483, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_484, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_485, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_486, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_487, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_488, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_489, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_490, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_491, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_492, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_493, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_494, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_495, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_496, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_497, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_498, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_499, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_500, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_501, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_502, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_503, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_217, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),X1,esk2_1(esk1_0)),X1)=esk2_1(esk1_0)|X1=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_504, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_505, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_506, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_507, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_508, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_515, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_516, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_517, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_518, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_519, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_520, negated_conjecture, (v1_funct_2(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_521, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_522, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_523, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_524, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_525, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_526, negated_conjecture, (v3_funct_2(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_569, negated_conjecture, (esk7_1(esk1_0)=esk2_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_570, negated_conjecture, (esk6_1(esk1_0)=esk2_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_571, negated_conjecture, (esk5_1(esk1_0)=esk2_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_572, negated_conjecture, (esk4_1(esk1_0)=esk2_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_573, negated_conjecture, (esk3_1(esk1_0)=esk2_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_574, negated_conjecture, (esk7_1(esk1_0)=esk2_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_575, negated_conjecture, (esk6_1(esk1_0)=esk2_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_231, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),X1,esk3_1(esk1_0)),X1)=esk3_1(esk1_0)|X1=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_245, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),X1,esk4_1(esk1_0)),X1)=esk4_1(esk1_0)|X1=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_259, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),X1,esk5_1(esk1_0)),X1)=esk5_1(esk1_0)|X1=esk5_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_273, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),X1,esk6_1(esk1_0)),X1)=esk6_1(esk1_0)|X1=esk6_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_287, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),X1,esk7_1(esk1_0)),X1)=esk7_1(esk1_0)|X1=esk7_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_223, negated_conjecture, (X1=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),X1,esk2_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_237, negated_conjecture, (X1=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),X1,esk3_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_251, negated_conjecture, (X1=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),X1,esk4_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_265, negated_conjecture, (X1=esk5_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),X1,esk5_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_279, negated_conjecture, (X1=esk6_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),X1,esk6_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_293, negated_conjecture, (X1=esk7_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),X1,esk7_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_563, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk2_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_510, negated_conjecture, (r2_analoaf(k3_analmetr(esk1_0),X1,X2,X3,esk6_1(esk1_0))|~r2_analoaf(esk1_0,X1,X2,X3,esk6_1(esk1_0))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_512, negated_conjecture, (r2_analoaf(k3_analmetr(esk1_0),X1,X2,X3,esk4_1(esk1_0))|~r2_analoaf(esk1_0,X1,X2,X3,esk4_1(esk1_0))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_513, negated_conjecture, (r2_analoaf(k3_analmetr(esk1_0),X1,X2,X3,esk3_1(esk1_0))|~r2_analoaf(esk1_0,X1,X2,X3,esk3_1(esk1_0))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_514, negated_conjecture, (r2_analoaf(k3_analmetr(esk1_0),X1,X2,X3,esk2_1(esk1_0))|~r2_analoaf(esk1_0,X1,X2,X3,esk2_1(esk1_0))|~m1_subset_1(X3,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_576, negated_conjecture, (esk5_1(esk1_0)=esk2_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_577, negated_conjecture, (esk4_1(esk1_0)=esk2_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_578, negated_conjecture, (esk3_1(esk1_0)=esk2_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk2_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_585, negated_conjecture, (esk7_1(esk1_0)=esk3_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_586, negated_conjecture, (esk6_1(esk1_0)=esk3_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_587, negated_conjecture, (esk5_1(esk1_0)=esk3_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_588, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_589, negated_conjecture, (esk3_1(esk1_0)=esk2_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_590, negated_conjecture, (esk7_1(esk1_0)=esk3_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_591, negated_conjecture, (esk6_1(esk1_0)=esk3_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_592, negated_conjecture, (esk5_1(esk1_0)=esk3_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_593, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_594, negated_conjecture, (esk3_1(esk1_0)=esk2_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_601, negated_conjecture, (esk7_1(esk1_0)=esk4_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_602, negated_conjecture, (esk6_1(esk1_0)=esk4_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_603, negated_conjecture, (esk5_1(esk1_0)=esk4_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_604, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_605, negated_conjecture, (esk4_1(esk1_0)=esk2_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_606, negated_conjecture, (esk7_1(esk1_0)=esk4_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_607, negated_conjecture, (esk6_1(esk1_0)=esk4_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_608, negated_conjecture, (esk5_1(esk1_0)=esk4_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_609, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_610, negated_conjecture, (esk4_1(esk1_0)=esk2_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_295, negated_conjecture, (X1=esk5_1(esk1_0)|X2=X3|r1_aff_1(esk1_0,X2,X3,esk2_1(esk1_0))|r1_aff_1(esk1_0,X2,X3,X4)|v7_struct_0(esk1_0)|~r1_aff_1(esk1_0,X2,X3,esk4_1(esk1_0))|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~r2_analoaf(esk1_0,esk2_1(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0))|~r2_analoaf(esk1_0,X4,X5,esk4_1(esk1_0),X1)|~r2_analoaf(esk1_0,X4,esk4_1(esk1_0),X5,X1)|~r2_analoaf(esk1_0,X2,esk2_1(esk1_0),X3,esk3_1(esk1_0))|~r2_analoaf(esk1_0,X2,X4,X3,X5)|~r2_analoaf(esk1_0,X2,X3,esk2_1(esk1_0),esk3_1(esk1_0))|~r2_analoaf(esk1_0,X2,X3,X4,X5)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_564, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk2_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_565, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk2_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_566, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk2_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_567, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk2_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_568, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk2_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_579, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk3_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_580, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_581, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk3_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_582, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_583, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk3_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_584, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_629, negated_conjecture, (esk7_1(esk1_0)=esk5_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_630, negated_conjecture, (esk6_1(esk1_0)=esk5_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_631, negated_conjecture, (esk5_1(esk1_0)=esk4_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_632, negated_conjecture, (esk5_1(esk1_0)=esk3_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_633, negated_conjecture, (esk5_1(esk1_0)=esk2_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_634, negated_conjecture, (esk7_1(esk1_0)=esk5_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_635, negated_conjecture, (esk6_1(esk1_0)=esk5_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_636, negated_conjecture, (esk5_1(esk1_0)=esk4_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_637, negated_conjecture, (esk5_1(esk1_0)=esk3_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_638, negated_conjecture, (esk5_1(esk1_0)=esk2_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk5_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_639, negated_conjecture, (esk7_1(esk1_0)=esk6_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_640, negated_conjecture, (esk6_1(esk1_0)=esk5_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_641, negated_conjecture, (esk6_1(esk1_0)=esk4_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_642, negated_conjecture, (esk6_1(esk1_0)=esk3_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_643, negated_conjecture, (esk6_1(esk1_0)=esk2_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_644, negated_conjecture, (esk7_1(esk1_0)=esk6_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_645, negated_conjecture, (esk6_1(esk1_0)=esk5_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_646, negated_conjecture, (esk6_1(esk1_0)=esk4_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_647, negated_conjecture, (esk6_1(esk1_0)=esk3_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_648, negated_conjecture, (esk6_1(esk1_0)=esk2_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_649, negated_conjecture, (esk7_1(esk1_0)=esk6_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_650, negated_conjecture, (esk7_1(esk1_0)=esk5_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_651, negated_conjecture, (esk7_1(esk1_0)=esk4_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_296, negated_conjecture, (X1=esk3_1(esk1_0)|X2=X3|r1_aff_1(esk1_0,X2,X3,esk2_1(esk1_0))|r1_aff_1(esk1_0,X2,X3,X4)|v7_struct_0(esk1_0)|~r1_aff_1(esk1_0,X2,X3,X1)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~r2_analoaf(esk1_0,esk2_1(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0))|~r2_analoaf(esk1_0,X4,X5,X1,esk5_1(esk1_0))|~r2_analoaf(esk1_0,X4,X1,X5,esk5_1(esk1_0))|~r2_analoaf(esk1_0,X2,esk2_1(esk1_0),X3,esk4_1(esk1_0))|~r2_analoaf(esk1_0,X2,X4,X3,X5)|~r2_analoaf(esk1_0,X2,X3,esk2_1(esk1_0),esk4_1(esk1_0))|~r2_analoaf(esk1_0,X2,X3,X4,X5)|~m1_subset_1(X5,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_595, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk4_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_596, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk4_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_597, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk4_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_598, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk4_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_599, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_600, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_611, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk5_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_612, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk5_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_613, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk5_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_614, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk5_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_615, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_652, negated_conjecture, (esk7_1(esk1_0)=esk3_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_653, negated_conjecture, (esk7_1(esk1_0)=esk2_1(esk1_0)|v1_funct_2(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_654, negated_conjecture, (esk7_1(esk1_0)=esk6_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_655, negated_conjecture, (esk7_1(esk1_0)=esk5_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_656, negated_conjecture, (esk7_1(esk1_0)=esk4_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_657, negated_conjecture, (esk7_1(esk1_0)=esk3_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_658, negated_conjecture, (esk7_1(esk1_0)=esk2_1(esk1_0)|v3_funct_2(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk7_1(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_616, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk5_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_617, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk6_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_618, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk6_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_619, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk6_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_620, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk6_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_621, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_622, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_623, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk7_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_624, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_625, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk7_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_626, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk7_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_627, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_628, negated_conjecture, (v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk7_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_527, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk2_1(esk1_0)),esk7_1(esk1_0))=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_528, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk2_1(esk1_0)),esk6_1(esk1_0))=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_529, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk2_1(esk1_0)),esk5_1(esk1_0))=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_530, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk2_1(esk1_0)),esk4_1(esk1_0))=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_531, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk2_1(esk1_0)),esk3_1(esk1_0))=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_532, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk2_1(esk1_0)),esk2_1(esk1_0))=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_533, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk3_1(esk1_0)),esk7_1(esk1_0))=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_534, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0)),esk6_1(esk1_0))=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_535, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk3_1(esk1_0)),esk5_1(esk1_0))=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_536, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0)),esk4_1(esk1_0))=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_537, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk3_1(esk1_0)),esk3_1(esk1_0))=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_538, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0)),esk2_1(esk1_0))=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_539, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk4_1(esk1_0)),esk7_1(esk1_0))=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_540, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk4_1(esk1_0)),esk6_1(esk1_0))=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_541, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk4_1(esk1_0)),esk5_1(esk1_0))=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_542, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk4_1(esk1_0)),esk4_1(esk1_0))=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_543, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0)),esk3_1(esk1_0))=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_544, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0)),esk2_1(esk1_0))=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_545, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk5_1(esk1_0)),esk7_1(esk1_0))=esk5_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_546, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk5_1(esk1_0)),esk6_1(esk1_0))=esk5_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_547, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk5_1(esk1_0)),esk5_1(esk1_0))=esk5_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_548, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk5_1(esk1_0)),esk4_1(esk1_0))=esk5_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_549, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0)),esk3_1(esk1_0))=esk5_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_550, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk5_1(esk1_0)),esk2_1(esk1_0))=esk5_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_551, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk6_1(esk1_0)),esk7_1(esk1_0))=esk6_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_552, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk6_1(esk1_0)),esk6_1(esk1_0))=esk6_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_553, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk6_1(esk1_0)),esk5_1(esk1_0))=esk6_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_554, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk6_1(esk1_0)),esk4_1(esk1_0))=esk6_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_555, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0)),esk3_1(esk1_0))=esk6_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_556, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0)),esk2_1(esk1_0))=esk6_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_557, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk7_1(esk1_0)),esk7_1(esk1_0))=esk7_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_558, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0)),esk6_1(esk1_0))=esk7_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_559, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk7_1(esk1_0)),esk5_1(esk1_0))=esk7_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_560, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk7_1(esk1_0)),esk4_1(esk1_0))=esk7_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_561, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0)),esk3_1(esk1_0))=esk7_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_562, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk14_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk7_1(esk1_0)),esk2_1(esk1_0))=esk7_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_689, negated_conjecture, (esk7_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk2_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_298, negated_conjecture, (r1_aff_1(esk1_0,esk2_1(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0))|r1_aff_1(esk1_0,esk2_1(esk1_0),esk4_1(esk1_0),X1)|v7_struct_0(esk1_0)|r2_analoaf(esk1_0,X1,esk3_1(esk1_0),X2,esk5_1(esk1_0))|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~r2_analoaf(esk1_0,esk2_1(esk1_0),X1,esk4_1(esk1_0),X2)|~r2_analoaf(esk1_0,esk2_1(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0))|~r2_analoaf(esk1_0,esk2_1(esk1_0),esk4_1(esk1_0),X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_690, negated_conjecture, (esk6_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk2_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_691, negated_conjecture, (esk5_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk2_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_692, negated_conjecture, (esk4_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk2_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_693, negated_conjecture, (esk3_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk2_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_694, negated_conjecture, (esk7_1(esk1_0)=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk3_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_695, negated_conjecture, (esk6_1(esk1_0)=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_696, negated_conjecture, (esk5_1(esk1_0)=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk3_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_697, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_698, negated_conjecture, (esk3_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_699, negated_conjecture, (esk7_1(esk1_0)=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk4_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_700, negated_conjecture, (esk6_1(esk1_0)=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk4_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_659, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk2_1(esk1_0)),esk7_1(esk1_0))=esk2_1(esk1_0)|esk7_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_660, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk2_1(esk1_0)),esk6_1(esk1_0))=esk2_1(esk1_0)|esk6_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_661, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk2_1(esk1_0)),esk5_1(esk1_0))=esk2_1(esk1_0)|esk5_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_662, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk2_1(esk1_0)),esk4_1(esk1_0))=esk2_1(esk1_0)|esk4_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_701, negated_conjecture, (esk5_1(esk1_0)=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk4_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_702, negated_conjecture, (esk4_1(esk1_0)=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_703, negated_conjecture, (esk4_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_704, negated_conjecture, (esk7_1(esk1_0)=esk5_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk5_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_705, negated_conjecture, (esk6_1(esk1_0)=esk5_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk5_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_706, negated_conjecture, (esk5_1(esk1_0)=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk5_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_707, negated_conjecture, (esk5_1(esk1_0)=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_708, negated_conjecture, (esk5_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk5_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_709, negated_conjecture, (esk7_1(esk1_0)=esk6_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk6_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_710, negated_conjecture, (esk6_1(esk1_0)=esk5_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk6_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_711, negated_conjecture, (esk6_1(esk1_0)=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk6_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_712, negated_conjecture, (esk6_1(esk1_0)=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_713, negated_conjecture, (esk6_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_714, negated_conjecture, (esk7_1(esk1_0)=esk6_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_715, negated_conjecture, (esk7_1(esk1_0)=esk5_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk7_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_716, negated_conjecture, (esk7_1(esk1_0)=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk7_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_717, negated_conjecture, (esk7_1(esk1_0)=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_718, negated_conjecture, (esk7_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0))|m1_subset_1(esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk7_1(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0))))))).#trainneg
cnf(i_0_663, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk2_1(esk1_0)),esk3_1(esk1_0))=esk2_1(esk1_0)|esk3_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_299, negated_conjecture, (X1=esk7_1(esk1_0)|X2=X3|r1_aff_1(esk1_0,X2,X3,esk2_1(esk1_0))|r1_aff_1(esk1_0,X2,X3,X4)|v7_struct_0(esk1_0)|~r1_aff_1(esk1_0,X2,X3,esk6_1(esk1_0))|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~r2_analoaf(esk1_0,esk2_1(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0))|~r2_analoaf(esk1_0,X4,X5,esk6_1(esk1_0),X1)|~r2_analoaf(esk1_0,X4,esk6_1(esk1_0),X5,X1)|~r2_analoaf(esk1_0,X2,esk2_1(esk1_0),X3,esk3_1(esk1_0))|~r2_analoaf(esk1_0,X2,X4,X3,X5)|~r2_analoaf(esk1_0,X2,X3,esk2_1(esk1_0),esk3_1(esk1_0))|~r2_analoaf(esk1_0,X2,X3,X4,X5)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_664, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk3_1(esk1_0)),esk7_1(esk1_0))=esk3_1(esk1_0)|esk7_1(esk1_0)=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_665, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0)),esk6_1(esk1_0))=esk3_1(esk1_0)|esk6_1(esk1_0)=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_666, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk3_1(esk1_0)),esk5_1(esk1_0))=esk3_1(esk1_0)|esk5_1(esk1_0)=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_667, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0)),esk4_1(esk1_0))=esk3_1(esk1_0)|esk4_1(esk1_0)=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_668, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0)),esk2_1(esk1_0))=esk3_1(esk1_0)|esk3_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_669, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk4_1(esk1_0)),esk7_1(esk1_0))=esk4_1(esk1_0)|esk7_1(esk1_0)=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_670, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk4_1(esk1_0)),esk6_1(esk1_0))=esk4_1(esk1_0)|esk6_1(esk1_0)=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_671, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk4_1(esk1_0)),esk5_1(esk1_0))=esk4_1(esk1_0)|esk5_1(esk1_0)=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_672, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0)),esk3_1(esk1_0))=esk4_1(esk1_0)|esk4_1(esk1_0)=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_673, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0)),esk2_1(esk1_0))=esk4_1(esk1_0)|esk4_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_674, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk5_1(esk1_0)),esk7_1(esk1_0))=esk5_1(esk1_0)|esk7_1(esk1_0)=esk5_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_675, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk5_1(esk1_0)),esk6_1(esk1_0))=esk5_1(esk1_0)|esk6_1(esk1_0)=esk5_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_676, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk5_1(esk1_0)),esk4_1(esk1_0))=esk5_1(esk1_0)|esk5_1(esk1_0)=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_677, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0)),esk3_1(esk1_0))=esk5_1(esk1_0)|esk5_1(esk1_0)=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_678, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk5_1(esk1_0)),esk2_1(esk1_0))=esk5_1(esk1_0)|esk5_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_679, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk7_1(esk1_0),esk6_1(esk1_0)),esk7_1(esk1_0))=esk6_1(esk1_0)|esk7_1(esk1_0)=esk6_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_680, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk6_1(esk1_0)),esk5_1(esk1_0))=esk6_1(esk1_0)|esk6_1(esk1_0)=esk5_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_681, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk6_1(esk1_0)),esk4_1(esk1_0))=esk6_1(esk1_0)|esk6_1(esk1_0)=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_682, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0)),esk3_1(esk1_0))=esk6_1(esk1_0)|esk6_1(esk1_0)=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_683, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0)),esk2_1(esk1_0))=esk6_1(esk1_0)|esk6_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_684, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0)),esk6_1(esk1_0))=esk7_1(esk1_0)|esk7_1(esk1_0)=esk6_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_685, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk5_1(esk1_0),esk7_1(esk1_0)),esk5_1(esk1_0))=esk7_1(esk1_0)|esk7_1(esk1_0)=esk5_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_686, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk4_1(esk1_0),esk7_1(esk1_0)),esk4_1(esk1_0))=esk7_1(esk1_0)|esk7_1(esk1_0)=esk4_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_687, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0)),esk3_1(esk1_0))=esk7_1(esk1_0)|esk7_1(esk1_0)=esk3_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_688, negated_conjecture, (k3_funct_2(u1_struct_0(k3_analmetr(esk1_0)),u1_struct_0(k3_analmetr(esk1_0)),esk15_3(k3_analmetr(esk1_0),esk2_1(esk1_0),esk7_1(esk1_0)),esk2_1(esk1_0))=esk7_1(esk1_0)|esk7_1(esk1_0)=esk2_1(esk1_0)|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_306, negated_conjecture, (r1_aff_1(esk1_0,esk2_1(esk1_0),esk3_1(esk1_0),esk4_1(esk1_0))|r1_aff_1(esk1_0,esk2_1(esk1_0),esk3_1(esk1_0),X1)|v7_struct_0(esk1_0)|r2_analoaf(esk1_0,X1,esk4_1(esk1_0),X2,esk5_1(esk1_0))|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~r2_analoaf(esk1_0,esk2_1(esk1_0),X1,esk3_1(esk1_0),X2)|~r2_analoaf(esk1_0,esk2_1(esk1_0),esk3_1(esk1_0),X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_310, negated_conjecture, (r1_aff_1(esk1_0,esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0))|r1_aff_1(esk1_0,esk2_1(esk1_0),esk3_1(esk1_0),X1)|v7_struct_0(esk1_0)|r2_analoaf(esk1_0,X1,esk6_1(esk1_0),X2,esk7_1(esk1_0))|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~r2_analoaf(esk1_0,esk2_1(esk1_0),X1,esk3_1(esk1_0),X2)|~r2_analoaf(esk1_0,esk2_1(esk1_0),esk3_1(esk1_0),X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_729, negated_conjecture, (r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0))|r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0),X1)|v7_struct_0(k3_analmetr(esk1_0))|r2_analoaf(k3_analmetr(esk1_0),X1,esk3_1(esk1_0),X2,esk7_1(esk1_0))|~r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),X1,esk6_1(esk1_0),X2)|~r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0),X1,X2)|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_743, negated_conjecture, (r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))|r2_analoaf(k3_analmetr(esk1_0),esk3_1(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0),esk7_1(esk1_0)))).#trainneg
cnf(i_0_748, negated_conjecture, (r2_analoaf(esk1_0,esk3_1(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0),esk7_1(esk1_0))|r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_300, negated_conjecture, (X1=esk3_1(esk1_0)|X2=X3|r1_aff_1(esk1_0,X2,X3,esk2_1(esk1_0))|r1_aff_1(esk1_0,X2,X3,X4)|v7_struct_0(esk1_0)|~r1_aff_1(esk1_0,X2,X3,X1)|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~r2_analoaf(esk1_0,esk2_1(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0))|~r2_analoaf(esk1_0,X4,X5,X1,esk7_1(esk1_0))|~r2_analoaf(esk1_0,X4,X1,X5,esk7_1(esk1_0))|~r2_analoaf(esk1_0,X2,esk2_1(esk1_0),X3,esk6_1(esk1_0))|~r2_analoaf(esk1_0,X2,X4,X3,X5)|~r2_analoaf(esk1_0,X2,X3,esk2_1(esk1_0),esk6_1(esk1_0))|~r2_analoaf(esk1_0,X2,X3,X4,X5)|~m1_subset_1(X5,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_741, negated_conjecture, (r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0))|r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0),X1)|v7_struct_0(k3_analmetr(esk1_0))|r2_analoaf(k3_analmetr(esk1_0),X1,esk3_1(esk1_0),X2,esk5_1(esk1_0))|~r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),X1,esk4_1(esk1_0),X2)|~r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0),X1,X2)|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0))))).#trainneg
cnf(i_0_754, negated_conjecture, (r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))|r2_analoaf(k3_analmetr(esk1_0),esk3_1(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0),esk5_1(esk1_0)))).#trainneg
cnf(i_0_759, negated_conjecture, (r2_analoaf(esk1_0,esk3_1(esk1_0),esk3_1(esk1_0),esk5_1(esk1_0),esk5_1(esk1_0))|r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_302, negated_conjecture, (r1_aff_1(esk1_0,esk2_1(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0))|r1_aff_1(esk1_0,esk2_1(esk1_0),esk6_1(esk1_0),X1)|v7_struct_0(esk1_0)|r2_analoaf(esk1_0,X1,esk3_1(esk1_0),X2,esk7_1(esk1_0))|~v2_diraf(esk1_0)|~v1_diraf(esk1_0)|~r2_analoaf(esk1_0,esk2_1(esk1_0),X1,esk6_1(esk1_0),X2)|~r2_analoaf(esk1_0,esk2_1(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0),esk7_1(esk1_0))|~r2_analoaf(esk1_0,esk2_1(esk1_0),esk6_1(esk1_0),X1,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v11_aff_2(esk1_0))).#trainneg
cnf(i_0_728, negated_conjecture, (r5_analmetr(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0),esk3_1(esk1_0))|r5_analmetr(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0),X1)|r2_analoaf(k3_analmetr(esk1_0),X1,esk3_1(esk1_0),X2,esk7_1(esk1_0))|v2_struct_0(k3_analmetr(esk1_0))|~r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),X1,esk6_1(esk1_0),X2)|~r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),esk6_1(esk1_0),X1,X2)|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0)))|~v8_conmetr(k3_analmetr(esk1_0))|~l1_analmetr(k3_analmetr(esk1_0))|~v3_analmetr(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_740, negated_conjecture, (r5_analmetr(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0),esk3_1(esk1_0))|r5_analmetr(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0),X1)|r2_analoaf(k3_analmetr(esk1_0),X1,esk3_1(esk1_0),X2,esk5_1(esk1_0))|v2_struct_0(k3_analmetr(esk1_0))|~r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),X1,esk4_1(esk1_0),X2)|~r2_analoaf(k3_analmetr(esk1_0),esk2_1(esk1_0),esk4_1(esk1_0),X1,X2)|~m1_subset_1(X2,u1_struct_0(k3_analmetr(esk1_0)))|~m1_subset_1(X1,u1_struct_0(k3_analmetr(esk1_0)))|~v8_conmetr(k3_analmetr(esk1_0))|~l1_analmetr(k3_analmetr(esk1_0))|~v3_analmetr(k3_analmetr(esk1_0)))).#trainneg
cnf(i_0_765, negated_conjecture, (r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0))|v7_struct_0(k3_analmetr(esk1_0))|r2_analoaf(k3_analmetr(esk1_0),esk6_1(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0),esk7_1(esk1_0)))).#trainneg
cnf(i_0_771, negated_conjecture, (r2_analoaf(esk1_0,esk6_1(esk1_0),esk6_1(esk1_0),esk7_1(esk1_0),esk7_1(esk1_0))|r1_aff_1(k3_analmetr(esk1_0),esk2_1(esk1_0),esk3_1(esk1_0),esk6_1(esk1_0))|v7_struct_0(k3_analmetr(esk1_0)))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 149
# Removed by relevancy pruning/SinE    : 128
# Initial clauses                      : 78
# Removed in clause preprocessing      : 5
# Initial clauses in saturation        : 73
# Processed clauses                    : 707
# ...of these trivial                  : 1
# ...subsumed                          : 12
# ...remaining for further processing  : 694
# Other redundant clauses eliminated   : 16
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 480
# Generated clauses                    : 707
# ...of the previous two non-trivial   : 684
# Contextual simplify-reflections      : 0
# Paramodulations                      : 701
# Factorizations                       : 0
# Equation resolutions                 : 16
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 205
#    Positive orientable unit clauses  : 30
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 2
#    Non-unit-clauses                  : 173
# Current number of unprocessed clauses: 50
# ...number of literals in the above   : 802
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1177
# Proof object given clauses           : 65
# Proof search given clauses           : 694
# Clause-clause subsumption calls (NU) : 21454
# Rec. Clause-clause subsumption calls : 6702
# Non-unit clause-clause subsumptions  : 11
# Unit Clause-clause subsumption calls : 688
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 13
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 47426

# -------------------------------------------------
# User time                : 0.071 s
# System time              : 0.000 s
# Total time               : 0.071 s
# Maximum resident set size: 8448 pages
