# local_data = 0x25d6dd0
#TFS# Evaluating axioms..
# ENIGMATIC: Connected to the TF server '127.0.0.1:8888'.
# ENIGMATIC: LightGBM model './big03/model.lgb' loaded with 17544 features 'C(l,x[c=3],r,h[b=4096],v,c[b=256],d,a):G:P'
#TF#SERVER: Lgb model filter skipped 1 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=0, query=83, total=97).
#TF#SERVER: Sent 40198 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (336 bytes expected).
#TF#SERVER: Received 336 bytes from the TF server.
# Preprocessing time       : 0.729 s
# ENIGMATIC: LightGBM parental guidance model #1 './froc_gnn_model2/model.lgb' loaded with 13963 features 'C(l,x[c=3],r,h[b=2048],v,c[b=256],d,a):G:F'
#TF#SERVER: Lgb model filter skipped 4 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=83, query=61, total=158).
#TF#SERVER: Sent 86726 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (580 bytes expected).
#TF#SERVER: Received 580 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 49 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=145, query=122, total=281).
#TF#SERVER: Sent 248422 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1072 bytes expected).
#TF#SERVER: Received 1072 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 345 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=263, query=683, total=960).
#TF#SERVER: Sent 1440803 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3788 bytes expected).
#TF#SERVER: Received 3788 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 359 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=339, query=679, total=1032).
#TF#SERVER: Sent 1478074 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4076 bytes expected).
#TF#SERVER: Received 4076 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 365 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=382, query=712, total=1108).
#TF#SERVER: Sent 1560745 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4380 bytes expected).
#TF#SERVER: Received 4380 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 410 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=415, query=615, total=1044).
#TF#SERVER: Sent 1491730 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4124 bytes expected).
#TF#SERVER: Received 4124 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 479 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=445, query=621, total=1080).
#TF#SERVER: Sent 1557884 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4268 bytes expected).
#TF#SERVER: Received 4268 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 474 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=463, query=661, total=1138).
#TF#SERVER: Sent 1662415 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4500 bytes expected).
#TF#SERVER: Received 4500 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 323 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=487, query=722, total=1223).
#TF#SERVER: Sent 1727066 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4840 bytes expected).
#TF#SERVER: Received 4840 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 356 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=509, query=696, total=1219).
#TF#SERVER: Sent 1745801 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4824 bytes expected).
#TF#SERVER: Received 4824 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 352 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=520, query=692, total=1226).
#TF#SERVER: Sent 1826184 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4852 bytes expected).
#TF#SERVER: Received 4852 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 463 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=540, query=599, total=1153).
#TF#SERVER: Sent 1646829 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4560 bytes expected).
#TF#SERVER: Received 4560 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 367 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=557, query=675, total=1246).
#TF#SERVER: Sent 1711538 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4932 bytes expected).
#TF#SERVER: Received 4932 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 343 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=576, query=721, total=1311).
#TF#SERVER: Sent 1737909 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5192 bytes expected).
#TF#SERVER: Received 5192 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 308 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=598, query=756, total=1368).
#TF#SERVER: Sent 2008194 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5420 bytes expected).
#TF#SERVER: Received 5420 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 404 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=624, query=805, total=1443).
#TF#SERVER: Sent 2064377 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5720 bytes expected).
#TF#SERVER: Received 5720 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 463 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=635, query=584, total=1233).
#TF#SERVER: Sent 1746273 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4880 bytes expected).
#TF#SERVER: Received 4880 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 312 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=649, query=936, total=1599).
#TF#SERVER: Sent 2378324 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6344 bytes expected).
#TF#SERVER: Received 6344 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 325 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=667, query=720, total=1401).
#TF#SERVER: Sent 1974093 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5552 bytes expected).
#TF#SERVER: Received 5552 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 307 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=676, query=804, total=1494).
#TF#SERVER: Sent 2186546 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5924 bytes expected).
#TF#SERVER: Received 5924 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 513 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=698, query=605, total=1317).
#TF#SERVER: Sent 1874686 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5216 bytes expected).
#TF#SERVER: Received 5216 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 404 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=729, query=628, total=1371).
#TF#SERVER: Sent 1899522 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5432 bytes expected).
#TF#SERVER: Received 5432 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 373 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=760, query=753, total=1527).
#TF#SERVER: Sent 2028191 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6056 bytes expected).
#TF#SERVER: Received 6056 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 402 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=787, query=673, total=1474).
#TF#SERVER: Sent 2246274 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5844 bytes expected).
#TF#SERVER: Received 5844 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 533 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=821, query=758, total=1593).
#TF#SERVER: Sent 2346211 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6320 bytes expected).
#TF#SERVER: Received 6320 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 390 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=860, query=650, total=1524).
#TF#SERVER: Sent 2272302 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6044 bytes expected).
#TF#SERVER: Received 6044 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 461 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=887, query=832, total=1733).
#TF#SERVER: Sent 2777761 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6880 bytes expected).
#TF#SERVER: Received 6880 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 475 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=910, query=625, total=1549).
#TF#SERVER: Sent 2275132 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6144 bytes expected).
#TF#SERVER: Received 6144 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 578 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=926, query=450, total=1390).
#TF#SERVER: Sent 1943247 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5508 bytes expected).
#TF#SERVER: Received 5508 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 335 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=946, query=694, total=1654).
#TF#SERVER: Sent 2423576 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6564 bytes expected).
#TF#SERVER: Received 6564 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 298 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=967, query=726, total=1707).
#TF#SERVER: Sent 2558062 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6776 bytes expected).
#TF#SERVER: Received 6776 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 354 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=985, query=815, total=1814).
#TF#SERVER: Sent 2755797 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7204 bytes expected).
#TF#SERVER: Received 7204 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 289 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1006, query=828, total=1848).
#TF#SERVER: Sent 2810955 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7340 bytes expected).
#TF#SERVER: Received 7340 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 506 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1023, query=587, total=1624).
#TF#SERVER: Sent 2356854 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6444 bytes expected).
#TF#SERVER: Received 6444 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 443 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1037, query=608, total=1659).
#TF#SERVER: Sent 2494644 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6584 bytes expected).
#TF#SERVER: Received 6584 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 492 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1048, query=655, total=1717).
#TF#SERVER: Sent 2696879 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6816 bytes expected).
#TF#SERVER: Received 6816 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 545 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1061, query=653, total=1728).
#TF#SERVER: Sent 2701718 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6860 bytes expected).
#TF#SERVER: Received 6860 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 601 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1081, query=594, total=1689).
#TF#SERVER: Sent 2557445 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6704 bytes expected).
#TF#SERVER: Received 6704 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 463 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1103, query=566, total=1683).
#TF#SERVER: Sent 2581772 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6680 bytes expected).
#TF#SERVER: Received 6680 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 510 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1123, query=866, total=2003).
#TF#SERVER: Sent 3008493 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7960 bytes expected).
#TF#SERVER: Received 7960 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 353 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1137, query=784, total=1935).
#TF#SERVER: Sent 2714804 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7688 bytes expected).
#TF#SERVER: Received 7688 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 362 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1157, query=689, total=1860).
#TF#SERVER: Sent 2850837 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7388 bytes expected).
#TF#SERVER: Received 7388 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 578 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1181, query=676, total=1871).
#TF#SERVER: Sent 2830477 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7432 bytes expected).
#TF#SERVER: Received 7432 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 537 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1194, query=652, total=1860).
#TF#SERVER: Sent 2750662 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7388 bytes expected).
#TF#SERVER: Received 7388 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 505 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1214, query=549, total=1777).
#TF#SERVER: Sent 2625335 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7056 bytes expected).
#TF#SERVER: Received 7056 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 547 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1224, query=724, total=1962).
#TF#SERVER: Sent 2944388 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7796 bytes expected).
#TF#SERVER: Received 7796 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 737 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1246, query=1011, total=2271).
#TF#SERVER: Sent 3732040 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9032 bytes expected).
#TF#SERVER: Received 9032 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 660 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1267, query=608, total=1889).
#TF#SERVER: Sent 2943229 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7504 bytes expected).
#TF#SERVER: Received 7504 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 542 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1291, query=593, total=1898).
#TF#SERVER: Sent 2795561 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7540 bytes expected).
#TF#SERVER: Received 7540 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 519 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1310, query=603, total=1927).
#TF#SERVER: Sent 2736161 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7656 bytes expected).
#TF#SERVER: Received 7656 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 662 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1329, query=573, total=1916).
#TF#SERVER: Sent 2671623 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7612 bytes expected).
#TF#SERVER: Received 7612 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 664 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1354, query=548, total=1916).
#TF#SERVER: Sent 2744760 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7612 bytes expected).
#TF#SERVER: Received 7612 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 510 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1378, query=553, total=1945).
#TF#SERVER: Sent 2986711 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7728 bytes expected).
#TF#SERVER: Received 7728 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 609 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1394, query=623, total=2031).
#TF#SERVER: Sent 2864131 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8072 bytes expected).
#TF#SERVER: Received 8072 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 446 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1418, query=607, total=2039).
#TF#SERVER: Sent 2989136 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8104 bytes expected).
#TF#SERVER: Received 8104 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 671 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1429, query=659, total=2102).
#TF#SERVER: Sent 3136407 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8356 bytes expected).
#TF#SERVER: Received 8356 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 329 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1435, query=752, total=2201).
#TF#SERVER: Sent 3370158 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8752 bytes expected).
#TF#SERVER: Received 8752 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 336 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1446, query=715, total=2175).
#TF#SERVER: Sent 3342251 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8648 bytes expected).
#TF#SERVER: Received 8648 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 571 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1459, query=549, total=2022).
#TF#SERVER: Sent 2933261 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8036 bytes expected).
#TF#SERVER: Received 8036 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 322 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1470, query=709, total=2193).
#TF#SERVER: Sent 3227116 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8720 bytes expected).
#TF#SERVER: Received 8720 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 486 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1477, query=762, total=2253).
#TF#SERVER: Sent 3444815 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8960 bytes expected).
#TF#SERVER: Received 8960 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 716 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1486, query=780, total=2280).
#TF#SERVER: Sent 3423482 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9068 bytes expected).
#TF#SERVER: Received 9068 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 604 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1501, query=496, total=2011).
#TF#SERVER: Sent 3062583 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7992 bytes expected).
#TF#SERVER: Received 7992 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 557 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1518, query=641, total=2173).
#TF#SERVER: Sent 3201175 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8640 bytes expected).
#TF#SERVER: Received 8640 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 512 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=580, total=2130).
#TF#SERVER: Sent 3146394 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8468 bytes expected).
#TF#SERVER: Received 8468 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 663 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=362, total=1912).
#TF#SERVER: Sent 2691870 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7596 bytes expected).
#TF#SERVER: Received 7596 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 523 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=588, total=2138).
#TF#SERVER: Sent 3105381 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8500 bytes expected).
#TF#SERVER: Received 8500 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 361 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=719, total=2269).
#TF#SERVER: Sent 3400607 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9024 bytes expected).
#TF#SERVER: Received 9024 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 592 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=824, total=2374).
#TF#SERVER: Sent 3674740 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9444 bytes expected).
#TF#SERVER: Received 9444 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 648 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=755, total=2305).
#TF#SERVER: Sent 3592383 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9168 bytes expected).
#TF#SERVER: Received 9168 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 515 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=859, total=2409).
#TF#SERVER: Sent 3735209 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9584 bytes expected).
#TF#SERVER: Received 9584 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 314 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=762, total=2312).
#TF#SERVER: Sent 3598538 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9196 bytes expected).
#TF#SERVER: Received 9196 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 444 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=631, total=2181).
#TF#SERVER: Sent 3304898 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8672 bytes expected).
#TF#SERVER: Received 8672 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 625 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=666, total=2216).
#TF#SERVER: Sent 3278204 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8812 bytes expected).
#TF#SERVER: Received 8812 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 571 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=643, total=2193).
#TF#SERVER: Sent 3369344 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8720 bytes expected).
#TF#SERVER: Received 8720 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 480 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=571, total=2121).
#TF#SERVER: Sent 3072966 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8432 bytes expected).
#TF#SERVER: Received 8432 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 553 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=849, total=2399).
#TF#SERVER: Sent 3714296 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9544 bytes expected).
#TF#SERVER: Received 9544 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 421 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=617, total=2167).
#TF#SERVER: Sent 3432782 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8616 bytes expected).
#TF#SERVER: Received 8616 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 643 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=505, total=2055).
#TF#SERVER: Sent 3045835 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8168 bytes expected).
#TF#SERVER: Received 8168 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 447 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=590, total=2140).
#TF#SERVER: Sent 3191618 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8508 bytes expected).
#TF#SERVER: Received 8508 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 730 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=617, total=2167).
#TF#SERVER: Sent 3312327 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8616 bytes expected).
#TF#SERVER: Received 8616 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 791 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=702, total=2252).
#TF#SERVER: Sent 3385703 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8956 bytes expected).
#TF#SERVER: Received 8956 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 560 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=483, total=2033).
#TF#SERVER: Sent 2991016 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8080 bytes expected).
#TF#SERVER: Received 8080 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 682 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=470, total=2020).
#TF#SERVER: Sent 2918944 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8028 bytes expected).
#TF#SERVER: Received 8028 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 499 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=532, total=2082).
#TF#SERVER: Sent 3098580 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8276 bytes expected).
#TF#SERVER: Received 8276 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 720 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=469, total=2019).
#TF#SERVER: Sent 3086769 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8024 bytes expected).
#TF#SERVER: Received 8024 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 649 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=514, total=2064).
#TF#SERVER: Sent 2984455 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8204 bytes expected).
#TF#SERVER: Received 8204 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 423 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=621, total=2171).
#TF#SERVER: Sent 3868592 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8632 bytes expected).
#TF#SERVER: Received 8632 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 431 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=652, total=2202).
#TF#SERVER: Sent 3868252 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8756 bytes expected).
#TF#SERVER: Received 8756 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 601 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=470, total=2020).
#TF#SERVER: Sent 3240990 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8028 bytes expected).
#TF#SERVER: Received 8028 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 462 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=684, total=2234).
#TF#SERVER: Sent 3744121 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8884 bytes expected).
#TF#SERVER: Received 8884 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 365 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=798, total=2348).
#TF#SERVER: Sent 3489289 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9340 bytes expected).
#TF#SERVER: Received 9340 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 441 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=811, total=2361).
#TF#SERVER: Sent 3877224 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9392 bytes expected).
#TF#SERVER: Received 9392 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 643 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=699, total=2249).
#TF#SERVER: Sent 3707847 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8944 bytes expected).
#TF#SERVER: Received 8944 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 695 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=605, total=2155).
#TF#SERVER: Sent 3695946 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8568 bytes expected).
#TF#SERVER: Received 8568 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 696 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=617, total=2167).
#TF#SERVER: Sent 3630623 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8616 bytes expected).
#TF#SERVER: Received 8616 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 580 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=512, total=2062).
#TF#SERVER: Sent 3316927 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8196 bytes expected).
#TF#SERVER: Received 8196 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 657 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=559, total=2109).
#TF#SERVER: Sent 3411688 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8384 bytes expected).
#TF#SERVER: Received 8384 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 623 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=674, total=2224).
#TF#SERVER: Sent 3618054 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8844 bytes expected).
#TF#SERVER: Received 8844 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 710 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=497, total=2047).
#TF#SERVER: Sent 3417551 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8136 bytes expected).
#TF#SERVER: Received 8136 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 608 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=459, total=2009).
#TF#SERVER: Sent 3288206 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7984 bytes expected).
#TF#SERVER: Received 7984 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 497 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=619, total=2169).
#TF#SERVER: Sent 3579726 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8624 bytes expected).
#TF#SERVER: Received 8624 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 956 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=767, total=2317).
#TF#SERVER: Sent 3877614 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9216 bytes expected).
#TF#SERVER: Received 9216 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 486 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=622, total=2172).
#TF#SERVER: Sent 3484750 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8636 bytes expected).
#TF#SERVER: Received 8636 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 738 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=506, total=2056).
#TF#SERVER: Sent 3301648 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8172 bytes expected).
#TF#SERVER: Received 8172 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 526 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=732, total=2282).
#TF#SERVER: Sent 3777776 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9076 bytes expected).
#TF#SERVER: Received 9076 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 493 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=567, total=2117).
#TF#SERVER: Sent 3453983 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8416 bytes expected).
#TF#SERVER: Received 8416 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 650 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=472, total=2022).
#TF#SERVER: Sent 3398811 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8036 bytes expected).
#TF#SERVER: Received 8036 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 593 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=617, total=2167).
#TF#SERVER: Sent 3598204 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8616 bytes expected).
#TF#SERVER: Received 8616 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 628 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=469, total=2019).
#TF#SERVER: Sent 3242461 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8024 bytes expected).
#TF#SERVER: Received 8024 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 711 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=666, total=2216).
#TF#SERVER: Sent 3740348 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8812 bytes expected).
#TF#SERVER: Received 8812 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 928 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=909, total=2459).
#TF#SERVER: Sent 4029558 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9784 bytes expected).
#TF#SERVER: Received 9784 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 602 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=491, total=2041).
#TF#SERVER: Sent 3218273 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8112 bytes expected).
#TF#SERVER: Received 8112 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 629 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=1536, query=517, total=2067).
#TF#SERVER: Sent 3282882 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8216 bytes expected).
#TF#SERVER: Received 8216 bytes from the TF server.

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(redefinition_k3_rlvect_1, axiom, ![X1, X2, X3]:(((((v2_rlvect_1(X1)&l1_algstr_0(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>k3_rlvect_1(X1,X2,X3)=k1_algstr_0(X1,X2,X3))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', redefinition_k3_rlvect_1)).
fof(commutativity_k3_rlvect_1, axiom, ![X1, X2, X3]:(((((v2_rlvect_1(X1)&l1_algstr_0(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>k3_rlvect_1(X1,X2,X3)=k3_rlvect_1(X1,X3,X2))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', commutativity_k3_rlvect_1)).
fof(d3_rlvect_1, axiom, ![X1]:((l1_algstr_0(X1)=>(v3_rlvect_1(X1)<=>![X2]:((m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((m1_subset_1(X3,u1_struct_0(X1))=>![X4]:((m1_subset_1(X4,u1_struct_0(X1))=>k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))))))))))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', d3_rlvect_1)).
fof(d14_algstr_0, axiom, ![X1]:((l2_algstr_0(X1)=>![X2]:((m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((m1_subset_1(X3,u1_struct_0(X1))=>k5_algstr_0(X1,X2,X3)=k1_algstr_0(X1,X2,k4_algstr_0(X1,X3)))))))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', d14_algstr_0)).
fof(t4_rlvect_1, axiom, ![X1]:((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&l2_algstr_0(X1))=>![X2]:((m1_subset_1(X2,u1_struct_0(X1))=>(k1_algstr_0(X1,X2,k4_struct_0(X1))=X2&k1_algstr_0(X1,k4_struct_0(X1),X2)=X2))))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', t4_rlvect_1)).
fof(t15_rlvect_1, axiom, ![X1]:((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&l2_algstr_0(X1))=>![X2]:((m1_subset_1(X2,u1_struct_0(X1))=>k5_algstr_0(X1,X2,X2)=k4_struct_0(X1))))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', t15_rlvect_1)).
fof(dt_k4_algstr_0, axiom, ![X1, X2]:(((l2_algstr_0(X1)&m1_subset_1(X2,u1_struct_0(X1)))=>m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1)))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', dt_k4_algstr_0)).
fof(dt_l2_algstr_0, axiom, ![X1]:((l2_algstr_0(X1)=>(l2_struct_0(X1)&l1_algstr_0(X1)))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', dt_l2_algstr_0)).
fof(dt_k5_algstr_0, axiom, ![X1, X2, X3]:((((l2_algstr_0(X1)&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1)))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', dt_k5_algstr_0)).
fof(existence_m1_subset_1, axiom, ![X1]:(?[X2]:(m1_subset_1(X2,X1))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', existence_m1_subset_1)).
fof(dt_k1_algstr_0, axiom, ![X1, X2, X3]:((((l1_algstr_0(X1)&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1)))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', dt_k1_algstr_0)).
fof(fraenkel_a_3_1_zmodul01, axiom, ![X1, X2, X3, X4]:(((((((((((((~(v2_struct_0(X2))&v13_algstr_0(X2))&v2_rlvect_1(X2))&v3_rlvect_1(X2))&v4_rlvect_1(X2))&v2_zmodul01(X2))&v3_zmodul01(X2))&v4_zmodul01(X2))&v5_zmodul01(X2))&l1_zmodul01(X2))&m1_subset_1(X3,u1_struct_0(X2)))&m1_zmodul01(X4,X2))=>(r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))<=>?[X5]:(((m1_subset_1(X5,u1_struct_0(X2))&X1=k3_rlvect_1(X2,X3,X5))&r1_struct_0(X4,X5)))))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', fraenkel_a_3_1_zmodul01)).
fof(dt_l1_zmodul01, axiom, ![X1]:((l1_zmodul01(X1)=>l2_algstr_0(X1))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', dt_l1_zmodul01)).
fof(t2_subset, axiom, ![X1, X2]:((m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2)))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', t2_subset)).
fof(d5_struct_0, axiom, ![X1]:((l1_struct_0(X1)=>![X2]:((r1_struct_0(X1,X2)<=>r2_hidden(X2,u1_struct_0(X1)))))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', d5_struct_0)).
fof(dt_l2_struct_0, axiom, ![X1]:((l2_struct_0(X1)=>l1_struct_0(X1))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', dt_l2_struct_0)).
fof(fc2_struct_0, axiom, ![X1]:(((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1))))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', fc2_struct_0)).
fof(t33_zmodul01, axiom, ![X1]:(((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v2_zmodul01(X1))&v3_zmodul01(X1))&v4_zmodul01(X1))&v5_zmodul01(X1))&l1_zmodul01(X1))=>![X2]:((m1_zmodul01(X2,X1)=>r1_struct_0(X2,k4_struct_0(X1)))))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', t33_zmodul01)).
fof(t36_zmodul01, axiom, ![X1]:(((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v2_zmodul01(X1))&v3_zmodul01(X1))&v4_zmodul01(X1))&v5_zmodul01(X1))&l1_zmodul01(X1))=>![X2]:((m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((m1_subset_1(X3,u1_struct_0(X1))=>![X4]:((m1_zmodul01(X4,X1)=>((r1_struct_0(X4,X2)&r1_struct_0(X4,X3))=>r1_struct_0(X4,k3_rlvect_1(X1,X2,X3))))))))))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', t36_zmodul01)).
fof(t1_subset, axiom, ![X1, X2]:((r2_hidden(X1,X2)=>m1_subset_1(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', t1_subset)).
fof(dt_k5_zmodul01, axiom, ![X1, X2, X3]:(((((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v2_zmodul01(X1))&v3_zmodul01(X1))&v4_zmodul01(X1))&v5_zmodul01(X1))&l1_zmodul01(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_zmodul01(X3,X1))=>m1_subset_1(k5_zmodul01(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1))))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', dt_k5_zmodul01)).
fof(d12_zmodul01, axiom, ![X1]:(((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v2_zmodul01(X1))&v3_zmodul01(X1))&v4_zmodul01(X1))&v5_zmodul01(X1))&l1_zmodul01(X1))=>![X2]:((m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((m1_zmodul01(X3,X1)=>k5_zmodul01(X1,X2,X3)=a_3_1_zmodul01(X1,X2,X3))))))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', d12_zmodul01)).
fof(t4_subset, axiom, ![X1, X2, X3]:(((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', t4_subset)).
fof(t28_zmodul01, axiom, ![X1]:(((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v2_zmodul01(X1))&v3_zmodul01(X1))&v4_zmodul01(X1))&v5_zmodul01(X1))&l1_zmodul01(X1))=>![X2]:((m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((m1_subset_1(X3,u1_struct_0(X1))=>![X4]:((m1_zmodul01(X4,X1)=>![X5]:((m1_subset_1(X5,u1_struct_0(X4))=>![X6]:((m1_subset_1(X6,u1_struct_0(X4))=>((X5=X2&X6=X3)=>k3_rlvect_1(X4,X5,X6)=k3_rlvect_1(X1,X2,X3)))))))))))))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', t28_zmodul01)).
fof(t25_zmodul01, axiom, ![X1]:(((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v2_zmodul01(X1))&v3_zmodul01(X1))&v4_zmodul01(X1))&v5_zmodul01(X1))&l1_zmodul01(X1))=>![X2]:((m1_zmodul01(X2,X1)=>![X3]:((m1_subset_1(X3,u1_struct_0(X2))=>m1_subset_1(X3,u1_struct_0(X1)))))))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', t25_zmodul01)).
fof(l78_zmodul01, conjecture, ![X1]:(((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v2_zmodul01(X1))&v3_zmodul01(X1))&v4_zmodul01(X1))&v5_zmodul01(X1))&l1_zmodul01(X1))=>![X2]:((m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((m1_zmodul01(X3,X1)=>(r1_struct_0(X3,X2)<=>k5_zmodul01(X1,X2,X3)=u1_struct_0(X3)))))))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', l78_zmodul01)).
fof(dt_m1_zmodul01, axiom, ![X1]:(((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v2_zmodul01(X1))&v3_zmodul01(X1))&v4_zmodul01(X1))&v5_zmodul01(X1))&l1_zmodul01(X1))=>![X2]:((m1_zmodul01(X2,X1)=>(((((((((~(v2_struct_0(X2))&v13_algstr_0(X2))&v2_rlvect_1(X2))&v3_rlvect_1(X2))&v4_rlvect_1(X2))&v2_zmodul01(X2))&v3_zmodul01(X2))&v4_zmodul01(X2))&v5_zmodul01(X2))&l1_zmodul01(X2)))))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', dt_m1_zmodul01)).
fof(d3_tarski, axiom, ![X1, X2]:((r1_tarski(X1,X2)<=>![X3]:((r2_hidden(X3,X1)=>r2_hidden(X3,X2))))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', d3_tarski)).
fof(dt_l1_algstr_0, axiom, ![X1]:((l1_algstr_0(X1)=>l1_struct_0(X1))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', dt_l1_algstr_0)).
fof(t3_subset, axiom, ![X1, X2]:((m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', t3_subset)).
fof(d10_xboole_0, axiom, ![X1, X2]:((X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1)))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', d10_xboole_0)).
fof(t31_zmodul01, axiom, ![X1]:(((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v2_zmodul01(X1))&v3_zmodul01(X1))&v4_zmodul01(X1))&v5_zmodul01(X1))&l1_zmodul01(X1))=>![X2]:((m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((m1_subset_1(X3,u1_struct_0(X1))=>![X4]:((m1_zmodul01(X4,X1)=>![X5]:((m1_subset_1(X5,u1_struct_0(X4))=>![X6]:((m1_subset_1(X6,u1_struct_0(X4))=>((X5=X2&X6=X3)=>k5_algstr_0(X4,X5,X6)=k5_algstr_0(X1,X2,X3)))))))))))))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', t31_zmodul01)).
fof(t7_boole, axiom, ![X1, X2]:(~((r2_hidden(X1,X2)&v1_xboole_0(X2)))), file('/scratch/mptp/hardpred/preds__0.005/l78_zmodul01', t7_boole)).
fof(c_0_33, plain, ![X169, X170, X171]:((~v2_rlvect_1(X169)|~l1_algstr_0(X169)|~m1_subset_1(X170,u1_struct_0(X169))|~m1_subset_1(X171,u1_struct_0(X169))|k3_rlvect_1(X169,X170,X171)=k1_algstr_0(X169,X170,X171))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k3_rlvect_1])])).
fof(c_0_34, plain, ![X172, X173, X174]:((~v2_rlvect_1(X172)|~l1_algstr_0(X172)|~m1_subset_1(X173,u1_struct_0(X172))|~m1_subset_1(X174,u1_struct_0(X172))|k3_rlvect_1(X172,X173,X174)=k3_rlvect_1(X172,X174,X173))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k3_rlvect_1])])).
cnf(c_0_35, plain, (k3_rlvect_1(X1,X2,X3)=k1_algstr_0(X1,X2,X3)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_36, plain, (k3_rlvect_1(X1,X2,X3)=k3_rlvect_1(X1,X3,X2)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_34])).
fof(c_0_37, plain, ![X146, X147, X148, X149]:(((~v3_rlvect_1(X146)|(~m1_subset_1(X147,u1_struct_0(X146))|(~m1_subset_1(X148,u1_struct_0(X146))|(~m1_subset_1(X149,u1_struct_0(X146))|k1_algstr_0(X146,k1_algstr_0(X146,X147,X148),X149)=k1_algstr_0(X146,X147,k1_algstr_0(X146,X148,X149)))))|~l1_algstr_0(X146))&((m1_subset_1(esk5_1(X146),u1_struct_0(X146))|v3_rlvect_1(X146)|~l1_algstr_0(X146))&((m1_subset_1(esk6_1(X146),u1_struct_0(X146))|v3_rlvect_1(X146)|~l1_algstr_0(X146))&((m1_subset_1(esk7_1(X146),u1_struct_0(X146))|v3_rlvect_1(X146)|~l1_algstr_0(X146))&(k1_algstr_0(X146,k1_algstr_0(X146,esk5_1(X146),esk6_1(X146)),esk7_1(X146))!=k1_algstr_0(X146,esk5_1(X146),k1_algstr_0(X146,esk6_1(X146),esk7_1(X146)))|v3_rlvect_1(X146)|~l1_algstr_0(X146))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_rlvect_1])])])])])).
fof(c_0_38, plain, ![X154, X155, X156]:((~l2_algstr_0(X154)|(~m1_subset_1(X155,u1_struct_0(X154))|(~m1_subset_1(X156,u1_struct_0(X154))|k5_algstr_0(X154,X155,X156)=k1_algstr_0(X154,X155,k4_algstr_0(X154,X156)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d14_algstr_0])])])).
cnf(c_0_39, plain, (k3_rlvect_1(X1,X2,X3)=k1_algstr_0(X1,X2,X3)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_35])).
cnf(c_0_40, plain, (k3_rlvect_1(X1,X2,X3)=k3_rlvect_1(X1,X3,X2)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_36])).
fof(c_0_41, plain, ![X1]:((((((~v2_struct_0(X1)&v13_algstr_0(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&l2_algstr_0(X1))=>![X2]:((m1_subset_1(X2,u1_struct_0(X1))=>(k1_algstr_0(X1,X2,k4_struct_0(X1))=X2&k1_algstr_0(X1,k4_struct_0(X1),X2)=X2))))), inference(fof_simplification,[status(thm)],[t4_rlvect_1])).
fof(c_0_42, plain, ![X1]:((((((~v2_struct_0(X1)&v13_algstr_0(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&l2_algstr_0(X1))=>![X2]:((m1_subset_1(X2,u1_struct_0(X1))=>k5_algstr_0(X1,X2,X2)=k4_struct_0(X1))))), inference(fof_simplification,[status(thm)],[t15_rlvect_1])).
cnf(c_0_43, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|~v3_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_44, plain, (k5_algstr_0(X1,X2,X3)=k1_algstr_0(X1,X2,k4_algstr_0(X1,X3))|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_45, plain, ![X144, X145]:((~l2_algstr_0(X144)|~m1_subset_1(X145,u1_struct_0(X144))|m1_subset_1(k4_algstr_0(X144,X145),u1_struct_0(X144)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_algstr_0])])).
fof(c_0_46, plain, ![X153]:(((l2_struct_0(X153)|~l2_algstr_0(X153))&(l1_algstr_0(X153)|~l2_algstr_0(X153)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l2_algstr_0])])])).
cnf(c_0_47, plain, (k3_rlvect_1(X1,X3,X2)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_40])])).
fof(c_0_48, plain, ![X124, X125]:(((k1_algstr_0(X124,X125,k4_struct_0(X124))=X125|~m1_subset_1(X125,u1_struct_0(X124))|(v2_struct_0(X124)|~v13_algstr_0(X124)|~v3_rlvect_1(X124)|~v4_rlvect_1(X124)|~l2_algstr_0(X124)))&(k1_algstr_0(X124,k4_struct_0(X124),X125)=X125|~m1_subset_1(X125,u1_struct_0(X124))|(v2_struct_0(X124)|~v13_algstr_0(X124)|~v3_rlvect_1(X124)|~v4_rlvect_1(X124)|~l2_algstr_0(X124))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_41])])])])).
fof(c_0_49, plain, ![X137, X138, X139]:((~l2_algstr_0(X137)|~m1_subset_1(X138,u1_struct_0(X137))|~m1_subset_1(X139,u1_struct_0(X137))|m1_subset_1(k5_algstr_0(X137,X138,X139),u1_struct_0(X137)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_algstr_0])])).
fof(c_0_50, plain, ![X126, X127]:((v2_struct_0(X126)|~v13_algstr_0(X126)|~v3_rlvect_1(X126)|~v4_rlvect_1(X126)|~l2_algstr_0(X126)|(~m1_subset_1(X127,u1_struct_0(X126))|k5_algstr_0(X126,X127,X127)=k4_struct_0(X126)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_42])])])).
cnf(c_0_51, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|~v3_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_43])).
cnf(c_0_52, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X3))=k5_algstr_0(X1,X2,X3)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_44])).
cnf(c_0_53, plain, (m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_54, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_55, plain, (k1_algstr_0(X1,X3,X2)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_47])])).
cnf(c_0_56, plain, (k1_algstr_0(X1,k4_struct_0(X1),X2)=X2|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_57, plain, (m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_58, plain, (v2_struct_0(X1)|k5_algstr_0(X1,X2,X2)=k4_struct_0(X1)|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_59, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X3),X4))=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X4)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52])])).
cnf(c_0_60, plain, (m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_53])).
cnf(c_0_61, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_54])).
cnf(c_0_62, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),X2)=k5_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_55])])).
cnf(c_0_63, plain, (k1_algstr_0(X1,k4_struct_0(X1),X2)=X2|v2_struct_0(X1)|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_56])).
cnf(c_0_64, plain, (m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_57])).
cnf(c_0_65, plain, (k5_algstr_0(X1,X2,X2)=k4_struct_0(X1)|v2_struct_0(X1)|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_58])).
fof(c_0_66, plain, ![X196]:(m1_subset_1(esk9_1(X196),X196)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[existence_m1_subset_1])])).
cnf(c_0_67, plain, (k1_algstr_0(X1,X2,k4_struct_0(X1))=X2|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_68, plain, ![X206, X207, X208]:((~l1_algstr_0(X206)|~m1_subset_1(X207,u1_struct_0(X206))|~m1_subset_1(X208,u1_struct_0(X206))|m1_subset_1(k1_algstr_0(X206,X207,X208),u1_struct_0(X206)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_algstr_0])])).
cnf(c_0_69, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X3),X4))=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X4)|~l2_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_rlvect_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_59, c_0_60]), c_0_61])).
cnf(c_0_70, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),X3)=k5_algstr_0(X1,X3,X2)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v2_rlvect_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_62, c_0_60]), c_0_61])).
cnf(c_0_71, plain, (k5_algstr_0(X1,k4_struct_0(X1),X2)=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_52])])).
cnf(c_0_72, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65])])).
cnf(c_0_73, plain, (m1_subset_1(esk9_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_74, plain, (k1_algstr_0(X1,X2,k4_struct_0(X1))=X2|v2_struct_0(X1)|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_67])).
cnf(c_0_75, plain, (m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_76, plain, (k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X4)=k1_algstr_0(X1,X2,k5_algstr_0(X1,X4,X3))|~l2_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70])])).
cnf(c_0_77, plain, (k5_algstr_0(X1,k4_struct_0(X1),X2)=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_71, c_0_72]), c_0_60])).
cnf(c_0_78, plain, (m1_subset_1(esk9_1(X1),X1)), inference(evalgc,[status(thm)],[c_0_73])).
cnf(c_0_79, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),X3))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_74])])).
cnf(c_0_80, plain, (k1_algstr_0(X1,X4,k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_55])])).
cnf(c_0_81, plain, (m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_75])).
cnf(c_0_82, plain, (k1_algstr_0(X1,k4_struct_0(X1),k5_algstr_0(X1,X3,X2))=k1_algstr_0(X1,k4_algstr_0(X1,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77])])).
cnf(c_0_83, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_78])])).
fof(c_0_84, plain, ![X1, X2, X3, X4]:(((((((((((((~v2_struct_0(X2)&v13_algstr_0(X2))&v2_rlvect_1(X2))&v3_rlvect_1(X2))&v4_rlvect_1(X2))&v2_zmodul01(X2))&v3_zmodul01(X2))&v4_zmodul01(X2))&v5_zmodul01(X2))&l1_zmodul01(X2))&m1_subset_1(X3,u1_struct_0(X2)))&m1_zmodul01(X4,X2))=>(r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))<=>?[X5]:(((m1_subset_1(X5,u1_struct_0(X2))&X1=k3_rlvect_1(X2,X3,X5))&r1_struct_0(X4,X5)))))), inference(fof_simplification,[status(thm)],[fraenkel_a_3_1_zmodul01])).
cnf(c_0_85, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X3,X2),X4)=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|~l1_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_55])])).
cnf(c_0_86, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),X3))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_79, c_0_72]), c_0_61])).
cnf(c_0_87, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))=k1_algstr_0(X1,X3,k1_algstr_0(X1,X4,X2))|~l1_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(csr,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_88, plain, (k1_algstr_0(X1,k4_struct_0(X1),k5_algstr_0(X1,X2,X3))=k1_algstr_0(X1,k4_algstr_0(X1,X3),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_82, c_0_83])).
fof(c_0_89, plain, ![X186, X187, X188, X189, X191]:(((((m1_subset_1(esk8_4(X186,X187,X188,X189),u1_struct_0(X187))|~r2_hidden(X186,a_3_1_zmodul01(X187,X188,X189))|(v2_struct_0(X187)|~v13_algstr_0(X187)|~v2_rlvect_1(X187)|~v3_rlvect_1(X187)|~v4_rlvect_1(X187)|~v2_zmodul01(X187)|~v3_zmodul01(X187)|~v4_zmodul01(X187)|~v5_zmodul01(X187)|~l1_zmodul01(X187)|~m1_subset_1(X188,u1_struct_0(X187))|~m1_zmodul01(X189,X187)))&(X186=k3_rlvect_1(X187,X188,esk8_4(X186,X187,X188,X189))|~r2_hidden(X186,a_3_1_zmodul01(X187,X188,X189))|(v2_struct_0(X187)|~v13_algstr_0(X187)|~v2_rlvect_1(X187)|~v3_rlvect_1(X187)|~v4_rlvect_1(X187)|~v2_zmodul01(X187)|~v3_zmodul01(X187)|~v4_zmodul01(X187)|~v5_zmodul01(X187)|~l1_zmodul01(X187)|~m1_subset_1(X188,u1_struct_0(X187))|~m1_zmodul01(X189,X187))))&(r1_struct_0(X189,esk8_4(X186,X187,X188,X189))|~r2_hidden(X186,a_3_1_zmodul01(X187,X188,X189))|(v2_struct_0(X187)|~v13_algstr_0(X187)|~v2_rlvect_1(X187)|~v3_rlvect_1(X187)|~v4_rlvect_1(X187)|~v2_zmodul01(X187)|~v3_zmodul01(X187)|~v4_zmodul01(X187)|~v5_zmodul01(X187)|~l1_zmodul01(X187)|~m1_subset_1(X188,u1_struct_0(X187))|~m1_zmodul01(X189,X187))))&(~m1_subset_1(X191,u1_struct_0(X187))|X186!=k3_rlvect_1(X187,X188,X191)|~r1_struct_0(X189,X191)|r2_hidden(X186,a_3_1_zmodul01(X187,X188,X189))|(v2_struct_0(X187)|~v13_algstr_0(X187)|~v2_rlvect_1(X187)|~v3_rlvect_1(X187)|~v4_rlvect_1(X187)|~v2_zmodul01(X187)|~v3_zmodul01(X187)|~v4_zmodul01(X187)|~v5_zmodul01(X187)|~l1_zmodul01(X187)|~m1_subset_1(X188,u1_struct_0(X187))|~m1_zmodul01(X189,X187))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_84])])])])])).
fof(c_0_90, plain, ![X177]:((~l1_zmodul01(X177)|l2_algstr_0(X177))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_zmodul01])])).
cnf(c_0_91, plain, (k1_algstr_0(X1,X3,k1_algstr_0(X1,X2,X4))=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|~l1_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_85])])).
cnf(c_0_92, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_63])])).
cnf(c_0_93, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X3,X2))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87])])).
cnf(c_0_94, plain, (k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))=k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_65])])).
cnf(c_0_95, plain, (r2_hidden(X3,a_3_1_zmodul01(X2,X4,X5))|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|X3!=k3_rlvect_1(X2,X4,X1)|~r1_struct_0(X5,X1)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_zmodul01(X5,X2)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_96, plain, (l2_algstr_0(X1)|~l1_zmodul01(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
fof(c_0_97, plain, ![X183, X184]:((~m1_subset_1(X183,X184)|(v1_xboole_0(X184)|r2_hidden(X183,X184)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_98, plain, (k1_algstr_0(X1,X3,k1_algstr_0(X1,k4_algstr_0(X1,X2),X4))=k5_algstr_0(X1,k1_algstr_0(X1,X3,X4),X2)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_91])])).
cnf(c_0_99, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_92, c_0_72]), c_0_61])).
cnf(c_0_100, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,X3,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_93, c_0_83]), c_0_61])).
cnf(c_0_101, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_94])])).
cnf(c_0_102, plain, (v2_struct_0(X2)|r2_hidden(X3,a_3_1_zmodul01(X2,X4,X5))|X3!=k3_rlvect_1(X2,X4,X1)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X5,X2)|~r1_struct_0(X5,X1)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))), inference(evalgc,[status(thm)],[c_0_95])).
cnf(c_0_103, plain, (l2_algstr_0(X1)|~l1_zmodul01(X1)), inference(evalgc,[status(thm)],[c_0_96])).
fof(c_0_104, plain, ![X122, X123]:(((~r1_struct_0(X122,X123)|r2_hidden(X123,u1_struct_0(X122))|~l1_struct_0(X122))&(~r2_hidden(X123,u1_struct_0(X122))|r1_struct_0(X122,X123)|~l1_struct_0(X122)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_struct_0])])])])).
cnf(c_0_105, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_97])).
fof(c_0_106, plain, ![X175]:((~l2_struct_0(X175)|l1_struct_0(X175))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l2_struct_0])])).
cnf(c_0_107, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X3),X4))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X4),X3)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_98, c_0_60]), c_0_81]), c_0_61])).
cnf(c_0_108, plain, (k1_algstr_0(X1,X3,X2)=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100])])).
cnf(c_0_109, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_101, c_0_83])).
cnf(c_0_110, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l1_algstr_0(X2)|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_39])])).
cnf(c_0_111, plain, (l1_algstr_0(X1)|~l1_zmodul01(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_61, c_0_103])])).
cnf(c_0_112, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4)),X5)=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k1_algstr_0(X1,X4,X5))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_51])])).
cnf(c_0_113, plain, (r1_struct_0(X2,X1)|~r2_hidden(X1,u1_struct_0(X2))|~l1_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_114, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_115, plain, (l1_struct_0(X1)|~l2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_116, plain, (l2_struct_0(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_117, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_94])])).
cnf(c_0_118, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_109])])).
fof(c_0_119, plain, ![X1]:(((~v2_struct_0(X1)&l1_struct_0(X1))=>~v1_xboole_0(u1_struct_0(X1)))), inference(fof_simplification,[status(thm)],[fc2_struct_0])).
cnf(c_0_120, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_110, c_0_111])).
fof(c_0_121, plain, ![X1]:(((((((((((~v2_struct_0(X1)&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v2_zmodul01(X1))&v3_zmodul01(X1))&v4_zmodul01(X1))&v5_zmodul01(X1))&l1_zmodul01(X1))=>![X2]:((m1_zmodul01(X2,X1)=>r1_struct_0(X2,k4_struct_0(X1)))))), inference(fof_simplification,[status(thm)],[t33_zmodul01])).
fof(c_0_122, plain, ![X1]:(((((((((((~v2_struct_0(X1)&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v2_zmodul01(X1))&v3_zmodul01(X1))&v4_zmodul01(X1))&v5_zmodul01(X1))&l1_zmodul01(X1))=>![X2]:((m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((m1_subset_1(X3,u1_struct_0(X1))=>![X4]:((m1_zmodul01(X4,X1)=>((r1_struct_0(X4,X2)&r1_struct_0(X4,X3))=>r1_struct_0(X4,k3_rlvect_1(X1,X2,X3))))))))))), inference(fof_simplification,[status(thm)],[t36_zmodul01])).
cnf(c_0_123, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4)),X5)=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k1_algstr_0(X1,X4,X5))|~l1_algstr_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)), inference(csr,[status(thm)],[c_0_112, c_0_81])).
cnf(c_0_124, plain, (X1=k3_rlvect_1(X2,X3,esk8_4(X1,X2,X3,X4))|v2_struct_0(X2)|~r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_zmodul01(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_125, plain, (r1_struct_0(X2,X1)|~l1_struct_0(X2)|~r2_hidden(X1,u1_struct_0(X2))), inference(evalgc,[status(thm)],[c_0_113])).
cnf(c_0_126, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_64])])).
cnf(c_0_127, plain, (l1_struct_0(X1)|~l2_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_115])).
cnf(c_0_128, plain, (l2_struct_0(X1)|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_116])).
cnf(c_0_129, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,X4),X4)=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_117])])).
cnf(c_0_130, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_118, c_0_60])).
fof(c_0_131, plain, ![X182]:((v2_struct_0(X182)|~l1_struct_0(X182)|~v1_xboole_0(u1_struct_0(X182)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_119])])).
cnf(c_0_132, plain, (k1_algstr_0(X1,X3,k1_algstr_0(X1,X2,k4_struct_0(X1)))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_85])])).
cnf(c_0_133, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_52])])).
cnf(c_0_134, plain, (k4_algstr_0(X1,k4_struct_0(X1))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_77])])).
fof(c_0_135, plain, ![X120, X121]:((v2_struct_0(X120)|~v13_algstr_0(X120)|~v2_rlvect_1(X120)|~v3_rlvect_1(X120)|~v4_rlvect_1(X120)|~v2_zmodul01(X120)|~v3_zmodul01(X120)|~v4_zmodul01(X120)|~v5_zmodul01(X120)|~l1_zmodul01(X120)|(~m1_zmodul01(X121,X120)|r1_struct_0(X121,k4_struct_0(X120))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_121])])])).
fof(c_0_136, plain, ![X140, X141, X142, X143]:((v2_struct_0(X140)|~v13_algstr_0(X140)|~v2_rlvect_1(X140)|~v3_rlvect_1(X140)|~v4_rlvect_1(X140)|~v2_zmodul01(X140)|~v3_zmodul01(X140)|~v4_zmodul01(X140)|~v5_zmodul01(X140)|~l1_zmodul01(X140)|(~m1_subset_1(X141,u1_struct_0(X140))|(~m1_subset_1(X142,u1_struct_0(X140))|(~m1_zmodul01(X143,X140)|(~r1_struct_0(X143,X141)|~r1_struct_0(X143,X142)|r1_struct_0(X143,k3_rlvect_1(X140,X141,X142)))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_122])])])).
cnf(c_0_137, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),k1_algstr_0(X1,X3,X4))=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_99])])).
cnf(c_0_138, plain, (k3_rlvect_1(X2,X3,esk8_4(X1,X2,X3,X4))=X1|v2_struct_0(X2)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))), inference(evalgc,[status(thm)],[c_0_124])).
cnf(c_0_139, plain, (m1_subset_1(esk8_4(X1,X2,X3,X4),u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_zmodul01(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_140, plain, (v1_xboole_0(u1_struct_0(X1))|r1_struct_0(X1,k5_algstr_0(X1,X2,X3))|~l2_algstr_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126])])).
cnf(c_0_141, plain, (l1_struct_0(X1)|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_128])])).
cnf(c_0_142, plain, (k5_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X2))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130])])).
cnf(c_0_143, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_131])).
cnf(c_0_144, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_struct_0(X1)))=k1_algstr_0(X1,X3,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_132, c_0_83]), c_0_81]), c_0_61])).
cnf(c_0_145, plain, (m1_subset_1(k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4)),u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_51])])).
cnf(c_0_146, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_133, c_0_60]), c_0_103])).
cnf(c_0_147, plain, (k4_algstr_0(X1,k4_struct_0(X1))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_134, c_0_83])).
cnf(c_0_148, plain, (v2_struct_0(X1)|r1_struct_0(X2,k4_struct_0(X1))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_zmodul01(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_135])).
cnf(c_0_149, plain, (k5_algstr_0(X1,X2,k4_struct_0(X1))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_74])])).
cnf(c_0_150, plain, (v2_struct_0(X1)|r1_struct_0(X4,k3_rlvect_1(X1,X2,X3))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_zmodul01(X4,X1)|~r1_struct_0(X4,X2)|~r1_struct_0(X4,X3)), inference(split_conjunct,[status(thm)],[c_0_136])).
cnf(c_0_151, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),k1_algstr_0(X1,X3,X4))=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_137, c_0_83]), c_0_61])).
cnf(c_0_152, plain, (k1_algstr_0(X1,X2,esk8_4(X3,X1,X2,X4))=X3|v2_struct_0(X1)|~l1_algstr_0(X1)|~r2_hidden(X3,a_3_1_zmodul01(X1,X2,X4))|~m1_zmodul01(X4,X1)|~m1_subset_1(esk8_4(X3,X1,X2,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_138])])).
cnf(c_0_153, plain, (v2_struct_0(X2)|m1_subset_1(esk8_4(X1,X2,X3,X4),u1_struct_0(X2))|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))), inference(evalgc,[status(thm)],[c_0_139])).
cnf(c_0_154, plain, (v1_xboole_0(u1_struct_0(X1))|r1_struct_0(X1,k5_algstr_0(X1,X2,X3))|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[c_0_140, c_0_141])).
cnf(c_0_155, plain, (k5_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X2))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_142, c_0_60])).
cnf(c_0_156, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_143])).
cnf(c_0_157, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_117])])).
cnf(c_0_158, plain, (m1_subset_1(k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4)),u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)), inference(csr,[status(thm)],[c_0_145, c_0_81])).
cnf(c_0_159, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_struct_0(X2))|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_struct_0(X2))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_147])])).
cnf(c_0_160, plain, (v2_struct_0(X1)|r1_struct_0(X2,k4_struct_0(X1))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_zmodul01(X2,X1)), inference(evalgc,[status(thm)],[c_0_148])).
cnf(c_0_161, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,X4),X3)=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X4)|~l2_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_107])])).
cnf(c_0_162, plain, (k5_algstr_0(X1,X2,k4_struct_0(X1))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_149, c_0_83])).
cnf(c_0_163, plain, (v2_struct_0(X1)|r1_struct_0(X4,k3_rlvect_1(X1,X2,X3))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_zmodul01(X4,X1)|~r1_struct_0(X4,X3)|~r1_struct_0(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_150])).
cnf(c_0_164, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),k1_algstr_0(X1,X4,X3))=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_55])])).
cnf(c_0_165, plain, (k1_algstr_0(X1,X2,esk8_4(X3,X1,X2,X4))=X3|v2_struct_0(X1)|~r2_hidden(X3,a_3_1_zmodul01(X1,X2,X4))|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_152, c_0_153]), c_0_111])).
cnf(c_0_166, plain, (v1_xboole_0(u1_struct_0(X1))|r1_struct_0(X1,k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_155])])).
cnf(c_0_167, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_141])])).
cnf(c_0_168, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_157, c_0_83])).
cnf(c_0_169, plain, (m1_subset_1(k1_algstr_0(X1,X2,k5_algstr_0(X1,X3,X4)),u1_struct_0(X1))|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_52])])).
cnf(c_0_170, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_struct_0(X1)))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_74])])).
cnf(c_0_171, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_struct_0(X2))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_159, c_0_83]), c_0_160]), c_0_103])).
cnf(c_0_172, plain, (k5_algstr_0(X1,X2,k4_struct_0(X1))=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_162])])).
cnf(c_0_173, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~l1_algstr_0(X2)|~r1_struct_0(X1,X4)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_39])])).
cnf(c_0_174, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),k1_algstr_0(X1,X3,X4))=k1_algstr_0(X1,k1_algstr_0(X1,X2,X4),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_164, c_0_61])).
cnf(c_0_175, plain, (m1_subset_1(k1_algstr_0(X1,X2,X4),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~r2_hidden(X4,a_3_1_zmodul01(X1,X3,X5))|~m1_zmodul01(X5,X1)|~m1_subset_1(esk8_4(X4,X1,X3,X5),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_165])])).
cnf(c_0_176, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X1,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X3,k4_struct_0(X2))|~m1_zmodul01(X3,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_74])])])).
cnf(c_0_177, plain, (r1_struct_0(X1,k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_166, c_0_83]), c_0_60]), c_0_167])).
cnf(c_0_178, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X3,X2),X3)=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_108])])).
cnf(c_0_179, plain, (m1_subset_1(k1_algstr_0(X1,X2,k5_algstr_0(X1,X3,X4)),u1_struct_0(X1))|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_rlvect_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_169, c_0_60]), c_0_61])).
cnf(c_0_180, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_struct_0(X1)))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_170, c_0_81]), c_0_61])).
cnf(c_0_181, plain, (r1_struct_0(X1,X4)|v2_struct_0(X2)|~r2_hidden(X4,a_3_1_zmodul01(X2,X3,X5))|~r1_struct_0(X1,esk8_4(X4,X2,X3,X5))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_zmodul01(X5,X2)|~m1_subset_1(esk8_4(X4,X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_138])])).
cnf(c_0_182, plain, (r1_struct_0(X1,esk8_4(X2,X3,X4,X1))|v2_struct_0(X3)|~r2_hidden(X2,a_3_1_zmodul01(X3,X4,X1))|~v13_algstr_0(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v2_zmodul01(X3)|~v3_zmodul01(X3)|~v4_zmodul01(X3)|~v5_zmodul01(X3)|~l1_zmodul01(X3)|~m1_subset_1(X4,u1_struct_0(X3))|~m1_zmodul01(X1,X3)), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_183, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k5_algstr_0(X2,X3,X5),X5)|~l2_algstr_0(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_172])])).
cnf(c_0_184, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~r1_struct_0(X1,X4)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_173, c_0_111])).
cnf(c_0_185, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X4),X3)=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_174])])).
cnf(c_0_186, plain, (m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~r2_hidden(X3,a_3_1_zmodul01(X1,X4,X5))|~m1_zmodul01(X5,X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_175, c_0_153]), c_0_111])).
cnf(c_0_187, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X1,X3))|v2_struct_0(X2)|~m1_zmodul01(X3,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_176, c_0_83]), c_0_160]), c_0_103])).
cnf(c_0_188, plain, (r1_struct_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_178])])).
cnf(c_0_189, plain, (m1_subset_1(k1_algstr_0(X1,X2,k4_struct_0(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_65])])).
cnf(c_0_190, plain, (k1_algstr_0(X1,X2,k4_struct_0(X1))=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_99])])).
cnf(c_0_191, plain, (r1_struct_0(X1,X2)|v2_struct_0(X3)|~r2_hidden(X2,a_3_1_zmodul01(X3,X4,X5))|~r1_struct_0(X1,esk8_4(X2,X3,X4,X5))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X3)|~m1_zmodul01(X5,X3)|~m1_subset_1(X4,u1_struct_0(X3))|~l1_zmodul01(X3)|~v5_zmodul01(X3)|~v4_zmodul01(X3)|~v3_zmodul01(X3)|~v2_zmodul01(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v13_algstr_0(X3)), inference(csr,[status(thm)],[c_0_181, c_0_153])).
cnf(c_0_192, plain, (v2_struct_0(X3)|r1_struct_0(X1,esk8_4(X2,X3,X4,X1))|~v13_algstr_0(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v2_zmodul01(X3)|~v3_zmodul01(X3)|~v4_zmodul01(X3)|~v5_zmodul01(X3)|~l1_zmodul01(X3)|~m1_zmodul01(X1,X3)|~m1_subset_1(X4,u1_struct_0(X3))|~r2_hidden(X2,a_3_1_zmodul01(X3,X4,X1))), inference(evalgc,[status(thm)],[c_0_182])).
cnf(c_0_193, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k5_algstr_0(X2,X3,X5),X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_183, c_0_103])).
cnf(c_0_194, plain, (r1_struct_0(X1,k1_algstr_0(X2,k1_algstr_0(X2,X3,X5),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~r1_struct_0(X1,X5)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_185])])).
cnf(c_0_195, plain, (m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_187])])).
cnf(c_0_196, plain, (r1_struct_0(X1,k1_algstr_0(X1,X3,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_100])])).
cnf(c_0_197, plain, (m1_subset_1(k1_algstr_0(X1,X2,k4_struct_0(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_78])])).
cnf(c_0_198, plain, (k1_algstr_0(X1,X2,k4_struct_0(X1))=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_190, c_0_83])).
cnf(c_0_199, plain, (r1_struct_0(X1,X2)|v2_struct_0(X3)|~r2_hidden(X2,a_3_1_zmodul01(X3,X4,X1))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X3)|~m1_subset_1(X4,u1_struct_0(X3))|~l1_zmodul01(X3)|~v5_zmodul01(X3)|~v4_zmodul01(X3)|~v3_zmodul01(X3)|~v2_zmodul01(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v13_algstr_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_192])])).
cnf(c_0_200, plain, (r2_hidden(k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3),a_3_1_zmodul01(X1,X2,X4))|v2_struct_0(X1)|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_193])])).
cnf(c_0_201, plain, (k1_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2),X3)=k1_algstr_0(X1,k4_algstr_0(X1,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,X2),X3),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_63, c_0_69])])).
cnf(c_0_202, plain, (r1_struct_0(X1,k1_algstr_0(X2,k1_algstr_0(X2,X3,X4),X5))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_194, c_0_195]), c_0_103])).
cnf(c_0_203, plain, (k5_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_155])])).
fof(c_0_204, plain, ![X178, X179]:((~r2_hidden(X178,X179)|m1_subset_1(X178,X179))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_205, plain, (r2_hidden(X2,u1_struct_0(X1))|~r1_struct_0(X1,X2)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_206, plain, (r1_struct_0(X1,k1_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3),X2),u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_86])])).
cnf(c_0_207, plain, (m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_198])])).
cnf(c_0_208, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_200])])).
cnf(c_0_209, plain, (k1_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2),X3)=k1_algstr_0(X1,k4_algstr_0(X1,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,X2),X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_201, c_0_83])).
cnf(c_0_210, plain, (r1_struct_0(X1,k1_algstr_0(X2,X4,X6))|v2_struct_0(X2)|~r2_hidden(X4,a_3_1_zmodul01(X2,X3,X5))|~r1_struct_0(X1,esk8_4(X4,X2,X3,X5))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X6))|~m1_zmodul01(X1,X2)|~m1_zmodul01(X5,X2)|~m1_subset_1(esk8_4(X4,X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_165])])).
cnf(c_0_211, plain, (r2_hidden(k3_rlvect_1(X1,X2,X3),a_3_1_zmodul01(X1,X2,X4))|v2_struct_0(X1)|~r1_struct_0(X4,X3)|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_102])])).
cnf(c_0_212, plain, (k4_algstr_0(X1,k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_203])])).
cnf(c_0_213, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_204])).
cnf(c_0_214, plain, (r2_hidden(X2,u1_struct_0(X1))|~l1_struct_0(X1)|~r1_struct_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_205])).
cnf(c_0_215, plain, (r1_struct_0(X1,k1_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_206, c_0_207])).
cnf(c_0_216, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3),X4))=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_86])])).
cnf(c_0_217, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_struct_0(X2))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_algstr_0(X2,X3),X3),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_209])])).
cnf(c_0_218, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~r2_hidden(X3,a_3_1_zmodul01(X2,X5,X6))|~r1_struct_0(X1,esk8_4(X3,X2,X5,X6))|~r1_struct_0(X1,k1_algstr_0(X2,X5,X4))|~m1_zmodul01(X1,X2)|~m1_zmodul01(X6,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_210, c_0_153])).
cnf(c_0_219, plain, (r2_hidden(k1_algstr_0(X1,X3,X2),a_3_1_zmodul01(X1,X2,X4))|v2_struct_0(X1)|~l1_algstr_0(X1)|~r1_struct_0(X4,X3)|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_47])])).
cnf(c_0_220, plain, (k4_algstr_0(X1,k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_212, c_0_60])).
cnf(c_0_221, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_213])).
cnf(c_0_222, plain, (r2_hidden(X1,u1_struct_0(X2))|~l2_algstr_0(X2)|~r1_struct_0(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_141])])).
cnf(c_0_223, plain, (r1_struct_0(X1,k1_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_81])])).
cnf(c_0_224, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3),X4))=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_216, c_0_61])).
cnf(c_0_225, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),X3))|v2_struct_0(X2)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_algstr_0(X2,X3),X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_217, c_0_83]), c_0_160]), c_0_103])).
cnf(c_0_226, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~r2_hidden(X3,a_3_1_zmodul01(X2,X5,X1))|~r1_struct_0(X1,k1_algstr_0(X2,X5,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_192])])).
cnf(c_0_227, plain, (r2_hidden(k1_algstr_0(X1,X2,X3),a_3_1_zmodul01(X1,X3,X4))|v2_struct_0(X1)|~r1_struct_0(X4,X2)|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_219, c_0_111])).
cnf(c_0_228, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),k4_algstr_0(X1,X2))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_220])])).
cnf(c_0_229, plain, (m1_subset_1(X1,u1_struct_0(X2))|~l2_algstr_0(X2)|~r1_struct_0(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_222])])).
cnf(c_0_230, plain, (r1_struct_0(X1,k1_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_223, c_0_207]), c_0_61])).
cnf(c_0_231, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_63])])).
cnf(c_0_232, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,k4_algstr_0(X2,X3)),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_108])])).
cnf(c_0_233, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k1_algstr_0(X2,X3,X5),X5)|~l2_algstr_0(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_162])])).
cnf(c_0_234, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_109])])).
cnf(c_0_235, plain, (r1_struct_0(X1,k1_algstr_0(X2,k1_algstr_0(X2,X3,X4),X5))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X4,X5))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_227])])).
cnf(c_0_236, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),k4_algstr_0(X1,X2))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_228, c_0_60])).
cnf(c_0_237, plain, (r1_struct_0(X1,k1_algstr_0(X2,X5,k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_108])])).
cnf(c_0_238, plain, (m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_230])])).
cnf(c_0_239, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_109])])).
cnf(c_0_240, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,k4_algstr_0(X2,X3)),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_232, c_0_60]), c_0_103])).
cnf(c_0_241, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k1_algstr_0(X2,X3,X5),X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_233, c_0_103])).
cnf(c_0_242, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_algstr_0(X1,X4)))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_52])])).
cnf(c_0_243, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X3,X2),X3)=k1_algstr_0(X1,k4_algstr_0(X1,X2),k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_70])])).
cnf(c_0_244, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),X4))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_236])])).
cnf(c_0_245, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k1_algstr_0(X2,X4,X5)))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X4,X3))|~r1_struct_0(X1,X5)|~m1_zmodul01(X1,X2)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_237, c_0_238]), c_0_103])).
cnf(c_0_246, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_239, c_0_60])).
cnf(c_0_247, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_238])])).
cnf(c_0_248, plain, (r2_hidden(k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3),a_3_1_zmodul01(X1,X2,X4))|v2_struct_0(X1)|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_241])])).
cnf(c_0_249, plain, (k1_algstr_0(X1,esk8_4(X3,X1,X2,X4),X2)=X3|v2_struct_0(X1)|~l1_algstr_0(X1)|~r2_hidden(X3,a_3_1_zmodul01(X1,X2,X4))|~m1_zmodul01(X4,X1)|~m1_subset_1(esk8_4(X3,X1,X2,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_47])])).
fof(c_0_250, plain, ![X1, X2, X3]:(((((((((((((~v2_struct_0(X1)&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v2_zmodul01(X1))&v3_zmodul01(X1))&v4_zmodul01(X1))&v5_zmodul01(X1))&l1_zmodul01(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_zmodul01(X3,X1))=>m1_subset_1(k5_zmodul01(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1))))), inference(fof_simplification,[status(thm)],[dt_k5_zmodul01])).
fof(c_0_251, plain, ![X1]:(((((((((((~v2_struct_0(X1)&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v2_zmodul01(X1))&v3_zmodul01(X1))&v4_zmodul01(X1))&v5_zmodul01(X1))&l1_zmodul01(X1))=>![X2]:((m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((m1_zmodul01(X3,X1)=>k5_zmodul01(X1,X2,X3)=a_3_1_zmodul01(X1,X2,X3))))))), inference(fof_simplification,[status(thm)],[d12_zmodul01])).
cnf(c_0_252, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_algstr_0(X1,X4)))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_rlvect_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_242, c_0_60]), c_0_81]), c_0_61])).
cnf(c_0_253, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),X2)=k1_algstr_0(X1,k4_algstr_0(X1,X3),k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_243, c_0_60])).
cnf(c_0_254, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X4),X3))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_244, c_0_60]), c_0_207]), c_0_103])).
cnf(c_0_255, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_algstr_0(X2,X3)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_246])])).
cnf(c_0_256, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_247, c_0_60]), c_0_103])).
cnf(c_0_257, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_248])])).
cnf(c_0_258, plain, (k1_algstr_0(X1,esk8_4(X2,X1,X3,X4),X3)=X2|v2_struct_0(X1)|~r2_hidden(X2,a_3_1_zmodul01(X1,X3,X4))|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_249, c_0_153]), c_0_111])).
fof(c_0_259, plain, ![X166, X167, X168]:((v2_struct_0(X166)|~v13_algstr_0(X166)|~v2_rlvect_1(X166)|~v3_rlvect_1(X166)|~v4_rlvect_1(X166)|~v2_zmodul01(X166)|~v3_zmodul01(X166)|~v4_zmodul01(X166)|~v5_zmodul01(X166)|~l1_zmodul01(X166)|~m1_subset_1(X167,u1_struct_0(X166))|~m1_zmodul01(X168,X166)|m1_subset_1(k5_zmodul01(X166,X167,X168),k1_zfmisc_1(u1_struct_0(X166))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_250])])).
fof(c_0_260, plain, ![X163, X164, X165]:((v2_struct_0(X163)|~v13_algstr_0(X163)|~v2_rlvect_1(X163)|~v3_rlvect_1(X163)|~v4_rlvect_1(X163)|~v2_zmodul01(X163)|~v3_zmodul01(X163)|~v4_zmodul01(X163)|~v5_zmodul01(X163)|~l1_zmodul01(X163)|(~m1_subset_1(X164,u1_struct_0(X163))|(~m1_zmodul01(X165,X163)|k5_zmodul01(X163,X164,X165)=a_3_1_zmodul01(X163,X164,X165))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_251])])])).
cnf(c_0_261, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_algstr_0(X2,X3),X4),u1_struct_0(X2))|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_209])])).
cnf(c_0_262, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)=k1_algstr_0(X1,X2,k5_algstr_0(X1,X3,X4))|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_rlvect_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_52])])).
cnf(c_0_263, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X3,X2),X3)=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_253])])).
cnf(c_0_264, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_70])])).
cnf(c_0_265, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_255, c_0_60]), c_0_256]), c_0_103])).
cnf(c_0_266, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~r2_hidden(X3,a_3_1_zmodul01(X2,X4,X5))|~r1_struct_0(X1,esk8_4(X3,X2,X4,X5))|~m1_zmodul01(X1,X2)|~m1_zmodul01(X5,X2)|~m1_subset_1(esk8_4(X3,X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_258])])).
fof(c_0_267, plain, ![X192, X193, X194]:((~r2_hidden(X192,X193)|~m1_subset_1(X193,k1_zfmisc_1(X194))|m1_subset_1(X192,X194))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_268, plain, (v2_struct_0(X1)|m1_subset_1(k5_zmodul01(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_zmodul01(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_259])).
cnf(c_0_269, plain, (v2_struct_0(X1)|k5_zmodul01(X1,X2,X3)=a_3_1_zmodul01(X1,X2,X3)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_zmodul01(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_260])).
cnf(c_0_270, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_algstr_0(X2,X3),X4),u1_struct_0(X2))|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_261, c_0_103])).
cnf(c_0_271, plain, (r1_struct_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_52])])).
cnf(c_0_272, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k5_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_109])])).
cnf(c_0_273, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),X2)=k4_algstr_0(X1,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_263, c_0_60])).
cnf(c_0_274, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_264, c_0_103])).
cnf(c_0_275, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_78])])).
cnf(c_0_276, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~r2_hidden(X3,a_3_1_zmodul01(X2,X4,X5))|~r1_struct_0(X1,esk8_4(X3,X2,X4,X5))|~m1_zmodul01(X1,X2)|~m1_zmodul01(X5,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_266, c_0_153])).
cnf(c_0_277, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_267])).
cnf(c_0_278, plain, (v2_struct_0(X1)|m1_subset_1(k5_zmodul01(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_268])).
cnf(c_0_279, plain, (k5_zmodul01(X1,X2,X3)=a_3_1_zmodul01(X1,X2,X3)|v2_struct_0(X1)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_269])).
cnf(c_0_280, plain, (r1_struct_0(X1,k5_algstr_0(X2,X4,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(k5_algstr_0(X2,X4,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_70])])).
cnf(c_0_281, plain, (r1_struct_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_271, c_0_83]), c_0_60])).
cnf(c_0_282, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k5_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_272, c_0_60])).
cnf(c_0_283, plain, (k5_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3),k1_algstr_0(X1,X2,X3))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_178])])).
cnf(c_0_284, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_86])])).
cnf(c_0_285, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_274, c_0_275])])).
cnf(c_0_286, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~r2_hidden(X3,a_3_1_zmodul01(X2,X4,X1))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_276, c_0_192])])).
cnf(c_0_287, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_277])).
cnf(c_0_288, plain, (m1_subset_1(a_3_1_zmodul01(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_279])])).
cnf(c_0_289, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_280, c_0_64]), c_0_103])).
cnf(c_0_290, plain, (m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_281])])).
cnf(c_0_291, plain, (k5_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X2))=k5_algstr_0(X1,k1_algstr_0(X1,X3,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_108])])).
cnf(c_0_292, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),X2)=k5_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_282])])).
cnf(c_0_293, plain, (k5_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),k1_algstr_0(X1,X3,X2))=k4_algstr_0(X1,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_283, c_0_238])).
cnf(c_0_294, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_284, c_0_103])).
cnf(c_0_295, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r2_hidden(X3,a_3_1_zmodul01(X2,X4,X1))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_286])])).
cnf(c_0_296, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_287, c_0_288])])).
cnf(c_0_297, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_290])])).
cnf(c_0_298, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_291])])).
cnf(c_0_299, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),X2)=k5_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_292, c_0_64])).
cnf(c_0_300, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X3,X2))=k4_algstr_0(X1,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_293, c_0_63])])).
cnf(c_0_301, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_294, c_0_207])])).
cnf(c_0_302, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r2_hidden(X3,a_3_1_zmodul01(X2,X4,X1))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_295, c_0_296])).
cnf(c_0_303, plain, (r2_hidden(k1_algstr_0(X1,X2,X3),a_3_1_zmodul01(X1,X2,X4))|v2_struct_0(X1)|~r1_struct_0(X4,X3)|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_120])])).
cnf(c_0_304, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_297, c_0_103])).
cnf(c_0_305, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_298, c_0_103])).
cnf(c_0_306, plain, (k5_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X3,X2))=k5_algstr_0(X1,k4_algstr_0(X1,X3),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_299, c_0_300])])).
cnf(c_0_307, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_301, c_0_103])).
cnf(c_0_308, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~r1_struct_0(X1,X3)|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_303])])).
cnf(c_0_309, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_algstr_0(X2,X4)))|v2_struct_0(X2)|~r1_struct_0(X1,X3)|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_304, c_0_305])])).
cnf(c_0_310, plain, (k5_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k4_algstr_0(X1,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_306, c_0_238])).
fof(c_0_311, plain, ![X1]:(((((((((((~v2_struct_0(X1)&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v2_zmodul01(X1))&v3_zmodul01(X1))&v4_zmodul01(X1))&v5_zmodul01(X1))&l1_zmodul01(X1))=>![X2]:((m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((m1_subset_1(X3,u1_struct_0(X1))=>![X4]:((m1_zmodul01(X4,X1)=>![X5]:((m1_subset_1(X5,u1_struct_0(X4))=>![X6]:((m1_subset_1(X6,u1_struct_0(X4))=>((X5=X2&X6=X3)=>k3_rlvect_1(X4,X5,X6)=k3_rlvect_1(X1,X2,X3)))))))))))))), inference(fof_simplification,[status(thm)],[t28_zmodul01])).
cnf(c_0_312, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,X5)|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_307, c_0_308])])).
cnf(c_0_313, plain, (r1_struct_0(X1,k4_algstr_0(X2,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_algstr_0(X2,X3),X4),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_273])])).
cnf(c_0_314, plain, (m1_subset_1(k5_algstr_0(X1,k4_algstr_0(X1,X2),X3),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_310])])).
fof(c_0_315, plain, ![X108, X109, X110, X111, X112, X113]:((v2_struct_0(X108)|~v13_algstr_0(X108)|~v2_rlvect_1(X108)|~v3_rlvect_1(X108)|~v4_rlvect_1(X108)|~v2_zmodul01(X108)|~v3_zmodul01(X108)|~v4_zmodul01(X108)|~v5_zmodul01(X108)|~l1_zmodul01(X108)|(~m1_subset_1(X109,u1_struct_0(X108))|(~m1_subset_1(X110,u1_struct_0(X108))|(~m1_zmodul01(X111,X108)|(~m1_subset_1(X112,u1_struct_0(X111))|(~m1_subset_1(X113,u1_struct_0(X111))|(X112!=X109|X113!=X110|k3_rlvect_1(X111,X112,X113)=k3_rlvect_1(X108,X109,X110))))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_311])])])).
fof(c_0_316, plain, ![X1]:(((((((((((~v2_struct_0(X1)&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v2_zmodul01(X1))&v3_zmodul01(X1))&v4_zmodul01(X1))&v5_zmodul01(X1))&l1_zmodul01(X1))=>![X2]:((m1_zmodul01(X2,X1)=>![X3]:((m1_subset_1(X3,u1_struct_0(X2))=>m1_subset_1(X3,u1_struct_0(X1)))))))), inference(fof_simplification,[status(thm)],[t25_zmodul01])).
cnf(c_0_317, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,X5)|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_312, c_0_195])).
cnf(c_0_318, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X4),X3))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_algstr_0(X2,X4),X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_313, c_0_60]), c_0_103])).
cnf(c_0_319, plain, (m1_subset_1(k5_algstr_0(X1,k4_algstr_0(X1,X2),X3),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_314, c_0_83]), c_0_238])).
cnf(c_0_320, plain, (v2_struct_0(X1)|k3_rlvect_1(X4,X5,X6)=k3_rlvect_1(X1,X2,X3)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_zmodul01(X4,X1)|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))|X5!=X2|X6!=X3), inference(split_conjunct,[status(thm)],[c_0_315])).
fof(c_0_321, plain, ![X134, X135, X136]:((v2_struct_0(X134)|~v13_algstr_0(X134)|~v2_rlvect_1(X134)|~v3_rlvect_1(X134)|~v4_rlvect_1(X134)|~v2_zmodul01(X134)|~v3_zmodul01(X134)|~v4_zmodul01(X134)|~v5_zmodul01(X134)|~l1_zmodul01(X134)|(~m1_zmodul01(X135,X134)|(~m1_subset_1(X136,u1_struct_0(X135))|m1_subset_1(X136,u1_struct_0(X134)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_316])])])).
fof(c_0_322, negated_conjecture, ~(![X1]:(((((((((((~v2_struct_0(X1)&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v2_zmodul01(X1))&v3_zmodul01(X1))&v4_zmodul01(X1))&v5_zmodul01(X1))&l1_zmodul01(X1))=>![X2]:((m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((m1_zmodul01(X3,X1)=>(r1_struct_0(X3,X2)<=>k5_zmodul01(X1,X2,X3)=u1_struct_0(X3))))))))), inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[l78_zmodul01])])).
fof(c_0_323, plain, ![X1]:(((((((((((~v2_struct_0(X1)&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v2_zmodul01(X1))&v3_zmodul01(X1))&v4_zmodul01(X1))&v5_zmodul01(X1))&l1_zmodul01(X1))=>![X2]:((m1_zmodul01(X2,X1)=>(((((((((~v2_struct_0(X2)&v13_algstr_0(X2))&v2_rlvect_1(X2))&v3_rlvect_1(X2))&v4_rlvect_1(X2))&v2_zmodul01(X2))&v3_zmodul01(X2))&v4_zmodul01(X2))&v5_zmodul01(X2))&l1_zmodul01(X2)))))), inference(fof_simplification,[status(thm)],[dt_m1_zmodul01])).
cnf(c_0_324, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k5_algstr_0(X2,X6,X5))|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_70])])).
cnf(c_0_325, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X4),X3))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_319])])).
cnf(c_0_326, plain, (k3_rlvect_1(X1,X2,X3)=k3_rlvect_1(X4,X5,X6)|v2_struct_0(X1)|X3!=X6|X2!=X5|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_zmodul01(X4,X1)|~m1_subset_1(X6,u1_struct_0(X4))|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_320])).
cnf(c_0_327, plain, (v2_struct_0(X1)|m1_subset_1(X3,u1_struct_0(X1))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_zmodul01(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(split_conjunct,[status(thm)],[c_0_321])).
fof(c_0_328, negated_conjecture, ((((((((((~v2_struct_0(esk1_0)&v13_algstr_0(esk1_0))&v2_rlvect_1(esk1_0))&v3_rlvect_1(esk1_0))&v4_rlvect_1(esk1_0))&v2_zmodul01(esk1_0))&v3_zmodul01(esk1_0))&v4_zmodul01(esk1_0))&v5_zmodul01(esk1_0))&l1_zmodul01(esk1_0))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_zmodul01(esk3_0,esk1_0)&((~r1_struct_0(esk3_0,esk2_0)|k5_zmodul01(esk1_0,esk2_0,esk3_0)!=u1_struct_0(esk3_0))&(r1_struct_0(esk3_0,esk2_0)|k5_zmodul01(esk1_0,esk2_0,esk3_0)=u1_struct_0(esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_322])])])).
fof(c_0_329, plain, ![X180, X181]:(((((((((((~v2_struct_0(X181)|~m1_zmodul01(X181,X180)|(v2_struct_0(X180)|~v13_algstr_0(X180)|~v2_rlvect_1(X180)|~v3_rlvect_1(X180)|~v4_rlvect_1(X180)|~v2_zmodul01(X180)|~v3_zmodul01(X180)|~v4_zmodul01(X180)|~v5_zmodul01(X180)|~l1_zmodul01(X180)))&(v13_algstr_0(X181)|~m1_zmodul01(X181,X180)|(v2_struct_0(X180)|~v13_algstr_0(X180)|~v2_rlvect_1(X180)|~v3_rlvect_1(X180)|~v4_rlvect_1(X180)|~v2_zmodul01(X180)|~v3_zmodul01(X180)|~v4_zmodul01(X180)|~v5_zmodul01(X180)|~l1_zmodul01(X180))))&(v2_rlvect_1(X181)|~m1_zmodul01(X181,X180)|(v2_struct_0(X180)|~v13_algstr_0(X180)|~v2_rlvect_1(X180)|~v3_rlvect_1(X180)|~v4_rlvect_1(X180)|~v2_zmodul01(X180)|~v3_zmodul01(X180)|~v4_zmodul01(X180)|~v5_zmodul01(X180)|~l1_zmodul01(X180))))&(v3_rlvect_1(X181)|~m1_zmodul01(X181,X180)|(v2_struct_0(X180)|~v13_algstr_0(X180)|~v2_rlvect_1(X180)|~v3_rlvect_1(X180)|~v4_rlvect_1(X180)|~v2_zmodul01(X180)|~v3_zmodul01(X180)|~v4_zmodul01(X180)|~v5_zmodul01(X180)|~l1_zmodul01(X180))))&(v4_rlvect_1(X181)|~m1_zmodul01(X181,X180)|(v2_struct_0(X180)|~v13_algstr_0(X180)|~v2_rlvect_1(X180)|~v3_rlvect_1(X180)|~v4_rlvect_1(X180)|~v2_zmodul01(X180)|~v3_zmodul01(X180)|~v4_zmodul01(X180)|~v5_zmodul01(X180)|~l1_zmodul01(X180))))&(v2_zmodul01(X181)|~m1_zmodul01(X181,X180)|(v2_struct_0(X180)|~v13_algstr_0(X180)|~v2_rlvect_1(X180)|~v3_rlvect_1(X180)|~v4_rlvect_1(X180)|~v2_zmodul01(X180)|~v3_zmodul01(X180)|~v4_zmodul01(X180)|~v5_zmodul01(X180)|~l1_zmodul01(X180))))&(v3_zmodul01(X181)|~m1_zmodul01(X181,X180)|(v2_struct_0(X180)|~v13_algstr_0(X180)|~v2_rlvect_1(X180)|~v3_rlvect_1(X180)|~v4_rlvect_1(X180)|~v2_zmodul01(X180)|~v3_zmodul01(X180)|~v4_zmodul01(X180)|~v5_zmodul01(X180)|~l1_zmodul01(X180))))&(v4_zmodul01(X181)|~m1_zmodul01(X181,X180)|(v2_struct_0(X180)|~v13_algstr_0(X180)|~v2_rlvect_1(X180)|~v3_rlvect_1(X180)|~v4_rlvect_1(X180)|~v2_zmodul01(X180)|~v3_zmodul01(X180)|~v4_zmodul01(X180)|~v5_zmodul01(X180)|~l1_zmodul01(X180))))&(v5_zmodul01(X181)|~m1_zmodul01(X181,X180)|(v2_struct_0(X180)|~v13_algstr_0(X180)|~v2_rlvect_1(X180)|~v3_rlvect_1(X180)|~v4_rlvect_1(X180)|~v2_zmodul01(X180)|~v3_zmodul01(X180)|~v4_zmodul01(X180)|~v5_zmodul01(X180)|~l1_zmodul01(X180))))&(l1_zmodul01(X181)|~m1_zmodul01(X181,X180)|(v2_struct_0(X180)|~v13_algstr_0(X180)|~v2_rlvect_1(X180)|~v3_rlvect_1(X180)|~v4_rlvect_1(X180)|~v2_zmodul01(X180)|~v3_zmodul01(X180)|~v4_zmodul01(X180)|~v5_zmodul01(X180)|~l1_zmodul01(X180))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_323])])])])).
cnf(c_0_330, plain, (k1_algstr_0(X1,k4_struct_0(X1),X3)=X3|v2_struct_0(X1)|~l2_algstr_0(X1)|~r2_hidden(X3,a_3_1_zmodul01(X1,X2,X4))|~m1_zmodul01(X4,X1)|~m1_subset_1(esk8_4(X3,X1,X2,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_165])])).
cnf(c_0_331, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_324, c_0_60]), c_0_103])).
cnf(c_0_332, plain, (k1_algstr_0(X1,X2,k5_algstr_0(X1,X3,X2))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_65])])).
cnf(c_0_333, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X4),X3))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_325, c_0_103])).
cnf(c_0_334, plain, (k3_rlvect_1(X1,X2,X3)=k3_rlvect_1(X4,X2,X3)|v2_struct_0(X1)|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_326])])])).
cnf(c_0_335, plain, (v2_struct_0(X1)|m1_subset_1(X3,u1_struct_0(X1))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_zmodul01(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))), inference(evalgc,[status(thm)],[c_0_327])).
cnf(c_0_336, negated_conjecture, (m1_zmodul01(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_328])).
cnf(c_0_337, negated_conjecture, (l1_zmodul01(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_328])).
cnf(c_0_338, negated_conjecture, (v5_zmodul01(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_328])).
cnf(c_0_339, negated_conjecture, (v4_zmodul01(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_328])).
cnf(c_0_340, negated_conjecture, (v3_zmodul01(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_328])).
cnf(c_0_341, negated_conjecture, (v2_zmodul01(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_328])).
cnf(c_0_342, negated_conjecture, (v4_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_328])).
cnf(c_0_343, negated_conjecture, (v3_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_328])).
cnf(c_0_344, negated_conjecture, (v2_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_328])).
cnf(c_0_345, negated_conjecture, (v13_algstr_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_328])).
cnf(c_0_346, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_328])).
cnf(c_0_347, plain, (v2_rlvect_1(X1)|v2_struct_0(X2)|~m1_zmodul01(X1,X2)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)), inference(split_conjunct,[status(thm)],[c_0_329])).
cnf(c_0_348, plain, (k1_algstr_0(X1,k4_struct_0(X1),X2)=X2|v2_struct_0(X1)|~r2_hidden(X2,a_3_1_zmodul01(X1,X3,X4))|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_330, c_0_153]), c_0_103])).
cnf(c_0_349, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k4_struct_0(X2),X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_331, c_0_332])])).
cnf(c_0_350, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,k4_struct_0(X2),X4),X3))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_333, c_0_220])])).
cnf(c_0_351, plain, (k5_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),k1_algstr_0(X1,X2,X3))=k4_algstr_0(X1,X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_293, c_0_55])])).
cnf(c_0_352, plain, (k5_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))=k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_77])])).
cnf(c_0_353, plain, (k3_rlvect_1(X1,X2,X3)=k3_rlvect_1(X4,X2,X3)|v2_struct_0(X1)|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_334, c_0_335]), c_0_335])).
cnf(c_0_354, negated_conjecture, (m1_zmodul01(esk3_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_336])).
cnf(c_0_355, negated_conjecture, (l1_zmodul01(esk1_0)), inference(evalgc,[status(thm)],[c_0_337])).
cnf(c_0_356, negated_conjecture, (v5_zmodul01(esk1_0)), inference(evalgc,[status(thm)],[c_0_338])).
cnf(c_0_357, negated_conjecture, (v4_zmodul01(esk1_0)), inference(evalgc,[status(thm)],[c_0_339])).
cnf(c_0_358, negated_conjecture, (v3_zmodul01(esk1_0)), inference(evalgc,[status(thm)],[c_0_340])).
cnf(c_0_359, negated_conjecture, (v2_zmodul01(esk1_0)), inference(evalgc,[status(thm)],[c_0_341])).
cnf(c_0_360, negated_conjecture, (v4_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_342])).
cnf(c_0_361, negated_conjecture, (v3_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_343])).
cnf(c_0_362, negated_conjecture, (v2_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_344])).
cnf(c_0_363, negated_conjecture, (v13_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_345])).
cnf(c_0_364, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_346])).
cnf(c_0_365, plain, (v2_struct_0(X2)|v2_rlvect_1(X1)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X1,X2)), inference(evalgc,[status(thm)],[c_0_347])).
cnf(c_0_366, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_86])])).
cnf(c_0_367, plain, (k1_algstr_0(X1,k4_struct_0(X1),X2)=X2|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_348, c_0_187])])).
cnf(c_0_368, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k4_struct_0(X2),X5)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_349, c_0_103])).
cnf(c_0_369, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,k4_struct_0(X2),X4),X3))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_350, c_0_60]), c_0_103])).
cnf(c_0_370, plain, (k5_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),k1_algstr_0(X1,X2,X3))=k4_algstr_0(X1,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_351, c_0_61])).
cnf(c_0_371, plain, (k5_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))=k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_352, c_0_83])).
cnf(c_0_372, negated_conjecture, (k3_rlvect_1(esk3_0,X1,X2)=k3_rlvect_1(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_353, c_0_354]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_373, negated_conjecture, (v2_rlvect_1(esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_365, c_0_354]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_374, plain, (v4_rlvect_1(X1)|v2_struct_0(X2)|~m1_zmodul01(X1,X2)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)), inference(split_conjunct,[status(thm)],[c_0_329])).
cnf(c_0_375, plain, (v3_rlvect_1(X1)|v2_struct_0(X2)|~m1_zmodul01(X1,X2)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)), inference(split_conjunct,[status(thm)],[c_0_329])).
cnf(c_0_376, plain, (v13_algstr_0(X1)|v2_struct_0(X2)|~m1_zmodul01(X1,X2)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)), inference(split_conjunct,[status(thm)],[c_0_329])).
cnf(c_0_377, plain, (v2_struct_0(X2)|~v2_struct_0(X1)|~m1_zmodul01(X1,X2)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)), inference(split_conjunct,[status(thm)],[c_0_329])).
cnf(c_0_378, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_366, c_0_103])).
cnf(c_0_379, negated_conjecture, (k1_algstr_0(esk1_0,k4_struct_0(esk1_0),X1)=X1|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_354]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_380, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X1)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_368, c_0_63])])])).
cnf(c_0_381, plain, (r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_369, c_0_370])])).
cnf(c_0_382, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X4),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_88])])).
cnf(c_0_383, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),X3)=k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_371, c_0_371])])).
cnf(c_0_384, negated_conjecture, (k3_rlvect_1(esk1_0,X1,X2)=k1_algstr_0(esk3_0,X1,X2)|~l1_algstr_0(esk3_0)|~m1_subset_1(X2,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_372]), c_0_373])])])).
cnf(c_0_385, plain, (v2_struct_0(X2)|v4_rlvect_1(X1)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X1,X2)), inference(evalgc,[status(thm)],[c_0_374])).
cnf(c_0_386, plain, (v2_struct_0(X2)|v3_rlvect_1(X1)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X1,X2)), inference(evalgc,[status(thm)],[c_0_375])).
cnf(c_0_387, plain, (v2_struct_0(X2)|v13_algstr_0(X1)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X1,X2)), inference(evalgc,[status(thm)],[c_0_376])).
cnf(c_0_388, plain, (v2_struct_0(X2)|~v2_struct_0(X1)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X1,X2)), inference(evalgc,[status(thm)],[c_0_377])).
cnf(c_0_389, negated_conjecture, (r1_struct_0(X1,k1_algstr_0(esk1_0,X2,X3))|~r1_struct_0(X1,X3)|~r1_struct_0(X1,X2)|~m1_zmodul01(X1,esk1_0)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_379]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
fof(c_0_390, plain, ![X128, X129, X130, X131, X132]:(((~r1_tarski(X128,X129)|(~r2_hidden(X130,X128)|r2_hidden(X130,X129)))&((r2_hidden(esk4_2(X131,X132),X131)|r1_tarski(X131,X132))&(~r2_hidden(esk4_2(X131,X132),X132)|r1_tarski(X131,X132))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
cnf(c_0_391, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X1)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_380, c_0_103])).
cnf(c_0_392, plain, (r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_381, c_0_238]), c_0_103])).
cnf(c_0_393, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X4,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_382, c_0_64]), c_0_103])).
cnf(c_0_394, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_algstr_0(X1,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_383])])).
fof(c_0_395, plain, ![X195]:((~l1_algstr_0(X195)|l1_struct_0(X195))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_algstr_0])])).
cnf(c_0_396, negated_conjecture, (k1_algstr_0(esk3_0,X1,esk8_4(X2,esk1_0,X1,X3))=X2|~l1_algstr_0(esk3_0)|~r2_hidden(X2,a_3_1_zmodul01(esk1_0,X1,X3))|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(esk8_4(X2,esk1_0,X1,X3),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_384]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_397, negated_conjecture, (v4_rlvect_1(esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_385, c_0_354]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_398, negated_conjecture, (v3_rlvect_1(esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_386, c_0_354]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_399, negated_conjecture, (v13_algstr_0(esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_354]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_400, negated_conjecture, (~v2_struct_0(esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_388, c_0_354]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_401, negated_conjecture, (r1_struct_0(X1,X2)|~r2_hidden(X2,a_3_1_zmodul01(esk1_0,X3,X4))|~r1_struct_0(X1,esk8_4(X2,esk1_0,X3,X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,esk1_0)|~m1_zmodul01(X4,esk1_0)|~m1_subset_1(esk8_4(X2,esk1_0,X3,X4),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_389, c_0_258]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_402, plain, (r2_hidden(esk4_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_390])).
cnf(c_0_403, plain, (r1_tarski(X1,X2)|~r2_hidden(esk4_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_390])).
cnf(c_0_404, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k1_algstr_0(X2,X3,X4),X5))|v2_struct_0(X2)|~r1_struct_0(X5,k4_algstr_0(X2,X4))|~r1_struct_0(X5,k4_algstr_0(X2,X3))|~r1_struct_0(X5,X1)|~m1_zmodul01(X5,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_392])])).
cnf(c_0_405, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_74])])).
cnf(c_0_406, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_algstr_0(X1,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_394, c_0_60])).
cnf(c_0_407, plain, (l1_struct_0(X1)|~l1_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_395])).
cnf(c_0_408, negated_conjecture, (r1_struct_0(esk3_0,X2)|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk3_0)|~r2_hidden(X2,a_3_1_zmodul01(esk1_0,X1,X3))|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(esk8_4(X2,esk1_0,X1,X3),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_396]), c_0_397]), c_0_398]), c_0_373]), c_0_399])]), c_0_400])])).
fof(c_0_409, plain, ![X202, X203]:(((~m1_subset_1(X202,k1_zfmisc_1(X203))|r1_tarski(X202,X203))&(~r1_tarski(X202,X203)|m1_subset_1(X202,k1_zfmisc_1(X203))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_410, negated_conjecture, (r1_struct_0(X1,X2)|~r2_hidden(X2,a_3_1_zmodul01(esk1_0,X3,X1))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,esk1_0)|~m1_subset_1(esk8_4(X2,esk1_0,X3,X1),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_192]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_411, plain, (r1_tarski(X1,X2)|r2_hidden(esk4_2(X1,X2),X1)), inference(evalgc,[status(thm)],[c_0_402])).
cnf(c_0_412, plain, (r1_tarski(X1,X2)|~r2_hidden(esk4_2(X1,X2),X2)), inference(evalgc,[status(thm)],[c_0_403])).
cnf(c_0_413, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k1_algstr_0(X2,X3,X4),X5))|v2_struct_0(X2)|~r1_struct_0(X5,k4_algstr_0(X2,X4))|~r1_struct_0(X5,k4_algstr_0(X2,X3))|~r1_struct_0(X5,X1)|~m1_zmodul01(X5,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[c_0_404, c_0_195])).
cnf(c_0_414, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_405, c_0_83]), c_0_60]), c_0_103])).
cnf(c_0_415, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),X3)=k5_algstr_0(X1,X2,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_406])])).
cnf(c_0_416, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_238])])).
cnf(c_0_417, plain, (l1_struct_0(X1)|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_407])).
cnf(c_0_418, negated_conjecture, (r1_struct_0(esk3_0,X1)|~l2_algstr_0(esk3_0)|~r2_hidden(X1,a_3_1_zmodul01(esk1_0,X2,X3))|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(esk8_4(X1,esk1_0,X2,X3),u1_struct_0(esk3_0))|~m1_subset_1(X2,u1_struct_0(esk3_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[c_0_408, c_0_61])).
cnf(c_0_419, plain, (m1_subset_1(esk8_4(X1,X2,X3,X4),u1_struct_0(X4))|v2_struct_0(X2)|~l2_algstr_0(X4)|~r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_192])])).
cnf(c_0_420, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_409])).
cnf(c_0_421, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X2,X1),X3)|r1_struct_0(X1,esk4_2(a_3_1_zmodul01(esk1_0,X2,X1),X3))|~r1_struct_0(X1,X2)|~m1_zmodul01(X1,esk1_0)|~m1_subset_1(esk8_4(esk4_2(a_3_1_zmodul01(esk1_0,X2,X1),X3),esk1_0,X2,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_410, c_0_411])])).
cnf(c_0_422, plain, (r1_tarski(a_3_1_zmodul01(X1,X2,X3),X4)|m1_subset_1(esk8_4(esk4_2(a_3_1_zmodul01(X1,X2,X3),X4),X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_411])])).
cnf(c_0_423, plain, (r1_tarski(X1,a_3_1_zmodul01(X2,k1_algstr_0(X2,X3,X4),X5))|v2_struct_0(X2)|~r1_struct_0(X5,esk4_2(X1,a_3_1_zmodul01(X2,k1_algstr_0(X2,X3,X4),X5)))|~r1_struct_0(X5,k4_algstr_0(X2,X4))|~r1_struct_0(X5,k4_algstr_0(X2,X3))|~m1_zmodul01(X5,X2)|~m1_subset_1(esk4_2(X1,a_3_1_zmodul01(X2,k1_algstr_0(X2,X3,X4),X5)),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_412, c_0_413])])).
cnf(c_0_424, plain, (r1_struct_0(X1,k4_algstr_0(X2,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_414, c_0_415])])).
cnf(c_0_425, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),X3))|v2_struct_0(X2)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_416, c_0_60]), c_0_103])).
cnf(c_0_426, plain, (r2_hidden(X1,u1_struct_0(X2))|~l1_algstr_0(X2)|~r1_struct_0(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_417])])).
cnf(c_0_427, negated_conjecture, (r1_struct_0(esk3_0,X1)|~l2_algstr_0(esk3_0)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(esk8_4(X1,esk1_0,X1,X2),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_418, c_0_187]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_428, plain, (m1_subset_1(esk8_4(X1,X2,X1,X3),u1_struct_0(X3))|v2_struct_0(X2)|~l2_algstr_0(X3)|~m1_zmodul01(X3,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_419, c_0_187])])).
cnf(c_0_429, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_420])).
cnf(c_0_430, plain, (r1_tarski(X1,u1_struct_0(X2))|~l2_algstr_0(X2)|~r1_struct_0(X2,esk4_2(X1,u1_struct_0(X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_412, c_0_222])])).
cnf(c_0_431, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,X2),X3)|r1_struct_0(X2,esk4_2(a_3_1_zmodul01(esk1_0,X1,X2),X3))|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_421, c_0_422]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_432, plain, (r1_tarski(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X4,esk4_2(X1,a_3_1_zmodul01(X2,X3,X4)))|~r1_struct_0(X4,k4_algstr_0(X2,k4_struct_0(X2)))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(esk4_2(X1,a_3_1_zmodul01(X2,X3,X4)),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_423, c_0_63])])).
cnf(c_0_433, plain, (r1_struct_0(X1,k4_algstr_0(X2,k4_struct_0(X2)))|v2_struct_0(X2)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_424, c_0_83]), c_0_425]), c_0_103])).
cnf(c_0_434, plain, (r1_tarski(X1,X2)|m1_subset_1(esk4_2(X1,X2),X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_411])])).
cnf(c_0_435, plain, (r1_tarski(X1,u1_struct_0(X2))|~l1_algstr_0(X2)|~r1_struct_0(X2,esk4_2(X1,u1_struct_0(X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_412, c_0_426])])).
cnf(c_0_436, negated_conjecture, (r1_struct_0(esk3_0,X1)|~l2_algstr_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_427, c_0_428]), c_0_354]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_437, plain, (m1_subset_1(X1,X2)|~r1_tarski(X3,X2)|~r2_hidden(X1,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_287, c_0_429])])).
cnf(c_0_438, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,X2),u1_struct_0(X2))|~l2_algstr_0(X2)|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_431])])).
fof(c_0_439, plain, ![X161, X162]:((((r1_tarski(X161,X162)|X161!=X162)&(r1_tarski(X162,X161)|X161!=X162))&(~r1_tarski(X161,X162)|~r1_tarski(X162,X161)|X161=X162))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_440, plain, (r1_tarski(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|~r1_struct_0(X4,esk4_2(X1,a_3_1_zmodul01(X2,X3,X4)))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(esk4_2(X1,a_3_1_zmodul01(X2,X3,X4)),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_432, c_0_83]), c_0_433]), c_0_103])).
cnf(c_0_441, plain, (r1_tarski(u1_struct_0(X1),X2)|m1_subset_1(esk4_2(u1_struct_0(X1),X2),u1_struct_0(X3))|v2_struct_0(X3)|~m1_zmodul01(X1,X3)|~l1_zmodul01(X3)|~v5_zmodul01(X3)|~v4_zmodul01(X3)|~v3_zmodul01(X3)|~v2_zmodul01(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v13_algstr_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_335, c_0_434])])).
cnf(c_0_442, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk3_0))|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(esk4_2(X1,u1_struct_0(esk3_0)),u1_struct_0(esk3_0))|~m1_subset_1(esk4_2(X1,u1_struct_0(esk3_0)),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_435, c_0_436])])).
cnf(c_0_443, negated_conjecture, (m1_subset_1(X1,u1_struct_0(X2))|~l2_algstr_0(X2)|~r2_hidden(X1,a_3_1_zmodul01(esk1_0,X3,X2))|~r1_struct_0(X2,X3)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X3,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_437, c_0_438])])).
cnf(c_0_444, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_439])).
cnf(c_0_445, plain, (r1_tarski(u1_struct_0(X1),a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|~r1_struct_0(X4,esk4_2(u1_struct_0(X1),a_3_1_zmodul01(X2,X3,X4)))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_440, c_0_441])])).
cnf(c_0_446, plain, (r1_tarski(u1_struct_0(X1),X2)|r1_struct_0(X1,esk4_2(u1_struct_0(X1),X2))|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_411])])).
cnf(c_0_447, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~m1_subset_1(esk4_2(X1,u1_struct_0(esk3_0)),u1_struct_0(esk3_0))|~m1_subset_1(esk4_2(X1,u1_struct_0(esk3_0)),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[c_0_442, c_0_61])).
cnf(c_0_448, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,X2),X3)|m1_subset_1(esk4_2(a_3_1_zmodul01(esk1_0,X1,X2),X3),u1_struct_0(X2))|~l2_algstr_0(X2)|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_443, c_0_411])])).
cnf(c_0_449, plain, (X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_444])).
cnf(c_0_450, plain, (r1_tarski(u1_struct_0(X1),a_3_1_zmodul01(X2,X3,X1))|v2_struct_0(X2)|~l1_struct_0(X1)|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_445, c_0_446])])).
cnf(c_0_451, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(esk4_2(a_3_1_zmodul01(esk1_0,X1,esk3_0),u1_struct_0(esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_447, c_0_448]), c_0_354])])])).
cnf(c_0_452, plain, (r1_tarski(a_3_1_zmodul01(X1,X2,X3),X4)|m1_subset_1(esk4_2(a_3_1_zmodul01(X1,X2,X3),X4),u1_struct_0(X1))|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_296, c_0_411])])).
cnf(c_0_453, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_409])).
cnf(c_0_454, plain, (a_3_1_zmodul01(X1,X2,X3)=u1_struct_0(X3)|v2_struct_0(X1)|~r1_tarski(a_3_1_zmodul01(X1,X2,X3),u1_struct_0(X3))|~l1_struct_0(X3)|~r1_struct_0(X3,k4_algstr_0(X1,X2))|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_449, c_0_450])])).
cnf(c_0_455, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_451, c_0_452]), c_0_354]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_456, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_328])).
cnf(c_0_457, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_453])).
cnf(c_0_458, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,esk3_0)=u1_struct_0(esk3_0)|~l2_algstr_0(esk3_0)|~l1_struct_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_454, c_0_455]), c_0_354]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_459, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_456])).
cnf(c_0_460, plain, (r1_tarski(k5_zmodul01(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_457, c_0_278])])).
cnf(c_0_461, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,esk3_0)=u1_struct_0(esk3_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[c_0_458, c_0_141])).
cnf(c_0_462, negated_conjecture, (r1_struct_0(X1,k4_algstr_0(esk1_0,k4_struct_0(esk1_0)))|~m1_zmodul01(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_433, c_0_459]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
fof(c_0_463, plain, ![X1]:(((((((((((~v2_struct_0(X1)&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v2_zmodul01(X1))&v3_zmodul01(X1))&v4_zmodul01(X1))&v5_zmodul01(X1))&l1_zmodul01(X1))=>![X2]:((m1_subset_1(X2,u1_struct_0(X1))=>![X3]:((m1_subset_1(X3,u1_struct_0(X1))=>![X4]:((m1_zmodul01(X4,X1)=>![X5]:((m1_subset_1(X5,u1_struct_0(X4))=>![X6]:((m1_subset_1(X6,u1_struct_0(X4))=>((X5=X2&X6=X3)=>k5_algstr_0(X4,X5,X6)=k5_algstr_0(X1,X2,X3)))))))))))))), inference(fof_simplification,[status(thm)],[t31_zmodul01])).
cnf(c_0_464, plain, (r1_tarski(a_3_1_zmodul01(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_460, c_0_279])])).
cnf(c_0_465, negated_conjecture, (a_3_1_zmodul01(esk1_0,k4_struct_0(esk1_0),esk3_0)=u1_struct_0(esk3_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_461, c_0_462]), c_0_354])])])).
fof(c_0_466, plain, ![X114, X115, X116, X117, X118, X119]:((v2_struct_0(X114)|~v13_algstr_0(X114)|~v2_rlvect_1(X114)|~v3_rlvect_1(X114)|~v4_rlvect_1(X114)|~v2_zmodul01(X114)|~v3_zmodul01(X114)|~v4_zmodul01(X114)|~v5_zmodul01(X114)|~l1_zmodul01(X114)|(~m1_subset_1(X115,u1_struct_0(X114))|(~m1_subset_1(X116,u1_struct_0(X114))|(~m1_zmodul01(X117,X114)|(~m1_subset_1(X118,u1_struct_0(X117))|(~m1_subset_1(X119,u1_struct_0(X117))|(X118!=X115|X119!=X116|k5_algstr_0(X117,X118,X119)=k5_algstr_0(X114,X115,X116))))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_463])])])).
cnf(c_0_467, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_464, c_0_465]), c_0_354]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_468, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_81])])).
cnf(c_0_469, plain, (v2_struct_0(X1)|k5_algstr_0(X4,X5,X6)=k5_algstr_0(X1,X2,X3)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_zmodul01(X4,X1)|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X4))|X5!=X2|X6!=X3), inference(split_conjunct,[status(thm)],[c_0_466])).
cnf(c_0_470, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_467, c_0_160]), c_0_354]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_471, plain, (v1_xboole_0(u1_struct_0(X1))|r1_struct_0(X1,k1_algstr_0(X1,X2,X3))|~l1_algstr_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_468])])).
cnf(c_0_472, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|k5_zmodul01(esk1_0,esk2_0,esk3_0)=u1_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_328])).
cnf(c_0_473, plain, (r1_struct_0(X1,k5_algstr_0(X1,X3,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,k4_algstr_0(X1,X2)),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_70])])).
cnf(c_0_474, plain, (k5_algstr_0(X1,X2,X3)=k5_algstr_0(X4,X5,X6)|v2_struct_0(X1)|X3!=X6|X2!=X5|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_zmodul01(X4,X1)|~m1_subset_1(X6,u1_struct_0(X4))|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_469])).
cnf(c_0_475, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_470, c_0_83]), c_0_360]), c_0_361]), c_0_363])]), c_0_364])])).
cnf(c_0_476, plain, (v1_xboole_0(u1_struct_0(X1))|r1_struct_0(X1,k1_algstr_0(X1,X2,X3))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(csr,[status(thm)],[c_0_471, c_0_417])).
cnf(c_0_477, negated_conjecture, (k5_zmodul01(esk1_0,esk2_0,esk3_0)=u1_struct_0(esk3_0)|r1_struct_0(esk3_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_472])).
cnf(c_0_478, plain, (r1_struct_0(X1,k5_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,k4_algstr_0(X1,X3)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_473, c_0_60])).
cnf(c_0_479, plain, (k5_algstr_0(X1,X2,X3)=k5_algstr_0(X4,X2,X3)|v2_struct_0(X1)|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(er,[status(thm)],[c_0_474])])])).
cnf(c_0_480, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~r2_hidden(X1,u1_struct_0(esk3_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_437, c_0_475])])).
cnf(c_0_481, plain, (v1_xboole_0(u1_struct_0(X1))|r1_struct_0(X1,X2)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_476, c_0_63])])).
cnf(c_0_482, plain, (l1_zmodul01(X1)|v2_struct_0(X2)|~m1_zmodul01(X1,X2)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)), inference(split_conjunct,[status(thm)],[c_0_329])).
cnf(c_0_483, negated_conjecture, (a_3_1_zmodul01(esk1_0,esk2_0,esk3_0)=u1_struct_0(esk3_0)|r1_struct_0(esk3_0,esk2_0)|v2_struct_0(esk1_0)|~m1_subset_1(esk2_0,u1_struct_0(esk1_0))|~v5_zmodul01(esk1_0)|~v4_zmodul01(esk1_0)|~v3_zmodul01(esk1_0)|~v2_zmodul01(esk1_0)|~v4_rlvect_1(esk1_0)|~v3_rlvect_1(esk1_0)|~v2_rlvect_1(esk1_0)|~v13_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_477, c_0_279]), c_0_354]), c_0_355])])])).
cnf(c_0_484, plain, (r1_struct_0(X1,k5_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_478, c_0_81])])).
cnf(c_0_485, plain, (k5_algstr_0(X1,X2,X3)=k5_algstr_0(X4,X2,X3)|v2_struct_0(X1)|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_479, c_0_335]), c_0_335])).
cnf(c_0_486, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~r1_struct_0(esk3_0,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_480, c_0_222])])).
cnf(c_0_487, plain, (r1_struct_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_481, c_0_83]), c_0_61]), c_0_167])).
cnf(c_0_488, plain, (v2_struct_0(X2)|l1_zmodul01(X1)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X1,X2)), inference(evalgc,[status(thm)],[c_0_482])).
cnf(c_0_489, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk3_0,X2)|~l1_algstr_0(esk3_0)|~r2_hidden(X2,a_3_1_zmodul01(esk1_0,X1,X3))|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(esk8_4(X2,esk1_0,X1,X3),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_476, c_0_396])])).
cnf(c_0_490, negated_conjecture, (a_3_1_zmodul01(esk1_0,esk2_0,esk3_0)=u1_struct_0(esk3_0)|r1_struct_0(esk3_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_483, c_0_459]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])).
cnf(c_0_491, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,X2),u1_struct_0(X2))|~l1_algstr_0(X2)|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_435, c_0_431])])).
cnf(c_0_492, plain, (r1_struct_0(X1,k5_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_484, c_0_60]), c_0_61])).
cnf(c_0_493, negated_conjecture, (k5_algstr_0(esk3_0,X1,X2)=k5_algstr_0(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_485, c_0_354]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_494, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_486, c_0_487]), c_0_397]), c_0_398]), c_0_399])]), c_0_400])])).
cnf(c_0_495, negated_conjecture, (l1_zmodul01(esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_488, c_0_354]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_496, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk3_0,X1)|~l1_algstr_0(esk3_0)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(esk8_4(X1,esk1_0,X1,X2),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_489, c_0_187]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_497, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk3_0))|r1_struct_0(esk3_0,esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_490]), c_0_354]), c_0_459]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_498, negated_conjecture, (a_3_1_zmodul01(esk1_0,X2,X1)=u1_struct_0(X1)|~l1_algstr_0(X1)|~r1_tarski(u1_struct_0(X1),a_3_1_zmodul01(esk1_0,X2,X1))|~r1_struct_0(X1,X2)|~m1_zmodul01(X1,esk1_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_449, c_0_491])])).
cnf(c_0_499, negated_conjecture, (r1_struct_0(esk3_0,k5_algstr_0(esk1_0,X1,X2))|~l2_algstr_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X2,u1_struct_0(esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_492, c_0_493]), c_0_397]), c_0_398]), c_0_373]), c_0_399])]), c_0_400])])).
cnf(c_0_500, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_494, c_0_103]), c_0_495])])])).
cnf(c_0_501, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk3_0,X1)|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_496, c_0_428]), c_0_354]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_502, plain, (m1_subset_1(X1,u1_struct_0(X2))|~l1_algstr_0(X2)|~r1_struct_0(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_426])])).
cnf(c_0_503, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|m1_subset_1(esk2_0,u1_struct_0(esk3_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_497])])).
fof(c_0_504, plain, ![X198, X199]:((~r2_hidden(X198,X199)|~v1_xboole_0(X199))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
cnf(c_0_505, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,X2)=u1_struct_0(X2)|~l1_algstr_0(X2)|~l1_struct_0(X2)|~r1_struct_0(X2,k4_algstr_0(esk1_0,X1))|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_498, c_0_450]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_506, negated_conjecture, (r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_414, c_0_499]), c_0_354]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_507, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_500, c_0_103]), c_0_355])])])).
cnf(c_0_508, negated_conjecture, (~r1_struct_0(esk3_0,esk2_0)|k5_zmodul01(esk1_0,esk2_0,esk3_0)!=u1_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_328])).
cnf(c_0_509, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk3_0,X1)|~l2_algstr_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[c_0_501, c_0_61])).
cnf(c_0_510, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk3_0))|~l1_algstr_0(esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_502, c_0_503])])).
cnf(c_0_511, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_504])).
cnf(c_0_512, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,X2)=u1_struct_0(X2)|~l1_algstr_0(X2)|~r1_struct_0(X2,k4_algstr_0(esk1_0,X1))|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[c_0_505, c_0_417])).
cnf(c_0_513, negated_conjecture, (r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(csr,[status(thm)],[c_0_506, c_0_507])).
cnf(c_0_514, negated_conjecture, (k5_zmodul01(esk1_0,esk2_0,esk3_0)!=u1_struct_0(esk3_0)|~r1_struct_0(esk3_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_508])).
cnf(c_0_515, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_509, c_0_103]), c_0_495])])])).
cnf(c_0_516, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_510, c_0_111]), c_0_495])])])).
cnf(c_0_517, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_511])).
cnf(c_0_518, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,esk3_0)=u1_struct_0(esk3_0)|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_512, c_0_513]), c_0_354])])])).
cnf(c_0_519, negated_conjecture, (v2_struct_0(esk1_0)|a_3_1_zmodul01(esk1_0,esk2_0,esk3_0)!=u1_struct_0(esk3_0)|~r1_struct_0(esk3_0,esk2_0)|~m1_subset_1(esk2_0,u1_struct_0(esk1_0))|~v5_zmodul01(esk1_0)|~v4_zmodul01(esk1_0)|~v3_zmodul01(esk1_0)|~v2_zmodul01(esk1_0)|~v4_rlvect_1(esk1_0)|~v3_rlvect_1(esk1_0)|~v2_rlvect_1(esk1_0)|~v13_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_514, c_0_279]), c_0_354]), c_0_355])])])).
cnf(c_0_520, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk3_0,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_515, c_0_516]), c_0_459])])])).
cnf(c_0_521, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|~v1_xboole_0(u1_struct_0(esk3_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_497])])).
cnf(c_0_522, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,esk3_0)=u1_struct_0(esk3_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_518, c_0_502]), c_0_61])).
cnf(c_0_523, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X2))|v2_struct_0(X1)|~l1_algstr_0(X2)|~m1_zmodul01(X2,X1)|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_502, c_0_160])])).
cnf(c_0_524, negated_conjecture, (a_3_1_zmodul01(esk1_0,esk2_0,esk3_0)!=u1_struct_0(esk3_0)|~r1_struct_0(esk3_0,esk2_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_519, c_0_459]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])).
cnf(c_0_525, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)), inference(csr,[status(thm)],[c_0_520, c_0_521])).
cnf(c_0_526, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,esk3_0)=u1_struct_0(esk3_0)|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_522, c_0_523]), c_0_354]), c_0_355]), c_0_356]), c_0_357]), c_0_358]), c_0_359]), c_0_360]), c_0_361]), c_0_362]), c_0_363])]), c_0_364])])).
cnf(c_0_527, negated_conjecture, (a_3_1_zmodul01(esk1_0,esk2_0,esk3_0)!=u1_struct_0(esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_524, c_0_525])])])).
cnf(c_0_528, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,esk3_0)=u1_struct_0(esk3_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(csr,[status(thm)],[c_0_526, c_0_61])).
cnf(c_0_529, negated_conjecture, (~l2_algstr_0(esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_527, c_0_528]), c_0_525]), c_0_459])])])).
cnf(c_0_530, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_529, c_0_103]), c_0_495])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 531
# Proof object clause steps            : 452
# Proof object formula steps           : 79
# Proof object conjectures             : 100
# Proof object clause conjectures      : 97
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 59
# Proof object initial formulas used   : 33
# Proof object generating inferences   : 225
# Proof object simplifying inferences  : 550
# Training examples: 286 positive, 1718 negative
# Training: Positive examples begin
cnf(c_0_514, negated_conjecture, (k5_zmodul01(esk1_0,esk2_0,esk3_0)!=u1_struct_0(esk3_0)|~r1_struct_0(esk3_0,esk2_0))).# trainpos
cnf(c_0_354, negated_conjecture, (m1_zmodul01(esk3_0,esk1_0))).# trainpos
cnf(c_0_477, negated_conjecture, (k5_zmodul01(esk1_0,esk2_0,esk3_0)=u1_struct_0(esk3_0)|r1_struct_0(esk3_0,esk2_0))).# trainpos
cnf(c_0_355, negated_conjecture, (l1_zmodul01(esk1_0))).# trainpos
cnf(c_0_103, plain, (l2_algstr_0(X1)|~l1_zmodul01(X1))).# trainpos
cnf(c_0_125, plain, (r1_struct_0(X2,X1)|~l1_struct_0(X2)|~r2_hidden(X1,u1_struct_0(X2)))).# trainpos
cnf(c_0_449, plain, (X1=X2|~r1_tarski(X2,X1)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_417, plain, (l1_struct_0(X1)|~l1_algstr_0(X1))).# trainpos
cnf(c_0_279, plain, (k5_zmodul01(X1,X2,X3)=a_3_1_zmodul01(X1,X2,X3)|v2_struct_0(X1)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_459, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_51, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|~v3_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_364, negated_conjecture, (~v2_struct_0(esk1_0))).# trainpos
cnf(c_0_363, negated_conjecture, (v13_algstr_0(esk1_0))).# trainpos
cnf(c_0_361, negated_conjecture, (v3_rlvect_1(esk1_0))).# trainpos
cnf(c_0_127, plain, (l1_struct_0(X1)|~l2_struct_0(X1))).# trainpos
cnf(c_0_412, plain, (r1_tarski(X1,X2)|~r2_hidden(esk4_2(X1,X2),X2))).# trainpos
cnf(c_0_335, plain, (v2_struct_0(X1)|m1_subset_1(X3,u1_struct_0(X1))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_zmodul01(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2)))).# trainpos
cnf(c_0_474, plain, (k5_algstr_0(X1,X2,X3)=k5_algstr_0(X4,X5,X6)|v2_struct_0(X1)|X3!=X6|X2!=X5|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_zmodul01(X4,X1)|~m1_subset_1(X6,u1_struct_0(X4))|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_39, plain, (k3_rlvect_1(X1,X2,X3)=k1_algstr_0(X1,X2,X3)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_214, plain, (r2_hidden(X2,u1_struct_0(X1))|~l1_struct_0(X1)|~r1_struct_0(X1,X2))).# trainpos
cnf(c_0_411, plain, (r1_tarski(X1,X2)|r2_hidden(esk4_2(X1,X2),X1))).# trainpos
cnf(c_0_61, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1))).# trainpos
cnf(c_0_360, negated_conjecture, (v4_rlvect_1(esk1_0))).# trainpos
cnf(c_0_74, plain, (k1_algstr_0(X1,X2,k4_struct_0(X1))=X2|v2_struct_0(X1)|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_362, negated_conjecture, (v2_rlvect_1(esk1_0))).# trainpos
cnf(c_0_102, plain, (v2_struct_0(X2)|r2_hidden(X3,a_3_1_zmodul01(X2,X4,X5))|X3!=k3_rlvect_1(X2,X4,X1)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X5,X2)|~r1_struct_0(X5,X1)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2)))).# trainpos
cnf(c_0_163, plain, (v2_struct_0(X1)|r1_struct_0(X4,k3_rlvect_1(X1,X2,X3))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_zmodul01(X4,X1)|~r1_struct_0(X4,X3)|~r1_struct_0(X4,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_128, plain, (l2_struct_0(X1)|~l2_algstr_0(X1))).# trainpos
cnf(c_0_63, plain, (k1_algstr_0(X1,k4_struct_0(X1),X2)=X2|v2_struct_0(X1)|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_326, plain, (k3_rlvect_1(X1,X2,X3)=k3_rlvect_1(X4,X5,X6)|v2_struct_0(X1)|X3!=X6|X2!=X5|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_zmodul01(X4,X1)|~m1_subset_1(X6,u1_struct_0(X4))|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_160, plain, (v2_struct_0(X1)|r1_struct_0(X2,k4_struct_0(X1))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_zmodul01(X2,X1))).# trainpos
cnf(c_0_278, plain, (v2_struct_0(X1)|m1_subset_1(k5_zmodul01(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v2_zmodul01(X1)|~v3_zmodul01(X1)|~v4_zmodul01(X1)|~v5_zmodul01(X1)|~l1_zmodul01(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_356, negated_conjecture, (v5_zmodul01(esk1_0))).# trainpos
cnf(c_0_357, negated_conjecture, (v4_zmodul01(esk1_0))).# trainpos
cnf(c_0_359, negated_conjecture, (v2_zmodul01(esk1_0))).# trainpos
cnf(c_0_358, negated_conjecture, (v3_zmodul01(esk1_0))).# trainpos
cnf(c_0_488, plain, (v2_struct_0(X2)|l1_zmodul01(X1)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X1,X2))).# trainpos
cnf(c_0_457, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_65, plain, (k5_algstr_0(X1,X2,X2)=k4_struct_0(X1)|v2_struct_0(X1)|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_40, plain, (k3_rlvect_1(X1,X2,X3)=k3_rlvect_1(X1,X3,X2)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_221, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_287, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_156, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).# trainpos
cnf(c_0_52, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X3))=k5_algstr_0(X1,X2,X3)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_517, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_114, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_429, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_64, plain, (m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_138, plain, (k3_rlvect_1(X2,X3,esk8_4(X1,X2,X3,X4))=X1|v2_struct_0(X2)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4)))).# trainpos
cnf(c_0_78, plain, (m1_subset_1(esk9_1(X1),X1))).# trainpos
cnf(c_0_81, plain, (m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_192, plain, (v2_struct_0(X3)|r1_struct_0(X1,esk8_4(X2,X3,X4,X1))|~v13_algstr_0(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v2_zmodul01(X3)|~v3_zmodul01(X3)|~v4_zmodul01(X3)|~v5_zmodul01(X3)|~l1_zmodul01(X3)|~m1_zmodul01(X1,X3)|~m1_subset_1(X4,u1_struct_0(X3))|~r2_hidden(X2,a_3_1_zmodul01(X3,X4,X1)))).# trainpos
cnf(c_0_386, plain, (v2_struct_0(X2)|v3_rlvect_1(X1)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X1,X2))).# trainpos
cnf(c_0_388, plain, (v2_struct_0(X2)|~v2_struct_0(X1)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X1,X2))).# trainpos
cnf(c_0_387, plain, (v2_struct_0(X2)|v13_algstr_0(X1)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X1,X2))).# trainpos
cnf(c_0_385, plain, (v2_struct_0(X2)|v4_rlvect_1(X1)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X1,X2))).# trainpos
cnf(c_0_60, plain, (m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_365, plain, (v2_struct_0(X2)|v2_rlvect_1(X1)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X1,X2))).# trainpos
cnf(c_0_153, plain, (v2_struct_0(X2)|m1_subset_1(esk8_4(X1,X2,X3,X4),u1_struct_0(X2))|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4)))).# trainpos
cnf(c_0_111, plain, (l1_algstr_0(X1)|~l1_zmodul01(X1))).# trainpos
cnf(c_0_483, negated_conjecture, (a_3_1_zmodul01(esk1_0,esk2_0,esk3_0)=u1_struct_0(esk3_0)|r1_struct_0(esk3_0,esk2_0)|v2_struct_0(esk1_0)|~m1_subset_1(esk2_0,u1_struct_0(esk1_0))|~v5_zmodul01(esk1_0)|~v4_zmodul01(esk1_0)|~v3_zmodul01(esk1_0)|~v2_zmodul01(esk1_0)|~v4_rlvect_1(esk1_0)|~v3_rlvect_1(esk1_0)|~v2_rlvect_1(esk1_0)|~v13_algstr_0(esk1_0))).# trainpos
cnf(c_0_519, negated_conjecture, (v2_struct_0(esk1_0)|a_3_1_zmodul01(esk1_0,esk2_0,esk3_0)!=u1_struct_0(esk3_0)|~r1_struct_0(esk3_0,esk2_0)|~m1_subset_1(esk2_0,u1_struct_0(esk1_0))|~v5_zmodul01(esk1_0)|~v4_zmodul01(esk1_0)|~v3_zmodul01(esk1_0)|~v2_zmodul01(esk1_0)|~v4_rlvect_1(esk1_0)|~v3_rlvect_1(esk1_0)|~v2_rlvect_1(esk1_0)|~v13_algstr_0(esk1_0))).# trainpos
cnf(c_0_426, plain, (r2_hidden(X1,u1_struct_0(X2))|~l1_algstr_0(X2)|~r1_struct_0(X2,X1))).# trainpos
cnf(c_0_110, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l1_algstr_0(X2)|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_400, negated_conjecture, (~v2_struct_0(esk3_0))).# trainpos
cnf(c_0_495, negated_conjecture, (l1_zmodul01(esk3_0))).# trainpos
cnf(c_0_211, plain, (r2_hidden(k3_rlvect_1(X1,X2,X3),a_3_1_zmodul01(X1,X2,X4))|v2_struct_0(X1)|~r1_struct_0(X4,X3)|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_47, plain, (k3_rlvect_1(X1,X3,X2)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_141, plain, (l1_struct_0(X1)|~l2_algstr_0(X1))).# trainpos
cnf(c_0_173, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~l1_algstr_0(X2)|~r1_struct_0(X1,X4)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_479, plain, (k5_algstr_0(X1,X2,X3)=k5_algstr_0(X4,X2,X3)|v2_struct_0(X1)|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_152, plain, (k1_algstr_0(X1,X2,esk8_4(X3,X1,X2,X4))=X3|v2_struct_0(X1)|~l1_algstr_0(X1)|~r2_hidden(X3,a_3_1_zmodul01(X1,X2,X4))|~m1_zmodul01(X4,X1)|~m1_subset_1(esk8_4(X3,X1,X2,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_398, negated_conjecture, (v3_rlvect_1(esk3_0))).# trainpos
cnf(c_0_399, negated_conjecture, (v13_algstr_0(esk3_0))).# trainpos
cnf(c_0_460, plain, (r1_tarski(k5_zmodul01(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_288, plain, (m1_subset_1(a_3_1_zmodul01(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_170, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_struct_0(X1)))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_422, plain, (r1_tarski(a_3_1_zmodul01(X1,X2,X3),X4)|m1_subset_1(esk8_4(esk4_2(a_3_1_zmodul01(X1,X2,X3),X4),X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_181, plain, (r1_struct_0(X1,X4)|v2_struct_0(X2)|~r2_hidden(X4,a_3_1_zmodul01(X2,X3,X5))|~r1_struct_0(X1,esk8_4(X4,X2,X3,X5))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_zmodul01(X5,X2)|~m1_subset_1(esk8_4(X4,X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_72, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_79, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),X3))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_334, plain, (k3_rlvect_1(X1,X2,X3)=k3_rlvect_1(X4,X2,X3)|v2_struct_0(X1)|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_446, plain, (r1_tarski(u1_struct_0(X1),X2)|r1_struct_0(X1,esk4_2(u1_struct_0(X1),X2))|~l1_struct_0(X1))).# trainpos
cnf(c_0_92, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_397, negated_conjecture, (v4_rlvect_1(esk3_0))).# trainpos
cnf(c_0_373, negated_conjecture, (v2_rlvect_1(esk3_0))).# trainpos
cnf(c_0_71, plain, (k5_algstr_0(X1,k4_struct_0(X1),X2)=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_242, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_algstr_0(X1,X4)))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_rlvect_1(X1))).# trainpos
cnf(c_0_112, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4)),X5)=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k1_algstr_0(X1,X4,X5))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1))).# trainpos
cnf(c_0_434, plain, (r1_tarski(X1,X2)|m1_subset_1(esk4_2(X1,X2),X1))).# trainpos
cnf(c_0_468, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_59, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X3),X4))=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X4)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_rlvect_1(X1))).# trainpos
cnf(c_0_126, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_437, plain, (m1_subset_1(X1,X2)|~r1_tarski(X3,X2)|~r2_hidden(X1,X3))).# trainpos
cnf(c_0_145, plain, (m1_subset_1(k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4)),u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1))).# trainpos
cnf(c_0_502, plain, (m1_subset_1(X1,u1_struct_0(X2))|~l1_algstr_0(X2)|~r1_struct_0(X2,X1))).# trainpos
cnf(c_0_83, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_199, plain, (r1_struct_0(X1,X2)|v2_struct_0(X3)|~r2_hidden(X2,a_3_1_zmodul01(X3,X4,X1))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X3)|~m1_subset_1(X4,u1_struct_0(X3))|~l1_zmodul01(X3)|~v5_zmodul01(X3)|~v4_zmodul01(X3)|~v3_zmodul01(X3)|~v2_zmodul01(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v13_algstr_0(X3))).# trainpos
cnf(c_0_55, plain, (k1_algstr_0(X1,X3,X2)=k1_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_249, plain, (k1_algstr_0(X1,esk8_4(X3,X1,X2,X4),X2)=X3|v2_struct_0(X1)|~l1_algstr_0(X1)|~r2_hidden(X3,a_3_1_zmodul01(X1,X2,X4))|~m1_zmodul01(X4,X1)|~m1_subset_1(esk8_4(X3,X1,X2,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_134, plain, (k4_algstr_0(X1,k4_struct_0(X1))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_296, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_222, plain, (r2_hidden(X1,u1_struct_0(X2))|~l2_algstr_0(X2)|~r1_struct_0(X2,X1))).# trainpos
cnf(c_0_303, plain, (r2_hidden(k1_algstr_0(X1,X2,X3),a_3_1_zmodul01(X1,X2,X4))|v2_struct_0(X1)|~r1_struct_0(X4,X3)|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_372, negated_conjecture, (k3_rlvect_1(esk3_0,X1,X2)=k3_rlvect_1(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0)))).# trainpos
cnf(c_0_493, negated_conjecture, (k5_algstr_0(esk3_0,X1,X2)=k5_algstr_0(esk1_0,X1,X2)|~m1_subset_1(X2,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0)))).# trainpos
cnf(c_0_176, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X1,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X3,k4_struct_0(X2))|~m1_zmodul01(X3,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_262, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)=k1_algstr_0(X1,X2,k5_algstr_0(X1,X3,X4))|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_rlvect_1(X1))).# trainpos
cnf(c_0_471, plain, (v1_xboole_0(u1_struct_0(X1))|r1_struct_0(X1,k1_algstr_0(X1,X2,X3))|~l1_algstr_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_435, plain, (r1_tarski(X1,u1_struct_0(X2))|~l1_algstr_0(X2)|~r1_struct_0(X2,esk4_2(X1,u1_struct_0(X2))))).# trainpos
cnf(c_0_330, plain, (k1_algstr_0(X1,k4_struct_0(X1),X3)=X3|v2_struct_0(X1)|~l2_algstr_0(X1)|~r2_hidden(X3,a_3_1_zmodul01(X1,X2,X4))|~m1_zmodul01(X4,X1)|~m1_subset_1(esk8_4(X3,X1,X2,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_167, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1))).# trainpos
cnf(c_0_284, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_464, plain, (r1_tarski(a_3_1_zmodul01(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_140, plain, (v1_xboole_0(u1_struct_0(X1))|r1_struct_0(X1,k5_algstr_0(X1,X2,X3))|~l2_algstr_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_133, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_137, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),k1_algstr_0(X1,X3,X4))=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_190, plain, (k1_algstr_0(X1,X2,k4_struct_0(X1))=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_219, plain, (r2_hidden(k1_algstr_0(X1,X3,X2),a_3_1_zmodul01(X1,X2,X4))|v2_struct_0(X1)|~l1_algstr_0(X1)|~r1_struct_0(X4,X3)|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_366, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_216, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3),X4))=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_441, plain, (r1_tarski(u1_struct_0(X1),X2)|m1_subset_1(esk4_2(u1_struct_0(X1),X2),u1_struct_0(X3))|v2_struct_0(X3)|~m1_zmodul01(X1,X3)|~l1_zmodul01(X3)|~v5_zmodul01(X3)|~v4_zmodul01(X3)|~v3_zmodul01(X3)|~v2_zmodul01(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v13_algstr_0(X3))).# trainpos
cnf(c_0_201, plain, (k1_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2),X3)=k1_algstr_0(X1,k4_algstr_0(X1,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,X2),X3),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_169, plain, (m1_subset_1(k1_algstr_0(X1,X2,k5_algstr_0(X1,X3,X4)),u1_struct_0(X1))|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_rlvect_1(X1))).# trainpos
cnf(c_0_175, plain, (m1_subset_1(k1_algstr_0(X1,X2,X4),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~r2_hidden(X4,a_3_1_zmodul01(X1,X3,X5))|~m1_zmodul01(X5,X1)|~m1_subset_1(esk8_4(X4,X1,X3,X5),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_497, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk3_0))|r1_struct_0(esk3_0,esk2_0))).# trainpos
cnf(c_0_229, plain, (m1_subset_1(X1,u1_struct_0(X2))|~l2_algstr_0(X2)|~r1_struct_0(X2,X1))).# trainpos
cnf(c_0_452, plain, (r1_tarski(a_3_1_zmodul01(X1,X2,X3),X4)|m1_subset_1(esk4_2(a_3_1_zmodul01(X1,X2,X3),X4),u1_struct_0(X1))|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_195, plain, (m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_481, plain, (v1_xboole_0(u1_struct_0(X1))|r1_struct_0(X1,X2)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_231, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_367, plain, (k1_algstr_0(X1,k4_struct_0(X1),X2)=X2|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_85, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X3,X2),X4)=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|~l1_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_384, negated_conjecture, (k3_rlvect_1(esk1_0,X1,X2)=k1_algstr_0(esk3_0,X1,X2)|~l1_algstr_0(esk3_0)|~m1_subset_1(X2,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0)))).# trainpos
cnf(c_0_523, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X2))|v2_struct_0(X1)|~l1_algstr_0(X2)|~m1_zmodul01(X2,X1)|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_80, plain, (k1_algstr_0(X1,X4,k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_430, plain, (r1_tarski(X1,u1_struct_0(X2))|~l2_algstr_0(X2)|~r1_struct_0(X2,esk4_2(X1,u1_struct_0(X2))))).# trainpos
cnf(c_0_261, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_algstr_0(X2,X3),X4),u1_struct_0(X2))|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_62, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),X2)=k5_algstr_0(X1,X2,X3)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_189, plain, (m1_subset_1(k1_algstr_0(X1,X2,k4_struct_0(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_159, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_struct_0(X2))|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_struct_0(X2))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_521, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|~v1_xboole_0(u1_struct_0(esk3_0)))).# trainpos
cnf(c_0_503, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|m1_subset_1(esk2_0,u1_struct_0(esk3_0)))).# trainpos
cnf(c_0_419, plain, (m1_subset_1(esk8_4(X1,X2,X3,X4),u1_struct_0(X4))|v2_struct_0(X2)|~l2_algstr_0(X4)|~r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_91, plain, (k1_algstr_0(X1,X3,k1_algstr_0(X1,X2,X4))=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|~l1_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_379, negated_conjecture, (k1_algstr_0(esk1_0,k4_struct_0(esk1_0),X1)=X1|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_280, plain, (r1_struct_0(X1,k5_algstr_0(X2,X4,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(k5_algstr_0(X2,X4,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_93, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X3,X2))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_132, plain, (k1_algstr_0(X1,X3,k1_algstr_0(X1,X2,k4_struct_0(X1)))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_197, plain, (m1_subset_1(k1_algstr_0(X1,X2,k4_struct_0(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_76, plain, (k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X4)=k1_algstr_0(X1,X2,k5_algstr_0(X1,X4,X3))|~l2_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_396, negated_conjecture, (k1_algstr_0(esk3_0,X1,esk8_4(X2,esk1_0,X1,X3))=X2|~l1_algstr_0(esk3_0)|~r2_hidden(X2,a_3_1_zmodul01(esk1_0,X1,X3))|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(esk8_4(X2,esk1_0,X1,X3),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk3_0)))).# trainpos
cnf(c_0_510, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk3_0))|~l1_algstr_0(esk3_0))).# trainpos
cnf(c_0_389, negated_conjecture, (r1_struct_0(X1,k1_algstr_0(esk1_0,X2,X3))|~r1_struct_0(X1,X3)|~r1_struct_0(X1,X2)|~m1_zmodul01(X1,esk1_0)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_108, plain, (k1_algstr_0(X1,X3,X2)=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_332, plain, (k1_algstr_0(X1,X2,k5_algstr_0(X1,X3,X2))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_82, plain, (k1_algstr_0(X1,k4_struct_0(X1),k5_algstr_0(X1,X3,X2))=k1_algstr_0(X1,k4_algstr_0(X1,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_489, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk3_0,X2)|~l1_algstr_0(esk3_0)|~r2_hidden(X2,a_3_1_zmodul01(esk1_0,X1,X3))|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(esk8_4(X2,esk1_0,X1,X3),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_207, plain, (m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_98, plain, (k1_algstr_0(X1,X3,k1_algstr_0(X1,k4_algstr_0(X1,X2),X4))=k5_algstr_0(X1,k1_algstr_0(X1,X3,X4),X2)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_428, plain, (m1_subset_1(esk8_4(X1,X2,X1,X3),u1_struct_0(X3))|v2_struct_0(X2)|~l2_algstr_0(X3)|~m1_zmodul01(X3,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_516, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk3_0)))).# trainpos
cnf(c_0_401, negated_conjecture, (r1_struct_0(X1,X2)|~r2_hidden(X2,a_3_1_zmodul01(esk1_0,X3,X4))|~r1_struct_0(X1,esk8_4(X2,esk1_0,X3,X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,esk1_0)|~m1_zmodul01(X4,esk1_0)|~m1_subset_1(esk8_4(X2,esk1_0,X3,X4),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_301, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_496, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk3_0,X1)|~l1_algstr_0(esk3_0)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(esk8_4(X1,esk1_0,X1,X2),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_161, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,X4),X3)=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X4)|~l2_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).# trainpos
cnf(c_0_94, plain, (k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))=k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_164, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),k1_algstr_0(X1,X4,X3))=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_501, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk3_0,X1)|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_410, negated_conjecture, (r1_struct_0(X1,X2)|~r2_hidden(X2,a_3_1_zmodul01(esk1_0,X3,X1))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,esk1_0)|~m1_subset_1(esk8_4(X2,esk1_0,X3,X1),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_101, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_185, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X4),X3)=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_117, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_515, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_118, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_234, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_129, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,X4),X4)=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_421, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X2,X1),X3)|r1_struct_0(X1,esk4_2(a_3_1_zmodul01(esk1_0,X2,X1),X3))|~r1_struct_0(X1,X2)|~m1_zmodul01(X1,esk1_0)|~m1_subset_1(esk8_4(esk4_2(a_3_1_zmodul01(esk1_0,X2,X1),X3),esk1_0,X2,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_157, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_239, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_272, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k5_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_194, plain, (r1_struct_0(X1,k1_algstr_0(X2,k1_algstr_0(X2,X3,X5),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~r1_struct_0(X1,X5)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_520, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk3_0,esk2_0))).# trainpos
cnf(c_0_149, plain, (k5_algstr_0(X1,X2,k4_struct_0(X1))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_178, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X3,X2),X3)=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_431, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,X2),X3)|r1_struct_0(X2,esk4_2(a_3_1_zmodul01(esk1_0,X1,X2),X3))|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_142, plain, (k5_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X2))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_292, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),X2)=k5_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_527, negated_conjecture, (a_3_1_zmodul01(esk1_0,esk2_0,esk3_0)!=u1_struct_0(esk3_0))).# trainpos
cnf(c_0_172, plain, (k5_algstr_0(X1,X2,k4_struct_0(X1))=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_203, plain, (k5_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_233, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k1_algstr_0(X2,X3,X5),X5)|~l2_algstr_0(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_291, plain, (k5_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X2))=k5_algstr_0(X1,k1_algstr_0(X1,X3,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_248, plain, (r2_hidden(k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3),a_3_1_zmodul01(X1,X2,X4))|v2_struct_0(X1)|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_212, plain, (k4_algstr_0(X1,k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_183, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k5_algstr_0(X2,X3,X5),X5)|~l2_algstr_0(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_257, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_200, plain, (r2_hidden(k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3),a_3_1_zmodul01(X1,X2,X4))|v2_struct_0(X1)|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_228, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),k4_algstr_0(X1,X2))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_166, plain, (v1_xboole_0(u1_struct_0(X1))|r1_struct_0(X1,k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_491, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,X2),u1_struct_0(X2))|~l1_algstr_0(X2)|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_210, plain, (r1_struct_0(X1,k1_algstr_0(X2,X4,X6))|v2_struct_0(X2)|~r2_hidden(X4,a_3_1_zmodul01(X2,X3,X5))|~r1_struct_0(X1,esk8_4(X4,X2,X3,X5))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X6))|~m1_zmodul01(X1,X2)|~m1_zmodul01(X5,X2)|~m1_subset_1(esk8_4(X4,X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_226, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~r2_hidden(X3,a_3_1_zmodul01(X2,X5,X1))|~r1_struct_0(X1,k1_algstr_0(X2,X5,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_208, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_188, plain, (r1_struct_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_266, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~r2_hidden(X3,a_3_1_zmodul01(X2,X4,X5))|~r1_struct_0(X1,esk8_4(X3,X2,X4,X5))|~m1_zmodul01(X1,X2)|~m1_zmodul01(X5,X2)|~m1_subset_1(esk8_4(X3,X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_286, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~r2_hidden(X3,a_3_1_zmodul01(X2,X4,X1))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_271, plain, (r1_struct_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_217, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_struct_0(X2))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_algstr_0(X2,X3),X3),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_196, plain, (r1_struct_0(X1,k1_algstr_0(X1,X3,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_235, plain, (r1_struct_0(X1,k1_algstr_0(X2,k1_algstr_0(X2,X3,X4),X5))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X4,X5))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_298, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_473, plain, (r1_struct_0(X1,k5_algstr_0(X1,X3,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,k4_algstr_0(X1,X2)),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_232, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,k4_algstr_0(X2,X3)),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_206, plain, (r1_struct_0(X1,k1_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3),X2),u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_352, plain, (k5_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))=k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_484, plain, (r1_struct_0(X1,k5_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_223, plain, (r1_struct_0(X1,k1_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_383, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),X3)=k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_238, plain, (m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_499, negated_conjecture, (r1_struct_0(esk3_0,k5_algstr_0(esk1_0,X1,X2))|~l2_algstr_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X2,u1_struct_0(esk3_0)))).# trainpos
cnf(c_0_408, negated_conjecture, (r1_struct_0(esk3_0,X2)|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk3_0)|~r2_hidden(X2,a_3_1_zmodul01(esk1_0,X1,X3))|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(esk8_4(X2,esk1_0,X1,X3),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_394, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_algstr_0(X1,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_247, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_415, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),X3)=k5_algstr_0(X1,X2,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_416, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_237, plain, (r1_struct_0(X1,k1_algstr_0(X2,X5,k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_255, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_algstr_0(X2,X3)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_244, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),X4))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_290, plain, (m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_297, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_275, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_264, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_285, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_382, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X4),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_309, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_algstr_0(X2,X4)))|v2_struct_0(X2)|~r1_struct_0(X1,X3)|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_405, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_295, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r2_hidden(X3,a_3_1_zmodul01(X2,X4,X1))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_498, negated_conjecture, (a_3_1_zmodul01(esk1_0,X2,X1)=u1_struct_0(X1)|~l1_algstr_0(X1)|~r1_tarski(u1_struct_0(X1),a_3_1_zmodul01(esk1_0,X2,X1))|~r1_struct_0(X1,X2)|~m1_zmodul01(X1,esk1_0)|~m1_subset_1(X2,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_308, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~r1_struct_0(X1,X3)|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_312, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,X5)|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_324, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k5_algstr_0(X2,X6,X5))|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_243, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X3,X2),X3)=k1_algstr_0(X1,k4_algstr_0(X1,X2),k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_349, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k4_struct_0(X2),X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_263, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X3,X2),X3)=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_380, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X1)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_313, plain, (r1_struct_0(X1,k4_algstr_0(X2,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_algstr_0(X2,X3),X4),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_283, plain, (k5_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3),k1_algstr_0(X1,X2,X3))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_300, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X3,X2))=k4_algstr_0(X1,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_351, plain, (k5_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),k1_algstr_0(X1,X2,X3))=k4_algstr_0(X1,X3)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_306, plain, (k5_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X3,X2))=k5_algstr_0(X1,k4_algstr_0(X1,X3),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_427, negated_conjecture, (r1_struct_0(esk3_0,X1)|~l2_algstr_0(esk3_0)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(esk8_4(X1,esk1_0,X1,X2),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_436, negated_conjecture, (r1_struct_0(esk3_0,X1)|~l2_algstr_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_314, plain, (m1_subset_1(k5_algstr_0(X1,k4_algstr_0(X1,X2),X3),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_442, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk3_0))|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(esk4_2(X1,u1_struct_0(esk3_0)),u1_struct_0(esk3_0))|~m1_subset_1(esk4_2(X1,u1_struct_0(esk3_0)),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_325, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X4),X3))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_424, plain, (r1_struct_0(X1,k4_algstr_0(X2,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_462, negated_conjecture, (r1_struct_0(X1,k4_algstr_0(esk1_0,k4_struct_0(esk1_0)))|~m1_zmodul01(X1,esk1_0))).# trainpos
cnf(c_0_350, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,k4_struct_0(X2),X4),X3))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_381, plain, (r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_438, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,X2),u1_struct_0(X2))|~l2_algstr_0(X2)|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_443, negated_conjecture, (m1_subset_1(X1,u1_struct_0(X2))|~l2_algstr_0(X2)|~r2_hidden(X1,a_3_1_zmodul01(esk1_0,X3,X2))|~r1_struct_0(X2,X3)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X3,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_404, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k1_algstr_0(X2,X3,X4),X5))|v2_struct_0(X2)|~r1_struct_0(X5,k4_algstr_0(X2,X4))|~r1_struct_0(X5,k4_algstr_0(X2,X3))|~r1_struct_0(X5,X1)|~m1_zmodul01(X5,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_448, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,X2),X3)|m1_subset_1(esk4_2(a_3_1_zmodul01(esk1_0,X1,X2),X3),u1_struct_0(X2))|~l2_algstr_0(X2)|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_423, plain, (r1_tarski(X1,a_3_1_zmodul01(X2,k1_algstr_0(X2,X3,X4),X5))|v2_struct_0(X2)|~r1_struct_0(X5,esk4_2(X1,a_3_1_zmodul01(X2,k1_algstr_0(X2,X3,X4),X5)))|~r1_struct_0(X5,k4_algstr_0(X2,X4))|~r1_struct_0(X5,k4_algstr_0(X2,X3))|~m1_zmodul01(X5,X2)|~m1_subset_1(esk4_2(X1,a_3_1_zmodul01(X2,k1_algstr_0(X2,X3,X4),X5)),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_451, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(esk4_2(a_3_1_zmodul01(esk1_0,X1,esk3_0),u1_struct_0(esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_432, plain, (r1_tarski(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X4,esk4_2(X1,a_3_1_zmodul01(X2,X3,X4)))|~r1_struct_0(X4,k4_algstr_0(X2,k4_struct_0(X2)))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(esk4_2(X1,a_3_1_zmodul01(X2,X3,X4)),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_455, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_445, plain, (r1_tarski(u1_struct_0(X1),a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|~r1_struct_0(X4,esk4_2(u1_struct_0(X1),a_3_1_zmodul01(X2,X3,X4)))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_450, plain, (r1_tarski(u1_struct_0(X1),a_3_1_zmodul01(X2,X3,X1))|v2_struct_0(X2)|~l1_struct_0(X1)|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).# trainpos
cnf(c_0_505, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,X2)=u1_struct_0(X2)|~l1_algstr_0(X2)|~l1_struct_0(X2)|~r1_struct_0(X2,k4_algstr_0(esk1_0,X1))|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_454, plain, (a_3_1_zmodul01(X1,X2,X3)=u1_struct_0(X3)|v2_struct_0(X1)|~r1_tarski(a_3_1_zmodul01(X1,X2,X3),u1_struct_0(X3))|~l1_struct_0(X3)|~r1_struct_0(X3,k4_algstr_0(X1,X2))|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_458, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,esk3_0)=u1_struct_0(esk3_0)|~l2_algstr_0(esk3_0)|~l1_struct_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_465, negated_conjecture, (a_3_1_zmodul01(esk1_0,k4_struct_0(esk1_0),esk3_0)=u1_struct_0(esk3_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_467, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_470, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_475, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0))).# trainpos
cnf(c_0_480, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~r2_hidden(X1,u1_struct_0(esk3_0)))).# trainpos
cnf(c_0_486, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~r1_struct_0(esk3_0,X1))).# trainpos
cnf(c_0_494, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk3_0)))).# trainpos
cnf(c_0_500, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk3_0)))).# trainpos
cnf(c_0_507, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk3_0)))).# trainpos
cnf(c_0_506, negated_conjecture, (r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk3_0)))).# trainpos
cnf(c_0_518, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,esk3_0)=u1_struct_0(esk3_0)|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk3_0)))).# trainpos
cnf(c_0_526, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,esk3_0)=u1_struct_0(esk3_0)|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_529, negated_conjecture, (~l2_algstr_0(esk3_0))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_61, plain, (o_0_0_xboole_0=k1_xboole_0)).#trainneg
cnf(i_0_83, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_74, plain, (k4_ordinal1=k5_numbers)).#trainneg
cnf(i_0_25, plain, (r2_hidden(X3,X2)|~r2_hidden(X3,X1)|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_70, plain, (v1_xboole_0(o_0_0_xboole_0))).#trainneg
cnf(i_0_78, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_39, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_82, plain, (m1_subset_1(u2_struct_0(X1),u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_64, plain, (u2_struct_0(X1)=k4_struct_0(X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_79, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_48, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_76, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_41, plain, (r1_tarski(X1,X2)|X1!=X2)).#trainneg
cnf(i_0_42, plain, (r1_tarski(X1,X2)|X1!=X2)).#trainneg
cnf(i_0_55, plain, (v2_struct_0(X2)|v2_zmodul01(X1)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X1,X2))).#trainneg
cnf(i_0_54, plain, (v2_struct_0(X2)|v3_zmodul01(X1)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X1,X2))).#trainneg
cnf(i_0_53, plain, (v2_struct_0(X2)|v4_zmodul01(X1)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X1,X2))).#trainneg
cnf(i_0_75, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_77, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_52, plain, (v2_struct_0(X2)|v5_zmodul01(X1)|~v13_algstr_0(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v2_zmodul01(X2)|~v3_zmodul01(X2)|~v4_zmodul01(X2)|~v5_zmodul01(X2)|~l1_zmodul01(X2)|~m1_zmodul01(X1,X2))).#trainneg
cnf(i_0_30, plain, (v3_rlvect_1(X1)|k1_algstr_0(X1,k1_algstr_0(X1,esk5_1(X1),esk6_1(X1)),esk7_1(X1))!=k1_algstr_0(X1,esk5_1(X1),k1_algstr_0(X1,esk6_1(X1),esk7_1(X1)))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_31, plain, (v3_rlvect_1(X1)|m1_subset_1(esk7_1(X1),u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_33, plain, (v3_rlvect_1(X1)|m1_subset_1(esk5_1(X1),u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_32, plain, (v3_rlvect_1(X1)|m1_subset_1(esk6_1(X1),u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_193, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v5_zmodul01(esk1_0)|~v4_zmodul01(esk1_0)|~v3_zmodul01(esk1_0)|~v2_zmodul01(esk1_0))).#trainneg
cnf(i_0_199, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k3_rlvect_1(X2,X5,X3)|~l1_algstr_0(X2)|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_224, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|~r2_hidden(X1,a_3_1_zmodul01(X2,X3,X5))|~r1_struct_0(X4,esk8_4(X1,X2,X3,X5))|~m1_zmodul01(X4,X2)|~m1_zmodul01(X5,X2)|~m1_subset_1(esk8_4(X1,X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_206, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_211, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(esk2_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_205, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k5_zmodul01(X2,X3,X4))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_232, plain, (m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X4))|v2_struct_0(X4)|~l1_algstr_0(X1)|~m1_zmodul01(X1,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X4)|~v5_zmodul01(X4)|~v4_zmodul01(X4)|~v3_zmodul01(X4)|~v2_zmodul01(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v13_algstr_0(X4))).#trainneg
cnf(i_0_249, plain, (m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X3))|v2_struct_0(X3)|~l2_algstr_0(X1)|~m1_zmodul01(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X3)|~v5_zmodul01(X3)|~v4_zmodul01(X3)|~v3_zmodul01(X3)|~v2_zmodul01(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v13_algstr_0(X3))).#trainneg
cnf(i_0_217, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X2))|v2_struct_0(X2)|~l2_struct_0(X1)|~m1_zmodul01(X1,X2)|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_251, negated_conjecture, (v2_zmodul01(esk3_0))).#trainneg
cnf(i_0_252, negated_conjecture, (v3_zmodul01(esk3_0))).#trainneg
cnf(i_0_216, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k4_struct_0(X1),u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_177, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(X1))|v2_struct_0(X1)|~m1_zmodul01(esk1_0,X1)|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_221, plain, (m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X4))|v2_struct_0(X4)|~l2_algstr_0(X1)|~m1_zmodul01(X1,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X4)|~v5_zmodul01(X4)|~v4_zmodul01(X4)|~v3_zmodul01(X4)|~v2_zmodul01(X4)|~v4_rlvect_1(X4)|~v3_rlvect_1(X4)|~v2_rlvect_1(X4)|~v13_algstr_0(X4))).#trainneg
cnf(i_0_248, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k4_algstr_0(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_230, plain, (v1_xboole_0(X1)|r2_hidden(esk9_1(X1),X1))).#trainneg
cnf(i_0_219, plain, (m1_subset_1(u2_struct_0(X1),u1_struct_0(X2))|v2_struct_0(X2)|~l2_struct_0(X1)|~m1_zmodul01(X1,X2)|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_210, plain, (r1_tarski(X1,X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_218, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(u2_struct_0(X1),u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_253, negated_conjecture, (v4_zmodul01(esk3_0))).#trainneg
cnf(i_0_255, negated_conjecture, (v5_zmodul01(esk3_0))).#trainneg
cnf(i_0_240, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_213, plain, (v1_xboole_0(k1_zfmisc_1(X1))|r2_hidden(X2,k1_zfmisc_1(X1))|~r1_tarski(X2,X1))).#trainneg
cnf(i_0_241, plain, (v1_xboole_0(k5_numbers)|r2_hidden(np__1,k5_numbers))).#trainneg
cnf(i_0_254, plain, (v1_xboole_0(k1_numbers)|r2_hidden(np__1,k1_numbers))).#trainneg
cnf(i_0_212, plain, (v1_xboole_0(k1_zfmisc_1(u1_struct_0(X1)))|r2_hidden(k5_zmodul01(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_510, plain, (r1_tarski(k1_xboole_0,X1))).#trainneg
cnf(i_0_372, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0))).#trainneg
cnf(i_0_312, plain, (esk8_4(X2,X1,k4_struct_0(X1),X3)=X2|v2_struct_0(X1)|~l2_algstr_0(X1)|~r2_hidden(X2,a_3_1_zmodul01(X1,k4_struct_0(X1),X3))|~m1_zmodul01(X3,X1)|~m1_subset_1(esk8_4(X2,X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_280, negated_conjecture, (r2_hidden(k3_rlvect_1(esk1_0,esk2_0,X1),u1_struct_0(esk3_0))|r1_struct_0(esk3_0,esk2_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_268, plain, (~v1_xboole_0(u1_struct_0(X1))|~l1_algstr_0(X1)|~r1_struct_0(X1,X2))).#trainneg
cnf(i_0_300, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk1_0))|r1_struct_0(esk3_0,esk2_0))).#trainneg
cnf(i_0_478, plain, (k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),k4_struct_0(X1))=k1_algstr_0(X1,X2,k4_algstr_0(X1,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_371, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|r2_hidden(esk2_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_331, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_384, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,k4_struct_0(X1)),X3)=k1_algstr_0(X1,X2,k4_algstr_0(X1,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_328, plain, (k1_algstr_0(X1,k4_struct_0(X1),X2)=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_490, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X3))|v2_struct_0(X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_zmodul01(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X3)|~v5_zmodul01(X3)|~v4_zmodul01(X3)|~v3_zmodul01(X3)|~v2_zmodul01(X3)|~v4_rlvect_1(X3)|~v4_rlvect_1(X1)|~v3_rlvect_1(X3)|~v3_rlvect_1(X1)|~v2_rlvect_1(X3)|~v13_algstr_0(X3)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_521, plain, (r1_tarski(k5_numbers,X1)|r2_hidden(np__1,k5_numbers))).#trainneg
cnf(i_0_298, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(u1_struct_0(esk1_0)))|r2_hidden(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_struct_0(esk3_0,esk2_0))).#trainneg
cnf(i_0_332, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X4,k4_struct_0(X2)))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X4,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_376, plain, (k5_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),X3)=k1_algstr_0(X1,X2,k4_algstr_0(X1,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,k4_algstr_0(X1,X3)),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_421, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,k4_struct_0(X1)),k1_algstr_0(X1,X3,X4))=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_272, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X3,X1)|~m1_zmodul01(X3,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_301, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_367, plain, (k1_algstr_0(X1,k4_struct_0(X1),k5_algstr_0(X1,X2,X3))=k5_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_279, plain, (m1_subset_1(esk8_4(k3_rlvect_1(X1,X2,X3),X1,X2,X4),u1_struct_0(X1))|v2_struct_0(X1)|~r1_struct_0(X4,X3)|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_403, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k1_algstr_0(X1,X4,k4_struct_0(X1)))=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4)),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_523, plain, (r1_tarski(k1_numbers,X1)|r2_hidden(np__1,k1_numbers))).#trainneg
cnf(i_0_462, plain, (v1_xboole_0(u1_struct_0(X1))|r1_struct_0(X1,k4_struct_0(X1))|~l2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_498, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k4_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_378, plain, (k5_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),X3)=k1_algstr_0(X1,X2,k4_algstr_0(X1,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_297, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_493, plain, (v1_xboole_0(u1_struct_0(X1))|r1_struct_0(X1,k4_algstr_0(X1,X2))|~l2_algstr_0(X1)|~l1_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_302, plain, (r2_hidden(k3_rlvect_1(X1,X2,X3),a_3_1_zmodul01(X1,X3,X4))|v2_struct_0(X1)|~r1_struct_0(X4,X2)|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_318, plain, (k5_zmodul01(X1,X2,X3)=u1_struct_0(X1)|v2_struct_0(X1)|~r1_tarski(u1_struct_0(X1),k5_zmodul01(X1,X2,X3))|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_390, plain, (r1_tarski(k5_zmodul01(X1,X2,X3),X4)|m1_subset_1(esk4_2(k5_zmodul01(X1,X2,X3),X4),u1_struct_0(X1))|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_511, negated_conjecture, (r1_tarski(u1_struct_0(esk1_0),X1)|r2_hidden(esk2_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_264, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|m1_subset_1(esk8_4(X1,esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~r2_hidden(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_354, plain, (k1_algstr_0(X1,X2,k5_algstr_0(X1,k4_struct_0(X1),X3))=k1_algstr_0(X1,X2,k4_algstr_0(X1,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_289, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,X3)|~l1_algstr_0(X2)|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_307, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,esk8_4(X1,X2,X3,X4),X5))|v2_struct_0(X2)|~r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|~r1_struct_0(X5,X3)|~m1_zmodul01(X5,X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(esk8_4(X1,X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_294, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k1_algstr_0(X2,X4,X5)))|v2_struct_0(X2)|~l1_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~r1_struct_0(X1,X5)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_361, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X4,X5)|~l2_algstr_0(X2)|~r1_struct_0(X3,k1_algstr_0(X2,X4,X5))|~m1_zmodul01(X3,X2)|~m1_subset_1(k1_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_382, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k1_algstr_0(X2,X3,X5),X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_algstr_0(X2,X6)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,k4_algstr_0(X2,X6)),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_519, plain, (v1_xboole_0(k1_zfmisc_1(X1))|r2_hidden(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_336, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),X1)|r1_struct_0(esk3_0,esk2_0)|m1_subset_1(esk8_4(esk4_2(u1_struct_0(esk3_0),X1),esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_479, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,k4_algstr_0(X1,X3)),X4)=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),k4_algstr_0(X1,X4))|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_319, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,k5_zmodul01(X2,X3,X4))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_330, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k1_algstr_0(X1,X3,k4_struct_0(X1)),X4))=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_419, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k1_algstr_0(X1,X4,X5))=k1_algstr_0(X1,X2,k1_algstr_0(X1,k1_algstr_0(X1,X3,X4),X5))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_476, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k5_algstr_0(X2,X3,X5),X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_algstr_0(X2,X5),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,k4_algstr_0(X2,X5),X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_423, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k1_algstr_0(X1,k4_struct_0(X1),X4))=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_333, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k1_algstr_0(X1,X4,k4_struct_0(X1))))=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_311, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,esk8_4(X4,X1,k1_algstr_0(X1,X2,X3),X5)))=X4|v2_struct_0(X1)|~l1_algstr_0(X1)|~r2_hidden(X4,a_3_1_zmodul01(X1,k1_algstr_0(X1,X2,X3),X5))|~m1_zmodul01(X5,X1)|~m1_subset_1(esk8_4(X4,X1,k1_algstr_0(X1,X2,X3),X5),u1_struct_0(X1))|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_271, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k1_algstr_0(X2,X3,X4),X5))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k1_algstr_0(X2,X4,X6))|~l1_algstr_0(X2)|~r1_struct_0(X5,X6)|~m1_zmodul01(X5,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_501, plain, (v1_xboole_0(u1_struct_0(X1))|r1_struct_0(X1,esk9_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_514, plain, (v1_xboole_0(u1_struct_0(X1))|r1_struct_0(X1,u2_struct_0(X1))|~l2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_364, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4)))=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_314, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,esk8_4(X3,X1,X2,X4),X5))=k1_algstr_0(X1,X3,X5)|v2_struct_0(X1)|~l1_algstr_0(X1)|~r2_hidden(X3,a_3_1_zmodul01(X1,X2,X4))|~m1_zmodul01(X4,X1)|~m1_subset_1(esk8_4(X3,X1,X2,X4),u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_412, plain, (r1_struct_0(X1,k1_algstr_0(X2,k1_algstr_0(X2,X3,X4),k1_algstr_0(X2,X5,X6)))|v2_struct_0(X2)|~l1_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k1_algstr_0(X2,X4,X5)))|~r1_struct_0(X1,X6)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,k1_algstr_0(X2,X4,X5)),u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_383, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X4,k4_algstr_0(X2,X5)))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X4,k4_algstr_0(X2,X5)),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_481, plain, (k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),k4_algstr_0(X1,X4))=k1_algstr_0(X1,X2,k5_algstr_0(X1,k4_algstr_0(X1,X3),X4))|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_335, plain, (r1_tarski(a_3_1_zmodul01(X1,X2,X3),X4)|m1_subset_1(esk8_4(esk4_2(a_3_1_zmodul01(X1,X2,X3),X4),X1,X2,X3),u1_struct_0(X5))|v2_struct_0(X5)|v2_struct_0(X1)|~m1_zmodul01(X1,X5)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X5)|~l1_zmodul01(X1)|~v5_zmodul01(X5)|~v5_zmodul01(X1)|~v4_zmodul01(X5)|~v4_zmodul01(X1)|~v3_zmodul01(X5)|~v3_zmodul01(X1)|~v2_zmodul01(X5)|~v2_zmodul01(X1)|~v4_rlvect_1(X5)|~v4_rlvect_1(X1)|~v3_rlvect_1(X5)|~v3_rlvect_1(X1)|~v2_rlvect_1(X5)|~v2_rlvect_1(X1)|~v13_algstr_0(X5)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_477, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X4),X5))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_algstr_0(X2,X4),X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_381, plain, (k1_algstr_0(X1,k4_struct_0(X1),k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,k4_algstr_0(X1,X4)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_356, plain, (k1_algstr_0(X1,X2,esk8_4(X3,X1,k4_struct_0(X1),X4))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~r2_hidden(X3,a_3_1_zmodul01(X1,k4_struct_0(X1),X4))|~m1_zmodul01(X4,X1)|~m1_subset_1(esk8_4(X3,X1,k4_struct_0(X1),X4),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_440, plain, (r1_tarski(k1_zfmisc_1(X2),X3)|m1_subset_1(X1,X2)|~r2_hidden(X1,esk4_2(k1_zfmisc_1(X2),X3)))).#trainneg
cnf(i_0_439, plain, (r1_tarski(esk4_2(k1_zfmisc_1(X1),X2),X1)|r1_tarski(k1_zfmisc_1(X1),X2))).#trainneg
cnf(i_0_444, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k4_struct_0(X2),u1_struct_0(X1))|v2_struct_0(X1)|~l2_struct_0(X2)|~m1_zmodul01(X2,X1)|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_458, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(X3,u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~r2_hidden(X3,a_3_1_zmodul01(X1,X2,X4))|~m1_zmodul01(X4,X1)|~m1_subset_1(esk8_4(X3,X1,X2,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_437, plain, (m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X3))|v2_struct_0(X3)|v2_struct_0(X4)|~l2_algstr_0(X1)|~m1_zmodul01(X4,X3)|~m1_zmodul01(X1,X4)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X3)|~l1_zmodul01(X4)|~v5_zmodul01(X3)|~v5_zmodul01(X4)|~v4_zmodul01(X3)|~v4_zmodul01(X4)|~v3_zmodul01(X3)|~v3_zmodul01(X4)|~v2_zmodul01(X3)|~v2_zmodul01(X4)|~v4_rlvect_1(X3)|~v4_rlvect_1(X4)|~v3_rlvect_1(X3)|~v3_rlvect_1(X4)|~v2_rlvect_1(X3)|~v2_rlvect_1(X4)|~v13_algstr_0(X3)|~v13_algstr_0(X4))).#trainneg
cnf(i_0_391, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k1_algstr_0(X2,X3,X4),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X2)|~m1_zmodul01(X2,X1)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_393, plain, (m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X4))|v2_struct_0(X4)|v2_struct_0(X5)|~l1_algstr_0(X1)|~m1_zmodul01(X5,X4)|~m1_zmodul01(X1,X5)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X4)|~l1_zmodul01(X5)|~v5_zmodul01(X4)|~v5_zmodul01(X5)|~v4_zmodul01(X4)|~v4_zmodul01(X5)|~v3_zmodul01(X4)|~v3_zmodul01(X5)|~v2_zmodul01(X4)|~v2_zmodul01(X5)|~v4_rlvect_1(X4)|~v4_rlvect_1(X5)|~v3_rlvect_1(X4)|~v3_rlvect_1(X5)|~v2_rlvect_1(X4)|~v2_rlvect_1(X5)|~v13_algstr_0(X4)|~v13_algstr_0(X5))).#trainneg
cnf(i_0_334, plain, (v1_xboole_0(u1_struct_0(X1))|r1_tarski(a_3_1_zmodul01(X1,X2,X3),X4)|r2_hidden(esk8_4(esk4_2(a_3_1_zmodul01(X1,X2,X3),X4),X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_446, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X3)|~l2_struct_0(X1)|~m1_zmodul01(X3,X2)|~m1_zmodul01(X1,X3)|~l1_zmodul01(X2)|~l1_zmodul01(X3)|~v5_zmodul01(X2)|~v5_zmodul01(X3)|~v4_zmodul01(X2)|~v4_zmodul01(X3)|~v3_zmodul01(X2)|~v3_zmodul01(X3)|~v2_zmodul01(X2)|~v2_zmodul01(X3)|~v4_rlvect_1(X2)|~v4_rlvect_1(X3)|~v3_rlvect_1(X2)|~v3_rlvect_1(X3)|~v2_rlvect_1(X2)|~v2_rlvect_1(X3)|~v13_algstr_0(X2)|~v13_algstr_0(X3))).#trainneg
cnf(i_0_277, plain, (m1_subset_1(k3_rlvect_1(X1,X2,X3),a_3_1_zmodul01(X1,X2,X4))|v2_struct_0(X1)|~r1_struct_0(X4,X3)|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_435, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k4_algstr_0(X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X2)|~m1_zmodul01(X2,X1)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_486, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k5_algstr_0(X2,X3,X4),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X2)|~m1_zmodul01(X2,X1)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_506, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(u2_struct_0(X2),u1_struct_0(X1))|v2_struct_0(X1)|~l2_struct_0(X2)|~m1_zmodul01(X2,X1)|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_508, plain, (m1_subset_1(u2_struct_0(X1),u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X3)|~l2_struct_0(X1)|~m1_zmodul01(X3,X2)|~m1_zmodul01(X1,X3)|~l1_zmodul01(X2)|~l1_zmodul01(X3)|~v5_zmodul01(X2)|~v5_zmodul01(X3)|~v4_zmodul01(X2)|~v4_zmodul01(X3)|~v3_zmodul01(X2)|~v3_zmodul01(X3)|~v2_zmodul01(X2)|~v2_zmodul01(X3)|~v4_rlvect_1(X2)|~v4_rlvect_1(X3)|~v3_rlvect_1(X2)|~v3_rlvect_1(X3)|~v2_rlvect_1(X2)|~v2_rlvect_1(X3)|~v13_algstr_0(X2)|~v13_algstr_0(X3))).#trainneg
cnf(i_0_488, plain, (m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X4))|v2_struct_0(X4)|v2_struct_0(X5)|~l2_algstr_0(X1)|~m1_zmodul01(X5,X4)|~m1_zmodul01(X1,X5)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X4)|~l1_zmodul01(X5)|~v5_zmodul01(X4)|~v5_zmodul01(X5)|~v4_zmodul01(X4)|~v4_zmodul01(X5)|~v3_zmodul01(X4)|~v3_zmodul01(X5)|~v2_zmodul01(X4)|~v2_zmodul01(X5)|~v4_rlvect_1(X4)|~v4_rlvect_1(X5)|~v3_rlvect_1(X4)|~v3_rlvect_1(X5)|~v2_rlvect_1(X4)|~v2_rlvect_1(X5)|~v13_algstr_0(X4)|~v13_algstr_0(X5))).#trainneg
cnf(i_0_450, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4)),u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_379, plain, (m1_subset_1(k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4),u1_struct_0(X1))|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,k4_algstr_0(X1,X4)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_469, plain, (m1_subset_1(k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X4),u1_struct_0(X1))|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,X3),X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_527, plain, (m1_subset_1(k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k1_algstr_0(X1,X4,X5))),u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_459, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4),u1_struct_0(X1))|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,k4_algstr_0(X1,X4)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_404, plain, (m1_subset_1(k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k1_algstr_0(X1,X4,X5)),u1_struct_0(X1))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_322, plain, (v1_xboole_0(k1_zfmisc_1(u1_struct_0(X1)))|r2_hidden(a_3_1_zmodul01(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_355, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),X3))=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_struct_0(X1)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_401, plain, (m1_subset_1(X3,u1_struct_0(X5))|v2_struct_0(X5)|v2_struct_0(X1)|~l1_algstr_0(X1)|~r2_hidden(X3,a_3_1_zmodul01(X1,X2,X4))|~m1_zmodul01(X1,X5)|~m1_zmodul01(X4,X1)|~m1_subset_1(esk8_4(X3,X1,X2,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X5)|~l1_zmodul01(X1)|~v5_zmodul01(X5)|~v5_zmodul01(X1)|~v4_zmodul01(X5)|~v4_zmodul01(X1)|~v3_zmodul01(X5)|~v3_zmodul01(X1)|~v2_zmodul01(X5)|~v2_zmodul01(X1)|~v4_rlvect_1(X5)|~v4_rlvect_1(X1)|~v3_rlvect_1(X5)|~v3_rlvect_1(X1)|~v2_rlvect_1(X5)|~v2_rlvect_1(X1)|~v13_algstr_0(X5)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_424, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k1_algstr_0(X1,k4_algstr_0(X1,X4),X5))=k1_algstr_0(X1,k1_algstr_0(X1,X2,k5_algstr_0(X1,X3,X4)),X5)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_380, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k1_algstr_0(X1,X3,k4_algstr_0(X1,X4)),X5))=k1_algstr_0(X1,k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4),X5)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,k4_algstr_0(X1,X4)),u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_470, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X3),X4),X5))=k1_algstr_0(X1,k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X4),X5)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,X3),X4),u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_rlvect_1(X1))).#trainneg
cnf(i_0_680, negated_conjecture, (r1_struct_0(esk3_0,k3_rlvect_1(esk1_0,esk2_0,X1))|r1_struct_0(esk3_0,esk2_0)|~l1_struct_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_685, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk3_0))|r1_struct_0(esk3_0,esk2_0)|~r2_hidden(X1,a_3_1_zmodul01(esk1_0,esk2_0,X2))|~r1_struct_0(esk3_0,esk8_4(X1,esk1_0,esk2_0,X2))|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(esk8_4(X1,esk1_0,esk2_0,X2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_550, plain, (m1_subset_1(esk8_4(X1,X2,X3,X4),u1_struct_0(X4))|v2_struct_0(X2)|~l1_algstr_0(X4)|~r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_641, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0))).#trainneg
cnf(i_0_642, negated_conjecture, (r1_struct_0(esk1_0,esk2_0)|~l1_algstr_0(esk1_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_824, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X2))|v2_struct_0(X2)|~l2_algstr_0(X1)|~m1_zmodul01(X1,X2)|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_662, negated_conjecture, (r2_hidden(k1_algstr_0(esk1_0,esk2_0,X1),u1_struct_0(esk3_0))|r1_struct_0(esk3_0,esk2_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_636, plain, (X1=k1_xboole_0|~r1_tarski(X1,k1_xboole_0))).#trainneg
cnf(i_0_773, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_559, plain, (r1_tarski(a_3_1_zmodul01(X2,X3,X1),X4)|r1_struct_0(X1,esk4_2(a_3_1_zmodul01(X2,X3,X1),X4))|v2_struct_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_861, plain, (v1_xboole_0(u1_struct_0(X1))|r1_struct_0(X1,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_769, plain, (k1_algstr_0(X1,k4_struct_0(X1),esk4_2(a_3_1_zmodul01(X1,X2,X3),X4))=esk4_2(a_3_1_zmodul01(X1,X2,X3),X4)|r1_tarski(a_3_1_zmodul01(X1,X2,X3),X4)|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2192, plain, (k1_algstr_0(X1,X2,X3)=k1_algstr_0(X1,X3,X2)|v2_struct_0(X1)|~r2_hidden(X3,a_3_1_zmodul01(X1,X2,X4))|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_639, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_1043, negated_conjecture, (r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_708, plain, (m1_subset_1(esk8_4(X1,X2,X1,X3),u1_struct_0(X2))|v2_struct_0(X2)|~m1_zmodul01(X3,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_2563, plain, (v1_xboole_0(u1_struct_0(X1))|r1_tarski(a_3_1_zmodul01(X1,X2,X3),X4)|r1_struct_0(X1,esk8_4(esk4_2(a_3_1_zmodul01(X1,X2,X3),X4),X1,X2,X3))|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_798, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l1_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_934, plain, (r1_tarski(X1,a_3_1_zmodul01(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r1_struct_0(X3,esk4_2(X1,a_3_1_zmodul01(X2,k4_struct_0(X2),X3)))|~m1_zmodul01(X3,X2)|~m1_subset_1(esk4_2(X1,a_3_1_zmodul01(X2,k4_struct_0(X2),X3)),u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_776, negated_conjecture, (u1_struct_0(esk1_0)=k1_xboole_0|r1_struct_0(esk1_0,esk2_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_702, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|r1_struct_0(esk3_0,esk2_0)|~r1_tarski(u1_struct_0(esk1_0),u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_980, plain, (m1_subset_1(esk8_4(X3,X1,X2,X5),u1_struct_0(X1))|v2_struct_0(X1)|~r2_hidden(X3,a_3_1_zmodul01(X1,X2,X4))|~r1_struct_0(X5,esk8_4(X3,X1,X2,X4))|~m1_zmodul01(X5,X1)|~m1_zmodul01(X4,X1)|~m1_subset_1(esk8_4(X3,X1,X2,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_785, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l1_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1219, negated_conjecture, (r2_hidden(k1_algstr_0(esk1_0,X1,esk2_0),u1_struct_0(esk3_0))|r1_struct_0(esk3_0,esk2_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_870, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~l1_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X4,k4_struct_0(X2)))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_723, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X3,X2),X4)=k1_algstr_0(X1,X2,k5_algstr_0(X1,X3,X4))|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_1177, plain, (k1_algstr_0(X1,k4_struct_0(X1),k5_algstr_0(X1,X2,X3))=k1_algstr_0(X1,X2,k4_algstr_0(X1,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1206, negated_conjecture, (r2_hidden(k3_rlvect_1(esk1_0,X1,esk2_0),u1_struct_0(esk3_0))|r1_struct_0(esk3_0,esk2_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_655, plain, (~v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)|~r1_struct_0(X1,X2))).#trainneg
cnf(i_0_659, plain, (m1_subset_1(esk8_4(k1_algstr_0(X1,X2,X3),X1,X2,X4),u1_struct_0(X1))|v2_struct_0(X1)|~r1_struct_0(X4,X3)|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_672, plain, (r2_hidden(k1_algstr_0(X1,X2,X3),a_3_1_zmodul01(X1,X2,X4))|v2_struct_0(X1)|~l2_algstr_0(X1)|~r1_struct_0(X4,k1_algstr_0(X1,k4_struct_0(X1),X3))|~m1_zmodul01(X4,X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1165, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,k4_struct_0(X1)))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_638, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_943, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk3_0,esk2_0)|m1_subset_1(u2_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_struct_0(esk3_0))).#trainneg
cnf(i_0_1369, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_679, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|~v1_xboole_0(u1_struct_0(esk3_0))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1354, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k1_algstr_0(X2,X5,X4)))|v2_struct_0(X2)|~l1_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~r1_struct_0(X1,X5)|~m1_zmodul01(X1,X2)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1331, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~l1_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_678, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|m1_subset_1(k3_rlvect_1(esk1_0,esk2_0,X1),u1_struct_0(esk3_0))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_850, plain, (X1=k5_numbers|r2_hidden(np__1,k5_numbers)|~r1_tarski(X1,k5_numbers))).#trainneg
cnf(i_0_946, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk3_0,esk2_0)|m1_subset_1(k4_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_struct_0(esk3_0))).#trainneg
cnf(i_0_1347, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1483, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k1_algstr_0(X2,X3,X5),X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_689, negated_conjecture, (k3_rlvect_1(esk1_0,X1,X2)=k1_algstr_0(esk3_0,X2,X1)|~l1_algstr_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X2,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_1765, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k4_struct_0(X1))=k5_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,X2),k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_949, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|m1_subset_1(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1))).#trainneg
cnf(i_0_686, negated_conjecture, (k3_rlvect_1(esk3_0,X2,X1)=k3_rlvect_1(esk1_0,X1,X2)|~l1_algstr_0(esk3_0)|~m1_subset_1(X2,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_722, plain, (k1_algstr_0(X1,X2,k5_algstr_0(X1,k4_struct_0(X1),X3))=k5_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1031, plain, (X1=k1_numbers|r2_hidden(np__1,k1_numbers)|~r1_tarski(X1,k1_numbers))).#trainneg
cnf(i_0_941, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk3_0,esk2_0)|m1_subset_1(esk9_1(u1_struct_0(esk3_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_696, negated_conjecture, (m1_subset_1(k5_algstr_0(esk1_0,X1,X2),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~m1_subset_1(X2,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_1284, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_707, plain, (v2_struct_0(X1)|~v1_xboole_0(a_3_1_zmodul01(X1,X2,X3))|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_970, plain, (k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1197, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k1_algstr_0(X1,k4_struct_0(X1),X4))=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1489, plain, (k1_zfmisc_1(X1)=k1_xboole_0|r2_hidden(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_857, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk3_0,k5_algstr_0(esk1_0,X1,X2))|~l2_algstr_0(esk3_0)|~m1_subset_1(X2,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_1976, plain, (r2_hidden(k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4)),a_3_1_zmodul01(X1,k1_algstr_0(X1,X2,X3),X5))|v2_struct_0(X1)|~r1_struct_0(X5,X4)|~m1_zmodul01(X5,X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_907, plain, (r1_struct_0(X1,k1_algstr_0(X2,k1_algstr_0(X2,X3,X4),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|~r1_struct_0(X1,k1_algstr_0(X2,X4,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(k1_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1659, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4))=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3469, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|~l1_struct_0(esk3_0))).#trainneg
cnf(i_0_3682, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_3471, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk3_0))|r1_struct_0(esk3_0,esk2_0)|~r2_hidden(X1,a_3_1_zmodul01(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk8_4(X1,esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3842, plain, (r1_struct_0(X1,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_3491, negated_conjecture, (m1_subset_1(k1_algstr_0(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3683, negated_conjecture, (r1_struct_0(esk1_0,esk2_0)|~l2_algstr_0(esk1_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_3927, negated_conjecture, (r1_struct_0(esk1_0,k4_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_3652, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|m1_subset_1(k1_algstr_0(esk1_0,esk2_0,X1),u1_struct_0(esk3_0))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3914, negated_conjecture, (k1_algstr_0(esk1_0,esk2_0,X1)=k1_algstr_0(esk1_0,X1,esk2_0)|r1_struct_0(esk3_0,esk2_0)|~r2_hidden(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_4814, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|m1_subset_1(X1,u1_struct_0(esk3_0))|~r2_hidden(X1,a_3_1_zmodul01(esk1_0,esk2_0,X2))|~r1_struct_0(esk3_0,esk8_4(X1,esk1_0,esk2_0,X2))|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(esk8_4(X1,esk1_0,esk2_0,X2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5177, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_3678, plain, (k1_numbers=k1_xboole_0|r2_hidden(np__1,k1_numbers))).#trainneg
cnf(i_0_3925, negated_conjecture, (m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_3655, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|m1_subset_1(k1_algstr_0(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4044, plain, (k1_algstr_0(X1,X3,k5_algstr_0(X1,X2,X4))=k1_algstr_0(X1,X2,k5_algstr_0(X1,X3,X4))|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_3509, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|v2_struct_0(X4)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X4)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v4_rlvect_1(X4)|~v3_rlvect_1(X2)|~v3_rlvect_1(X4)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)|~v13_algstr_0(X4))).#trainneg
cnf(i_0_3746, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))=k1_algstr_0(X1,X4,k1_algstr_0(X1,X3,X2))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_4690, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_3679, plain, (k1_xboole_0=k5_numbers|r2_hidden(np__1,k5_numbers))).#trainneg
cnf(i_0_3813, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3),k1_algstr_0(X1,X2,X4))=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_4819, plain, (r1_struct_0(X1,k1_algstr_0(X2,X4,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~l1_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X4,k4_algstr_0(X2,X3)),u1_struct_0(X2))|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_5701, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_4643, plain, (r1_tarski(X1,u1_struct_0(X2))|v2_struct_0(X2)|~l2_algstr_0(X2)|~m1_subset_1(esk4_2(X1,u1_struct_0(X2)),u1_struct_0(X2))|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_3998, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|v2_struct_0(X4)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X4)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,k4_struct_0(X2)),u1_struct_0(X4))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v4_rlvect_1(X4)|~v3_rlvect_1(X2)|~v3_rlvect_1(X4)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)|~v13_algstr_0(X4))).#trainneg
cnf(i_0_5234, plain, (k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_5309, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k1_algstr_0(X2,X4,X5)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X2))|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_5429, plain, (k1_zfmisc_1(X1)=k1_xboole_0|m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4951, negated_conjecture, (u1_struct_0(esk1_0)=k5_numbers|r2_hidden(esk2_0,u1_struct_0(esk1_0))|r2_hidden(np__1,k5_numbers))).#trainneg
cnf(i_0_5024, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,k5_algstr_0(X2,X3,X6))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_3887, plain, (v1_xboole_0(u1_struct_0(X1))|r1_tarski(a_3_1_zmodul01(X1,X2,X3),X4)|r1_struct_0(X1,esk4_2(a_3_1_zmodul01(X1,X2,X3),X4))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(esk4_2(a_3_1_zmodul01(X1,X2,X3),X4),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_4117, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_algstr_0(X2,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_struct_0(X2))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_4439, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,X3,X4))=k1_algstr_0(X1,X3,k5_algstr_0(X1,X4,X2))|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_5761, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_5760, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|~l1_algstr_0(esk3_0))).#trainneg
cnf(i_0_5772, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_6310, negated_conjecture, (r2_hidden(X1,a_3_1_zmodul01(esk1_0,X2,X3))|X1!=k1_algstr_0(esk1_0,X2,X4)|~r1_struct_0(X3,X4)|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8151, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X4)|~l2_algstr_0(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_6378, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|m1_subset_1(X1,u1_struct_0(esk3_0))|~r2_hidden(X1,a_3_1_zmodul01(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk8_4(X1,esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5796, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,a_3_1_zmodul01(esk1_0,X2,X3))|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(esk8_4(X1,esk1_0,X2,X3),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6308, negated_conjecture, (r2_hidden(k1_algstr_0(esk1_0,X1,X2),a_3_1_zmodul01(esk1_0,X1,X3))|~r1_struct_0(X3,X2)|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8301, negated_conjecture, (u1_struct_0(esk1_0)=k5_numbers|r2_hidden(np__1,k5_numbers)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8302, negated_conjecture, (u1_struct_0(esk1_0)=k5_numbers|r2_hidden(np__1,k5_numbers)|r1_struct_0(esk1_0,esk2_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_7574, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X3,X2),X4)=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_8153, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l1_algstr_0(X2)|~l2_algstr_0(X4)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X4))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_8048, plain, (r1_tarski(u1_struct_0(X1),u1_struct_0(X2))|v2_struct_0(X2)|~l2_algstr_0(X2)|~m1_zmodul01(X1,X2)|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_637, plain, (v1_xboole_0(k1_zfmisc_1(X1))|r2_hidden(k1_xboole_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5201, negated_conjecture, (u1_struct_0(esk1_0)=k1_numbers|r2_hidden(esk2_0,u1_struct_0(esk1_0))|r2_hidden(np__1,k1_numbers))).#trainneg
cnf(i_0_5911, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X3,X4),X2)=k1_algstr_0(X1,X3,k1_algstr_0(X1,X2,X4))|~l1_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_5945, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,k1_algstr_0(X2,X3,X6))|~l1_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_8319, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1250, negated_conjecture, (X1=u1_struct_0(esk1_0)|r2_hidden(esk2_0,u1_struct_0(esk1_0))|~r1_tarski(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8367, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X3,X4),k4_algstr_0(X1,X2))=k1_algstr_0(X1,X3,k5_algstr_0(X1,X4,X2))|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_4001, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|m1_subset_1(k1_algstr_0(esk1_0,X1,esk2_0),u1_struct_0(esk3_0))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3977, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|r2_hidden(esk2_0,u1_struct_0(esk1_0))|r1_struct_0(esk3_0,esk2_0))).#trainneg
cnf(i_0_854, negated_conjecture, (k1_zfmisc_1(u1_struct_0(esk1_0))=k1_xboole_0|r2_hidden(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_struct_0(esk3_0,esk2_0))).#trainneg
cnf(i_0_3494, plain, (r1_tarski(a_3_1_zmodul01(X1,X2,X3),X4)|m1_subset_1(esk8_4(esk4_2(a_3_1_zmodul01(X1,X2,X3),X4),X1,X2,X3),u1_struct_0(X3))|v2_struct_0(X1)|~l1_algstr_0(X3)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_9579, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,u1_struct_0(X3))|~m1_zmodul01(X3,X2)|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_9578, plain, (r2_hidden(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,u1_struct_0(X3))|~m1_zmodul01(X3,X2)|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_10503, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|r1_struct_0(esk3_0,esk2_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9681, negated_conjecture, (u1_struct_0(esk1_0)=k1_numbers|r2_hidden(np__1,k1_numbers)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9070, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|v2_struct_0(X4)|X1!=k1_algstr_0(X2,X3,k4_struct_0(X4))|~l2_algstr_0(X4)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v4_rlvect_1(X4)|~v3_rlvect_1(X2)|~v3_rlvect_1(X4)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)|~v13_algstr_0(X4))).#trainneg
cnf(i_0_8649, negated_conjecture, (r2_hidden(X1,a_3_1_zmodul01(esk1_0,X2,X3))|~r2_hidden(X1,a_3_1_zmodul01(esk1_0,X2,X4))|~r1_struct_0(X3,esk8_4(X1,esk1_0,X2,X4))|~m1_zmodul01(X3,esk1_0)|~m1_zmodul01(X4,esk1_0)|~m1_subset_1(esk8_4(X1,esk1_0,X2,X4),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9099, negated_conjecture, (m1_subset_1(esk8_4(k1_algstr_0(esk1_0,X1,X2),esk1_0,X1,X3),u1_struct_0(esk1_0))|~r1_struct_0(X3,X2)|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10050, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|v2_struct_0(X4)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X4)|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v4_rlvect_1(X4)|~v3_rlvect_1(X2)|~v3_rlvect_1(X4)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)|~v13_algstr_0(X4))).#trainneg
cnf(i_0_9580, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|r1_struct_0(esk3_0,esk2_0)|~m1_zmodul01(esk1_0,esk3_0))).#trainneg
cnf(i_0_9682, negated_conjecture, (u1_struct_0(esk1_0)=k1_numbers|r2_hidden(np__1,k1_numbers)|r1_struct_0(esk1_0,esk2_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_9320, plain, (k1_algstr_0(X1,X3,k1_algstr_0(X1,X2,X4))=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_8802, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3),X2)=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_8700, plain, (k1_algstr_0(X1,X4,k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_9678, plain, (k1_zfmisc_1(X1)=k1_xboole_0|r2_hidden(k1_xboole_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_10504, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|r1_struct_0(esk1_0,esk2_0)|r1_struct_0(esk3_0,esk2_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_10112, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k5_algstr_0(X2,X4,X5)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_4895, plain, (k5_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3),X2)=k1_algstr_0(X1,k4_algstr_0(X1,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_7437, plain, (r1_struct_0(X1,k1_algstr_0(X2,k1_algstr_0(X2,X4,X3),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,k1_algstr_0(X2,X4,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(k1_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_3810, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),k1_algstr_0(X1,X2,X4))=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X4)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_10059, negated_conjecture, (u1_struct_0(X1)=u1_struct_0(esk1_0)|r2_hidden(esk2_0,u1_struct_0(esk1_0))|~m1_zmodul01(X1,esk1_0))).#trainneg
cnf(i_0_4072, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),k5_algstr_0(X1,X2,X4))=k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),X4)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_1036, plain, (r2_hidden(k4_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_844, plain, (a_3_1_zmodul01(X1,X2,X3)=u1_struct_0(X1)|v2_struct_0(X1)|~r1_tarski(u1_struct_0(X1),a_3_1_zmodul01(X1,X2,X3))|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_10572, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r2_hidden(esk2_0,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_12707, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|r2_hidden(esk2_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4648, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk3_0))|r2_hidden(u2_struct_0(esk3_0),u1_struct_0(esk1_0))|r1_struct_0(esk3_0,esk2_0)|~l2_struct_0(esk3_0))).#trainneg
cnf(i_0_10546, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(X1))|r1_struct_0(esk3_0,esk2_0)|v2_struct_0(X1)|~m1_zmodul01(esk3_0,X1)|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_12567, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X4,X3),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X4,X5))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_12761, plain, (r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X4,X3),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X4,X5))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_10637, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_12623, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X4,X3),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,k1_algstr_0(X2,X4,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k5_algstr_0(X2,X4,X3),X5),u1_struct_0(X2))|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(k1_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_12307, plain, (k1_zfmisc_1(X1)=k1_xboole_0|m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4952, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk3_0))|r2_hidden(k4_struct_0(esk3_0),u1_struct_0(esk1_0))|r1_struct_0(esk3_0,esk2_0)|~l2_struct_0(esk3_0))).#trainneg
cnf(i_0_10576, negated_conjecture, (r1_struct_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~r2_hidden(X2,a_3_1_zmodul01(esk1_0,X3,X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,esk1_0)|~m1_zmodul01(X4,esk1_0)|~m1_subset_1(esk8_4(X2,esk1_0,X3,X4),u1_struct_0(esk1_0))|~m1_subset_1(esk8_4(X2,esk1_0,X3,X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_13076, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X4),X3)=k1_algstr_0(X1,X2,k1_algstr_0(X1,X3,X4))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_11291, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X4,X3))=k1_algstr_0(X1,X3,k1_algstr_0(X1,X2,X4))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_13092, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk3_0))|~l1_algstr_0(esk3_0))).#trainneg
cnf(i_0_13098, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13093, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_13125, negated_conjecture, (r1_struct_0(esk3_0,esk2_0)|r1_struct_0(X1,esk2_0)|v2_struct_0(X1)|~l1_struct_0(X1)|~m1_zmodul01(esk3_0,X1)|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_13091, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),X1)|r2_hidden(esk2_0,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_13818, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_13275, plain, (k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3)=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_13099, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|r1_struct_0(esk1_0,esk2_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_13090, negated_conjecture, (u1_struct_0(esk3_0)=k1_xboole_0|r2_hidden(esk2_0,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_13820, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_13875, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_14180, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,X3,X2))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_704, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|r1_struct_0(esk3_0,esk2_0)|~r2_hidden(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_13128, negated_conjecture, (r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|r1_struct_0(esk3_0,esk2_0)|~l2_algstr_0(esk1_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,esk2_0))|~m1_subset_1(k4_algstr_0(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13234, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_15169, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk3_0))|~l1_algstr_0(esk3_0))).#trainneg
cnf(i_0_15176, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_15452, negated_conjecture, (X1=u1_struct_0(esk3_0)|r2_hidden(esk2_0,u1_struct_0(esk3_0))|~r1_tarski(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_15160, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk3_0,k4_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16778, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k4_struct_0(X1))=k5_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_16101, plain, (k1_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X2))=k5_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_16085, plain, (k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3)=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2565, negated_conjecture, (m1_subset_1(k4_struct_0(X1),u1_struct_0(esk1_0))|~l2_struct_0(X1)|~m1_zmodul01(X1,esk3_0))).#trainneg
cnf(i_0_15184, negated_conjecture, (k4_algstr_0(esk1_0,k4_struct_0(esk1_0))=k4_struct_0(esk1_0)|~l2_algstr_0(esk1_0)|~m1_subset_1(k4_algstr_0(esk1_0,k4_struct_0(esk1_0)),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_15767, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,X4),X4)=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_16019, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_16254, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X4,X3))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X4,X3),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_2595, negated_conjecture, (m1_subset_1(u2_struct_0(X1),u1_struct_0(esk1_0))|~l2_struct_0(X1)|~m1_zmodul01(X1,esk3_0))).#trainneg
cnf(i_0_15447, negated_conjecture, (u1_struct_0(esk3_0)=k1_numbers|r2_hidden(esk2_0,u1_struct_0(esk3_0))|r2_hidden(np__1,k1_numbers))).#trainneg
cnf(i_0_16689, plain, (k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),k4_algstr_0(X1,X2))=k5_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_15962, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),k4_algstr_0(X1,X3))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X4),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_2, negated_conjecture, (k5_zmodul01(esk1_0,esk2_0,esk3_0)!=u1_struct_0(esk3_0))).#trainneg
cnf(i_0_18332, negated_conjecture, (u1_struct_0(esk3_0)=k1_numbers|r2_hidden(np__1,k1_numbers)|~v1_xboole_0(u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_17061, negated_conjecture, (m1_subset_1(k4_struct_0(esk3_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_17331, plain, (k5_algstr_0(X1,X2,k4_struct_0(X1))=k5_algstr_0(X1,k1_algstr_0(X1,X3,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_17783, plain, (k1_algstr_0(X1,X2,k5_algstr_0(X1,X3,X2))=k5_algstr_0(X1,k1_algstr_0(X1,X3,X4),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_17880, plain, (k1_algstr_0(X1,k5_algstr_0(X1,X2,X4),X4)=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_17059, negated_conjecture, (u1_struct_0(X1)=u1_struct_0(esk3_0)|r2_hidden(esk2_0,u1_struct_0(esk3_0))|~m1_zmodul01(X1,esk3_0))).#trainneg
cnf(i_0_19622, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k4_struct_0(X2),k4_algstr_0(X2,X3))|~l2_algstr_0(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_18911, plain, (k5_algstr_0(X1,k4_algstr_0(X1,X2),k4_struct_0(X1))=k5_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_18793, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),k4_algstr_0(X1,X3))=k5_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1253, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))|m1_subset_1(X1,X2)|~r2_hidden(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_17672, negated_conjecture, (k4_algstr_0(esk1_0,k4_struct_0(esk1_0))=k4_struct_0(esk1_0)|~l2_algstr_0(esk1_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_19224, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r2_hidden(k4_struct_0(esk3_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_18501, plain, (k5_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X2))=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_18317, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~r1_struct_0(X1,X4)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_17263, plain, (k5_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),X2)=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_19029, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,k4_struct_0(X2))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_18690, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X4),k4_algstr_0(X1,X4))=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_15290, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_15455, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk3_0))|m1_subset_1(X1,X2)|~r2_hidden(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_15448, negated_conjecture, (u1_struct_0(esk3_0)=k5_numbers|r2_hidden(esk2_0,u1_struct_0(esk3_0))|r2_hidden(np__1,k5_numbers))).#trainneg
cnf(i_0_17751, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,X4),X4)=k5_algstr_0(X1,k5_algstr_0(X1,X2,k4_algstr_0(X1,X3)),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,k4_algstr_0(X1,X3)),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_17836, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X5),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_21415, plain, (k5_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2),k4_algstr_0(X1,X2))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_21788, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X3,X2),k4_algstr_0(X1,X2))=k1_algstr_0(X1,X3,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_17157, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_22040, negated_conjecture, (u1_struct_0(esk3_0)=k5_numbers|r2_hidden(np__1,k5_numbers)|~v1_xboole_0(u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_20727, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_18149, negated_conjecture, (r1_struct_0(X1,k5_algstr_0(esk1_0,X2,X3))|~r1_struct_0(X1,k4_algstr_0(esk1_0,X3))|~r1_struct_0(X1,X2)|~m1_zmodul01(X1,esk1_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_20416, plain, (k1_algstr_0(X1,X2,k5_algstr_0(X1,X3,X2))=k1_algstr_0(X1,k5_algstr_0(X1,X3,X4),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_20899, negated_conjecture, (k1_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0))=k4_algstr_0(esk1_0,k4_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_15294, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k4_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_19458, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,k4_struct_0(X1)),X2)=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_699, negated_conjecture, (k5_algstr_0(esk1_0,X1,X1)=k4_struct_0(esk3_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_24110, negated_conjecture, (k4_algstr_0(esk1_0,k4_struct_0(esk1_0))=k4_struct_0(esk3_0)|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_22438, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_22690, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,k4_struct_0(X1),X2))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_22500, plain, (r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X3,X4),k4_algstr_0(X2,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_22494, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X4,X3),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_24169, negated_conjecture, (m1_subset_1(k4_struct_0(esk3_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_24051, negated_conjecture, (k5_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0))=k4_struct_0(esk1_0)|~l2_algstr_0(esk1_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_24094, negated_conjecture, (m1_subset_1(X1,u1_struct_0(X2))|~l1_algstr_0(X2)|~r2_hidden(X1,a_3_1_zmodul01(esk1_0,X3,X2))|~r1_struct_0(X2,X3)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_23206, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~l2_algstr_0(X2)|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_22491, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_23870, plain, (r1_struct_0(X1,k5_algstr_0(X1,X2,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_21357, negated_conjecture, (k5_algstr_0(esk1_0,X1,X1)=k4_struct_0(esk1_0)|~l2_algstr_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_23835, negated_conjecture, (r1_struct_0(esk1_0,k5_algstr_0(esk1_0,k4_algstr_0(esk1_0,k4_struct_0(esk1_0)),k4_struct_0(esk1_0)))|~l2_algstr_0(esk1_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_21149, plain, (k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3)=k4_algstr_0(X1,k4_algstr_0(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_19026, plain, (k1_algstr_0(X1,X2,k5_algstr_0(X1,X3,X2))=k5_algstr_0(X1,X3,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_24948, plain, (r1_struct_0(X1,k1_algstr_0(X1,X2,k4_struct_0(X1)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_24773, negated_conjecture, (m1_subset_1(k4_struct_0(esk3_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_24228, negated_conjecture, (k4_struct_0(esk3_0)=k4_struct_0(esk1_0)|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_24243, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|X3!=k1_algstr_0(X2,X5,k4_struct_0(X1))|~l2_algstr_0(X1)|~r1_struct_0(X1,k1_algstr_0(X2,X5,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_25050, negated_conjecture, (r1_struct_0(esk3_0,k5_algstr_0(esk1_0,X1,X1))|~l2_algstr_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_25223, plain, (k4_algstr_0(X1,k4_algstr_0(X1,X2))=k5_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_25697, plain, (r1_struct_0(X1,k5_algstr_0(X2,X4,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X4,X3))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X4,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_24239, plain, (r1_struct_0(X1,k1_algstr_0(X2,k3_rlvect_1(X2,X3,X4),X5))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_24943, negated_conjecture, (r1_struct_0(esk1_0,X1)|~l2_algstr_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_24353, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_17156, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~l2_algstr_0(X4)|~l2_algstr_0(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X4))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_13296, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,k4_struct_0(X1)),X3)=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_18543, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)=k4_algstr_0(X1,k4_algstr_0(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_18055, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X3,X2),k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_24944, negated_conjecture, (r1_struct_0(esk1_0,k4_algstr_0(esk1_0,k4_struct_0(esk1_0)))|~l2_algstr_0(esk1_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_15273, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_25845, plain, (r1_struct_0(X1,k4_algstr_0(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_25922, plain, (r1_struct_0(X1,k5_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~r1_struct_0(X1,k4_algstr_0(X1,X3))|~m1_zmodul01(X1,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_26194, plain, (r1_struct_0(X1,k5_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,X3),X2),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_27116, plain, (k4_algstr_0(X1,k4_algstr_0(X1,X2))=k5_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_25893, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_26564, plain, (r1_struct_0(X1,k3_rlvect_1(X2,X3,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k3_rlvect_1(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_25969, negated_conjecture, (k4_struct_0(esk3_0)=k4_struct_0(esk1_0)|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_26250, negated_conjecture, (r1_struct_0(esk3_0,k4_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_26249, negated_conjecture, (r1_struct_0(esk3_0,k4_struct_0(esk1_0))|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_25784, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r2_hidden(X4,a_3_1_zmodul01(X2,X3,X1))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_26213, plain, (r1_struct_0(X1,k5_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X3),X2),u1_struct_0(X1))|~m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_26848, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k1_algstr_0(X2,X4,X5)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X4,k4_struct_0(X2)))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X4,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(k1_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_27600, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X3,X2),k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_27372, negated_conjecture, (m1_subset_1(k4_algstr_0(esk1_0,k4_struct_0(esk1_0)),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_27380, negated_conjecture, (r1_struct_0(esk1_0,k4_struct_0(esk3_0))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_27114, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k4_algstr_0(X2,k4_algstr_0(X2,X3))|~l2_algstr_0(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_26370, plain, (r1_struct_0(X1,k1_algstr_0(X2,k1_algstr_0(X2,k4_struct_0(X2),k4_struct_0(X2)),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_15866, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)=k5_algstr_0(X1,k1_algstr_0(X1,X4,X2),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_1252, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))|r2_hidden(X1,X2)|~r2_hidden(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_19334, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,k5_algstr_0(X2,X3,X5))|~l2_algstr_0(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_25027, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_24963, plain, (r1_struct_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~r2_hidden(X2,a_3_1_zmodul01(X1,k4_struct_0(X1),X3))|~m1_zmodul01(X3,X1)|~m1_subset_1(esk8_4(X2,X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_28050, negated_conjecture, (r1_struct_0(esk3_0,k4_struct_0(esk3_0))|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_29606, plain, (k1_algstr_0(X1,X2,k5_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_27659, plain, (r1_struct_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2))|v2_struct_0(X1)|~r1_struct_0(X1,k4_algstr_0(X1,X3))|~m1_zmodul01(X1,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_27928, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_28095, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X4,X3)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_29603, plain, (k1_algstr_0(X1,X2,k5_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)))=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_28094, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X3,X4)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_29323, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k4_struct_0(X2))|~l2_algstr_0(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_29713, negated_conjecture, (r1_struct_0(esk3_0,k4_struct_0(esk1_0))|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_28682, negated_conjecture, (m1_subset_1(k4_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_29794, plain, (r1_struct_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~r2_hidden(X2,a_3_1_zmodul01(X1,X3,X4))|~m1_zmodul01(X4,X1)|~m1_subset_1(esk8_4(X2,X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_29913, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)=k4_algstr_0(X1,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_29212, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X3,X2),X3)=k5_algstr_0(X1,k1_algstr_0(X1,X4,X2),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_30810, negated_conjecture, (r1_struct_0(X1,k4_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~m1_zmodul01(X1,esk1_0)|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_31613, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k4_struct_0(X2),X3)|~l2_algstr_0(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_29687, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)=k4_algstr_0(X1,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_27938, plain, (r1_struct_0(X1,X4)|v2_struct_0(X2)|~r2_hidden(X4,a_3_1_zmodul01(X2,X3,X5))|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|~r1_struct_0(X1,esk8_4(X4,X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_zmodul01(X5,X2)|~m1_subset_1(esk8_4(X4,X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_28995, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X4,X2),X4)=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_31832, plain, (r1_struct_0(X1,k1_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~r1_struct_0(X4,X2)|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_31835, plain, (r1_struct_0(X1,k3_rlvect_1(X1,X2,X3))|v2_struct_0(X1)|~r1_struct_0(X4,X2)|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_31839, plain, (r1_struct_0(X1,X2)|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_31834, plain, (r1_struct_0(X1,k3_rlvect_1(X1,X2,X3))|v2_struct_0(X1)|~r1_struct_0(X4,X3)|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_32735, negated_conjecture, (r1_struct_0(X1,k4_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~m1_zmodul01(X1,esk1_0))).#trainneg
cnf(i_0_31765, negated_conjecture, (m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_31831, plain, (r1_struct_0(X1,k1_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~r1_struct_0(X4,X3)|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_32781, negated_conjecture, (r2_hidden(X1,a_3_1_zmodul01(esk1_0,X1,X2))|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_31827, plain, (r1_tarski(a_3_1_zmodul01(X1,X2,X3),X4)|r1_struct_0(X1,esk4_2(a_3_1_zmodul01(X1,X2,X3),X4))|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_33138, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,k4_struct_0(X1)))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_33500, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_33135, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,k4_struct_0(X1)))=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_33083, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_algstr_0(X2,k4_struct_0(X2))))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_31631, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k1_algstr_0(X2,X4,X5)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_33046, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,k4_struct_0(X2)),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_30830, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k5_algstr_0(X2,X3,X4)))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_33729, negated_conjecture, (r1_struct_0(esk1_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_35670, plain, (k5_algstr_0(X1,X3,X3)=k5_algstr_0(X1,X2,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_35668, plain, (k1_algstr_0(X1,X2,k5_algstr_0(X1,X4,X4))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_35680, plain, (k5_algstr_0(X1,X3,X3)=k4_algstr_0(X1,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_30828, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r2_hidden(X4,a_3_1_zmodul01(X2,X3,X5))|~r1_struct_0(X1,esk8_4(X4,X2,X3,X5))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_zmodul01(X5,X2)|~m1_subset_1(esk8_4(X4,X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_27936, plain, (r1_struct_0(X1,k1_algstr_0(X2,X4,X3))|v2_struct_0(X2)|~l1_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X4,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_35665, plain, (k1_algstr_0(X1,X2,k5_algstr_0(X1,X4,X4))=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_32228, plain, (r1_struct_0(X1,k1_algstr_0(X2,k1_algstr_0(X2,k4_algstr_0(X2,X5),X5),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_33514, plain, (k1_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2),X2)=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_36548, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_36451, plain, (k5_algstr_0(X1,X2,X2)=k4_algstr_0(X1,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_37318, plain, (k1_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2),k1_algstr_0(X1,X3,X2))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_37383, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_37423, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X4),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_33739, negated_conjecture, (~v1_xboole_0(a_3_1_zmodul01(esk1_0,X1,X2))|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_30847, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X3,X4)))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_30854, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,k1_algstr_0(X2,X4,X3),X5)))|~r1_struct_0(X1,k1_algstr_0(X2,X4,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_15454, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk3_0))|r2_hidden(X1,X2)|~r2_hidden(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_35652, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|X3!=k1_algstr_0(X2,k4_algstr_0(X2,X4),k4_struct_0(X1))|~l2_algstr_0(X1)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X5,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_30855, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X4,k1_algstr_0(X2,X3,X5))))|~r1_struct_0(X1,k1_algstr_0(X2,X4,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_27844, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k5_algstr_0(X1,X2,k4_struct_0(X1)))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_19304, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,esk8_4(X5,X2,X3,X6))|~r2_hidden(X5,a_3_1_zmodul01(X2,X3,X6))|~m1_zmodul01(X4,X2)|~m1_zmodul01(X6,X2)|~m1_subset_1(esk8_4(X5,X2,X3,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_25093, negated_conjecture, (k1_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0))=k4_struct_0(esk1_0)|~l2_algstr_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_25147, negated_conjecture, (r1_struct_0(esk1_0,k5_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0)))|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_25944, plain, (r1_struct_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X2),X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_23037, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_algstr_0(X2,X4)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_36506, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_algstr_0(X2,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_13822, plain, (r1_struct_0(X1,k1_algstr_0(X2,X5,k5_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l1_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_25964, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r2_hidden(k4_struct_0(esk3_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_39320, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_39449, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X4,X3))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_37900, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,k1_algstr_0(X2,X4,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(k1_algstr_0(X2,X4,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_38150, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,k1_algstr_0(X2,X4,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_38179, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X3)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_37767, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,k4_struct_0(X2)),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_39531, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_38172, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_39114, negated_conjecture, (k1_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0))=k4_struct_0(esk1_0)|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_26243, negated_conjecture, (r1_struct_0(esk3_0,k4_algstr_0(esk1_0,k4_struct_0(esk1_0)))|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_39641, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k5_algstr_0(X2,X4,X5)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X4,X3))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_38655, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|X3!=k5_algstr_0(X2,k4_struct_0(X1),X4)|~l2_algstr_0(X1)|~l2_algstr_0(X2)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_23836, plain, (r1_struct_0(X1,k5_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~r2_hidden(X2,a_3_1_zmodul01(X1,X3,X4))|~m1_zmodul01(X4,X1)|~m1_subset_1(esk8_4(X2,X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_24634, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k1_algstr_0(X2,X3,X5),k4_struct_0(X2))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_37355, plain, (k1_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X3),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_9679, plain, (r1_tarski(k1_zfmisc_1(X1),X2)|r2_hidden(k1_xboole_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_33733, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_25397, plain, (r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_15347, plain, (k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1))=k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_35701, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~m1_subset_1(esk4_2(X1,u1_struct_0(esk1_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_39650, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|v2_struct_0(X2)|~r2_hidden(X5,a_3_1_zmodul01(X2,X4,X6))|~r1_struct_0(X1,esk8_4(X5,X2,X4,X6))|~r1_struct_0(X1,k1_algstr_0(X2,X4,X3))|~m1_zmodul01(X1,X2)|~m1_zmodul01(X6,X2)|~m1_subset_1(esk8_4(X5,X2,X4,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_25025, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_25947, plain, (r1_struct_0(X1,k1_algstr_0(X1,X2,k4_algstr_0(X1,X3)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_15172, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(esk3_0)|~m1_zmodul01(esk3_0,X1)|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_27118, plain, (k4_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_38342, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l1_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X4,X3)))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_35702, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~m1_subset_1(esk4_2(X1,u1_struct_0(esk1_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_39158, negated_conjecture, (m1_subset_1(k5_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0)),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_41904, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~r2_hidden(X4,a_3_1_zmodul01(X2,X5,X1))|~r1_struct_0(X1,k1_algstr_0(X2,X5,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_41405, plain, (r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~l1_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X4,X3))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_40797, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,k4_struct_0(X2))|~r2_hidden(X5,a_3_1_zmodul01(X2,X3,X6))|~r1_struct_0(X4,esk8_4(X5,X2,X3,X6))|~m1_zmodul01(X4,X2)|~m1_zmodul01(X6,X2)|~m1_subset_1(esk8_4(X5,X2,X3,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_41911, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_42214, negated_conjecture, (r1_struct_0(X1,esk2_0)|v2_struct_0(X1)|~l1_algstr_0(esk3_0)|~l1_struct_0(X1)|~m1_zmodul01(esk3_0,X1)|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_40691, plain, (r1_struct_0(X1,k5_algstr_0(X1,X2,X2))|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_40075, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X4),X3))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_40268, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~r2_hidden(k4_struct_0(X2),a_3_1_zmodul01(X2,X4,X1))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_41900, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~m1_zmodul01(X1,esk1_0))).#trainneg
cnf(i_0_40456, negated_conjecture, (m1_subset_1(k4_algstr_0(esk1_0,k4_struct_0(esk1_0)),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_40671, plain, (r1_struct_0(X1,k1_algstr_0(X1,X2,X3))|v2_struct_0(X1)|X2!=k4_algstr_0(X1,X3)|~l2_algstr_0(X1)|~m1_zmodul01(X1,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_39850, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~r2_hidden(X3,a_3_1_zmodul01(X2,X4,X1))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_40469, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_27903, plain, (k5_algstr_0(X1,X2,k4_struct_0(X1))=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_20291, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X5),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_42464, negated_conjecture, (r1_tarski(u1_struct_0(X1),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~m1_zmodul01(X1,esk1_0))).#trainneg
cnf(i_0_37762, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k4_algstr_0(X2,k4_struct_0(X2)))|~l2_algstr_0(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_39951, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X4,X3))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_32778, plain, (r2_hidden(k1_algstr_0(X1,k4_struct_0(X1),X2),a_3_1_zmodul01(X1,X2,X3))|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_40687, negated_conjecture, (r1_struct_0(esk1_0,k5_algstr_0(esk1_0,X1,X1))|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_42313, negated_conjecture, (k4_algstr_0(esk1_0,k4_struct_0(esk3_0))=k4_struct_0(esk1_0)|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_12794, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X3,X2),X4)=k5_algstr_0(X1,k5_algstr_0(X1,X3,X4),X2)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_rlvect_1(X1)|~v2_rlvect_1(X1))).#trainneg
cnf(i_0_17469, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),k4_algstr_0(X1,X2))=k5_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_13116, negated_conjecture, (r1_struct_0(X1,k3_rlvect_1(esk1_0,X2,X3))|~r1_struct_0(X1,X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,esk1_0)|~m1_subset_1(esk8_4(k3_rlvect_1(esk1_0,X2,X3),esk1_0,X2,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12883, plain, (r1_struct_0(X1,k1_algstr_0(X2,k1_algstr_0(X2,X3,X5),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,k1_algstr_0(X2,X4,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(k1_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_38982, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),X4))|~r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_22374, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X5,X3))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_6303, negated_conjecture, (k1_algstr_0(esk1_0,X1,k4_struct_0(esk1_0))=X1|~l2_algstr_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_39196, negated_conjecture, (r1_struct_0(esk1_0,k1_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0)))|~l2_algstr_0(esk1_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_44357, negated_conjecture, (r1_struct_0(X1,k3_rlvect_1(esk1_0,X2,X3))|~r1_struct_0(X1,X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,esk1_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_43064, plain, (r1_struct_0(X1,k1_algstr_0(X1,X2,X3))|v2_struct_0(X1)|X2!=k4_algstr_0(X1,X3)|~m1_zmodul01(X1,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_42797, plain, (r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~r1_struct_0(X1,X3)|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_43697, plain, (r1_struct_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2))|v2_struct_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_42619, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_42850, negated_conjecture, (r1_struct_0(esk1_0,k5_algstr_0(esk1_0,X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_42629, plain, (r1_struct_0(X1,k1_algstr_0(X2,X4,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_42848, negated_conjecture, (r1_struct_0(X1,esk2_0)|v2_struct_0(X1)|~l1_algstr_0(esk3_0)|~l1_algstr_0(X1)|~m1_zmodul01(esk3_0,X1)|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_42632, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k1_algstr_0(X2,X4,X5)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k1_algstr_0(X2,X3,X4),X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k1_algstr_0(X2,X3,X4),X5),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_43450, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X5,X3))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_42621, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_43673, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),k4_struct_0(X2)))|v2_struct_0(X2)|~r2_hidden(X4,a_3_1_zmodul01(X2,X3,X1))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_42961, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~r2_hidden(X1,u1_struct_0(X2))|~m1_zmodul01(X2,esk1_0))).#trainneg
cnf(i_0_42576, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_44505, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X4),X3))|~r1_struct_0(X1,k5_algstr_0(X2,X4,k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_43197, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~r1_struct_0(X4,k4_struct_0(X2))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_43701, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,k4_struct_0(X2),X3)))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_44241, plain, (r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,k4_struct_0(X2),X4),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_43125, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~r2_hidden(X3,a_3_1_zmodul01(X2,X4,X5))|~r1_struct_0(X1,esk8_4(X3,X2,X4,X5))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_zmodul01(X5,X2)|~m1_subset_1(esk8_4(X3,X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_44513, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X5)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_43530, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~r2_hidden(X1,u1_struct_0(X2))|~m1_zmodul01(X2,esk1_0))).#trainneg
cnf(i_0_25149, negated_conjecture, (r1_struct_0(esk1_0,k5_algstr_0(esk1_0,k4_struct_0(esk3_0),k4_struct_0(esk1_0)))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_43120, plain, (r1_struct_0(X1,k1_algstr_0(X2,k1_algstr_0(X2,X3,X4),k4_struct_0(X2)))|v2_struct_0(X2)|~r1_struct_0(X1,X4)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_46729, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_45283, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_45526, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_algstr_0(X2,X3),X4),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_46210, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),k4_struct_0(X2)))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X3,X4)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_45312, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X4),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_46211, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),k4_struct_0(X2)))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X4,X3)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_46242, negated_conjecture, (m1_subset_1(k4_struct_0(esk3_0),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_46758, negated_conjecture, (m1_subset_1(k4_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_46328, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_45381, plain, (r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_45494, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X5),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_algstr_0(X2,X4)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,k4_algstr_0(X2,X4)),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_45493, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X5),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_algstr_0(X2,X4)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,k4_algstr_0(X2,X4)),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_46684, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_46731, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X5,k4_struct_0(X2)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_45569, negated_conjecture, (r1_struct_0(esk1_0,k4_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_44939, negated_conjecture, (m1_subset_1(k1_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0)),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_43115, plain, (r1_tarski(a_3_1_zmodul01(X2,X3,X1),X4)|r1_struct_0(X1,k1_algstr_0(X2,esk4_2(a_3_1_zmodul01(X2,X3,X1),X4),k4_struct_0(X2)))|v2_struct_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_32336, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X4,X3),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_26418, plain, (r1_struct_0(X1,k1_algstr_0(X2,k1_algstr_0(X2,X3,k4_struct_0(X2)),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X4,X5))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_27934, plain, (r1_struct_0(X1,k3_rlvect_1(X2,X4,X3))|v2_struct_0(X2)|~l1_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k3_rlvect_1(X2,X4,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_33212, negated_conjecture, (r1_struct_0(X1,X2)|~r2_hidden(X2,a_3_1_zmodul01(esk1_0,X3,X4))|~r1_struct_0(X1,esk8_4(X2,esk1_0,X3,X4))|~r1_struct_0(X1,k4_struct_0(esk1_0))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,esk1_0)|~m1_zmodul01(X4,esk1_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_25976, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|X3!=k1_algstr_0(X2,k4_algstr_0(X2,X5),k4_struct_0(X1))|~l2_algstr_0(X1)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X4,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_42958, negated_conjecture, (u1_struct_0(esk1_0)=u1_struct_0(X1)|~l1_algstr_0(esk1_0)|~r1_tarski(u1_struct_0(esk1_0),u1_struct_0(X1))|~m1_zmodul01(X1,esk1_0))).#trainneg
cnf(i_0_33205, plain, (r1_struct_0(X1,X2)|v2_struct_0(X3)|~l2_algstr_0(X3)|~r2_hidden(X2,a_3_1_zmodul01(X3,X4,X5))|~r1_struct_0(X1,k1_algstr_0(X3,k4_struct_0(X3),X4))|~r1_struct_0(X1,esk8_4(X2,X3,X4,X5))|~m1_zmodul01(X1,X3)|~m1_zmodul01(X5,X3)|~m1_subset_1(X4,u1_struct_0(X3))|~l1_zmodul01(X3)|~v5_zmodul01(X3)|~v4_zmodul01(X3)|~v3_zmodul01(X3)|~v2_zmodul01(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v2_rlvect_1(X3)|~v13_algstr_0(X3))).#trainneg
cnf(i_0_26091, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|X3!=k1_algstr_0(X2,X5,k4_struct_0(X1))|~l2_algstr_0(X1)|~r1_struct_0(X1,X4)|~r1_struct_0(X1,X5)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_12797, plain, (r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X4,X3),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,k5_algstr_0(X2,X4,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k5_algstr_0(X2,X4,X3),X5),u1_struct_0(X2))|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(k5_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_21120, plain, (v1_xboole_0(u1_struct_0(X1))|r1_struct_0(X1,k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_43527, negated_conjecture, (u1_struct_0(esk1_0)=u1_struct_0(X1)|~r1_tarski(u1_struct_0(esk1_0),u1_struct_0(X1))|~l2_algstr_0(esk1_0)|~m1_zmodul01(X1,esk1_0))).#trainneg
cnf(i_0_43812, negated_conjecture, (r1_struct_0(esk1_0,k1_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk3_0)))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_20382, plain, (k1_algstr_0(X1,k5_algstr_0(X1,X2,X4),X4)=k5_algstr_0(X1,k5_algstr_0(X1,X2,k4_algstr_0(X1,X3)),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,k4_algstr_0(X1,X3)),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_27198, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X3,X2),X3)=k4_algstr_0(X1,k4_algstr_0(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_48452, negated_conjecture, (r1_struct_0(X1,X2)|~r2_hidden(X2,a_3_1_zmodul01(esk1_0,X3,X4))|~r1_struct_0(X1,esk8_4(X2,esk1_0,X3,X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,esk1_0)|~m1_zmodul01(X4,esk1_0)|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_46999, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~r2_hidden(k4_struct_0(X2),a_3_1_zmodul01(X2,X3,X1))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_48431, plain, (r1_struct_0(X1,k1_algstr_0(X1,X2,X3))|v2_struct_0(X1)|X2!=k4_algstr_0(X1,X4)|~l2_algstr_0(X1)|~r1_struct_0(X1,k5_algstr_0(X1,X3,X4))|~m1_zmodul01(X1,X1)|~m1_subset_1(k4_algstr_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_49597, plain, (k4_algstr_0(X1,k4_algstr_0(X1,X2))=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_48657, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_48659, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,k4_struct_0(X2)),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_algstr_0(X2,k4_struct_0(X2)),X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_47010, plain, (r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_47957, negated_conjecture, (r1_struct_0(esk1_0,k4_struct_0(esk3_0))|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_46996, negated_conjecture, (r1_struct_0(esk1_0,k4_algstr_0(esk1_0,k4_struct_0(esk1_0)))|~m1_zmodul01(esk1_0,esk1_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_48438, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|X3!=k5_algstr_0(X2,k4_struct_0(X1),X5)|~l2_algstr_0(X1)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X4,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_49613, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k4_algstr_0(X2,k4_algstr_0(X2,X3))|~l2_algstr_0(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_49614, plain, (r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_47290, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),k4_struct_0(X2)))|v2_struct_0(X2)|~r2_hidden(X4,a_3_1_zmodul01(X2,X3,X5))|~r1_struct_0(X1,esk8_4(X4,X2,X3,X5))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_zmodul01(X5,X2)|~m1_subset_1(esk8_4(X4,X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_47085, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,k4_algstr_0(X2,X5)))|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_47161, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,k4_struct_0(X2),X4),X3))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k1_algstr_0(X2,k4_struct_0(X2),X4),X3),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_47934, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~r2_hidden(X5,a_3_1_zmodul01(X2,k4_struct_0(X2),X1))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_42960, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~r2_hidden(X1,u1_struct_0(X2))|~m1_zmodul01(X2,esk1_0))).#trainneg
cnf(i_0_46994, negated_conjecture, (r1_struct_0(esk3_0,k4_algstr_0(esk1_0,k4_struct_0(esk1_0)))|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_47571, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_47104, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~r1_struct_0(X1,X5)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_47940, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_48689, plain, (r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X3,X4),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_47580, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_47685, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|~r1_struct_0(X1,k4_struct_0(X2))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_46991, plain, (r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_49450, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k1_algstr_0(X2,X6,X5),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_49606, plain, (r2_hidden(k4_algstr_0(X1,k4_algstr_0(X1,X2)),a_3_1_zmodul01(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_43529, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~r2_hidden(X1,u1_struct_0(X2))|~m1_zmodul01(X2,esk1_0))).#trainneg
cnf(i_0_47516, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(k4_struct_0(esk3_0),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_47939, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_47116, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),k3_rlvect_1(X2,X3,X4)))|v2_struct_0(X2)|~r1_struct_0(X1,X4)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_48554, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|X3!=k5_algstr_0(X2,k4_struct_0(X1),X5)|~l2_algstr_0(X1)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_47181, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X4),X3))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_48434, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|X3!=k1_algstr_0(X2,k4_struct_0(X1),k4_algstr_0(X2,X5))|~l2_algstr_0(X1)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X4,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_47909, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X5,X6),X6))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_47520, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(k4_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_47911, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X6,X5),X6))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_47281, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k5_algstr_0(X2,X3,X4)))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_49783, negated_conjecture, (r1_struct_0(X1,X2)|~r2_hidden(X2,a_3_1_zmodul01(esk1_0,X3,X1))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,esk1_0)|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_50129, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|X3!=k5_algstr_0(X2,k4_struct_0(X1),X4)|~l2_algstr_0(X1)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_51628, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,k4_algstr_0(X2,X3))))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_51187, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_51594, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~r2_hidden(k1_algstr_0(X2,X6,X5),a_3_1_zmodul01(X2,X6,X4))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_51188, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_51384, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_algstr_0(X2,X4)))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_51600, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k1_algstr_0(X2,k4_algstr_0(X2,X6),X5))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,k4_algstr_0(X2,X6),X5),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_51107, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,k4_struct_0(X2)),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_51625, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,k4_algstr_0(X2,X3)),X4))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_51133, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X5),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,k4_algstr_0(X2,X4)))|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_52111, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X5,k4_algstr_0(X2,X6)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~r1_struct_0(X1,X6)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X5,k4_algstr_0(X2,X6)),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_51853, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,k4_struct_0(X2),X4),X3))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_51797, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k3_rlvect_1(X2,X5,X6))|~r1_struct_0(X4,X6)|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(k3_rlvect_1(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_51460, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,k4_struct_0(X2)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_50025, negated_conjecture, (m1_subset_1(k4_algstr_0(esk1_0,k4_struct_0(esk1_0)),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~m1_zmodul01(esk1_0,esk1_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_49950, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X4),X3))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_52320, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X5),X6))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X6))|~r1_struct_0(X1,X5)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_algstr_0(X2,X5),X6),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_49874, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_48026, plain, (r1_tarski(a_3_1_zmodul01(X2,X3,X1),X4)|r1_struct_0(X1,k5_algstr_0(X2,esk4_2(a_3_1_zmodul01(X2,X3,X1),X4),X5))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(esk4_2(a_3_1_zmodul01(X2,X3,X1),X4),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_48048, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k1_algstr_0(X2,X3,X5),k4_struct_0(X2))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_48587, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|X3!=k5_algstr_0(X2,k4_struct_0(X2),k4_algstr_0(X2,X5))|~l2_algstr_0(X1)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X5,k4_struct_0(X1)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X5,k4_struct_0(X1)),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_1487, plain, (v1_xboole_0(k1_zfmisc_1(X1))|m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_50666, negated_conjecture, (m1_subset_1(k4_algstr_0(esk1_0,k4_struct_0(esk1_0)),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_49917, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X2)=k1_algstr_0(X1,X3,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_49591, plain, (k4_algstr_0(X1,k4_algstr_0(X1,X3))=k1_algstr_0(X1,k5_algstr_0(X1,X2,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_52650, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~r2_hidden(k1_algstr_0(X2,X3,X4),a_3_1_zmodul01(X2,X4,X1))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_54628, plain, (r1_struct_0(X1,k1_algstr_0(X2,X4,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r2_hidden(k1_algstr_0(X2,X3,X4),a_3_1_zmodul01(X2,X3,X1))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_53162, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_54041, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~r2_hidden(X5,a_3_1_zmodul01(X2,k4_struct_0(X2),X4))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_53054, plain, (r1_struct_0(X1,k4_algstr_0(X2,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_52966, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k1_algstr_0(X2,k4_algstr_0(X2,X6),X5))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,k4_algstr_0(X2,X6),X5),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_47757, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k1_algstr_0(X2,k4_struct_0(X2),X5)))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,k1_algstr_0(X2,k4_struct_0(X2),X5)),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_3677, plain, (esk4_2(k1_zfmisc_1(k1_xboole_0),X1)=k1_xboole_0|r1_tarski(k1_zfmisc_1(k1_xboole_0),X1))).#trainneg
cnf(i_0_51200, negated_conjecture, (r1_struct_0(esk1_0,k1_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0)))|~m1_zmodul01(esk1_0,esk1_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_50613, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|v2_struct_0(X2)|~r2_hidden(X3,a_3_1_zmodul01(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,esk8_4(X3,X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,X5)|~m1_zmodul01(X1,X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(esk8_4(X3,X2,k4_struct_0(X2),X4),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_49882, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k1_algstr_0(X2,X5,k4_struct_0(X2))))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,k1_algstr_0(X2,X5,k4_struct_0(X2))),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_51501, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~l1_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k1_algstr_0(X2,X5,X6),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_46251, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l1_algstr_0(X2)|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0))).#trainneg
cnf(i_0_48173, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),k4_algstr_0(X2,X3)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_48585, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|X3!=k5_algstr_0(X2,X5,k4_struct_0(X2))|~l2_algstr_0(X1)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X5,k4_struct_0(X1)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X5,k4_struct_0(X1)),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_47001, plain, (r1_struct_0(X1,k4_algstr_0(X1,X2))|v2_struct_0(X1)|~m1_zmodul01(X1,X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_54713, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),k4_algstr_0(X1,X3))=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_51403, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_46252, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l2_algstr_0(X2)|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0))).#trainneg
cnf(i_0_46770, negated_conjecture, (m1_subset_1(k5_algstr_0(esk1_0,k4_struct_0(esk3_0),k4_struct_0(esk1_0)),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_25776, plain, (k5_algstr_0(X1,X2,k4_struct_0(X1))=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_56142, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~r2_hidden(X3,a_3_1_zmodul01(X2,k4_struct_0(X2),X1))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_55674, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_56493, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_algstr_0(X2,X6)))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,k4_algstr_0(X2,X6)),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_56652, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k5_algstr_0(X2,k4_struct_0(X2),X6))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_55896, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_56492, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_algstr_0(X2,X6)))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,k4_algstr_0(X2,X6)),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_57789, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,k4_struct_0(X1)),X2)=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_56076, plain, (r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)|~r2_hidden(k1_xboole_0,X1))).#trainneg
cnf(i_0_56070, negated_conjecture, (r1_tarski(k1_zfmisc_1(k1_xboole_0),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10649, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|v2_struct_0(X4)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X4)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v4_rlvect_1(X4)|~v3_rlvect_1(X2)|~v3_rlvect_1(X4)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)|~v13_algstr_0(X4))).#trainneg
cnf(i_0_11678, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)))=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_15219, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k4_algstr_0(X1,k4_struct_0(X1)))=k5_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_51235, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k5_algstr_0(X2,X5,k4_struct_0(X2)))|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_46214, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),k4_struct_0(X2)))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k3_rlvect_1(X2,X4,X3)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_58926, plain, (k1_zfmisc_1(k1_xboole_0)=k1_xboole_0|~r2_hidden(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_56071, negated_conjecture, (r1_tarski(k1_zfmisc_1(k1_xboole_0),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_58141, plain, (r1_struct_0(X1,k4_algstr_0(X2,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r2_hidden(k5_algstr_0(X2,X3,X4),a_3_1_zmodul01(X2,X3,X1))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_58003, negated_conjecture, (r2_hidden(X1,a_3_1_zmodul01(esk1_0,X2,X3))|~r1_struct_0(X3,k4_algstr_0(esk1_0,X2))|~r1_struct_0(X3,X1)|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_58131, plain, (k5_algstr_0(X1,k4_algstr_0(X1,X3),k5_algstr_0(X1,X2,X3))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_58139, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X4),k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_58474, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X6,X5))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_58068, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,k4_algstr_0(X2,X6))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_58192, plain, (r1_struct_0(X1,k4_algstr_0(X2,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),X4))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_algstr_0(X2,X3),X4),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_58061, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_58930, plain, (X1=k1_zfmisc_1(k1_xboole_0)|~r1_tarski(X1,k1_zfmisc_1(k1_xboole_0))|~r2_hidden(k1_xboole_0,X1))).#trainneg
cnf(i_0_58935, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)=u1_struct_0(esk1_0)|~r1_tarski(u1_struct_0(esk1_0),k1_zfmisc_1(k1_xboole_0))|~l2_algstr_0(esk1_0)|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_58067, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_58222, plain, (k5_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k4_struct_0(X1),X2))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_58269, plain, (k5_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X3),k5_algstr_0(X1,X2,X3))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_58015, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k4_struct_0(X2),X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_58183, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),k4_struct_0(X1))=k1_algstr_0(X1,X2,k4_algstr_0(X1,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_60892, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,X6)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X6)|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_60598, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X6,X5))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_61760, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k4_algstr_0(X2,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_60149, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|~r1_struct_0(X4,X1)|~r1_struct_0(X4,X3)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_60160, plain, (m1_subset_1(X1,k1_xboole_0)|~r1_tarski(X1,k1_xboole_0)|~r2_hidden(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_60167, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)=u1_struct_0(esk1_0)|~l1_algstr_0(esk1_0)|~r1_tarski(u1_struct_0(esk1_0),k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_61743, plain, (r2_hidden(k5_algstr_0(X1,k4_struct_0(X1),X2),a_3_1_zmodul01(X1,X3,X4))|v2_struct_0(X1)|~r1_struct_0(X4,k4_algstr_0(X1,X3))|~r1_struct_0(X4,k4_algstr_0(X1,X2))|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_61016, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,X6)|~r2_hidden(X5,a_3_1_zmodul01(X2,X6,X4))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_60148, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X4,X1)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_61022, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,X5)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_60896, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,k1_algstr_0(X2,X6,X7))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X7)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X6,X7),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_60998, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|v2_struct_0(X4)|X1!=k5_algstr_0(X2,X5,X6)|~l2_algstr_0(X4)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v4_rlvect_1(X4)|~v3_rlvect_1(X2)|~v3_rlvect_1(X4)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)|~v13_algstr_0(X4))).#trainneg
cnf(i_0_60870, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|v2_struct_0(X4)|X1!=k1_algstr_0(X2,X5,X6)|~l2_algstr_0(X4)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v4_rlvect_1(X4)|~v3_rlvect_1(X2)|~v3_rlvect_1(X4)|~v2_rlvect_1(X2)|~v13_algstr_0(X2)|~v13_algstr_0(X4))).#trainneg
cnf(i_0_61110, plain, (k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_60901, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,k5_algstr_0(X2,X6,X7))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,k4_algstr_0(X2,X7))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X6,X7),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_60161, plain, (r1_tarski(X1,k1_xboole_0)|~r2_hidden(k1_xboole_0,k1_xboole_0)|~m1_subset_1(X1,k1_xboole_0))).#trainneg
cnf(i_0_60912, plain, (k1_zfmisc_1(k1_xboole_0)=k1_numbers|r2_hidden(np__1,k1_numbers)|~r2_hidden(k1_xboole_0,k1_numbers))).#trainneg
cnf(i_0_60449, plain, (k5_algstr_0(X1,k4_algstr_0(X1,X2),k4_struct_0(X1))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_61376, plain, (k5_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2),k4_struct_0(X1))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_63401, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X2,X3))=k4_algstr_0(X1,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_63453, plain, (k1_xboole_0=X1|~r2_hidden(k1_xboole_0,k1_xboole_0)|~m1_subset_1(X1,k1_xboole_0))).#trainneg
cnf(i_0_60913, plain, (k1_zfmisc_1(k1_xboole_0)=k5_numbers|r2_hidden(np__1,k5_numbers)|~r2_hidden(k1_xboole_0,k5_numbers))).#trainneg
cnf(i_0_63252, plain, (k5_algstr_0(X1,k4_algstr_0(X1,X3),X2)=k4_algstr_0(X1,k1_algstr_0(X1,X3,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_62228, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k4_algstr_0(X2,k4_algstr_0(X2,X5)),X5)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_62084, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~r2_hidden(X6,a_3_1_zmodul01(X2,X5,X4))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_62076, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k4_algstr_0(X2,k1_algstr_0(X2,X5,X6)))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_62082, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k4_algstr_0(X2,k5_algstr_0(X2,X6,X7)))|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X7))|~r1_struct_0(X4,k5_algstr_0(X2,X6,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X6,X7),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_64022, plain, (k1_algstr_0(X1,k4_struct_0(X1),k5_algstr_0(X1,k4_struct_0(X1),X2))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_62230, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k1_algstr_0(X2,k4_struct_0(X2),X5),X5)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_62170, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk3_0))|m1_subset_1(u1_struct_0(esk3_0),k1_xboole_0)|~r2_hidden(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_62227, plain, (r1_tarski(a_3_1_zmodul01(X2,X5,X4),X6)|r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,esk4_2(a_3_1_zmodul01(X2,X5,X4),X6),X5)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_62073, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k4_algstr_0(X2,k1_algstr_0(X2,X6,X5)))|~r1_struct_0(X4,k1_algstr_0(X2,X6,k4_struct_0(X2)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X6,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_62698, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|~l2_algstr_0(X4)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_63302, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k1_algstr_0(X2,X5,X6))|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_62194, plain, (r1_struct_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2))|v2_struct_0(X1)|~r1_struct_0(X4,k4_algstr_0(X1,X3))|~r1_struct_0(X4,k4_algstr_0(X1,X2))|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_63245, plain, (r1_struct_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X3),k1_algstr_0(X1,X3,X2)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_64845, plain, (k5_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X3,X2))=k5_algstr_0(X1,k4_algstr_0(X1,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_65523, negated_conjecture, (m1_subset_1(u1_struct_0(esk3_0),k1_xboole_0)|~v1_xboole_0(u1_struct_0(esk3_0))|~r2_hidden(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_62171, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))|m1_subset_1(u1_struct_0(esk1_0),k1_xboole_0)|~r2_hidden(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_64187, plain, (k5_algstr_0(X1,k4_algstr_0(X1,X3),X2)=k4_algstr_0(X1,k1_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_64607, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k4_algstr_0(X2,k4_algstr_0(X2,k4_algstr_0(X2,X5))))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_64199, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_65520, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k4_algstr_0(X2,X5),X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_64730, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k1_algstr_0(X2,X5,X6))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X5),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_64172, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),k1_algstr_0(X1,X2,X3))=k5_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_62242, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,k4_struct_0(X2))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_63085, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,k4_struct_0(X2))|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_62245, plain, (r1_tarski(X1,a_3_1_zmodul01(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|~r1_struct_0(X4,esk4_2(X1,a_3_1_zmodul01(X2,k4_algstr_0(X2,X3),X4)))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(esk4_2(X1,a_3_1_zmodul01(X2,k4_algstr_0(X2,X3),X4)),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_62070, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k4_algstr_0(X2,k4_algstr_0(X2,X6)))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_58933, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0))|~r2_hidden(k1_xboole_0,X2))).#trainneg
cnf(i_0_63867, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r2_hidden(k4_algstr_0(X2,X3),a_3_1_zmodul01(X2,k4_struct_0(X2),X1))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_62114, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|X1!=k4_algstr_0(X2,X5)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_64609, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k4_algstr_0(X2,k1_algstr_0(X2,k4_struct_0(X2),X5)))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_62219, plain, (r2_hidden(k4_algstr_0(X1,X2),a_3_1_zmodul01(X1,X3,X4))|v2_struct_0(X1)|~l2_algstr_0(X1)|~r1_struct_0(X4,k4_algstr_0(X1,X3))|~r1_struct_0(X4,k4_algstr_0(X1,X2))|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_61979, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,X6)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X4,X6)|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_64096, negated_conjecture, (k5_algstr_0(esk1_0,X1,k1_algstr_0(esk3_0,X1,X2))=k4_algstr_0(esk3_0,X2)|~l2_algstr_0(esk3_0)|~m1_subset_1(k1_algstr_0(esk3_0,X1,X2),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X2,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_60445, plain, (k5_algstr_0(X1,k4_algstr_0(X1,X2),k4_algstr_0(X1,X3))=k4_algstr_0(X1,k5_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_66310, plain, (k5_algstr_0(X1,k4_algstr_0(X1,X3),X2)=k5_algstr_0(X1,k4_algstr_0(X1,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_67872, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X5,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_66465, negated_conjecture, (m1_subset_1(u1_struct_0(esk1_0),k1_xboole_0)|~v1_xboole_0(u1_struct_0(esk1_0))|~r2_hidden(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_66466, negated_conjecture, (r1_struct_0(esk1_0,esk2_0)|m1_subset_1(u1_struct_0(esk1_0),k1_xboole_0)|~r2_hidden(k1_xboole_0,k1_xboole_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_66833, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),k4_struct_0(X2)))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X4,X3))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_66532, plain, (r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X4,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r2_hidden(k4_algstr_0(X2,X3),a_3_1_zmodul01(X2,X4,X1))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_67870, plain, (r1_struct_0(X1,k4_algstr_0(X1,X2))|v2_struct_0(X1)|~r1_struct_0(X4,k4_algstr_0(X1,X3))|~r1_struct_0(X4,k4_algstr_0(X1,X2))|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_67241, plain, (r1_struct_0(X1,k5_algstr_0(X1,X2,k4_struct_0(X1)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_67771, negated_conjecture, (r2_hidden(X1,a_3_1_zmodul01(esk1_0,X2,X3))|X1!=k5_algstr_0(esk1_0,X2,X4)|~r1_struct_0(X3,k4_algstr_0(esk1_0,X4))|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_67869, plain, (r1_struct_0(X1,k5_algstr_0(X1,k4_algstr_0(X1,X2),X3))|v2_struct_0(X1)|~r1_struct_0(X4,k4_algstr_0(X1,X3))|~r1_struct_0(X4,k4_algstr_0(X1,X2))|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_67878, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k4_algstr_0(X2,X5),X6)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_65309, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|X1!=k4_algstr_0(X2,X4)|~l2_algstr_0(X2)|~r1_struct_0(X3,k5_algstr_0(X2,X5,X4))|~r1_struct_0(X3,k4_algstr_0(X2,X5))|~m1_zmodul01(X3,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_64023, plain, (k5_algstr_0(X1,k4_algstr_0(X1,X2),k5_algstr_0(X1,k4_struct_0(X1),X2))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_58932, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0))|~r2_hidden(k1_xboole_0,X2))).#trainneg
cnf(i_0_48849, negated_conjecture, (m1_subset_1(k1_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk3_0)),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_62413, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,k1_algstr_0(X2,X7,X6))|~l1_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X7)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_64173, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,k5_algstr_0(X1,X2,X4),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_62074, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k4_algstr_0(X2,k1_algstr_0(X2,X6,X5)))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X6,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_69613, plain, (k5_algstr_0(X1,k4_algstr_0(X1,X2),k4_algstr_0(X1,X2))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_60902, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,k5_algstr_0(X2,X6,X7))|~r1_struct_0(X4,k1_algstr_0(X2,X6,X5))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,k4_algstr_0(X2,X7))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X6,X7),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_62370, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_60904, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_58927, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)=u1_struct_0(esk3_0)|r2_hidden(esk2_0,u1_struct_0(esk3_0))|~r2_hidden(k1_xboole_0,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_63011, plain, (k4_algstr_0(X1,k1_algstr_0(X1,X2,k4_struct_0(X1)))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_61021, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,k4_struct_0(X2))|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_61014, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k5_algstr_0(X2,X5,X6),X7)|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X7))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_64229, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k4_algstr_0(X1,X3))=k5_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_60898, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,k1_algstr_0(X2,X6,X7))|~r1_struct_0(X4,k1_algstr_0(X2,X6,X5))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X7)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X6,X7),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_64164, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X4),k1_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k4_algstr_0(X1,X3),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_61005, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k1_algstr_0(X2,X5,X6),X6)|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_62851, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,k4_struct_0(X1),X2))=k1_algstr_0(X1,k4_algstr_0(X1,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_72833, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k1_algstr_0(X2,X5,X6),X6)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_73415, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),X3)=k5_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_58928, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)=u1_struct_0(esk1_0)|r2_hidden(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k1_xboole_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_57248, negated_conjecture, (m1_subset_1(k4_algstr_0(esk1_0,k4_struct_0(esk1_0)),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_71539, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X3,k4_struct_0(X2)))|~r1_struct_0(X4,X1)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X3,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_71286, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k4_algstr_0(X2,X5),X6)|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_62072, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k4_algstr_0(X2,k1_algstr_0(X2,X6,X5)))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X6,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_63190, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,X2,k4_struct_0(X1)))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_60601, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k4_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,X3))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_61880, plain, (k1_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X2))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_62852, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_algstr_0(X1,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_63240, plain, (k5_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_algstr_0(X1,X3),k1_algstr_0(X1,X3,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_62241, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,k4_algstr_0(X2,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X5)|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_57363, negated_conjecture, (r1_struct_0(esk1_0,k1_algstr_0(esk1_0,k4_struct_0(esk3_0),k4_struct_0(esk1_0)))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_60651, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k4_algstr_0(X2,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_60723, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|~r2_hidden(X1,a_3_1_zmodul01(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X1)|~m1_zmodul01(X4,X2)|~m1_zmodul01(X6,X2)|~m1_subset_1(esk8_4(X1,X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_62075, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k4_algstr_0(X2,k1_algstr_0(X2,X5,X6)))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_65717, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|~l2_algstr_0(X4)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_64020, plain, (k1_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2),k4_struct_0(X1))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_62043, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,k5_algstr_0(X2,X5,X6)))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_64658, plain, (r1_struct_0(X1,k5_algstr_0(X1,k4_algstr_0(X1,X2),X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_33798, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_algstr_0(X2,k4_struct_0(X2))))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r2_hidden(k1_algstr_0(X2,X3,X4),a_3_1_zmodul01(X2,X4,X1))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_37750, plain, (k5_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X2))=k1_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_13392, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k1_algstr_0(X2,X5,X3),X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_76200, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),k4_algstr_0(X2,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_76881, plain, (r1_tarski(a_3_1_zmodul01(X1,X2,X3),X4)|r2_hidden(esk4_2(a_3_1_zmodul01(X1,X2,X3),X4),a_3_1_zmodul01(X1,X5,X6))|v2_struct_0(X1)|~r1_struct_0(X6,esk4_2(a_3_1_zmodul01(X1,X2,X3),X4))|~r1_struct_0(X6,k4_algstr_0(X1,X5))|~m1_zmodul01(X6,X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_53555, plain, (r1_struct_0(X1,k5_algstr_0(X2,X4,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_62502, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_56077, plain, (r1_tarski(k1_zfmisc_1(k1_xboole_0),X1)|m1_subset_1(k1_xboole_0,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_76785, negated_conjecture, (m1_subset_1(k1_algstr_0(esk1_0,k4_struct_0(esk3_0),k4_struct_0(esk1_0)),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_56724, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,k4_struct_0(X2))|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_50117, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|X3!=k5_algstr_0(X2,k4_struct_0(X1),X4)|~l2_algstr_0(X1)|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_26420, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,k4_struct_0(X2)),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X4,X5))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_47825, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_algstr_0(X2,X5)))|~r1_struct_0(X1,X5)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_64623, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X6!=k1_algstr_0(X2,X5,k4_struct_0(X4))|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X4)|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_60162, plain, (m1_subset_1(X1,k1_xboole_0)|~r2_hidden(k1_xboole_0,k1_xboole_0)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_xboole_0))).#trainneg
cnf(i_0_19346, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k5_algstr_0(X2,X3,X6),k4_algstr_0(X2,X6))|~l2_algstr_0(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_49538, plain, (r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r2_hidden(k1_algstr_0(X2,X3,X4),a_3_1_zmodul01(X2,X3,X1))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_63121, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_46992, plain, (r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X4,X3),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_17341, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),k1_algstr_0(X1,k4_struct_0(X1),X3))=k5_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_77400, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k4_struct_0(X1))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_76879, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k4_algstr_0(X2,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_66519, plain, (k4_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X2),X2))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_79537, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,X6)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_78400, plain, (r1_tarski(a_3_1_zmodul01(X1,X2,X3),a_3_1_zmodul01(X1,X4,X5))|v2_struct_0(X1)|~r1_struct_0(X5,esk4_2(a_3_1_zmodul01(X1,X2,X3),a_3_1_zmodul01(X1,X4,X5)))|~r1_struct_0(X5,k4_algstr_0(X1,X4))|~m1_zmodul01(X5,X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_78438, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_78533, negated_conjecture, (r2_hidden(X1,a_3_1_zmodul01(esk1_0,X2,X3))|X1!=k1_algstr_0(esk1_0,X4,X5)|~r1_struct_0(X3,k4_algstr_0(esk1_0,X2))|~r1_struct_0(X3,X4)|~r1_struct_0(X3,X5)|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_79447, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,X6)|~l1_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_63455, plain, (r2_hidden(X1,k1_xboole_0)|~r2_hidden(k1_xboole_0,k1_xboole_0)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_xboole_0))).#trainneg
cnf(i_0_79076, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X4)|~r1_struct_0(X4,k4_algstr_0(X2,k1_algstr_0(X2,X5,k4_struct_0(X4))))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_78431, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,k4_algstr_0(X2,X6))|~r1_struct_0(X4,k1_algstr_0(X2,X6,X5))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_78726, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k1_algstr_0(X2,X6,X5),X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_78977, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_algstr_0(X2,X5)))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_78277, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~l1_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_79442, negated_conjecture, (r2_hidden(X1,a_3_1_zmodul01(esk1_0,X2,X3))|X1!=k5_algstr_0(esk1_0,X4,X5)|~r1_struct_0(X3,k4_algstr_0(esk1_0,X2))|~r1_struct_0(X3,k4_algstr_0(esk1_0,X5))|~r1_struct_0(X3,X4)|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_56072, plain, (r1_tarski(k1_zfmisc_1(k1_xboole_0),u1_struct_0(X1))|~l1_algstr_0(X1)|~r1_struct_0(X1,k1_xboole_0))).#trainneg
cnf(i_0_28676, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(k4_algstr_0(esk1_0,k4_struct_0(esk1_0)),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_78426, plain, (r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_16971, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X4),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_63246, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),k1_algstr_0(X1,X3,X2))=k4_algstr_0(X1,k4_algstr_0(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_33728, negated_conjecture, (r1_struct_0(esk1_0,k3_rlvect_1(esk1_0,X1,X2))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_56073, plain, (r1_tarski(k1_zfmisc_1(k1_xboole_0),u1_struct_0(X1))|~l2_algstr_0(X1)|~r1_struct_0(X1,k1_xboole_0))).#trainneg
cnf(i_0_71676, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k4_algstr_0(X2,X5),k4_struct_0(X2))|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_29218, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),k4_algstr_0(X1,X3))=k5_algstr_0(X1,k1_algstr_0(X1,X4,X2),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_82494, plain, (r1_struct_0(X1,k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),k4_algstr_0(X1,X3)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_81329, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k4_algstr_0(X2,k4_algstr_0(X2,X5))|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_80767, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,X2),a_3_1_zmodul01(esk1_0,X3,X2))|~r1_struct_0(X2,k4_algstr_0(esk1_0,X3))|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10581, plain, (r1_tarski(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|esk4_2(X1,a_3_1_zmodul01(X2,X3,X4))!=k1_algstr_0(X2,X3,k4_struct_0(X4))|~l2_algstr_0(X4)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_68447, plain, (r1_struct_0(X1,k4_algstr_0(X2,k5_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_46767, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l1_algstr_0(X2)|~l2_algstr_0(esk1_0)|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0))).#trainneg
cnf(i_0_51195, negated_conjecture, (r1_struct_0(esk3_0,k1_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0)))|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_56720, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,k5_algstr_0(X2,X3,X6))|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_44178, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),k4_algstr_0(X1,X2))=k5_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X3),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_46768, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~l2_algstr_0(X2)|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0))).#trainneg
cnf(i_0_39660, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k5_algstr_0(X2,X4,X5)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X1,k1_algstr_0(X2,X4,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_21299, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),k5_algstr_0(X1,X3,X2))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_27654, plain, (m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X1)|~l1_algstr_0(X1)|~m1_zmodul01(X1,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_80765, plain, (r1_tarski(a_3_1_zmodul01(X1,X2,X3),a_3_1_zmodul01(X1,X4,X3))|v2_struct_0(X1)|~r1_struct_0(X3,k4_algstr_0(X1,X4))|~r1_struct_0(X3,X2)|~m1_zmodul01(X3,X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_49883, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k1_algstr_0(X2,X5,k4_struct_0(X2))))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,k1_algstr_0(X2,X5,k4_struct_0(X2))),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_78639, plain, (X1=k1_zfmisc_1(k1_xboole_0)|m1_subset_1(k1_xboole_0,k1_zfmisc_1(k1_xboole_0))|~r1_tarski(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_63484, negated_conjecture, (k4_algstr_0(esk1_0,k4_struct_0(esk1_0))=k4_struct_0(esk3_0)|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~m1_subset_1(k1_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0)),u1_struct_0(esk3_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_33069, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X2))=k4_algstr_0(X1,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_84768, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X4),X3))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_83283, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,X2)=a_3_1_zmodul01(esk1_0,X3,X2)|~r1_tarski(a_3_1_zmodul01(esk1_0,X1,X2),a_3_1_zmodul01(esk1_0,X3,X2))|~r1_struct_0(X2,k4_algstr_0(esk1_0,X1))|~r1_struct_0(X2,X3)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_84816, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|~r2_hidden(X1,a_3_1_zmodul01(X2,X5,X4))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_85249, plain, (k5_algstr_0(X1,X2,X2)=k1_algstr_0(X1,X3,k4_algstr_0(X1,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_83231, plain, (r1_struct_0(X1,k5_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X2)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_41706, plain, (r1_struct_0(X1,k1_algstr_0(X1,X2,k4_algstr_0(X1,X2)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_33731, negated_conjecture, (m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(X1))|~l1_algstr_0(X1)|~l2_algstr_0(esk1_0)|~m1_zmodul01(X1,esk1_0))).#trainneg
cnf(i_0_85019, negated_conjecture, (k4_algstr_0(esk1_0,k4_struct_0(esk1_0))=k4_struct_0(esk3_0)|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~m1_subset_1(k4_algstr_0(esk1_0,k4_struct_0(esk1_0)),u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_51189, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X4,X3),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_17282, plain, (k5_algstr_0(X1,k4_algstr_0(X1,X2),k4_struct_0(X1))=k5_algstr_0(X1,k5_algstr_0(X1,X3,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_31553, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X4,k1_algstr_0(X2,X5,X3))))|~r1_struct_0(X1,k1_algstr_0(X2,X4,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_86767, negated_conjecture, (r1_struct_0(X1,k4_algstr_0(esk3_0,k4_struct_0(esk3_0)))|~m1_zmodul01(X1,esk3_0))).#trainneg
cnf(i_0_85662, negated_conjecture, (r1_tarski(k1_zfmisc_1(k1_xboole_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~m1_subset_1(k1_xboole_0,u1_struct_0(esk3_0))|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_87296, plain, (r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X4,X3),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_85735, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,X2)=a_3_1_zmodul01(esk1_0,X3,X2)|~r1_struct_0(X2,k4_algstr_0(esk1_0,X1))|~r1_struct_0(X2,k4_algstr_0(esk1_0,X3))|~r1_struct_0(X2,X3)|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_85734, plain, (r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_87096, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X4),k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r2_hidden(k5_algstr_0(X2,X3,X4),a_3_1_zmodul01(X2,X3,X1))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_87094, plain, (r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,k4_algstr_0(X2,X4),k4_struct_0(X2)),k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_85680, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X4),X3))|~r1_struct_0(X1,k1_algstr_0(X2,X4,k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X4,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_85719, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~r2_hidden(k4_algstr_0(X2,X4),a_3_1_zmodul01(X2,X3,X1))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_85702, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X4,X3)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_86347, plain, (m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X2)),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_86517, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk3_0,k4_struct_0(esk1_0))|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_85674, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X4,k4_struct_0(X2)),X3))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_85703, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X3,X4)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_85972, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X5,X5),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_85975, plain, (r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X4,X4),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_85676, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X4,k4_struct_0(X2)),X3))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_87428, plain, (k5_algstr_0(X1,k4_algstr_0(X1,X3),k5_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k4_algstr_0(X1,X2),k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_89237, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X4,X3))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_87772, negated_conjecture, (r1_struct_0(X1,k4_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~m1_zmodul01(X1,esk3_0))).#trainneg
cnf(i_0_88646, negated_conjecture, (m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_88442, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k5_algstr_0(X2,X3,X4)))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_88439, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~r2_hidden(X4,a_3_1_zmodul01(X2,X3,X5))|~r1_struct_0(X1,esk8_4(X4,X2,X3,X5))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_zmodul01(X5,X2)|~m1_subset_1(esk8_4(X4,X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_88228, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r2_hidden(k4_struct_0(X2),a_3_1_zmodul01(X2,X3,X1))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_87837, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k5_algstr_0(X2,X3,X5),k4_struct_0(X2))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_87083, plain, (k5_algstr_0(X1,k5_algstr_0(X1,k4_algstr_0(X1,X3),k4_struct_0(X1)),k5_algstr_0(X1,X2,X3))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_33732, negated_conjecture, (m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(X1))|~l2_algstr_0(esk1_0)|~l2_algstr_0(X1)|~m1_zmodul01(X1,esk1_0))).#trainneg
cnf(i_0_87777, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)=u1_struct_0(esk3_0)|~r1_tarski(u1_struct_0(esk3_0),k1_zfmisc_1(k1_xboole_0))|~l2_algstr_0(esk3_0)|~m1_subset_1(k1_xboole_0,u1_struct_0(esk3_0))|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_87078, plain, (r1_struct_0(X1,k5_algstr_0(X1,k4_algstr_0(X1,X3),k4_struct_0(X1)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_87838, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k4_struct_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,k5_algstr_0(X2,X3,X5)))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_58938, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_86594, negated_conjecture, (r1_struct_0(esk1_0,k5_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk3_0)))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_50155, negated_conjecture, (r1_struct_0(X1,k5_algstr_0(esk1_0,X2,k4_algstr_0(esk1_0,X3)))|~r1_struct_0(X1,X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,esk1_0)|~m1_subset_1(k4_algstr_0(esk1_0,X3),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_70248, plain, (r2_hidden(k1_algstr_0(X1,k5_algstr_0(X1,X3,X5),X5),a_3_1_zmodul01(X1,k4_algstr_0(X1,X2),X4))|v2_struct_0(X1)|~l2_algstr_0(X1)|~r1_struct_0(X4,k1_algstr_0(X1,X3,X2))|~m1_zmodul01(X4,X1)|~m1_subset_1(k1_algstr_0(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_65344, plain, (k4_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2))=k4_algstr_0(X1,k4_algstr_0(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_79514, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k4_algstr_0(X2,X5),X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_31548, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k5_algstr_0(X2,X3,X4)))|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_60170, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_71225, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k4_algstr_0(X2,X5),X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_37307, plain, (k1_algstr_0(X1,X2,k5_algstr_0(X1,k4_struct_0(X1),X2))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_90088, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~r2_hidden(X4,a_3_1_zmodul01(X2,X3,X1))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_89971, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_89933, plain, (r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X4)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_89964, negated_conjecture, (m1_subset_1(k4_struct_0(esk3_0),u1_struct_0(X1))|~l1_algstr_0(X1)|~l2_algstr_0(esk3_0)|~m1_zmodul01(X1,esk3_0))).#trainneg
cnf(i_0_90987, negated_conjecture, (r1_struct_0(esk1_0,k4_algstr_0(esk1_0,k4_struct_0(esk3_0)))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_90538, plain, (k5_algstr_0(X1,k5_algstr_0(X1,k5_algstr_0(X1,X4,X2),X4),k5_algstr_0(X1,X3,X2))=k4_algstr_0(X1,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_82503, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k5_algstr_0(X2,X3,X5),k4_algstr_0(X2,X5))|~l2_algstr_0(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_62605, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k4_algstr_0(X2,X5)|~l2_algstr_0(X4)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X4))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_26567, plain, (r1_struct_0(X1,k1_algstr_0(X2,X5,k3_rlvect_1(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k3_rlvect_1(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_92212, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k5_algstr_0(X1,k4_algstr_0(X1,X2),X3))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_71027, plain, (k4_algstr_0(X1,k5_algstr_0(X1,X2,X2))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_92229, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k3_rlvect_1(X2,X4,X3)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_92215, plain, (r1_tarski(a_3_1_zmodul01(X2,X3,X1),X4)|r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,esk4_2(a_3_1_zmodul01(X2,X3,X1),X4)))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_89965, negated_conjecture, (m1_subset_1(k4_struct_0(esk3_0),u1_struct_0(X1))|~l2_algstr_0(esk3_0)|~l2_algstr_0(X1)|~m1_zmodul01(X1,esk3_0))).#trainneg
cnf(i_0_92530, negated_conjecture, (m1_subset_1(k4_algstr_0(esk1_0,k4_struct_0(esk3_0)),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_92285, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X4,X5),X3))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k1_algstr_0(X2,X4,X5),X3),u1_struct_0(X2))|~m1_subset_1(k1_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_92228, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k3_rlvect_1(X2,X3,X4)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_92279, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k1_algstr_0(X2,X4,X5)))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,k1_algstr_0(X2,X4,X5)),u1_struct_0(X2))|~m1_subset_1(k1_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_92931, plain, (r2_hidden(k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),k4_algstr_0(X1,X3)),a_3_1_zmodul01(X1,X2,X4))|v2_struct_0(X1)|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_45480, negated_conjecture, (r1_struct_0(X1,k4_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(X1,k4_algstr_0(esk1_0,X2))|~r1_struct_0(X1,X2)|~m1_zmodul01(X1,esk1_0)|~m1_subset_1(k4_algstr_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(k4_algstr_0(esk1_0,X2),u1_struct_0(esk3_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_21306, plain, (k5_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k4_struct_0(X1),X2))=k5_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_15592, plain, (k1_algstr_0(X1,X3,k5_algstr_0(X1,X2,X3))=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_64195, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X4),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X3,X4)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_51932, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,k4_struct_0(X2),k4_struct_0(X2)),X5))|~r1_struct_0(X1,k5_algstr_0(X2,X3,k4_algstr_0(X2,X5)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_27456, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_71152, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k5_algstr_0(X2,X6,X7),X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X6,k4_algstr_0(X2,X5)))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,k4_algstr_0(X2,X7))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X6,X7),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_52299, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X5,X6),k4_algstr_0(X2,X6)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_15136, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk3_0,esk8_4(X1,X2,X1,esk3_0))|v2_struct_0(X2)|~l2_algstr_0(esk3_0)|~m1_zmodul01(esk3_0,X2)|~m1_subset_1(esk8_4(X1,X2,X1,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_58115, plain, (k5_algstr_0(X1,k4_algstr_0(X1,X3),k5_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_95746, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X5,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_95017, plain, (k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2))=k5_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_94385, plain, (r1_tarski(u1_struct_0(X1),a_3_1_zmodul01(X1,k1_algstr_0(X1,X2,X3),X4))|v2_struct_0(X1)|~r1_struct_0(X4,esk4_2(u1_struct_0(X1),a_3_1_zmodul01(X1,k1_algstr_0(X1,X2,X3),X4)))|~r1_struct_0(X4,k4_algstr_0(X1,X3))|~r1_struct_0(X4,k4_algstr_0(X1,X2))|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_94203, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2)|r1_struct_0(esk3_0,esk4_2(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(esk4_2(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_31638, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k1_algstr_0(X2,X3,X5),X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_algstr_0(X2,X6)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_57784, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),X2)=k5_algstr_0(X1,k5_algstr_0(X1,X4,X3),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_45393, negated_conjecture, (r1_struct_0(esk1_0,k1_algstr_0(esk1_0,k4_struct_0(esk1_0),X1))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_18115, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),X2)=k1_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_97312, plain, (r1_tarski(u1_struct_0(X1),a_3_1_zmodul01(X1,X2,X3))|v2_struct_0(X1)|~r1_struct_0(X3,esk4_2(u1_struct_0(X1),a_3_1_zmodul01(X1,X2,X3)))|~r1_struct_0(X3,k4_algstr_0(X1,X2))|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_97315, plain, (r1_tarski(k5_zmodul01(X1,X2,X3),a_3_1_zmodul01(X1,X4,X5))|v2_struct_0(X1)|~r1_struct_0(X5,esk4_2(k5_zmodul01(X1,X2,X3),a_3_1_zmodul01(X1,X4,X5)))|~r1_struct_0(X5,k4_algstr_0(X1,X4))|~m1_zmodul01(X5,X1)|~m1_zmodul01(X3,X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_97310, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),a_3_1_zmodul01(esk1_0,X2,esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(esk4_2(a_3_1_zmodul01(esk1_0,X1,esk3_0),a_3_1_zmodul01(esk1_0,X2,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_97302, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2)|m1_subset_1(esk4_2(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2),u1_struct_0(esk3_0))|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(esk4_2(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_96973, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,k4_struct_0(X2),X4),X3))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_97540, plain, (r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X4,X3),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_50663, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l1_algstr_0(X2)|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0))).#trainneg
cnf(i_0_97430, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k1_algstr_0(X2,X3,X5),X6)|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_98017, plain, (k5_algstr_0(X1,k4_algstr_0(X1,X3),k5_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k5_algstr_0(X1,X4,X2),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_97583, plain, (r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X5,X4),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_algstr_0(X2,X3),X4),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_96768, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X4,X5)))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_96765, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X5,X4)))|~r1_struct_0(X1,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X4),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_98491, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X4,X3),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_96936, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k5_algstr_0(X2,k4_struct_0(X2),X5)))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_96966, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,k4_struct_0(X2),X4),X3))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_50664, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l2_algstr_0(X2)|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0))).#trainneg
cnf(i_0_96975, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r2_hidden(k5_algstr_0(X2,k4_struct_0(X2),X4),a_3_1_zmodul01(X2,X3,X1))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_96065, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r2_hidden(k1_algstr_0(X2,X4,X3),a_3_1_zmodul01(X2,X4,X1))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_61980, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,X6)|~r1_struct_0(X4,X6)|~r1_struct_0(X4,X5)|~r1_struct_0(X4,X3)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_11485, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X3,X4),X2)=k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_98967, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k4_algstr_0(X2,k5_algstr_0(X2,X5,X6)))|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_99143, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~m1_zmodul01(esk1_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_100652, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r2_hidden(k5_algstr_0(X2,X4,X3),a_3_1_zmodul01(X2,X4,X1))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_98965, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k5_algstr_0(X2,X3,X4)))|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_98831, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),a_3_1_zmodul01(esk1_0,X2,esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_98817, negated_conjecture, (r1_tarski(u1_struct_0(X1),a_3_1_zmodul01(X1,X2,esk3_0))|v2_struct_0(X1)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(X1,X2))|~m1_zmodul01(esk3_0,X1)|~m1_subset_1(esk4_2(u1_struct_0(X1),a_3_1_zmodul01(X1,X2,esk3_0)),u1_struct_0(esk3_0))|~m1_subset_1(esk4_2(u1_struct_0(X1),a_3_1_zmodul01(X1,X2,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_98824, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r2_hidden(X1,a_3_1_zmodul01(esk1_0,X2,esk3_0))|~r1_struct_0(esk3_0,X2)|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_98920, plain, (r1_struct_0(X1,k4_algstr_0(X2,k5_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_99379, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,X7)|~r2_hidden(X5,a_3_1_zmodul01(X2,X3,X6))|~r1_struct_0(X4,esk8_4(X5,X2,X3,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X7))|~m1_zmodul01(X4,X2)|~m1_zmodul01(X6,X2)|~m1_subset_1(esk8_4(X5,X2,X3,X6),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X7),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_98957, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k5_algstr_0(X2,X5,X6),X5)|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_102029, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k1_algstr_0(X2,X3,X5),X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_98819, plain, (r1_tarski(u1_struct_0(X1),a_3_1_zmodul01(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X3)|~l2_algstr_0(X3)|~r1_struct_0(X3,k4_algstr_0(X1,X2))|~m1_zmodul01(X3,X1)|~m1_subset_1(esk4_2(u1_struct_0(X1),a_3_1_zmodul01(X1,X2,X3)),u1_struct_0(X3))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v4_rlvect_1(X3)|~v3_rlvect_1(X1)|~v3_rlvect_1(X3)|~v2_rlvect_1(X1)|~v13_algstr_0(X1)|~v13_algstr_0(X3))).#trainneg
cnf(i_0_98934, plain, (r1_struct_0(X1,k4_algstr_0(X2,k5_algstr_0(X2,k4_struct_0(X2),X3)))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_51682, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~l1_algstr_0(X2)|~l2_algstr_0(esk1_0)|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0))).#trainneg
cnf(i_0_99130, negated_conjecture, (r2_hidden(k4_algstr_0(esk1_0,k4_struct_0(esk3_0)),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_zmodul01(esk1_0,esk1_0))).#trainneg
cnf(i_0_99246, plain, (r2_hidden(k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4),a_3_1_zmodul01(X1,X2,X5))|v2_struct_0(X1)|~r1_struct_0(X5,k4_algstr_0(X1,X4))|~r1_struct_0(X5,X3)|~m1_zmodul01(X5,X1)|~m1_subset_1(k4_algstr_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_102168, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,k1_algstr_0(X2,k1_algstr_0(X2,X5,X6),X7)))|~r1_struct_0(X4,k1_algstr_0(X2,X6,X7))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X6,X7),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_99416, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k4_struct_0(X2),X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_99675, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X4),k5_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_101919, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),k1_algstr_0(X1,k1_algstr_0(X1,X4,X2),X3))=k4_algstr_0(X1,X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_102959, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2)|r2_hidden(esk4_2(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_103000, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_103005, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k5_algstr_0(X2,X3,X4),X5))|v2_struct_0(X2)|~r1_struct_0(X5,k1_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X5,k4_algstr_0(X2,X3))|~r1_struct_0(X5,X1)|~m1_zmodul01(X5,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_51683, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~l2_algstr_0(X2)|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0))).#trainneg
cnf(i_0_103140, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,X6)|~r2_hidden(X5,a_3_1_zmodul01(X2,X3,X4))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_102955, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,esk3_0)=a_3_1_zmodul01(esk1_0,X2,esk3_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,X2)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_103482, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k5_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_103194, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k5_algstr_0(X2,X5,X6),X5)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_102578, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|~r2_hidden(X1,u1_struct_0(X4))|~l1_struct_0(X4)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_103487, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k5_algstr_0(X2,k4_struct_0(X2),X3),X4))|v2_struct_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X4,X1)|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_102963, negated_conjecture, (r2_hidden(k1_algstr_0(esk1_0,X1,X2),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~r1_struct_0(esk3_0,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_104293, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k4_algstr_0(X2,X5),k5_algstr_0(X2,X6,X5))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X6,X5),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_31855, negated_conjecture, (r1_struct_0(esk3_0,k3_rlvect_1(esk1_0,X1,X2))|~l2_algstr_0(esk3_0)|~r1_struct_0(X3,X2)|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(esk8_4(k3_rlvect_1(esk1_0,X1,X2),esk1_0,X1,X3),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_91979, plain, (k1_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2),k5_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_66866, plain, (r2_hidden(k1_algstr_0(X1,k4_algstr_0(X1,X2),X2),a_3_1_zmodul01(X1,X3,X4))|v2_struct_0(X1)|~r1_struct_0(X4,k4_algstr_0(X1,X3))|~m1_zmodul01(X4,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_92895, plain, (k5_algstr_0(X1,k5_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X2),k5_algstr_0(X1,X3,X2))=k4_algstr_0(X1,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_92007, plain, (k5_algstr_0(X1,k4_struct_0(X1),k5_algstr_0(X1,k4_struct_0(X1),X2))=k4_algstr_0(X1,k4_algstr_0(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_43699, plain, (r1_struct_0(X1,k1_algstr_0(X2,k1_algstr_0(X2,k4_struct_0(X2),X3),X4))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_60900, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,k4_struct_0(X2))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_41209, plain, (X1=k1_zfmisc_1(X2)|r2_hidden(k1_xboole_0,k1_zfmisc_1(X2))|~r1_tarski(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_90950, negated_conjecture, (m1_subset_1(k5_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk3_0)),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_49878, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_105087, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,X2)=u1_struct_0(X2)|~l2_algstr_0(X2)|~l1_struct_0(X2)|~r1_struct_0(X2,k4_algstr_0(esk1_0,X1))|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_105067, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,esk3_0)=u1_struct_0(esk3_0)|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk3_0))|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_105928, negated_conjecture, (r1_struct_0(esk3_0,k1_algstr_0(esk1_0,X1,X2))|~l2_algstr_0(esk3_0)|~l1_struct_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~r1_struct_0(esk3_0,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_105937, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r2_hidden(X1,a_3_1_zmodul01(esk1_0,X2,X3))|~r1_struct_0(esk3_0,esk8_4(X1,esk1_0,X2,X3))|~r1_struct_0(esk3_0,X2)|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(esk8_4(X1,esk1_0,X2,X3),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_105212, plain, (r1_tarski(X1,a_3_1_zmodul01(X2,k5_algstr_0(X2,X3,X4),X5))|v2_struct_0(X2)|~r1_struct_0(X5,esk4_2(X1,a_3_1_zmodul01(X2,k5_algstr_0(X2,X3,X4),X5)))|~r1_struct_0(X5,k1_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X5,k4_algstr_0(X2,X3))|~m1_zmodul01(X5,X2)|~m1_subset_1(esk4_2(X1,a_3_1_zmodul01(X2,k5_algstr_0(X2,X3,X4),X5)),u1_struct_0(X2))|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_105397, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~m1_zmodul01(esk1_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_106107, negated_conjecture, (r2_hidden(k4_algstr_0(esk1_0,k4_struct_0(esk1_0)),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~r1_struct_0(esk3_0,k4_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_105091, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2)|r1_struct_0(esk3_0,esk4_2(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2))|~l2_algstr_0(esk3_0)|~l1_struct_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_106503, plain, (k5_algstr_0(X1,k4_struct_0(X1),k5_algstr_0(X1,k4_struct_0(X1),X2))=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_105089, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2)|m1_subset_1(esk4_2(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_105078, negated_conjecture, (a_3_1_zmodul01(esk1_0,k5_algstr_0(esk1_0,X1,X2),X3)=u1_struct_0(X3)|~l1_algstr_0(X3)|~r1_struct_0(X3,k5_algstr_0(esk1_0,X1,X2))|~r1_struct_0(X3,k4_algstr_0(esk1_0,X1))|~r1_struct_0(X3,X2)|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(k5_algstr_0(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_57595, plain, (r2_hidden(k5_algstr_0(X1,k5_algstr_0(X1,X4,X2),X4),a_3_1_zmodul01(X1,k4_struct_0(X1),X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~r1_struct_0(X3,k4_algstr_0(X1,X2))|~m1_zmodul01(X3,X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_87993, negated_conjecture, (a_3_1_zmodul01(esk1_0,k4_algstr_0(esk1_0,X1),X2)=a_3_1_zmodul01(esk1_0,X3,X2)|~r1_struct_0(X2,k4_algstr_0(esk1_0,X3))|~r1_struct_0(X2,k4_algstr_0(esk1_0,X1))|~r1_struct_0(X2,X3)|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_43456, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X5)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_53386, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k5_algstr_0(X2,X3,X5),X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_81694, plain, (m1_subset_1(X1,u1_struct_0(X2))|~l1_algstr_0(X2)|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0))|~r1_struct_0(X2,k1_xboole_0))).#trainneg
cnf(i_0_100906, negated_conjecture, (r2_hidden(k4_algstr_0(esk1_0,k4_struct_0(esk1_0)),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_35448, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X4),X4))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_107978, negated_conjecture, (r1_struct_0(esk3_0,X1)|~l2_algstr_0(esk3_0)|~r2_hidden(X1,a_3_1_zmodul01(esk1_0,X2,X3))|~r1_struct_0(esk3_0,esk8_4(X1,esk1_0,X2,X3))|~r1_struct_0(esk3_0,X2)|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(esk8_4(X1,esk1_0,X2,X3),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_108665, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,X2)=u1_struct_0(X2)|~l1_algstr_0(X2)|~r1_tarski(a_3_1_zmodul01(esk1_0,X1,X2),u1_struct_0(X2))|~r1_struct_0(X2,k5_algstr_0(esk1_0,X3,X4))|~r1_struct_0(X2,k4_algstr_0(esk1_0,X1))|~r1_struct_0(X2,k4_algstr_0(esk1_0,X3))|~r1_struct_0(X2,X4)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(k5_algstr_0(esk1_0,X3,X4),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_109037, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_107820, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,esk3_0)=u1_struct_0(esk3_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_107822, negated_conjecture, (a_3_1_zmodul01(esk1_0,k4_algstr_0(esk1_0,X1),esk3_0)=u1_struct_0(esk3_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_82296, plain, (m1_subset_1(X1,u1_struct_0(X2))|~l2_algstr_0(X2)|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0))|~r1_struct_0(X2,k1_xboole_0))).#trainneg
cnf(i_0_109255, negated_conjecture, (r1_struct_0(esk1_0,k4_algstr_0(esk1_0,k4_struct_0(esk1_0)))|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~l1_struct_0(esk1_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_107824, negated_conjecture, (a_3_1_zmodul01(esk1_0,k1_algstr_0(esk1_0,X1,X2),esk3_0)=u1_struct_0(esk3_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k1_algstr_0(esk1_0,X1,X2))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~m1_subset_1(k1_algstr_0(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_107825, negated_conjecture, (a_3_1_zmodul01(esk1_0,k5_algstr_0(esk1_0,X1,X2),esk3_0)=u1_struct_0(esk3_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k5_algstr_0(esk1_0,X1,X2))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X2)|~m1_subset_1(k5_algstr_0(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_33860, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,k4_struct_0(X1)))=k5_algstr_0(X1,k1_algstr_0(X1,X3,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_30889, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,k4_struct_0(X2),X4)))|~r1_struct_0(X1,k5_algstr_0(X2,X4,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X4,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_44013, plain, (r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X3,k4_struct_0(X2)),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_41212, plain, (r2_hidden(k1_xboole_0,k1_zfmisc_1(X3))|m1_subset_1(X1,X2)|~r2_hidden(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_105386, negated_conjecture, (r2_hidden(k1_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0)),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~m1_zmodul01(esk1_0,esk1_0))).#trainneg
cnf(i_0_18909, plain, (k5_algstr_0(X1,k5_algstr_0(X1,k5_algstr_0(X1,X2,X4),k4_algstr_0(X1,X4)),k1_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_51041, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),k5_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~r1_struct_0(X1,X5)|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_58937, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_105392, negated_conjecture, (r2_hidden(k5_algstr_0(esk1_0,k4_struct_0(esk1_0),k4_struct_0(esk1_0)),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~m1_zmodul01(esk1_0,esk1_0))).#trainneg
cnf(i_0_43348, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k4_struct_0(X1),X2))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_110236, negated_conjecture, (r1_struct_0(esk3_0,X1)|~l2_algstr_0(esk3_0)|~r2_hidden(X1,a_3_1_zmodul01(esk1_0,X2,esk3_0))|~r1_struct_0(esk3_0,X2)|~m1_subset_1(esk8_4(X1,esk1_0,X2,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_110413, negated_conjecture, (r1_struct_0(esk3_0,X1)|~l2_algstr_0(esk3_0)|~r2_hidden(X1,u1_struct_0(esk3_0))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,X2)|~m1_subset_1(k4_algstr_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_110401, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),a_3_1_zmodul01(esk1_0,X2,esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_110235, negated_conjecture, (r1_struct_0(esk3_0,X1)|~l2_algstr_0(esk3_0)|~r2_hidden(X1,a_3_1_zmodul01(esk1_0,X2,X3))|~r1_struct_0(esk3_0,X2)|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(esk8_4(X1,esk1_0,X2,X3),u1_struct_0(esk1_0))|~m1_subset_1(esk8_4(X1,esk1_0,X2,X3),u1_struct_0(esk3_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_112917, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k4_struct_0(X1),X2))=k5_algstr_0(X1,k1_algstr_0(X1,X3,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_60169, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~r2_hidden(X1,k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_90945, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k1_xboole_0))|m1_subset_1(esk9_1(k1_zfmisc_1(k1_xboole_0)),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_64744, plain, (r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_68780, plain, (k5_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,k4_struct_0(X1),X3))=k5_algstr_0(X1,k4_algstr_0(X1,X3),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_62311, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,X5)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_47339, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X4,X4),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_71358, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(esk2_0,u1_struct_0(esk3_0))|~r2_hidden(k1_xboole_0,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_94219, plain, (r1_tarski(X1,a_3_1_zmodul01(X2,X3,X6))|v2_struct_0(X2)|~r2_hidden(X3,a_3_1_zmodul01(X2,X4,X5))|~r1_struct_0(X6,k4_algstr_0(X2,esk8_4(X3,X2,X4,X5)))|~r1_struct_0(X6,esk4_2(X1,a_3_1_zmodul01(X2,X3,X6)))|~r1_struct_0(X6,k4_algstr_0(X2,X4))|~m1_zmodul01(X6,X2)|~m1_zmodul01(X5,X2)|~m1_subset_1(esk4_2(X1,a_3_1_zmodul01(X2,X3,X6)),u1_struct_0(X2))|~m1_subset_1(esk8_4(X3,X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_30888, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X4,k4_struct_0(X2))))|~r1_struct_0(X1,k5_algstr_0(X2,X4,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X4,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_57738, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k5_algstr_0(X2,X6,X5),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_51935, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,k4_struct_0(X2)),X5))|~r1_struct_0(X1,k5_algstr_0(X2,X3,k4_algstr_0(X2,X5)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_74375, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k1_xboole_0,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_91792, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k1_xboole_0))|m1_subset_1(esk9_1(k1_zfmisc_1(k1_xboole_0)),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_18906, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,k4_struct_0(X1)),k1_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_56619, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k5_algstr_0(X2,X5,k4_struct_0(X2)))|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_82424, plain, (r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X3,X4),k4_algstr_0(X2,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_35387, plain, (m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,X3),X3),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_78432, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|X3!=k5_algstr_0(X2,k4_struct_0(X1),k4_algstr_0(X2,X5))|~l2_algstr_0(X1)|~r1_struct_0(X1,k1_algstr_0(X2,X5,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_58721, negated_conjecture, (r2_hidden(X1,a_3_1_zmodul01(esk1_0,X2,X3))|X1!=k5_algstr_0(esk1_0,X2,k4_algstr_0(esk1_0,X4))|~r1_struct_0(X3,k4_algstr_0(esk1_0,X5))|~r1_struct_0(X3,X5)|~r1_struct_0(X3,X4)|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(k4_algstr_0(esk1_0,X5),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_108459, plain, (k4_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2))=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_115051, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),k4_struct_0(X2)))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X4,X3))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_113568, negated_conjecture, (r1_struct_0(esk3_0,X1)|X1!=k1_algstr_0(esk1_0,X2,k4_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X2)|~m1_subset_1(esk8_4(X1,esk1_0,X2,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_115066, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk3_0))|~l1_algstr_0(esk1_0)|~r2_hidden(k1_xboole_0,u1_struct_0(esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_115067, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk3_0))|~l2_algstr_0(esk1_0)|~r2_hidden(k1_xboole_0,u1_struct_0(esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_116215, plain, (m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,X2),X2),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_115329, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X6),X5))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,k4_algstr_0(X2,X6),X5),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_113557, negated_conjecture, (r1_struct_0(esk3_0,k3_rlvect_1(esk1_0,X1,X2))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~r1_struct_0(esk3_0,X2)|~m1_subset_1(esk8_4(k3_rlvect_1(esk1_0,X1,X2),esk1_0,X1,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_31623, plain, (m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X4))|v2_struct_0(X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_zmodul01(X1,X4)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X4)|~v5_zmodul01(X4)|~v4_zmodul01(X4)|~v3_zmodul01(X4)|~v2_zmodul01(X4)|~v4_rlvect_1(X4)|~v4_rlvect_1(X1)|~v3_rlvect_1(X4)|~v3_rlvect_1(X1)|~v2_rlvect_1(X4)|~v2_rlvect_1(X1)|~v13_algstr_0(X4)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_64734, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k4_algstr_0(X2,X5),X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,k1_algstr_0(X2,X5,X6)))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_56633, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k1_algstr_0(X2,k4_struct_0(X2),X6))|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_88943, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X4,k1_algstr_0(X2,X3,X5))))|~r1_struct_0(X1,k1_algstr_0(X2,X4,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_115409, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))|~l2_algstr_0(esk3_0)|~r2_hidden(k1_xboole_0,u1_struct_0(esk1_0))|~m1_subset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_115410, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))|~l1_algstr_0(esk3_0)|~r2_hidden(k1_xboole_0,u1_struct_0(esk1_0))|~m1_subset_1(u1_struct_0(esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_46205, plain, (r1_tarski(a_3_1_zmodul01(X2,X3,X1),X4)|r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),k4_struct_0(X2)))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,esk4_2(a_3_1_zmodul01(X2,X3,X1),X4)))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_30890, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k5_algstr_0(X2,X4,k4_struct_0(X2))))|~r1_struct_0(X1,k5_algstr_0(X2,X4,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X4,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_86045, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X6,X6),X5))|~r1_struct_0(X1,k5_algstr_0(X2,X3,k4_algstr_0(X2,X5)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_78725, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k5_algstr_0(X2,X5,X6),X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_53187, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k1_algstr_0(X2,X6,k4_struct_0(X2)))|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X6,X5))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X6,X5),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_115076, negated_conjecture, (X1=k1_xboole_0|r2_hidden(esk2_0,u1_struct_0(esk3_0))|~r2_hidden(k1_xboole_0,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_50478, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),k1_algstr_0(X2,X4,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_46989, plain, (r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_25045, plain, (m1_subset_1(k5_algstr_0(X1,X2,X2),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_19217, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~l2_algstr_0(X4)|~l2_algstr_0(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X4))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_118447, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_115424, negated_conjecture, (X1=k1_xboole_0|r2_hidden(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(k1_xboole_0,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_117545, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X6,X5))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_117084, negated_conjecture, (r1_struct_0(esk3_0,k3_rlvect_1(esk1_0,X1,X2))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~r1_struct_0(esk3_0,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_116885, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),k4_struct_0(X2)))|v2_struct_0(X2)|~r2_hidden(X4,a_3_1_zmodul01(X2,X3,X1))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_116923, plain, (m1_subset_1(k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_62850, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,k4_struct_0(X1),X2))=k4_algstr_0(X1,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_48017, plain, (r1_tarski(a_3_1_zmodul01(X2,X5,X4),X6)|r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,esk4_2(a_3_1_zmodul01(X2,X5,X4),X6))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(esk4_2(a_3_1_zmodul01(X2,X5,X4),X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_86490, plain, (r1_struct_0(X1,k5_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2),k4_algstr_0(X1,X2)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_87035, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X4),X3))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_algstr_0(X2,X4),X3),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_71716, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k5_algstr_0(X2,X5,X6),X7)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,X7))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_11346, plain, (k1_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),k5_algstr_0(X1,X3,X4))=k5_algstr_0(X1,k1_algstr_0(X1,X3,X2),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_86732, negated_conjecture, (k5_algstr_0(esk1_0,X1,X1)=k4_struct_0(esk3_0)|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_118565, negated_conjecture, (m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_43350, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k4_struct_0(X1),X2))=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_121090, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X4,X3),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,k5_algstr_0(X2,X4,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(k5_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_120917, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k4_struct_0(X2),X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_119573, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k4_algstr_0(X2,k4_algstr_0(X2,X6)))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_120079, plain, (m1_subset_1(k4_algstr_0(X1,k4_algstr_0(X1,X2)),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_119807, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k1_algstr_0(X2,X7,X5))|~r1_struct_0(X4,k4_algstr_0(X2,X7))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_121822, negated_conjecture, (m1_subset_1(u1_struct_0(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_119833, negated_conjecture, (r1_struct_0(esk3_0,k1_algstr_0(esk3_0,X1,X2))|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~r1_struct_0(esk3_0,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_119550, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X7))|~r1_struct_0(X4,k4_algstr_0(X2,X7))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_101859, plain, (r2_hidden(k1_algstr_0(X1,k1_algstr_0(X1,X4,X2),X3),a_3_1_zmodul01(X1,X4,X5))|v2_struct_0(X1)|~l2_algstr_0(X1)|~r1_struct_0(X5,k1_algstr_0(X1,X2,X3))|~m1_zmodul01(X5,X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_66616, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X3,X4)))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_121789, negated_conjecture, (r1_struct_0(X1,k4_struct_0(esk3_0))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_zmodul01(X1,esk1_0)|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_118574, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|~r1_tarski(u1_struct_0(esk1_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_69166, plain, (m1_subset_1(k5_algstr_0(X1,X2,k4_struct_0(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_26547, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k3_rlvect_1(X2,X5,X6))|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k3_rlvect_1(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_53735, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~r1_struct_0(X1,X6)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_115070, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|r2_hidden(esk2_0,u1_struct_0(esk3_0))|~r2_hidden(k1_xboole_0,u1_struct_0(esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_118620, negated_conjecture, (r1_struct_0(esk3_0,k4_algstr_0(esk1_0,k4_struct_0(esk1_0)))|~l2_algstr_0(esk3_0)|~r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~r1_struct_0(esk3_0,k4_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_51047, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k5_algstr_0(X2,X5,X6),k4_struct_0(X2))|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_46213, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),k4_struct_0(X2)))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k3_rlvect_1(X2,X3,X4)))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_119578, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k4_algstr_0(X2,k4_algstr_0(X2,X5)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_112642, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k1_xboole_0))|r2_hidden(esk9_1(k1_zfmisc_1(k1_xboole_0)),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_18323, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,X2),X3)|m1_subset_1(esk4_2(a_3_1_zmodul01(esk1_0,X1,X2),X3),u1_struct_0(X2))|~l1_algstr_0(X2)|~r1_struct_0(X2,X1)|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_32870, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,k4_algstr_0(X1,X2),X2))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_125393, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2)|r1_struct_0(esk3_0,esk4_2(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2))|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(esk4_2(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_125074, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~r1_struct_0(X1,k1_algstr_0(X2,X6,X5))|~r1_struct_0(X1,k4_algstr_0(X2,X6))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X6,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_122866, plain, (r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_123780, negated_conjecture, (r1_struct_0(esk3_0,k5_algstr_0(esk3_0,X2,X1))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk3_0,X1))|~r1_struct_0(esk3_0,X2)|~m1_subset_1(k4_algstr_0(esk3_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk3_0))|~m1_subset_1(X2,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_50917, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X5),X5))|~r1_struct_0(X1,k4_struct_0(X2))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_123677, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_struct_0(esk1_0)|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_95371, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),k1_algstr_0(X1,X2,k4_struct_0(X1)))=k4_algstr_0(X1,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_72499, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X5),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,k1_algstr_0(X2,X3,X5)))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_35786, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k5_algstr_0(X1,X3,X3))=k5_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_63335, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X2,X3))=k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_127732, plain, (k5_algstr_0(X1,X3,k1_algstr_0(X1,X3,X2))=k5_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_125872, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),u1_struct_0(esk3_0))|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(esk4_2(a_3_1_zmodul01(esk1_0,X1,esk3_0),u1_struct_0(esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_127734, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X6,k1_algstr_0(X2,X6,X5)))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_125874, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),a_3_1_zmodul01(esk1_0,X2,esk3_0))|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(esk4_2(a_3_1_zmodul01(esk1_0,X1,esk3_0),a_3_1_zmodul01(esk1_0,X2,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_125862, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2)|m1_subset_1(esk4_2(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2),u1_struct_0(esk3_0))|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(esk4_2(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_127733, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X2,X3))=k5_algstr_0(X1,X4,k1_algstr_0(X1,X4,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_129058, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),u1_struct_0(esk3_0))|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_129296, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2)|r2_hidden(esk4_2(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2),u1_struct_0(esk3_0))|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(esk4_2(a_3_1_zmodul01(esk1_0,X1,esk3_0),X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_128772, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_128997, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X3,X2))=k5_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_126461, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|~l2_struct_0(esk1_0)|~r1_tarski(u1_struct_0(esk1_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_123680, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_zmodul01(esk1_0,esk1_0))).#trainneg
cnf(i_0_128571, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X4,k1_algstr_0(X2,X4,X3)))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_128569, plain, (r1_struct_0(X1,k5_algstr_0(X2,X4,k1_algstr_0(X2,X4,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_128805, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~r1_struct_0(X1,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_129500, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X5,k1_algstr_0(X2,X5,X4)))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_128552, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X6,k1_algstr_0(X2,X6,X5))|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_130256, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,esk3_0)=u1_struct_0(esk3_0)|~l1_algstr_0(esk3_0)|~l1_struct_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_130876, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~r1_struct_0(X1,k1_algstr_0(X2,X6,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X6,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_131134, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|~l2_struct_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_zmodul01(esk1_0,esk3_0))).#trainneg
cnf(i_0_129896, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k4_algstr_0(X1,X3),k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_129295, negated_conjecture, (r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),a_3_1_zmodul01(esk1_0,X2,esk3_0))|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_129504, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X5,k1_algstr_0(X2,X5,X4)))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_130843, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X1,k1_algstr_0(X2,X5,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_131157, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_129193, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X6,k5_algstr_0(X2,X6,X5)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_129487, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X6,k1_algstr_0(X2,X6,X5)))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_126467, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|~r1_tarski(u1_struct_0(esk1_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_126713, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),k1_algstr_0(X1,k4_struct_0(X1),X2))=k4_algstr_0(X1,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_131101, plain, (k5_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2),k5_algstr_0(X1,X3,X2))=k5_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_132486, plain, (k5_algstr_0(X1,X4,k1_algstr_0(X1,X4,X2))=k5_algstr_0(X1,k5_algstr_0(X1,X3,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_132997, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,esk3_0)=a_3_1_zmodul01(esk1_0,X2,esk3_0)|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,X2)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_132194, negated_conjecture, (a_3_1_zmodul01(esk1_0,k4_algstr_0(esk1_0,X1),esk3_0)=u1_struct_0(esk3_0)|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_132669, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X5),k4_struct_0(X2)))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_132197, negated_conjecture, (a_3_1_zmodul01(esk1_0,k5_algstr_0(esk1_0,X1,X2),esk3_0)=u1_struct_0(esk3_0)|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k5_algstr_0(esk1_0,X1,X2))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X2)|~m1_subset_1(k5_algstr_0(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_132313, plain, (r1_struct_0(X1,k5_algstr_0(X1,X3,k1_algstr_0(X1,X3,X2)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_132993, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,esk3_0)=a_3_1_zmodul01(esk1_0,X2,esk3_0)|~l1_algstr_0(esk3_0)|~r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),a_3_1_zmodul01(esk1_0,X2,esk3_0))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_132196, negated_conjecture, (a_3_1_zmodul01(esk1_0,k1_algstr_0(esk1_0,X1,X2),esk3_0)=u1_struct_0(esk3_0)|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k1_algstr_0(esk1_0,X1,X2))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~m1_subset_1(k1_algstr_0(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_132273, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~m1_zmodul01(esk1_0,esk3_0))).#trainneg
cnf(i_0_132782, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X5),k4_struct_0(X2)))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_132702, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X6),k4_struct_0(X2)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~r1_struct_0(X1,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_135504, negated_conjecture, (r1_struct_0(esk3_0,X1)|~l1_algstr_0(esk3_0)|~r2_hidden(X1,u1_struct_0(esk3_0))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,X2)|~m1_subset_1(k4_algstr_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_136381, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X6))|~r1_struct_0(X1,k1_algstr_0(X2,X6,X5))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_134685, plain, (r1_struct_0(X1,k5_algstr_0(X2,X4,k1_algstr_0(X2,X4,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_135229, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_135491, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),a_3_1_zmodul01(esk1_0,X2,esk3_0))|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_134991, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X6,X5),X6))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_134766, plain, (k5_algstr_0(X1,X4,k1_algstr_0(X1,X4,X3))=k5_algstr_0(X1,k5_algstr_0(X1,X2,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_137821, plain, (r1_struct_0(X1,k5_algstr_0(X2,X5,k1_algstr_0(X2,X5,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_137835, plain, (k5_algstr_0(X1,X3,k1_algstr_0(X1,X3,X2))=k1_algstr_0(X1,k4_algstr_0(X1,X2),k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_136276, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|~l2_algstr_0(esk1_0)|~m1_zmodul01(esk1_0,esk3_0))).#trainneg
cnf(i_0_136154, plain, (r1_struct_0(X1,k5_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~r2_hidden(X3,a_3_1_zmodul01(X1,X2,X4))|~m1_zmodul01(X4,X1)|~m1_subset_1(esk8_4(X3,X1,X2,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_132193, negated_conjecture, (a_3_1_zmodul01(esk1_0,k4_algstr_0(esk1_0,X1),esk3_0)=u1_struct_0(esk3_0)|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k1_algstr_0(esk1_0,k4_struct_0(esk1_0),X1))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_137863, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X2,k4_struct_0(X1)))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_132198, negated_conjecture, (a_3_1_zmodul01(esk1_0,k5_algstr_0(esk1_0,X1,X2),esk3_0)=u1_struct_0(esk3_0)|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k1_algstr_0(esk1_0,k4_struct_0(esk1_0),X2))|~r1_struct_0(esk3_0,k5_algstr_0(esk1_0,X1,X2))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~m1_subset_1(k5_algstr_0(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_135004, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X5,X4),X5))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_137837, plain, (k5_algstr_0(X1,X3,k1_algstr_0(X1,X3,X2))=k1_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_139482, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|~m1_zmodul01(esk1_0,esk3_0))).#trainneg
cnf(i_0_138588, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~r2_hidden(X5,a_3_1_zmodul01(X2,X4,X1))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_139271, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_algstr_0(X2,X5),k4_struct_0(X2)))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_138842, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X6,k1_algstr_0(X2,X6,X5)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_141006, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),k4_algstr_0(X2,X5)))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_139146, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X5),k4_struct_0(X2)))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_135242, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X4),X4),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_134743, plain, (r1_struct_0(X1,k5_algstr_0(X2,X5,k1_algstr_0(X2,X5,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),X4))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_algstr_0(X2,X3),X4),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_135457, negated_conjecture, (a_3_1_zmodul01(esk1_0,k4_algstr_0(esk1_0,X1),esk3_0)=a_3_1_zmodul01(esk1_0,X2,esk3_0)|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X2)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_126464, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l2_struct_0(esk1_0)|~l2_algstr_0(esk3_0)|~r2_hidden(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_124748, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|~r1_tarski(u1_struct_0(esk1_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_132454, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X6,k1_algstr_0(X2,X6,X5)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,k4_algstr_0(X2,X5)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_140294, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~r2_hidden(k5_algstr_0(X2,X4,X5),a_3_1_zmodul01(X2,X4,X1))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_134626, plain, (r1_struct_0(X1,k5_algstr_0(X2,X5,k1_algstr_0(X2,X5,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r2_hidden(k5_algstr_0(X2,X3,X4),a_3_1_zmodul01(X2,X3,X1))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_135008, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X5,X4),X5))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_143320, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X5),X4))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_algstr_0(X2,X5),X4),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_143164, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X4),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_143323, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X5),X4))|~r1_struct_0(X1,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_algstr_0(X2,X5),X4),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_142626, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X5),X4))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_142629, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X5),X4))|~r1_struct_0(X1,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X5,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_126463, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~l2_struct_0(esk1_0)|~l2_algstr_0(esk3_0)|~r2_hidden(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_131136, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|~l1_algstr_0(esk1_0)|~r1_tarski(u1_struct_0(esk1_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~m1_zmodul01(esk1_0,esk1_0))).#trainneg
cnf(i_0_137122, plain, (k5_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X2)=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_135520, negated_conjecture, (r1_struct_0(esk1_0,X1)|~l1_algstr_0(esk3_0)|~r2_hidden(X1,u1_struct_0(esk3_0))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,X2)|~m1_subset_1(k4_algstr_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_133676, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X5,k5_algstr_0(X2,X5,X6)))|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_algstr_0(X2,X6)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_133701, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_algstr_0(X2,X5)))|~r1_struct_0(X1,k5_algstr_0(X2,X5,k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_135547, negated_conjecture, (r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~l1_algstr_0(esk3_0)|~r2_hidden(k5_algstr_0(esk1_0,k4_algstr_0(esk1_0,X2),X1),u1_struct_0(esk3_0))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,X2)|~m1_subset_1(k4_algstr_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_129187, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X6,k1_algstr_0(X2,X5,X6)))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_135934, plain, (m1_subset_1(k5_algstr_0(X1,X2,k1_algstr_0(X1,X2,X3)),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_135024, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X7,X6),X7))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~r1_struct_0(X1,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_132935, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X3,X2))=k5_algstr_0(X1,k4_algstr_0(X1,X3),k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_144932, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X5),k4_struct_0(X2)))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_144937, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_algstr_0(X2,X5),k4_struct_0(X2)))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_144938, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),k4_algstr_0(X2,X5)))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_145950, plain, (r1_struct_0(X1,k5_algstr_0(X1,k4_algstr_0(X1,X2),k4_struct_0(X1)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_144935, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,k4_struct_0(X2)),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_125132, negated_conjecture, (m1_subset_1(k4_algstr_0(esk1_0,k4_struct_0(esk1_0)),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_struct_0(esk1_0))|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_145908, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X6),k4_struct_0(X2)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~r1_struct_0(X1,k1_algstr_0(X2,X6,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X6,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_131435, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k4_struct_0(X2),k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_144936, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k1_algstr_0(X2,k4_algstr_0(X2,X6),X6),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_140878, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),k4_algstr_0(X2,X5)))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_126912, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_128921, plain, (k5_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),k1_algstr_0(X1,X3,X2))=k5_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_135552, negated_conjecture, (r2_hidden(X1,a_3_1_zmodul01(esk1_0,X2,esk3_0))|~l1_algstr_0(esk3_0)|~r2_hidden(X1,u1_struct_0(esk3_0))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X3))|~r1_struct_0(esk3_0,X3)|~m1_subset_1(k4_algstr_0(esk1_0,X3),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_134176, plain, (k5_algstr_0(X1,k4_algstr_0(X1,X3),k1_algstr_0(X1,k4_struct_0(X1),X2))=k4_algstr_0(X1,k1_algstr_0(X1,X3,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_144978, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|m1_subset_1(u2_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_struct_0(esk3_0)|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_140301, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X4),X5))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_algstr_0(X2,X4),X5),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_142619, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,k4_struct_0(X2),X5),X4))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_137855, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X4,k1_algstr_0(X2,X4,X3)))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_133358, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_133345, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X4,X3),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_147192, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X6),X5))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_131682, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X6))|~r1_struct_0(X1,k1_algstr_0(X2,X6,X5))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_134161, plain, (k5_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),k1_algstr_0(X1,k4_struct_0(X1),X2))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_146596, negated_conjecture, (m1_subset_1(k4_algstr_0(esk3_0,k4_struct_0(esk3_0)),u1_struct_0(esk1_0))|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~m1_zmodul01(esk3_0,esk3_0))).#trainneg
cnf(i_0_137822, plain, (r1_struct_0(X1,k5_algstr_0(X2,X5,k1_algstr_0(X2,X5,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_130251, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,esk3_0)=u1_struct_0(esk3_0)|~l1_algstr_0(esk3_0)|~r1_tarski(u1_struct_0(esk3_0),a_3_1_zmodul01(esk1_0,X1,esk3_0))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_144933, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k5_algstr_0(X2,X6,X5),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_132237, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X5),X6))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X6))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_124678, negated_conjecture, (r1_struct_0(X1,k4_struct_0(esk3_0))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_zmodul01(X1,esk1_0))).#trainneg
cnf(i_0_148276, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(u2_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_139217, plain, (r1_struct_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X3),k4_struct_0(X1)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_139179, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X6),k4_struct_0(X2)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~r1_struct_0(X1,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_129978, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X3),k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_135009, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X5,X4),X5))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_140005, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,k4_struct_0(X1),X2))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_149989, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X6),X5))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,k4_algstr_0(X2,X6),X5),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_137009, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,k4_struct_0(X2)),X5))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_131927, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X4),k4_struct_0(X2)))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_126203, negated_conjecture, (r1_struct_0(esk3_0,k5_algstr_0(esk3_0,X1,X2))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(k4_algstr_0(esk3_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_150165, negated_conjecture, (r1_struct_0(esk1_0,u2_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_131399, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k4_algstr_0(X2,k1_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_144638, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_145875, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X5),k4_struct_0(X2)))|~r1_struct_0(X1,k1_algstr_0(X2,X5,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_133870, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,k5_algstr_0(X2,X5,X6)))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X1,k4_algstr_0(X2,X6))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X2))|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_129003, plain, (k5_algstr_0(X1,X2,k5_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_128527, plain, (k1_algstr_0(X1,k5_algstr_0(X1,X3,k1_algstr_0(X1,X3,X2)),X2)=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_153059, plain, (k5_algstr_0(X1,X3,k5_algstr_0(X1,X3,X2))=k1_algstr_0(X1,k4_struct_0(X1),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_153060, plain, (k5_algstr_0(X1,X4,k5_algstr_0(X1,X4,X2))=k5_algstr_0(X1,k1_algstr_0(X1,X3,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_153057, plain, (k5_algstr_0(X1,X4,k5_algstr_0(X1,X4,X2))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_153061, plain, (k5_algstr_0(X1,X3,k5_algstr_0(X1,X3,X2))=k1_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_153058, plain, (k5_algstr_0(X1,X4,k5_algstr_0(X1,X4,X2))=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_153062, plain, (k5_algstr_0(X1,X4,k5_algstr_0(X1,X4,X2))=k1_algstr_0(X1,k4_algstr_0(X1,X3),k1_algstr_0(X1,X3,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_159077, plain, (r1_struct_0(X1,k5_algstr_0(X2,X5,k5_algstr_0(X2,X5,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_151784, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X6),X5))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_153063, plain, (r1_struct_0(X1,k5_algstr_0(X1,X3,k5_algstr_0(X1,X3,X2)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_159308, plain, (k5_algstr_0(X1,X2,k5_algstr_0(X1,X2,k4_struct_0(X1)))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_158741, plain, (k5_algstr_0(X1,X2,k5_algstr_0(X1,X2,X3))=k1_algstr_0(X1,k5_algstr_0(X1,X4,X4),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_154267, negated_conjecture, (m1_subset_1(u2_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_struct_0(esk3_0))).#trainneg
cnf(i_0_154289, negated_conjecture, (m1_subset_1(k4_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_160256, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k5_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~r1_struct_0(X1,X4)|~r1_struct_0(X1,X5)|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_160283, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X6),X5))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_161599, plain, (r1_struct_0(X1,k5_algstr_0(X2,X4,k5_algstr_0(X2,X4,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_160165, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k5_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X4,X5))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_160286, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X6),X5))|~r1_struct_0(X4,k1_algstr_0(X2,X6,k4_struct_0(X2)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X6,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_161579, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X6,k5_algstr_0(X2,X6,X5))|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_161600, plain, (r1_struct_0(X1,k5_algstr_0(X2,X4,k5_algstr_0(X2,X4,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_154294, negated_conjecture, (m1_subset_1(k4_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_struct_0(esk3_0))).#trainneg
cnf(i_0_161592, plain, (k5_algstr_0(X1,X3,k5_algstr_0(X1,X3,X2))=k4_algstr_0(X1,k4_algstr_0(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_162576, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_algstr_0(X2,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_160390, plain, (m1_subset_1(k5_algstr_0(X1,X2,k5_algstr_0(X1,X2,X3)),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_163142, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,k4_algstr_0(X2,X6))|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_154258, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),X1)|m1_subset_1(esk4_2(u1_struct_0(esk3_0),X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_154552, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~r1_struct_0(X1,k5_algstr_0(X2,X5,X6))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_163275, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_160292, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k1_algstr_0(X2,k4_struct_0(X2),X6),X5))|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_160284, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k5_algstr_0(X2,k4_struct_0(X2),X6),X5))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_154298, negated_conjecture, (m1_subset_1(k4_struct_0(X1),u1_struct_0(esk1_0))|~l2_algstr_0(X1)|~m1_zmodul01(X1,esk3_0))).#trainneg
cnf(i_0_161611, plain, (k5_algstr_0(X1,X2,k5_algstr_0(X1,X2,X3))=k5_algstr_0(X1,X4,k5_algstr_0(X1,X4,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_164485, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k1_algstr_0(X2,X6,X5))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_165274, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X4,k5_algstr_0(X2,X6,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_164960, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),a_3_1_zmodul01(esk1_0,X1,X2))|~r1_struct_0(X2,esk4_2(u1_struct_0(esk3_0),a_3_1_zmodul01(esk1_0,X1,X2)))|~r1_struct_0(X2,k4_algstr_0(esk1_0,X1))|~m1_zmodul01(X2,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_164403, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k1_algstr_0(X2,k4_algstr_0(X2,X6),X5))|~r1_struct_0(X4,k1_algstr_0(X2,X6,k4_struct_0(X2)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_158904, plain, (k5_algstr_0(X1,X2,k4_algstr_0(X1,k4_struct_0(X1)))=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_156735, plain, (k5_algstr_0(X1,X2,k4_algstr_0(X1,k4_struct_0(X1)))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_166897, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),a_3_1_zmodul01(esk1_0,X1,esk3_0))|~l1_struct_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_164962, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_166840, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X7,k1_algstr_0(X2,X7,X6)))|~r1_struct_0(X4,k5_algstr_0(X2,X6,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_166841, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X7,k1_algstr_0(X2,X6,X7)))|~r1_struct_0(X4,k5_algstr_0(X2,X6,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_165388, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k1_algstr_0(X2,X5,X6))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_154775, plain, (k1_algstr_0(X1,k4_struct_0(X1),k1_algstr_0(X1,X2,X3))=k5_algstr_0(X1,X2,k4_algstr_0(X1,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_165849, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X5,k5_algstr_0(X2,X5,X4)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_164964, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0))).#trainneg
cnf(i_0_162173, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(u2_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_struct_0(esk3_0))).#trainneg
cnf(i_0_153280, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),k4_algstr_0(X2,X5)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_158237, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X6,k5_algstr_0(X2,X6,X3))|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_165845, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X5,k5_algstr_0(X2,X5,X4)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_161107, plain, (k5_algstr_0(X1,X2,k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_169224, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,esk3_0)=u1_struct_0(esk3_0)|~r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),u1_struct_0(esk3_0))|~l1_struct_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_170256, plain, (k5_algstr_0(X1,X2,k4_algstr_0(X1,X3))=k1_algstr_0(X1,X2,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_169237, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|~r1_tarski(u1_struct_0(esk1_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_170739, negated_conjecture, (u1_struct_0(esk3_0)=u1_struct_0(esk1_0)|~l1_algstr_0(esk1_0)|~r1_tarski(u1_struct_0(esk1_0),u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_171854, plain, (k5_algstr_0(X1,X2,k4_algstr_0(X1,k4_algstr_0(X1,X2)))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_169226, negated_conjecture, (r2_hidden(X1,a_3_1_zmodul01(esk1_0,X2,esk3_0))|~r2_hidden(X1,u1_struct_0(esk3_0))|~l1_struct_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_169899, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k4_struct_0(X1))=k5_algstr_0(X1,X2,k4_algstr_0(X1,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_158500, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X4,X5),X5))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_172358, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_172168, plain, (r1_struct_0(X1,k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k4_algstr_0(X1,X3)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_171917, plain, (k4_algstr_0(X1,k5_algstr_0(X1,X2,X3))=k1_algstr_0(X1,k4_algstr_0(X1,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_172298, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,X3,X2))=k5_algstr_0(X1,k1_algstr_0(X1,X3,X4),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_173904, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_169240, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~r2_hidden(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_170746, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r1_struct_0(esk1_0,u2_struct_0(esk3_0))|~l2_struct_0(esk3_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_173897, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r2_hidden(X5,a_3_1_zmodul01(X2,X6,X4))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_174538, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_algstr_0(X2,X5),X6))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173818, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173845, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X6),k4_struct_0(X2)))|~r1_struct_0(X4,k1_algstr_0(X2,X6,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X6,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173820, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X6),k4_struct_0(X2)))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173839, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X7,k5_algstr_0(X2,X7,X6)))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173819, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X7,k1_algstr_0(X2,X7,X6)))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173826, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_algstr_0(X2,X6),k4_struct_0(X2)))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_170742, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~r2_hidden(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_162179, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(k4_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_173844, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X4,k1_algstr_0(X2,X6,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X6,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173824, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,k4_struct_0(X2)),X6))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_175049, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X4),k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173830, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),k4_algstr_0(X2,X6)))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173827, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),k4_algstr_0(X2,X6)))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_174535, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),X4))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_174274, plain, (r1_struct_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X2),X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_174460, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r2_hidden(k1_algstr_0(X2,k4_algstr_0(X2,X4),X5),a_3_1_zmodul01(X2,X3,X1))|~r1_struct_0(X1,k5_algstr_0(X2,X4,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173835, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_algstr_0(X2,k4_struct_0(X2)),X6))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173890, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173834, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k1_algstr_0(X2,X7,X6),X7))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_169239, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~r2_hidden(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_163500, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(k4_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_struct_0(esk3_0))).#trainneg
cnf(i_0_173874, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X7))|~r1_struct_0(X4,X7)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_175924, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X6,k4_struct_0(X2)))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_178352, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k5_algstr_0(X2,X6,k4_struct_0(X2)))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173887, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X7))|~r1_struct_0(X4,X7)|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_174388, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k5_algstr_0(X2,X5,X6))|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_algstr_0(X2,X5),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173831, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k1_algstr_0(X2,X6,X7),X7))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_174500, plain, (r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,k4_algstr_0(X2,X3),X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173880, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173822, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k5_algstr_0(X2,X7,X6),X7))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173886, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X7))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~r1_struct_0(X4,X7)|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_174884, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X4),k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_170741, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~r2_hidden(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_176021, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r1_struct_0(esk1_0,k4_struct_0(esk3_0))|~l2_struct_0(esk3_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_174673, plain, (k4_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)))=k1_algstr_0(X1,k4_algstr_0(X1,X2),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_179134, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k5_algstr_0(X2,X5,k5_algstr_0(X2,X6,X7)))|~r1_struct_0(X4,k5_algstr_0(X2,X6,k4_struct_0(X2)))|~r1_struct_0(X4,k4_algstr_0(X2,X7))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X6,X7),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_178999, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k5_algstr_0(X2,X5,k5_algstr_0(X2,X6,X7)))|~r1_struct_0(X4,k1_algstr_0(X2,X6,k4_struct_0(X2)))|~r1_struct_0(X4,k4_algstr_0(X2,X7))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X6,X7),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_178972, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_180220, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r2_hidden(k5_algstr_0(X2,X6,X7),a_3_1_zmodul01(X2,X6,X4))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X7))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173825, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k1_algstr_0(X2,k4_algstr_0(X2,X7),X7),X6))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173668, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X4,k4_struct_0(X2)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X4,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173816, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,k1_algstr_0(X2,k4_struct_0(X2),X6)))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_179146, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k5_algstr_0(X2,X5,k4_struct_0(X2)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_174426, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,k4_algstr_0(X2,X5),X6),X4))|~r1_struct_0(X1,k5_algstr_0(X2,X5,X6))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_178989, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k5_algstr_0(X2,X5,k4_algstr_0(X2,X6)))|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_176009, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~r1_struct_0(esk3_0,X1))).#trainneg
cnf(i_0_178717, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r2_hidden(k1_algstr_0(X2,X6,X7),a_3_1_zmodul01(X2,X6,X4))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X7))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_172854, plain, (r1_struct_0(X1,k5_algstr_0(X1,X2,k4_algstr_0(X1,X3)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_172226, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r2_hidden(X3,a_3_1_zmodul01(X2,k4_algstr_0(X2,X4),X1))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173660, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k5_algstr_0(X2,X4,X5))|~r1_struct_0(X1,X5)|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_172152, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k4_algstr_0(X1,X3))=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_174458, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,k4_algstr_0(X2,X4),X5),X3))|~r1_struct_0(X1,k5_algstr_0(X2,X4,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_177042, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1))).#trainneg
cnf(i_0_173849, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k1_algstr_0(X2,X5,X6),X6))|~r1_struct_0(X4,k4_struct_0(X2))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,k1_algstr_0(X2,X5,X6),X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_184571, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k1_algstr_0(X2,X5,X4))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X4),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_184378, negated_conjecture, (r1_struct_0(esk3_0,k1_algstr_0(esk1_0,X1,X2))|~l1_algstr_0(esk3_0)|~r2_hidden(X1,u1_struct_0(esk3_0))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,X2)|~m1_subset_1(k4_algstr_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_184545, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k1_algstr_0(X2,X4,X5))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_182533, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k5_algstr_0(X2,X5,k1_algstr_0(X2,X6,X7)))|~r1_struct_0(X4,k1_algstr_0(X2,X6,k4_struct_0(X2)))|~r1_struct_0(X4,X7)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X6,X7),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_184377, negated_conjecture, (r1_struct_0(esk3_0,k1_algstr_0(esk1_0,X1,X2))|~l2_algstr_0(esk3_0)|~r2_hidden(X1,u1_struct_0(esk3_0))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,X2)|~m1_subset_1(k4_algstr_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_185265, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X4),X3))|~r1_struct_0(X1,k5_algstr_0(X2,X4,k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_182732, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X6),X7))|~r1_struct_0(X4,k1_algstr_0(X2,X6,k4_struct_0(X2)))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X7))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,k4_algstr_0(X2,X6),X7),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_177043, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1))).#trainneg
cnf(i_0_185442, negated_conjecture, (m1_subset_1(k4_algstr_0(esk1_0,k4_struct_0(esk1_0)),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l1_algstr_0(esk3_0))).#trainneg
cnf(i_0_182507, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,k4_algstr_0(X2,X6)))|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_183056, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X4,k4_struct_0(X2)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X4,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_182544, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k5_algstr_0(X2,X5,k1_algstr_0(X2,X6,X7)))|~r1_struct_0(X4,X6)|~r1_struct_0(X4,X7)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X6,X7),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_182730, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X6),X7))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X7))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,k4_algstr_0(X2,X6),X7),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_183295, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k5_algstr_0(X2,X5,k4_struct_0(X2)))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_182510, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X7,k1_algstr_0(X2,X7,X6)))|~r1_struct_0(X4,k5_algstr_0(X2,X5,k4_algstr_0(X2,X6)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_182532, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k5_algstr_0(X2,X5,k1_algstr_0(X2,X6,X7)))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X7))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X6,X7),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_182428, negated_conjecture, (r2_hidden(X1,a_3_1_zmodul01(esk1_0,X2,X3))|X1!=k1_algstr_0(esk1_0,X2,X4)|~r1_struct_0(X3,k5_algstr_0(esk1_0,X4,X5))|~r1_struct_0(X3,X5)|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_182546, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k5_algstr_0(X2,X5,k3_rlvect_1(X2,X6,X7)))|~r1_struct_0(X4,X7)|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k3_rlvect_1(X2,X6,X7),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_184087, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k1_algstr_0(X2,k4_algstr_0(X2,X6),X7))|~r1_struct_0(X4,k1_algstr_0(X2,X6,k4_struct_0(X2)))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X7))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,k4_algstr_0(X2,X6),X7),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_178765, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~r1_struct_0(esk3_0,X1))).#trainneg
cnf(i_0_183747, negated_conjecture, (m1_subset_1(k4_algstr_0(esk3_0,k4_struct_0(esk3_0)),u1_struct_0(esk1_0))|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~m1_zmodul01(esk3_0,esk3_0))).#trainneg
cnf(i_0_182531, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k5_algstr_0(X2,X5,k1_algstr_0(X2,X6,X7)))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X4,X7)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X6,X7),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_184085, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k1_algstr_0(X2,k4_algstr_0(X2,X6),X7))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X7))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,k4_algstr_0(X2,X6),X7),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_175579, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,X3,X2))=k5_algstr_0(X1,k1_algstr_0(X1,X4,X3),X4)|v2_struct_0(X1)|~l1_algstr_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_188840, negated_conjecture, (r2_hidden(X1,a_3_1_zmodul01(esk1_0,X2,X3))|X1!=k1_algstr_0(esk1_0,X2,X4)|~r2_hidden(X4,a_3_1_zmodul01(esk1_0,X5,X3))|~r1_struct_0(X3,X5)|~m1_zmodul01(X3,esk1_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_166866, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~r2_hidden(k4_struct_0(X2),a_3_1_zmodul01(X2,X6,X4))|~r1_struct_0(X4,k5_algstr_0(X2,X6,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_165406, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k1_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_127328, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k4_struct_0(X2),X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X5),X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_189146, negated_conjecture, (r1_struct_0(esk1_0,X1)|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~l1_struct_0(esk1_0)|~r1_struct_0(esk3_0,X1))).#trainneg
cnf(i_0_186989, negated_conjecture, (m1_subset_1(k4_algstr_0(esk1_0,k4_struct_0(esk1_0)),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_166894, negated_conjecture, (r1_tarski(u1_struct_0(esk3_0),a_3_1_zmodul01(esk1_0,X1,esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~m1_subset_1(esk4_2(u1_struct_0(esk3_0),a_3_1_zmodul01(esk1_0,X1,esk3_0)),u1_struct_0(esk3_0))|~m1_subset_1(esk4_2(u1_struct_0(esk3_0),a_3_1_zmodul01(esk1_0,X1,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_145862, plain, (k5_algstr_0(X1,k5_algstr_0(X1,k4_algstr_0(X1,X3),k4_struct_0(X1)),X2)=k4_algstr_0(X1,k1_algstr_0(X1,X3,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_144260, plain, (k4_algstr_0(X1,k5_algstr_0(X1,X2,k4_struct_0(X1)))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,k4_struct_0(X1)),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_183057, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X4),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_126563, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,k5_algstr_0(X2,X5,X6)))|~r1_struct_0(X1,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~r1_struct_0(X1,k4_algstr_0(X2,X6))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X5,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_130073, plain, (k5_algstr_0(X1,k4_algstr_0(X1,X2),k5_algstr_0(X1,X3,X2))=k5_algstr_0(X1,X4,k1_algstr_0(X1,X4,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_191432, negated_conjecture, (a_3_1_zmodul01(esk1_0,X1,esk3_0)=u1_struct_0(esk3_0)|~r1_tarski(a_3_1_zmodul01(esk1_0,X1,esk3_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_193102, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X7),k5_algstr_0(X2,X6,X7)))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_178766, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1))).#trainneg
cnf(i_0_191282, negated_conjecture, (r1_struct_0(esk1_0,k4_algstr_0(esk3_0,k4_struct_0(esk3_0)))|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~m1_zmodul01(esk3_0,esk3_0))).#trainneg
cnf(i_0_191908, plain, (k5_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2),X3)=k4_algstr_0(X1,k1_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_192199, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X6),X5))|~r1_struct_0(X4,k5_algstr_0(X2,X6,k4_struct_0(X2)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X6,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_193217, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X6),k5_algstr_0(X2,X5,X6)))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_192136, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X5),X4))|~r1_struct_0(X1,k5_algstr_0(X2,X5,k4_struct_0(X2)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X5,k4_struct_0(X2)),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_193047, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X6),k5_algstr_0(X2,X5,X6)))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_191267, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k4_struct_0(X2),X5)|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_128977, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),k5_algstr_0(X1,X2,k4_struct_0(X1)))=k5_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_173786, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k5_algstr_0(X2,X5,X6))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,k4_struct_0(X2)))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_182536, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k5_algstr_0(X2,X5,k5_algstr_0(X2,X6,X7)))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X6))|~r1_struct_0(X4,k4_algstr_0(X2,X7))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X6,X7),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_187439, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,k4_struct_0(X2)))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_180859, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1))).#trainneg
cnf(i_0_191281, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk1_0,u2_struct_0(esk3_0))|~l2_struct_0(esk3_0)|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_131397, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k1_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X5),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_195000, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X5),k4_algstr_0(X2,k4_struct_0(X2))))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_138548, negated_conjecture, (r1_struct_0(esk3_0,k5_algstr_0(esk1_0,X1,k1_algstr_0(esk1_0,X1,X2)))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_137137, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k1_algstr_0(X1,X4,k4_algstr_0(X1,X4)),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_131688, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X6,k4_algstr_0(X2,X5)))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X6))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_137540, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X5),X6))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X6))|~r1_struct_0(X1,X5)|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_algstr_0(X2,X5),X6),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_18067, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X3,X2),k1_algstr_0(X1,k4_struct_0(X1),X2))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_24496, plain, (k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),k1_algstr_0(X1,k4_struct_0(X1),X3))=k1_algstr_0(X1,X2,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_197851, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X6,k1_algstr_0(X2,X6,X5))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X3),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_196790, negated_conjecture, (r1_struct_0(esk1_0,k4_struct_0(esk3_0))|~l2_struct_0(esk3_0)|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_196789, negated_conjecture, (r2_hidden(u2_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_struct_0(esk3_0)|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~m1_zmodul01(esk1_0,esk1_0))).#trainneg
cnf(i_0_198407, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_197003, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X5),k1_algstr_0(X2,k4_algstr_0(X2,X6),X6)))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_198220, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k1_algstr_0(X2,X7,k4_algstr_0(X2,X7)),X6))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_199609, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X4),u1_struct_0(X2))|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_32063, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,X4,k4_algstr_0(X1,X4)))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_53176, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k1_algstr_0(X2,X6,k5_algstr_0(X2,X7,X5)))|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X6,X7))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X6,X7),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_173823, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k5_algstr_0(X2,X7,X7),X6))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_194178, negated_conjecture, (r1_struct_0(esk1_0,X1)|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~l1_struct_0(esk1_0)|~r1_struct_0(esk3_0,X1))).#trainneg
cnf(i_0_163002, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k1_algstr_0(X2,X6,k4_struct_0(X2)))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_140911, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),k4_algstr_0(X2,X6)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~r1_struct_0(X1,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_16319, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k5_algstr_0(X1,X2,X3))=k1_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_200266, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_202932, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_202803, negated_conjecture, (r1_struct_0(esk1_0,X1)|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_200252, negated_conjecture, (r2_hidden(k4_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_struct_0(esk3_0)|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~m1_zmodul01(esk1_0,esk1_0))).#trainneg
cnf(i_0_65703, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|~l2_algstr_0(X4)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X3,k4_struct_0(X2)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_30884, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,k4_algstr_0(X2,X4),X3)))|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k4_algstr_0(X2,X4),X3),u1_struct_0(X2))|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_178896, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X7))|~r1_struct_0(X4,X7)|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_165411, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k1_algstr_0(X2,k1_algstr_0(X2,k4_struct_0(X2),X6),X5))|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,k4_struct_0(X2),X6),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_10826, plain, (k1_algstr_0(X1,k5_algstr_0(X1,X2,X5),esk8_4(X3,X1,X2,X4))=k5_algstr_0(X1,X3,X5)|v2_struct_0(X1)|~l2_algstr_0(X1)|~r2_hidden(X3,a_3_1_zmodul01(X1,X2,X4))|~m1_zmodul01(X4,X1)|~m1_subset_1(esk8_4(X3,X1,X2,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_206249, plain, (r1_struct_0(X1,k5_algstr_0(X2,X5,X4))|v2_struct_0(X2)|~r2_hidden(X5,a_3_1_zmodul01(X2,X3,X6))|~r1_struct_0(X1,esk8_4(X5,X2,X3,X6))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_zmodul01(X6,X2)|~m1_subset_1(esk8_4(X5,X2,X3,X6),u1_struct_0(X2))|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_204404, plain, (k5_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),k4_struct_0(X1)),X2)=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_205861, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,k1_algstr_0(X2,X3,X6))|~l1_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,X7))|~r1_struct_0(X4,X7)|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_196771, negated_conjecture, (r1_struct_0(esk1_0,X1)|~l1_algstr_0(esk1_0)|~l1_algstr_0(esk3_0)|~l1_struct_0(esk1_0)|~r1_struct_0(esk3_0,X1))).#trainneg
cnf(i_0_206429, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X6,X5)|~r2_hidden(X6,a_3_1_zmodul01(X2,X3,X7))|~r1_struct_0(X4,esk8_4(X6,X2,X3,X7))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_zmodul01(X7,X2)|~m1_subset_1(esk8_4(X6,X2,X3,X7),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_109038, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~r1_struct_0(X1,X5)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_58069, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k4_algstr_0(X2,X6),X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_23830, plain, (m1_subset_1(k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_65334, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),k5_algstr_0(X1,k4_struct_0(X1),X2))=k5_algstr_0(X1,k4_algstr_0(X1,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_207990, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_180860, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1))).#trainneg
cnf(i_0_207697, negated_conjecture, (r1_struct_0(esk1_0,k4_algstr_0(esk1_0,k4_struct_0(esk1_0)))|~l1_algstr_0(esk1_0)|~l1_algstr_0(esk3_0))).#trainneg
cnf(i_0_207838, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,X6)|~r2_hidden(X5,a_3_1_zmodul01(X2,X3,X4))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_207496, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,X3)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_struct_0(X2))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_207992, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X5),k4_struct_0(X2)))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_207991, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X6,k1_algstr_0(X2,X6,X5)))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_207505, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,k5_algstr_0(X2,X3,X6))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,X7))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~r1_struct_0(X4,X7)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_208004, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),k4_algstr_0(X2,X5)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_208132, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k4_algstr_0(X2,X5),X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X6),X5))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_208000, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),k4_algstr_0(X2,X5)))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_208019, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X5),k4_struct_0(X2)))|~r1_struct_0(X1,k1_algstr_0(X2,X5,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_207999, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X5),k4_struct_0(X2)))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_208241, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k1_algstr_0(X2,X5,X6),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_208538, plain, (k1_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X3),k4_algstr_0(X1,X2))=k5_algstr_0(X1,k4_algstr_0(X1,X3),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X3),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_208013, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X6,k5_algstr_0(X2,X6,X5)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_207996, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,k4_struct_0(X2)),X5))|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_209468, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,X3)|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_210754, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,k4_algstr_0(X2,X5)))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_207530, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,k4_struct_0(X2))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_209127, negated_conjecture, (r1_struct_0(esk1_0,X1)|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~l1_struct_0(esk1_0)|~r1_struct_0(esk3_0,X1))).#trainneg
cnf(i_0_208018, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X1,k1_algstr_0(X2,X5,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_207333, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k4_algstr_0(X2,k4_struct_0(X2))|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_207667, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k1_algstr_0(X2,X5,X3),X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,X7))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~r1_struct_0(X4,X7)|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_177332, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X3),k5_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_56599, plain, (r2_hidden(k1_algstr_0(X1,X2,k5_algstr_0(X1,X3,X4)),a_3_1_zmodul01(X1,X2,X5))|v2_struct_0(X1)|~r1_struct_0(X5,k4_algstr_0(X1,X4))|~r1_struct_0(X5,X3)|~m1_zmodul01(X5,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_64206, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,X4),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),k1_algstr_0(X2,X3,X4)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_60607, plain, (r2_hidden(k1_algstr_0(X1,X2,k4_algstr_0(X1,X3)),a_3_1_zmodul01(X1,X4,X5))|v2_struct_0(X1)|~r1_struct_0(X5,k5_algstr_0(X1,X2,X3))|~r1_struct_0(X5,k4_algstr_0(X1,X4))|~m1_zmodul01(X5,X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_211958, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,k4_struct_0(X2)))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_212352, negated_conjecture, (r1_struct_0(esk1_0,X1)|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_212237, negated_conjecture, (r1_struct_0(esk1_0,k4_algstr_0(esk1_0,k4_struct_0(esk1_0)))|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_51927, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X6,X5),X6))|~r1_struct_0(X1,k5_algstr_0(X2,X3,k4_algstr_0(X2,X5)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_205381, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X5,k4_algstr_0(X2,X6))|~l2_algstr_0(X4)|~r1_struct_0(X4,k5_algstr_0(X2,X6,k4_struct_0(X2)))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X4))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_66860, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X4)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_137042, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_algstr_0(X2,k4_struct_0(X2)),X6))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~r1_struct_0(X1,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_72736, plain, (k5_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),k1_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k4_algstr_0(X1,X3),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_98969, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k5_algstr_0(X2,X3,X5),k4_struct_0(X2))|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_192358, plain, (k4_algstr_0(X1,k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_31639, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X5))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X4,k4_algstr_0(X2,X5)))|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_15246, plain, (r2_hidden(k4_struct_0(X1),a_3_1_zmodul01(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~r1_struct_0(X3,k4_algstr_0(X1,X2))|~m1_zmodul01(X3,X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_26354, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k1_algstr_0(X2,X6,k4_struct_0(X2)))|~r1_struct_0(X4,X5)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_119799, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k4_algstr_0(X2,k5_algstr_0(X2,X6,X5)))|~r1_struct_0(X4,k1_algstr_0(X2,X6,k4_struct_0(X2)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X6,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_66775, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X4,X3))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_40692, plain, (r2_hidden(k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),k4_struct_0(X1)),a_3_1_zmodul01(X1,X2,X4))|v2_struct_0(X1)|~r1_struct_0(X4,X3)|~m1_zmodul01(X4,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_18907, plain, (k5_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X2)),k1_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_154273, negated_conjecture, (m1_subset_1(k4_algstr_0(esk3_0,X1),u1_struct_0(esk1_0))|~l2_algstr_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_78440, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k4_algstr_0(X2,X3),X4))|v2_struct_0(X2)|X1!=k4_struct_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X3,X5))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_170257, plain, (k5_algstr_0(X1,X2,k4_algstr_0(X1,X3))=k1_algstr_0(X1,X3,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_216653, negated_conjecture, (r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_218810, plain, (r1_struct_0(X1,k1_algstr_0(X2,X4,X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r2_hidden(X3,a_3_1_zmodul01(X2,k4_algstr_0(X2,X4),X1))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_218401, negated_conjecture, (k1_algstr_0(esk1_0,X1,k4_algstr_0(esk1_0,X1))=k4_struct_0(esk1_0)|~l2_algstr_0(esk1_0)|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_216694, negated_conjecture, (r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,k4_algstr_0(esk1_0,X1)))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_175979, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|m1_subset_1(esk9_1(u1_struct_0(esk3_0)),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_202687, negated_conjecture, (r1_struct_0(esk1_0,k4_algstr_0(esk3_0,k4_struct_0(esk3_0)))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_zmodul01(esk3_0,esk3_0))).#trainneg
cnf(i_0_216664, plain, (m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|v2_struct_0(X1)|~r1_struct_0(X3,k4_algstr_0(X1,X2))|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_216654, negated_conjecture, (r1_struct_0(esk3_0,k4_struct_0(esk1_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(esk8_4(k4_struct_0(esk1_0),esk1_0,X1,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_216672, plain, (r1_struct_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2))|v2_struct_0(X1)|~r1_struct_0(X3,k4_algstr_0(X1,X2))|~m1_zmodul01(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_zmodul01(X1)|~v5_zmodul01(X1)|~v4_zmodul01(X1)|~v3_zmodul01(X1)|~v2_zmodul01(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_218404, negated_conjecture, (k1_algstr_0(esk1_0,X1,k4_algstr_0(esk1_0,X1))=k4_struct_0(esk3_0)|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_218680, plain, (k1_algstr_0(X1,X2,k5_algstr_0(X1,k4_algstr_0(X1,X2),X3))=k5_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,k4_algstr_0(X1,X2),X3),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_218785, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k4_algstr_0(X1,X3))=k4_algstr_0(X1,k1_algstr_0(X1,X3,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_219716, negated_conjecture, (r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_219083, negated_conjecture, (r1_struct_0(esk3_0,k1_algstr_0(esk1_0,X1,X2))|~l1_algstr_0(esk3_0)|~r2_hidden(X2,u1_struct_0(esk3_0))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_221247, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k4_algstr_0(X1,X3))=k4_algstr_0(X1,k1_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_177012, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|m1_subset_1(esk9_1(u1_struct_0(esk3_0)),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0))).#trainneg
cnf(i_0_207696, negated_conjecture, (r1_struct_0(esk1_0,k4_algstr_0(esk3_0,k4_struct_0(esk3_0)))|~l1_algstr_0(esk1_0)|~l1_algstr_0(esk3_0)|~m1_zmodul01(esk3_0,esk3_0))).#trainneg
cnf(i_0_221033, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,k1_algstr_0(X1,X2,X3)))=k5_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_219046, negated_conjecture, (r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k5_algstr_0(esk1_0,X1,X2))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X2)|~m1_subset_1(k5_algstr_0(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_219044, negated_conjecture, (r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k1_algstr_0(esk1_0,X1,X2))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~m1_subset_1(k1_algstr_0(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_219045, negated_conjecture, (r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k1_algstr_0(esk1_0,k4_algstr_0(esk1_0,X1),X2))|~r1_struct_0(esk3_0,k5_algstr_0(esk1_0,X1,X2))|~m1_subset_1(k1_algstr_0(esk1_0,k4_algstr_0(esk1_0,X1),X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_219082, negated_conjecture, (r1_struct_0(esk3_0,k1_algstr_0(esk1_0,X1,X2))|~l2_algstr_0(esk3_0)|~r2_hidden(X2,u1_struct_0(esk3_0))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_218646, plain, (k5_algstr_0(X1,X2,k5_algstr_0(X1,X2,X3))=k1_algstr_0(X1,X3,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_221569, negated_conjecture, (r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k1_algstr_0(esk1_0,X1,X2))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~m1_subset_1(k4_algstr_0(esk1_0,k1_algstr_0(esk1_0,X1,X2)),u1_struct_0(esk1_0))|~m1_subset_1(k1_algstr_0(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_221571, negated_conjecture, (r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k5_algstr_0(esk1_0,X1,X2))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~r1_struct_0(esk3_0,X2)|~m1_subset_1(k4_algstr_0(esk1_0,k5_algstr_0(esk1_0,X1,X2)),u1_struct_0(esk1_0))|~m1_subset_1(k5_algstr_0(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_221565, negated_conjecture, (r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_189140, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~r1_struct_0(esk3_0,esk4_2(X1,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_212236, negated_conjecture, (r1_struct_0(esk1_0,k4_algstr_0(esk3_0,k4_struct_0(esk3_0)))|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_zmodul01(esk3_0,esk3_0))).#trainneg
cnf(i_0_218536, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X2),X3))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_algstr_0(X1,X2),X3),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_97037, plain, (r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),k1_algstr_0(X2,X3,X4)))|~r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_119562, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_142623, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X5,k4_struct_0(X2)),X4))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_78417, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k4_algstr_0(X2,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_194172, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,esk4_2(X1,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_139425, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X3,X2))=k1_algstr_0(X1,k4_algstr_0(X1,X3),k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_66517, plain, (k5_algstr_0(X1,k4_algstr_0(X1,k1_algstr_0(X1,X3,X2)),k4_algstr_0(X1,X2))=k4_algstr_0(X1,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_226966, plain, (r1_struct_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X2),k4_struct_0(X1)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_226825, plain, (k5_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X3),k4_struct_0(X1)),X2)=k4_algstr_0(X1,k1_algstr_0(X1,X3,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_226913, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_algstr_0(X2,X6),k4_struct_0(X2)))|~r1_struct_0(X4,k1_algstr_0(X2,X6,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X6,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_53226, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k1_algstr_0(X2,X5,k4_algstr_0(X2,X6)))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_struct_0(X2))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_196765, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,esk4_2(X1,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_224661, negated_conjecture, (r1_tarski(k1_zfmisc_1(k1_xboole_0),u1_struct_0(esk1_0))|~l1_algstr_0(esk3_0)|~l2_algstr_0(esk1_0)|~r1_struct_0(esk3_0,k1_xboole_0))).#trainneg
cnf(i_0_64189, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k1_algstr_0(X2,X3,X5),k4_algstr_0(X2,X5))|~l2_algstr_0(X2)|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_13898, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k5_algstr_0(X2,X5,X6))|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_algstr_0(X2,X6)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_134730, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X7,k1_algstr_0(X2,X7,X6)))|~r1_struct_0(X4,k4_algstr_0(X2,k5_algstr_0(X2,X5,X6)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_62020, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,k1_algstr_0(X2,X5,X6)))|~r1_struct_0(X4,k4_algstr_0(X2,k4_algstr_0(X2,X6)))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_208010, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X6),k1_algstr_0(X2,X6,X5)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_76502, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X4),k1_algstr_0(X2,X4,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_51934, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,k4_struct_0(X2),k4_struct_0(X2)),X5))|~r1_struct_0(X1,k5_algstr_0(X2,X3,k4_algstr_0(X2,X5)))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_209121, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,esk4_2(X1,u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_226108, negated_conjecture, (r1_tarski(k1_zfmisc_1(k1_xboole_0),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k1_xboole_0))).#trainneg
cnf(i_0_102044, plain, (r1_struct_0(X1,k1_algstr_0(X2,X3,k1_algstr_0(X2,X4,X5)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k1_algstr_0(X2,X3,X4),X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_164435, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k4_algstr_0(X2,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k1_algstr_0(X2,k4_struct_0(X2),X5))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_146404, plain, (m1_subset_1(k5_algstr_0(X1,k4_algstr_0(X1,X2),k4_struct_0(X1)),u1_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_219041, negated_conjecture, (r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k1_algstr_0(esk1_0,k4_struct_0(esk1_0),X1))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_194058, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X6),k4_algstr_0(X2,k4_struct_0(X2))))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_149982, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~r2_hidden(k5_algstr_0(X2,X6,X5),a_3_1_zmodul01(X2,X6,X4))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_50456, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),k5_algstr_0(X2,X4,X3)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k4_algstr_0(X2,X3))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_228856, negated_conjecture, (r1_tarski(k1_zfmisc_1(k1_xboole_0),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l1_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k1_xboole_0))).#trainneg
cnf(i_0_230638, negated_conjecture, (r1_tarski(k1_zfmisc_1(k1_xboole_0),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k1_xboole_0))).#trainneg
cnf(i_0_78965, plain, (r1_struct_0(X1,k1_algstr_0(X2,k5_algstr_0(X2,X3,k4_struct_0(X2)),k4_algstr_0(X2,X4)))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_13360, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,k4_algstr_0(X1,X4)),X3)=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_101252, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k1_algstr_0(X2,X4,X5),X3))|~r1_struct_0(X1,k4_algstr_0(X2,X5))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,k1_algstr_0(X2,X4,X5),X3),u1_struct_0(X2))|~m1_subset_1(k1_algstr_0(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_202698, negated_conjecture, (r1_struct_0(esk1_0,k4_algstr_0(esk3_0,X1))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_228866, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)=u1_struct_0(esk1_0)|~l1_algstr_0(esk3_0)|~r1_tarski(u1_struct_0(esk1_0),k1_zfmisc_1(k1_xboole_0))|~l2_algstr_0(esk1_0)|~r1_struct_0(esk3_0,k1_xboole_0))).#trainneg
cnf(i_0_66624, plain, (k4_algstr_0(X1,k1_algstr_0(X1,k5_algstr_0(X1,X3,X2),X2))=k4_algstr_0(X1,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_231305, negated_conjecture, (r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,k3_rlvect_1(esk1_0,X1,X2)))|~r1_struct_0(esk3_0,X2)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(k4_algstr_0(esk1_0,k3_rlvect_1(esk1_0,X1,X2)),u1_struct_0(esk1_0))|~m1_subset_1(k3_rlvect_1(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_231617, negated_conjecture, (r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,k5_algstr_0(esk1_0,X1,X2)))|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X2))|~r1_struct_0(esk3_0,X1)|~m1_subset_1(k4_algstr_0(esk1_0,k5_algstr_0(esk1_0,X1,X2)),u1_struct_0(esk1_0))|~m1_subset_1(k5_algstr_0(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_233878, plain, (k4_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X2))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_231291, negated_conjecture, (r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k4_algstr_0(esk1_0,X1))|~m1_subset_1(k1_algstr_0(esk1_0,k4_struct_0(esk1_0),X1),u1_struct_0(esk3_0))|~m1_subset_1(k4_algstr_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_231685, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~r1_struct_0(X4,k5_algstr_0(X2,X5,k5_algstr_0(X2,X6,X7)))|~r1_struct_0(X4,k4_algstr_0(X2,X7))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X6,X7),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_232848, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X3)=k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_235534, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X2)=k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_235540, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X2)=k4_algstr_0(X1,k4_algstr_0(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_238536, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,k4_algstr_0(X2,X6)))|~r1_struct_0(X4,k5_algstr_0(X2,X5,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_238533, plain, (k4_algstr_0(X1,k4_algstr_0(X1,k4_algstr_0(X1,X2)))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_212247, negated_conjecture, (r1_struct_0(esk1_0,k4_algstr_0(esk3_0,X1))|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_230648, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)=u1_struct_0(esk1_0)|~r1_tarski(u1_struct_0(esk1_0),k1_zfmisc_1(k1_xboole_0))|~l2_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k1_xboole_0))).#trainneg
cnf(i_0_237811, plain, (k1_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),k5_algstr_0(X1,X2,X3))=k1_algstr_0(X1,k4_algstr_0(X1,X3),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_235567, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_algstr_0(X2,k4_struct_0(X2)),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_235549, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,X3),X3)=k1_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_239939, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k5_algstr_0(X2,X5,X6),X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_239579, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),k4_struct_0(X1))=k1_algstr_0(X1,k4_algstr_0(X1,X3),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_60888, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k5_algstr_0(X2,X5,X6),X6)|~r1_struct_0(X4,k1_algstr_0(X2,X5,k4_struct_0(X2)))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_208891, plain, (k5_algstr_0(X1,k4_algstr_0(X1,X2),k5_algstr_0(X1,k4_struct_0(X1),X2))=k4_algstr_0(X1,k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_82386, plain, (r1_struct_0(X1,k5_algstr_0(X2,k5_algstr_0(X2,X4,X5),k4_algstr_0(X2,X5)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r2_hidden(k1_algstr_0(X2,X3,X4),a_3_1_zmodul01(X2,X3,X1))|~m1_zmodul01(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_38061, plain, (k1_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X2),k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_178735, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r2_hidden(esk9_1(u1_struct_0(esk3_0)),u1_struct_0(esk1_0))|~l2_algstr_0(esk1_0))).#trainneg
cnf(i_0_231907, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)=u1_struct_0(esk1_0)|~l1_algstr_0(esk1_0)|~l1_algstr_0(esk3_0)|~r1_tarski(u1_struct_0(esk1_0),k1_zfmisc_1(k1_xboole_0))|~r1_struct_0(esk3_0,k1_xboole_0))).#trainneg
cnf(i_0_242051, plain, (k5_algstr_0(X1,k4_algstr_0(X1,X3),k4_struct_0(X1))=k1_algstr_0(X1,k4_algstr_0(X1,X2),k5_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_241539, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),k4_algstr_0(X2,X5)))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_241755, plain, (k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),k4_struct_0(X1))=k1_algstr_0(X1,k4_algstr_0(X1,X3),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_241557, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X6,k5_algstr_0(X2,X6,X5)))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_244380, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),k5_algstr_0(X1,X3,X2))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_244352, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,k4_algstr_0(X1,X4),k5_algstr_0(X1,X4,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_247384, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,k4_algstr_0(X1,X3),k4_algstr_0(X1,k4_struct_0(X1)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_247212, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),k5_algstr_0(X1,X3,X2))=k1_algstr_0(X1,k4_algstr_0(X1,X2),k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_246059, plain, (r1_struct_0(X1,k4_algstr_0(X2,X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_246238, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k5_algstr_0(X1,k4_struct_0(X1),X3))=k4_algstr_0(X1,k1_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_244378, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X4),k5_algstr_0(X1,X4,X2))=k5_algstr_0(X1,k5_algstr_0(X1,X3,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_243606, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk1_0,esk9_1(u1_struct_0(esk3_0)))|~l2_algstr_0(esk1_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_231918, negated_conjecture, (k1_zfmisc_1(k1_xboole_0)=u1_struct_0(esk1_0)|~l1_algstr_0(esk1_0)|~r1_tarski(u1_struct_0(esk1_0),k1_zfmisc_1(k1_xboole_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,k1_xboole_0))).#trainneg
cnf(i_0_251092, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k5_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_249524, plain, (r1_struct_0(X1,k4_algstr_0(X2,X3))|v2_struct_0(X2)|~r2_hidden(k4_struct_0(X2),a_3_1_zmodul01(X2,X4,X1))|~r1_struct_0(X1,k5_algstr_0(X2,X4,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_246255, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,k4_struct_0(X1),X3))=k4_algstr_0(X1,k5_algstr_0(X1,X2,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_247221, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_algstr_0(X2,X6),k5_algstr_0(X2,X6,X5)))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_244383, plain, (r1_struct_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X3),k5_algstr_0(X1,X3,X2)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_246270, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,k4_algstr_0(X1,X3),k4_struct_0(X1)))=k4_algstr_0(X1,k1_algstr_0(X1,X3,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_253385, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_algstr_0(X2,X5),k4_algstr_0(X2,k4_struct_0(X2))))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_251116, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k5_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X4),X4),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_180829, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r2_hidden(esk9_1(u1_struct_0(esk3_0)),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0))).#trainneg
cnf(i_0_207695, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk1_0,u2_struct_0(esk3_0))|~l2_struct_0(esk3_0)|~l1_algstr_0(esk1_0)|~l1_algstr_0(esk3_0))).#trainneg
cnf(i_0_246835, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X3),k5_algstr_0(X1,X3,X2))=k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_245856, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),k4_algstr_0(X1,X3))=k4_algstr_0(X1,k4_algstr_0(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_247456, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X2,X3))=k1_algstr_0(X1,k4_struct_0(X1),k5_algstr_0(X1,k4_struct_0(X1),X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_246264, plain, (k1_algstr_0(X1,k4_algstr_0(X1,X2),k1_algstr_0(X1,X3,X2))=k4_algstr_0(X1,k4_algstr_0(X1,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_245688, plain, (k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),k4_algstr_0(X1,X2))=k4_algstr_0(X1,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_246038, plain, (k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),k5_algstr_0(X1,X3,X4))=k1_algstr_0(X1,X2,k4_algstr_0(X1,X4))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_258866, plain, (r1_struct_0(X1,k4_algstr_0(X1,k4_algstr_0(X1,X2)))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_255862, negated_conjecture, (r1_struct_0(esk1_0,k4_struct_0(esk3_0))|~l2_struct_0(esk3_0)|~l1_algstr_0(esk1_0)|~l1_algstr_0(esk3_0))).#trainneg
cnf(i_0_255858, negated_conjecture, (r2_hidden(u2_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_struct_0(esk3_0)|~l1_algstr_0(esk1_0)|~l1_algstr_0(esk3_0)|~m1_zmodul01(esk1_0,esk1_0))).#trainneg
cnf(i_0_259800, plain, (k1_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X3),k4_algstr_0(X1,X2))=k4_algstr_0(X1,k1_algstr_0(X1,X3,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X3,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_260351, plain, (k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),k4_algstr_0(X1,k4_struct_0(X1)))=k1_algstr_0(X1,X2,k4_algstr_0(X1,X3))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_259846, plain, (k1_algstr_0(X1,k5_algstr_0(X1,X3,X4),k5_algstr_0(X1,X2,X3))=k1_algstr_0(X1,X2,k4_algstr_0(X1,X4))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k5_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_119813, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k5_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k1_algstr_0(X2,X5,X7))|~r1_struct_0(X4,k4_algstr_0(X2,X7))|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X7,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_253409, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_algstr_0(X2,X5),k1_algstr_0(X2,k4_algstr_0(X2,X6),X6)))|~m1_zmodul01(X4,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_255843, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk1_0,esk9_1(u1_struct_0(esk3_0)))|~l1_algstr_0(esk1_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_56629, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k5_algstr_0(X2,X6,k4_struct_0(X2)))|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,X5,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k1_algstr_0(X2,X5,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_260898, plain, (r1_struct_0(X1,k4_algstr_0(X2,k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k4_struct_0(X2),X3))|~r1_struct_0(X1,k4_algstr_0(X2,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,k4_struct_0(X2),X3),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_60619, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X5,k4_struct_0(X2))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,X5,k4_struct_0(X2)))|~r1_struct_0(X4,k4_algstr_0(X2,X3))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_126155, plain, (r1_struct_0(X1,k4_algstr_0(X2,k4_algstr_0(X2,X3)))|v2_struct_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~r1_struct_0(X1,k4_struct_0(X2))|~r1_struct_0(X1,X4)|~m1_zmodul01(X1,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_63446, plain, (k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),k5_algstr_0(X1,X2,k4_struct_0(X1)))=k4_algstr_0(X1,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_56691, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,k4_struct_0(X2),X3))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X4,X5)|~l2_algstr_0(X2)|~r1_struct_0(X3,k4_algstr_0(X2,X5))|~r1_struct_0(X3,X4)|~m1_zmodul01(X3,X2)|~m1_subset_1(k4_struct_0(X2),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_132851, plain, (k5_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X2),k1_algstr_0(X1,X3,X2))=k5_algstr_0(X1,k4_algstr_0(X1,X3),k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,k4_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_260716, negated_conjecture, (r2_hidden(k4_struct_0(esk3_0),u1_struct_0(esk1_0))|~l2_struct_0(esk3_0)|~l1_algstr_0(esk1_0)|~l1_algstr_0(esk3_0)|~m1_zmodul01(esk1_0,esk1_0))).#trainneg
cnf(i_0_99370, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,k5_algstr_0(X2,X3,X5),X6)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,X6))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_42942, plain, (r1_struct_0(X1,k1_algstr_0(X2,k4_struct_0(X2),k1_algstr_0(X2,X3,X4)))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,k4_algstr_0(X2,X5),X4))|~r1_struct_0(X1,X5)|~r1_struct_0(X1,X3)|~m1_zmodul01(X1,X2)|~m1_subset_1(k1_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_137267, plain, (k5_algstr_0(X1,X2,k1_algstr_0(X1,X3,X2))=k5_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_165831, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X3,X4),X4))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X6,k5_algstr_0(X2,X6,X5)))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_227241, plain, (k5_algstr_0(X1,k4_algstr_0(X1,X2),k5_algstr_0(X1,X3,X2))=k1_algstr_0(X1,k4_algstr_0(X1,X3),k4_struct_0(X1))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_267227, plain, (k5_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X2)=k5_algstr_0(X1,k5_algstr_0(X1,X3,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_267233, plain, (k5_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),X2)=k1_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X2))|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_18902, plain, (k5_algstr_0(X1,k1_algstr_0(X1,X2,k4_struct_0(X1)),k1_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_77935, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k4_struct_0(X2),X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k1_algstr_0(X2,k4_algstr_0(X2,X3),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_212235, negated_conjecture, (v1_xboole_0(u1_struct_0(esk3_0))|r1_struct_0(esk1_0,u2_struct_0(esk3_0))|~l2_struct_0(esk3_0)|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_270217, plain, (k5_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),k4_algstr_0(X1,X2)),k5_algstr_0(X1,X3,X2))=k4_algstr_0(X1,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_271108, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,k4_struct_0(X2),k5_algstr_0(X2,X3,X5))|~r1_struct_0(X4,k4_algstr_0(X2,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_269656, plain, (k5_algstr_0(X1,k4_algstr_0(X1,k4_struct_0(X1)),k1_algstr_0(X1,X2,X3))=k5_algstr_0(X1,k4_algstr_0(X1,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_24232, plain, (r1_tarski(a_3_1_zmodul01(X2,X3,X1),X4)|r1_struct_0(X1,k1_algstr_0(X2,esk4_2(a_3_1_zmodul01(X2,X3,X1),X4),X5))|v2_struct_0(X2)|~r1_struct_0(X1,k1_algstr_0(X2,X3,X5))|~m1_zmodul01(X1,X2)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_174251, plain, (k1_algstr_0(X1,k5_algstr_0(X1,X2,X3),k1_algstr_0(X1,k4_algstr_0(X1,X2),X3))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_44229, plain, (k5_algstr_0(X1,k5_algstr_0(X1,k4_struct_0(X1),X3),X2)=k5_algstr_0(X1,k4_algstr_0(X1,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_77055, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,X5)|~l2_algstr_0(X2)|~r1_struct_0(X4,k4_algstr_0(X2,k1_algstr_0(X2,k4_struct_0(X2),X6)))|~r1_struct_0(X4,k5_algstr_0(X2,X6,X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,X6,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_271111, negated_conjecture, (m1_subset_1(u2_struct_0(esk3_0),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0))).#trainneg
cnf(i_0_51533, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k5_algstr_0(X2,X3,k4_algstr_0(X2,k4_algstr_0(X2,X5)))|~l2_algstr_0(X2)|~r1_struct_0(X4,k5_algstr_0(X2,k4_struct_0(X2),X5))|~m1_zmodul01(X4,X2)|~m1_subset_1(k4_algstr_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_15262, plain, (k1_algstr_0(X1,X2,k1_algstr_0(X1,k4_algstr_0(X1,X2),X3))=k1_algstr_0(X1,k4_struct_0(X1),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_115328, plain, (r2_hidden(X1,a_3_1_zmodul01(X2,X3,X4))|v2_struct_0(X2)|X1!=k1_algstr_0(X2,X3,k4_algstr_0(X2,X5))|~r1_struct_0(X4,k5_algstr_0(X2,k4_algstr_0(X2,X6),X5))|~r1_struct_0(X4,X6)|~m1_zmodul01(X4,X2)|~m1_subset_1(k5_algstr_0(X2,k4_algstr_0(X2,X6),X5),u1_struct_0(X2))|~m1_subset_1(k4_algstr_0(X2,X6),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X6,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_63885, plain, (k5_algstr_0(X1,k1_algstr_0(X1,k4_algstr_0(X1,X2),X3),X3)=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_271117, negated_conjecture, (r2_hidden(u2_struct_0(esk3_0),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_zmodul01(esk1_0,esk1_0))).#trainneg
cnf(i_0_185441, negated_conjecture, (m1_subset_1(k4_algstr_0(esk3_0,k4_struct_0(esk3_0)),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l1_algstr_0(esk3_0)|~m1_zmodul01(esk3_0,esk3_0))).#trainneg
cnf(i_0_64228, plain, (k1_algstr_0(X1,k1_algstr_0(X1,X2,X3),k4_algstr_0(X1,X3))=k5_algstr_0(X1,k1_algstr_0(X1,X2,X4),X4)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_276065, plain, (k5_algstr_0(X1,k4_algstr_0(X1,k4_algstr_0(X1,X3)),k1_algstr_0(X1,X2,X3))=k4_algstr_0(X1,X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k1_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_275309, negated_conjecture, (r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~r1_struct_0(esk3_0,X1)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk3_0)))).#trainneg
cnf(i_0_95337, plain, (r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X4,k4_struct_0(X2)),X3))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,X4,X3))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X4,X3),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_63470, plain, (r2_hidden(np__1,k1_numbers)|m1_subset_1(X1,k1_numbers)|~r1_tarski(X1,k1_xboole_0)|~r2_hidden(k1_xboole_0,k1_numbers))).#trainneg
cnf(i_0_186988, negated_conjecture, (m1_subset_1(k4_algstr_0(esk3_0,k4_struct_0(esk3_0)),u1_struct_0(esk1_0))|~l1_algstr_0(esk1_0)|~l2_algstr_0(esk3_0)|~m1_zmodul01(esk3_0,esk3_0))).#trainneg
cnf(i_0_256259, plain, (k4_algstr_0(X1,k1_algstr_0(X1,k4_struct_0(X1),X3))=k5_algstr_0(X1,k5_algstr_0(X1,X2,X3),X2)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_278257, negated_conjecture, (r2_hidden(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~l2_algstr_0(esk3_0)|~m1_subset_1(k4_struct_0(esk1_0),u1_struct_0(esk3_0))|~m1_subset_1(k4_struct_0(esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_279908, plain, (k4_algstr_0(X1,k1_algstr_0(X1,X2,k4_struct_0(X1)))=k5_algstr_0(X1,k5_algstr_0(X1,X3,X2),X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_278075, plain, (k5_algstr_0(X1,k4_algstr_0(X1,k4_algstr_0(X1,X2)),k1_algstr_0(X1,X2,X3))=k4_algstr_0(X1,X3)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_156309, plain, (r1_struct_0(X1,k5_algstr_0(X2,X3,k4_struct_0(X2)))|v2_struct_0(X2)|~l2_algstr_0(X2)|~r1_struct_0(X1,k5_algstr_0(X2,k1_algstr_0(X2,X4,X5),X5))|~r1_struct_0(X1,k5_algstr_0(X2,X3,X4))|~m1_zmodul01(X1,X2)|~m1_subset_1(k5_algstr_0(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~l1_zmodul01(X2)|~v5_zmodul01(X2)|~v4_zmodul01(X2)|~v3_zmodul01(X2)|~v2_zmodul01(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v2_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 44
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 84
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 84
# Processed clauses                    : 3013
# ...of these trivial                  : 28
# ...subsumed                          : 981
# ...remaining for further processing  : 2004
# Other redundant clauses eliminated   : 108
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 176
# Backward-rewritten                   : 40
# Generated clauses                    : 282140
# ...of the previous two non-redundant : 131117
# ...frozen by parental guidance       : 128356
# ...of these subsequently unfrozen    : 0
# ...aggressively subsumed             : 21478
# Contextual simplify-reflections      : 1610
# Paramodulations                      : 281915
# Factorizations                       : 0
# NegExts                              : 0
# Equation resolutions                 : 227
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1784
#    Positive orientable unit clauses  : 32
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 1747
# Current number of unprocessed clauses: 127941
# ...number of literals in the above   : 1831521
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2220
# Proof object given clauses           : 286
# Proof search given clauses           : 2004
# Clause-clause subsumption calls (NU) : 27147821
# Rec. Clause-clause subsumption calls : 374099
# Non-unit clause-clause subsumptions  : 24241
# Unit Clause-clause subsumption calls : 1385
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 4
# BW rewrite match successes           : 2
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 13561610

# -------------------------------------------------
# User time                : 198.775 s
# System time              : 2.035 s
# Total time               : 200.810 s
# Maximum resident set size: 607460 pages
198.77user 2.12system 7:02.36elapsed 47%CPU (0avgtext+0avgdata 607892maxresident)k
24inputs+1848outputs (0major+276612minor)pagefaults 0swaps
