# local_data = 0x1a77fb0
#TFS# Evaluating axioms..
# ENIGMATIC: Connected to the TF server '127.0.0.1:8888'.
# ENIGMATIC: LightGBM model './big03/model.lgb' loaded with 17544 features 'C(l,x[c=3],r,h[b=4096],v,c[b=256],d,a):G:P'
#TF#SERVER: Lgb model filter skipped 1 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=0, query=587, total=588).
#TF#SERVER: Sent 82018 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2352 bytes expected).
#TF#SERVER: Received 2352 bytes from the TF server.
# Preprocessing time       : 0.793 s
# ENIGMATIC: LightGBM parental guidance model #1 './froc_gnn_model2/model.lgb' loaded with 13963 features 'C(l,x[c=3],r,h[b=2048],v,c[b=256],d,a):G:F'
#TF#SERVER: Lgb model filter skipped 11 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=587, query=466, total=1054).
#TF#SERVER: Sent 158666 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4216 bytes expected).
#TF#SERVER: Received 4216 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 55 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=929, query=721, total=1651).
#TF#SERVER: Sent 363104 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6604 bytes expected).
#TF#SERVER: Received 6604 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 172 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1212, query=853, total=2066).
#TF#SERVER: Sent 598169 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8264 bytes expected).
#TF#SERVER: Received 8264 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 104 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1352, query=922, total=2275).
#TF#SERVER: Sent 704914 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9100 bytes expected).
#TF#SERVER: Received 9100 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 94 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1458, query=963, total=2422).
#TF#SERVER: Sent 642754 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9688 bytes expected).
#TF#SERVER: Received 9688 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 170 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=854, total=2391).
#TF#SERVER: Sent 787588 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9564 bytes expected).
#TF#SERVER: Received 9564 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 320 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=709, total=2246).
#TF#SERVER: Sent 695891 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8984 bytes expected).
#TF#SERVER: Received 8984 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 184 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=858, total=2395).
#TF#SERVER: Sent 831441 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9580 bytes expected).
#TF#SERVER: Received 9580 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 162 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=922, total=2459).
#TF#SERVER: Sent 907931 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9836 bytes expected).
#TF#SERVER: Received 9836 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 385 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=682, total=2219).
#TF#SERVER: Sent 800674 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8876 bytes expected).
#TF#SERVER: Received 8876 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 203 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=846, total=2383).
#TF#SERVER: Sent 809153 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9532 bytes expected).
#TF#SERVER: Received 9532 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 138 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=899, total=2436).
#TF#SERVER: Sent 743858 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9744 bytes expected).
#TF#SERVER: Received 9744 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 177 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=931, total=2468).
#TF#SERVER: Sent 891389 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9872 bytes expected).
#TF#SERVER: Received 9872 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 225 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=838, total=2375).
#TF#SERVER: Sent 876268 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9500 bytes expected).
#TF#SERVER: Received 9500 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 195 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=850, total=2387).
#TF#SERVER: Sent 910669 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9548 bytes expected).
#TF#SERVER: Received 9548 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 210 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=816, total=2353).
#TF#SERVER: Sent 911720 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9412 bytes expected).
#TF#SERVER: Received 9412 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 201 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=845, total=2382).
#TF#SERVER: Sent 939616 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9528 bytes expected).
#TF#SERVER: Received 9528 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 254 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=814, total=2351).
#TF#SERVER: Sent 969994 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9404 bytes expected).
#TF#SERVER: Received 9404 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 219 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=844, total=2381).
#TF#SERVER: Sent 819197 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9524 bytes expected).
#TF#SERVER: Received 9524 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 322 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=761, total=2298).
#TF#SERVER: Sent 722354 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9192 bytes expected).
#TF#SERVER: Received 9192 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 193 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=854, total=2391).
#TF#SERVER: Sent 882266 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9564 bytes expected).
#TF#SERVER: Received 9564 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 234 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=802, total=2339).
#TF#SERVER: Sent 891643 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9356 bytes expected).
#TF#SERVER: Received 9356 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 176 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=884, total=2421).
#TF#SERVER: Sent 945372 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9684 bytes expected).
#TF#SERVER: Received 9684 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 204 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=896, total=2433).
#TF#SERVER: Sent 955180 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9732 bytes expected).
#TF#SERVER: Received 9732 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 234 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=800, total=2337).
#TF#SERVER: Sent 777586 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9348 bytes expected).
#TF#SERVER: Received 9348 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 199 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=828, total=2365).
#TF#SERVER: Sent 854224 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9460 bytes expected).
#TF#SERVER: Received 9460 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 111 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=964, total=2501).
#TF#SERVER: Sent 936134 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10004 bytes expected).
#TF#SERVER: Received 10004 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 122 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=928, total=2465).
#TF#SERVER: Sent 810413 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9860 bytes expected).
#TF#SERVER: Received 9860 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 322 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=704, total=2241).
#TF#SERVER: Sent 751578 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8964 bytes expected).
#TF#SERVER: Received 8964 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 241 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=838, total=2375).
#TF#SERVER: Sent 770625 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9500 bytes expected).
#TF#SERVER: Received 9500 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 260 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=772, total=2309).
#TF#SERVER: Sent 746771 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9236 bytes expected).
#TF#SERVER: Received 9236 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 344 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=751, total=2288).
#TF#SERVER: Sent 733765 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9152 bytes expected).
#TF#SERVER: Received 9152 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 360 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=746, total=2283).
#TF#SERVER: Sent 781954 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9132 bytes expected).
#TF#SERVER: Received 9132 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 323 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=736, total=2273).
#TF#SERVER: Sent 728763 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9092 bytes expected).
#TF#SERVER: Received 9092 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 338 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=741, total=2278).
#TF#SERVER: Sent 719141 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9112 bytes expected).
#TF#SERVER: Received 9112 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 393 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=694, total=2231).
#TF#SERVER: Sent 676831 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8924 bytes expected).
#TF#SERVER: Received 8924 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 396 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=682, total=2219).
#TF#SERVER: Sent 619694 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8876 bytes expected).
#TF#SERVER: Received 8876 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 365 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=666, total=2203).
#TF#SERVER: Sent 707182 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8812 bytes expected).
#TF#SERVER: Received 8812 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 180 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=867, total=2404).
#TF#SERVER: Sent 1031501 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9616 bytes expected).
#TF#SERVER: Received 9616 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 165 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=937, total=2474).
#TF#SERVER: Sent 844971 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9896 bytes expected).
#TF#SERVER: Received 9896 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 174 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=932, total=2469).
#TF#SERVER: Sent 761607 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9876 bytes expected).
#TF#SERVER: Received 9876 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 156 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=992, total=2529).
#TF#SERVER: Sent 722821 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (10116 bytes expected).
#TF#SERVER: Received 10116 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 199 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=852, total=2389).
#TF#SERVER: Sent 789036 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9556 bytes expected).
#TF#SERVER: Received 9556 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 245 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=797, total=2334).
#TF#SERVER: Sent 740237 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9336 bytes expected).
#TF#SERVER: Received 9336 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 251 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=810, total=2347).
#TF#SERVER: Sent 773667 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9388 bytes expected).
#TF#SERVER: Received 9388 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 318 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=790, total=2327).
#TF#SERVER: Sent 832812 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9308 bytes expected).
#TF#SERVER: Received 9308 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 340 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=704, total=2241).
#TF#SERVER: Sent 787935 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8964 bytes expected).
#TF#SERVER: Received 8964 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 395 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=676, total=2213).
#TF#SERVER: Sent 778841 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8852 bytes expected).
#TF#SERVER: Received 8852 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 351 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=692, total=2229).
#TF#SERVER: Sent 695871 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8916 bytes expected).
#TF#SERVER: Received 8916 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 462 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=576, total=2113).
#TF#SERVER: Sent 722767 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8452 bytes expected).
#TF#SERVER: Received 8452 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 301 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=724, total=2261).
#TF#SERVER: Sent 784753 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9044 bytes expected).
#TF#SERVER: Received 9044 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 443 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=581, total=2118).
#TF#SERVER: Sent 710591 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8472 bytes expected).
#TF#SERVER: Received 8472 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 152 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=888, total=2425).
#TF#SERVER: Sent 945130 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9700 bytes expected).
#TF#SERVER: Received 9700 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 258 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=766, total=2303).
#TF#SERVER: Sent 683066 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9212 bytes expected).
#TF#SERVER: Received 9212 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 256 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=957, total=2494).
#TF#SERVER: Sent 787673 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9976 bytes expected).
#TF#SERVER: Received 9976 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 343 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=781, total=2318).
#TF#SERVER: Sent 596577 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9272 bytes expected).
#TF#SERVER: Received 9272 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 360 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=825, total=2362).
#TF#SERVER: Sent 808801 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9448 bytes expected).
#TF#SERVER: Received 9448 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 389 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=657, total=2194).
#TF#SERVER: Sent 640028 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8776 bytes expected).
#TF#SERVER: Received 8776 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 344 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=754, total=2291).
#TF#SERVER: Sent 694159 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9164 bytes expected).
#TF#SERVER: Received 9164 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 450 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=647, total=2184).
#TF#SERVER: Sent 1022316 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8736 bytes expected).
#TF#SERVER: Received 8736 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 400 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=661, total=2198).
#TF#SERVER: Sent 1090811 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8792 bytes expected).
#TF#SERVER: Received 8792 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 323 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=768, total=2305).
#TF#SERVER: Sent 1187033 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9220 bytes expected).
#TF#SERVER: Received 9220 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 314 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=759, total=2296).
#TF#SERVER: Sent 1085021 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9184 bytes expected).
#TF#SERVER: Received 9184 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 216 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=830, total=2367).
#TF#SERVER: Sent 1130589 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9468 bytes expected).
#TF#SERVER: Received 9468 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 307 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=803, total=2340).
#TF#SERVER: Sent 1140010 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9360 bytes expected).
#TF#SERVER: Received 9360 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 87 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=938, total=2475).
#TF#SERVER: Sent 1018825 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9900 bytes expected).
#TF#SERVER: Received 9900 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 196 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=831, total=2368).
#TF#SERVER: Sent 1053674 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9472 bytes expected).
#TF#SERVER: Received 9472 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 309 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=741, total=2278).
#TF#SERVER: Sent 1118382 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9112 bytes expected).
#TF#SERVER: Received 9112 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 275 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=777, total=2314).
#TF#SERVER: Sent 1125717 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9256 bytes expected).
#TF#SERVER: Received 9256 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 353 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=685, total=2222).
#TF#SERVER: Sent 1044328 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8888 bytes expected).
#TF#SERVER: Received 8888 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 339 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=736, total=2273).
#TF#SERVER: Sent 1033862 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9092 bytes expected).
#TF#SERVER: Received 9092 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 329 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=832, total=2369).
#TF#SERVER: Sent 1027616 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9476 bytes expected).
#TF#SERVER: Received 9476 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 261 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=769, total=2306).
#TF#SERVER: Sent 1029496 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9224 bytes expected).
#TF#SERVER: Received 9224 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 302 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=792, total=2329).
#TF#SERVER: Sent 1091773 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9316 bytes expected).
#TF#SERVER: Received 9316 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 408 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=622, total=2159).
#TF#SERVER: Sent 972772 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8636 bytes expected).
#TF#SERVER: Received 8636 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 295 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=749, total=2286).
#TF#SERVER: Sent 958962 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9144 bytes expected).
#TF#SERVER: Received 9144 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 412 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=682, total=2219).
#TF#SERVER: Sent 1152881 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8876 bytes expected).
#TF#SERVER: Received 8876 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 282 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=753, total=2290).
#TF#SERVER: Sent 1192421 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9160 bytes expected).
#TF#SERVER: Received 9160 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 359 clauses
#TF#SERVER: Sending data to the TF server (conj=1, context=1536, query=690, total=2227).
#TF#SERVER: Sent 1101784 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8908 bytes expected).
#TF#SERVER: Received 8908 bytes from the TF server.

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:((((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:((m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', redefinition_m2_subset_1)).
fof(fc6_ordinal1, axiom, (~(v1_xboole_0(k4_ordinal1))&v3_ordinal1(k4_ordinal1)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', fc6_ordinal1)).
fof(fc7_card_1, axiom, ~(v1_finset_1(k4_ordinal1)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', fc7_card_1)).
fof(redefinition_k4_nat_1, axiom, ![X1, X2]:(((m1_subset_1(X1,k5_numbers)&v7_ordinal1(X2))=>k4_nat_1(X1,X2)=k3_xcmplx_0(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', redefinition_k4_nat_1)).
fof(commutativity_k4_nat_1, axiom, ![X1, X2]:(((m1_subset_1(X1,k5_numbers)&v7_ordinal1(X2))=>k4_nat_1(X1,X2)=k4_nat_1(X2,X1))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', commutativity_k4_nat_1)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', redefinition_k5_numbers)).
fof(dt_k4_nat_1, axiom, ![X1, X2]:(((m1_subset_1(X1,k5_numbers)&v7_ordinal1(X2))=>m2_subset_1(k4_nat_1(X1,X2),k1_numbers,k5_numbers))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', dt_k4_nat_1)).
fof(cc2_finset_1, axiom, ![X1]:((v1_finset_1(X1)=>![X2]:((m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_finset_1(X2))))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', cc2_finset_1)).
fof(cc3_xreal_0, axiom, ![X1]:((v1_xreal_0(X1)=>v1_xcmplx_0(X1))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', cc3_xreal_0)).
fof(cc2_xreal_0, axiom, ![X1]:((v7_ordinal1(X1)=>v1_xreal_0(X1))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', cc2_xreal_0)).
fof(cc8_ordinal1, axiom, ![X1]:((m1_subset_1(X1,k4_ordinal1)=>v7_ordinal1(X1))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', cc8_ordinal1)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', dt_k5_numbers)).
fof(cc1_finset_1, axiom, ![X1]:((v1_xboole_0(X1)=>v1_finset_1(X1))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', cc1_finset_1)).
fof(t3_arithm, axiom, ![X1]:((v1_xcmplx_0(X1)=>k3_xcmplx_0(np__1,X1)=X1)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', t3_arithm)).
fof(cc5_ordinal1, axiom, ![X1]:((v3_ordinal1(X1)=>![X2]:((m1_subset_1(X2,X1)=>v3_ordinal1(X2))))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', cc5_ordinal1)).
fof(d2_nat_d, axiom, ![X1]:((v7_ordinal1(X1)=>![X2]:((v7_ordinal1(X2)=>![X3]:((v7_ordinal1(X3)=>(X3=k2_nat_d(X1,X2)<=>~((![X4]:((v7_ordinal1(X4)=>~((X1=k2_xcmplx_0(k3_xcmplx_0(X2,X4),X3)&~(r1_xxreal_0(X2,X3))))))&~((X3=k6_numbers&X2=k6_numbers))))))))))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', d2_nat_d)).
fof(cc6_card_1, axiom, ![X1]:(((v3_ordinal1(X1)&v1_finset_1(X1))=>v7_ordinal1(X1))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', cc6_card_1)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', spc1_numerals)).
fof(cc4_card_1, axiom, ![X1]:((m1_subset_1(X1,k4_ordinal1)=>v1_finset_1(X1))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', cc4_card_1)).
fof(t9_euler_2, axiom, ![X1]:((v7_ordinal1(X1)=>![X2]:((v7_ordinal1(X2)=>![X3]:((v7_ordinal1(X3)=>k4_nat_d(k3_xcmplx_0(X1,X2),X3)=k4_nat_d(k4_nat_1(k4_nat_d(X1,X3),k4_nat_d(X2,X3)),X3))))))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', t9_euler_2)).
fof(commutativity_k2_xcmplx_0, axiom, ![X1, X2]:(((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', commutativity_k2_xcmplx_0)).
fof(dt_k4_nat_d, axiom, ![X1, X2]:(((v7_ordinal1(X1)&v7_ordinal1(X2))=>m1_subset_1(k4_nat_d(X1,X2),k5_numbers))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', dt_k4_nat_d)).
fof(redefinition_k2_nat_d, axiom, ![X1, X2]:(((v7_ordinal1(X1)&v7_ordinal1(X2))=>k2_nat_d(X1,X2)=k6_int_1(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', redefinition_k2_nat_d)).
fof(redefinition_k4_nat_d, axiom, ![X1, X2]:(((v7_ordinal1(X1)&v7_ordinal1(X2))=>k4_nat_d(X1,X2)=k6_int_1(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', redefinition_k4_nat_d)).
fof(fc1_nat_1, axiom, ![X1, X2]:(((v7_ordinal1(X1)&v7_ordinal1(X2))=>v7_ordinal1(k2_xcmplx_0(X1,X2)))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', fc1_nat_1)).
fof(redefinition_k1_nat_1, axiom, ![X1, X2]:(((v7_ordinal1(X1)&m1_subset_1(X2,k5_numbers))=>k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', redefinition_k1_nat_1)).
fof(commutativity_k1_nat_1, axiom, ![X1, X2]:(((v7_ordinal1(X1)&m1_subset_1(X2,k5_numbers))=>k1_nat_1(X1,X2)=k1_nat_1(X2,X1))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', commutativity_k1_nat_1)).
fof(fc5_xreal_0, axiom, ![X1, X2]:(((v1_xreal_0(X1)&v1_xreal_0(X2))=>v1_xreal_0(k2_xcmplx_0(X1,X2)))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', fc5_xreal_0)).
fof(fc7_xreal_0, axiom, ![X1, X2]:(((v1_xreal_0(X1)&v1_xreal_0(X2))=>v1_xreal_0(k6_xcmplx_0(X1,X2)))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', fc7_xreal_0)).
fof(dt_k2_nat_d, axiom, ![X1, X2]:(((v7_ordinal1(X1)&v7_ordinal1(X2))=>v7_ordinal1(k2_nat_d(X1,X2)))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', dt_k2_nat_d)).
fof(t6_boole, axiom, ![X1]:((v1_xboole_0(X1)=>X1=k1_xboole_0)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', t6_boole)).
fof(fc8_xreal_0, axiom, ![X1, X2]:(((v1_xreal_0(X1)&v1_xreal_0(X2))=>v1_xreal_0(k7_xcmplx_0(X1,X2)))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', fc8_xreal_0)).
fof(rqRealAdd__k2_xcmplx_0__r1_r1_r2, axiom, k2_xcmplx_0(np__1,np__1)=np__2, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqRealAdd__k2_xcmplx_0__r1_r1_r2)).
fof(cc1_xreal_0, axiom, ![X1]:((m1_subset_1(X1,k1_numbers)=>v1_xreal_0(X1))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', cc1_xreal_0)).
fof(fc3_xreal_0, axiom, ![X1]:((v1_xreal_0(X1)=>(v1_xcmplx_0(k4_xcmplx_0(X1))&v1_xreal_0(k4_xcmplx_0(X1))))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', fc3_xreal_0)).
fof(involutiveness_k4_xcmplx_0, axiom, ![X1]:((v1_xcmplx_0(X1)=>k4_xcmplx_0(k4_xcmplx_0(X1))=X1)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', involutiveness_k4_xcmplx_0)).
fof(rqRealDiff__k6_xcmplx_0__r1_r257_rm256, axiom, k6_xcmplx_0(np__1,np__257)=k4_xcmplx_0(np__256), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqRealDiff__k6_xcmplx_0__r1_r257_rm256)).
fof(rqRealAdd__k2_xcmplx_0__r256_r1_r257, axiom, k2_xcmplx_0(np__256,np__1)=np__257, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqRealAdd__k2_xcmplx_0__r256_r1_r257)).
fof(commutativity_k3_xcmplx_0, axiom, ![X1, X2]:(((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', commutativity_k3_xcmplx_0)).
fof(spc257_numerals, axiom, (((v2_xxreal_0(np__257)&m2_subset_1(np__257,k1_numbers,k5_numbers))&m1_subset_1(np__257,k5_numbers))&m1_subset_1(np__257,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', spc257_numerals)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', redefinition_k6_numbers)).
fof(dt_m2_subset_1, axiom, ![X1, X2]:((((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:((m2_subset_1(X3,X1,X2)=>m1_subset_1(X3,X1))))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', dt_m2_subset_1)).
fof(rqRealDiv__k7_xcmplx_0__r2_r2_r1, axiom, k7_xcmplx_0(np__2,np__2)=np__1, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqRealDiv__k7_xcmplx_0__r2_r2_r1)).
fof(spc256_numerals, axiom, (((v2_xxreal_0(np__256)&m2_subset_1(np__256,k1_numbers,k5_numbers))&m1_subset_1(np__256,k5_numbers))&m1_subset_1(np__256,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', spc256_numerals)).
fof(cc7_ordinal1, axiom, ![X1]:((v1_xboole_0(X1)=>v7_ordinal1(X1))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', cc7_ordinal1)).
fof(t1_arithm, axiom, ![X1]:((v1_xcmplx_0(X1)=>k2_xcmplx_0(X1,k6_numbers)=X1)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', t1_arithm)).
fof(spc0_boole, axiom, v1_xboole_0(np__0), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', spc0_boole)).
fof(fc6_xreal_0, axiom, ![X1, X2]:(((v1_xreal_0(X1)&v1_xreal_0(X2))=>v1_xreal_0(k3_xcmplx_0(X1,X2)))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', fc6_xreal_0)).
fof(rqRealDiff__k6_xcmplx_0__r256_r257_rm1, axiom, k6_xcmplx_0(np__256,np__257)=k4_xcmplx_0(np__1), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqRealDiff__k6_xcmplx_0__r256_r257_rm1)).
fof(rqRealNeg__k4_xcmplx_0__rm256_r256, axiom, k4_xcmplx_0(k4_xcmplx_0(np__256))=np__256, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqRealNeg__k4_xcmplx_0__rm256_r256)).
fof(fc2_nat_1, axiom, ![X1, X2]:(((v7_ordinal1(X1)&v7_ordinal1(X2))=>v7_ordinal1(k3_xcmplx_0(X1,X2)))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', fc2_nat_1)).
fof(rqRealMult__k3_xcmplx_0__r81_r81_r6561, axiom, k3_xcmplx_0(np__81,np__81)=np__6561, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqRealMult__k3_xcmplx_0__r81_r81_r6561)).
fof(spc136_numerals, axiom, (((v2_xxreal_0(np__136)&m2_subset_1(np__136,k1_numbers,k5_numbers))&m1_subset_1(np__136,k5_numbers))&m1_subset_1(np__136,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', spc136_numerals)).
fof(rqRealMult__k3_xcmplx_0__r257_r0_r0, axiom, k3_xcmplx_0(np__257,np__0)=np__0, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqRealMult__k3_xcmplx_0__r257_r0_r0)).
fof(spc0_numerals, axiom, ((m2_subset_1(np__0,k1_numbers,k5_numbers)&m1_subset_1(np__0,k5_numbers))&m1_subset_1(np__0,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', spc0_numerals)).
fof(rqRealDiff__k6_xcmplx_0__r6561_r6425_r136, axiom, k6_xcmplx_0(np__6561,np__6425)=np__136, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqRealDiff__k6_xcmplx_0__r6561_r6425_r136)).
fof(spc81_numerals, axiom, (((v2_xxreal_0(np__81)&m2_subset_1(np__81,k1_numbers,k5_numbers))&m1_subset_1(np__81,k5_numbers))&m1_subset_1(np__81,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', spc81_numerals)).
fof(rqRealDiff__k6_xcmplx_0__r6561_r136_r6425, axiom, k6_xcmplx_0(np__6561,np__136)=np__6425, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqRealDiff__k6_xcmplx_0__r6561_r136_r6425)).
fof(spc7_arithm, axiom, ![X1, X2, X3]:((((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))&v1_xcmplx_0(X3))=>k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', spc7_arithm)).
fof(spc2_arithm, axiom, ![X1]:((v1_xcmplx_0(X1)=>k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', spc2_arithm)).
fof(rqLessOrEqual__r1_xxreal_0__r257_r1, axiom, ~(r1_xxreal_0(np__257,np__1)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqLessOrEqual__r1_xxreal_0__r257_r1)).
fof(spc25_numerals, axiom, (((v2_xxreal_0(np__25)&m2_subset_1(np__25,k1_numbers,k5_numbers))&m1_subset_1(np__25,k5_numbers))&m1_subset_1(np__25,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', spc25_numerals)).
fof(rqRealMult__k3_xcmplx_0__r257_r25_r6425, axiom, k3_xcmplx_0(np__257,np__25)=np__6425, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqRealMult__k3_xcmplx_0__r257_r25_r6425)).
fof(rqRealNeg__k4_xcmplx_0__rm1_r1, axiom, k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqRealNeg__k4_xcmplx_0__rm1_r1)).
fof(dt_k4_xcmplx_0, axiom, ![X1]:((v1_xcmplx_0(X1)=>v1_xcmplx_0(k4_xcmplx_0(X1)))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', dt_k4_xcmplx_0)).
fof(rqLessOrEqual__r1_xxreal_0__r257_r136, axiom, ~(r1_xxreal_0(np__257,np__136)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqLessOrEqual__r1_xxreal_0__r257_r136)).
fof(l87_pepin, axiom, k4_nat_d(k13_newton(np__3,np__4),k4_pepin(np__3))=np__81, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', l87_pepin)).
fof(t53_pepin, axiom, k4_pepin(np__3)=np__257, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', t53_pepin)).
fof(rqLessOrEqual__r1_xxreal_0__r257_r81, axiom, ~(r1_xxreal_0(np__257,np__81)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqLessOrEqual__r1_xxreal_0__r257_r81)).
fof(rqRealNeg__k4_xcmplx_0__rm6561_r6561, axiom, k4_xcmplx_0(k4_xcmplx_0(np__6561))=np__6561, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqRealNeg__k4_xcmplx_0__rm6561_r6561)).
fof(rqRealMult__k3_xcmplx_0__r1_r81_r81, axiom, k3_xcmplx_0(np__1,np__81)=np__81, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqRealMult__k3_xcmplx_0__r1_r81_r81)).
fof(rqRealAdd__k2_xcmplx_0__r0_r81_r81, axiom, k2_xcmplx_0(np__0,np__81)=np__81, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqRealAdd__k2_xcmplx_0__r0_r81_r81)).
fof(rqRealAdd__k2_xcmplx_0__r6425_r136_r6561, axiom, k2_xcmplx_0(np__6425,np__136)=np__6561, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqRealAdd__k2_xcmplx_0__r6425_r136_r6561)).
fof(spc6561_numerals, axiom, (((v2_xxreal_0(np__6561)&m2_subset_1(np__6561,k1_numbers,k5_numbers))&m1_subset_1(np__6561,k5_numbers))&m1_subset_1(np__6561,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', spc6561_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r2_r1_r3, axiom, k2_xcmplx_0(np__2,np__1)=np__3, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqRealAdd__k2_xcmplx_0__r2_r1_r3)).
fof(l88_pepin, conjecture, k4_nat_d(k13_newton(np__3,np__8),k4_pepin(np__3))=np__136, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', l88_pepin)).
fof(redefinition_k13_newton, axiom, ![X1, X2]:(((m1_subset_1(X1,k5_numbers)&m1_subset_1(X2,k5_numbers))=>k13_newton(X1,X2)=k1_newton(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', redefinition_k13_newton)).
fof(rqRealAdd__k2_xcmplx_0__r2_r2_r4, axiom, k2_xcmplx_0(np__2,np__2)=np__4, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqRealAdd__k2_xcmplx_0__r2_r2_r4)).
fof(spc3_numerals, axiom, (((v2_xxreal_0(np__3)&m2_subset_1(np__3,k1_numbers,k5_numbers))&m1_subset_1(np__3,k5_numbers))&m1_subset_1(np__3,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', spc3_numerals)).
fof(spc4_numerals, axiom, (((v2_xxreal_0(np__4)&m2_subset_1(np__4,k1_numbers,k5_numbers))&m1_subset_1(np__4,k5_numbers))&m1_subset_1(np__4,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', spc4_numerals)).
fof(t8_newton, axiom, ![X1]:((v7_ordinal1(X1)=>![X2]:((v7_ordinal1(X2)=>![X3]:((v1_xcmplx_0(X3)=>k1_newton(X3,k2_xcmplx_0(X1,X2))=k3_xcmplx_0(k1_newton(X3,X1),k1_newton(X3,X2)))))))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', t8_newton)).
fof(spc8_numerals, axiom, (((v2_xxreal_0(np__8)&m2_subset_1(np__8,k1_numbers,k5_numbers))&m1_subset_1(np__8,k5_numbers))&m1_subset_1(np__8,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', spc8_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r4_r4_r8, axiom, k2_xcmplx_0(np__4,np__4)=np__8, file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', rqRealAdd__k2_xcmplx_0__r4_r4_r8)).
fof(fc4_newton, axiom, ![X1, X2]:(((v7_ordinal1(X1)&v7_ordinal1(X2))=>v7_ordinal1(k1_newton(X1,X2)))), file('/scratch/mptp/hardpred/preds__0.005/l88_pepin', fc4_newton)).
fof(c_0_84, plain, ![X1, X2]:((((~v1_xboole_0(X1)&~v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:((m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))))), inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])).
fof(c_0_85, plain, (~v1_xboole_0(k4_ordinal1)&v3_ordinal1(k4_ordinal1)), inference(fof_simplification,[status(thm)],[fc6_ordinal1])).
fof(c_0_86, plain, ~v1_finset_1(k4_ordinal1), inference(fof_simplification,[status(thm)],[fc7_card_1])).
fof(c_0_87, plain, ![X166, X167]:((~m1_subset_1(X166,k5_numbers)|~v7_ordinal1(X167)|k4_nat_1(X166,X167)=k3_xcmplx_0(X166,X167))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_nat_1])])).
fof(c_0_88, plain, ![X199, X200]:((~m1_subset_1(X199,k5_numbers)|~v7_ordinal1(X200)|k4_nat_1(X199,X200)=k4_nat_1(X200,X199))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k4_nat_1])])).
fof(c_0_89, plain, ![X253, X254, X255]:(((~m2_subset_1(X255,X253,X254)|m1_subset_1(X255,X254)|(v1_xboole_0(X253)|v1_xboole_0(X254)|~m1_subset_1(X254,k1_zfmisc_1(X253))))&(~m1_subset_1(X255,X254)|m2_subset_1(X255,X253,X254)|(v1_xboole_0(X253)|v1_xboole_0(X254)|~m1_subset_1(X254,k1_zfmisc_1(X253)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_84])])])])).
cnf(c_0_90, plain, (~v1_xboole_0(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_91, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
fof(c_0_92, plain, ![X251, X252]:((~m1_subset_1(X251,k5_numbers)|~v7_ordinal1(X252)|m2_subset_1(k4_nat_1(X251,X252),k1_numbers,k5_numbers))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_nat_1])])).
fof(c_0_93, plain, ![X296, X297]:((~v1_finset_1(X296)|(~m1_subset_1(X297,k1_zfmisc_1(X296))|v1_finset_1(X297)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_finset_1])])])).
cnf(c_0_94, plain, (~v1_finset_1(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_86])).
fof(c_0_95, plain, ![X159]:((~v1_xreal_0(X159)|v1_xcmplx_0(X159))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_xreal_0])])).
fof(c_0_96, plain, ![X160]:((~v7_ordinal1(X160)|v1_xreal_0(X160))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_xreal_0])])).
cnf(c_0_97, plain, (k4_nat_1(X1,X2)=k3_xcmplx_0(X1,X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_98, plain, (k4_nat_1(X1,X2)=k4_nat_1(X2,X1)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_88])).
fof(c_0_99, plain, ![X158]:((~m1_subset_1(X158,k4_ordinal1)|v7_ordinal1(X158))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_ordinal1])])).
cnf(c_0_100, plain, (m1_subset_1(X1,X3)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_101, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
cnf(c_0_102, plain, (~v1_xboole_0(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_103, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_91])).
cnf(c_0_104, plain, (m2_subset_1(k4_nat_1(X1,X2),k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_105, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_106, plain, (~v1_finset_1(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_94])).
fof(c_0_107, plain, ![X293]:((~v1_xboole_0(X293)|v1_finset_1(X293))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_finset_1])])).
fof(c_0_108, plain, ![X168]:((~v1_xcmplx_0(X168)|k3_xcmplx_0(np__1,X168)=X168)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_arithm])])).
cnf(c_0_109, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_110, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_111, plain, (k4_nat_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_97])).
cnf(c_0_112, plain, (k4_nat_1(X1,X2)=k4_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_113, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_114, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_100])).
cnf(c_0_115, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_101])).
cnf(c_0_116, plain, (~v1_xboole_0(k5_numbers)), inference(rw,[status(thm)],[c_0_102, c_0_103])).
cnf(c_0_117, plain, (m2_subset_1(k4_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_104])).
cnf(c_0_118, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_119, plain, (~v1_finset_1(k5_numbers)), inference(rw,[status(thm)],[c_0_106, c_0_103])).
cnf(c_0_120, plain, (v1_finset_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_107])).
fof(c_0_121, plain, ![X234, X235]:((~v3_ordinal1(X234)|(~m1_subset_1(X235,X234)|v3_ordinal1(X235)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_ordinal1])])])).
cnf(c_0_122, plain, (v3_ordinal1(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_85])).
fof(c_0_123, plain, ![X1]:((v7_ordinal1(X1)=>![X2]:((v7_ordinal1(X2)=>![X3]:((v7_ordinal1(X3)=>(X3=k2_nat_d(X1,X2)<=>~((![X4]:((v7_ordinal1(X4)=>~((X1=k2_xcmplx_0(k3_xcmplx_0(X2,X4),X3)&~r1_xxreal_0(X2,X3)))))&~((X3=k6_numbers&X2=k6_numbers))))))))))), inference(fof_simplification,[status(thm)],[d2_nat_d])).
cnf(c_0_124, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_108])).
cnf(c_0_125, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_109])).
cnf(c_0_126, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_110])).
cnf(c_0_127, plain, (k4_nat_1(X2,X1)=k3_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112])])).
cnf(c_0_128, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_113])).
cnf(c_0_129, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115]), c_0_116])])).
cnf(c_0_130, plain, (m2_subset_1(k3_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_111])])).
cnf(c_0_131, plain, (~v1_finset_1(k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_115]), c_0_119])])).
cnf(c_0_132, plain, (v1_finset_1(X1)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_120])).
fof(c_0_133, plain, ![X180]:((~v3_ordinal1(X180)|~v1_finset_1(X180)|v7_ordinal1(X180))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc6_card_1])])).
cnf(c_0_134, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_135, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_136, plain, (v3_ordinal1(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_122])).
fof(c_0_137, plain, ![X188]:((~m1_subset_1(X188,k4_ordinal1)|v1_finset_1(X188))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_card_1])])).
fof(c_0_138, plain, ![X241, X242, X243, X245]:(((((X243=k6_numbers|v7_ordinal1(esk2_3(X241,X242,X243))|X243!=k2_nat_d(X241,X242)|~v7_ordinal1(X243)|~v7_ordinal1(X242)|~v7_ordinal1(X241))&(X242=k6_numbers|v7_ordinal1(esk2_3(X241,X242,X243))|X243!=k2_nat_d(X241,X242)|~v7_ordinal1(X243)|~v7_ordinal1(X242)|~v7_ordinal1(X241)))&(((X243=k6_numbers|X241=k2_xcmplx_0(k3_xcmplx_0(X242,esk2_3(X241,X242,X243)),X243)|X243!=k2_nat_d(X241,X242)|~v7_ordinal1(X243)|~v7_ordinal1(X242)|~v7_ordinal1(X241))&(X242=k6_numbers|X241=k2_xcmplx_0(k3_xcmplx_0(X242,esk2_3(X241,X242,X243)),X243)|X243!=k2_nat_d(X241,X242)|~v7_ordinal1(X243)|~v7_ordinal1(X242)|~v7_ordinal1(X241)))&((X243=k6_numbers|~r1_xxreal_0(X242,X243)|X243!=k2_nat_d(X241,X242)|~v7_ordinal1(X243)|~v7_ordinal1(X242)|~v7_ordinal1(X241))&(X242=k6_numbers|~r1_xxreal_0(X242,X243)|X243!=k2_nat_d(X241,X242)|~v7_ordinal1(X243)|~v7_ordinal1(X242)|~v7_ordinal1(X241)))))&((~v7_ordinal1(X245)|(X241!=k2_xcmplx_0(k3_xcmplx_0(X242,X245),X243)|r1_xxreal_0(X242,X243))|X243=k2_nat_d(X241,X242)|~v7_ordinal1(X243)|~v7_ordinal1(X242)|~v7_ordinal1(X241))&(X243!=k6_numbers|X242!=k6_numbers|X243=k2_nat_d(X241,X242)|~v7_ordinal1(X243)|~v7_ordinal1(X242)|~v7_ordinal1(X241))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_123])])])])])).
cnf(c_0_139, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_124])).
cnf(c_0_140, plain, (v1_xcmplx_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126])])).
cnf(c_0_141, plain, (k3_xcmplx_0(X2,X1)=k3_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_127])])).
cnf(c_0_142, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_128, c_0_103])])).
cnf(c_0_143, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k3_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130])])).
cnf(c_0_144, plain, (~v1_xboole_0(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_132])])).
cnf(c_0_145, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1)), inference(split_conjunct,[status(thm)],[c_0_133])).
cnf(c_0_146, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_134])).
cnf(c_0_147, plain, (m1_subset_1(np__1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_135])).
cnf(c_0_148, plain, (v3_ordinal1(k5_numbers)), inference(rw,[status(thm)],[c_0_136, c_0_103])).
cnf(c_0_149, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_137])).
fof(c_0_150, plain, ![X182, X183, X184]:((~v7_ordinal1(X182)|(~v7_ordinal1(X183)|(~v7_ordinal1(X184)|k4_nat_d(k3_xcmplx_0(X182,X183),X184)=k4_nat_d(k4_nat_1(k4_nat_d(X182,X184),k4_nat_d(X183,X184)),X184))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_euler_2])])])).
cnf(c_0_151, plain, (r1_xxreal_0(X3,X4)|X4=k2_nat_d(X2,X3)|~v7_ordinal1(X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(X3,X1),X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_138])).
cnf(c_0_152, plain, (k3_xcmplx_0(np__1,X1)=X1|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140])])).
cnf(c_0_153, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_141, c_0_142]), c_0_142])).
cnf(c_0_154, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)), inference(sr,[status(thm)],[c_0_143, c_0_144])).
cnf(c_0_155, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1)), inference(evalgc,[status(thm)],[c_0_145])).
cnf(c_0_156, plain, (v3_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_147]), c_0_148])])])).
cnf(c_0_157, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_149])).
fof(c_0_158, plain, ![X177, X178]:((~v1_xcmplx_0(X177)|~v1_xcmplx_0(X178)|k2_xcmplx_0(X177,X178)=k2_xcmplx_0(X178,X177))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k2_xcmplx_0])])).
cnf(c_0_159, plain, (k4_nat_d(k3_xcmplx_0(X1,X2),X3)=k4_nat_d(k4_nat_1(k4_nat_d(X1,X3),k4_nat_d(X2,X3)),X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)), inference(split_conjunct,[status(thm)],[c_0_150])).
fof(c_0_160, plain, ![X195, X196]:((~v7_ordinal1(X195)|~v7_ordinal1(X196)|m1_subset_1(k4_nat_d(X195,X196),k5_numbers))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_nat_d])])).
fof(c_0_161, plain, ![X189, X190]:((~v7_ordinal1(X189)|~v7_ordinal1(X190)|k2_nat_d(X189,X190)=k6_int_1(X189,X190))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_nat_d])])).
fof(c_0_162, plain, ![X214, X215]:((~v7_ordinal1(X214)|~v7_ordinal1(X215)|k4_nat_d(X214,X215)=k6_int_1(X214,X215))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_nat_d])])).
cnf(c_0_163, plain, (X4=k2_nat_d(X2,X3)|r1_xxreal_0(X3,X4)|X2!=k2_xcmplx_0(k3_xcmplx_0(X3,X1),X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_151])).
cnf(c_0_164, plain, (k3_xcmplx_0(X1,np__1)=X1|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_153]), c_0_147])])])).
cnf(c_0_165, plain, (m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_152]), c_0_147])])])).
cnf(c_0_166, plain, (v7_ordinal1(np__1)|~v1_finset_1(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_156])])).
cnf(c_0_167, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_157, c_0_103])).
fof(c_0_168, plain, ![X217, X218]:((~v7_ordinal1(X217)|~v7_ordinal1(X218)|v7_ordinal1(k2_xcmplx_0(X217,X218)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_nat_1])])).
cnf(c_0_169, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_158])).
fof(c_0_170, plain, ![X197, X198]:((~v7_ordinal1(X197)|~m1_subset_1(X198,k5_numbers)|k1_nat_1(X197,X198)=k2_xcmplx_0(X197,X198))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_nat_1])])).
fof(c_0_171, plain, ![X221, X222]:((~v7_ordinal1(X221)|~m1_subset_1(X222,k5_numbers)|k1_nat_1(X221,X222)=k1_nat_1(X222,X221))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k1_nat_1])])).
cnf(c_0_172, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X3),k4_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(X1,X2),X3)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_159])).
cnf(c_0_173, plain, (m1_subset_1(k4_nat_d(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_160])).
cnf(c_0_174, plain, (k2_nat_d(X1,X2)=k6_int_1(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_161])).
cnf(c_0_175, plain, (k4_nat_d(X1,X2)=k6_int_1(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_162])).
cnf(c_0_176, plain, (k2_nat_d(k2_xcmplx_0(k3_xcmplx_0(X2,X3),X1),X2)=X1|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(k3_xcmplx_0(X2,X3),X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_163])])).
cnf(c_0_177, plain, (k3_xcmplx_0(X1,np__1)=X1|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_164, c_0_165])).
cnf(c_0_178, plain, (v7_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_167]), c_0_147])])])).
cnf(c_0_179, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_168])).
cnf(c_0_180, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_169])).
cnf(c_0_181, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_170])).
cnf(c_0_182, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_171])).
fof(c_0_183, plain, ![X257, X258]:((~v1_xreal_0(X257)|~v1_xreal_0(X258)|v1_xreal_0(k2_xcmplx_0(X257,X258)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc5_xreal_0])])).
fof(c_0_184, plain, ![X201, X202]:((~v1_xreal_0(X201)|~v1_xreal_0(X202)|v1_xreal_0(k6_xcmplx_0(X201,X202)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc7_xreal_0])])).
cnf(c_0_185, plain, (k4_nat_d(k3_xcmplx_0(k4_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(k4_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X1,X2),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_111])])).
cnf(c_0_186, plain, (m1_subset_1(k4_nat_d(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_173])).
cnf(c_0_187, plain, (k6_int_1(X1,X2)=k2_nat_d(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_174])).
cnf(c_0_188, plain, (k6_int_1(X1,X2)=k4_nat_d(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_175])).
fof(c_0_189, plain, ![X230, X231]:((~v7_ordinal1(X230)|~v7_ordinal1(X231)|v7_ordinal1(k2_nat_d(X230,X231)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_nat_d])])).
cnf(c_0_190, plain, (k2_nat_d(k2_xcmplx_0(X1,X2),X1)=X2|r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_177]), c_0_178])])])).
cnf(c_0_191, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_179])).
cnf(c_0_192, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_140])])).
cnf(c_0_193, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_181])).
cnf(c_0_194, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_182])).
fof(c_0_195, plain, ![X213]:((~v1_xboole_0(X213)|X213=k1_xboole_0)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_196, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X3,X2),k4_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(k4_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X3,X2),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_112])])).
fof(c_0_197, plain, ![X232, X233]:((~v1_xreal_0(X232)|~v1_xreal_0(X233)|v1_xreal_0(k7_xcmplx_0(X232,X233)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc8_xreal_0])])).
cnf(c_0_198, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_183])).
cnf(c_0_199, plain, (k2_xcmplx_0(np__1,np__1)=np__2), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r1_r2])).
fof(c_0_200, plain, ![X162]:((~m1_subset_1(X162,k1_numbers)|v1_xreal_0(X162))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_xreal_0])])).
fof(c_0_201, plain, ![X179]:(((v1_xcmplx_0(k4_xcmplx_0(X179))|~v1_xreal_0(X179))&(v1_xreal_0(k4_xcmplx_0(X179))|~v1_xreal_0(X179)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc3_xreal_0])])])).
fof(c_0_202, plain, ![X256]:((~v1_xcmplx_0(X256)|k4_xcmplx_0(k4_xcmplx_0(X256))=X256)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k4_xcmplx_0])])).
cnf(c_0_203, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_184])).
cnf(c_0_204, plain, (k6_xcmplx_0(np__1,np__257)=k4_xcmplx_0(np__256)), inference(split_conjunct,[status(thm)],[rqRealDiff__k6_xcmplx_0__r1_r257_rm256])).
cnf(c_0_205, plain, (k2_xcmplx_0(np__256,np__1)=np__257), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r256_r1_r257])).
cnf(c_0_206, plain, (k4_nat_d(k3_xcmplx_0(k4_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(k4_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_185, c_0_186])).
cnf(c_0_207, plain, (k4_nat_d(X1,X2)=k2_nat_d(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_188])])).
cnf(c_0_208, plain, (v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_189])).
cnf(c_0_209, plain, (k2_nat_d(k2_xcmplx_0(X1,X2),X1)=X2|r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_190, c_0_191])).
cnf(c_0_210, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_140])])).
fof(c_0_211, plain, ![X175, X176]:((~v1_xcmplx_0(X175)|~v1_xcmplx_0(X176)|k3_xcmplx_0(X175,X176)=k3_xcmplx_0(X176,X175))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k3_xcmplx_0])])).
cnf(c_0_212, plain, (k1_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_194])])).
cnf(c_0_213, plain, (m1_subset_1(np__257,k5_numbers)), inference(split_conjunct,[status(thm)],[spc257_numerals])).
cnf(c_0_214, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_195])).
cnf(c_0_215, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
cnf(c_0_216, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(k4_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)), inference(csr,[status(thm)],[c_0_196, c_0_186])).
fof(c_0_217, plain, ![X1, X2]:((((~v1_xboole_0(X1)&~v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:((m2_subset_1(X3,X1,X2)=>m1_subset_1(X3,X1))))), inference(fof_simplification,[status(thm)],[dt_m2_subset_1])).
cnf(c_0_218, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_197])).
cnf(c_0_219, plain, (k7_xcmplx_0(np__2,np__2)=np__1), inference(split_conjunct,[status(thm)],[rqRealDiv__k7_xcmplx_0__r2_r2_r1])).
cnf(c_0_220, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_198])).
cnf(c_0_221, plain, (k2_xcmplx_0(np__1,np__1)=np__2), inference(evalgc,[status(thm)],[c_0_199])).
cnf(c_0_222, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_200])).
cnf(c_0_223, plain, (m1_subset_1(np__1,k1_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_224, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_201])).
cnf(c_0_225, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_202])).
cnf(c_0_226, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_203])).
cnf(c_0_227, plain, (k6_xcmplx_0(np__1,np__257)=k4_xcmplx_0(np__256)), inference(evalgc,[status(thm)],[c_0_204])).
cnf(c_0_228, plain, (k2_xcmplx_0(np__256,np__1)=np__257), inference(evalgc,[status(thm)],[c_0_205])).
cnf(c_0_229, plain, (m1_subset_1(np__256,k1_numbers)), inference(split_conjunct,[status(thm)],[spc256_numerals])).
cnf(c_0_230, plain, (k4_nat_d(k3_xcmplx_0(k4_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(k2_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_207])])).
cnf(c_0_231, plain, (v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_208])).
cnf(c_0_232, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(k4_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_207])])).
cnf(c_0_233, plain, (k2_nat_d(k2_xcmplx_0(X2,X1),X1)=X2|r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_210])])).
cnf(c_0_234, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_211])).
cnf(c_0_235, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_212])])).
cnf(c_0_236, plain, (m1_subset_1(np__257,k5_numbers)), inference(evalgc,[status(thm)],[c_0_213])).
fof(c_0_237, plain, ![X187]:((~v1_xboole_0(X187)|v7_ordinal1(X187))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc7_ordinal1])])).
fof(c_0_238, plain, ![X191]:((~v1_xcmplx_0(X191)|k2_xcmplx_0(X191,k6_numbers)=X191)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_arithm])])).
cnf(c_0_239, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_214])).
cnf(c_0_240, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_215])).
cnf(c_0_241, plain, (v1_xboole_0(np__0)), inference(split_conjunct,[status(thm)],[spc0_boole])).
cnf(c_0_242, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(k2_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_207])])).
fof(c_0_243, plain, ![X290, X291, X292]:((v1_xboole_0(X290)|v1_xboole_0(X291)|~m1_subset_1(X291,k1_zfmisc_1(X290))|(~m2_subset_1(X292,X290,X291)|m1_subset_1(X292,X290)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_217])])])).
fof(c_0_244, plain, ![X249, X250]:((~v1_xreal_0(X249)|~v1_xreal_0(X250)|v1_xreal_0(k3_xcmplx_0(X249,X250)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc6_xreal_0])])).
cnf(c_0_245, plain, (k6_xcmplx_0(np__256,np__257)=k4_xcmplx_0(np__1)), inference(split_conjunct,[status(thm)],[rqRealDiff__k6_xcmplx_0__r256_r257_rm1])).
cnf(c_0_246, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_218])).
cnf(c_0_247, plain, (k7_xcmplx_0(np__2,np__2)=np__1), inference(evalgc,[status(thm)],[c_0_219])).
cnf(c_0_248, plain, (v1_xreal_0(np__2)|~v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_221])])).
cnf(c_0_249, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_222])).
cnf(c_0_250, plain, (m1_subset_1(np__1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_223])).
cnf(c_0_251, plain, (k4_xcmplx_0(k4_xcmplx_0(np__256))=np__256), inference(split_conjunct,[status(thm)],[rqRealNeg__k4_xcmplx_0__rm256_r256])).
cnf(c_0_252, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_224])).
cnf(c_0_253, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_225])).
cnf(c_0_254, plain, (v1_xreal_0(k4_xcmplx_0(np__256))|~v1_xreal_0(np__257)|~v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_227])])).
cnf(c_0_255, plain, (v1_xreal_0(np__257)|~v1_xreal_0(np__1)|~v1_xreal_0(np__256)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_228])])).
cnf(c_0_256, plain, (m1_subset_1(np__256,k1_numbers)), inference(evalgc,[status(thm)],[c_0_229])).
cnf(c_0_257, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_207])])).
cnf(c_0_258, plain, (k4_nat_d(k3_xcmplx_0(k4_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_230, c_0_231])).
cnf(c_0_259, plain, (k4_nat_d(k3_xcmplx_0(X1,k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k4_nat_d(X3,X2))|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_233])])).
fof(c_0_260, plain, ![X259, X260]:((~v7_ordinal1(X259)|~v7_ordinal1(X260)|v7_ordinal1(k3_xcmplx_0(X259,X260)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc2_nat_1])])).
cnf(c_0_261, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_234])).
cnf(c_0_262, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_235, c_0_142]), c_0_142])).
cnf(c_0_263, plain, (v3_ordinal1(np__257)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_236]), c_0_148])])])).
cnf(c_0_264, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_237])).
cnf(c_0_265, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_238])).
cnf(c_0_266, plain, (X1=k6_numbers|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_239, c_0_240])).
cnf(c_0_267, plain, (v1_xboole_0(np__0)), inference(evalgc,[status(thm)],[c_0_241])).
cnf(c_0_268, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)), inference(csr,[status(thm)],[c_0_242, c_0_231])).
cnf(c_0_269, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m2_subset_1(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_243])).
cnf(c_0_270, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_244])).
cnf(c_0_271, plain, (k3_xcmplx_0(np__81,np__81)=np__6561), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r81_r81_r6561])).
cnf(c_0_272, plain, (m1_subset_1(np__136,k5_numbers)), inference(split_conjunct,[status(thm)],[spc136_numerals])).
cnf(c_0_273, plain, (k6_xcmplx_0(np__256,np__257)=k4_xcmplx_0(np__1)), inference(evalgc,[status(thm)],[c_0_245])).
cnf(c_0_274, plain, (v1_xreal_0(np__1)|~v1_xreal_0(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_247])])).
cnf(c_0_275, plain, (v1_xreal_0(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_248, c_0_249]), c_0_250])])])).
cnf(c_0_276, plain, (k4_xcmplx_0(k4_xcmplx_0(np__256))=np__256), inference(evalgc,[status(thm)],[c_0_251])).
cnf(c_0_277, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_253])])).
cnf(c_0_278, plain, (v1_xcmplx_0(k4_xcmplx_0(np__256))|~v1_xreal_0(np__257)|~v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_254])])).
cnf(c_0_279, plain, (v1_xreal_0(np__257)|~v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_249]), c_0_256])])])).
cnf(c_0_280, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_207])])).
cnf(c_0_281, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_207])])).
cnf(c_0_282, plain, (k4_nat_d(k4_nat_1(X1,k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_233])])).
cnf(c_0_283, plain, (k4_nat_d(k3_xcmplx_0(X1,k4_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k4_nat_d(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_259, c_0_191])).
cnf(c_0_284, plain, (v7_ordinal1(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_260])).
cnf(c_0_285, plain, (k4_nat_d(k3_xcmplx_0(k4_nat_d(X3,X2),k2_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X3,X2),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_127])])).
cnf(c_0_286, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_140])])).
cnf(c_0_287, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_249])])).
cnf(c_0_288, plain, (X1=k2_nat_d(X2,X3)|r1_xxreal_0(X3,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(X3,X4))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X4)|~m1_subset_1(k3_xcmplx_0(X3,X4),k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_262])])).
cnf(c_0_289, plain, (k3_xcmplx_0(np__257,np__0)=np__0), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r257_r0_r0])).
cnf(c_0_290, plain, (v7_ordinal1(np__257)|~v1_finset_1(np__257)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_263])])).
cnf(c_0_291, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_264])).
cnf(c_0_292, plain, (m1_subset_1(np__0,k5_numbers)), inference(split_conjunct,[status(thm)],[spc0_numerals])).
cnf(c_0_293, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_265])).
cnf(c_0_294, plain, (k6_numbers=np__0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_267])])).
cnf(c_0_295, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X3),X1),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_209])])).
cnf(c_0_296, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_269])).
cnf(c_0_297, plain, (k6_xcmplx_0(np__6561,np__6425)=np__136), inference(split_conjunct,[status(thm)],[rqRealDiff__k6_xcmplx_0__r6561_r6425_r136])).
cnf(c_0_298, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_270])).
cnf(c_0_299, plain, (k3_xcmplx_0(np__81,np__81)=np__6561), inference(evalgc,[status(thm)],[c_0_271])).
cnf(c_0_300, plain, (m1_subset_1(np__81,k1_numbers)), inference(split_conjunct,[status(thm)],[spc81_numerals])).
cnf(c_0_301, plain, (k6_xcmplx_0(np__6561,np__136)=np__6425), inference(split_conjunct,[status(thm)],[rqRealDiff__k6_xcmplx_0__r6561_r136_r6425])).
cnf(c_0_302, plain, (m1_subset_1(np__136,k5_numbers)), inference(evalgc,[status(thm)],[c_0_272])).
fof(c_0_303, plain, ![X261, X262, X263]:((~v1_xcmplx_0(X261)|~v1_xcmplx_0(X262)|~v1_xcmplx_0(X263)|k3_xcmplx_0(k3_xcmplx_0(X261,X262),X263)=k3_xcmplx_0(X261,k3_xcmplx_0(X262,X263)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc7_arithm])])).
fof(c_0_304, plain, ![X267]:((~v1_xcmplx_0(X267)|k3_xcmplx_0(X267,k4_xcmplx_0(np__1))=k4_xcmplx_0(X267))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc2_arithm])])).
cnf(c_0_305, plain, (v1_xreal_0(k4_xcmplx_0(np__1))|~v1_xreal_0(np__257)|~v1_xreal_0(np__256)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_273])])).
cnf(c_0_306, plain, (v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_274, c_0_275])])])).
cnf(c_0_307, plain, (v1_xreal_0(np__256)|~v1_xreal_0(k4_xcmplx_0(np__256))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_276])])).
cnf(c_0_308, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_277, c_0_249])])).
cnf(c_0_309, plain, (v1_xcmplx_0(k4_xcmplx_0(np__256))|~v1_xreal_0(np__1)), inference(csr,[status(thm)],[c_0_278, c_0_279])).
cnf(c_0_310, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_201])).
cnf(c_0_311, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X3,X2)),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_xcmplx_0(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_233])])).
cnf(c_0_312, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X3,X2)),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_xcmplx_0(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_281, c_0_233])])).
cnf(c_0_313, plain, (k4_nat_d(k4_nat_1(X1,k4_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_282, c_0_191])).
cnf(c_0_314, plain, (k2_nat_d(k3_xcmplx_0(X1,k4_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k3_xcmplx_0(X1,k4_nat_d(X2,X3)))|~v7_ordinal1(k4_nat_d(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_283])])).
cnf(c_0_315, plain, (v7_ordinal1(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_284])).
cnf(c_0_316, plain, (k4_nat_d(k3_xcmplx_0(k4_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_285, c_0_186]), c_0_231])).
cnf(c_0_317, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_287])])).
cnf(c_0_318, plain, (X1=k2_nat_d(X2,X3)|r1_xxreal_0(X3,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(X3,X4))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X4)|~m1_subset_1(k3_xcmplx_0(X3,X4),k5_numbers)), inference(csr,[status(thm)],[c_0_288, c_0_165])).
cnf(c_0_319, plain, (k3_xcmplx_0(np__257,np__0)=np__0), inference(evalgc,[status(thm)],[c_0_289])).
cnf(c_0_320, plain, (v7_ordinal1(np__257)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_167]), c_0_236])])])).
cnf(c_0_321, plain, (v7_ordinal1(np__0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_267])])).
cnf(c_0_322, plain, (m1_subset_1(np__0,k5_numbers)), inference(evalgc,[status(thm)],[c_0_292])).
cnf(c_0_323, plain, (k2_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_293, c_0_294])])).
cnf(c_0_324, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X3),X1),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)), inference(csr,[status(thm)],[c_0_295, c_0_191])).
cnf(c_0_325, plain, (k3_xcmplx_0(np__1,X1)=X1|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_287])])).
cnf(c_0_326, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(X1,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_296, c_0_115]), c_0_116])])).
cnf(c_0_327, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_152]), c_0_147])])])).
cnf(c_0_328, plain, (k6_xcmplx_0(np__6561,np__6425)=np__136), inference(evalgc,[status(thm)],[c_0_297])).
cnf(c_0_329, plain, (v1_xreal_0(np__6561)|~v1_xreal_0(np__81)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_299])])).
cnf(c_0_330, plain, (m1_subset_1(np__81,k1_numbers)), inference(evalgc,[status(thm)],[c_0_300])).
cnf(c_0_331, plain, (k6_xcmplx_0(np__6561,np__136)=np__6425), inference(evalgc,[status(thm)],[c_0_301])).
cnf(c_0_332, plain, (v3_ordinal1(np__136)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_302]), c_0_148])])])).
cnf(c_0_333, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)), inference(split_conjunct,[status(thm)],[c_0_303])).
cnf(c_0_334, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_304])).
cnf(c_0_335, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xreal_0(np__257)|~v1_xreal_0(np__256)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_305])])).
cnf(c_0_336, plain, (v1_xreal_0(np__257)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_279, c_0_306])])])).
cnf(c_0_337, plain, (v1_xreal_0(np__256)|~v1_xcmplx_0(k4_xcmplx_0(np__256))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_307, c_0_308]), c_0_276]), c_0_256])])])).
cnf(c_0_338, plain, (v1_xcmplx_0(k4_xcmplx_0(np__256))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_309, c_0_306])])])).
cnf(c_0_339, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_310])).
cnf(c_0_340, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X3,X2),k4_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(k4_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X1,X2),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_112])])).
cnf(c_0_341, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X3,X2)),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)), inference(csr,[status(thm)],[c_0_311, c_0_191])).
cnf(c_0_342, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X3,X2)),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)), inference(csr,[status(thm)],[c_0_312, c_0_191])).
cnf(c_0_343, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X2,X3),X1),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X2,X3),k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_313, c_0_112])])).
cnf(c_0_344, plain, (k2_nat_d(k3_xcmplx_0(X1,k4_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k4_nat_d(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_314, c_0_315])).
cnf(c_0_345, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(X3,X2),k4_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(k2_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k4_nat_d(X1,X2),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_316, c_0_317])])).
cnf(c_0_346, plain, (X1=k2_nat_d(X2,np__257)|r1_xxreal_0(np__257,X1)|X2!=k2_xcmplx_0(X1,np__0)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_319]), c_0_320]), c_0_321]), c_0_322])])])).
cnf(c_0_347, plain, (k2_xcmplx_0(X1,np__0)=X1|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_140])])).
fof(c_0_348, plain, ~r1_xxreal_0(np__257,np__1), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r257_r1])).
cnf(c_0_349, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,k4_nat_d(X3,X2)),X1),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,k4_nat_d(X3,X2))|~v7_ordinal1(k4_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_324])])).
cnf(c_0_350, plain, (k3_xcmplx_0(X1,np__1)=X1|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_317]), c_0_178])])])).
cnf(c_0_351, plain, (m1_subset_1(X1,k1_numbers)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_326, c_0_327]), c_0_144])])).
cnf(c_0_352, plain, (v1_xreal_0(np__136)|~v1_xreal_0(np__6425)|~v1_xreal_0(np__6561)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_328])])).
cnf(c_0_353, plain, (v1_xreal_0(np__6561)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_249]), c_0_330])])])).
cnf(c_0_354, plain, (v1_xreal_0(np__6425)|~v1_xreal_0(np__136)|~v1_xreal_0(np__6561)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_331])])).
cnf(c_0_355, plain, (v7_ordinal1(np__136)|~v1_finset_1(np__136)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_332])])).
cnf(c_0_356, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_333])).
cnf(c_0_357, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_334])).
cnf(c_0_358, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xreal_0(np__256)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_335, c_0_336])])])).
cnf(c_0_359, plain, (v1_xreal_0(np__256)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_337, c_0_338])])])).
cnf(c_0_360, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k3_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_339])])).
cnf(c_0_361, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(k4_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)), inference(csr,[status(thm)],[c_0_340, c_0_186])).
cnf(c_0_362, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X2)),X2)=k4_nat_d(k4_nat_1(X1,X3),X2)|r1_xxreal_0(X2,X3)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_341, c_0_233])])).
cnf(c_0_363, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,X3)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_342, c_0_233])])).
cnf(c_0_364, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X3,X2),X1),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)), inference(csr,[status(thm)],[c_0_343, c_0_186])).
cnf(c_0_365, plain, (k2_nat_d(k3_xcmplx_0(X1,k2_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k2_nat_d(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_207])])).
cnf(c_0_366, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X3,X2),k1_numbers)), inference(csr,[status(thm)],[c_0_345, c_0_231])).
cnf(c_0_367, plain, (k2_nat_d(X1,np__257)=X1|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_346, c_0_347])])])).
cnf(c_0_368, plain, (~r1_xxreal_0(np__257,np__1)), inference(split_conjunct,[status(thm)],[c_0_348])).
cnf(c_0_369, plain, (m1_subset_1(np__25,k5_numbers)), inference(split_conjunct,[status(thm)],[spc25_numerals])).
cnf(c_0_370, plain, (k4_nat_d(k2_xcmplx_0(X1,k4_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__1,X2),X1)|r1_xxreal_0(X1,k4_nat_d(X2,X1))|~v7_ordinal1(k4_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k2_xcmplx_0(X1,k4_nat_d(X2,X1)),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_349, c_0_350]), c_0_178])])])).
cnf(c_0_371, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_351, c_0_191])])).
cnf(c_0_372, plain, (v1_xreal_0(np__136)|~v1_xreal_0(np__6425)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_352, c_0_353])])])).
cnf(c_0_373, plain, (v1_xreal_0(np__6425)|~v1_xreal_0(np__136)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_354, c_0_353])])])).
cnf(c_0_374, plain, (v7_ordinal1(np__136)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_355, c_0_167]), c_0_302])])])).
cnf(c_0_375, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(k4_xcmplx_0(np__1),X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_357])])).
cnf(c_0_376, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_358, c_0_359])])).
cnf(c_0_377, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_357, c_0_360])])).
cnf(c_0_378, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_207])])).
cnf(c_0_379, plain, (m1_subset_1(np__81,k5_numbers)), inference(split_conjunct,[status(thm)],[spc81_numerals])).
cnf(c_0_380, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X2)),X2)=k4_nat_d(k4_nat_1(X1,X3),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_362, c_0_191])).
cnf(c_0_381, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_363, c_0_191])).
cnf(c_0_382, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X3,X2),X1),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_364, c_0_207])])).
cnf(c_0_383, plain, (k2_nat_d(k3_xcmplx_0(X1,k2_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_365, c_0_231])).
cnf(c_0_384, plain, (k4_nat_d(k3_xcmplx_0(X1,k4_nat_d(X2,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X2,np__257),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_367]), c_0_320])])])).
cnf(c_0_385, plain, (~r1_xxreal_0(np__257,np__1)), inference(evalgc,[status(thm)],[c_0_368])).
cnf(c_0_386, plain, (k3_xcmplx_0(np__257,np__25)=np__6425), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r257_r25_r6425])).
cnf(c_0_387, plain, (m1_subset_1(np__25,k5_numbers)), inference(evalgc,[status(thm)],[c_0_369])).
cnf(c_0_388, plain, (k4_nat_d(k2_xcmplx_0(X1,k4_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__1,X2),X1)|r1_xxreal_0(X1,k4_nat_d(X2,X1))|~v7_ordinal1(k4_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(csr,[status(thm)],[c_0_370, c_0_371])).
cnf(c_0_389, plain, (v1_xcmplx_0(np__136)|~v1_xreal_0(np__6425)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_372])])).
cnf(c_0_390, plain, (v1_xreal_0(np__6425)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_126]), c_0_374])])])).
cnf(c_0_391, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(k4_xcmplx_0(np__1),X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_375, c_0_376])])])).
cnf(c_0_392, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_377, c_0_306])])])).
cnf(c_0_393, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1), inference(split_conjunct,[status(thm)],[rqRealNeg__k4_xcmplx_0__rm1_r1])).
fof(c_0_394, plain, ![X174]:((~v1_xcmplx_0(X174)|v1_xcmplx_0(k4_xcmplx_0(X174)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_xcmplx_0])])).
cnf(c_0_395, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)), inference(csr,[status(thm)],[c_0_378, c_0_231])).
cnf(c_0_396, plain, (m1_subset_1(np__81,k5_numbers)), inference(evalgc,[status(thm)],[c_0_379])).
cnf(c_0_397, plain, (k4_nat_d(k4_nat_1(X1,X3),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_380, c_0_381])])).
cnf(c_0_398, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X3,X2),k2_xcmplx_0(X1,X2)),X2)=k4_nat_d(k4_nat_1(X1,X3),X2)|r1_xxreal_0(X2,X3)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_233])])).
cnf(c_0_399, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(X2,X3)),X3)=k2_nat_d(k3_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,X1)|r1_xxreal_0(X3,X2)|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_233])])).
cnf(c_0_400, plain, (k4_nat_d(k4_nat_d(X1,np__257),np__257)=k4_nat_d(k3_xcmplx_0(np__1,X1),np__257)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X1,np__257),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_384, c_0_325]), c_0_178])]), c_0_385])])).
cnf(c_0_401, plain, (k3_xcmplx_0(np__257,np__25)=np__6425), inference(evalgc,[status(thm)],[c_0_386])).
cnf(c_0_402, plain, (v3_ordinal1(np__25)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_387]), c_0_148])])])).
cnf(c_0_403, plain, (k2_nat_d(k2_xcmplx_0(X1,k4_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__1,X2),X1)|r1_xxreal_0(X1,k4_nat_d(X2,X1))|~v7_ordinal1(k2_xcmplx_0(X1,k4_nat_d(X2,X1)))|~v7_ordinal1(k4_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_388])])).
cnf(c_0_404, plain, (v1_xcmplx_0(np__136)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_389, c_0_390])])])).
fof(c_0_405, plain, ~r1_xxreal_0(np__257,np__136), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r257_r136])).
cnf(c_0_406, plain, (v1_xcmplx_0(np__257)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_336])])).
cnf(c_0_407, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_392])])).
cnf(c_0_408, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1), inference(evalgc,[status(thm)],[c_0_393])).
cnf(c_0_409, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_394])).
cnf(c_0_410, plain, (k4_nat_d(k13_newton(np__3,np__4),k4_pepin(np__3))=np__81), inference(split_conjunct,[status(thm)],[l87_pepin])).
cnf(c_0_411, plain, (k4_pepin(np__3)=np__257), inference(split_conjunct,[status(thm)],[t53_pepin])).
cnf(c_0_412, plain, (k4_nat_d(k3_xcmplx_0(X3,X1),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_395])])).
cnf(c_0_413, plain, (v3_ordinal1(np__81)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_396]), c_0_148])])])).
cnf(c_0_414, plain, (k4_nat_d(k4_nat_1(X2,X1),X3)=k4_nat_d(k3_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,X2)|r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_397, c_0_112])])).
cnf(c_0_415, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X2)),X2)=k4_nat_d(k4_nat_1(X3,X1),X2)|r1_xxreal_0(X2,X3)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)), inference(csr,[status(thm)],[c_0_398, c_0_191])).
cnf(c_0_416, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X2)),X2)=k2_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,X3)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)), inference(csr,[status(thm)],[c_0_399, c_0_191])).
cnf(c_0_417, plain, (k4_nat_d(k2_nat_d(X1,np__257),np__257)=k4_nat_d(k3_xcmplx_0(np__1,X1),np__257)|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X1,np__257),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_207]), c_0_320])])])).
fof(c_0_418, plain, ~r1_xxreal_0(np__257,np__81), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r257_r81])).
cnf(c_0_419, plain, (k2_nat_d(k2_xcmplx_0(np__6425,X1),np__257)=X1|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(np__6425,X1))|~v7_ordinal1(np__257)|~v7_ordinal1(np__25)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_401])])).
cnf(c_0_420, plain, (v7_ordinal1(np__25)|~v1_finset_1(np__25)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_402])])).
cnf(c_0_421, plain, (k2_nat_d(k2_xcmplx_0(X1,k4_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__1,X2),X1)|r1_xxreal_0(X1,k4_nat_d(X2,X1))|~v7_ordinal1(k4_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(csr,[status(thm)],[c_0_403, c_0_191])).
cnf(c_0_422, plain, (k2_nat_d(k2_xcmplx_0(np__0,X1),np__257)=X1|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(np__0,X1))|~v7_ordinal1(np__257)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_319]), c_0_321])])])).
cnf(c_0_423, plain, (k2_xcmplx_0(np__136,np__0)=np__136), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_404])])).
cnf(c_0_424, plain, (~r1_xxreal_0(np__257,np__136)), inference(split_conjunct,[status(thm)],[c_0_405])).
cnf(c_0_425, plain, (k2_xcmplx_0(np__257,X1)=k2_xcmplx_0(X1,np__257)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_406])])).
cnf(c_0_426, plain, (k4_xcmplx_0(k4_xcmplx_0(np__6561))=np__6561), inference(split_conjunct,[status(thm)],[rqRealNeg__k4_xcmplx_0__rm6561_r6561])).
cnf(c_0_427, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k3_xcmplx_0(X1,np__1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_357, c_0_407]), c_0_408]), c_0_376])])])).
cnf(c_0_428, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_409])).
cnf(c_0_429, plain, (k4_nat_d(k13_newton(np__3,np__4),k4_pepin(np__3))=np__81), inference(evalgc,[status(thm)],[c_0_410])).
cnf(c_0_430, plain, (k4_pepin(np__3)=np__257), inference(evalgc,[status(thm)],[c_0_411])).
cnf(c_0_431, plain, (k4_nat_d(k3_xcmplx_0(np__1,X1),X2)=k4_nat_d(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_412, c_0_177]), c_0_178])])])).
cnf(c_0_432, plain, (k3_xcmplx_0(np__1,np__81)=np__81), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r1_r81_r81])).
cnf(c_0_433, plain, (v7_ordinal1(np__81)|~v1_finset_1(np__81)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_413])])).
cnf(c_0_434, plain, (k4_nat_d(k4_nat_1(X1,X2),X3)=k4_nat_d(k3_xcmplx_0(X2,X1),X3)|r1_xxreal_0(X3,X2)|r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(csr,[status(thm)],[c_0_414, c_0_165])).
cnf(c_0_435, plain, (k4_nat_d(k4_nat_1(X3,X1),X2)=k2_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,X3)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_415, c_0_416])])).
cnf(c_0_436, plain, (k4_nat_d(k3_xcmplx_0(np__1,X1),np__257)=k4_nat_d(X1,np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_417, c_0_367])])).
cnf(c_0_437, plain, (k2_xcmplx_0(np__0,np__81)=np__81), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r0_r81_r81])).
cnf(c_0_438, plain, (~r1_xxreal_0(np__257,np__81)), inference(split_conjunct,[status(thm)],[c_0_418])).
cnf(c_0_439, plain, (k2_nat_d(k2_xcmplx_0(np__6425,X1),np__257)=X1|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(np__6425,X1))|~v7_ordinal1(np__25)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_419, c_0_320])])])).
cnf(c_0_440, plain, (v7_ordinal1(np__25)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_420, c_0_167]), c_0_387])])])).
cnf(c_0_441, plain, (k2_xcmplx_0(np__6425,np__136)=np__6561), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r6425_r136_r6561])).
cnf(c_0_442, plain, (v7_ordinal1(np__6561)|~v7_ordinal1(np__81)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_299])])).
cnf(c_0_443, plain, (k2_nat_d(k2_xcmplx_0(X1,k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__1,X2),X1)|r1_xxreal_0(X1,k2_nat_d(X2,X1))|~v7_ordinal1(k2_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_421, c_0_207])])).
cnf(c_0_444, plain, (k2_nat_d(k2_xcmplx_0(np__0,X1),np__257)=X1|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(np__0,X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_422, c_0_320])])])).
cnf(c_0_445, plain, (k2_xcmplx_0(np__0,np__136)=np__136), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_423]), c_0_321]), c_0_374])])])).
cnf(c_0_446, plain, (~r1_xxreal_0(np__257,np__136)), inference(evalgc,[status(thm)],[c_0_424])).
cnf(c_0_447, plain, (k2_nat_d(k2_xcmplx_0(X1,np__257),np__257)=X1|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_425]), c_0_320])])])).
cnf(c_0_448, plain, (k2_xcmplx_0(np__136,X1)=k2_xcmplx_0(X1,np__136)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_404])])).
cnf(c_0_449, plain, (k4_xcmplx_0(k4_xcmplx_0(np__6561))=np__6561), inference(evalgc,[status(thm)],[c_0_426])).
cnf(c_0_450, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k3_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_427, c_0_428])).
cnf(c_0_451, plain, (v1_xcmplx_0(np__6561)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_353])])).
cnf(c_0_452, plain, (m1_subset_1(np__6561,k5_numbers)), inference(split_conjunct,[status(thm)],[spc6561_numerals])).
cnf(c_0_453, plain, (k4_nat_d(k13_newton(np__3,np__4),np__257)=np__81), inference(rw,[status(thm)],[c_0_429, c_0_430])).
cnf(c_0_454, plain, (k2_nat_d(k3_xcmplx_0(np__1,X1),X2)=k4_nat_d(X1,X2)|~v7_ordinal1(k3_xcmplx_0(np__1,X1))|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_431])])).
cnf(c_0_455, plain, (k3_xcmplx_0(np__1,np__81)=np__81), inference(evalgc,[status(thm)],[c_0_432])).
cnf(c_0_456, plain, (v7_ordinal1(np__81)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_433, c_0_167]), c_0_396])])])).
cnf(c_0_457, plain, (k4_nat_d(k3_xcmplx_0(X2,X1),X3)=k2_nat_d(k3_xcmplx_0(X2,X1),X3)|r1_xxreal_0(X3,X2)|r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_434, c_0_435])])).
cnf(c_0_458, plain, (k4_nat_d(k3_xcmplx_0(np__1,X1),np__257)=k4_nat_d(X1,np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)), inference(csr,[status(thm)],[c_0_436, c_0_351])).
cnf(c_0_459, plain, (k2_xcmplx_0(np__0,np__81)=np__81), inference(evalgc,[status(thm)],[c_0_437])).
cnf(c_0_460, plain, (~r1_xxreal_0(np__257,np__81)), inference(evalgc,[status(thm)],[c_0_438])).
cnf(c_0_461, plain, (k2_nat_d(k2_xcmplx_0(np__6425,X1),np__257)=X1|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(np__6425,X1))|~v7_ordinal1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_439, c_0_440])])).
cnf(c_0_462, plain, (k2_xcmplx_0(np__6425,np__136)=np__6561), inference(evalgc,[status(thm)],[c_0_441])).
cnf(c_0_463, plain, (v7_ordinal1(np__6561)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_442, c_0_142]), c_0_396])])])).
cnf(c_0_464, plain, (k2_nat_d(k2_xcmplx_0(X1,k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__1,X2),X1)|r1_xxreal_0(X1,k2_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(csr,[status(thm)],[c_0_443, c_0_231])).
cnf(c_0_465, plain, (k2_nat_d(np__136,np__257)=np__136), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_444, c_0_445]), c_0_374])]), c_0_446])])).
cnf(c_0_466, plain, (k2_nat_d(k2_xcmplx_0(np__257,np__136),np__257)=np__136), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_447, c_0_448]), c_0_374]), c_0_320])]), c_0_446])])).
cnf(c_0_467, plain, (k3_xcmplx_0(np__1,np__136)=np__136), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_404])])).
cnf(c_0_468, plain, (k3_xcmplx_0(np__6561,np__1)=np__6561), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_449, c_0_450]), c_0_451])])])).
cnf(c_0_469, plain, (m1_subset_1(np__6561,k5_numbers)), inference(evalgc,[status(thm)],[c_0_452])).
cnf(c_0_470, plain, (m1_subset_1(np__136,k1_numbers)), inference(split_conjunct,[status(thm)],[spc136_numerals])).
cnf(c_0_471, plain, (k2_xcmplx_0(np__2,np__1)=np__3), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r2_r1_r3])).
cnf(c_0_472, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,np__257),np__81),np__257)=k4_nat_d(k3_xcmplx_0(X1,k13_newton(np__3,np__4)),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(np__257)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_453])])).
cnf(c_0_473, plain, (k4_nat_d(np__81,X1)=k2_nat_d(np__81,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_454, c_0_455]), c_0_456])])])).
cnf(c_0_474, plain, (k2_nat_d(k3_xcmplx_0(np__1,X1),np__257)=k4_nat_d(X1,np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_457, c_0_458]), c_0_320]), c_0_178])]), c_0_385])])).
cnf(c_0_475, plain, (k2_nat_d(np__81,np__257)=np__81), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_444, c_0_459]), c_0_456])]), c_0_460])])).
cnf(c_0_476, plain, (k2_nat_d(np__6561,np__257)=np__136), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_461, c_0_462]), c_0_463]), c_0_374])]), c_0_446])])).
cnf(c_0_477, plain, (k4_nat_d(np__136,np__257)=np__136), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_464, c_0_465]), c_0_466]), c_0_467]), c_0_320]), c_0_374])]), c_0_446])])).
cnf(c_0_478, plain, (k3_xcmplx_0(np__1,np__6561)=np__6561), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_468]), c_0_147]), c_0_469])])])).
cnf(c_0_479, plain, (m1_subset_1(np__136,k1_numbers)), inference(evalgc,[status(thm)],[c_0_470])).
cnf(c_0_480, plain, (k2_xcmplx_0(np__2,np__1)=np__3), inference(evalgc,[status(thm)],[c_0_471])).
fof(c_0_481, negated_conjecture, k4_nat_d(k13_newton(np__3,np__8),k4_pepin(np__3))!=np__136, inference(fof_simplification,[status(thm)],[inference(assume_negation,[status(cth)],[l88_pepin])])).
cnf(c_0_482, plain, (k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),k13_newton(np__3,np__4)),np__257)=k4_nat_d(k4_nat_1(np__81,np__81),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(np__257)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_472, c_0_453])])).
cnf(c_0_483, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(np__81,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__81),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_473]), c_0_456])])])).
cnf(c_0_484, plain, (k4_nat_d(np__81,np__257)=np__81), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_474, c_0_455]), c_0_475]), c_0_456])]), c_0_460])])).
cnf(c_0_485, plain, (k4_nat_d(np__6561,np__257)=np__136), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_417, c_0_476]), c_0_477]), c_0_478]), c_0_463]), c_0_479])])])).
fof(c_0_486, plain, ![X169, X170]:((~m1_subset_1(X169,k5_numbers)|~m1_subset_1(X170,k5_numbers)|k13_newton(X169,X170)=k1_newton(X169,X170))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k13_newton])])).
cnf(c_0_487, plain, (v1_xreal_0(np__3)|~v1_xreal_0(np__1)|~v1_xreal_0(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_480])])).
cnf(c_0_488, plain, (k2_xcmplx_0(np__2,np__2)=np__4), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r2_r2_r4])).
cnf(c_0_489, negated_conjecture, (k4_nat_d(k13_newton(np__3,np__8),k4_pepin(np__3))!=np__136), inference(split_conjunct,[status(thm)],[c_0_481])).
cnf(c_0_490, plain, (m1_subset_1(np__3,k5_numbers)), inference(split_conjunct,[status(thm)],[spc3_numerals])).
cnf(c_0_491, plain, (k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),k13_newton(np__3,np__4)),np__257)=k4_nat_d(k4_nat_1(np__81,np__81),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_482, c_0_320])])])).
cnf(c_0_492, plain, (k4_nat_d(k4_nat_1(np__81,np__81),np__257)=np__136), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_483, c_0_484]), c_0_475]), c_0_299]), c_0_485]), c_0_320]), c_0_456])])])).
cnf(c_0_493, plain, (k13_newton(X1,X2)=k1_newton(X1,X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_486])).
cnf(c_0_494, plain, (m1_subset_1(np__4,k5_numbers)), inference(split_conjunct,[status(thm)],[spc4_numerals])).
fof(c_0_495, plain, ![X210, X211, X212]:((~v7_ordinal1(X210)|(~v7_ordinal1(X211)|(~v1_xcmplx_0(X212)|k1_newton(X212,k2_xcmplx_0(X210,X211))=k3_xcmplx_0(k1_newton(X212,X210),k1_newton(X212,X211)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_newton])])])).
cnf(c_0_496, plain, (v1_xreal_0(np__3)|~v1_xreal_0(np__1)), inference(csr,[status(thm)],[c_0_487, c_0_248])).
cnf(c_0_497, plain, (k2_xcmplx_0(np__2,np__2)=np__4), inference(evalgc,[status(thm)],[c_0_488])).
cnf(c_0_498, plain, (v7_ordinal1(np__2)|~v7_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_221])])).
cnf(c_0_499, negated_conjecture, (k4_nat_d(k13_newton(np__3,np__8),k4_pepin(np__3))!=np__136), inference(evalgc,[status(thm)],[c_0_489])).
cnf(c_0_500, plain, (m1_subset_1(np__8,k5_numbers)), inference(split_conjunct,[status(thm)],[spc8_numerals])).
cnf(c_0_501, plain, (m1_subset_1(np__3,k5_numbers)), inference(evalgc,[status(thm)],[c_0_490])).
cnf(c_0_502, plain, (k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),k13_newton(np__3,np__4)),np__257)=np__136|~v7_ordinal1(k13_newton(np__3,np__4))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_491, c_0_492])])).
cnf(c_0_503, plain, (k13_newton(X1,X2)=k1_newton(X1,X2)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_493])).
cnf(c_0_504, plain, (m1_subset_1(np__4,k5_numbers)), inference(evalgc,[status(thm)],[c_0_494])).
cnf(c_0_505, plain, (k1_newton(X3,k2_xcmplx_0(X1,X2))=k3_xcmplx_0(k1_newton(X3,X1),k1_newton(X3,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v1_xcmplx_0(X3)), inference(split_conjunct,[status(thm)],[c_0_495])).
cnf(c_0_506, plain, (k2_xcmplx_0(np__4,np__4)=np__8), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r4_r4_r8])).
cnf(c_0_507, plain, (v1_xcmplx_0(np__3)|~v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_496])])).
cnf(c_0_508, plain, (v7_ordinal1(np__4)|~v7_ordinal1(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_497])])).
cnf(c_0_509, plain, (v7_ordinal1(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_498, c_0_142]), c_0_147])])])).
cnf(c_0_510, negated_conjecture, (k4_nat_d(k13_newton(np__3,np__8),np__257)!=np__136), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_499, c_0_430])])).
cnf(c_0_511, plain, (m1_subset_1(np__8,k5_numbers)), inference(evalgc,[status(thm)],[c_0_500])).
fof(c_0_512, plain, ![X223, X224]:((~v7_ordinal1(X223)|~v7_ordinal1(X224)|v7_ordinal1(k1_newton(X223,X224)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc4_newton])])).
cnf(c_0_513, plain, (v3_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_501]), c_0_148])])])).
cnf(c_0_514, plain, (k4_nat_d(k3_xcmplx_0(k1_newton(np__3,np__4),k1_newton(np__3,np__4)),np__257)=np__136|~v7_ordinal1(k1_newton(np__3,np__4))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_502, c_0_503]), c_0_504]), c_0_501])])])).
cnf(c_0_515, plain, (k3_xcmplx_0(k1_newton(X3,X1),k1_newton(X3,X2))=k1_newton(X3,k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xcmplx_0(X3)), inference(evalgc,[status(thm)],[c_0_505])).
cnf(c_0_516, plain, (k2_xcmplx_0(np__4,np__4)=np__8), inference(evalgc,[status(thm)],[c_0_506])).
cnf(c_0_517, plain, (v1_xcmplx_0(np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_507, c_0_249]), c_0_250])])])).
cnf(c_0_518, plain, (v7_ordinal1(np__4)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_508, c_0_509])])])).
cnf(c_0_519, negated_conjecture, (k4_nat_d(k1_newton(np__3,np__8),np__257)!=np__136), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_510, c_0_503]), c_0_511]), c_0_501])])])).
cnf(c_0_520, plain, (v7_ordinal1(k1_newton(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_512])).
cnf(c_0_521, plain, (v7_ordinal1(np__3)|~v1_finset_1(np__3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_513])])).
cnf(c_0_522, plain, (~v7_ordinal1(k1_newton(np__3,np__4))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_514, c_0_515]), c_0_516]), c_0_517]), c_0_518])]), c_0_519])])).
cnf(c_0_523, plain, (v7_ordinal1(k1_newton(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_520])).
cnf(c_0_524, plain, (v7_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_521, c_0_167]), c_0_501])])])).
cnf(c_0_525, plain, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_522, c_0_523]), c_0_518]), c_0_524])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 526
# Proof object clause steps            : 391
# Proof object formula steps           : 135
# Proof object conjectures             : 6
# Proof object clause conjectures      : 4
# Proof object formula conjectures     : 2
# Proof object initial clauses used    : 89
# Proof object initial formulas used   : 84
# Proof object generating inferences   : 153
# Proof object simplifying inferences  : 222
# Training examples: 259 positive, 3868 negative
# Training: Positive examples begin
cnf(c_0_499, negated_conjecture, (k4_nat_d(k13_newton(np__3,np__8),k4_pepin(np__3))!=np__136)).# trainpos
cnf(c_0_430, plain, (k4_pepin(np__3)=np__257)).# trainpos
cnf(c_0_429, plain, (k4_nat_d(k13_newton(np__3,np__4),k4_pepin(np__3))=np__81)).# trainpos
cnf(c_0_240, plain, (k1_xboole_0=k6_numbers)).# trainpos
cnf(c_0_501, plain, (m1_subset_1(np__3,k5_numbers))).# trainpos
cnf(c_0_172, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X3),k4_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(X1,X2),X3)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_239, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_128, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).# trainpos
cnf(c_0_103, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_503, plain, (k13_newton(X1,X2)=k1_newton(X1,X2)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_102, plain, (~v1_xboole_0(k4_ordinal1))).# trainpos
cnf(c_0_302, plain, (m1_subset_1(np__136,k5_numbers))).# trainpos
cnf(c_0_163, plain, (X4=k2_nat_d(X2,X3)|r1_xxreal_0(X3,X4)|X2!=k2_xcmplx_0(k3_xcmplx_0(X3,X1),X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_511, plain, (m1_subset_1(np__8,k5_numbers))).# trainpos
cnf(c_0_193, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_155, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1))).# trainpos
cnf(c_0_446, plain, (~r1_xxreal_0(np__257,np__136))).# trainpos
cnf(c_0_147, plain, (m1_subset_1(np__1,k5_numbers))).# trainpos
cnf(c_0_291, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_236, plain, (m1_subset_1(np__257,k5_numbers))).# trainpos
cnf(c_0_387, plain, (m1_subset_1(np__25,k5_numbers))).# trainpos
cnf(c_0_396, plain, (m1_subset_1(np__81,k5_numbers))).# trainpos
cnf(c_0_136, plain, (v3_ordinal1(k4_ordinal1))).# trainpos
cnf(c_0_267, plain, (v1_xboole_0(np__0))).# trainpos
cnf(c_0_221, plain, (k2_xcmplx_0(np__1,np__1)=np__2)).# trainpos
cnf(c_0_516, plain, (k2_xcmplx_0(np__4,np__4)=np__8)).# trainpos
cnf(c_0_126, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_125, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_480, plain, (k2_xcmplx_0(np__2,np__1)=np__3)).# trainpos
cnf(c_0_146, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_228, plain, (k2_xcmplx_0(np__256,np__1)=np__257)).# trainpos
cnf(c_0_106, plain, (~v1_finset_1(k4_ordinal1))).# trainpos
cnf(c_0_114, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_191, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_504, plain, (m1_subset_1(np__4,k5_numbers))).# trainpos
cnf(c_0_194, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_455, plain, (k3_xcmplx_0(np__1,np__81)=np__81)).# trainpos
cnf(c_0_231, plain, (v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_293, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_322, plain, (m1_subset_1(np__0,k5_numbers))).# trainpos
cnf(c_0_479, plain, (m1_subset_1(np__136,k1_numbers))).# trainpos
cnf(c_0_469, plain, (m1_subset_1(np__6561,k5_numbers))).# trainpos
cnf(c_0_523, plain, (v7_ordinal1(k1_newton(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_249, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_157, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1))).# trainpos
cnf(c_0_139, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_186, plain, (m1_subset_1(k4_nat_d(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_328, plain, (k6_xcmplx_0(np__6561,np__6425)=np__136)).# trainpos
cnf(c_0_115, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_187, plain, (k6_int_1(X1,X2)=k2_nat_d(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_459, plain, (k2_xcmplx_0(np__0,np__81)=np__81)).# trainpos
cnf(c_0_250, plain, (m1_subset_1(np__1,k1_numbers))).# trainpos
cnf(c_0_408, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1)).# trainpos
cnf(c_0_296, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_460, plain, (~r1_xxreal_0(np__257,np__81))).# trainpos
cnf(c_0_132, plain, (v1_finset_1(X1)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_118, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_188, plain, (k6_int_1(X1,X2)=k4_nat_d(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_111, plain, (k4_nat_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_315, plain, (v7_ordinal1(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_385, plain, (~r1_xxreal_0(np__257,np__1))).# trainpos
cnf(c_0_112, plain, (k4_nat_1(X1,X2)=k4_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_401, plain, (k3_xcmplx_0(np__257,np__25)=np__6425)).# trainpos
cnf(c_0_497, plain, (k2_xcmplx_0(np__2,np__2)=np__4)).# trainpos
cnf(c_0_247, plain, (k7_xcmplx_0(np__2,np__2)=np__1)).# trainpos
cnf(c_0_462, plain, (k2_xcmplx_0(np__6425,np__136)=np__6561)).# trainpos
cnf(c_0_330, plain, (m1_subset_1(np__81,k1_numbers))).# trainpos
cnf(c_0_180, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_515, plain, (k3_xcmplx_0(k1_newton(X3,X1),k1_newton(X3,X2))=k1_newton(X3,k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xcmplx_0(X3))).# trainpos
cnf(c_0_331, plain, (k6_xcmplx_0(np__6561,np__136)=np__6425)).# trainpos
cnf(c_0_299, plain, (k3_xcmplx_0(np__81,np__81)=np__6561)).# trainpos
cnf(c_0_319, plain, (k3_xcmplx_0(np__257,np__0)=np__0)).# trainpos
cnf(c_0_256, plain, (m1_subset_1(np__256,k1_numbers))).# trainpos
cnf(c_0_117, plain, (m2_subset_1(k4_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_339, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_428, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_252, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_449, plain, (k4_xcmplx_0(k4_xcmplx_0(np__6561))=np__6561)).# trainpos
cnf(c_0_276, plain, (k4_xcmplx_0(k4_xcmplx_0(np__256))=np__256)).# trainpos
cnf(c_0_273, plain, (k6_xcmplx_0(np__256,np__257)=k4_xcmplx_0(np__1))).# trainpos
cnf(c_0_298, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_261, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_220, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_227, plain, (k6_xcmplx_0(np__1,np__257)=k4_xcmplx_0(np__256))).# trainpos
cnf(c_0_226, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_253, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_357, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(np__1))=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_246, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_356, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_294, plain, (k6_numbers=np__0)).# trainpos
cnf(c_0_131, plain, (~v1_finset_1(k1_numbers))).# trainpos
cnf(c_0_510, negated_conjecture, (k4_nat_d(k13_newton(np__3,np__8),np__257)!=np__136)).# trainpos
cnf(c_0_207, plain, (k4_nat_d(X1,X2)=k2_nat_d(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_142, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_321, plain, (v7_ordinal1(np__0))).# trainpos
cnf(c_0_129, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_442, plain, (v7_ordinal1(np__6561)|~v7_ordinal1(np__81))).# trainpos
cnf(c_0_498, plain, (v7_ordinal1(np__2)|~v7_ordinal1(np__1))).# trainpos
cnf(c_0_332, plain, (v3_ordinal1(np__136))).# trainpos
cnf(c_0_472, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,np__257),np__81),np__257)=k4_nat_d(k3_xcmplx_0(X1,k13_newton(np__3,np__4)),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_513, plain, (v3_ordinal1(np__3))).# trainpos
cnf(c_0_402, plain, (v3_ordinal1(np__25))).# trainpos
cnf(c_0_140, plain, (v1_xcmplx_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_156, plain, (v3_ordinal1(np__1))).# trainpos
cnf(c_0_263, plain, (v3_ordinal1(np__257))).# trainpos
cnf(c_0_413, plain, (v3_ordinal1(np__81))).# trainpos
cnf(c_0_508, plain, (v7_ordinal1(np__4)|~v7_ordinal1(np__2))).# trainpos
cnf(c_0_340, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X3,X2),k4_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(k4_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X1,X2),k5_numbers))).# trainpos
cnf(c_0_196, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X3,X2),k4_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(k4_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X3,X2),k5_numbers))).# trainpos
cnf(c_0_176, plain, (k2_nat_d(k2_xcmplx_0(k3_xcmplx_0(X2,X3),X1),X2)=X1|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(k3_xcmplx_0(X2,X3),X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_329, plain, (v1_xreal_0(np__6561)|~v1_xreal_0(np__81))).# trainpos
cnf(c_0_212, plain, (k1_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_287, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_127, plain, (k4_nat_1(X2,X1)=k3_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_326, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(X1,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_248, plain, (v1_xreal_0(np__2)|~v1_xreal_0(np__1))).# trainpos
cnf(c_0_130, plain, (m2_subset_1(k3_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_185, plain, (k4_nat_d(k3_xcmplx_0(k4_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(k4_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X1,X2),k5_numbers))).# trainpos
cnf(c_0_352, plain, (v1_xreal_0(np__136)|~v1_xreal_0(np__6425)|~v1_xreal_0(np__6561))).# trainpos
cnf(c_0_375, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(k4_xcmplx_0(np__1),X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_360, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k3_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_277, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).# trainpos
cnf(c_0_274, plain, (v1_xreal_0(np__1)|~v1_xreal_0(np__2))).# trainpos
cnf(c_0_354, plain, (v1_xreal_0(np__6425)|~v1_xreal_0(np__136)|~v1_xreal_0(np__6561))).# trainpos
cnf(c_0_305, plain, (v1_xreal_0(k4_xcmplx_0(np__1))|~v1_xreal_0(np__257)|~v1_xreal_0(np__256))).# trainpos
cnf(c_0_255, plain, (v1_xreal_0(np__257)|~v1_xreal_0(np__1)|~v1_xreal_0(np__256))).# trainpos
cnf(c_0_487, plain, (v1_xreal_0(np__3)|~v1_xreal_0(np__1)|~v1_xreal_0(np__2))).# trainpos
cnf(c_0_307, plain, (v1_xreal_0(np__256)|~v1_xreal_0(k4_xcmplx_0(np__256)))).# trainpos
cnf(c_0_254, plain, (v1_xreal_0(k4_xcmplx_0(np__256))|~v1_xreal_0(np__257)|~v1_xreal_0(np__1))).# trainpos
cnf(c_0_519, negated_conjecture, (k4_nat_d(k1_newton(np__3,np__8),np__257)!=np__136)).# trainpos
cnf(c_0_144, plain, (~v1_xboole_0(k1_numbers))).# trainpos
cnf(c_0_509, plain, (v7_ordinal1(np__2))).# trainpos
cnf(c_0_257, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_192, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_463, plain, (v7_ordinal1(np__6561))).# trainpos
cnf(c_0_232, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(k4_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_323, plain, (k2_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_152, plain, (k3_xcmplx_0(np__1,X1)=X1|~v7_ordinal1(X1))).# trainpos
cnf(c_0_290, plain, (v7_ordinal1(np__257)|~v1_finset_1(np__257))).# trainpos
cnf(c_0_230, plain, (k4_nat_d(k3_xcmplx_0(k4_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(k2_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_355, plain, (v7_ordinal1(np__136)|~v1_finset_1(np__136))).# trainpos
cnf(c_0_353, plain, (v1_xreal_0(np__6561))).# trainpos
cnf(c_0_275, plain, (v1_xreal_0(np__2))).# trainpos
cnf(c_0_378, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_521, plain, (v7_ordinal1(np__3)|~v1_finset_1(np__3))).# trainpos
cnf(c_0_420, plain, (v7_ordinal1(np__25)|~v1_finset_1(np__25))).# trainpos
cnf(c_0_242, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(k2_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_235, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_166, plain, (v7_ordinal1(np__1)|~v1_finset_1(np__1))).# trainpos
cnf(c_0_433, plain, (v7_ordinal1(np__81)|~v1_finset_1(np__81))).# trainpos
cnf(c_0_286, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_507, plain, (v1_xcmplx_0(np__3)|~v1_xreal_0(np__1))).# trainpos
cnf(c_0_141, plain, (k3_xcmplx_0(X2,X1)=k3_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_422, plain, (k2_nat_d(k2_xcmplx_0(np__0,X1),np__257)=X1|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(np__0,X1))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_335, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xreal_0(np__257)|~v1_xreal_0(np__256))).# trainpos
cnf(c_0_325, plain, (k3_xcmplx_0(np__1,X1)=X1|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_482, plain, (k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),k13_newton(np__3,np__4)),np__257)=k4_nat_d(k4_nat_1(np__81,np__81),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(np__257))).# trainpos
cnf(c_0_377, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(np__1))).# trainpos
cnf(c_0_279, plain, (v1_xreal_0(np__257)|~v1_xreal_0(np__1))).# trainpos
cnf(c_0_419, plain, (k2_nat_d(k2_xcmplx_0(np__6425,X1),np__257)=X1|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(np__6425,X1))|~v7_ordinal1(np__257)|~v7_ordinal1(np__25)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_278, plain, (v1_xcmplx_0(k4_xcmplx_0(np__256))|~v1_xreal_0(np__257)|~v1_xreal_0(np__1))).# trainpos
cnf(c_0_308, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).# trainpos
cnf(c_0_143, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k3_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_178, plain, (v7_ordinal1(np__1))).# trainpos
cnf(c_0_374, plain, (v7_ordinal1(np__136))).# trainpos
cnf(c_0_524, plain, (v7_ordinal1(np__3))).# trainpos
cnf(c_0_320, plain, (v7_ordinal1(np__257))).# trainpos
cnf(c_0_456, plain, (v7_ordinal1(np__81))).# trainpos
cnf(c_0_518, plain, (v7_ordinal1(np__4))).# trainpos
cnf(c_0_210, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_280, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_165, plain, (m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_281, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_347, plain, (k2_xcmplx_0(X1,np__0)=X1|~v7_ordinal1(X1))).# trainpos
cnf(c_0_517, plain, (v1_xcmplx_0(np__3))).# trainpos
cnf(c_0_440, plain, (v7_ordinal1(np__25))).# trainpos
cnf(c_0_164, plain, (k3_xcmplx_0(X1,np__1)=X1|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_373, plain, (v1_xreal_0(np__6425)|~v1_xreal_0(np__136))).# trainpos
cnf(c_0_451, plain, (v1_xcmplx_0(np__6561))).# trainpos
cnf(c_0_306, plain, (v1_xreal_0(np__1))).# trainpos
cnf(c_0_327, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_285, plain, (k4_nat_d(k3_xcmplx_0(k4_nat_d(X3,X2),k2_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X3,X2),k5_numbers))).# trainpos
cnf(c_0_372, plain, (v1_xreal_0(np__136)|~v1_xreal_0(np__6425))).# trainpos
cnf(c_0_317, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_412, plain, (k4_nat_d(k3_xcmplx_0(X3,X1),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_288, plain, (X1=k2_nat_d(X2,X3)|r1_xxreal_0(X3,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(X3,X4))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X4)|~m1_subset_1(k3_xcmplx_0(X3,X4),k5_numbers)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_351, plain, (m1_subset_1(X1,k1_numbers)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_389, plain, (v1_xcmplx_0(np__136)|~v1_xreal_0(np__6425))).# trainpos
cnf(c_0_190, plain, (k2_nat_d(k2_xcmplx_0(X1,X2),X1)=X2|r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_431, plain, (k4_nat_d(k3_xcmplx_0(np__1,X1),X2)=k4_nat_d(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_392, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_390, plain, (v1_xreal_0(np__6425))).# trainpos
cnf(c_0_346, plain, (X1=k2_nat_d(X2,np__257)|r1_xxreal_0(np__257,X1)|X2!=k2_xcmplx_0(X1,np__0)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_336, plain, (v1_xreal_0(np__257))).# trainpos
cnf(c_0_350, plain, (k3_xcmplx_0(X1,np__1)=X1|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_345, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(X3,X2),k4_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(k2_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k4_nat_d(X1,X2),k1_numbers))).# trainpos
cnf(c_0_404, plain, (v1_xcmplx_0(np__136))).# trainpos
cnf(c_0_406, plain, (v1_xcmplx_0(np__257))).# trainpos
cnf(c_0_233, plain, (k2_nat_d(k2_xcmplx_0(X2,X1),X1)=X2|r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_367, plain, (k2_nat_d(X1,np__257)=X1|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_295, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X3),X1),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_371, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_423, plain, (k2_xcmplx_0(np__136,np__0)=np__136)).# trainpos
cnf(c_0_467, plain, (k3_xcmplx_0(np__1,np__136)=np__136)).# trainpos
cnf(c_0_425, plain, (k2_xcmplx_0(np__257,X1)=k2_xcmplx_0(X1,np__257)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_311, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X3,X2)),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_xcmplx_0(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_282, plain, (k4_nat_d(k4_nat_1(X1,k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_349, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,k4_nat_d(X3,X2)),X1),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,k4_nat_d(X3,X2))|~v7_ordinal1(k4_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_448, plain, (k2_xcmplx_0(np__136,X1)=k2_xcmplx_0(X1,np__136)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_312, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X3,X2)),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_xcmplx_0(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_445, plain, (k2_xcmplx_0(np__0,np__136)=np__136)).# trainpos
cnf(c_0_370, plain, (k4_nat_d(k2_xcmplx_0(X1,k4_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__1,X2),X1)|r1_xxreal_0(X1,k4_nat_d(X2,X1))|~v7_ordinal1(k4_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k2_xcmplx_0(X1,k4_nat_d(X2,X1)),k1_numbers))).# trainpos
cnf(c_0_447, plain, (k2_nat_d(k2_xcmplx_0(X1,np__257),np__257)=X1|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_362, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X2)),X2)=k4_nat_d(k4_nat_1(X1,X3),X2)|r1_xxreal_0(X2,X3)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_259, plain, (k4_nat_d(k3_xcmplx_0(X1,k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k4_nat_d(X3,X2))|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_466, plain, (k2_nat_d(k2_xcmplx_0(np__257,np__136),np__257)=np__136)).# trainpos
cnf(c_0_403, plain, (k2_nat_d(k2_xcmplx_0(X1,k4_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__1,X2),X1)|r1_xxreal_0(X1,k4_nat_d(X2,X1))|~v7_ordinal1(k2_xcmplx_0(X1,k4_nat_d(X2,X1)))|~v7_ordinal1(k4_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_337, plain, (v1_xreal_0(np__256)|~v1_xcmplx_0(k4_xcmplx_0(np__256)))).# trainpos
cnf(c_0_314, plain, (k2_nat_d(k3_xcmplx_0(X1,k4_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k3_xcmplx_0(X1,k4_nat_d(X2,X3)))|~v7_ordinal1(k4_nat_d(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_439, plain, (k2_nat_d(k2_xcmplx_0(np__6425,X1),np__257)=X1|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(np__6425,X1))|~v7_ordinal1(np__25)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_444, plain, (k2_nat_d(k2_xcmplx_0(np__0,X1),np__257)=X1|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(np__0,X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_343, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X2,X3),X1),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X2,X3),k5_numbers))).# trainpos
cnf(c_0_465, plain, (k2_nat_d(np__136,np__257)=np__136)).# trainpos
cnf(c_0_476, plain, (k2_nat_d(np__6561,np__257)=np__136)).# trainpos
cnf(c_0_475, plain, (k2_nat_d(np__81,np__257)=np__81)).# trainpos
cnf(c_0_382, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X3,X2),X1),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_363, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,X3)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_397, plain, (k4_nat_d(k4_nat_1(X1,X3),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_398, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X3,X2),k2_xcmplx_0(X1,X2)),X2)=k4_nat_d(k4_nat_1(X1,X3),X2)|r1_xxreal_0(X2,X3)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_414, plain, (k4_nat_d(k4_nat_1(X2,X1),X3)=k4_nat_d(k3_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,X2)|r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_443, plain, (k2_nat_d(k2_xcmplx_0(X1,k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__1,X2),X1)|r1_xxreal_0(X1,k2_nat_d(X2,X1))|~v7_ordinal1(k2_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_477, plain, (k4_nat_d(np__136,np__257)=np__136)).# trainpos
cnf(c_0_365, plain, (k2_nat_d(k3_xcmplx_0(X1,k2_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k2_nat_d(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_491, plain, (k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),k13_newton(np__3,np__4)),np__257)=k4_nat_d(k4_nat_1(np__81,np__81),np__257)|~v7_ordinal1(k13_newton(np__3,np__4)))).# trainpos
cnf(c_0_399, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(X2,X3)),X3)=k2_nat_d(k3_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,X1)|r1_xxreal_0(X3,X2)|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_435, plain, (k4_nat_d(k4_nat_1(X3,X1),X2)=k2_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,X3)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).# trainpos
cnf(c_0_457, plain, (k4_nat_d(k3_xcmplx_0(X2,X1),X3)=k2_nat_d(k3_xcmplx_0(X2,X1),X3)|r1_xxreal_0(X3,X2)|r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_338, plain, (v1_xcmplx_0(k4_xcmplx_0(np__256)))).# trainpos
cnf(c_0_359, plain, (v1_xreal_0(np__256))).# trainpos
cnf(c_0_454, plain, (k2_nat_d(k3_xcmplx_0(np__1,X1),X2)=k4_nat_d(X1,X2)|~v7_ordinal1(k3_xcmplx_0(np__1,X1))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_358, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xreal_0(np__256))).# trainpos
cnf(c_0_473, plain, (k4_nat_d(np__81,X1)=k2_nat_d(np__81,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_483, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(np__81,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__81),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_391, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(k4_xcmplx_0(np__1),X2))=k3_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_407, plain, (k3_xcmplx_0(k4_xcmplx_0(X1),X2)=k3_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_427, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k3_xcmplx_0(X1,np__1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_468, plain, (k3_xcmplx_0(np__6561,np__1)=np__6561)).# trainpos
cnf(c_0_384, plain, (k4_nat_d(k3_xcmplx_0(X1,k4_nat_d(X2,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X2,np__257),k1_numbers))).# trainpos
cnf(c_0_478, plain, (k3_xcmplx_0(np__1,np__6561)=np__6561)).# trainpos
cnf(c_0_400, plain, (k4_nat_d(k4_nat_d(X1,np__257),np__257)=k4_nat_d(k3_xcmplx_0(np__1,X1),np__257)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X1,np__257),k1_numbers))).# trainpos
cnf(c_0_417, plain, (k4_nat_d(k2_nat_d(X1,np__257),np__257)=k4_nat_d(k3_xcmplx_0(np__1,X1),np__257)|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X1,np__257),k1_numbers))).# trainpos
cnf(c_0_485, plain, (k4_nat_d(np__6561,np__257)=np__136)).# trainpos
cnf(c_0_436, plain, (k4_nat_d(k3_xcmplx_0(np__1,X1),np__257)=k4_nat_d(X1,np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_474, plain, (k2_nat_d(k3_xcmplx_0(np__1,X1),np__257)=k4_nat_d(X1,np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_484, plain, (k4_nat_d(np__81,np__257)=np__81)).# trainpos
cnf(c_0_492, plain, (k4_nat_d(k4_nat_1(np__81,np__81),np__257)=np__136)).# trainpos
cnf(c_0_502, plain, (k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),k13_newton(np__3,np__4)),np__257)=np__136|~v7_ordinal1(k13_newton(np__3,np__4)))).# trainpos
cnf(c_0_514, plain, (k4_nat_d(k3_xcmplx_0(k1_newton(np__3,np__4),k1_newton(np__3,np__4)),np__257)=np__136|~v7_ordinal1(k1_newton(np__3,np__4)))).# trainpos
cnf(c_0_522, plain, (~v7_ordinal1(k1_newton(np__3,np__4)))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_250, plain, (k1_nat_1(k2_newton(np__2,k2_newton(np__2,X1)),np__1)=k4_pepin(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_387, plain, (X1=k2_nat_d(X3,X2)|X2!=k6_numbers|X1!=k6_numbers|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6, plain, (v2_xxreal_0(np__3))).#trainneg
cnf(i_0_477, plain, (k4_xcmplx_0(np__0)=np__0)).#trainneg
cnf(i_0_526, plain, (~v1_xboole_0(np__3))).#trainneg
cnf(i_0_586, plain, (~v1_xboole_0(np__8))).#trainneg
cnf(i_0_437, plain, (~v1_xboole_0(np__136))).#trainneg
cnf(i_0_351, plain, (~v1_xboole_0(np__81))).#trainneg
cnf(i_0_334, plain, (~v1_xboole_0(np__257))).#trainneg
cnf(i_0_420, plain, (~v1_xboole_0(np__6561))).#trainneg
cnf(i_0_425, plain, (~v1_xboole_0(np__6425))).#trainneg
cnf(i_0_384, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_35, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_89, plain, (m1_subset_1(np__2,k5_numbers))).#trainneg
cnf(i_0_467, plain, (k2_xcmplx_0(np__3,np__5)=np__8)).#trainneg
cnf(i_0_75, plain, (v2_xxreal_0(np__8))).#trainneg
cnf(i_0_79, plain, (v2_xxreal_0(np__136))).#trainneg
cnf(i_0_383, plain, (~v1_xboole_0(np__79))).#trainneg
cnf(i_0_557, plain, (~v1_xboole_0(np__256))).#trainneg
cnf(i_0_344, plain, (~v1_xboole_0(np__231))).#trainneg
cnf(i_0_406, plain, (~v1_xboole_0(np__6289))).#trainneg
cnf(i_0_435, plain, (~v1_xboole_0(np__23))).#trainneg
cnf(i_0_431, plain, (~v1_xboole_0(np__56))).#trainneg
cnf(i_0_424, plain, (~v1_xboole_0(np__6536))).#trainneg
cnf(i_0_452, plain, (~v1_xboole_0(np__6168))).#trainneg
cnf(i_0_390, plain, (X1=k6_numbers|X1!=k2_nat_d(X3,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X2,X1))).#trainneg
cnf(i_0_42, plain, (m1_subset_1(k4_pepin(X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_333, plain, (k2_xcmplx_0(np__3,np__1)=np__4)).#trainneg
cnf(i_0_389, plain, (X1=k6_numbers|X2!=k2_nat_d(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_241, plain, (k2_xcmplx_0(np__1,np__256)=np__257)).#trainneg
cnf(i_0_522, plain, (v1_xxreal_0(esk4_0))).#trainneg
cnf(i_0_455, plain, (k2_xcmplx_0(np__1,np__7)=np__8)).#trainneg
cnf(i_0_511, plain, (r1_xxreal_0(np__3,np__3))).#trainneg
cnf(i_0_36, plain, (v1_xreal_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_98, plain, (v2_xxreal_0(np__4))).#trainneg
cnf(i_0_491, plain, (~v1_xboole_0(np__135))).#trainneg
cnf(i_0_551, plain, (~v1_xboole_0(np__2))).#trainneg
cnf(i_0_433, plain, (~v1_xboole_0(np__25))).#trainneg
cnf(i_0_458, plain, (~v1_xboole_0(np__24))).#trainneg
cnf(i_0_418, plain, (~v1_xboole_0(np__78))).#trainneg
cnf(i_0_419, plain, (~v1_xboole_0(np__77))).#trainneg
cnf(i_0_397, plain, (~v1_xboole_0(np__6560))).#trainneg
cnf(i_0_421, plain, (~v1_xboole_0(np__6559))).#trainneg
cnf(i_0_482, plain, (k2_xcmplx_0(np__2,np__6)=np__8)).#trainneg
cnf(i_0_404, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_414, plain, (k2_xcmplx_0(np__1,np__2)=np__3)).#trainneg
cnf(i_0_533, plain, (k2_xcmplx_0(np__257,np__0)=np__257)).#trainneg
cnf(i_0_292, plain, (v1_xxreal_0(esk1_0))).#trainneg
cnf(i_0_255, plain, (k2_newton(X1,X2)=k1_newton(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_196, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_386, plain, (~r1_xxreal_0(np__8,np__3))).#trainneg
cnf(i_0_21, plain, (v2_xxreal_0(np__81))).#trainneg
cnf(i_0_426, plain, (~v1_xboole_0(np__6423))).#trainneg
cnf(i_0_442, plain, (~v1_xboole_0(np__6422))).#trainneg
cnf(i_0_559, plain, (~v1_xboole_0(np__6421))).#trainneg
cnf(i_0_430, plain, (~v1_xboole_0(np__6169))).#trainneg
cnf(i_0_505, plain, (~v1_xboole_0(np__21))).#trainneg
cnf(i_0_436, plain, (~v1_xboole_0(np__176))).#trainneg
cnf(i_0_450, plain, (~v1_xboole_0(np__175))).#trainneg
cnf(i_0_439, plain, (~v1_xboole_0(np__132))).#trainneg
cnf(i_0_252, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_489, plain, (~v7_ordinal1(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_512, plain, (k2_xcmplx_0(np__3,np__0)=np__3)).#trainneg
cnf(i_0_281, plain, (v1_finset_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_372, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_380, plain, (~r1_xxreal_0(np__136,np__3))).#trainneg
cnf(i_0_570, plain, (v3_ordinal1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_14, plain, (v2_xxreal_0(np__257))).#trainneg
cnf(i_0_463, plain, (~v1_xboole_0(np__121))).#trainneg
cnf(i_0_440, plain, (~v1_xboole_0(np__111))).#trainneg
cnf(i_0_515, plain, (~v1_xboole_0(np__22))).#trainneg
cnf(i_0_428, plain, (~v1_xboole_0(np__6344))).#trainneg
cnf(i_0_423, plain, (~v1_xboole_0(np__6557))).#trainneg
cnf(i_0_427, plain, (~v1_xboole_0(np__6400))).#trainneg
cnf(i_0_396, plain, (~v1_xboole_0(np__232))).#trainneg
cnf(i_0_438, plain, (~v1_xboole_0(np__134))).#trainneg
cnf(i_0_265, plain, (~r1_xxreal_0(np__6425,np__136))).#trainneg
cnf(i_0_572, plain, (~v1_int_1(X1)|~v1_abian(X1)|~v1_abian(k2_xcmplx_0(X1,np__1)))).#trainneg
cnf(i_0_3, plain, (m1_subset_1(np__3,k1_numbers))).#trainneg
cnf(i_0_216, plain, (~r1_xxreal_0(np__6561,np__136))).#trainneg
cnf(i_0_476, plain, (k2_xcmplx_0(np__6,np__2)=np__8)).#trainneg
cnf(i_0_27, plain, (m1_subset_1(np__6425,k5_numbers))).#trainneg
cnf(i_0_16, plain, (v1_int_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_315, plain, (k3_xcmplx_0(X1,k6_numbers)=k6_numbers|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_200, plain, (v2_xxreal_0(np__6561))).#trainneg
cnf(i_0_417, plain, (~v1_xboole_0(np__80))).#trainneg
cnf(i_0_429, plain, (~v1_xboole_0(np__6304))).#trainneg
cnf(i_0_453, plain, (~v1_xboole_0(np__6424))).#trainneg
cnf(i_0_422, plain, (~v1_xboole_0(np__6558))).#trainneg
cnf(i_0_464, plain, (~v1_xboole_0(np__252))).#trainneg
cnf(i_0_280, plain, (~v1_xboole_0(np__6305))).#trainneg
cnf(i_0_432, plain, (~v1_xboole_0(np__55))).#trainneg
cnf(i_0_378, plain, (~v1_xboole_0(np__120))).#trainneg
cnf(i_0_318, plain, (k3_xcmplx_0(np__2,np__4)=np__8)).#trainneg
cnf(i_0_257, plain, (~r1_xxreal_0(np__136,np__25))).#trainneg
cnf(i_0_473, plain, (m2_subset_1(k13_newton(X1,X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_577, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_398, plain, (k2_xcmplx_0(np__81,np__0)=np__81)).#trainneg
cnf(i_0_536, plain, (v3_ordinal1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_25, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_568, plain, (~v1_xboole_0(np__255))).#trainneg
cnf(i_0_434, plain, (~v1_xboole_0(np__253))).#trainneg
cnf(i_0_457, plain, (~v1_xboole_0(np__254))).#trainneg
cnf(i_0_456, plain, (~v1_xboole_0(np__6480))).#trainneg
cnf(i_0_484, plain, (~r1_xxreal_0(np__3,np__1))).#trainneg
cnf(i_0_530, plain, (~r1_xxreal_0(np__3,np__2))).#trainneg
cnf(i_0_403, plain, (~r1_xxreal_0(np__3,np__0))).#trainneg
cnf(i_0_465, plain, (~r1_xxreal_0(np__8,np__4))).#trainneg
cnf(i_0_534, plain, (k2_xcmplx_0(np__0,np__257)=np__257)).#trainneg
cnf(i_0_496, plain, (k3_xcmplx_0(np__81,np__1)=np__81)).#trainneg
cnf(i_0_317, plain, (k3_xcmplx_0(np__4,np__2)=np__8)).#trainneg
cnf(i_0_449, plain, (k2_xcmplx_0(np__7,np__1)=np__8)).#trainneg
cnf(i_0_493, plain, (r1_xxreal_0(np__0,np__257))).#trainneg
cnf(i_0_88, plain, (m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_407, plain, (k2_xcmplx_0(np__0,np__3)=np__3)).#trainneg
cnf(i_0_371, plain, (~r1_xxreal_0(np__136,np__4))).#trainneg
cnf(i_0_39, plain, (m1_subset_1(np__256,k5_numbers))).#trainneg
cnf(i_0_339, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_29, plain, (v2_xxreal_0(np__6425))).#trainneg
cnf(i_0_486, plain, (~r1_xxreal_0(np__8,np__1))).#trainneg
cnf(i_0_555, plain, (~r1_xxreal_0(np__8,np__6))).#trainneg
cnf(i_0_443, plain, (~r1_xxreal_0(np__8,np__2))).#trainneg
cnf(i_0_411, plain, (~r1_xxreal_0(np__8,np__0))).#trainneg
cnf(i_0_558, plain, (~r1_xxreal_0(np__8,np__5))).#trainneg
cnf(i_0_501, plain, (~r1_xxreal_0(np__8,np__7))).#trainneg
cnf(i_0_195, plain, (~r1_xxreal_0(np__136,np__81))).#trainneg
cnf(i_0_304, plain, (~r1_xxreal_0(np__136,np__1))).#trainneg
cnf(i_0_72, plain, (m1_subset_1(np__8,k1_numbers))).#trainneg
cnf(i_0_273, plain, (k2_xcmplx_0(np__5,np__3)=np__8)).#trainneg
cnf(i_0_481, plain, (v1_card_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_211, plain, (~r1_xxreal_0(np__257,np__25))).#trainneg
cnf(i_0_33, plain, (v2_xxreal_0(np__79))).#trainneg
cnf(i_0_346, plain, (~r1_xxreal_0(np__136,np__2))).#trainneg
cnf(i_0_286, plain, (~r1_xxreal_0(np__136,np__0))).#trainneg
cnf(i_0_521, plain, (~r1_xxreal_0(np__4,np__3))).#trainneg
cnf(i_0_471, plain, (~r1_xxreal_0(np__4,np__2))).#trainneg
cnf(i_0_459, plain, (~r1_xxreal_0(np__4,np__0))).#trainneg
cnf(i_0_303, plain, (~r1_xxreal_0(np__81,np__3))).#trainneg
cnf(i_0_307, plain, (~r1_xxreal_0(np__81,np__4))).#trainneg
cnf(i_0_296, plain, (~r1_xxreal_0(np__81,np__1))).#trainneg
cnf(i_0_382, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_416, plain, (k4_xcmplx_0(k4_xcmplx_0(np__2))=np__2)).#trainneg
cnf(i_0_210, plain, (~r1_xxreal_0(np__6425,np__257))).#trainneg
cnf(i_0_285, plain, (~r1_xxreal_0(np__257,np__4))).#trainneg
cnf(i_0_245, plain, (k4_xcmplx_0(k4_xcmplx_0(np__3))=np__3)).#trainneg
cnf(i_0_135, plain, (m1_subset_1(np__25,k1_numbers))).#trainneg
cnf(i_0_140, plain, (m1_subset_1(np__21,k5_numbers))).#trainneg
cnf(i_0_361, plain, (~r1_xxreal_0(np__25,np__4))).#trainneg
cnf(i_0_295, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_373, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_41, plain, (v2_xxreal_0(np__256))).#trainneg
cnf(i_0_209, plain, (~r1_xxreal_0(np__257,np__3))).#trainneg
cnf(i_0_332, plain, (~r1_xxreal_0(np__6561,np__3))).#trainneg
cnf(i_0_282, plain, (~r1_xxreal_0(np__6561,np__8))).#trainneg
cnf(i_0_276, plain, (~r1_xxreal_0(np__81,np__2))).#trainneg
cnf(i_0_224, plain, (~r1_xxreal_0(np__81,np__25))).#trainneg
cnf(i_0_239, plain, (~r1_xxreal_0(np__81,np__0))).#trainneg
cnf(i_0_400, plain, (~r1_xxreal_0(np__81,np__5))).#trainneg
cnf(i_0_589, plain, (~v2_xxreal_0(X1)|~v1_xxreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_254, plain, (k2_xcmplx_0(np__3,np__3)=np__6)).#trainneg
cnf(i_0_405, plain, (k2_xcmplx_0(np__1,np__3)=np__4)).#trainneg
cnf(i_0_381, plain, (k2_xcmplx_0(np__3,np__4)=np__7)).#trainneg
cnf(i_0_247, plain, (m1_subset_1(np__5,k5_numbers))).#trainneg
cnf(i_0_325, plain, (k2_xcmplx_0(np__3,np__2)=np__5)).#trainneg
cnf(i_0_553, plain, (k2_xcmplx_0(np__257,k4_xcmplx_0(np__1))=np__256)).#trainneg
cnf(i_0_275, plain, (~r1_xxreal_0(np__257,np__2))).#trainneg
cnf(i_0_58, plain, (~r1_xxreal_0(np__6561,np__257))).#trainneg
cnf(i_0_95, plain, (m1_subset_1(np__4,k1_numbers))).#trainneg
cnf(i_0_591, plain, (~v2_xxreal_0(X1)|~v1_xxreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_46, plain, (v2_xxreal_0(np__231))).#trainneg
cnf(i_0_52, plain, (v2_xxreal_0(np__6289))).#trainneg
cnf(i_0_408, plain, (~r1_xxreal_0(np__6425,np__3))).#trainneg
cnf(i_0_552, plain, (~r1_xxreal_0(np__256,np__3))).#trainneg
cnf(i_0_513, plain, (~r1_xxreal_0(np__25,np__3))).#trainneg
cnf(i_0_338, plain, (~r1_xxreal_0(np__257,np__256))).#trainneg
cnf(i_0_240, plain, (~r1_xxreal_0(np__257,np__0))).#trainneg
cnf(i_0_312, plain, (~r1_xxreal_0(np__6561,np__4))).#trainneg
cnf(i_0_63, plain, (~r1_xxreal_0(np__6561,np__81))).#trainneg
cnf(i_0_314, plain, (~r1_xxreal_0(np__6561,np__1))).#trainneg
cnf(i_0_168, plain, (m1_subset_1(np__22,k5_numbers))).#trainneg
cnf(i_0_541, plain, (r1_xxreal_0(np__2,np__6561))).#trainneg
cnf(i_0_69, plain, (m1_subset_1(np__6,k5_numbers))).#trainneg
cnf(i_0_544, plain, (k7_xcmplx_0(np__4,np__2)=np__2)).#trainneg
cnf(i_0_172, plain, (k2_xcmplx_0(np__136,np__6425)=np__6561)).#trainneg
cnf(i_0_272, plain, (~r1_xxreal_0(np__6425,np__25))).#trainneg
cnf(i_0_502, plain, (v1_abian(k3_xcmplx_0(np__2,X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_548, plain, (m2_subset_1(k1_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_320, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_283, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_56, plain, (v2_xxreal_0(np__23))).#trainneg
cnf(i_0_62, plain, (v2_xxreal_0(np__56))).#trainneg
cnf(i_0_228, plain, (~r1_xxreal_0(np__6561,np__6425))).#trainneg
cnf(i_0_331, plain, (~r1_xxreal_0(np__6561,np__256))).#trainneg
cnf(i_0_267, plain, (~r1_xxreal_0(np__6561,np__2))).#trainneg
cnf(i_0_223, plain, (~r1_xxreal_0(np__6561,np__25))).#trainneg
cnf(i_0_227, plain, (~r1_xxreal_0(np__6561,np__6560))).#trainneg
cnf(i_0_253, plain, (~r1_xxreal_0(np__6561,np__0))).#trainneg
cnf(i_0_238, plain, (~r1_xxreal_0(np__1,np__0))).#trainneg
cnf(i_0_395, plain, (~r1_xxreal_0(np__6425,np__4))).#trainneg
cnf(i_0_11, plain, (m1_subset_1(np__257,k1_numbers))).#trainneg
cnf(i_0_262, plain, (m1_subset_1(np__7,k5_numbers))).#trainneg
cnf(i_0_221, plain, (~r1_xxreal_0(np__6425,np__81))).#trainneg
cnf(i_0_341, plain, (m1_subset_1(np__255,k5_numbers))).#trainneg
cnf(i_0_503, plain, (m1_subset_1(esk3_1(X1),X1))).#trainneg
cnf(i_0_500, plain, (~r1_xxreal_0(np__25,np__2))).#trainneg
cnf(i_0_583, plain, (k2_xcmplx_0(np__4,np__0)=np__4)).#trainneg
cnf(i_0_410, plain, (k7_xcmplx_0(k6_numbers,X1)=k6_numbers|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_293, plain, (v1_xcmplx_0(esk1_0))).#trainneg
cnf(i_0_375, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_67, plain, (v2_xxreal_0(np__6536))).#trainneg
cnf(i_0_71, plain, (v2_xxreal_0(np__6))).#trainneg
cnf(i_0_402, plain, (~r1_xxreal_0(np__6425,np__1))).#trainneg
cnf(i_0_370, plain, (~r1_xxreal_0(np__6425,np__2))).#trainneg
cnf(i_0_311, plain, (~r1_xxreal_0(np__6425,np__0))).#trainneg
cnf(i_0_535, plain, (~r1_xxreal_0(np__256,np__4))).#trainneg
cnf(i_0_305, plain, (~r1_xxreal_0(np__256,np__81))).#trainneg
cnf(i_0_454, plain, (~r1_xxreal_0(np__256,np__1))).#trainneg
cnf(i_0_516, plain, (~r1_xxreal_0(np__256,np__2))).#trainneg
cnf(i_0_485, plain, (~r1_xxreal_0(np__256,np__0))).#trainneg
cnf(i_0_509, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_474, plain, (~r1_xxreal_0(np__25,np__1))).#trainneg
cnf(i_0_335, plain, (k2_xcmplx_0(np__1,np__21)=np__22)).#trainneg
cnf(i_0_523, plain, (v1_xboole_0(esk4_0))).#trainneg
cnf(i_0_392, plain, (X1=k6_numbers|k2_xcmplx_0(k3_xcmplx_0(X3,esk2_3(X2,X3,X1)),X1)=X2|X1!=k2_nat_d(X2,X3)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_413, plain, (k2_xcmplx_0(np__1,np__6)=np__7)).#trainneg
cnf(i_0_446, plain, (k2_xcmplx_0(np__1,np__4)=np__5)).#trainneg
cnf(i_0_547, plain, (r1_xxreal_0(np__6561,np__6561))).#trainneg
cnf(i_0_460, plain, (k3_xcmplx_0(np__2,np__2)=np__4)).#trainneg
cnf(i_0_560, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_83, plain, (v2_xxreal_0(np__6168))).#trainneg
cnf(i_0_87, plain, (v2_xxreal_0(np__135))).#trainneg
cnf(i_0_571, plain, (~r1_xxreal_0(np__256,np__255))).#trainneg
cnf(i_0_549, plain, (~r1_xxreal_0(np__6,np__4))).#trainneg
cnf(i_0_573, plain, (~r1_xxreal_0(np__6,np__2))).#trainneg
cnf(i_0_543, plain, (~r1_xxreal_0(np__6,np__5))).#trainneg
cnf(i_0_508, plain, (~r1_xxreal_0(np__2,np__1))).#trainneg
cnf(i_0_362, plain, (~r1_xxreal_0(np__2,np__0))).#trainneg
cnf(i_0_345, plain, (~r1_xxreal_0(np__25,np__0))).#trainneg
cnf(i_0_401, plain, (~r1_xxreal_0(np__6560,np__1))).#trainneg
cnf(i_0_391, plain, (X1=k6_numbers|k2_xcmplx_0(k3_xcmplx_0(X1,esk2_3(X2,X1,X3)),X3)=X2|X3!=k2_nat_d(X2,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_92, plain, (k3_xcmplx_0(np__25,np__257)=np__6425)).#trainneg
cnf(i_0_112, plain, (m1_subset_1(np__6560,k5_numbers))).#trainneg
cnf(i_0_139, plain, (m1_subset_1(np__21,k1_numbers))).#trainneg
cnf(i_0_447, plain, (k2_xcmplx_0(np__1,np__5)=np__6)).#trainneg
cnf(i_0_326, plain, (~r1_xxreal_0(np__255,np__81))).#trainneg
cnf(i_0_91, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_26, plain, (m1_subset_1(np__6425,k1_numbers))).#trainneg
cnf(i_0_409, plain, (k2_xcmplx_0(np__2,np__5)=np__7)).#trainneg
cnf(i_0_374, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_138, plain, (v2_xxreal_0(np__25))).#trainneg
cnf(i_0_102, plain, (v2_xxreal_0(np__24))).#trainneg
cnf(i_0_369, plain, (~r1_xxreal_0(np__6560,np__2))).#trainneg
cnf(i_0_310, plain, (~r1_xxreal_0(np__6560,np__0))).#trainneg
cnf(i_0_495, plain, (~r1_xxreal_0(np__5,np__4))).#trainneg
cnf(i_0_563, plain, (~r1_xxreal_0(np__5,np__2))).#trainneg
cnf(i_0_565, plain, (~r1_xxreal_0(np__5,np__0))).#trainneg
cnf(i_0_584, plain, (~r1_xxreal_0(np__7,np__4))).#trainneg
cnf(i_0_546, plain, (~r1_xxreal_0(np__7,np__6))).#trainneg
cnf(i_0_531, plain, (~r1_xxreal_0(np__7,np__2))).#trainneg
cnf(i_0_294, plain, (v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_479, plain, (k2_xcmplx_0(np__6561,np__0)=np__6561)).#trainneg
cnf(i_0_578, plain, (v3_xxreal_0(X1)|v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5, plain, (m2_subset_1(np__3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_492, plain, (k2_xcmplx_0(np__2,np__4)=np__6)).#trainneg
cnf(i_0_366, plain, (m1_subset_1(np__6305,k5_numbers))).#trainneg
cnf(i_0_580, plain, (k4_xcmplx_0(k4_xcmplx_0(np__4))=np__4)).#trainneg
cnf(i_0_352, plain, (k2_xcmplx_0(np__4,np__1)=np__5)).#trainneg
cnf(i_0_65, plain, (m1_subset_1(np__6536,k5_numbers))).#trainneg
cnf(i_0_323, plain, (k2_xcmplx_0(np__2,np__3)=np__5)).#trainneg
cnf(i_0_106, plain, (v2_xxreal_0(np__78))).#trainneg
cnf(i_0_110, plain, (v2_xxreal_0(np__77))).#trainneg
cnf(i_0_114, plain, (v2_xxreal_0(np__6560))).#trainneg
cnf(i_0_528, plain, (~r1_xxreal_0(np__7,np__0))).#trainneg
cnf(i_0_566, plain, (~r1_xxreal_0(np__255,np__4))).#trainneg
cnf(i_0_564, plain, (~r1_xxreal_0(np__255,np__1))).#trainneg
cnf(i_0_529, plain, (~r1_xxreal_0(np__255,np__2))).#trainneg
cnf(i_0_478, plain, (~r1_xxreal_0(np__255,np__0))).#trainneg
cnf(i_0_118, plain, (v2_xxreal_0(np__6559))).#trainneg
cnf(i_0_122, plain, (v2_xxreal_0(np__6423))).#trainneg
cnf(i_0_126, plain, (v2_xxreal_0(np__6422))).#trainneg
cnf(i_0_179, plain, (m1_subset_1(np__6557,k5_numbers))).#trainneg
cnf(i_0_218, plain, (m1_subset_1(np__6558,k5_numbers))).#trainneg
cnf(i_0_354, plain, (m1_subset_1(np__6480,k5_numbers))).#trainneg
cnf(i_0_291, plain, (v1_xreal_0(esk1_0))).#trainneg
cnf(i_0_575, plain, (k6_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_108, plain, (m1_subset_1(np__77,k5_numbers))).#trainneg
cnf(i_0_412, plain, (k2_xcmplx_0(np__6,np__1)=np__7)).#trainneg
cnf(i_0_206, plain, (m1_subset_1(np__6304,k5_numbers))).#trainneg
cnf(i_0_202, plain, (m1_subset_1(np__80,k5_numbers))).#trainneg
cnf(i_0_31, plain, (m1_subset_1(np__79,k5_numbers))).#trainneg
cnf(i_0_148, plain, (m1_subset_1(np__175,k5_numbers))).#trainneg
cnf(i_0_130, plain, (v2_xxreal_0(np__6421))).#trainneg
cnf(i_0_134, plain, (v2_xxreal_0(np__6169))).#trainneg
cnf(i_0_142, plain, (v2_xxreal_0(np__21))).#trainneg
cnf(i_0_146, plain, (v2_xxreal_0(np__176))).#trainneg
cnf(i_0_150, plain, (v2_xxreal_0(np__175))).#trainneg
cnf(i_0_154, plain, (v2_xxreal_0(np__133))).#trainneg
cnf(i_0_158, plain, (v2_xxreal_0(np__132))).#trainneg
cnf(i_0_162, plain, (v2_xxreal_0(np__121))).#trainneg
cnf(i_0_166, plain, (v2_xxreal_0(np__111))).#trainneg
cnf(i_0_170, plain, (v2_xxreal_0(np__22))).#trainneg
cnf(i_0_177, plain, (v2_xxreal_0(np__6344))).#trainneg
cnf(i_0_348, plain, (m1_subset_1(np__254,k5_numbers))).#trainneg
cnf(i_0_104, plain, (m1_subset_1(np__78,k5_numbers))).#trainneg
cnf(i_0_100, plain, (m1_subset_1(np__24,k5_numbers))).#trainneg
cnf(i_0_297, plain, (k2_xcmplx_0(np__1,np__22)=np__23)).#trainneg
cnf(i_0_246, plain, (m1_subset_1(np__5,k1_numbers))).#trainneg
cnf(i_0_480, plain, (k2_xcmplx_0(np__0,np__6561)=np__6561)).#trainneg
cnf(i_0_54, plain, (m1_subset_1(np__23,k5_numbers))).#trainneg
cnf(i_0_116, plain, (m1_subset_1(np__6559,k5_numbers))).#trainneg
cnf(i_0_181, plain, (v2_xxreal_0(np__6557))).#trainneg
cnf(i_0_185, plain, (v2_xxreal_0(np__6400))).#trainneg
cnf(i_0_190, plain, (v2_xxreal_0(np__232))).#trainneg
cnf(i_0_194, plain, (v2_xxreal_0(np__134))).#trainneg
cnf(i_0_204, plain, (v2_xxreal_0(np__80))).#trainneg
cnf(i_0_208, plain, (v2_xxreal_0(np__6304))).#trainneg
cnf(i_0_215, plain, (v2_xxreal_0(np__6424))).#trainneg
cnf(i_0_220, plain, (v2_xxreal_0(np__6558))).#trainneg
cnf(i_0_236, plain, (v2_xxreal_0(np__252))).#trainneg
cnf(i_0_249, plain, (v2_xxreal_0(np__5))).#trainneg
cnf(i_0_264, plain, (v2_xxreal_0(np__7))).#trainneg
cnf(i_0_340, plain, (m1_subset_1(np__255,k1_numbers))).#trainneg
cnf(i_0_167, plain, (m1_subset_1(np__22,k1_numbers))).#trainneg
cnf(i_0_268, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_44, plain, (m1_subset_1(np__231,k5_numbers))).#trainneg
cnf(i_0_50, plain, (m1_subset_1(np__6289,k5_numbers))).#trainneg
cnf(i_0_60, plain, (m1_subset_1(np__56,k5_numbers))).#trainneg
cnf(i_0_85, plain, (m1_subset_1(np__135,k5_numbers))).#trainneg
cnf(i_0_120, plain, (m1_subset_1(np__6423,k5_numbers))).#trainneg
cnf(i_0_124, plain, (m1_subset_1(np__6422,k5_numbers))).#trainneg
cnf(i_0_368, plain, (v2_xxreal_0(np__6305))).#trainneg
cnf(i_0_290, plain, (v2_xxreal_0(np__55))).#trainneg
cnf(i_0_302, plain, (v2_xxreal_0(np__120))).#trainneg
cnf(i_0_343, plain, (v2_xxreal_0(np__255))).#trainneg
cnf(i_0_330, plain, (v2_xxreal_0(np__253))).#trainneg
cnf(i_0_350, plain, (v2_xxreal_0(np__254))).#trainneg
cnf(i_0_356, plain, (v2_xxreal_0(np__6480))).#trainneg
cnf(i_0_197, plain, (m1_subset_1(np__6561,k1_numbers))).#trainneg
cnf(i_0_30, plain, (m1_subset_1(np__79,k1_numbers))).#trainneg
cnf(i_0_43, plain, (m1_subset_1(np__231,k1_numbers))).#trainneg
cnf(i_0_49, plain, (m1_subset_1(np__6289,k1_numbers))).#trainneg
cnf(i_0_128, plain, (m1_subset_1(np__6421,k5_numbers))).#trainneg
cnf(i_0_132, plain, (m1_subset_1(np__6169,k5_numbers))).#trainneg
cnf(i_0_144, plain, (m1_subset_1(np__176,k5_numbers))).#trainneg
cnf(i_0_156, plain, (m1_subset_1(np__132,k5_numbers))).#trainneg
cnf(i_0_160, plain, (m1_subset_1(np__121,k5_numbers))).#trainneg
cnf(i_0_164, plain, (m1_subset_1(np__111,k5_numbers))).#trainneg
cnf(i_0_175, plain, (m1_subset_1(np__6344,k5_numbers))).#trainneg
cnf(i_0_183, plain, (m1_subset_1(np__6400,k5_numbers))).#trainneg
cnf(i_0_188, plain, (m1_subset_1(np__232,k5_numbers))).#trainneg
cnf(i_0_192, plain, (m1_subset_1(np__134,k5_numbers))).#trainneg
cnf(i_0_213, plain, (m1_subset_1(np__6424,k5_numbers))).#trainneg
cnf(i_0_270, plain, (k3_xcmplx_0(np__2,np__3)=np__6)).#trainneg
cnf(i_0_519, plain, (k6_xcmplx_0(np__81,np__3)=np__78)).#trainneg
cnf(i_0_588, plain, (k6_xcmplx_0(np__257,np__3)=np__254)).#trainneg
cnf(i_0_53, plain, (m1_subset_1(np__23,k1_numbers))).#trainneg
cnf(i_0_59, plain, (m1_subset_1(np__56,k1_numbers))).#trainneg
cnf(i_0_64, plain, (m1_subset_1(np__6536,k1_numbers))).#trainneg
cnf(i_0_68, plain, (m1_subset_1(np__6,k1_numbers))).#trainneg
cnf(i_0_80, plain, (m1_subset_1(np__6168,k1_numbers))).#trainneg
cnf(i_0_234, plain, (m1_subset_1(np__252,k5_numbers))).#trainneg
cnf(i_0_288, plain, (m1_subset_1(np__55,k5_numbers))).#trainneg
cnf(i_0_300, plain, (m1_subset_1(np__120,k5_numbers))).#trainneg
cnf(i_0_328, plain, (m1_subset_1(np__253,k5_numbers))).#trainneg
cnf(i_0_483, plain, (k2_xcmplx_0(np__4,np__2)=np__6)).#trainneg
cnf(i_0_81, plain, (m1_subset_1(np__6168,k5_numbers))).#trainneg
cnf(i_0_337, plain, (k3_xcmplx_0(np__0,np__6561)=np__0)).#trainneg
cnf(i_0_336, plain, (k3_xcmplx_0(np__6561,np__0)=np__0)).#trainneg
cnf(i_0_269, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_111, plain, (m1_subset_1(np__6560,k1_numbers))).#trainneg
cnf(i_0_261, plain, (m1_subset_1(np__7,k1_numbers))).#trainneg
cnf(i_0_506, plain, (v5_finset_1(k1_zfmisc_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_518, plain, (k4_xcmplx_0(k4_xcmplx_0(np__23))=np__23)).#trainneg
cnf(i_0_554, plain, (k4_xcmplx_0(k4_xcmplx_0(np__6))=np__6)).#trainneg
cnf(i_0_514, plain, (k6_xcmplx_0(np__6561,np__3)=np__6558)).#trainneg
cnf(i_0_278, plain, (k3_xcmplx_0(np__6,np__4)=np__24)).#trainneg
cnf(i_0_360, plain, (k3_xcmplx_0(np__0,np__257)=np__0)).#trainneg
cnf(i_0_84, plain, (m1_subset_1(np__135,k1_numbers))).#trainneg
cnf(i_0_99, plain, (m1_subset_1(np__24,k1_numbers))).#trainneg
cnf(i_0_103, plain, (m1_subset_1(np__78,k1_numbers))).#trainneg
cnf(i_0_107, plain, (m1_subset_1(np__77,k1_numbers))).#trainneg
cnf(i_0_115, plain, (m1_subset_1(np__6559,k1_numbers))).#trainneg
cnf(i_0_494, plain, (k6_xcmplx_0(np__1,np__4)=k4_xcmplx_0(np__3))).#trainneg
cnf(i_0_74, plain, (m2_subset_1(np__8,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_527, plain, (k6_xcmplx_0(np__3,np__6)=k4_xcmplx_0(np__3))).#trainneg
cnf(i_0_582, plain, (m2_subset_1(k2_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_448, plain, (k2_xcmplx_0(np__5,np__1)=np__6)).#trainneg
cnf(i_0_562, plain, (r1_xxreal_0(np__6560,np__6561))).#trainneg
cnf(i_0_205, plain, (m1_subset_1(np__6304,k1_numbers))).#trainneg
cnf(i_0_217, plain, (m1_subset_1(np__6558,k1_numbers))).#trainneg
cnf(i_0_78, plain, (m2_subset_1(np__136,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_347, plain, (m1_subset_1(np__254,k1_numbers))).#trainneg
cnf(i_0_510, plain, (k6_xcmplx_0(np__81,np__4)=np__77)).#trainneg
cnf(i_0_556, plain, (k6_xcmplx_0(np__81,np__1)=np__80)).#trainneg
cnf(i_0_468, plain, (k6_xcmplx_0(np__81,np__2)=np__79)).#trainneg
cnf(i_0_587, plain, (k6_xcmplx_0(np__6561,np__4)=np__6557)).#trainneg
cnf(i_0_520, plain, (k6_xcmplx_0(np__6561,np__81)=np__6480)).#trainneg
cnf(i_0_119, plain, (m1_subset_1(np__6423,k1_numbers))).#trainneg
cnf(i_0_123, plain, (m1_subset_1(np__6422,k1_numbers))).#trainneg
cnf(i_0_127, plain, (m1_subset_1(np__6421,k1_numbers))).#trainneg
cnf(i_0_131, plain, (m1_subset_1(np__6169,k1_numbers))).#trainneg
cnf(i_0_143, plain, (m1_subset_1(np__176,k1_numbers))).#trainneg
cnf(i_0_365, plain, (m1_subset_1(np__6305,k1_numbers))).#trainneg
cnf(i_0_284, plain, (v1_xcmplx_0(k1_newton(X1,X2))|~v7_ordinal1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_507, plain, (v1_int_1(k6_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_178, plain, (m1_subset_1(np__6557,k1_numbers))).#trainneg
cnf(i_0_353, plain, (m1_subset_1(np__6480,k1_numbers))).#trainneg
cnf(i_0_201, plain, (m1_subset_1(np__80,k1_numbers))).#trainneg
cnf(i_0_147, plain, (m1_subset_1(np__175,k1_numbers))).#trainneg
cnf(i_0_13, plain, (m2_subset_1(np__257,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_576, plain, (k6_xcmplx_0(np__6561,np__257)=np__6304)).#trainneg
cnf(i_0_90, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_487, plain, (v1_xreal_0(k1_newton(X1,X2))|~v7_ordinal1(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_574, plain, (k6_xcmplx_0(np__6561,np__256)=np__6305)).#trainneg
cnf(i_0_540, plain, (k6_xcmplx_0(np__6561,np__2)=np__6559)).#trainneg
cnf(i_0_488, plain, (k6_xcmplx_0(np__6561,np__25)=np__6536)).#trainneg
cnf(i_0_499, plain, (k6_xcmplx_0(np__256,np__81)=np__175)).#trainneg
cnf(i_0_151, plain, (m1_subset_1(np__133,k1_numbers))).#trainneg
cnf(i_0_152, plain, (m1_subset_1(np__133,k5_numbers))).#trainneg
cnf(i_0_155, plain, (m1_subset_1(np__132,k1_numbers))).#trainneg
cnf(i_0_159, plain, (m1_subset_1(np__121,k1_numbers))).#trainneg
cnf(i_0_163, plain, (m1_subset_1(np__111,k1_numbers))).#trainneg
cnf(i_0_174, plain, (m1_subset_1(np__6344,k1_numbers))).#trainneg
cnf(i_0_182, plain, (m1_subset_1(np__6400,k1_numbers))).#trainneg
cnf(i_0_187, plain, (m1_subset_1(np__232,k1_numbers))).#trainneg
cnf(i_0_191, plain, (m1_subset_1(np__134,k1_numbers))).#trainneg
cnf(i_0_212, plain, (m1_subset_1(np__6424,k1_numbers))).#trainneg
cnf(i_0_233, plain, (m1_subset_1(np__252,k1_numbers))).#trainneg
cnf(i_0_287, plain, (m1_subset_1(np__55,k1_numbers))).#trainneg
cnf(i_0_299, plain, (m1_subset_1(np__120,k1_numbers))).#trainneg
cnf(i_0_327, plain, (m1_subset_1(np__253,k1_numbers))).#trainneg
cnf(i_0_137, plain, (m2_subset_1(np__25,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_24, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_497, plain, (k6_xcmplx_0(np__2,np__5)=k4_xcmplx_0(np__3))).#trainneg
cnf(i_0_567, plain, (k6_xcmplx_0(np__257,np__6425)=k4_xcmplx_0(np__6168))).#trainneg
cnf(i_0_585, plain, (k6_xcmplx_0(np__2,np__6561)=k4_xcmplx_0(np__6559))).#trainneg
cnf(i_0_97, plain, (m2_subset_1(np__4,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_20, plain, (m2_subset_1(np__81,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_199, plain, (m2_subset_1(np__6561,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_28, plain, (m2_subset_1(np__6425,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_32, plain, (m2_subset_1(np__79,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_40, plain, (m2_subset_1(np__256,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_45, plain, (m2_subset_1(np__231,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_173, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_322, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_517, plain, (v1_int_1(k6_int_1(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_504, plain, (m1_subset_1(k2_newton(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_363, plain, (k7_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_475, plain, (v1_int_1(k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_470, plain, (k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k7_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_263, plain, (m2_subset_1(np__7,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_51, plain, (m2_subset_1(np__6289,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_545, plain, (v1_int_1(k3_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_55, plain, (m2_subset_1(np__23,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_61, plain, (m2_subset_1(np__56,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_66, plain, (m2_subset_1(np__6536,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_70, plain, (m2_subset_1(np__6,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_82, plain, (m2_subset_1(np__6168,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_86, plain, (m2_subset_1(np__135,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_101, plain, (m2_subset_1(np__24,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_105, plain, (m2_subset_1(np__78,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_248, plain, (m2_subset_1(np__5,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_394, plain, (X1=k6_numbers|v7_ordinal1(esk2_3(X2,X3,X1))|X1!=k2_nat_d(X2,X3)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_393, plain, (X1=k6_numbers|v7_ordinal1(esk2_3(X2,X1,X3))|X3!=k2_nat_d(X2,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_266, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_342, plain, (m2_subset_1(np__255,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_169, plain, (m2_subset_1(np__22,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_219, plain, (m2_subset_1(np__6558,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_367, plain, (m2_subset_1(np__6305,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_113, plain, (m2_subset_1(np__6560,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_349, plain, (m2_subset_1(np__254,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_355, plain, (m2_subset_1(np__6480,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_444, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_109, plain, (m2_subset_1(np__77,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_117, plain, (m2_subset_1(np__6559,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_121, plain, (m2_subset_1(np__6423,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_125, plain, (m2_subset_1(np__6422,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_129, plain, (m2_subset_1(np__6421,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_133, plain, (m2_subset_1(np__6169,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_141, plain, (m2_subset_1(np__21,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_145, plain, (m2_subset_1(np__176,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_149, plain, (m2_subset_1(np__175,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_153, plain, (m2_subset_1(np__133,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_207, plain, (m2_subset_1(np__6304,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_180, plain, (m2_subset_1(np__6557,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_203, plain, (m2_subset_1(np__80,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_157, plain, (m2_subset_1(np__132,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_161, plain, (m2_subset_1(np__121,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_165, plain, (m2_subset_1(np__111,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_176, plain, (m2_subset_1(np__6344,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_184, plain, (m2_subset_1(np__6400,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_189, plain, (m2_subset_1(np__232,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_193, plain, (m2_subset_1(np__134,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_214, plain, (m2_subset_1(np__6424,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_235, plain, (m2_subset_1(np__252,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_289, plain, (m2_subset_1(np__55,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_301, plain, (m2_subset_1(np__120,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_329, plain, (m2_subset_1(np__253,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1421, plain, (~v7_ordinal1(k5_numbers))).#trainneg
cnf(i_0_1460, plain, (~m1_subset_1(k5_numbers,k5_numbers))).#trainneg
cnf(i_0_1506, plain, (~v1_xboole_0(esk4_0)|~v2_xxreal_0(esk4_0))).#trainneg
cnf(i_0_1538, plain, (~r2_hidden(X1,k5_numbers)|~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_1426, plain, (v7_ordinal1(np__257)|~v7_ordinal1(np__1)|~v7_ordinal1(np__256))).#trainneg
cnf(i_0_1553, plain, (X1=esk4_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1419, plain, (v3_ordinal1(np__8))).#trainneg
cnf(i_0_1507, plain, (~v1_xboole_0(esk1_0)|~v2_xxreal_0(esk1_0))).#trainneg
cnf(i_0_1441, plain, (~v1_abian(np__3)|~v1_abian(np__2)|~v1_int_1(np__2))).#trainneg
cnf(i_0_1442, plain, (~v1_abian(np__257)|~v1_abian(np__256)|~v1_int_1(np__256))).#trainneg
cnf(i_0_1443, plain, (~v1_abian(np__2)|~v1_abian(np__1)|~v1_int_1(np__1))).#trainneg
cnf(i_0_1444, plain, (~v1_abian(np__4)|~v1_abian(np__3)|~v1_int_1(np__3))).#trainneg
cnf(i_0_1570, plain, (v7_ordinal1(esk1_0))).#trainneg
cnf(i_0_1425, plain, (v7_ordinal1(np__3)|~v7_ordinal1(np__1)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_1540, plain, (v7_ordinal1(esk4_0))).#trainneg
cnf(i_0_1515, plain, (v7_ordinal1(np__6561)|~v7_ordinal1(np__6425)|~v7_ordinal1(np__136))).#trainneg
cnf(i_0_1495, plain, (v7_ordinal1(np__6)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_1439, plain, (k1_nat_1(np__1,k2_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(k2_newton(np__2,k2_newton(np__2,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1393, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1400, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1418, plain, (v3_ordinal1(k4_pepin(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1554, plain, (X1=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1389, plain, (k4_nat_d(k4_nat_1(np__81,k4_nat_d(X1,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),X1),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1413, plain, (v3_ordinal1(np__2))).#trainneg
cnf(i_0_1469, plain, (~v1_abian(np__8)|~v1_abian(np__7)|~v1_int_1(np__7))).#trainneg
cnf(i_0_1586, plain, (~v1_abian(np__5)|~v1_abian(np__4)|~v1_int_1(np__4))).#trainneg
cnf(i_0_1597, plain, (~v1_abian(np__7)|~v1_abian(np__6)|~v1_int_1(np__6))).#trainneg
cnf(i_0_1728, plain, (~v1_abian(np__6)|~v1_abian(np__5)|~v1_int_1(np__5))).#trainneg
cnf(i_0_1435, plain, (v7_ordinal1(np__4)|~v7_ordinal1(np__1)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_1566, plain, (v7_ordinal1(np__6)|~v7_ordinal1(np__5)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1424, plain, (v7_ordinal1(np__8)|~v7_ordinal1(np__6)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_1433, plain, (v7_ordinal1(np__8)|~v7_ordinal1(np__5)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_1508, plain, (~v1_xboole_0(X1)|~v7_ordinal1(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_1572, plain, (X1=esk1_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1391, plain, (k4_nat_d(k1_newton(np__3,np__4),np__257)=np__81|~m1_subset_1(np__4,k5_numbers))).#trainneg
cnf(i_0_1436, plain, (v3_ordinal1(np__4))).#trainneg
cnf(i_0_1449, plain, (v3_ordinal1(np__6425))).#trainneg
cnf(i_0_1452, plain, (v3_ordinal1(np__0))).#trainneg
cnf(i_0_1456, plain, (v3_ordinal1(np__6561))).#trainneg
cnf(i_0_1472, plain, (v3_ordinal1(np__256))).#trainneg
cnf(i_0_1485, plain, (v3_ordinal1(np__21))).#trainneg
cnf(i_0_1498, plain, (v3_ordinal1(np__5))).#trainneg
cnf(i_0_1510, plain, (v3_ordinal1(np__22))).#trainneg
cnf(i_0_1582, plain, (v7_ordinal1(np__6)|~v7_ordinal1(np__4)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_1503, plain, (v7_ordinal1(np__5)|~v7_ordinal1(np__2)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_1427, plain, (v7_ordinal1(np__8)|~v7_ordinal1(np__7)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1489, plain, (v7_ordinal1(np__8)|~v7_ordinal1(np__4)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_1432, plain, (v7_ordinal1(np__8)|~v7_ordinal1(np__4))).#trainneg
cnf(i_0_1409, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1699, plain, (v1_xcmplx_0(np__0)|~v1_int_1(np__0))).#trainneg
cnf(i_0_1511, plain, (v3_ordinal1(np__6))).#trainneg
cnf(i_0_1524, plain, (v3_ordinal1(np__7))).#trainneg
cnf(i_0_1525, plain, (v3_ordinal1(np__255))).#trainneg
cnf(i_0_1564, plain, (v3_ordinal1(np__6560))).#trainneg
cnf(i_0_1583, plain, (v3_ordinal1(np__6305))).#trainneg
cnf(i_0_1587, plain, (v3_ordinal1(np__6536))).#trainneg
cnf(i_0_1591, plain, (v3_ordinal1(np__6557))).#trainneg
cnf(i_0_1592, plain, (v3_ordinal1(np__6558))).#trainneg
cnf(i_0_1547, plain, (v7_ordinal1(np__7)|~v7_ordinal1(np__6)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1394, plain, (k2_xcmplx_0(k2_newton(np__2,k2_newton(np__2,X1)),np__1)=k4_pepin(X1)|~v7_ordinal1(k2_newton(np__2,k2_newton(np__2,X1)))|~v7_ordinal1(X1)|~m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_1518, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1406, plain, (k1_nat_1(k1_newton(np__2,k2_newton(np__2,X1)),np__1)=k4_pepin(X1)|~v7_ordinal1(k2_newton(np__2,X1))|~v7_ordinal1(X1)|~m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_1512, plain, (v7_ordinal1(np__6425)|~v7_ordinal1(np__25)|~v7_ordinal1(np__257))).#trainneg
cnf(i_0_1568, plain, (v7_ordinal1(np__7)|~v7_ordinal1(np__5)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_1497, plain, (v7_ordinal1(np__7)|~v7_ordinal1(np__4)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_1407, plain, (k1_nat_1(k2_newton(np__2,k1_newton(np__2,X1)),np__1)=k4_pepin(X1)|~v7_ordinal1(X1)|~m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_1549, plain, (v7_ordinal1(np__5)|~v7_ordinal1(np__4)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1853, plain, (v1_int_1(k4_nat_d(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1437, plain, (k1_nat_1(np__1,k2_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(np__1)|~v7_ordinal1(X1)|~m1_subset_1(k2_newton(np__2,k2_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_1593, plain, (v3_ordinal1(np__6480))).#trainneg
cnf(i_0_1595, plain, (v3_ordinal1(np__77))).#trainneg
cnf(i_0_1598, plain, (v3_ordinal1(np__6304))).#trainneg
cnf(i_0_1599, plain, (v3_ordinal1(np__80))).#trainneg
cnf(i_0_1600, plain, (v3_ordinal1(np__79))).#trainneg
cnf(i_0_1601, plain, (v3_ordinal1(np__175))).#trainneg
cnf(i_0_1602, plain, (v3_ordinal1(np__254))).#trainneg
cnf(i_0_1603, plain, (v3_ordinal1(np__78))).#trainneg
cnf(i_0_1519, plain, (m2_subset_1(k1_nat_1(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1453, plain, (m2_subset_1(k1_newton(X1,X2),k1_numbers,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1983, plain, (v1_xboole_0(k1_numbers)|m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1614, plain, (m2_subset_1(k4_nat_1(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1851, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1446, plain, (X1=k2_nat_d(X2,np__1)|r1_xxreal_0(np__1,X1)|X2!=k2_xcmplx_0(np__81,X1)|~v7_ordinal1(np__1)|~v7_ordinal1(np__81)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1604, plain, (v3_ordinal1(np__24))).#trainneg
cnf(i_0_1610, plain, (v3_ordinal1(np__23))).#trainneg
cnf(i_0_1612, plain, (v3_ordinal1(np__6559))).#trainneg
cnf(i_0_1642, plain, (v3_ordinal1(np__231))).#trainneg
cnf(i_0_1643, plain, (v3_ordinal1(np__6289))).#trainneg
cnf(i_0_1644, plain, (v3_ordinal1(np__56))).#trainneg
cnf(i_0_1645, plain, (v3_ordinal1(np__135))).#trainneg
cnf(i_0_1646, plain, (v3_ordinal1(np__6423))).#trainneg
cnf(i_0_1557, plain, (X1=k2_nat_d(X2,k1_newton(X3,X4))|r1_xxreal_0(k1_newton(X3,X4),X1)|X2!=k2_xcmplx_0(k1_newton(X3,k2_xcmplx_0(X4,X5)),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,X4))|~v7_ordinal1(k1_newton(X3,X5))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_1493, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1718, plain, (v7_ordinal1(np__24)|~v7_ordinal1(np__4)|~v7_ordinal1(np__6))).#trainneg
cnf(i_0_1513, plain, (X1=k2_nat_d(X2,np__257)|r1_xxreal_0(np__257,X1)|X2!=k2_xcmplx_0(np__6425,X1)|~v7_ordinal1(np__257)|~v7_ordinal1(np__25)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1535, plain, (k3_xcmplx_0(np__1,esk1_0)=esk1_0)).#trainneg
cnf(i_0_1551, plain, (X1=k2_nat_d(X2,np__2)|r1_xxreal_0(np__2,X1)|X2!=k2_xcmplx_0(np__4,X1)|~v7_ordinal1(np__2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1678, plain, (X1=k2_nat_d(X2,np__2)|r1_xxreal_0(np__2,X1)|X2!=k2_xcmplx_0(np__6,X1)|~v7_ordinal1(np__2)|~v7_ordinal1(np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1590, plain, (X1=k2_nat_d(X2,np__81)|r1_xxreal_0(np__81,X1)|X2!=k2_xcmplx_0(np__6561,X1)|~v7_ordinal1(np__81)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1607, plain, (X1=k2_nat_d(X2,np__257)|r1_xxreal_0(np__257,X1)|X2!=k2_xcmplx_0(np__0,X1)|~v7_ordinal1(np__257)|~v7_ordinal1(np__0)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1912, plain, (v1_int_1(np__2)|~v1_int_1(np__1))).#trainneg
cnf(i_0_1451, plain, (X1=k2_nat_d(X2,np__2)|r1_xxreal_0(np__2,X1)|X2!=k2_xcmplx_0(np__8,X1)|~v7_ordinal1(np__2)|~v7_ordinal1(np__4)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1647, plain, (v3_ordinal1(np__6422))).#trainneg
cnf(i_0_1652, plain, (v3_ordinal1(np__6421))).#trainneg
cnf(i_0_1653, plain, (v3_ordinal1(np__6169))).#trainneg
cnf(i_0_1654, plain, (v3_ordinal1(np__176))).#trainneg
cnf(i_0_1655, plain, (v3_ordinal1(np__132))).#trainneg
cnf(i_0_1656, plain, (v3_ordinal1(np__121))).#trainneg
cnf(i_0_1657, plain, (v3_ordinal1(np__111))).#trainneg
cnf(i_0_1658, plain, (v3_ordinal1(np__6344))).#trainneg
cnf(i_0_1539, plain, (v7_ordinal1(np__22)|~v7_ordinal1(np__21)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1605, plain, (v7_ordinal1(np__23)|~v7_ordinal1(np__22)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1483, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_1935, plain, (X1=k2_nat_d(X2,X3)|r1_xxreal_0(X3,X1)|X2!=k2_xcmplx_0(k4_xcmplx_0(X3),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k4_xcmplx_0(np__1))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1981, plain, (m2_subset_1(k2_nat_1(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1757, plain, (k2_xcmplx_0(X1,k1_newton(X2,X3))=k2_xcmplx_0(k1_newton(X2,X3),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1886, plain, (m1_subset_1(k1_newton(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1694, plain, (X1=k2_nat_d(X2,np__6561)|r1_xxreal_0(np__6561,X1)|X2!=k2_xcmplx_0(np__0,X1)|~v7_ordinal1(np__6561)|~v7_ordinal1(np__0)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1463, plain, (X1=k2_nat_d(X2,np__81)|r1_xxreal_0(np__81,X1)|X2!=k2_xcmplx_0(np__81,X1)|~v7_ordinal1(np__81)|~v7_ordinal1(np__1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1467, plain, (X1=k2_nat_d(X2,np__4)|r1_xxreal_0(np__4,X1)|X2!=k2_xcmplx_0(np__8,X1)|~v7_ordinal1(np__4)|~v7_ordinal1(np__2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1563, plain, (X1=k2_nat_d(X2,np__25)|r1_xxreal_0(np__25,X1)|X2!=k2_xcmplx_0(np__6425,X1)|~v7_ordinal1(np__25)|~v7_ordinal1(np__257)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1973, plain, (v1_int_1(np__6561)|~v1_int_1(np__81))).#trainneg
cnf(i_0_1454, plain, (~v1_xreal_0(X1)|~v7_ordinal1(X1)|~v2_xxreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1659, plain, (v3_ordinal1(np__6400))).#trainneg
cnf(i_0_1660, plain, (v3_ordinal1(np__232))).#trainneg
cnf(i_0_1661, plain, (v3_ordinal1(np__134))).#trainneg
cnf(i_0_1662, plain, (v3_ordinal1(np__6424))).#trainneg
cnf(i_0_1685, plain, (v3_ordinal1(np__252))).#trainneg
cnf(i_0_1686, plain, (v3_ordinal1(np__55))).#trainneg
cnf(i_0_1687, plain, (v3_ordinal1(np__120))).#trainneg
cnf(i_0_1688, plain, (v3_ordinal1(np__253))).#trainneg
cnf(i_0_1627, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_1692, plain, (X1=k2_nat_d(X2,np__0)|r1_xxreal_0(np__0,X1)|X2!=k2_xcmplx_0(np__0,X1)|~v7_ordinal1(np__0)|~v7_ordinal1(np__6561)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1721, plain, (X1=k2_nat_d(X2,np__0)|r1_xxreal_0(np__0,X1)|X2!=k2_xcmplx_0(np__0,X1)|~v7_ordinal1(np__0)|~v7_ordinal1(np__257)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1782, plain, (k3_xcmplx_0(X1,k1_newton(X2,X3))=k3_xcmplx_0(k1_newton(X2,X3),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1980, plain, (k2_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1636, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1719, plain, (X1=k2_nat_d(X2,np__6)|r1_xxreal_0(np__6,X1)|X2!=k2_xcmplx_0(np__24,X1)|~v7_ordinal1(np__6)|~v7_ordinal1(np__4)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1992, plain, (X1=k2_nat_d(X2,k3_xcmplx_0(X3,X4))|r1_xxreal_0(k3_xcmplx_0(X3,X4),X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(X3,k3_xcmplx_0(X4,X5)),X1)|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v7_ordinal1(k3_xcmplx_0(X3,X4))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_1999, plain, (k3_xcmplx_0(np__2,k3_xcmplx_0(np__3,X1))=k3_xcmplx_0(np__6,X1)|~v1_xcmplx_0(np__3)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1697, plain, (k3_xcmplx_0(np__1,k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_1516, plain, (v1_abian(np__8)|~v1_int_1(np__4))).#trainneg
cnf(i_0_1736, plain, (v1_xreal_0(np__6425)|~v1_xreal_0(np__257)|~v1_xreal_0(np__25))).#trainneg
cnf(i_0_1690, plain, (v3_ordinal1(np__6168))).#trainneg
cnf(i_0_1789, plain, (v3_ordinal1(np__133))).#trainneg
cnf(i_0_1840, plain, (v1_xreal_0(np__8)|~v1_xreal_0(np__4))).#trainneg
cnf(i_0_2007, plain, (k3_xcmplx_0(np__4,k3_xcmplx_0(np__2,X1))=k3_xcmplx_0(np__8,X1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1998, plain, (k3_xcmplx_0(np__2,k3_xcmplx_0(np__4,X1))=k3_xcmplx_0(np__8,X1)|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1780, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k3_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1638, plain, (v1_xcmplx_0(np__2)|~v1_xcmplx_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1494, plain, (~v1_xxreal_0(X1)|~v1_xreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X1))|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_1679, plain, (v1_abian(np__6)|~v1_int_1(np__3))).#trainneg
cnf(i_0_1641, plain, (v1_xcmplx_0(np__3)|~v1_xcmplx_0(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_1943, plain, (k3_xcmplx_0(np__2,k7_xcmplx_0(np__3,X1))=k7_xcmplx_0(np__6,X1)|~v1_xcmplx_0(np__3)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1620, plain, (k3_xcmplx_0(np__1,k4_xcmplx_0(X1))=k4_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1854, plain, (v1_int_1(k2_nat_d(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1623, plain, (v1_xcmplx_0(np__2)|~v1_xreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1924, plain, (v1_int_1(np__8)|~v1_int_1(np__4))).#trainneg
cnf(i_0_1847, plain, (v1_xreal_0(np__6)|~v1_xreal_0(np__3))).#trainneg
cnf(i_0_1931, plain, (v1_int_1(np__6)|~v1_int_1(np__3))).#trainneg
cnf(i_0_1626, plain, (v1_xcmplx_0(np__3)|~v1_xreal_0(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_1552, plain, (v1_abian(np__4)|~v1_int_1(np__2))).#trainneg
cnf(i_0_1941, plain, (k3_xcmplx_0(np__2,k7_xcmplx_0(np__2,X1))=k7_xcmplx_0(np__4,X1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1879, plain, (v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1698, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_int_1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1621, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1932, plain, (v1_abian(k4_xcmplx_0(np__2))|~v1_int_1(k4_xcmplx_0(np__1))|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_1639, plain, (v1_xcmplx_0(np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1996, plain, (k3_xcmplx_0(np__25,k3_xcmplx_0(np__257,X1))=k3_xcmplx_0(np__6425,X1)|~v1_xcmplx_0(np__257)|~v1_xcmplx_0(np__25)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2004, plain, (k3_xcmplx_0(np__257,k3_xcmplx_0(np__25,X1))=k3_xcmplx_0(np__6425,X1)|~v1_xcmplx_0(np__25)|~v1_xcmplx_0(np__257)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1700, plain, (v1_xcmplx_0(np__2)|~v1_int_1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1938, plain, (k3_xcmplx_0(np__0,k7_xcmplx_0(np__257,X1))=k7_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__257)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1705, plain, (v1_xcmplx_0(np__3)|~v1_int_1(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_1764, plain, (v1_int_1(np__6425)|~v1_int_1(np__136)|~v1_int_1(np__6561))).#trainneg
cnf(i_0_1804, plain, (v1_xreal_0(np__8)|~v1_xreal_0(np__1)|~v1_xreal_0(np__7))).#trainneg
cnf(i_0_1632, plain, (v1_int_1(np__3)|~v1_int_1(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_1668, plain, (v1_xreal_0(np__3)|~v1_xreal_0(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_1950, plain, (k3_xcmplx_0(np__81,k7_xcmplx_0(np__81,X1))=k7_xcmplx_0(np__6561,X1)|~v1_xcmplx_0(np__81)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1951, plain, (k3_xcmplx_0(np__4,k7_xcmplx_0(np__2,X1))=k7_xcmplx_0(np__8,X1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1640, plain, (v1_xcmplx_0(np__4)|~v1_xcmplx_0(k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_1939, plain, (k3_xcmplx_0(k1_newton(X1,X2),k7_xcmplx_0(k1_newton(X1,X3),X4))=k7_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X2,X3)),X4)|~v1_xcmplx_0(k1_newton(X1,X3))|~v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1949, plain, (k3_xcmplx_0(np__81,k7_xcmplx_0(np__1,X1))=k7_xcmplx_0(np__81,X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__81)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2000, plain, (k3_xcmplx_0(np__6,k3_xcmplx_0(np__4,X1))=k3_xcmplx_0(np__24,X1)|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(np__6)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1937, plain, (k3_xcmplx_0(np__0,k7_xcmplx_0(np__6561,X1))=k7_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__6561)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1945, plain, (k3_xcmplx_0(np__1,k7_xcmplx_0(np__81,X1))=k7_xcmplx_0(np__81,X1)|~v1_xcmplx_0(np__81)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1942, plain, (k3_xcmplx_0(np__2,k7_xcmplx_0(np__4,X1))=k7_xcmplx_0(np__8,X1)|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1761, plain, (v1_int_1(np__136)|~v1_int_1(np__6425)|~v1_int_1(np__6561))).#trainneg
cnf(i_0_1820, plain, (v1_xreal_0(np__6561)|~v1_xreal_0(np__136)|~v1_xreal_0(np__6425))).#trainneg
cnf(i_0_1805, plain, (v1_xreal_0(np__6)|~v1_xreal_0(np__1)|~v1_xreal_0(np__5))).#trainneg
cnf(i_0_1624, plain, (v1_xcmplx_0(np__1)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1940, plain, (k3_xcmplx_0(np__25,k7_xcmplx_0(np__257,X1))=k7_xcmplx_0(np__6425,X1)|~v1_xcmplx_0(np__257)|~v1_xcmplx_0(np__25)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1779, plain, (k3_xcmplx_0(X1,k4_xcmplx_0(X2))=k3_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_int_1(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1946, plain, (k3_xcmplx_0(np__6561,k7_xcmplx_0(np__0,X1))=k7_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__6561)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1986, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,k4_xcmplx_0(np__1)))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1625, plain, (v1_xcmplx_0(np__4)|~v1_xreal_0(k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_1671, plain, (v1_xcmplx_0(np__6561)|~v1_xcmplx_0(k4_xcmplx_0(np__6561)))).#trainneg
cnf(i_0_1702, plain, (v1_xcmplx_0(np__1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1898, plain, (v1_int_1(np__3)|~v1_int_1(np__1)|~v1_int_1(np__2))).#trainneg
cnf(i_0_1806, plain, (v1_xreal_0(np__8)|~v1_xreal_0(np__3)|~v1_xreal_0(np__5))).#trainneg
cnf(i_0_1629, plain, (v1_int_1(np__2)|~v1_int_1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1737, plain, (v1_xreal_0(np__4)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_1947, plain, (k3_xcmplx_0(np__257,k7_xcmplx_0(np__0,X1))=k7_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__257)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1704, plain, (v1_xcmplx_0(np__4)|~v1_int_1(k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_1948, plain, (k3_xcmplx_0(np__257,k7_xcmplx_0(np__25,X1))=k7_xcmplx_0(np__6425,X1)|~v1_xcmplx_0(np__25)|~v1_xcmplx_0(np__257)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1759, plain, (v1_int_1(k4_xcmplx_0(np__1))|~v1_int_1(np__257)|~v1_int_1(np__256))).#trainneg
cnf(i_0_1903, plain, (v1_int_1(np__257)|~v1_int_1(np__1)|~v1_int_1(np__256))).#trainneg
cnf(i_0_1830, plain, (v1_xreal_0(np__4)|~v1_xreal_0(np__3)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_1813, plain, (v1_xreal_0(np__8)|~v1_xreal_0(np__6)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_1630, plain, (v1_int_1(np__1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1944, plain, (k3_xcmplx_0(np__6,k7_xcmplx_0(np__4,X1))=k7_xcmplx_0(np__24,X1)|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(np__6)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1672, plain, (v1_xcmplx_0(np__6561)|~v1_xreal_0(k4_xcmplx_0(np__6561)))).#trainneg
cnf(i_0_1716, plain, (v1_xcmplx_0(np__6)|~v1_xcmplx_0(k4_xcmplx_0(np__6)))).#trainneg
cnf(i_0_1665, plain, (v1_xreal_0(np__2)|~v1_xreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1896, plain, (v1_int_1(np__4)|~v1_int_1(np__2))).#trainneg
cnf(i_0_1904, plain, (v1_int_1(np__6561)|~v1_int_1(np__136)|~v1_int_1(np__6425))).#trainneg
cnf(i_0_1631, plain, (v1_int_1(np__4)|~v1_int_1(k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_1667, plain, (v1_xreal_0(np__4)|~v1_xreal_0(k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_1666, plain, (v1_xreal_0(np__1)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1952, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(k4_xcmplx_0(np__1),X2))=k7_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1802, plain, (v1_int_1(k4_xcmplx_0(np__3))|~v1_int_1(np__5)|~v1_int_1(np__2))).#trainneg
cnf(i_0_1816, plain, (v1_xreal_0(np__5)|~v1_xreal_0(np__3)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_1717, plain, (v1_xcmplx_0(np__6)|~v1_xreal_0(k4_xcmplx_0(np__6)))).#trainneg
cnf(i_0_1914, plain, (v1_int_1(np__4)|~v1_int_1(np__3)|~v1_int_1(np__1))).#trainneg
cnf(i_0_1669, plain, (v1_int_1(np__6561)|~v1_int_1(k4_xcmplx_0(np__6561)))).#trainneg
cnf(i_0_1670, plain, (v1_xreal_0(np__6561)|~v1_xreal_0(k4_xcmplx_0(np__6561)))).#trainneg
cnf(i_0_1856, plain, (v1_xreal_0(k4_xcmplx_0(np__3))|~v1_xreal_0(np__5)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_1815, plain, (v1_xreal_0(np__6)|~v1_xreal_0(np__4)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_1866, plain, (v1_xreal_0(np__6304)|~v1_xreal_0(np__257)|~v1_xreal_0(np__6561))).#trainneg
cnf(i_0_1703, plain, (v1_xcmplx_0(np__6561)|~v1_int_1(k4_xcmplx_0(np__6561)))).#trainneg
cnf(i_0_1963, plain, (v1_int_1(np__6425)|~v1_int_1(np__257)|~v1_int_1(np__25))).#trainneg
cnf(i_0_1897, plain, (v1_int_1(np__8)|~v1_int_1(np__6)|~v1_int_1(np__2))).#trainneg
cnf(i_0_1823, plain, (v1_xreal_0(np__23)|~v1_xreal_0(np__22)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_1673, plain, (v1_int_1(np__256)|~v1_int_1(k4_xcmplx_0(np__256)))).#trainneg
cnf(i_0_1675, plain, (v1_xcmplx_0(np__256)|~v1_xcmplx_0(k4_xcmplx_0(np__256)))).#trainneg
cnf(i_0_1676, plain, (v1_xcmplx_0(np__256)|~v1_xreal_0(k4_xcmplx_0(np__256)))).#trainneg
cnf(i_0_1701, plain, (v1_xcmplx_0(np__256)|~v1_int_1(k4_xcmplx_0(np__256)))).#trainneg
cnf(i_0_1708, plain, (v1_int_1(np__23)|~v1_int_1(k4_xcmplx_0(np__23)))).#trainneg
cnf(i_0_1709, plain, (v1_xcmplx_0(np__23)|~v1_int_1(k4_xcmplx_0(np__23)))).#trainneg
cnf(i_0_1710, plain, (v1_xreal_0(np__23)|~v1_xreal_0(k4_xcmplx_0(np__23)))).#trainneg
cnf(i_0_1711, plain, (v1_xcmplx_0(np__23)|~v1_xcmplx_0(k4_xcmplx_0(np__23)))).#trainneg
cnf(i_0_1760, plain, (v1_int_1(k4_xcmplx_0(np__3))|~v1_int_1(np__4)|~v1_int_1(np__1))).#trainneg
cnf(i_0_1861, plain, (v1_xreal_0(k4_xcmplx_0(np__3))|~v1_xreal_0(np__4)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_1818, plain, (v1_xreal_0(np__7)|~v1_xreal_0(np__1)|~v1_xreal_0(np__6))).#trainneg
cnf(i_0_1870, plain, (v1_xreal_0(np__6558)|~v1_xreal_0(np__3)|~v1_xreal_0(np__6561))).#trainneg
cnf(i_0_1829, plain, (v1_xreal_0(np__5)|~v1_xreal_0(np__4)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_1846, plain, (v1_xreal_0(np__7)|~v1_xreal_0(np__4)|~v1_xreal_0(np__3))).#trainneg
cnf(i_0_1811, plain, (v1_xreal_0(np__7)|~v1_xreal_0(np__5)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_1872, plain, (v1_xreal_0(np__254)|~v1_xreal_0(np__3)|~v1_xreal_0(np__257))).#trainneg
cnf(i_0_1824, plain, (v1_xreal_0(np__22)|~v1_xreal_0(np__21)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_1712, plain, (v1_xcmplx_0(np__23)|~v1_xreal_0(k4_xcmplx_0(np__23)))).#trainneg
cnf(i_0_1713, plain, (v1_int_1(np__6)|~v1_int_1(k4_xcmplx_0(np__6)))).#trainneg
cnf(i_0_1714, plain, (v1_xcmplx_0(np__6)|~v1_int_1(k4_xcmplx_0(np__6)))).#trainneg
cnf(i_0_1715, plain, (v1_xreal_0(np__6)|~v1_xreal_0(k4_xcmplx_0(np__6)))).#trainneg
cnf(i_0_1740, plain, (v1_xreal_0(np__24)|~v1_xreal_0(np__4)|~v1_xreal_0(np__6))).#trainneg
cnf(i_0_1762, plain, (v1_int_1(np__6480)|~v1_int_1(np__81)|~v1_int_1(np__6561))).#trainneg
cnf(i_0_1763, plain, (v1_int_1(np__6557)|~v1_int_1(np__4)|~v1_int_1(np__6561))).#trainneg
cnf(i_0_1765, plain, (v1_int_1(np__6558)|~v1_int_1(np__3)|~v1_int_1(np__6561))).#trainneg
cnf(i_0_1766, plain, (v1_int_1(np__254)|~v1_int_1(np__3)|~v1_int_1(np__257))).#trainneg
cnf(i_0_1864, plain, (v1_xreal_0(np__6305)|~v1_xreal_0(np__256)|~v1_xreal_0(np__6561))).#trainneg
cnf(i_0_1889, plain, (v1_int_1(np__6)|~v1_int_1(np__1)|~v1_int_1(np__5))).#trainneg
cnf(i_0_1899, plain, (v1_int_1(np__6)|~v1_int_1(np__4)|~v1_int_1(np__2))).#trainneg
cnf(i_0_1867, plain, (v1_xreal_0(np__6480)|~v1_xreal_0(np__81)|~v1_xreal_0(np__6561))).#trainneg
cnf(i_0_1862, plain, (v1_xreal_0(np__6536)|~v1_xreal_0(np__25)|~v1_xreal_0(np__6561))).#trainneg
cnf(i_0_1890, plain, (v1_int_1(np__8)|~v1_int_1(np__3)|~v1_int_1(np__5))).#trainneg
cnf(i_0_1876, plain, (v1_xreal_0(np__78)|~v1_xreal_0(np__3)|~v1_xreal_0(np__81))).#trainneg
cnf(i_0_1770, plain, (v1_int_1(np__78)|~v1_int_1(np__3)|~v1_int_1(np__81))).#trainneg
cnf(i_0_1888, plain, (v1_int_1(np__8)|~v1_int_1(np__1)|~v1_int_1(np__7))).#trainneg
cnf(i_0_1767, plain, (v1_int_1(np__79)|~v1_int_1(np__2)|~v1_int_1(np__81))).#trainneg
cnf(i_0_1907, plain, (v1_int_1(np__23)|~v1_int_1(np__22)|~v1_int_1(np__1))).#trainneg
cnf(i_0_1908, plain, (v1_int_1(np__22)|~v1_int_1(np__21)|~v1_int_1(np__1))).#trainneg
cnf(i_0_1768, plain, (v1_int_1(np__80)|~v1_int_1(np__1)|~v1_int_1(np__81))).#trainneg
cnf(i_0_1769, plain, (v1_int_1(np__77)|~v1_int_1(np__4)|~v1_int_1(np__81))).#trainneg
cnf(i_0_1777, plain, (v1_int_1(np__6304)|~v1_int_1(np__257)|~v1_int_1(np__6561))).#trainneg
cnf(i_0_1784, plain, (v1_int_1(np__6305)|~v1_int_1(np__256)|~v1_int_1(np__6561))).#trainneg
cnf(i_0_1785, plain, (v1_int_1(np__6559)|~v1_int_1(np__2)|~v1_int_1(np__6561))).#trainneg
cnf(i_0_1863, plain, (v1_xreal_0(np__6559)|~v1_xreal_0(np__2)|~v1_xreal_0(np__6561))).#trainneg
cnf(i_0_1902, plain, (v1_int_1(np__7)|~v1_int_1(np__1)|~v1_int_1(np__6))).#trainneg
cnf(i_0_1868, plain, (v1_xreal_0(np__6557)|~v1_xreal_0(np__4)|~v1_xreal_0(np__6561))).#trainneg
cnf(i_0_1900, plain, (v1_int_1(np__5)|~v1_int_1(np__3)|~v1_int_1(np__2))).#trainneg
cnf(i_0_1873, plain, (v1_xreal_0(np__79)|~v1_xreal_0(np__2)|~v1_xreal_0(np__81))).#trainneg
cnf(i_0_1895, plain, (v1_int_1(np__7)|~v1_int_1(np__5)|~v1_int_1(np__2))).#trainneg
cnf(i_0_1913, plain, (v1_int_1(np__5)|~v1_int_1(np__4)|~v1_int_1(np__1))).#trainneg
cnf(i_0_1859, plain, (v1_xreal_0(np__175)|~v1_xreal_0(np__81)|~v1_xreal_0(np__256))).#trainneg
cnf(i_0_1930, plain, (v1_int_1(np__7)|~v1_int_1(np__4)|~v1_int_1(np__3))).#trainneg
cnf(i_0_1786, plain, (v1_int_1(np__6536)|~v1_int_1(np__25)|~v1_int_1(np__6561))).#trainneg
cnf(i_0_1787, plain, (v1_int_1(np__175)|~v1_int_1(np__81)|~v1_int_1(np__256))).#trainneg
cnf(i_0_1874, plain, (v1_xreal_0(np__80)|~v1_xreal_0(np__1)|~v1_xreal_0(np__81))).#trainneg
cnf(i_0_1875, plain, (v1_xreal_0(np__77)|~v1_xreal_0(np__4)|~v1_xreal_0(np__81))).#trainneg
cnf(i_0_1967, plain, (v1_int_1(np__24)|~v1_int_1(np__4)|~v1_int_1(np__6))).#trainneg
cnf(i_0_1848, plain, (v1_int_1(k4_xcmplx_0(np__6168))|~v1_int_1(np__6425)|~v1_int_1(np__257))).#trainneg
cnf(i_0_1849, plain, (v1_int_1(k4_xcmplx_0(np__256))|~v1_int_1(np__257)|~v1_int_1(np__1))).#trainneg
cnf(i_0_1850, plain, (v1_int_1(k4_xcmplx_0(np__6559))|~v1_int_1(np__6561)|~v1_int_1(np__2))).#trainneg
cnf(i_0_1857, plain, (v1_xreal_0(k4_xcmplx_0(np__6559))|~v1_xreal_0(np__6561)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_1871, plain, (v1_xreal_0(k4_xcmplx_0(np__6168))|~v1_xreal_0(np__6425)|~v1_xreal_0(np__257))).#trainneg
cnf(i_0_1997, plain, (k3_xcmplx_0(np__2,k3_xcmplx_0(np__2,X1))=k3_xcmplx_0(np__4,X1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2006, plain, (k3_xcmplx_0(np__81,k3_xcmplx_0(np__81,X1))=k3_xcmplx_0(np__6561,X1)|~v1_xcmplx_0(np__81)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1993, plain, (k3_xcmplx_0(np__0,k3_xcmplx_0(np__6561,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__6561)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1994, plain, (k3_xcmplx_0(np__0,k3_xcmplx_0(np__257,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__257)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2001, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(np__81,X1))=k3_xcmplx_0(np__81,X1)|~v1_xcmplx_0(np__81)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2002, plain, (k3_xcmplx_0(np__6561,k3_xcmplx_0(np__0,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__6561)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2003, plain, (k3_xcmplx_0(np__257,k3_xcmplx_0(np__0,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(np__257)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2005, plain, (k3_xcmplx_0(np__81,k3_xcmplx_0(np__1,X1))=k3_xcmplx_0(np__81,X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(np__81)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2708, plain, (~v7_ordinal1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_2141, plain, (v7_ordinal1(np__8)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_37, plain, (k1_xboole_0=np__0)).#trainneg
cnf(i_0_391, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk2_3(X2,X1,X3)),X3)=X2|X1=np__0|X3!=k2_nat_d(X2,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2087, plain, (esk4_0=esk1_0)).#trainneg
cnf(i_0_2122, plain, (esk1_0=np__0)).#trainneg
cnf(i_0_2010, plain, (~v7_ordinal1(k1_numbers))).#trainneg
cnf(i_0_2149, plain, (~v1_xboole_0(np__7))).#trainneg
cnf(i_0_2150, plain, (~v1_xboole_0(np__5))).#trainneg
cnf(i_0_2165, plain, (~v1_xboole_0(np__133))).#trainneg
cnf(i_0_2182, plain, (~v1_xboole_0(np__6))).#trainneg
cnf(i_0_2712, plain, (~v7_ordinal1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2715, plain, (~v7_ordinal1(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_2018, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_390, plain, (X1=np__0|X1!=k2_nat_d(X2,X3)|~r1_xxreal_0(X3,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2714, plain, (~v7_ordinal1(k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_2015, negated_conjecture, (k2_nat_d(k13_newton(np__3,np__8),np__257)!=np__136|~v7_ordinal1(k13_newton(np__3,np__8))|~v7_ordinal1(np__257))).#trainneg
cnf(i_0_2709, plain, (~v7_ordinal1(k4_xcmplx_0(np__6)))).#trainneg
cnf(i_0_3491, plain, (v1_xreal_0(np__8)|~v1_xreal_0(np__3))).#trainneg
cnf(i_0_392, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,esk2_3(X2,X1,X3)),X3)=X2|X3=np__0|X3!=k2_nat_d(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2191, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_2195, plain, (~v1_xboole_0(np__4))).#trainneg
cnf(i_0_2707, plain, (~v2_xxreal_0(np__0))).#trainneg
cnf(i_0_2011, plain, (~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2710, plain, (~v7_ordinal1(k4_xcmplx_0(np__23)))).#trainneg
cnf(i_0_2711, plain, (~v7_ordinal1(k4_xcmplx_0(np__256)))).#trainneg
cnf(i_0_2713, plain, (~v7_ordinal1(k4_xcmplx_0(np__6561)))).#trainneg
cnf(i_0_389, plain, (X1=np__0|X2!=k2_nat_d(X3,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2083, plain, (v7_ordinal1(np__257)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_2307, plain, (k2_xcmplx_0(k2_newton(np__2,k1_newton(np__2,X1)),np__1)=k4_pepin(X1)|~v7_ordinal1(k2_newton(np__2,k1_newton(np__2,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2100, plain, (v7_ordinal1(np__8)|~v1_finset_1(np__8))).#trainneg
cnf(i_0_393, plain, (X1=np__0|v7_ordinal1(esk2_3(X2,X1,X3))|X3!=k2_nat_d(X2,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2275, plain, (v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_2949, plain, (~v1_xreal_0(k4_xcmplx_0(np__2))|~v2_xxreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_2950, plain, (~v1_xreal_0(k4_xcmplx_0(np__6))|~v2_xxreal_0(k4_xcmplx_0(np__6)))).#trainneg
cnf(i_0_2951, plain, (~v1_xreal_0(k4_xcmplx_0(np__23))|~v2_xxreal_0(k4_xcmplx_0(np__23)))).#trainneg
cnf(i_0_2952, plain, (~v1_xreal_0(k4_xcmplx_0(np__256))|~v2_xxreal_0(k4_xcmplx_0(np__256)))).#trainneg
cnf(i_0_2953, plain, (~v1_xreal_0(k4_xcmplx_0(np__1))|~v2_xxreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2264, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(k4_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2357, plain, (k1_nat_1(k1_newton(np__2,k1_newton(np__2,X1)),np__1)=k4_pepin(X1)|~v7_ordinal1(k1_newton(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2113, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2117, plain, (r1_xxreal_0(X1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_387, plain, (X1=k2_nat_d(X2,X3)|X3!=np__0|X1!=np__0|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2230, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(k4_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2016, plain, (m1_subset_1(k2_nat_d(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2413, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_394, plain, (X1=np__0|v7_ordinal1(esk2_3(X2,X3,X1))|X1!=k2_nat_d(X2,X3)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2960, plain, (v1_xcmplx_0(np__3)|~v1_int_1(np__3))).#trainneg
cnf(i_0_2954, plain, (~v1_xreal_0(k4_xcmplx_0(np__6561))|~v2_xxreal_0(k4_xcmplx_0(np__6561)))).#trainneg
cnf(i_0_2955, plain, (~v1_xreal_0(k4_xcmplx_0(np__4))|~v2_xxreal_0(k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_2956, plain, (~v1_xreal_0(k4_xcmplx_0(np__3))|~v2_xxreal_0(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_3036, plain, (~v1_abian(np__3)|~v1_int_1(np__3)|~v1_int_1(np__2))).#trainneg
cnf(i_0_2273, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2270, plain, (v7_ordinal1(np__8)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_2308, plain, (k2_xcmplx_0(k1_newton(np__2,k2_newton(np__2,X1)),np__1)=k4_pepin(X1)|~v7_ordinal1(k1_newton(np__2,k2_newton(np__2,X1)))|~v7_ordinal1(k2_newton(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2138, plain, (v7_ordinal1(np__6)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_2358, plain, (v7_ordinal1(np__6425)|~v7_ordinal1(np__257))).#trainneg
cnf(i_0_2118, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_575, plain, (k6_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2870, plain, (v1_xcmplx_0(np__8)|~v1_xreal_0(np__4))).#trainneg
cnf(i_0_3000, plain, (v1_xreal_0(np__8)|~v1_int_1(np__4))).#trainneg
cnf(i_0_3030, plain, (v1_xcmplx_0(np__6)|~v1_xreal_0(np__3))).#trainneg
cnf(i_0_2207, plain, (v7_ordinal1(np__4)|~v1_finset_1(np__4))).#trainneg
cnf(i_0_2208, plain, (v7_ordinal1(np__6425)|~v1_finset_1(np__6425))).#trainneg
cnf(i_0_2211, plain, (v7_ordinal1(np__256)|~v1_finset_1(np__256))).#trainneg
cnf(i_0_2212, plain, (v7_ordinal1(np__21)|~v1_finset_1(np__21))).#trainneg
cnf(i_0_2109, plain, (k1_nat_1(np__1,k2_newton(np__2,k1_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(k2_newton(np__2,k1_newton(np__2,X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2013, plain, (k2_nat_d(k4_nat_1(k4_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(k4_nat_1(k4_nat_d(X1,X2),k4_nat_d(X3,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2410, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k1_newton(X1,X2),k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2014, plain, (k2_nat_d(k13_newton(np__3,np__4),np__257)=np__81|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(np__257))).#trainneg
cnf(i_0_2756, plain, (k3_xcmplx_0(k1_newton(X1,X3),k1_newton(X1,X2))=k1_newton(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2221, plain, (k2_nat_d(k4_nat_1(k4_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(k4_nat_1(k4_nat_d(X1,X2),k4_nat_d(X3,X2)))|~v7_ordinal1(k4_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2119, plain, (v7_ordinal1(k4_pepin(X1))|~v1_finset_1(k4_pepin(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2213, plain, (v7_ordinal1(np__5)|~v1_finset_1(np__5))).#trainneg
cnf(i_0_410, plain, (k7_xcmplx_0(np__0,X1)=np__0|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3031, plain, (v1_xreal_0(np__6)|~v1_int_1(np__3))).#trainneg
cnf(i_0_3489, plain, (v1_xreal_0(np__3)|~v1_int_1(np__1))).#trainneg
cnf(i_0_2214, plain, (v7_ordinal1(np__22)|~v1_finset_1(np__22))).#trainneg
cnf(i_0_2294, plain, (v7_ordinal1(np__6)|~v1_finset_1(np__6))).#trainneg
cnf(i_0_2295, plain, (v7_ordinal1(np__7)|~v1_finset_1(np__7))).#trainneg
cnf(i_0_2296, plain, (v7_ordinal1(np__255)|~v1_finset_1(np__255))).#trainneg
cnf(i_0_2297, plain, (v7_ordinal1(np__6560)|~v1_finset_1(np__6560))).#trainneg
cnf(i_0_2080, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k13_newton(X1,X2),k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2245, plain, (k2_nat_d(k4_nat_1(k4_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(k4_nat_1(k4_nat_d(X1,X2),k4_nat_d(X3,X2)))|~v7_ordinal1(k4_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2379, plain, (k2_xcmplx_0(np__1,k2_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(np__1)|~v7_ordinal1(X1)|~m1_subset_1(k2_newton(np__2,k2_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_2765, plain, (X1=k2_nat_d(X2,X3)|r1_xxreal_0(X3,X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(k1_newton(X4,X5),X3),X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4)|~v7_ordinal1(k1_newton(X4,X5))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X5))).#trainneg
cnf(i_0_2279, plain, (k2_nat_d(k2_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X2,X3)),X4),k1_newton(X1,X2))=X4|r1_xxreal_0(k1_newton(X1,X2),X4)|~v1_xcmplx_0(X1)|~v7_ordinal1(k2_xcmplx_0(k1_newton(X1,k2_xcmplx_0(X2,X3)),X4))|~v7_ordinal1(k1_newton(X1,X2))|~v7_ordinal1(k1_newton(X1,X3))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2110, plain, (k1_nat_1(np__1,k1_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(k1_newton(np__2,k2_newton(np__2,X1)))|~v7_ordinal1(k2_newton(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2274, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2298, plain, (v7_ordinal1(np__6305)|~v1_finset_1(np__6305))).#trainneg
cnf(i_0_315, plain, (k3_xcmplx_0(X1,np__0)=np__0|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3681, plain, (v1_xcmplx_0(np__3)|~v1_int_1(np__1))).#trainneg
cnf(i_0_3188, plain, (v1_xreal_0(np__8)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_3719, plain, (v1_int_1(np__8)|~v1_int_1(np__2))).#trainneg
cnf(i_0_2299, plain, (v7_ordinal1(np__6536)|~v1_finset_1(np__6536))).#trainneg
cnf(i_0_2300, plain, (v7_ordinal1(np__6557)|~v1_finset_1(np__6557))).#trainneg
cnf(i_0_2301, plain, (v7_ordinal1(np__6558)|~v1_finset_1(np__6558))).#trainneg
cnf(i_0_2383, plain, (v7_ordinal1(np__6480)|~v1_finset_1(np__6480))).#trainneg
cnf(i_0_2384, plain, (v7_ordinal1(np__77)|~v1_finset_1(np__77))).#trainneg
cnf(i_0_2646, plain, (k2_nat_d(X1,np__257)=np__81|X1!=np__81|~v7_ordinal1(np__257)|~v7_ordinal1(np__81)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2393, plain, (k2_xcmplx_0(np__1,k1_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(k2_newton(np__2,X1))|~v7_ordinal1(np__1)|~v7_ordinal1(X1)|~m1_subset_1(k1_newton(np__2,k2_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_2203, plain, (k2_nat_d(k1_newton(np__3,np__4),np__257)=np__81|~v7_ordinal1(k1_newton(np__3,np__4))|~v7_ordinal1(np__257))).#trainneg
cnf(i_0_2392, plain, (k2_xcmplx_0(np__1,k2_newton(np__2,k1_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(np__1)|~v7_ordinal1(X1)|~m1_subset_1(k2_newton(np__2,k1_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_2091, plain, (k2_nat_d(k4_nat_1(k4_nat_d(X1,np__257),np__81),np__257)=k4_nat_d(k3_xcmplx_0(X1,k13_newton(np__3,np__4)),np__257)|~v7_ordinal1(k4_nat_1(k4_nat_d(X1,np__257),np__81))|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2385, plain, (v7_ordinal1(np__6304)|~v1_finset_1(np__6304))).#trainneg
cnf(i_0_2029, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k1_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2386, plain, (v7_ordinal1(np__80)|~v1_finset_1(np__80))).#trainneg
cnf(i_0_2387, plain, (v7_ordinal1(np__79)|~v1_finset_1(np__79))).#trainneg
cnf(i_0_2388, plain, (v7_ordinal1(np__175)|~v1_finset_1(np__175))).#trainneg
cnf(i_0_2389, plain, (v7_ordinal1(np__254)|~v1_finset_1(np__254))).#trainneg
cnf(i_0_2390, plain, (v7_ordinal1(np__78)|~v1_finset_1(np__78))).#trainneg
cnf(i_0_2545, plain, (v7_ordinal1(np__24)|~v1_finset_1(np__24))).#trainneg
cnf(i_0_2546, plain, (v7_ordinal1(np__23)|~v1_finset_1(np__23))).#trainneg
cnf(i_0_2547, plain, (v7_ordinal1(np__6559)|~v1_finset_1(np__6559))).#trainneg
cnf(i_0_2647, plain, (k2_nat_d(X1,np__257)=np__3|X1!=np__3|~v7_ordinal1(np__257)|~v7_ordinal1(np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2229, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,np__257),np__81),np__257)=k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),X1),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(k4_nat_d(X1,np__257))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2132, plain, (k4_nat_d(k4_nat_1(np__81,k2_nat_d(X1,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),X1),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2364, plain, (k1_nat_1(np__1,k2_newton(np__2,k1_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(np__1)|~v7_ordinal1(X1)|~m1_subset_1(k2_newton(np__2,k1_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_3404, plain, (X1=k2_nat_d(X2,X3)|r1_xxreal_0(X3,X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X4),X3),X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(X4)|~v7_ordinal1(k4_xcmplx_0(X4))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2282, plain, (k2_nat_d(k2_xcmplx_0(np__8,X1),np__2)=X1|r1_xxreal_0(np__2,X1)|~v7_ordinal1(k2_xcmplx_0(np__8,X1))|~v7_ordinal1(np__2)|~v7_ordinal1(np__4)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2356, plain, (k1_nat_1(np__1,k1_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(k2_newton(np__2,X1))|~v7_ordinal1(np__1)|~v7_ordinal1(X1)|~m1_subset_1(k1_newton(np__2,k2_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_3722, plain, (v1_int_1(np__8)|~v1_int_1(np__3)|~v7_ordinal1(np__5))).#trainneg
cnf(i_0_2548, plain, (v7_ordinal1(np__231)|~v1_finset_1(np__231))).#trainneg
cnf(i_0_2042, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k2_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2549, plain, (v7_ordinal1(np__6289)|~v1_finset_1(np__6289))).#trainneg
cnf(i_0_2550, plain, (v7_ordinal1(np__56)|~v1_finset_1(np__56))).#trainneg
cnf(i_0_2551, plain, (v7_ordinal1(np__135)|~v1_finset_1(np__135))).#trainneg
cnf(i_0_2552, plain, (v7_ordinal1(np__6423)|~v1_finset_1(np__6423))).#trainneg
cnf(i_0_2648, plain, (v1_int_1(np__2)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_2650, plain, (v7_ordinal1(np__6422)|~v1_finset_1(np__6422))).#trainneg
cnf(i_0_2651, plain, (v7_ordinal1(np__6421)|~v1_finset_1(np__6421))).#trainneg
cnf(i_0_2652, plain, (v7_ordinal1(np__6169)|~v1_finset_1(np__6169))).#trainneg
cnf(i_0_2225, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,np__257),np__81),np__257)=k4_nat_d(k3_xcmplx_0(k1_newton(np__3,np__4),X1),np__257)|~v7_ordinal1(k1_newton(np__3,np__4))|~v7_ordinal1(k4_nat_d(X1,np__257))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2583, plain, (X1=k2_nat_d(X2,k1_newton(X3,np__257))|r1_xxreal_0(k1_newton(X3,np__257),X1)|X2!=k2_xcmplx_0(k1_newton(X3,np__257),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,np__257))|~v7_ordinal1(k1_newton(X3,np__0))|~v7_ordinal1(np__257)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2563, plain, (X1=k2_nat_d(X2,k1_newton(X3,np__2))|r1_xxreal_0(k1_newton(X3,np__2),X1)|X2!=k2_xcmplx_0(k1_newton(X3,np__4),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,np__2))|~v7_ordinal1(np__2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3120, plain, (k2_xcmplx_0(X1,np__0)=k2_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_2355, plain, (m2_subset_1(k4_pepin(X1),k1_numbers,k5_numbers)|~v7_ordinal1(k1_newton(np__2,k2_newton(np__2,X1)))|~v7_ordinal1(k2_newton(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2202, plain, (k4_nat_d(k4_nat_1(np__81,k4_nat_d(X1,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(k1_newton(np__3,np__4),X1),np__257)|~v7_ordinal1(k1_newton(np__3,np__4))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2281, plain, (k2_nat_d(k2_xcmplx_0(np__4,X1),np__2)=X1|r1_xxreal_0(np__2,X1)|~v7_ordinal1(k2_xcmplx_0(np__4,X1))|~v7_ordinal1(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2763, plain, (k2_nat_d(k2_xcmplx_0(k3_xcmplx_0(k1_newton(X2,X3),X1),X4),X1)=X4|r1_xxreal_0(X1,X4)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v7_ordinal1(k2_xcmplx_0(k3_xcmplx_0(k1_newton(X2,X3),X1),X4))|~v7_ordinal1(k1_newton(X2,X3))|~v7_ordinal1(X4)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_2752, plain, (k2_nat_d(k2_xcmplx_0(k3_xcmplx_0(X3,k1_newton(X1,X2)),X4),k1_newton(X1,X2))=X4|r1_xxreal_0(k1_newton(X1,X2),X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v7_ordinal1(k2_xcmplx_0(k3_xcmplx_0(X3,k1_newton(X1,X2)),X4))|~v7_ordinal1(k1_newton(X1,X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3746, plain, (v1_int_1(np__8)|~v1_int_1(np__3)|~v1_int_1(np__1))).#trainneg
cnf(i_0_2653, plain, (v7_ordinal1(np__176)|~v1_finset_1(np__176))).#trainneg
cnf(i_0_2069, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k4_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2654, plain, (v7_ordinal1(np__132)|~v1_finset_1(np__132))).#trainneg
cnf(i_0_2655, plain, (v7_ordinal1(np__121)|~v1_finset_1(np__121))).#trainneg
cnf(i_0_2656, plain, (v7_ordinal1(np__111)|~v1_finset_1(np__111))).#trainneg
cnf(i_0_2657, plain, (v7_ordinal1(np__6344)|~v1_finset_1(np__6344))).#trainneg
cnf(i_0_2658, plain, (v7_ordinal1(np__22)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_2659, plain, (v7_ordinal1(np__23)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_2706, plain, (v1_int_1(np__6561)|~v7_ordinal1(np__81))).#trainneg
cnf(i_0_2717, plain, (v7_ordinal1(np__6400)|~v1_finset_1(np__6400))).#trainneg
cnf(i_0_2559, plain, (X1=k2_nat_d(X2,k1_newton(X3,np__0))|r1_xxreal_0(k1_newton(X3,np__0),X1)|X2!=k2_xcmplx_0(k1_newton(X3,np__257),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,np__0))|~v7_ordinal1(k1_newton(X3,np__257))|~v7_ordinal1(np__257)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2400, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k1_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3682, plain, (v1_xcmplx_0(k4_xcmplx_0(np__3))|~v1_xreal_0(np__4)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_2544, plain, (k2_nat_d(X1,np__1)=np__0|X1!=np__81|~v7_ordinal1(np__1)|~v7_ordinal1(np__81)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2702, plain, (k2_nat_d(X1,np__81)=np__0|X1!=np__81|~v7_ordinal1(np__81)|~v7_ordinal1(np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3332, plain, (X1=k2_nat_d(X2,X3)|r1_xxreal_0(X3,X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(k4_xcmplx_0(X4),X3),X1)|~v1_int_1(X4)|~v1_xcmplx_0(X3)|~v7_ordinal1(k4_xcmplx_0(X4))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2256, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,np__257),np__81),np__257)=k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),X1),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(np__81)|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3060, plain, (v1_int_1(X1)|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2718, plain, (v7_ordinal1(np__232)|~v1_finset_1(np__232))).#trainneg
cnf(i_0_2719, plain, (v7_ordinal1(np__134)|~v1_finset_1(np__134))).#trainneg
cnf(i_0_2720, plain, (v7_ordinal1(np__6424)|~v1_finset_1(np__6424))).#trainneg
cnf(i_0_2721, plain, (v7_ordinal1(np__252)|~v1_finset_1(np__252))).#trainneg
cnf(i_0_2722, plain, (v7_ordinal1(np__55)|~v1_finset_1(np__55))).#trainneg
cnf(i_0_2723, plain, (v7_ordinal1(np__120)|~v1_finset_1(np__120))).#trainneg
cnf(i_0_2724, plain, (v7_ordinal1(np__253)|~v1_finset_1(np__253))).#trainneg
cnf(i_0_2867, plain, (v1_xreal_0(np__6425)|~v1_xreal_0(np__257))).#trainneg
cnf(i_0_2868, plain, (v7_ordinal1(np__6168)|~v1_finset_1(np__6168))).#trainneg
cnf(i_0_2989, plain, (k3_xcmplx_0(np__1,np__1)=np__1|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2579, plain, (X1=k2_nat_d(X2,k1_newton(X3,np__1))|r1_xxreal_0(k1_newton(X3,np__1),X1)|X2!=k2_xcmplx_0(k1_newton(X3,np__2),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,np__1))|~v7_ordinal1(np__1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3392, plain, (X1=k2_nat_d(X2,k4_xcmplx_0(X3))|r1_xxreal_0(k4_xcmplx_0(X3),X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(X4,k4_xcmplx_0(X3)),X1)|~v1_xcmplx_0(X4)|~v1_xreal_0(X3)|~v7_ordinal1(k4_xcmplx_0(X3))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_3532, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2283, plain, (k2_nat_d(k2_xcmplx_0(np__6,X1),np__2)=X1|r1_xxreal_0(np__2,X1)|~v7_ordinal1(k2_xcmplx_0(np__6,X1))|~v7_ordinal1(np__2)|~v7_ordinal1(np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2589, plain, (X1=k2_nat_d(X2,k1_newton(X3,np__4))|r1_xxreal_0(k1_newton(X3,np__4),X1)|X2!=k2_xcmplx_0(k1_newton(X3,np__8),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,np__4))|~v7_ordinal1(np__4)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2560, plain, (X1=k2_nat_d(X2,k1_newton(X3,np__0))|r1_xxreal_0(k1_newton(X3,np__0),X1)|X2!=k2_xcmplx_0(k1_newton(X3,np__81),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,np__0))|~v7_ordinal1(k1_newton(X3,np__81))|~v7_ordinal1(np__81)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2251, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,np__257),np__81),np__257)=k4_nat_d(k3_xcmplx_0(k1_newton(np__3,np__4),X1),np__257)|~v7_ordinal1(k1_newton(np__3,np__4))|~v7_ordinal1(np__81)|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2096, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,np__257),np__81),np__257)=k4_nat_d(k3_xcmplx_0(X1,k13_newton(np__3,np__4)),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3539, plain, (v1_xcmplx_0(np__4)|~v1_xreal_0(np__3)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_2869, plain, (v7_ordinal1(np__133)|~v1_finset_1(np__133))).#trainneg
cnf(i_0_2097, plain, (k4_nat_d(k3_xcmplx_0(X1,k13_newton(np__3,np__4)),np__257)=k4_nat_d(k3_xcmplx_0(k4_nat_d(X1,np__257),np__81),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(np__257)|~v7_ordinal1(np__81)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X1,np__257),k5_numbers))).#trainneg
cnf(i_0_2945, plain, (v1_xcmplx_0(np__2)|~v1_int_1(np__2))).#trainneg
cnf(i_0_3142, plain, (v1_xcmplx_0(np__1)|~v1_int_1(np__1))).#trainneg
cnf(i_0_3219, plain, (v1_xcmplx_0(np__4)|~v1_int_1(np__4))).#trainneg
cnf(i_0_3297, plain, (v1_xreal_0(np__6)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_2585, plain, (X1=k2_nat_d(X2,k1_newton(X3,np__81))|r1_xxreal_0(k1_newton(X3,np__81),X1)|X2!=k2_xcmplx_0(k1_newton(X3,np__81),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,np__81))|~v7_ordinal1(k1_newton(X3,np__0))|~v7_ordinal1(np__81)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2985, plain, (k3_xcmplx_0(np__1,np__2)=np__2|~v1_xreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_2790, plain, (k2_xcmplx_0(X1,np__2)=k2_xcmplx_0(np__2,X1)|~v1_xcmplx_0(k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2235, plain, (k4_nat_d(k4_nat_1(np__81,k4_nat_d(X1,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(X1,k13_newton(np__3,np__4)),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(np__81)|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3738, plain, (v1_xreal_0(np__5)|~v1_int_1(np__3)|~v1_int_1(np__2))).#trainneg
cnf(i_0_3484, plain, (v1_xcmplx_0(np__6561)|~v1_int_1(np__6561))).#trainneg
cnf(i_0_2125, plain, (k2_nat_d(k4_nat_1(np__81,k4_nat_d(X1,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),X1),np__257)|~v7_ordinal1(k4_nat_1(np__81,k4_nat_d(X1,np__257)))|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3572, plain, (v1_xcmplx_0(np__6)|~v1_int_1(np__6))).#trainneg
cnf(i_0_3579, plain, (v1_xreal_0(np__4)|~v1_int_1(np__2))).#trainneg
cnf(i_0_3627, plain, (v1_xreal_0(np__4)|~v1_int_1(np__1))).#trainneg
cnf(i_0_2984, plain, (k3_xcmplx_0(np__1,np__0)=np__0|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_2263, plain, (k4_nat_d(k4_nat_1(np__81,k4_nat_d(X1,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(X1,k13_newton(np__3,np__4)),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(k4_nat_d(X1,np__257))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3121, plain, (k2_xcmplx_0(X1,np__2)=k2_xcmplx_0(np__2,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_3081, plain, (k2_xcmplx_0(X1,np__0)=k2_xcmplx_0(np__0,X1)|~v1_int_1(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2285, plain, (k2_nat_d(k2_xcmplx_0(np__81,X1),np__1)=X1|r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(np__81,X1))|~v7_ordinal1(np__1)|~v7_ordinal1(np__81)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2959, plain, (v1_xcmplx_0(np__3)|~m1_subset_1(k4_xcmplx_0(np__3),k1_numbers))).#trainneg
cnf(i_0_2127, plain, (k4_nat_d(k4_nat_1(np__81,k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),X1),np__257)),np__257)=k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),k4_nat_1(np__81,k4_nat_d(X1,np__257))),np__257)|~v7_ordinal1(k4_nat_1(np__81,k4_nat_d(X1,np__257)))|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3656, plain, (v1_xcmplx_0(np__256)|~v1_int_1(np__256))).#trainneg
cnf(i_0_3647, plain, (v1_xreal_0(np__23)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_3668, plain, (v1_xcmplx_0(np__23)|~v1_int_1(np__23))).#trainneg
cnf(i_0_3698, plain, (v1_xreal_0(np__22)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_2412, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2290, plain, (k2_nat_d(k2_xcmplx_0(np__81,X1),np__81)=X1|r1_xxreal_0(np__81,X1)|~v7_ordinal1(k2_xcmplx_0(np__81,X1))|~v7_ordinal1(np__81)|~v7_ordinal1(np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2309, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2292, plain, (k2_nat_d(k2_xcmplx_0(np__8,X1),np__4)=X1|r1_xxreal_0(np__4,X1)|~v7_ordinal1(k2_xcmplx_0(np__8,X1))|~v7_ordinal1(np__4)|~v7_ordinal1(np__2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2204, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,np__257),np__81),np__257)=k4_nat_d(k3_xcmplx_0(X1,k1_newton(np__3,np__4)),np__257)|~v7_ordinal1(k1_newton(np__3,np__4))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3712, plain, (v1_xreal_0(np__6480)|~v1_int_1(np__81))).#trainneg
cnf(i_0_3718, plain, (v1_xreal_0(np__6)|~v1_int_1(np__2))).#trainneg
cnf(i_0_3720, plain, (v1_xcmplx_0(np__6480)|~v1_xreal_0(np__81))).#trainneg
cnf(i_0_3744, plain, (v1_xreal_0(np__5)|~v1_int_1(np__1))).#trainneg
cnf(i_0_3745, plain, (v1_int_1(np__6)|~v1_int_1(np__1))).#trainneg
cnf(i_0_2728, plain, (v1_xreal_0(np__6)|~v1_int_1(k4_xcmplx_0(np__6)))).#trainneg
cnf(i_0_2729, plain, (v1_xreal_0(np__23)|~v1_int_1(k4_xcmplx_0(np__23)))).#trainneg
cnf(i_0_2992, plain, (k3_xcmplx_0(np__1,np__3)=np__3|~v1_xreal_0(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_2586, plain, (X1=k2_nat_d(X2,k1_newton(X3,np__4))|r1_xxreal_0(k1_newton(X3,np__4),X1)|X2!=k2_xcmplx_0(k1_newton(X3,np__4),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,np__4))|~v7_ordinal1(k1_newton(X3,np__0))|~v7_ordinal1(np__4)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2856, plain, (k3_xcmplx_0(np__1,np__0)=np__0|~v1_int_1(np__0))).#trainneg
cnf(i_0_2278, plain, (k2_nat_d(k2_xcmplx_0(np__0,X1),np__0)=X1|r1_xxreal_0(np__0,X1)|~v7_ordinal1(k2_xcmplx_0(np__0,X1))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2582, plain, (X1=k2_nat_d(X2,k1_newton(X3,np__6561))|r1_xxreal_0(k1_newton(X3,np__6561),X1)|X2!=k2_xcmplx_0(k1_newton(X3,np__6561),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,np__6561))|~v7_ordinal1(k1_newton(X3,np__0))|~v7_ordinal1(np__6561)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3192, plain, (v1_xreal_0(np__3)|~m1_subset_1(k4_xcmplx_0(np__3),k1_numbers))).#trainneg
cnf(i_0_3291, plain, (v1_xreal_0(np__136)|~v1_int_1(np__6425)|~v1_int_1(np__6561))).#trainneg
cnf(i_0_2133, plain, (k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),X1),np__257)=k4_nat_d(k3_xcmplx_0(np__81,k4_nat_d(X1,np__257)),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(k4_nat_d(X1,np__257))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2734, plain, (v1_xreal_0(np__3)|~v1_int_1(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_2730, plain, (v1_xreal_0(np__256)|~v1_int_1(k4_xcmplx_0(np__256)))).#trainneg
cnf(i_0_2731, plain, (v1_xreal_0(np__1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2733, plain, (v1_xreal_0(np__4)|~v1_int_1(k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_2754, plain, (X1=k2_nat_d(X2,k1_newton(X3,X4))|r1_xxreal_0(k1_newton(X3,X4),X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(X5,k1_newton(X3,X4)),X1)|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,X4))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_3049, plain, (k3_xcmplx_0(np__2,np__1)=np__2|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_3407, plain, (k3_xcmplx_0(X1,np__0)=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_2429, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2291, plain, (k2_nat_d(k2_xcmplx_0(np__6561,X1),np__81)=X1|r1_xxreal_0(np__81,X1)|~v7_ordinal1(k2_xcmplx_0(np__6561,X1))|~v7_ordinal1(np__81)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2861, plain, (k3_xcmplx_0(np__1,np__1)=np__1|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2258, plain, (k4_nat_d(k4_nat_1(np__81,k4_nat_d(X1,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(X1,k1_newton(np__3,np__4)),np__257)|~v7_ordinal1(k1_newton(np__3,np__4))|~v7_ordinal1(k4_nat_d(X1,np__257))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3796, plain, (k3_xcmplx_0(np__4,np__3)=k3_xcmplx_0(np__2,np__6)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(np__3))).#trainneg
cnf(i_0_3521, plain, (v1_xcmplx_0(np__1)|~v1_int_1(np__257)|~v1_int_1(np__256))).#trainneg
cnf(i_0_2205, plain, (k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),k1_newton(np__3,np__4)),np__257)=k4_nat_d(k4_nat_1(np__81,np__81),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(k1_newton(np__3,np__4))|~v7_ordinal1(np__257))).#trainneg
cnf(i_0_2991, plain, (k3_xcmplx_0(np__1,np__4)=np__4|~v1_xreal_0(k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_2415, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k4_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3125, plain, (k2_xcmplx_0(X1,np__1)=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2857, plain, (k3_xcmplx_0(np__1,np__2)=np__2|~v1_int_1(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_2231, plain, (k4_nat_d(k4_nat_1(np__81,k4_nat_d(X1,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(X1,k1_newton(np__3,np__4)),np__257)|~v7_ordinal1(k1_newton(np__3,np__4))|~v7_ordinal1(np__81)|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2596, plain, (X1=k2_nat_d(X2,k1_newton(X3,np__3))|r1_xxreal_0(k1_newton(X3,np__3),X1)|X2!=k2_xcmplx_0(k1_newton(X3,np__6),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,np__3))|~v7_ordinal1(np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2687, plain, (k1_newton(X1,X2)=k2_nat_d(X3,X4)|r1_xxreal_0(X4,k1_newton(X1,X2))|X3!=k2_xcmplx_0(k1_newton(X1,X2),k3_xcmplx_0(X4,X5))|~v1_xcmplx_0(k3_xcmplx_0(X4,X5))|~v1_xcmplx_0(X1)|~v7_ordinal1(k1_newton(X1,X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X5)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3322, plain, (X1=k2_nat_d(X2,k4_xcmplx_0(X3))|r1_xxreal_0(k4_xcmplx_0(X3),X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(X4,k4_xcmplx_0(X3)),X1)|~v1_int_1(X3)|~v1_xcmplx_0(X4)|~v7_ordinal1(k4_xcmplx_0(X3))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_2922, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1489, plain, (v7_ordinal1(np__8))).#trainneg
cnf(i_0_2087, plain, (esk4_0=np__0)).#trainneg
cnf(i_0_1535, plain, (k3_xcmplx_0(np__1,np__0)=np__0)).#trainneg
cnf(i_0_4141, plain, (~v2_xxreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_4359, plain, (~v2_xxreal_0(k4_xcmplx_0(np__6561)))).#trainneg
cnf(i_0_3927, plain, (~m1_subset_1(k4_xcmplx_0(np__2),k5_numbers))).#trainneg
cnf(i_0_3945, plain, (~m1_subset_1(k4_xcmplx_0(np__1),k5_numbers))).#trainneg
cnf(i_0_4036, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4602, plain, (v7_ordinal1(np__256))).#trainneg
cnf(i_0_4583, plain, (v1_xcmplx_0(np__8))).#trainneg
cnf(i_0_291, plain, (v1_xreal_0(np__0))).#trainneg
cnf(i_0_292, plain, (v1_xxreal_0(np__0))).#trainneg
cnf(i_0_3946, plain, (~m1_subset_1(k4_xcmplx_0(np__3),k5_numbers))).#trainneg
cnf(i_0_4007, plain, (~m1_subset_1(k4_xcmplx_0(np__4),k5_numbers))).#trainneg
cnf(i_0_4009, plain, (~m1_subset_1(k4_xcmplx_0(np__6),k5_numbers))).#trainneg
cnf(i_0_4026, plain, (~m1_subset_1(k4_xcmplx_0(np__23),k5_numbers))).#trainneg
cnf(i_0_4027, plain, (~m1_subset_1(k4_xcmplx_0(np__256),k5_numbers))).#trainneg
cnf(i_0_4235, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4599, plain, (v7_ordinal1(np__6425))).#trainneg
cnf(i_0_4661, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1582, plain, (v7_ordinal1(np__6))).#trainneg
cnf(i_0_3928, negated_conjecture, (k2_nat_d(k1_newton(np__3,np__8),np__257)!=np__136|~v7_ordinal1(k1_newton(np__3,np__8))|~v7_ordinal1(np__257))).#trainneg
cnf(i_0_4029, plain, (X1=np__0|~r1_xxreal_0(X1,k2_nat_d(X2,X1))|~v7_ordinal1(k2_nat_d(X2,X1))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4573, plain, (r1_xxreal_0(np__8,np__8)|~v1_xreal_0(np__3))).#trainneg
cnf(i_0_1678, plain, (X1=k2_nat_d(X2,np__2)|r1_xxreal_0(np__2,X1)|X2!=k2_xcmplx_0(np__6,X1)|~v7_ordinal1(np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1815, plain, (v1_xreal_0(np__6))).#trainneg
cnf(i_0_4028, plain, (~m1_subset_1(k4_xcmplx_0(np__6561),k5_numbers))).#trainneg
cnf(i_0_4145, plain, (~v1_xreal_0(np__6)|~v2_xxreal_0(k4_xcmplx_0(np__6)))).#trainneg
cnf(i_0_4149, plain, (~v1_xreal_0(np__23)|~v2_xxreal_0(k4_xcmplx_0(np__23)))).#trainneg
cnf(i_0_4005, plain, (k2_nat_d(X1,X2)=np__0|~r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4138, plain, (v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_5068, plain, (X1=k2_nat_d(X2,X3)|r1_xxreal_0(X3,X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(X4,X3),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X4)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_4297, plain, (X1=k2_nat_d(X2,X3)|X3!=np__0|X1!=np__0|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3981, plain, (k2_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(X3,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6323, plain, (v1_xcmplx_0(np__6)|~v1_int_1(np__3))).#trainneg
cnf(i_0_1813, plain, (v1_xreal_0(np__8)|~v1_xreal_0(np__6))).#trainneg
cnf(i_0_1568, plain, (v7_ordinal1(np__7)|~v7_ordinal1(np__5))).#trainneg
cnf(i_0_1737, plain, (v1_xreal_0(np__4))).#trainneg
cnf(i_0_4589, plain, (v1_xcmplx_0(np__6))).#trainneg
cnf(i_0_4605, plain, (v7_ordinal1(np__21))).#trainneg
cnf(i_0_4150, plain, (~v1_int_1(np__23)|~v2_xxreal_0(k4_xcmplx_0(np__23)))).#trainneg
cnf(i_0_4153, plain, (~v1_xreal_0(np__256)|~v2_xxreal_0(k4_xcmplx_0(np__256)))).#trainneg
cnf(i_0_4154, plain, (~v1_int_1(np__256)|~v2_xxreal_0(k4_xcmplx_0(np__256)))).#trainneg
cnf(i_0_4158, plain, (~v1_xreal_0(np__1)|~v2_xxreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_4265, plain, (r1_xxreal_0(np__0,np__0))).#trainneg
cnf(i_0_4070, plain, (k2_xcmplx_0(X1,np__0)=X1|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4210, plain, (k2_nat_d(k4_nat_1(k2_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(k4_nat_1(k2_nat_d(X1,X2),k4_nat_d(X3,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_3953, plain, (k2_nat_d(k4_nat_1(k2_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(k4_nat_1(k2_nat_d(X1,X2),k4_nat_d(X3,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1551, plain, (X1=k2_nat_d(X2,np__2)|r1_xxreal_0(np__2,X1)|X2!=k2_xcmplx_0(np__4,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4740, plain, (v7_ordinal1(np__5))).#trainneg
cnf(i_0_4749, plain, (v7_ordinal1(np__22))).#trainneg
cnf(i_0_4755, plain, (v7_ordinal1(np__7))).#trainneg
cnf(i_0_4363, plain, (~v1_xreal_0(np__4)|~v2_xxreal_0(k4_xcmplx_0(np__4)))).#trainneg
cnf(i_0_4367, plain, (~v1_xreal_0(np__3)|~v2_xxreal_0(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_4368, plain, (~v1_int_1(np__3)|~v2_xxreal_0(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_4406, plain, (k2_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(X3,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4003, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(X3,X2),k4_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(k4_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X3,X2),k5_numbers))).#trainneg
cnf(i_0_4345, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7091, plain, (v1_xcmplx_0(np__1)|~v1_int_1(np__257)|~v7_ordinal1(np__256))).#trainneg
cnf(i_0_4744, plain, (v1_xreal_0(np__3)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1451, plain, (X1=k2_nat_d(X2,np__2)|r1_xxreal_0(np__2,X1)|X2!=k2_xcmplx_0(np__8,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4758, plain, (v7_ordinal1(np__255))).#trainneg
cnf(i_0_4761, plain, (v7_ordinal1(np__6560))).#trainneg
cnf(i_0_4954, plain, (v7_ordinal1(np__6305))).#trainneg
cnf(i_0_4148, plain, (~m1_subset_1(k4_xcmplx_0(np__23),k1_numbers)|~v2_xxreal_0(k4_xcmplx_0(np__23)))).#trainneg
cnf(i_0_4152, plain, (~m1_subset_1(k4_xcmplx_0(np__256),k1_numbers)|~v2_xxreal_0(k4_xcmplx_0(np__256)))).#trainneg
cnf(i_0_4284, plain, (k2_nat_d(X1,X2)=np__0|X2!=np__0|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4095, plain, (k2_xcmplx_0(k1_newton(np__2,k1_newton(np__2,X1)),np__1)=k4_pepin(X1)|~v7_ordinal1(k1_newton(np__2,k1_newton(np__2,X1)))|~v7_ordinal1(k1_newton(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5052, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(X3,X2),k4_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X1,X2),k5_numbers)|~m1_subset_1(k2_nat_d(X3,X2),k5_numbers))).#trainneg
cnf(i_0_6744, plain, (k3_xcmplx_0(np__1,np__3)=np__3|~v1_xreal_0(np__3))).#trainneg
cnf(i_0_5902, plain, (v1_xcmplx_0(np__6425)|~v1_xreal_0(np__257))).#trainneg
cnf(i_0_1467, plain, (X1=k2_nat_d(X2,np__4)|r1_xxreal_0(np__4,X1)|X2!=k2_xcmplx_0(np__8,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4970, plain, (v7_ordinal1(np__6536))).#trainneg
cnf(i_0_4973, plain, (v7_ordinal1(np__6557))).#trainneg
cnf(i_0_4976, plain, (v7_ordinal1(np__6558))).#trainneg
cnf(i_0_4366, plain, (~m1_subset_1(k4_xcmplx_0(np__3),k1_numbers)|~v2_xxreal_0(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_4979, plain, (v7_ordinal1(np__6480))).#trainneg
cnf(i_0_6185, plain, (v1_xcmplx_0(np__4)|~v1_int_1(np__2))).#trainneg
cnf(i_0_6186, plain, (v1_xcmplx_0(np__4)|~v1_int_1(np__1))).#trainneg
cnf(i_0_5384, plain, (k3_xcmplx_0(X1,np__1)=X1|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3999, plain, (k4_nat_d(k4_nat_1(np__81,k2_nat_d(X1,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),X1),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4278, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4982, plain, (v7_ordinal1(np__77))).#trainneg
cnf(i_0_5174, plain, (v7_ordinal1(np__6304))).#trainneg
cnf(i_0_5190, plain, (v7_ordinal1(np__80))).#trainneg
cnf(i_0_5193, plain, (v7_ordinal1(np__79))).#trainneg
cnf(i_0_5196, plain, (v7_ordinal1(np__175))).#trainneg
cnf(i_0_5199, plain, (v7_ordinal1(np__254))).#trainneg
cnf(i_0_5202, plain, (v7_ordinal1(np__78))).#trainneg
cnf(i_0_5205, plain, (v7_ordinal1(np__24))).#trainneg
cnf(i_0_4443, plain, (k2_xcmplx_0(np__1,k2_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(k2_newton(np__2,k2_newton(np__2,X1)))|~v7_ordinal1(X1)|~m1_subset_1(k2_newton(np__2,k2_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_3991, plain, (k4_nat_d(k4_nat_1(np__81,k2_nat_d(X1,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(k1_newton(np__3,np__4),X1),np__257)|~v7_ordinal1(k1_newton(np__3,np__4))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_5560, plain, (k2_nat_d(np__4,np__2)=np__0|~v7_ordinal1(np__4))).#trainneg
cnf(i_0_4234, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,np__257),np__81),np__257)=k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),X1),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6207, plain, (k2_xcmplx_0(X1,np__2)=k2_xcmplx_0(np__2,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4576, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3548, plain, (v1_int_1(np__1)|~v1_int_1(np__257)|~v1_int_1(np__256))).#trainneg
cnf(i_0_1694, plain, (X1=k2_nat_d(X2,np__6561)|r1_xxreal_0(np__6561,X1)|X2!=k2_xcmplx_0(np__0,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5208, plain, (v7_ordinal1(np__23))).#trainneg
cnf(i_0_5211, plain, (v7_ordinal1(np__6559))).#trainneg
cnf(i_0_5397, plain, (v7_ordinal1(np__231))).#trainneg
cnf(i_0_5405, plain, (v7_ordinal1(np__6289))).#trainneg
cnf(i_0_5408, plain, (v7_ordinal1(np__56))).#trainneg
cnf(i_0_5411, plain, (v7_ordinal1(np__135))).#trainneg
cnf(i_0_5414, plain, (v7_ordinal1(np__6423))).#trainneg
cnf(i_0_5418, plain, (v7_ordinal1(np__6422))).#trainneg
cnf(i_0_4444, plain, (k2_xcmplx_0(np__1,k2_newton(np__2,k1_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(k2_newton(np__2,k1_newton(np__2,X1)))|~v7_ordinal1(X1)|~m1_subset_1(k2_newton(np__2,k1_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_6738, plain, (r1_xxreal_0(X1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_5963, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6708, plain, (k3_xcmplx_0(np__1,np__2)=np__2|~v1_xcmplx_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_6201, plain, (k3_xcmplx_0(np__1,np__2)=np__2)).#trainneg
cnf(i_0_4223, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,np__257),np__81),np__257)=k4_nat_d(k3_xcmplx_0(k1_newton(np__3,np__4),X1),np__257)|~v7_ordinal1(k1_newton(np__3,np__4))|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4246, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3551, plain, (v1_xcmplx_0(np__1)|~v1_xreal_0(np__257)|~v1_xreal_0(np__256))).#trainneg
cnf(i_0_1692, plain, (X1=k2_nat_d(X2,np__0)|r1_xxreal_0(np__0,X1)|X2!=k2_xcmplx_0(np__0,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_5421, plain, (v7_ordinal1(np__6421))).#trainneg
cnf(i_0_5424, plain, (v7_ordinal1(np__6169))).#trainneg
cnf(i_0_8487, plain, (~v2_xxreal_0(k4_xcmplx_0(np__3)))).#trainneg
cnf(i_0_8260, plain, (~v2_xxreal_0(k4_xcmplx_0(np__256)))).#trainneg
cnf(i_0_9507, plain, (~v1_xcmplx_0(k4_xcmplx_0(np__23))|~v2_xxreal_0(k4_xcmplx_0(np__23)))).#trainneg
cnf(i_0_7551, plain, (m1_subset_1(k1_newton(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7984, plain, (X1=k2_nat_d(X2,X3)|r1_xxreal_0(X3,X1)|X2!=k2_xcmplx_0(X3,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2003, plain, (k3_xcmplx_0(np__257,k3_xcmplx_0(np__0,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__257)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5632, plain, (v7_ordinal1(np__176))).#trainneg
cnf(i_0_5640, plain, (v7_ordinal1(np__132))).#trainneg
cnf(i_0_5643, plain, (v7_ordinal1(np__121))).#trainneg
cnf(i_0_5646, plain, (v7_ordinal1(np__111))).#trainneg
cnf(i_0_5649, plain, (v7_ordinal1(np__6344))).#trainneg
cnf(i_0_5660, plain, (v7_ordinal1(np__6400))).#trainneg
cnf(i_0_7737, plain, (k2_xcmplx_0(np__0,X1)=X1|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2015, negated_conjecture, (k2_nat_d(k13_newton(np__3,np__8),np__257)!=np__136|~v7_ordinal1(k13_newton(np__3,np__8)))).#trainneg
cnf(i_0_7369, plain, (r1_xxreal_0(np__8,np__8))).#trainneg
cnf(i_0_9426, plain, (v1_int_1(np__1)|~v1_int_1(np__257))).#trainneg
cnf(i_0_2002, plain, (k3_xcmplx_0(np__6561,k3_xcmplx_0(np__0,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__6561)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5849, plain, (v7_ordinal1(np__232))).#trainneg
cnf(i_0_5886, plain, (v7_ordinal1(np__134))).#trainneg
cnf(i_0_5889, plain, (v7_ordinal1(np__6424))).#trainneg
cnf(i_0_5892, plain, (v7_ordinal1(np__252))).#trainneg
cnf(i_0_5895, plain, (v7_ordinal1(np__55))).#trainneg
cnf(i_0_5898, plain, (v7_ordinal1(np__120))).#trainneg
cnf(i_0_5901, plain, (v7_ordinal1(np__253))).#trainneg
cnf(i_0_5908, plain, (v7_ordinal1(np__6168))).#trainneg
cnf(i_0_9347, plain, (X1=k2_nat_d(X2,X3)|r1_xxreal_0(X3,X1)|X2!=k2_xcmplx_0(X1,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_9645, plain, (r1_xxreal_0(np__0,np__136))).#trainneg
cnf(i_0_9083, plain, (k4_nat_d(k3_xcmplx_0(X1,np__1),X2)=k4_nat_d(X1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8452, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_7383, plain, (r1_xxreal_0(np__136,np__136))).#trainneg
cnf(i_0_9649, plain, (v1_xcmplx_0(np__1)|~v1_xreal_0(np__257))).#trainneg
cnf(i_0_1994, plain, (k3_xcmplx_0(np__0,k3_xcmplx_0(np__257,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__257)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6134, plain, (v7_ordinal1(np__133))).#trainneg
cnf(i_0_6683, plain, (v1_xcmplx_0(np__6480))).#trainneg
cnf(i_0_2648, plain, (v1_int_1(np__2))).#trainneg
cnf(i_0_2706, plain, (v1_int_1(np__6561))).#trainneg
cnf(i_0_7555, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9278, plain, (v1_int_1(np__0))).#trainneg
cnf(i_0_9643, plain, (r1_xxreal_0(np__0,np__81))).#trainneg
cnf(i_0_9289, plain, (np__0=np__2|X1!=np__0|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1993, plain, (k3_xcmplx_0(np__0,k3_xcmplx_0(np__6561,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__6561)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3698, plain, (v1_xreal_0(np__22))).#trainneg
cnf(i_0_3647, plain, (v1_xreal_0(np__23))).#trainneg
cnf(i_0_3539, plain, (v1_xcmplx_0(np__4))).#trainneg
cnf(i_0_8014, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_8846, negated_conjecture, (np__136!=np__0|np__257!=np__0|~v7_ordinal1(k1_newton(np__3,np__8)))).#trainneg
cnf(i_0_2647, plain, (k2_nat_d(X1,np__257)=np__3|X1!=np__3|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9302, plain, (X1=k2_nat_d(X2,np__0)|r1_xxreal_0(np__0,X1)|X2!=k2_xcmplx_0(X1,np__0)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_7550, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3715, plain, (v1_xreal_0(np__254)|~v1_int_1(np__3)|~v1_int_1(np__257))).#trainneg
cnf(i_0_1947, plain, (k3_xcmplx_0(np__257,k7_xcmplx_0(np__0,X1))=k7_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__257)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7751, plain, (k2_xcmplx_0(np__8,np__0)=np__8)).#trainneg
cnf(i_0_7753, plain, (k3_xcmplx_0(np__1,np__8)=np__8)).#trainneg
cnf(i_0_1870, plain, (v1_xreal_0(np__6558)|~v1_xreal_0(np__3))).#trainneg
cnf(i_0_1868, plain, (v1_xreal_0(np__6557)|~v1_xreal_0(np__4))).#trainneg
cnf(i_0_1866, plain, (v1_xreal_0(np__6304)|~v1_xreal_0(np__257))).#trainneg
cnf(i_0_1864, plain, (v1_xreal_0(np__6305)|~v1_xreal_0(np__256))).#trainneg
cnf(i_0_1863, plain, (v1_xreal_0(np__6559)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_7507, plain, (k2_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|~v7_ordinal1(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(X3,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3722, plain, (v1_int_1(np__8)|~v1_int_1(np__3))).#trainneg
cnf(i_0_3724, plain, (v1_xreal_0(np__78)|~v1_int_1(np__3)|~v1_int_1(np__81))).#trainneg
cnf(i_0_1946, plain, (k3_xcmplx_0(np__6561,k7_xcmplx_0(np__0,X1))=k7_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__6561)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7752, plain, (k3_xcmplx_0(np__8,np__0)=np__0)).#trainneg
cnf(i_0_7765, plain, (k3_xcmplx_0(np__3,np__0)=np__0)).#trainneg
cnf(i_0_9646, plain, (r1_xxreal_0(np__0,np__8))).#trainneg
cnf(i_0_1862, plain, (v1_xreal_0(np__6536)|~v1_xreal_0(np__25))).#trainneg
cnf(i_0_1873, plain, (v1_xreal_0(np__79)|~v1_xreal_0(np__81))).#trainneg
cnf(i_0_1816, plain, (v1_xreal_0(np__5)|~v1_xreal_0(np__3))).#trainneg
cnf(i_0_1811, plain, (v1_xreal_0(np__7)|~v1_xreal_0(np__5))).#trainneg
cnf(i_0_6677, plain, (v1_xreal_0(np__6480)|~v7_ordinal1(np__81))).#trainneg
cnf(i_0_2014, plain, (k2_nat_d(k13_newton(np__3,np__4),np__257)=np__81|~v7_ordinal1(k13_newton(np__3,np__4)))).#trainneg
cnf(i_0_8736, plain, (X1=k2_nat_d(X2,X3)|r1_xxreal_0(X3,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(X4,X3))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X4)|~m1_subset_1(k3_xcmplx_0(X4,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7688, plain, (X1=k2_nat_d(X2,X3)|r1_xxreal_0(X3,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(X3,X4))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X4)|~m1_subset_1(k3_xcmplx_0(X3,X4),k1_numbers))).#trainneg
cnf(i_0_9332, plain, (X1=k2_nat_d(X2,np__257)|r1_xxreal_0(np__257,X1)|X2!=k2_xcmplx_0(X1,np__6425)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_7370, plain, (r1_xxreal_0(X1,np__8)|r1_xxreal_0(np__8,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3729, plain, (v1_xreal_0(np__80)|~v1_int_1(np__1)|~v1_int_1(np__81))).#trainneg
cnf(i_0_1937, plain, (k3_xcmplx_0(np__0,k7_xcmplx_0(np__6561,X1))=k7_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__6561)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9647, plain, (r1_xxreal_0(np__0,np__3))).#trainneg
cnf(i_0_7963, plain, (v1_abian(np__2)|~v1_int_1(np__1))).#trainneg
cnf(i_0_1874, plain, (v1_xreal_0(np__80)|~v1_xreal_0(np__81))).#trainneg
cnf(i_0_8286, plain, (k2_xcmplx_0(np__0,X1)=X1|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2203, plain, (k2_nat_d(k1_newton(np__3,np__4),np__257)=np__81|~v7_ordinal1(k1_newton(np__3,np__4)))).#trainneg
cnf(i_0_7384, plain, (r1_xxreal_0(X1,np__136)|r1_xxreal_0(np__136,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3730, plain, (v1_xreal_0(np__77)|~v1_int_1(np__4)|~v1_int_1(np__81))).#trainneg
cnf(i_0_1938, plain, (k3_xcmplx_0(np__0,k7_xcmplx_0(np__257,X1))=k7_xcmplx_0(np__0,X1)|~v1_xcmplx_0(np__257)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4132, plain, (k2_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_4269, plain, (r1_xxreal_0(np__2,np__2))).#trainneg
cnf(i_0_1740, plain, (v1_xreal_0(np__24)|~v1_xreal_0(np__4))).#trainneg
cnf(i_0_1875, plain, (v1_xreal_0(np__77)|~v1_xreal_0(np__81))).#trainneg
cnf(i_0_1876, plain, (v1_xreal_0(np__78)|~v1_xreal_0(np__81))).#trainneg
cnf(i_0_1872, plain, (v1_xreal_0(np__254)|~v1_xreal_0(np__257))).#trainneg
cnf(i_0_7781, plain, (k2_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(X3,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_9336, plain, (X1=k2_nat_d(X2,np__81)|r1_xxreal_0(np__81,X1)|X2!=k2_xcmplx_0(X1,np__81)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_8100, plain, (X1=k2_nat_d(X2,X3)|r1_xxreal_0(X3,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(X4,X3))|~v7_ordinal1(k3_xcmplx_0(X4,X3))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_8465, plain, (X1=k2_nat_d(X2,np__2)|r1_xxreal_0(np__2,X1)|X2!=k2_xcmplx_0(X1,np__4)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_7387, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__3,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10547, plain, (np__0=np__2|np__0!=np__6561)).#trainneg
cnf(i_0_1857, plain, (v1_xreal_0(k4_xcmplx_0(np__6559))|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_10617, negated_conjecture, (np__136!=np__0|np__257!=np__0)).#trainneg
cnf(i_0_3719, plain, (v1_int_1(np__8))).#trainneg
cnf(i_0_4149, plain, (~v2_xxreal_0(k4_xcmplx_0(np__23)))).#trainneg
cnf(i_0_3036, plain, (~v1_abian(np__3)|~v1_int_1(np__3))).#trainneg
cnf(i_0_1441, plain, (~v1_abian(np__3)|~v1_abian(np__2))).#trainneg
cnf(i_0_9977, plain, (~v1_abian(np__1)|~v1_abian(np__0)|~v1_int_1(np__0))).#trainneg
cnf(i_0_11841, plain, (r1_xxreal_0(np__2,np__3))).#trainneg
cnf(i_0_9709, plain, (m1_subset_1(k1_newton(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1865, plain, (v1_xreal_0(np__136))).#trainneg
cnf(i_0_10642, plain, (np__0=np__3|np__257!=np__0|X1!=np__3|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1856, plain, (v1_xreal_0(k4_xcmplx_0(np__3))|~v1_xreal_0(np__5))).#trainneg
cnf(i_0_10043, plain, (v1_int_1(np__1))).#trainneg
cnf(i_0_1899, plain, (v1_int_1(np__6))).#trainneg
cnf(i_0_1896, plain, (v1_int_1(np__4))).#trainneg
cnf(i_0_1552, plain, (v1_abian(np__4))).#trainneg
cnf(i_0_10397, plain, (v1_xcmplx_0(np__6425))).#trainneg
cnf(i_0_9931, plain, (k2_nat_d(np__6561,np__6425)=np__136)).#trainneg
cnf(i_0_9956, plain, (k2_nat_d(np__3,np__3)=np__0)).#trainneg
cnf(i_0_9898, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_9901, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_9900, plain, (k4_nat_d(k3_xcmplx_0(k4_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_10781, plain, (k2_xcmplx_0(np__0,np__8)=np__8)).#trainneg
cnf(i_0_11288, plain, (r1_xxreal_0(np__3,np__8))).#trainneg
cnf(i_0_10552, plain, (np__0=np__2|np__136!=np__0|~r1_xxreal_0(np__2,np__136))).#trainneg
cnf(i_0_2702, plain, (k2_nat_d(X1,np__81)=np__0|X1!=np__81|~v7_ordinal1(np__81)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10586, plain, (v1_xcmplx_0(np__22))).#trainneg
cnf(i_0_10590, plain, (v1_xcmplx_0(np__23))).#trainneg
cnf(i_0_10815, plain, (v1_xcmplx_0(np__6558))).#trainneg
cnf(i_0_10819, plain, (v1_xcmplx_0(np__6557))).#trainneg
cnf(i_0_10823, plain, (v1_xcmplx_0(np__6304))).#trainneg
cnf(i_0_10831, plain, (v1_xcmplx_0(np__6559))).#trainneg
cnf(i_0_10969, plain, (v1_xcmplx_0(np__5))).#trainneg
cnf(i_0_10487, plain, (r1_xxreal_0(np__257,np__257))).#trainneg
cnf(i_0_9904, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_11532, plain, (r1_xxreal_0(np__3,np__136))).#trainneg
cnf(i_0_10553, plain, (np__0=np__2|np__8!=np__0|~r1_xxreal_0(np__2,np__8))).#trainneg
cnf(i_0_2544, plain, (k2_nat_d(X1,np__1)=np__0|X1!=np__81|~v7_ordinal1(np__81)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10973, plain, (v1_xcmplx_0(np__7))).#trainneg
cnf(i_0_11580, plain, (v1_xcmplx_0(np__24))).#trainneg
cnf(i_0_11592, plain, (v1_xcmplx_0(np__254))).#trainneg
cnf(i_0_1900, plain, (v1_int_1(np__5)|~v1_int_1(np__3))).#trainneg
cnf(i_0_1895, plain, (v1_int_1(np__7)|~v1_int_1(np__5))).#trainneg
cnf(i_0_1785, plain, (v1_int_1(np__6559)|~v1_int_1(np__6561))).#trainneg
cnf(i_0_1767, plain, (v1_int_1(np__79)|~v1_int_1(np__81))).#trainneg
cnf(i_0_9925, plain, (k2_nat_d(np__3,np__2)=np__1)).#trainneg
cnf(i_0_9724, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11749, plain, (X1=k2_nat_d(X2,np__1)|r1_xxreal_0(np__1,X1)|X2!=k2_xcmplx_0(X1,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_9947, plain, (k2_nat_d(np__257,np__257)=np__0)).#trainneg
cnf(i_0_10932, plain, (k3_xcmplx_0(np__0,np__8)=np__0)).#trainneg
cnf(i_0_10554, plain, (np__0=np__2|np__0!=np__3|~r1_xxreal_0(np__2,np__3))).#trainneg
cnf(i_0_2579, plain, (X1=k2_nat_d(X2,k1_newton(X3,np__1))|r1_xxreal_0(k1_newton(X3,np__1),X1)|X2!=k2_xcmplx_0(k1_newton(X3,np__2),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,np__1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10788, plain, (k2_nat_d(np__8,np__8)=np__0)).#trainneg
cnf(i_0_1786, plain, (v1_int_1(np__6536)|~v1_int_1(np__25))).#trainneg
cnf(i_0_1784, plain, (v1_int_1(np__6305)|~v1_int_1(np__256))).#trainneg
cnf(i_0_1777, plain, (v1_int_1(np__6304)|~v1_int_1(np__257))).#trainneg
cnf(i_0_1765, plain, (v1_int_1(np__6558)|~v1_int_1(np__3))).#trainneg
cnf(i_0_1764, plain, (v1_int_1(np__6425)|~v1_int_1(np__136))).#trainneg
cnf(i_0_9886, plain, (k4_nat_d(k4_nat_1(X2,k4_nat_d(X3,X1)),X1)=k4_nat_d(k3_xcmplx_0(X3,k2_xcmplx_0(X1,X2)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_9905, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X3),X1),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_11754, plain, (X1=k2_nat_d(X2,X3)|r1_xxreal_0(X3,X1)|X2!=k2_xcmplx_0(X1,k3_xcmplx_0(X3,X4))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_11286, plain, (r1_xxreal_0(np__8,np__136)|r1_xxreal_0(np__136,np__8))).#trainneg
cnf(i_0_10493, plain, (np__0=np__2|np__6480!=np__0|~r1_xxreal_0(np__2,np__6480))).#trainneg
cnf(i_0_2393, plain, (k2_xcmplx_0(np__1,k1_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(k2_newton(np__2,X1))|~v7_ordinal1(X1)|~m1_subset_1(k1_newton(np__2,k2_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_10907, plain, (v1_xreal_0(np__78)|~v1_int_1(np__3))).#trainneg
cnf(i_0_9928, plain, (k2_nat_d(np__8,np__6)=np__2)).#trainneg
cnf(i_0_9957, plain, (k2_nat_d(np__5,np__3)=np__2)).#trainneg
cnf(i_0_1763, plain, (v1_int_1(np__6557)|~v1_int_1(np__4))).#trainneg
cnf(i_0_1761, plain, (v1_int_1(np__136)|~v1_int_1(np__6425))).#trainneg
cnf(i_0_10621, negated_conjecture, (k1_newton(np__3,np__8)!=np__3|np__136!=np__3|~v7_ordinal1(k1_newton(np__3,np__8)))).#trainneg
cnf(i_0_10629, negated_conjecture, (k13_newton(np__3,np__8)!=np__3|np__136!=np__3|~v7_ordinal1(k13_newton(np__3,np__8)))).#trainneg
cnf(i_0_10827, plain, (v1_xcmplx_0(np__6305)|~v1_xreal_0(np__256))).#trainneg
cnf(i_0_11853, plain, (r1_xxreal_0(np__1,np__3))).#trainneg
cnf(i_0_11723, plain, (X1=k2_nat_d(X2,np__3)|r1_xxreal_0(np__3,X1)|X2!=k2_xcmplx_0(X1,np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10875, plain, (k2_nat_d(k4_nat_1(k2_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k4_nat_1(k2_nat_d(X1,X2),X3))|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_9903, plain, (k2_nat_d(k4_nat_1(k4_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k4_nat_1(k4_nat_d(X1,X2),X3))|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_9689, plain, (m1_subset_1(k2_nat_d(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11722, plain, (X1=k2_nat_d(X2,np__8)|r1_xxreal_0(np__8,X1)|X2!=k2_xcmplx_0(X1,np__8)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_11740, plain, (X1=k2_nat_d(X2,np__2)|r1_xxreal_0(np__2,X1)|X2!=k2_xcmplx_0(X1,np__8)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10494, plain, (np__0=np__2|np__254!=np__0|~r1_xxreal_0(np__2,np__254))).#trainneg
cnf(i_0_2392, plain, (k2_xcmplx_0(np__1,k2_newton(np__2,k1_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(X1)|~m1_subset_1(k2_newton(np__2,k1_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_9958, plain, (k2_nat_d(np__4,np__3)=np__1)).#trainneg
cnf(i_0_11234, plain, (r1_xxreal_0(np__7,np__8))).#trainneg
cnf(i_0_11236, plain, (r1_xxreal_0(np__5,np__8))).#trainneg
cnf(i_0_10961, plain, (v1_xcmplx_0(np__6536)|~v1_xreal_0(np__25))).#trainneg
cnf(i_0_1443, plain, (~v1_abian(np__2)|~v1_abian(np__1))).#trainneg
cnf(i_0_1597, plain, (~v1_abian(np__7)|~v1_abian(np__6))).#trainneg
cnf(i_0_1586, plain, (~v1_abian(np__5)|~v1_abian(np__4))).#trainneg
cnf(i_0_13416, negated_conjecture, (k1_newton(np__3,np__8)!=np__3|np__136!=np__3)).#trainneg
cnf(i_0_12531, negated_conjecture, (r1_xxreal_0(np__257,k13_newton(np__3,np__8))|k13_newton(np__3,np__8)!=np__136|~v7_ordinal1(k13_newton(np__3,np__8)))).#trainneg
cnf(i_0_12915, plain, (k2_nat_d(X1,np__1)=np__0|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13129, plain, (k4_nat_d(k3_xcmplx_0(X3,k2_xcmplx_0(X2,k2_nat_d(X1,X2))),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1898, plain, (v1_int_1(np__3))).#trainneg
cnf(i_0_12109, plain, (np__136=np__0|np__6425!=np__0)).#trainneg
cnf(i_0_2356, plain, (k1_nat_1(np__1,k1_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(k2_newton(np__2,X1))|~v7_ordinal1(X1)|~m1_subset_1(k1_newton(np__2,k2_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_13014, plain, (np__0=np__3|np__257!=np__3)).#trainneg
cnf(i_0_12076, plain, (np__0=np__3|np__257!=np__0)).#trainneg
cnf(i_0_13791, plain, (np__0=np__1|np__0!=np__3)).#trainneg
cnf(i_0_7963, plain, (v1_abian(np__2))).#trainneg
cnf(i_0_1913, plain, (v1_int_1(np__5))).#trainneg
cnf(i_0_1516, plain, (v1_abian(np__8))).#trainneg
cnf(i_0_11992, plain, (k2_nat_d(X1,np__0)=X1|r1_xxreal_0(np__0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11942, plain, (k2_nat_d(np__81,np__81)=np__0)).#trainneg
cnf(i_0_12520, negated_conjecture, (r1_xxreal_0(np__257,k1_newton(np__3,np__8))|k1_newton(np__3,np__8)!=np__136|~v7_ordinal1(k1_newton(np__3,np__8)))).#trainneg
cnf(i_0_11882, plain, (k3_xcmplx_0(np__1,np__257)=np__257)).#trainneg
cnf(i_0_11890, plain, (k2_xcmplx_0(X1,np__257)=k2_xcmplx_0(np__257,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13167, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(X2,k2_xcmplx_0(X3,X1)),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11901, plain, (k4_nat_d(k4_nat_1(X1,k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_12012, plain, (r1_xxreal_0(X1,np__136)|r1_xxreal_0(np__136,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_12098, plain, (v1_int_1(np__136))).#trainneg
cnf(i_0_12749, plain, (np__0=np__1|np__0!=np__2)).#trainneg
cnf(i_0_2364, plain, (k1_nat_1(np__1,k2_newton(np__2,k1_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(X1)|~m1_subset_1(k2_newton(np__2,k1_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_13352, plain, (np__0=np__2|np__6!=np__0)).#trainneg
cnf(i_0_13419, plain, (v1_xcmplx_0(np__6305))).#trainneg
cnf(i_0_13817, plain, (v1_xcmplx_0(np__6536))).#trainneg
cnf(i_0_10965, plain, (v1_xcmplx_0(np__79)|~v1_xreal_0(np__81))).#trainneg
cnf(i_0_11311, plain, (v1_xcmplx_0(np__80)|~v1_xreal_0(np__81))).#trainneg
cnf(i_0_11584, plain, (v1_xcmplx_0(np__77)|~v1_xreal_0(np__81))).#trainneg
cnf(i_0_11588, plain, (v1_xcmplx_0(np__78)|~v1_xreal_0(np__81))).#trainneg
cnf(i_0_11909, plain, (k4_nat_d(k3_xcmplx_0(X1,k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11907, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11917, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X3,X2)),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_xcmplx_0(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_12554, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,np__257),X2),np__257)=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|r1_xxreal_0(np__257,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13128, plain, (k4_nat_d(k3_xcmplx_0(X3,k2_xcmplx_0(X2,k2_nat_d(X1,X2))),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11886, plain, (k2_xcmplx_0(np__257,X1)=k2_xcmplx_0(X1,np__257)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13072, plain, (np__8=np__0|X1!=np__0|~r1_xxreal_0(np__8,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2379, plain, (k2_xcmplx_0(np__1,k2_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(X1)|~m1_subset_1(k2_newton(np__2,k2_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_11268, plain, (r1_xxreal_0(np__2,np__8))).#trainneg
cnf(i_0_11271, plain, (r1_xxreal_0(np__6,np__8))).#trainneg
cnf(i_0_3729, plain, (v1_xreal_0(np__80)|~v1_int_1(np__81))).#trainneg
cnf(i_0_1908, plain, (v1_int_1(np__22)|~v1_int_1(np__21))).#trainneg
cnf(i_0_1907, plain, (v1_int_1(np__23)|~v1_int_1(np__22))).#trainneg
cnf(i_0_1903, plain, (v1_int_1(np__257)|~v1_int_1(np__256))).#trainneg
cnf(i_0_1902, plain, (v1_int_1(np__7)|~v1_int_1(np__6))).#trainneg
cnf(i_0_12861, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,k2_nat_d(X3,X2)),X1),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,k2_nat_d(X3,X2))|~v7_ordinal1(k2_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11912, plain, (k4_nat_d(k3_xcmplx_0(X1,k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11919, plain, (k4_nat_d(k3_xcmplx_0(k4_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X3,X2)),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_xcmplx_0(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_12557, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,np__257),X2),np__257)=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|r1_xxreal_0(np__257,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_12535, plain, (k4_nat_d(k4_nat_1(X1,k4_nat_d(X2,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8236, plain, (r1_xxreal_0(X1,np__8)|r1_xxreal_0(np__8,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_10553, plain, (np__0=np__2|np__8!=np__0)).#trainneg
cnf(i_0_2290, plain, (k2_nat_d(k2_xcmplx_0(np__81,X1),np__81)=X1|r1_xxreal_0(np__81,X1)|~v7_ordinal1(np__81)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14065, plain, (np__6425!=np__0)).#trainneg
cnf(i_0_3036, plain, (~v1_abian(np__3))).#trainneg
cnf(i_0_1443, plain, (~v1_abian(np__1))).#trainneg
cnf(i_0_13847, plain, (k2_nat_d(X1,np__257)=np__136|X1!=np__136|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14432, plain, (r1_xxreal_0(np__2,np__136))).#trainneg
cnf(i_0_13834, plain, (k2_nat_d(np__136,np__136)=np__0)).#trainneg
cnf(i_0_15447, plain, (np__8=np__0|np__136!=np__0|~r1_xxreal_0(np__8,np__136))).#trainneg
cnf(i_0_2285, plain, (k2_nat_d(k2_xcmplx_0(np__81,X1),np__1)=X1|r1_xxreal_0(np__1,X1)|~v7_ordinal1(np__81)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1765, plain, (v1_int_1(np__6558))).#trainneg
cnf(i_0_1469, plain, (~v1_abian(np__7)|~v1_int_1(np__7))).#trainneg
cnf(i_0_14128, plain, (X1=np__0|r1_xxreal_0(np__0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14668, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14438, plain, (r1_xxreal_0(np__1,np__136))).#trainneg
cnf(i_0_15475, plain, (v1_int_1(np__257))).#trainneg
cnf(i_0_13843, plain, (k2_nat_d(X1,np__136)=np__0|X1!=np__136|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8817, plain, (r1_xxreal_0(X1,np__3)|r1_xxreal_0(np__3,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_15388, plain, (np__8=np__0|np__6480!=np__0|~r1_xxreal_0(np__8,np__6480))).#trainneg
cnf(i_0_1463, plain, (X1=k2_nat_d(X2,np__81)|r1_xxreal_0(np__81,X1)|X2!=k2_xcmplx_0(np__81,X1)|~v7_ordinal1(np__81)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10907, plain, (v1_xreal_0(np__78))).#trainneg
cnf(i_0_1679, plain, (v1_abian(np__6))).#trainneg
cnf(i_0_1764, plain, (v1_int_1(np__6425))).#trainneg
cnf(i_0_14569, plain, (v1_xcmplx_0(np__79))).#trainneg
cnf(i_0_14572, plain, (v1_xcmplx_0(np__80))).#trainneg
cnf(i_0_14575, plain, (v1_xcmplx_0(np__77))).#trainneg
cnf(i_0_14578, plain, (v1_xcmplx_0(np__78))).#trainneg
cnf(i_0_13940, plain, (k2_nat_d(k2_xcmplx_0(np__257,X1),np__257)=X1|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15382, plain, (X1=k2_nat_d(X2,np__257)|r1_xxreal_0(np__257,X1)|X2!=k2_xcmplx_0(X1,np__257)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15121, plain, (k2_nat_d(X1,X2)=np__136|r1_xxreal_0(X2,np__136)|X1!=k2_xcmplx_0(X2,np__136)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13862, plain, (X1=k2_nat_d(X2,np__136)|r1_xxreal_0(np__136,X1)|X2!=k2_xcmplx_0(X1,np__136)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_15389, plain, (np__8=np__0|np__254!=np__0|~r1_xxreal_0(np__8,np__254))).#trainneg
cnf(i_0_1437, plain, (k1_nat_1(np__1,k2_newton(np__2,k2_newton(np__2,X1)))=k4_pepin(X1)|~v7_ordinal1(X1)|~m1_subset_1(k2_newton(np__2,k2_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_15469, plain, (v1_xreal_0(np__80))).#trainneg
cnf(i_0_15471, plain, (v1_int_1(np__22))).#trainneg
cnf(i_0_11280, plain, (r1_xxreal_0(np__1,np__8))).#trainneg
cnf(i_0_1768, plain, (v1_int_1(np__80)|~v1_int_1(np__81))).#trainneg
cnf(i_0_1967, plain, (v1_int_1(np__24)|~v1_int_1(np__4))).#trainneg
cnf(i_0_3730, plain, (v1_xreal_0(np__77)|~v1_int_1(np__81))).#trainneg
cnf(i_0_1769, plain, (v1_int_1(np__77)|~v1_int_1(np__81))).#trainneg
cnf(i_0_13821, negated_conjecture, (r1_xxreal_0(np__257,k13_newton(np__3,np__8))|k13_newton(np__3,np__8)!=np__136|~m1_subset_1(k13_newton(np__3,np__8),k5_numbers))).#trainneg
cnf(i_0_15090, plain, (k2_nat_d(k2_xcmplx_0(np__136,X1),np__136)=X1|r1_xxreal_0(np__136,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13859, plain, (X1=k2_nat_d(X2,np__1)|r1_xxreal_0(np__1,X1)|X2!=k2_xcmplx_0(X1,np__136)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_15459, plain, (k2_newton(np__2,k2_newton(np__2,X1))=k2_nat_d(k4_pepin(X1),np__1)|r1_xxreal_0(np__1,k2_newton(np__2,k2_newton(np__2,X1)))|~v7_ordinal1(k2_newton(np__2,k2_newton(np__2,X1)))|~v7_ordinal1(X1)|~m1_subset_1(k2_newton(np__2,k2_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_14052, plain, (k4_nat_d(k2_xcmplx_0(X1,k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__1,X2),X1)|r1_xxreal_0(X1,k2_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k2_xcmplx_0(X1,k2_nat_d(X2,X1)),k1_numbers))).#trainneg
cnf(i_0_14359, plain, (k4_nat_d(k3_xcmplx_0(X3,k2_xcmplx_0(k2_nat_d(X1,X2),X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_4250, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15390, plain, (np__8=np__0|np__253!=np__0|~r1_xxreal_0(np__8,np__253))).#trainneg
cnf(i_0_1446, plain, (X1=k2_nat_d(X2,np__1)|r1_xxreal_0(np__1,X1)|X2!=k2_xcmplx_0(np__81,X1)|~v7_ordinal1(np__81)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_11281, plain, (r1_xxreal_0(np__8,np__6561))).#trainneg
cnf(i_0_11285, plain, (r1_xxreal_0(np__4,np__8))).#trainneg
cnf(i_0_11511, plain, (r1_xxreal_0(np__25,np__136))).#trainneg
cnf(i_0_12736, plain, (v1_xreal_0(np__79)|~v1_int_1(np__81))).#trainneg
cnf(i_0_13087, plain, (v1_xreal_0(np__6536)|~v1_int_1(np__25))).#trainneg
cnf(i_0_13088, plain, (v1_xreal_0(np__6305)|~v1_int_1(np__256))).#trainneg
cnf(i_0_1770, plain, (v1_int_1(np__78)|~v1_int_1(np__81))).#trainneg
cnf(i_0_1766, plain, (v1_int_1(np__254)|~v1_int_1(np__257))).#trainneg
cnf(i_0_9884, plain, (k4_nat_d(k4_nat_1(X2,k4_nat_d(X3,X1)),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_16204, negated_conjecture, (k1_newton(np__3,np__8)!=np__136|~v7_ordinal1(k1_newton(np__3,np__8)))).#trainneg
cnf(i_0_16945, negated_conjecture, (k1_newton(np__3,np__8)!=k2_xcmplx_0(np__257,np__136)|~v7_ordinal1(k1_newton(np__3,np__8)))).#trainneg
cnf(i_0_16218, negated_conjecture, (k13_newton(np__3,np__8)!=np__136|~v7_ordinal1(k13_newton(np__3,np__8)))).#trainneg
cnf(i_0_17519, plain, (r1_xxreal_0(np__2,np__4))).#trainneg
cnf(i_0_16963, negated_conjecture, (k13_newton(np__3,np__8)!=k2_xcmplx_0(np__257,np__136)|~v7_ordinal1(k13_newton(np__3,np__8)))).#trainneg
cnf(i_0_17505, plain, (r1_xxreal_0(np__2,np__6))).#trainneg
cnf(i_0_16216, plain, (np__136=np__0|np__136!=np__257)).#trainneg
cnf(i_0_2596, plain, (X1=k2_nat_d(X2,k1_newton(X3,np__3))|r1_xxreal_0(k1_newton(X3,np__3),X1)|X2!=k2_xcmplx_0(k1_newton(X3,np__6),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,np__3))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10552, plain, (np__0=np__2|np__136!=np__0)).#trainneg
cnf(i_0_1777, plain, (v1_int_1(np__6304))).#trainneg
cnf(i_0_1907, plain, (v1_int_1(np__23))).#trainneg
cnf(i_0_17146, plain, (v1_xreal_0(np__77))).#trainneg
cnf(i_0_17537, plain, (v1_xreal_0(np__79))).#trainneg
cnf(i_0_17539, plain, (v1_xreal_0(np__6536))).#trainneg
cnf(i_0_17541, plain, (v1_xreal_0(np__6305))).#trainneg
cnf(i_0_16336, plain, (k2_nat_d(X1,np__6561)=np__136|X1!=np__136|~v7_ordinal1(X1))).#trainneg
cnf(i_0_17668, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X2,X3)),X2)=k4_nat_d(k4_nat_1(X1,X3),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16387, plain, (k2_nat_d(np__81,np__1)=np__0)).#trainneg
cnf(i_0_17451, plain, (k4_nat_d(k2_xcmplx_0(k2_nat_d(X1,X2),X2),X2)=k4_nat_d(k3_xcmplx_0(np__1,X1),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_xcmplx_0(k2_nat_d(X1,X2),X2),k1_numbers))).#trainneg
cnf(i_0_4252, plain, (r1_xxreal_0(X1,np__6561)|r1_xxreal_0(np__6561,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16231, plain, (np__136=np__0|np__257!=np__0|X1!=np__136|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2283, plain, (k2_nat_d(k2_xcmplx_0(np__6,X1),np__2)=X1|r1_xxreal_0(np__2,X1)|~v7_ordinal1(k2_xcmplx_0(np__6,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16957, plain, (np__136=np__0|k2_xcmplx_0(np__257,np__136)!=np__257)).#trainneg
cnf(i_0_11523, plain, (r1_xxreal_0(np__136,np__6425))).#trainneg
cnf(i_0_17514, plain, (r1_xxreal_0(np__1,np__2))).#trainneg
cnf(i_0_17450, plain, (k4_nat_d(k2_xcmplx_0(k2_nat_d(X1,X2),X2),X2)=k4_nat_d(k3_xcmplx_0(np__1,X1),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_xcmplx_0(k2_nat_d(X1,X2),X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_17884, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X3,X1),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15391, plain, (np__8=np__0|np__255!=np__0|~r1_xxreal_0(np__8,np__255))).#trainneg
cnf(i_0_2646, plain, (k2_nat_d(X1,np__257)=np__81|X1!=np__81|~v7_ordinal1(np__81)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11525, plain, (r1_xxreal_0(np__136,np__6561))).#trainneg
cnf(i_0_11527, plain, (r1_xxreal_0(np__136,np__257))).#trainneg
cnf(i_0_13418, negated_conjecture, (k13_newton(np__3,np__8)!=np__3|np__136!=np__3|~m1_subset_1(k13_newton(np__3,np__8),k5_numbers))).#trainneg
cnf(i_0_3651, plain, (v1_xreal_0(np__256)|~m1_subset_1(k4_xcmplx_0(np__256),k1_numbers))).#trainneg
cnf(i_0_3655, plain, (v1_xcmplx_0(np__256)|~m1_subset_1(k4_xcmplx_0(np__256),k1_numbers))).#trainneg
cnf(i_0_16824, plain, (X1=k2_nat_d(X2,np__257)|r1_xxreal_0(np__257,X1)|X2!=k2_xcmplx_0(np__257,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_17106, plain, (k2_nat_d(k2_xcmplx_0(X1,np__136),np__136)=X1|r1_xxreal_0(np__136,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__136))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16823, plain, (k2_nat_d(X1,np__257)=np__136|X1!=k2_xcmplx_0(np__257,np__136)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16475, plain, (k2_nat_d(k4_nat_1(X1,k2_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k4_nat_1(X1,k2_nat_d(X2,X3)))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15392, plain, (np__8=np__0|np__120!=np__0|~r1_xxreal_0(np__8,np__120))).#trainneg
cnf(i_0_11528, plain, (r1_xxreal_0(np__81,np__136))).#trainneg
cnf(i_0_11529, plain, (r1_xxreal_0(np__4,np__136))).#trainneg
cnf(i_0_11840, plain, (r1_xxreal_0(np__3,np__25))).#trainneg
cnf(i_0_3747, plain, (v1_xcmplx_0(np__175)|~v1_xreal_0(np__81)|~v1_xreal_0(np__256))).#trainneg
cnf(i_0_3751, plain, (v1_xreal_0(np__175)|~v1_int_1(np__81)|~v1_int_1(np__256))).#trainneg
cnf(i_0_11019, plain, (v1_int_1(np__81)|~v7_ordinal1(k13_newton(np__3,np__4)))).#trainneg
cnf(i_0_11434, plain, (v1_int_1(np__81)|~v7_ordinal1(k1_newton(np__3,np__4)))).#trainneg
cnf(i_0_18006, plain, (k4_nat_d(k3_xcmplx_0(X1,k2_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k2_nat_d(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_17081, plain, (X1=np__136|r1_xxreal_0(np__257,X1)|k2_xcmplx_0(X1,np__257)!=np__136|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16812, plain, (k2_nat_d(X1,np__257)=np__0|X1!=np__257|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16522, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)=k4_nat_d(k3_xcmplx_0(X2,k2_xcmplx_0(X1,X3)),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18080, negated_conjecture, (k1_newton(np__3,np__8)!=k2_xcmplx_0(np__257,np__136))).#trainneg
cnf(i_0_18708, plain, (np__257!=np__0)).#trainneg
cnf(i_0_19233, negated_conjecture, (k13_newton(np__3,np__8)!=np__3|np__136!=np__3)).#trainneg
cnf(i_0_19162, plain, (np__81=np__0|np__257!=np__81)).#trainneg
cnf(i_0_2559, plain, (X1=k2_nat_d(X2,k1_newton(X3,np__0))|r1_xxreal_0(k1_newton(X3,np__0),X1)|X2!=k2_xcmplx_0(k1_newton(X3,np__257),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,np__0))|~v7_ordinal1(k1_newton(X3,np__257))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_18078, negated_conjecture, (k1_newton(np__3,np__8)!=np__136)).#trainneg
cnf(i_0_18082, negated_conjecture, (k13_newton(np__3,np__8)!=np__136|~m1_subset_1(k13_newton(np__3,np__8),k5_numbers))).#trainneg
cnf(i_0_18084, negated_conjecture, (k13_newton(np__3,np__8)!=k2_xcmplx_0(np__257,np__136)|~m1_subset_1(k13_newton(np__3,np__8),k5_numbers))).#trainneg
cnf(i_0_19725, plain, (X1=np__136|r1_xxreal_0(np__257,X1)|k2_xcmplx_0(X1,np__257)!=np__136|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20250, plain, (k4_nat_d(k3_xcmplx_0(np__1,k2_xcmplx_0(X1,X2)),X2)=k4_nat_d(k2_xcmplx_0(X1,X2),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers))).#trainneg
cnf(i_0_18777, plain, (k2_nat_d(np__7,np__2)=np__1)).#trainneg
cnf(i_0_18764, plain, (k2_nat_d(np__6,np__2)=np__0)).#trainneg
cnf(i_0_19163, negated_conjecture, (k13_newton(np__3,np__8)!=np__81|np__136!=np__81|~v7_ordinal1(k13_newton(np__3,np__8)))).#trainneg
cnf(i_0_18162, plain, (np__136=np__0|np__0!=np__6561|X1!=np__136|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18595, plain, (r1_xxreal_0(np__0,np__6561))).#trainneg
cnf(i_0_19809, plain, (np__0=X1|r1_xxreal_0(np__257,X1)|X1!=np__257|~v7_ordinal1(X1))).#trainneg
cnf(i_0_19573, plain, (v1_int_1(np__81))).#trainneg
cnf(i_0_18410, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),k2_xcmplx_0(X2,X3)),X2)=k4_nat_d(k4_nat_1(X1,X3),X2)|r1_xxreal_0(X2,X3)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_19428, plain, (np__136=X1|r1_xxreal_0(np__257,X1)|X1!=k2_xcmplx_0(np__257,np__136)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15321, plain, (k4_nat_d(k2_xcmplx_0(X1,k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__1),X1)|r1_xxreal_0(X1,k2_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k2_xcmplx_0(X1,k2_nat_d(X2,X1)),k1_numbers))).#trainneg
cnf(i_0_19749, negated_conjecture, (k13_newton(np__3,np__8)!=np__257|np__136!=np__0|~v7_ordinal1(k13_newton(np__3,np__8)))).#trainneg
cnf(i_0_9902, plain, (k4_nat_d(k3_xcmplx_0(k4_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X3),X1),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_18047, plain, (X1=np__0|X1!=np__136|~r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(np__257,np__136))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13960, plain, (m1_subset_1(k2_xcmplx_0(X1,np__257),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14204, plain, (X1=k2_nat_d(X2,np__1)|r1_xxreal_0(np__1,X1)|X2!=k2_xcmplx_0(X1,np__257)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_19520, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X2,X3)),X2)=k2_nat_d(k4_nat_1(X1,X3),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k4_nat_1(X1,X3))|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_12540, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(X2,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4942, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_19811, plain, (np__136=np__0|X1!=np__136|X1!=np__257|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2278, plain, (k2_nat_d(k2_xcmplx_0(np__0,X1),np__0)=X1|r1_xxreal_0(np__0,X1)|~v7_ordinal1(k2_xcmplx_0(np__0,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9894, plain, (k4_nat_d(k3_xcmplx_0(X2,k4_nat_d(X3,X1)),X1)=k4_nat_d(k3_xcmplx_0(X3,k2_xcmplx_0(X1,X2)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13212, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,k4_nat_d(X3,X2)),X1),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|r1_xxreal_0(X2,k4_nat_d(X3,X2))|~v7_ordinal1(k4_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_11879, plain, (k2_xcmplx_0(X1,np__136)=k2_xcmplx_0(np__136,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12537, plain, (k4_nat_d(k4_nat_1(X1,k4_nat_d(X2,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(X2,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_20276, negated_conjecture, (k13_newton(np__3,np__8)!=k2_xcmplx_0(np__257,np__136))).#trainneg
cnf(i_0_20274, negated_conjecture, (k13_newton(np__3,np__8)!=np__136)).#trainneg
cnf(i_0_21971, plain, (np__136=np__0|np__136!=np__8|np__8!=np__257)).#trainneg
cnf(i_0_21972, plain, (np__136=np__0|np__136!=np__3|np__257!=np__3)).#trainneg
cnf(i_0_2263, plain, (k4_nat_d(k4_nat_1(np__81,k4_nat_d(X1,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(X1,k13_newton(np__3,np__4)),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(k4_nat_d(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21181, plain, (k2_nat_d(np__3,np__257)=np__3)).#trainneg
cnf(i_0_20751, plain, (v1_xreal_0(np__81))).#trainneg
cnf(i_0_21872, plain, (r1_xxreal_0(np__2,np__5))).#trainneg
cnf(i_0_21336, plain, (X1=np__0|X1!=np__136|~r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21165, plain, (k2_nat_d(np__0,np__257)=np__0)).#trainneg
cnf(i_0_22389, plain, (k2_xcmplx_0(np__136,np__257)=k2_xcmplx_0(np__257,np__136))).#trainneg
cnf(i_0_4944, plain, (r1_xxreal_0(X1,np__6561)|r1_xxreal_0(np__6561,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_20603, plain, (np__136=np__0|np__0!=np__6561|np__136!=np__8)).#trainneg
cnf(i_0_2258, plain, (k4_nat_d(k4_nat_1(np__81,k4_nat_d(X1,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(X1,k1_newton(np__3,np__4)),np__257)|~v7_ordinal1(k1_newton(np__3,np__4))|~v7_ordinal1(k4_nat_d(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21167, plain, (k2_nat_d(k2_xcmplx_0(X1,np__0),np__257)=X1|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__0))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20678, plain, (k2_nat_d(k2_xcmplx_0(X1,np__6425),np__257)=X1|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__6425))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_22188, plain, (k4_nat_d(k3_xcmplx_0(X1,k2_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(X2,k2_xcmplx_0(X3,X1)),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20674, plain, (k2_nat_d(np__6425,np__257)=np__0)).#trainneg
cnf(i_0_21518, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X2)),X2)=k2_nat_d(k4_nat_1(X1,X3),X2)|r1_xxreal_0(X2,X3)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k4_nat_1(X1,X3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_7367, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20604, plain, (np__136=np__0|np__0!=np__6561|np__136!=np__3)).#trainneg
cnf(i_0_2256, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,np__257),np__81),np__257)=k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),X1),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(np__81)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1770, plain, (v1_int_1(np__78))).#trainneg
cnf(i_0_1769, plain, (v1_int_1(np__77))).#trainneg
cnf(i_0_1768, plain, (v1_int_1(np__80))).#trainneg
cnf(i_0_21317, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X3),X1),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_22882, plain, (v1_xcmplx_0(np__81))).#trainneg
cnf(i_0_23533, plain, (k2_nat_d(np__4,np__257)=np__4)).#trainneg
cnf(i_0_23013, plain, (np__136!=np__257)).#trainneg
cnf(i_0_23860, plain, (np__136=np__0|np__136!=np__6425)).#trainneg
cnf(i_0_2251, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,np__257),np__81),np__257)=k4_nat_d(k3_xcmplx_0(k1_newton(np__3,np__4),X1),np__257)|~v7_ordinal1(k1_newton(np__3,np__4))|~v7_ordinal1(np__81)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_22807, plain, (np__136=np__3|np__6561!=np__3)).#trainneg
cnf(i_0_23852, plain, (np__0=np__3|np__6425!=np__3)).#trainneg
cnf(i_0_20602, plain, (np__136=np__0|np__0!=np__6561)).#trainneg
cnf(i_0_23009, plain, (np__136!=np__6425|~r1_xxreal_0(np__257,np__6425))).#trainneg
cnf(i_0_1767, plain, (v1_int_1(np__79))).#trainneg
cnf(i_0_1762, plain, (v1_int_1(np__6480))).#trainneg
cnf(i_0_19150, negated_conjecture, (k1_newton(np__3,np__8)!=np__81|np__136!=np__81|~v7_ordinal1(k1_newton(np__3,np__8)))).#trainneg
cnf(i_0_19230, plain, (v1_xreal_0(np__256)|~v1_xcmplx_0(np__256))).#trainneg
cnf(i_0_24153, plain, (r1_xxreal_0(np__1,np__256))).#trainneg
cnf(i_0_23019, plain, (X1=np__0|X1!=np__136|~r1_xxreal_0(np__257,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_24156, plain, (r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_24147, plain, (r1_xxreal_0(np__1,np__6)|r1_xxreal_0(np__6,np__1))).#trainneg
cnf(i_0_24620, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X2,X3)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_24161, plain, (r1_xxreal_0(np__1,np__4)|r1_xxreal_0(np__4,np__1))).#trainneg
cnf(i_0_24110, plain, (r1_xxreal_0(np__1,np__7)|r1_xxreal_0(np__7,np__1))).#trainneg
cnf(i_0_24112, plain, (r1_xxreal_0(np__1,np__5)|r1_xxreal_0(np__5,np__1))).#trainneg
cnf(i_0_22713, plain, (X1!=np__136|~r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7390, plain, (r1_xxreal_0(X1,np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21952, plain, (np__136=np__0|np__136!=np__2|np__257!=np__2)).#trainneg
cnf(i_0_2202, plain, (k4_nat_d(k4_nat_1(np__81,k4_nat_d(X1,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(k1_newton(np__3,np__4),X1),np__257)|~v7_ordinal1(k1_newton(np__3,np__4))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_22812, plain, (np__136=np__81|np__81!=np__6561)).#trainneg
cnf(i_0_22815, plain, (np__136=np__0|np__257!=np__6561)).#trainneg
cnf(i_0_20728, plain, (np__25=np__0|np__25!=np__257)).#trainneg
cnf(i_0_19735, negated_conjecture, (k1_newton(np__3,np__8)!=np__257|np__136!=np__0|~v7_ordinal1(k1_newton(np__3,np__8)))).#trainneg
cnf(i_0_3751, plain, (v1_xreal_0(np__175)|~v1_int_1(np__256))).#trainneg
cnf(i_0_1787, plain, (v1_int_1(np__175)|~v1_int_1(np__256))).#trainneg
cnf(i_0_23858, plain, (r1_xxreal_0(np__257,np__6425))).#trainneg
cnf(i_0_24114, plain, (r1_xxreal_0(np__0,np__1))).#trainneg
cnf(i_0_7393, plain, (r1_xxreal_0(X1,np__81)|r1_xxreal_0(np__81,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21955, plain, (np__136=np__0|np__136!=np__6|np__6!=np__257)).#trainneg
cnf(i_0_2204, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,np__257),np__81),np__257)=k4_nat_d(k3_xcmplx_0(X1,k1_newton(np__3,np__4)),np__257)|~v7_ordinal1(k1_newton(np__3,np__4))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20729, plain, (np__0=np__2|np__257!=np__2)).#trainneg
cnf(i_0_20738, plain, (np__256=np__0|np__256!=np__257)).#trainneg
cnf(i_0_20741, plain, (np__0=np__1|np__257!=np__1)).#trainneg
cnf(i_0_3747, plain, (v1_xcmplx_0(np__175)|~v1_xreal_0(np__256))).#trainneg
cnf(i_0_1859, plain, (v1_xreal_0(np__175)|~v1_xreal_0(np__256))).#trainneg
cnf(i_0_20540, negated_conjecture, (k13_newton(np__3,np__8)!=np__81|np__136!=np__81|~m1_subset_1(k13_newton(np__3,np__8),k5_numbers))).#trainneg
cnf(i_0_21095, negated_conjecture, (k13_newton(np__3,np__8)!=np__257|np__136!=np__0|~m1_subset_1(k13_newton(np__3,np__8),k5_numbers))).#trainneg
cnf(i_0_24602, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X2,np__136)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__136),X2)|r1_xxreal_0(X2,np__136)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_23046, plain, (X1!=np__0|~r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_23612, plain, (k2_xcmplx_0(X1,np__6425)=X1|r1_xxreal_0(np__257,k2_xcmplx_0(X1,np__6425))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__6425))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20460, plain, (X1=np__0|X1!=np__1|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15320, plain, (k4_nat_d(k2_xcmplx_0(X1,k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__1),X1)|r1_xxreal_0(X1,k2_nat_d(X2,X1))|~v7_ordinal1(k2_xcmplx_0(X1,k2_nat_d(X2,X1)))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_7396, plain, (r1_xxreal_0(X1,np__4)|r1_xxreal_0(np__4,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21961, plain, (np__136=np__0|np__136!=np__256|np__256!=np__257)).#trainneg
cnf(i_0_2205, plain, (k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),k1_newton(np__3,np__4)),np__257)=k4_nat_d(k4_nat_1(np__81,np__81),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(k1_newton(np__3,np__4)))).#trainneg
cnf(i_0_20746, plain, (np__0=np__4|np__257!=np__4)).#trainneg
cnf(i_0_23857, plain, (np__81=np__0|np__6425!=np__81)).#trainneg
cnf(i_0_7392, plain, (r1_xxreal_0(np__81,np__81))).#trainneg
cnf(i_0_7395, plain, (r1_xxreal_0(np__4,np__4))).#trainneg
cnf(i_0_18884, plain, (k4_nat_d(k2_xcmplx_0(X2,k2_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__1,X1),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_xcmplx_0(X2,k2_nat_d(X1,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X1,X2),k1_numbers))).#trainneg
cnf(i_0_9790, plain, (k2_nat_d(X1,X1)=np__0|r1_xxreal_0(X1,np__0)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_26805, plain, (r1_xxreal_0(np__0,np__4))).#trainneg
cnf(i_0_25490, plain, (X1!=np__136|~r1_xxreal_0(np__257,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_26855, plain, (r1_xxreal_0(np__3,np__4))).#trainneg
cnf(i_0_26575, plain, (X1!=np__0|~r1_xxreal_0(np__257,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_25902, plain, (k4_nat_d(k4_nat_1(X3,X1),X2)=k4_nat_d(k4_nat_1(X1,X3),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_24736, plain, (k2_nat_d(k4_nat_1(X1,X2),X3)=k4_nat_d(k3_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,X2)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k4_nat_1(X1,X2))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_24857, plain, (X1!=np__4|~r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_23009, plain, (np__136!=np__6425)).#trainneg
cnf(i_0_26656, plain, (np__136=np__0|np__136!=np__1)).#trainneg
cnf(i_0_25054, plain, (np__6561!=np__3)).#trainneg
cnf(i_0_25116, plain, (np__0!=np__6561)).#trainneg
cnf(i_0_25779, plain, (np__257!=np__6561)).#trainneg
cnf(i_0_26067, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),k2_xcmplx_0(X3,X2)),X2)=k4_nat_d(k4_nat_1(X3,X1),X2)|r1_xxreal_0(X2,X3)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_26657, plain, (np__8=np__0|np__8!=np__1)).#trainneg
cnf(i_0_27061, plain, (np__136=np__0|np__136!=np__6561)).#trainneg
cnf(i_0_11956, plain, (k2_nat_d(np__6,np__4)=np__2)).#trainneg
cnf(i_0_11884, plain, (k3_xcmplx_0(np__257,X1)=k3_xcmplx_0(X1,np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_12862, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,k2_nat_d(X3,X2)),X1),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|r1_xxreal_0(X2,k2_nat_d(X3,X2))|~v7_ordinal1(k2_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_26323, negated_conjecture, (k13_newton(np__3,np__8)!=np__81|np__136!=np__81)).#trainneg
cnf(i_0_2231, plain, (k4_nat_d(k4_nat_1(np__81,k4_nat_d(X1,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(X1,k1_newton(np__3,np__4)),np__257)|~v7_ordinal1(k1_newton(np__3,np__4))|~v7_ordinal1(np__81)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_22845, plain, (X1=np__0|X1!=np__3|~r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_22919, plain, (X1=np__0|X1!=np__81|~r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_12560, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(X1,np__257),X2),np__257)=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|r1_xxreal_0(np__257,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15088, plain, (k2_nat_d(k2_xcmplx_0(np__136,X1),X1)=np__136|r1_xxreal_0(X1,np__136)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_26325, negated_conjecture, (k13_newton(np__3,np__8)!=np__257|np__136!=np__0)).#trainneg
cnf(i_0_2235, plain, (k4_nat_d(k4_nat_1(np__81,k4_nat_d(X1,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(X1,k13_newton(np__3,np__4)),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(np__81)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_25773, plain, (np__257!=np__81|np__81!=np__6561)).#trainneg
cnf(i_0_26658, plain, (np__0=np__3|np__3!=np__1)).#trainneg
cnf(i_0_25428, plain, (np__6480!=np__136|~r1_xxreal_0(np__257,np__6480))).#trainneg
cnf(i_0_25429, plain, (np__254!=np__136|~r1_xxreal_0(np__257,np__254))).#trainneg
cnf(i_0_25430, plain, (np__253!=np__136|~r1_xxreal_0(np__257,np__253))).#trainneg
cnf(i_0_25431, plain, (np__255!=np__136|~r1_xxreal_0(np__257,np__255))).#trainneg
cnf(i_0_12541, plain, (k4_nat_d(k3_xcmplx_0(X1,k2_nat_d(X2,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_20424, plain, (k4_nat_d(k3_xcmplx_0(np__1,k2_xcmplx_0(X2,X1)),X2)=k4_nat_d(k2_xcmplx_0(X2,X1),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_16930, plain, (k4_nat_d(k2_xcmplx_0(X1,k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__1,X2),X1)|r1_xxreal_0(X1,k2_nat_d(X2,X1))|~v7_ordinal1(k2_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_7750, plain, (r1_xxreal_0(X1,np__256)|r1_xxreal_0(np__256,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21968, plain, (np__136=np__0|np__136!=np__81|np__257!=np__81)).#trainneg
cnf(i_0_2133, plain, (k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),X1),np__257)=k4_nat_d(k3_xcmplx_0(np__81,k4_nat_d(X1,np__257)),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(k4_nat_d(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_26607, plain, (np__5=np__0|np__5!=np__1)).#trainneg
cnf(i_0_26641, plain, (np__6=np__0|np__6!=np__1)).#trainneg
cnf(i_0_25432, plain, (np__120!=np__136|~r1_xxreal_0(np__257,np__120))).#trainneg
cnf(i_0_25433, plain, (np__55!=np__136|~r1_xxreal_0(np__257,np__55))).#trainneg
cnf(i_0_25434, plain, (np__6305!=np__136|~r1_xxreal_0(np__257,np__6305))).#trainneg
cnf(i_0_25435, plain, (np__7!=np__136|~r1_xxreal_0(np__257,np__7))).#trainneg
cnf(i_0_25437, plain, (np__5!=np__136|~r1_xxreal_0(np__257,np__5))).#trainneg
cnf(i_0_19322, plain, (X1=np__0|r1_xxreal_0(np__136,X1)|k2_xcmplx_0(X1,np__136)!=np__136|~v7_ordinal1(k2_xcmplx_0(X1,np__136))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_25982, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X2,X3)),X2)=k4_nat_d(k4_nat_1(X3,X1),X2)|r1_xxreal_0(X2,X3)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_28421, plain, (np__257!=np__81)).#trainneg
cnf(i_0_29474, plain, (k4_nat_d(k2_xcmplx_0(np__257,np__136),np__257)=k4_nat_d(np__136,np__257))).#trainneg
cnf(i_0_28569, plain, (k4_nat_d(np__3,np__2)=np__1)).#trainneg
cnf(i_0_28560, plain, (k4_nat_d(np__81,np__81)=np__0)).#trainneg
cnf(i_0_27487, plain, (np__257!=np__4)).#trainneg
cnf(i_0_28565, plain, (k4_nat_d(np__136,np__136)=np__0)).#trainneg
cnf(i_0_29039, plain, (v1_xxreal_0(np__3)|np__3!=np__1)).#trainneg
cnf(i_0_29064, plain, (np__3=np__2|np__3!=np__1)).#trainneg
cnf(i_0_28995, plain, (np__3!=np__1)).#trainneg
cnf(i_0_27483, plain, (np__6425!=np__4)).#trainneg
cnf(i_0_28352, plain, (np__6425!=np__3)).#trainneg
cnf(i_0_28356, plain, (np__257!=np__3)).#trainneg
cnf(i_0_28417, plain, (np__6425!=np__81)).#trainneg
cnf(i_0_28515, plain, (k4_nat_d(np__0,np__257)=np__0)).#trainneg
cnf(i_0_29423, plain, (k4_nat_d(np__257,np__257)=k4_nat_d(np__0,np__257))).#trainneg
cnf(i_0_27493, plain, (X1!=np__4|~r1_xxreal_0(np__257,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_28427, plain, (X1!=np__81|~r1_xxreal_0(np__257,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_28362, plain, (X1!=np__3|~r1_xxreal_0(np__257,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_28567, plain, (k4_nat_d(np__8,np__6)=np__2)).#trainneg
cnf(i_0_28571, plain, (k4_nat_d(np__3,np__3)=np__0)).#trainneg
cnf(i_0_28568, plain, (k4_nat_d(np__8,np__8)=np__0)).#trainneg
cnf(i_0_27964, plain, (X1=k2_nat_d(X2,X3)|r1_xxreal_0(X3,X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(X3,np__257),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_29488, plain, (k4_nat_d(k3_xcmplx_0(np__1,X1),np__257)=k4_nat_d(k2_xcmplx_0(np__257,np__136),np__257)|X1!=np__136|~v7_ordinal1(X1))).#trainneg
cnf(i_0_27943, plain, (X1=k2_nat_d(X2,X3)|r1_xxreal_0(X3,X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(np__257,X3),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2127, plain, (k4_nat_d(k4_nat_1(np__81,k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),X1),np__257)),np__257)=k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),k4_nat_1(np__81,k4_nat_d(X1,np__257))),np__257)|~v7_ordinal1(k4_nat_1(np__81,k4_nat_d(X1,np__257)))|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_25438, plain, (np__252!=np__136|~r1_xxreal_0(np__257,np__252))).#trainneg
cnf(i_0_25440, plain, (np__6558!=np__136|~r1_xxreal_0(np__257,np__6558))).#trainneg
cnf(i_0_25441, plain, (np__6424!=np__136|~r1_xxreal_0(np__257,np__6424))).#trainneg
cnf(i_0_25442, plain, (np__6304!=np__136|~r1_xxreal_0(np__257,np__6304))).#trainneg
cnf(i_0_25443, plain, (np__80!=np__136|~r1_xxreal_0(np__257,np__80))).#trainneg
cnf(i_0_11889, plain, (k3_xcmplx_0(X1,np__257)=k3_xcmplx_0(np__257,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_18074, plain, (k2_xcmplx_0(np__257,np__136)=np__136|r1_xxreal_0(np__257,k2_xcmplx_0(np__257,np__136))|~v7_ordinal1(k2_xcmplx_0(np__257,np__136)))).#trainneg
cnf(i_0_12963, plain, (k2_nat_d(X1,np__1)=X1|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21911, plain, (np__136=np__0|np__6480!=np__136|np__6480!=np__257)).#trainneg
cnf(i_0_2125, plain, (k2_nat_d(k4_nat_1(np__81,k4_nat_d(X1,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),X1),np__257)|~v7_ordinal1(k4_nat_1(np__81,k4_nat_d(X1,np__257)))|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_25444, plain, (np__134!=np__136|~r1_xxreal_0(np__257,np__134))).#trainneg
cnf(i_0_25445, plain, (np__232!=np__136|~r1_xxreal_0(np__257,np__232))).#trainneg
cnf(i_0_25446, plain, (np__6400!=np__136|~r1_xxreal_0(np__257,np__6400))).#trainneg
cnf(i_0_25447, plain, (np__6557!=np__136|~r1_xxreal_0(np__257,np__6557))).#trainneg
cnf(i_0_25448, plain, (np__6344!=np__136|~r1_xxreal_0(np__257,np__6344))).#trainneg
cnf(i_0_23614, plain, (X1=np__136|r1_xxreal_0(np__257,X1)|k2_xcmplx_0(X1,np__6425)!=np__136|~v7_ordinal1(k2_xcmplx_0(X1,np__6425))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9731, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_14588, plain, (k2_nat_d(k4_nat_1(X1,k4_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k4_nat_1(X1,k4_nat_d(X2,X3)))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11913, plain, (k4_nat_d(k3_xcmplx_0(X1,k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X3,X2),k1_numbers))).#trainneg
cnf(i_0_24713, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X3,X2),k2_xcmplx_0(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_xcmplx_0(X3,X2))|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_27938, plain, (k2_nat_d(k2_xcmplx_0(k3_xcmplx_0(np__257,X1),X2),X1)=X2|r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(k3_xcmplx_0(np__257,X1),X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_30808, plain, (k3_xcmplx_0(np__257,np__1)=np__257)).#trainneg
cnf(i_0_30636, plain, (k4_nat_d(X1,np__257)=np__136|X1!=np__136|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_30910, plain, (X1=np__0|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21912, plain, (np__136=np__0|np__254!=np__136|np__254!=np__257)).#trainneg
cnf(i_0_2091, plain, (k2_nat_d(k4_nat_1(k4_nat_d(X1,np__257),np__81),np__257)=k4_nat_d(k3_xcmplx_0(X1,k13_newton(np__3,np__4)),np__257)|~v7_ordinal1(k4_nat_1(k4_nat_d(X1,np__257),np__81))|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_27857, plain, (np__81=np__0|np__81!=np__6561)).#trainneg
cnf(i_0_27870, plain, (np__0=np__2|np__0!=np__4)).#trainneg
cnf(i_0_25450, plain, (np__22!=np__136|~r1_xxreal_0(np__257,np__22))).#trainneg
cnf(i_0_25451, plain, (np__111!=np__136|~r1_xxreal_0(np__257,np__111))).#trainneg
cnf(i_0_25452, plain, (np__121!=np__136|~r1_xxreal_0(np__257,np__121))).#trainneg
cnf(i_0_25453, plain, (np__132!=np__136|~r1_xxreal_0(np__257,np__132))).#trainneg
cnf(i_0_25454, plain, (np__133!=np__136|~r1_xxreal_0(np__257,np__133))).#trainneg
cnf(i_0_31961, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),k2_xcmplx_0(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_31880, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X2,X3)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_7763, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_21913, plain, (np__136=np__0|np__253!=np__136|np__253!=np__257)).#trainneg
cnf(i_0_25724, plain, (np__6480=np__6425|np__81!=np__6561)).#trainneg
cnf(i_0_26638, plain, (np__0=np__2|np__1!=np__2)).#trainneg
cnf(i_0_26651, plain, (np__0=np__6561|np__6561!=np__1)).#trainneg
cnf(i_0_25455, plain, (np__175!=np__136|~r1_xxreal_0(np__257,np__175))).#trainneg
cnf(i_0_25456, plain, (np__176!=np__136|~r1_xxreal_0(np__257,np__176))).#trainneg
cnf(i_0_25457, plain, (np__21!=np__136|~r1_xxreal_0(np__257,np__21))).#trainneg
cnf(i_0_25458, plain, (np__6169!=np__136|~r1_xxreal_0(np__257,np__6169))).#trainneg
cnf(i_0_25459, plain, (np__6421!=np__136|~r1_xxreal_0(np__257,np__6421))).#trainneg
cnf(i_0_31529, plain, (k2_nat_d(k3_xcmplx_0(X1,k4_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k3_xcmplx_0(X1,k4_nat_d(X2,X3)))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X2,X3),k1_numbers))).#trainneg
cnf(i_0_31118, plain, (X1=np__136|r1_xxreal_0(np__257,X1)|k2_xcmplx_0(X1,np__6425)!=np__136|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7774, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_21914, plain, (np__136=np__0|np__255!=np__136|np__255!=np__257)).#trainneg
cnf(i_0_2096, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,np__257),np__81),np__257)=k4_nat_d(k3_xcmplx_0(X1,k13_newton(np__3,np__4)),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_26655, plain, (np__0=np__4|np__4!=np__1)).#trainneg
cnf(i_0_21915, plain, (np__136=np__0|np__120!=np__136|np__120!=np__257)).#trainneg
cnf(i_0_21916, plain, (np__136=np__0|np__55!=np__136|np__55!=np__257)).#trainneg
cnf(i_0_25460, plain, (np__6422!=np__136|~r1_xxreal_0(np__257,np__6422))).#trainneg
cnf(i_0_25461, plain, (np__6423!=np__136|~r1_xxreal_0(np__257,np__6423))).#trainneg
cnf(i_0_25462, plain, (np__6559!=np__136|~r1_xxreal_0(np__257,np__6559))).#trainneg
cnf(i_0_25463, plain, (np__6560!=np__136|~r1_xxreal_0(np__257,np__6560))).#trainneg
cnf(i_0_25464, plain, (np__77!=np__136|~r1_xxreal_0(np__257,np__77))).#trainneg
cnf(i_0_31719, plain, (k4_nat_d(k3_xcmplx_0(k4_nat_d(X2,X3),X1),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X2,X3),k1_numbers))).#trainneg
cnf(i_0_31504, plain, (k2_nat_d(k4_nat_1(k4_nat_d(X2,X3),X1),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k4_nat_1(k4_nat_d(X2,X3),X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X2,X3),k5_numbers))).#trainneg
cnf(i_0_31227, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X2,X3)),X2)=k2_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_31503, plain, (k2_nat_d(k3_xcmplx_0(k4_nat_d(X2,X3),X1),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k3_xcmplx_0(k4_nat_d(X2,X3),X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X2,X3),k5_numbers))).#trainneg
cnf(i_0_33677, plain, (np__8=np__4|np__4!=np__1)).#trainneg
cnf(i_0_2097, plain, (k4_nat_d(k3_xcmplx_0(X1,k13_newton(np__3,np__4)),np__257)=k4_nat_d(k3_xcmplx_0(k4_nat_d(X1,np__257),np__81),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(np__81)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_33228, plain, (np__1!=np__2)).#trainneg
cnf(i_0_33661, plain, (np__4!=np__1)).#trainneg
cnf(i_0_34275, plain, (k4_nat_d(k4_nat_1(X1,X3),X2)=k2_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_32719, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X1,X3)),X1)=k4_nat_d(k3_xcmplx_0(X3,X2),X1)|r1_xxreal_0(X1,X3)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_25465, plain, (np__78!=np__136|~r1_xxreal_0(np__257,np__78))).#trainneg
cnf(i_0_25466, plain, (np__24!=np__136|~r1_xxreal_0(np__257,np__24))).#trainneg
cnf(i_0_25470, plain, (np__135!=np__136|~r1_xxreal_0(np__257,np__135))).#trainneg
cnf(i_0_25471, plain, (np__6168!=np__136|~r1_xxreal_0(np__257,np__6168))).#trainneg
cnf(i_0_25472, plain, (np__136!=np__6|~r1_xxreal_0(np__257,np__6))).#trainneg
cnf(i_0_7824, plain, (r1_xxreal_0(X1,np__6425)|r1_xxreal_0(np__6425,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21917, plain, (np__136=np__0|np__6305!=np__136|np__6305!=np__257)).#trainneg
cnf(i_0_1563, plain, (X1=k2_nat_d(X2,np__25)|r1_xxreal_0(np__25,X1)|X2!=k2_xcmplx_0(np__6425,X1)|~v7_ordinal1(np__25)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_25473, plain, (np__6536!=np__136|~r1_xxreal_0(np__257,np__6536))).#trainneg
cnf(i_0_25474, plain, (np__56!=np__136|~r1_xxreal_0(np__257,np__56))).#trainneg
cnf(i_0_25475, plain, (np__23!=np__136|~r1_xxreal_0(np__257,np__23))).#trainneg
cnf(i_0_25476, plain, (np__6289!=np__136|~r1_xxreal_0(np__257,np__6289))).#trainneg
cnf(i_0_25477, plain, (np__231!=np__136|~r1_xxreal_0(np__257,np__231))).#trainneg
cnf(i_0_31293, plain, (k2_nat_d(k3_xcmplx_0(k2_nat_d(X2,X3),X1),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X2,X3),k1_numbers))).#trainneg
cnf(i_0_34190, plain, (k2_nat_d(k4_nat_1(k2_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X3,X2),X1),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k4_nat_1(k2_nat_d(X1,X2),X3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_30816, plain, (k3_xcmplx_0(X1,np__257)=k3_xcmplx_0(np__257,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_15113, plain, (m1_subset_1(k2_xcmplx_0(X1,np__136),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_32068, plain, (k2_nat_d(k2_xcmplx_0(k3_xcmplx_0(np__257,np__257),X1),np__257)=X1|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(k3_xcmplx_0(np__257,np__257),X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7940, plain, (r1_xxreal_0(X1,np__25)|r1_xxreal_0(np__25,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21918, plain, (np__136=np__0|np__7!=np__136|np__7!=np__257)).#trainneg
cnf(i_0_1513, plain, (X1=k2_nat_d(X2,np__257)|r1_xxreal_0(np__257,X1)|X2!=k2_xcmplx_0(np__6425,X1)|~v7_ordinal1(np__25)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_25479, plain, (np__79!=np__136|~r1_xxreal_0(np__257,np__79))).#trainneg
cnf(i_0_25482, plain, (np__136!=np__6561|~r1_xxreal_0(np__257,np__6561))).#trainneg
cnf(i_0_25488, plain, (np__136!=np__8|~r1_xxreal_0(np__257,np__8))).#trainneg
cnf(i_0_25720, plain, (np__81!=np__6561|~r1_xxreal_0(np__81,np__81))).#trainneg
cnf(i_0_22352, plain, (k4_nat_d(k2_xcmplx_0(X1,k4_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__1),X1)|r1_xxreal_0(X1,k4_nat_d(X2,X1))|~v7_ordinal1(k4_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k2_xcmplx_0(X1,k4_nat_d(X2,X1)),k1_numbers))).#trainneg
cnf(i_0_21873, plain, (r1_xxreal_0(np__0,np__2))).#trainneg
cnf(i_0_14455, plain, (k2_nat_d(k4_nat_1(k2_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X3,X2)),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k4_nat_1(k2_nat_d(X1,X2),X3))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_12543, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(X2,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(X2,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_33956, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X3,X2),X1),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k2_nat_d(X1,X2),k1_numbers))).#trainneg
cnf(i_0_15100, plain, (k2_nat_d(X1,np__257)=np__136|X1!=np__6561|~v7_ordinal1(X1))).#trainneg
cnf(i_0_37153, negated_conjecture, (k13_newton(np__3,np__8)!=np__6561|~v7_ordinal1(k13_newton(np__3,np__8)))).#trainneg
cnf(i_0_37143, plain, (np__136=np__0|np__6425!=np__6561)).#trainneg
cnf(i_0_1388, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,np__257),np__81),np__257)=k4_nat_d(k3_xcmplx_0(X1,k13_newton(np__3,np__4)),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_26513, plain, (np__6480!=np__0|~r1_xxreal_0(np__257,np__6480))).#trainneg
cnf(i_0_26514, plain, (np__254!=np__0|~r1_xxreal_0(np__257,np__254))).#trainneg
cnf(i_0_35338, plain, (k2_nat_d(k3_xcmplx_0(k4_nat_d(X3,X2),k4_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|r1_xxreal_0(X2,k4_nat_d(X3,X2))|r1_xxreal_0(X2,k4_nat_d(X1,X2))|~v7_ordinal1(k4_nat_d(X1,X2))|~v7_ordinal1(k4_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_35151, plain, (k2_nat_d(k4_nat_1(X1,X2),X3)=k2_nat_d(k3_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,X2)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k4_nat_1(X1,X2))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_35352, plain, (k4_nat_d(k3_xcmplx_0(X1,X2),X3)=k2_nat_d(k3_xcmplx_0(X2,X1),X3)|r1_xxreal_0(X3,X2)|r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7944, plain, (r1_xxreal_0(X1,np__6)|r1_xxreal_0(np__6,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_37151, plain, (np__136=np__4|np__4!=np__6561)).#trainneg
cnf(i_0_1389, plain, (k4_nat_d(k4_nat_1(np__81,k4_nat_d(X1,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),X1),np__257)|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21920, plain, (np__136=np__0|np__5!=np__136|np__5!=np__257)).#trainneg
cnf(i_0_21921, plain, (np__136=np__0|np__252!=np__136|np__252!=np__257)).#trainneg
cnf(i_0_21923, plain, (np__136=np__0|np__6558!=np__136|np__6558!=np__257)).#trainneg
cnf(i_0_26515, plain, (np__253!=np__0|~r1_xxreal_0(np__257,np__253))).#trainneg
cnf(i_0_26516, plain, (np__255!=np__0|~r1_xxreal_0(np__257,np__255))).#trainneg
cnf(i_0_26517, plain, (np__120!=np__0|~r1_xxreal_0(np__257,np__120))).#trainneg
cnf(i_0_26518, plain, (np__55!=np__0|~r1_xxreal_0(np__257,np__55))).#trainneg
cnf(i_0_26519, plain, (np__6305!=np__0|~r1_xxreal_0(np__257,np__6305))).#trainneg
cnf(i_0_35328, plain, (k2_nat_d(k3_xcmplx_0(k2_nat_d(X3,X2),k2_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|r1_xxreal_0(X2,k2_nat_d(X3,X2))|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_nat_d(X3,X2))|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_35221, plain, (k2_nat_d(k3_xcmplx_0(k4_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|r1_xxreal_0(X2,k4_nat_d(X3,X2))|r1_xxreal_0(X2,k4_nat_d(X1,X2))|~v7_ordinal1(k4_nat_d(X3,X2))|~v7_ordinal1(k4_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_35213, plain, (k2_nat_d(k3_xcmplx_0(k2_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_35222, plain, (k2_nat_d(k3_xcmplx_0(k4_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X2)|r1_xxreal_0(X2,k4_nat_d(X1,X2))|r1_xxreal_0(X2,X3)|~v7_ordinal1(k4_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_35233, plain, (k2_nat_d(k3_xcmplx_0(X1,k4_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(X2,k2_xcmplx_0(X1,X3)),X3)|r1_xxreal_0(X3,k4_nat_d(X2,X3))|r1_xxreal_0(X3,X1)|~v7_ordinal1(k4_nat_d(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21924, plain, (np__136=np__0|np__6424!=np__136|np__6424!=np__257)).#trainneg
cnf(i_0_1607, plain, (X1=k2_nat_d(X2,np__257)|r1_xxreal_0(np__257,X1)|X2!=k2_xcmplx_0(np__0,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_21925, plain, (np__136=np__0|np__6304!=np__136|np__6304!=np__257)).#trainneg
cnf(i_0_21926, plain, (np__136=np__0|np__80!=np__136|np__80!=np__257)).#trainneg
cnf(i_0_21927, plain, (np__136=np__0|np__134!=np__136|np__134!=np__257)).#trainneg
cnf(i_0_26520, plain, (np__7!=np__0|~r1_xxreal_0(np__257,np__7))).#trainneg
cnf(i_0_26522, plain, (np__5!=np__0|~r1_xxreal_0(np__257,np__5))).#trainneg
cnf(i_0_26523, plain, (np__252!=np__0|~r1_xxreal_0(np__257,np__252))).#trainneg
cnf(i_0_26525, plain, (np__6558!=np__0|~r1_xxreal_0(np__257,np__6558))).#trainneg
cnf(i_0_26526, plain, (np__6424!=np__0|~r1_xxreal_0(np__257,np__6424))).#trainneg
cnf(i_0_35226, plain, (k2_nat_d(k3_xcmplx_0(X1,k2_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(X2,k2_xcmplx_0(X1,X3)),X3)|r1_xxreal_0(X3,k2_nat_d(X2,X3))|r1_xxreal_0(X3,X1)|~v7_ordinal1(k2_nat_d(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_35210, plain, (k2_nat_d(k3_xcmplx_0(k2_nat_d(X1,X2),k2_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|r1_xxreal_0(X2,k2_nat_d(X3,X2))|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_nat_d(X3,X2))|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_35214, plain, (k2_nat_d(k3_xcmplx_0(k2_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X3),X1),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_36335, plain, (r1_xxreal_0(np__25,np__257))).#trainneg
cnf(i_0_36099, plain, (k2_nat_d(k2_xcmplx_0(k3_xcmplx_0(X1,np__257),X2),X1)=X2|r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(k3_xcmplx_0(X1,np__257),X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_35223, plain, (k2_nat_d(k3_xcmplx_0(k4_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X3),X1),X2)|r1_xxreal_0(X2,k4_nat_d(X1,X2))|r1_xxreal_0(X2,X3)|~v7_ordinal1(k4_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_35229, plain, (k2_nat_d(k3_xcmplx_0(X1,k2_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X3,X1),X2),X3)|r1_xxreal_0(X3,k2_nat_d(X2,X3))|r1_xxreal_0(X3,X1)|~v7_ordinal1(k2_nat_d(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8017, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_21928, plain, (np__136=np__0|np__232!=np__136|np__232!=np__257)).#trainneg
cnf(i_0_37431, plain, (np__6425!=np__6561)).#trainneg
cnf(i_0_38325, plain, (k2_nat_d(k3_xcmplx_0(X2,X1),X3)=k2_nat_d(k3_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,X2)|r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_39611, plain, (k2_nat_d(k3_xcmplx_0(X3,k2_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k2_nat_d(X1,X2),k1_numbers))).#trainneg
cnf(i_0_37429, negated_conjecture, (k1_newton(np__3,np__8)!=np__6561|~v7_ordinal1(k1_newton(np__3,np__8)))).#trainneg
cnf(i_0_37430, negated_conjecture, (k13_newton(np__3,np__8)!=np__6561|~m1_subset_1(k13_newton(np__3,np__8),k5_numbers))).#trainneg
cnf(i_0_37391, plain, (k4_nat_d(np__136,X1)=k2_nat_d(np__136,X1)|r1_xxreal_0(X1,np__136)|r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21929, plain, (np__136=np__0|np__6400!=np__136|np__6400!=np__257)).#trainneg
cnf(i_0_2560, plain, (X1=k2_nat_d(X2,k1_newton(X3,np__0))|r1_xxreal_0(k1_newton(X3,np__0),X1)|X2!=k2_xcmplx_0(k1_newton(X3,np__81),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,np__0))|~v7_ordinal1(k1_newton(X3,np__81))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_38461, plain, (np__4!=np__6561)).#trainneg
cnf(i_0_26527, plain, (np__6304!=np__0|~r1_xxreal_0(np__257,np__6304))).#trainneg
cnf(i_0_40123, plain, (k2_nat_d(X1,np__257)=np__0|X1!=np__0|~v7_ordinal1(X1))).#trainneg
cnf(i_0_41273, plain, (k2_nat_d(k3_xcmplx_0(X3,k4_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X3),X1),X2)|r1_xxreal_0(X2,k4_nat_d(X1,X2))|r1_xxreal_0(X2,X3)|~v7_ordinal1(k4_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k4_nat_d(X1,X2),k1_numbers))).#trainneg
cnf(i_0_38871, plain, (k2_nat_d(k3_xcmplx_0(k2_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X3,X2)),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_xcmplx_0(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_37392, plain, (k4_nat_d(np__8,X1)=k2_nat_d(np__8,X1)|r1_xxreal_0(X1,np__8)|r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21930, plain, (np__136=np__0|np__6557!=np__136|np__6557!=np__257)).#trainneg
cnf(i_0_2291, plain, (k2_nat_d(k2_xcmplx_0(np__6561,X1),np__81)=X1|r1_xxreal_0(np__81,X1)|~v7_ordinal1(k2_xcmplx_0(np__6561,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_26528, plain, (np__80!=np__0|~r1_xxreal_0(np__257,np__80))).#trainneg
cnf(i_0_26529, plain, (np__134!=np__0|~r1_xxreal_0(np__257,np__134))).#trainneg
cnf(i_0_26530, plain, (np__232!=np__0|~r1_xxreal_0(np__257,np__232))).#trainneg
cnf(i_0_26531, plain, (np__6400!=np__0|~r1_xxreal_0(np__257,np__6400))).#trainneg
cnf(i_0_26532, plain, (np__6557!=np__0|~r1_xxreal_0(np__257,np__6557))).#trainneg
cnf(i_0_39864, plain, (k2_nat_d(k3_xcmplx_0(X3,k4_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X2)|r1_xxreal_0(X2,k4_nat_d(X1,X2))|r1_xxreal_0(X2,X3)|~v7_ordinal1(k4_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k4_nat_d(X1,X2),k1_numbers))).#trainneg
cnf(i_0_37393, plain, (k4_nat_d(np__3,X1)=k2_nat_d(np__3,X1)|r1_xxreal_0(X1,np__3)|r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21931, plain, (np__136=np__0|np__6344!=np__136|np__6344!=np__257)).#trainneg
cnf(i_0_1590, plain, (X1=k2_nat_d(X2,np__81)|r1_xxreal_0(np__81,X1)|X2!=k2_xcmplx_0(np__6561,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_21933, plain, (np__136=np__0|np__22!=np__136|np__22!=np__257)).#trainneg
cnf(i_0_21934, plain, (np__136=np__0|np__111!=np__136|np__111!=np__257)).#trainneg
cnf(i_0_21935, plain, (np__136=np__0|np__121!=np__136|np__121!=np__257)).#trainneg
cnf(i_0_26533, plain, (np__6344!=np__0|~r1_xxreal_0(np__257,np__6344))).#trainneg
cnf(i_0_26535, plain, (np__22!=np__0|~r1_xxreal_0(np__257,np__22))).#trainneg
cnf(i_0_26536, plain, (np__111!=np__0|~r1_xxreal_0(np__257,np__111))).#trainneg
cnf(i_0_26537, plain, (np__121!=np__0|~r1_xxreal_0(np__257,np__121))).#trainneg
cnf(i_0_26538, plain, (np__132!=np__0|~r1_xxreal_0(np__257,np__132))).#trainneg
cnf(i_0_40114, plain, (k2_nat_d(k3_xcmplx_0(k4_nat_d(X2,X3),X1),X3)=k4_nat_d(k3_xcmplx_0(X2,k2_xcmplx_0(X1,X3)),X3)|r1_xxreal_0(X3,k4_nat_d(X2,X3))|r1_xxreal_0(X3,X1)|~v7_ordinal1(k4_nat_d(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21936, plain, (np__136=np__0|np__132!=np__136|np__132!=np__257)).#trainneg
cnf(i_0_2589, plain, (X1=k2_nat_d(X2,k1_newton(X3,np__4))|r1_xxreal_0(k1_newton(X3,np__4),X1)|X2!=k2_xcmplx_0(k1_newton(X3,np__8),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,np__4))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_21937, plain, (np__136=np__0|np__133!=np__136|np__133!=np__257)).#trainneg
cnf(i_0_21938, plain, (np__136=np__0|np__175!=np__136|np__175!=np__257)).#trainneg
cnf(i_0_21939, plain, (np__136=np__0|np__176!=np__136|np__176!=np__257)).#trainneg
cnf(i_0_26539, plain, (np__133!=np__0|~r1_xxreal_0(np__257,np__133))).#trainneg
cnf(i_0_26540, plain, (np__175!=np__0|~r1_xxreal_0(np__257,np__175))).#trainneg
cnf(i_0_26541, plain, (np__176!=np__0|~r1_xxreal_0(np__257,np__176))).#trainneg
cnf(i_0_26542, plain, (np__21!=np__0|~r1_xxreal_0(np__257,np__21))).#trainneg
cnf(i_0_26543, plain, (np__6169!=np__0|~r1_xxreal_0(np__257,np__6169))).#trainneg
cnf(i_0_36384, plain, (k2_nat_d(k2_xcmplx_0(X1,k4_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__1),X1)|r1_xxreal_0(X1,k4_nat_d(X2,X1))|~v7_ordinal1(k2_xcmplx_0(X1,k4_nat_d(X2,X1)))|~v7_ordinal1(k4_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_38355, plain, (k2_nat_d(k3_xcmplx_0(X1,np__1),X2)=k4_nat_d(X1,X2)|r1_xxreal_0(X2,np__1)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11922, plain, (k2_nat_d(k4_nat_1(k4_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X3,X2)),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k4_nat_1(k4_nat_d(X1,X2),X3))|~v7_ordinal1(k2_xcmplx_0(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_16464, plain, (X1=np__0|r1_xxreal_0(np__0,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_8021, plain, (r1_xxreal_0(X1,np__6)|r1_xxreal_0(np__6,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_21940, plain, (np__136=np__0|np__21!=np__136|np__21!=np__257)).#trainneg
cnf(i_0_21941, plain, (np__136=np__0|np__6169!=np__136|np__6169!=np__257)).#trainneg
cnf(i_0_21942, plain, (np__136=np__0|np__6421!=np__136|np__6421!=np__257)).#trainneg
cnf(i_0_21943, plain, (np__136=np__0|np__6422!=np__136|np__6422!=np__257)).#trainneg
cnf(i_0_26544, plain, (np__6421!=np__0|~r1_xxreal_0(np__257,np__6421))).#trainneg
cnf(i_0_26545, plain, (np__6422!=np__0|~r1_xxreal_0(np__257,np__6422))).#trainneg
cnf(i_0_26546, plain, (np__6423!=np__0|~r1_xxreal_0(np__257,np__6423))).#trainneg
cnf(i_0_26547, plain, (np__6559!=np__0|~r1_xxreal_0(np__257,np__6559))).#trainneg
cnf(i_0_26548, plain, (np__6560!=np__0|~r1_xxreal_0(np__257,np__6560))).#trainneg
cnf(i_0_18016, plain, (k4_nat_d(k3_xcmplx_0(k4_nat_d(X2,X3),X1),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k4_nat_d(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_23084, plain, (v7_ordinal1(k2_xcmplx_0(np__257,np__136)))).#trainneg
cnf(i_0_31943, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,np__136),k2_xcmplx_0(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__136),X1)|r1_xxreal_0(X1,np__136)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_21944, plain, (np__136=np__0|np__6423!=np__136|np__6423!=np__257)).#trainneg
cnf(i_0_2292, plain, (k2_nat_d(k2_xcmplx_0(np__8,X1),np__4)=X1|r1_xxreal_0(np__4,X1)|~v7_ordinal1(k2_xcmplx_0(np__8,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21945, plain, (np__136=np__0|np__6559!=np__136|np__6559!=np__257)).#trainneg
cnf(i_0_21946, plain, (np__136=np__0|np__6560!=np__136|np__6560!=np__257)).#trainneg
cnf(i_0_21947, plain, (np__136=np__0|np__77!=np__136|np__77!=np__257)).#trainneg
cnf(i_0_26549, plain, (np__77!=np__0|~r1_xxreal_0(np__257,np__77))).#trainneg
cnf(i_0_26550, plain, (np__78!=np__0|~r1_xxreal_0(np__257,np__78))).#trainneg
cnf(i_0_26551, plain, (np__24!=np__0|~r1_xxreal_0(np__257,np__24))).#trainneg
cnf(i_0_26555, plain, (np__135!=np__0|~r1_xxreal_0(np__257,np__135))).#trainneg
cnf(i_0_26556, plain, (np__6168!=np__0|~r1_xxreal_0(np__257,np__6168))).#trainneg
cnf(i_0_12555, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,np__257),X2),np__257)=k4_nat_d(k3_xcmplx_0(X2,X1),np__257)|r1_xxreal_0(np__257,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_42058, negated_conjecture, (k13_newton(np__3,np__8)!=np__6561)).#trainneg
cnf(i_0_42055, negated_conjecture, (k1_newton(np__3,np__8)!=np__6561)).#trainneg
cnf(i_0_44606, plain, (m1_subset_1(k2_xcmplx_0(np__257,np__136),k5_numbers))).#trainneg
cnf(i_0_44939, plain, (k2_nat_d(np__8,np__4)=np__0)).#trainneg
cnf(i_0_44736, plain, (k4_nat_d(k4_nat_1(X2,np__136),X1)=k4_nat_d(k3_xcmplx_0(X2,np__136),X1)|r1_xxreal_0(X1,np__136)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_41784, plain, (k2_nat_d(k3_xcmplx_0(X1,np__1),X2)=k2_nat_d(X1,X2)|r1_xxreal_0(X2,np__1)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_42227, negated_conjecture, (k13_newton(np__3,np__8)!=np__0|np__136!=np__0|~v7_ordinal1(k13_newton(np__3,np__8)))).#trainneg
cnf(i_0_21948, plain, (np__136=np__0|np__78!=np__136|np__78!=np__257)).#trainneg
cnf(i_0_2282, plain, (k2_nat_d(k2_xcmplx_0(np__8,X1),np__2)=X1|r1_xxreal_0(np__2,X1)|~v7_ordinal1(k2_xcmplx_0(np__8,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21949, plain, (np__136=np__0|np__24!=np__136|np__24!=np__257)).#trainneg
cnf(i_0_21951, plain, (np__136=np__0|np__25!=np__136|np__25!=np__257)).#trainneg
cnf(i_0_21953, plain, (np__136=np__0|np__135!=np__136|np__135!=np__257)).#trainneg
cnf(i_0_26557, plain, (np__6!=np__0|~r1_xxreal_0(np__257,np__6))).#trainneg
cnf(i_0_26558, plain, (np__6536!=np__0|~r1_xxreal_0(np__257,np__6536))).#trainneg
cnf(i_0_26559, plain, (np__56!=np__0|~r1_xxreal_0(np__257,np__56))).#trainneg
cnf(i_0_26560, plain, (np__23!=np__0|~r1_xxreal_0(np__257,np__23))).#trainneg
cnf(i_0_26561, plain, (np__6289!=np__0|~r1_xxreal_0(np__257,np__6289))).#trainneg
cnf(i_0_42984, plain, (k2_nat_d(np__6561,np__81)=np__0)).#trainneg
cnf(i_0_43976, plain, (X1=np__0|r1_xxreal_0(X1,np__1)|r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,k4_nat_d(X2,X1))|~v7_ordinal1(k3_xcmplx_0(X2,np__1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_43758, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__257,np__136),np__1),np__257)=np__136|~v7_ordinal1(k2_xcmplx_0(np__257,np__136)))).#trainneg
cnf(i_0_44846, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__136,X1),k2_xcmplx_0(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__136),X1)|r1_xxreal_0(X1,np__136)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8244, plain, (r1_xxreal_0(X1,np__4)|r1_xxreal_0(np__4,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_21954, plain, (np__136=np__0|np__6168!=np__136|np__6168!=np__257)).#trainneg
cnf(i_0_1719, plain, (X1=k2_nat_d(X2,np__6)|r1_xxreal_0(np__6,X1)|X2!=k2_xcmplx_0(np__24,X1)|~v7_ordinal1(np__6)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_21956, plain, (np__136=np__0|np__6536!=np__136|np__6536!=np__257)).#trainneg
cnf(i_0_21957, plain, (np__136=np__0|np__56!=np__136|np__56!=np__257)).#trainneg
cnf(i_0_21958, plain, (np__136=np__0|np__23!=np__136|np__23!=np__257)).#trainneg
cnf(i_0_26562, plain, (np__231!=np__0|~r1_xxreal_0(np__257,np__231))).#trainneg
cnf(i_0_26564, plain, (np__79!=np__0|~r1_xxreal_0(np__257,np__79))).#trainneg
cnf(i_0_26573, plain, (np__8!=np__0|~r1_xxreal_0(np__257,np__8))).#trainneg
cnf(i_0_26649, plain, (np__6425!=np__1|~r1_xxreal_0(np__2,np__6425))).#trainneg
cnf(i_0_43884, plain, (k4_nat_d(k3_xcmplx_0(np__136,np__1),np__257)=np__136)).#trainneg
cnf(i_0_43894, plain, (k4_nat_d(k3_xcmplx_0(X1,np__1),np__257)=np__136|X1!=np__136|~v7_ordinal1(X1))).#trainneg
cnf(i_0_12116, plain, (X1=np__0|X1!=np__136|~r1_xxreal_0(np__6425,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_8257, plain, (r1_xxreal_0(X1,np__21)|r1_xxreal_0(np__21,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21959, plain, (np__136=np__0|np__6289!=np__136|np__6289!=np__257)).#trainneg
cnf(i_0_3407, plain, (k3_xcmplx_0(X1,np__0)=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_21960, plain, (np__136=np__0|np__231!=np__136|np__231!=np__257)).#trainneg
cnf(i_0_21962, plain, (np__136=np__0|np__79!=np__136|np__79!=np__257)).#trainneg
cnf(i_0_26653, plain, (np__257!=np__1|~r1_xxreal_0(np__2,np__257))).#trainneg
cnf(i_0_27431, plain, (np__6480!=np__4|~r1_xxreal_0(np__257,np__6480))).#trainneg
cnf(i_0_27432, plain, (np__254!=np__4|~r1_xxreal_0(np__257,np__254))).#trainneg
cnf(i_0_27433, plain, (np__253!=np__4|~r1_xxreal_0(np__257,np__253))).#trainneg
cnf(i_0_27434, plain, (np__255!=np__4|~r1_xxreal_0(np__257,np__255))).#trainneg
cnf(i_0_14326, plain, (k2_nat_d(k4_nat_1(X1,k4_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(X2,k2_xcmplx_0(X1,X3)),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k4_nat_1(X1,k4_nat_d(X2,X3)))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_24555, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),X2),X2)=k4_nat_d(k3_xcmplx_0(X1,np__0),X2)|r1_xxreal_0(X2,np__0)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_31306, plain, (k4_nat_1(X1,k4_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,np__257),X2),np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,k4_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k4_nat_1(X1,k4_nat_d(X2,np__257)))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14358, plain, (k4_nat_d(k3_xcmplx_0(X3,k2_xcmplx_0(k2_nat_d(X1,X2),X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_35174, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(k3_xcmplx_0(X3,X4),X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,k4_nat_1(X3,X4)),X2)|r1_xxreal_0(X2,X4)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k4_nat_1(X3,X4))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_8478, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__5,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3120, plain, (k2_xcmplx_0(X1,np__0)=k2_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_25169, negated_conjecture, (k1_newton(np__3,np__8)!=np__81|np__136!=np__81)).#trainneg
cnf(i_0_25831, negated_conjecture, (k1_newton(np__3,np__8)!=np__257|np__136!=np__0)).#trainneg
cnf(i_0_27435, plain, (np__120!=np__4|~r1_xxreal_0(np__257,np__120))).#trainneg
cnf(i_0_27436, plain, (np__55!=np__4|~r1_xxreal_0(np__257,np__55))).#trainneg
cnf(i_0_27437, plain, (np__6305!=np__4|~r1_xxreal_0(np__257,np__6305))).#trainneg
cnf(i_0_27438, plain, (np__7!=np__4|~r1_xxreal_0(np__257,np__7))).#trainneg
cnf(i_0_27440, plain, (np__5!=np__4|~r1_xxreal_0(np__257,np__5))).#trainneg
cnf(i_0_46318, plain, (k2_nat_d(k3_xcmplx_0(np__136,np__1),np__257)=np__136)).#trainneg
cnf(i_0_46528, plain, (X1=np__0|X1!=np__136|~r1_xxreal_0(np__6425,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_45615, plain, (k2_nat_d(np__8,np__2)=np__0)).#trainneg
cnf(i_0_47498, plain, (k4_nat_d(k2_xcmplx_0(k2_nat_d(X1,X2),X2),X2)=k4_nat_d(k3_xcmplx_0(X1,np__1),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_xcmplx_0(k2_nat_d(X1,X2),X2),k1_numbers))).#trainneg
cnf(i_0_46438, plain, (k2_nat_d(k3_xcmplx_0(X1,np__1),np__257)=np__136|r1_xxreal_0(np__257,X1)|X1!=np__136|~v7_ordinal1(X1))).#trainneg
cnf(i_0_47224, plain, (k4_nat_1(X1,k2_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,np__257),X2),np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,k2_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k4_nat_1(X1,k2_nat_d(X2,np__257)))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_46437, plain, (k2_nat_d(k3_xcmplx_0(np__1,X1),np__257)=np__136|r1_xxreal_0(np__257,X1)|X1!=np__136|~v7_ordinal1(X1))).#trainneg
cnf(i_0_46645, plain, (k4_nat_1(X1,k4_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(X2,k2_xcmplx_0(X1,np__257)),np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,k4_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k4_nat_1(X1,k4_nat_d(X2,np__257)))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_46627, plain, (k3_xcmplx_0(np__0,np__3)=np__0)).#trainneg
cnf(i_0_8482, plain, (r1_xxreal_0(X1,np__22)|r1_xxreal_0(np__22,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3796, plain, (k3_xcmplx_0(np__4,np__3)=k3_xcmplx_0(np__2,np__6)|~v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_27441, plain, (np__252!=np__4|~r1_xxreal_0(np__257,np__252))).#trainneg
cnf(i_0_27443, plain, (np__6558!=np__4|~r1_xxreal_0(np__257,np__6558))).#trainneg
cnf(i_0_27444, plain, (np__6424!=np__4|~r1_xxreal_0(np__257,np__6424))).#trainneg
cnf(i_0_27445, plain, (np__6304!=np__4|~r1_xxreal_0(np__257,np__6304))).#trainneg
cnf(i_0_27446, plain, (np__80!=np__4|~r1_xxreal_0(np__257,np__80))).#trainneg
cnf(i_0_47497, plain, (k4_nat_d(k2_xcmplx_0(k2_nat_d(X1,X2),X2),X2)=k4_nat_d(k3_xcmplx_0(X1,np__1),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_xcmplx_0(k2_nat_d(X1,X2),X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_46620, plain, (k3_xcmplx_0(np__0,np__1)=np__0)).#trainneg
cnf(i_0_46893, plain, (k2_nat_d(k4_nat_1(X1,k2_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(X2,k2_xcmplx_0(X1,X3)),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k4_nat_1(X1,k2_nat_d(X2,X3)))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_47080, plain, (k4_nat_d(k3_xcmplx_0(X2,k2_xcmplx_0(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__0),X2)|r1_xxreal_0(X2,np__0)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers))).#trainneg
cnf(i_0_45421, plain, (k2_nat_d(X1,np__257)=k3_xcmplx_0(X1,np__1)|r1_xxreal_0(np__257,k3_xcmplx_0(X1,np__1))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k3_xcmplx_0(X1,np__1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_47230, plain, (k4_nat_1(X1,k4_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,np__257),X2),np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,k4_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_1(X1,k4_nat_d(X2,np__257)),k5_numbers))).#trainneg
cnf(i_0_20750, plain, (np__0=X1|r1_xxreal_0(np__257,X1)|X1!=np__257|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_8486, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__7,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1999, plain, (k3_xcmplx_0(np__2,k3_xcmplx_0(np__3,X1))=k3_xcmplx_0(np__6,X1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_27447, plain, (np__134!=np__4|~r1_xxreal_0(np__257,np__134))).#trainneg
cnf(i_0_27448, plain, (np__232!=np__4|~r1_xxreal_0(np__257,np__232))).#trainneg
cnf(i_0_27449, plain, (np__6400!=np__4|~r1_xxreal_0(np__257,np__6400))).#trainneg
cnf(i_0_27450, plain, (np__6557!=np__4|~r1_xxreal_0(np__257,np__6557))).#trainneg
cnf(i_0_27451, plain, (np__6344!=np__4|~r1_xxreal_0(np__257,np__6344))).#trainneg
cnf(i_0_14198, plain, (k2_nat_d(k2_xcmplx_0(np__257,X1),np__1)=X1|r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(np__257,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11940, plain, (k2_nat_d(np__6561,np__6561)=np__0)).#trainneg
cnf(i_0_50372, plain, (k2_nat_d(np__257,np__1)=np__0)).#trainneg
cnf(i_0_50175, plain, (np__8=np__0|np__136!=np__8|np__8!=np__257)).#trainneg
cnf(i_0_50155, plain, (np__6=np__0|np__136!=np__6|np__6!=np__257)).#trainneg
cnf(i_0_1943, plain, (k3_xcmplx_0(np__2,k7_xcmplx_0(np__3,X1))=k7_xcmplx_0(np__6,X1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_25832, plain, (v1_xreal_0(np__175))).#trainneg
cnf(i_0_26308, plain, (v1_xcmplx_0(np__175))).#trainneg
cnf(i_0_27453, plain, (np__22!=np__4|~r1_xxreal_0(np__257,np__22))).#trainneg
cnf(i_0_27454, plain, (np__111!=np__4|~r1_xxreal_0(np__257,np__111))).#trainneg
cnf(i_0_27455, plain, (np__121!=np__4|~r1_xxreal_0(np__257,np__121))).#trainneg
cnf(i_0_27456, plain, (np__132!=np__4|~r1_xxreal_0(np__257,np__132))).#trainneg
cnf(i_0_27457, plain, (np__133!=np__4|~r1_xxreal_0(np__257,np__133))).#trainneg
cnf(i_0_49057, plain, (k2_nat_d(k2_xcmplx_0(k2_nat_d(X1,X2),X2),X2)=k4_nat_d(k3_xcmplx_0(X1,np__1),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_xcmplx_0(k2_nat_d(X1,X2),X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_47971, plain, (k2_xcmplx_0(X1,np__0)=k2_xcmplx_0(np__0,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_47970, plain, (k2_xcmplx_0(X1,np__0)=k2_xcmplx_0(np__0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50427, plain, (X1!=np__0|~r1_xxreal_0(np__6561,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50178, plain, (np__0=X1|X1!=np__136|X1!=np__257|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_8832, plain, (r1_xxreal_0(X1,np__255)|r1_xxreal_0(np__255,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50063, plain, (np__6480=np__0|np__6480!=np__136|np__6480!=np__257)).#trainneg
cnf(i_0_50066, plain, (np__254=np__0|np__254!=np__136|np__254!=np__257)).#trainneg
cnf(i_0_50069, plain, (np__253=np__0|np__253!=np__136|np__253!=np__257)).#trainneg
cnf(i_0_27458, plain, (np__175!=np__4|~r1_xxreal_0(np__257,np__175))).#trainneg
cnf(i_0_27459, plain, (np__176!=np__4|~r1_xxreal_0(np__257,np__176))).#trainneg
cnf(i_0_27460, plain, (np__21!=np__4|~r1_xxreal_0(np__257,np__21))).#trainneg
cnf(i_0_27461, plain, (np__6169!=np__4|~r1_xxreal_0(np__257,np__6169))).#trainneg
cnf(i_0_27462, plain, (np__6421!=np__4|~r1_xxreal_0(np__257,np__6421))).#trainneg
cnf(i_0_47966, plain, (k2_xcmplx_0(np__0,np__4)=np__4)).#trainneg
cnf(i_0_49440, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X3),k2_xcmplx_0(X1,X3)),X3)=k2_nat_d(k4_nat_1(X1,X2),X3)|r1_xxreal_0(X3,X1)|r1_xxreal_0(X3,X2)|~v7_ordinal1(k4_nat_1(X1,X2))|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_49532, plain, (k2_nat_d(k3_xcmplx_0(X1,k2_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(X2,k2_xcmplx_0(X1,X3)),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k3_xcmplx_0(X1,k2_nat_d(X2,X3)))|~v7_ordinal1(k2_nat_d(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_15977, plain, (k4_nat_d(k4_nat_1(np__0,X1),np__257)=k4_nat_d(k3_xcmplx_0(np__257,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_36555, plain, (k4_nat_d(k2_xcmplx_0(X1,k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__1),X1)|r1_xxreal_0(X1,k2_nat_d(X2,X1))|~v7_ordinal1(k2_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_35269, plain, (k2_nat_d(k4_nat_1(X1,X2),X3)=k2_nat_d(k3_xcmplx_0(X2,X1),X3)|r1_xxreal_0(X3,X2)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k4_nat_1(X1,X2))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50760, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__1),X2)=k2_nat_d(k2_xcmplx_0(X1,X2),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_51575, plain, (k2_nat_d(X1,np__257)=np__4|X1!=np__4|~v7_ordinal1(X1))).#trainneg
cnf(i_0_51414, plain, (k4_pepin(X1)=np__0|k4_pepin(X1)!=np__136|k4_pepin(X1)!=np__257|~v7_ordinal1(X1))).#trainneg
cnf(i_0_51413, plain, (k4_nat_d(X1,X2)=np__0|k4_nat_d(X1,X2)!=np__136|k4_nat_d(X1,X2)!=np__257|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3682, plain, (v1_xcmplx_0(k4_xcmplx_0(np__3))|~v1_xreal_0(np__4))).#trainneg
cnf(i_0_1675, plain, (v1_xcmplx_0(np__256))).#trainneg
cnf(i_0_27463, plain, (np__6422!=np__4|~r1_xxreal_0(np__257,np__6422))).#trainneg
cnf(i_0_27464, plain, (np__6423!=np__4|~r1_xxreal_0(np__257,np__6423))).#trainneg
cnf(i_0_27465, plain, (np__6559!=np__4|~r1_xxreal_0(np__257,np__6559))).#trainneg
cnf(i_0_27466, plain, (np__6560!=np__4|~r1_xxreal_0(np__257,np__6560))).#trainneg
cnf(i_0_27467, plain, (np__77!=np__4|~r1_xxreal_0(np__257,np__77))).#trainneg
cnf(i_0_51340, plain, (X1!=np__0|~r1_xxreal_0(np__6561,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_51691, plain, (k2_nat_d(k4_nat_1(X3,X1),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k4_nat_1(X3,X1))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_51521, plain, (k2_nat_d(np__4,np__4)=np__0)).#trainneg
cnf(i_0_50866, plain, (k2_nat_d(k2_xcmplx_0(X2,k2_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__1),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_xcmplx_0(X2,k2_nat_d(X1,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X1,X2),k1_numbers))).#trainneg
cnf(i_0_51690, plain, (k4_nat_d(k4_nat_1(X1,X3),X2)=k2_nat_d(k4_nat_1(X3,X1),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k4_nat_1(X3,X1))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_52398, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__1),X1)=k4_nat_d(k2_xcmplx_0(X1,X2),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_50941, plain, (X1=k2_nat_d(X2,np__6561)|r1_xxreal_0(np__6561,X1)|X2!=k2_xcmplx_0(X1,np__0)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8836, plain, (r1_xxreal_0(X1,np__6560)|r1_xxreal_0(np__6560,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_51415, plain, (k13_newton(X1,X2)=np__0|k13_newton(X1,X2)!=np__136|k13_newton(X1,X2)!=np__257|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_50072, plain, (np__255=np__0|np__255!=np__136|np__255!=np__257)).#trainneg
cnf(i_0_50075, plain, (np__120=np__0|np__120!=np__136|np__120!=np__257)).#trainneg
cnf(i_0_50078, plain, (np__55=np__0|np__55!=np__136|np__55!=np__257)).#trainneg
cnf(i_0_27468, plain, (np__78!=np__4|~r1_xxreal_0(np__257,np__78))).#trainneg
cnf(i_0_27469, plain, (np__24!=np__4|~r1_xxreal_0(np__257,np__24))).#trainneg
cnf(i_0_27473, plain, (np__135!=np__4|~r1_xxreal_0(np__257,np__135))).#trainneg
cnf(i_0_27474, plain, (np__6168!=np__4|~r1_xxreal_0(np__257,np__6168))).#trainneg
cnf(i_0_27475, plain, (np__6!=np__4|~r1_xxreal_0(np__257,np__6))).#trainneg
cnf(i_0_50844, plain, (k4_nat_d(k3_xcmplx_0(X1,np__1),np__257)=k2_nat_d(k2_xcmplx_0(X1,np__257),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50759, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__1),X1)=k2_nat_d(k2_xcmplx_0(X2,X1),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(X2,X1))|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_52828, plain, (k4_nat_d(k2_xcmplx_0(X1,X2),X2)=k2_nat_d(k2_xcmplx_0(X1,X2),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers))).#trainneg
cnf(i_0_54039, plain, (k4_nat_d(k2_xcmplx_0(X1,X2),X2)=k4_nat_d(k2_xcmplx_0(X2,X1),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers))).#trainneg
cnf(i_0_54496, plain, (k4_nat_d(k2_xcmplx_0(X1,X2),X1)=k2_nat_d(k2_xcmplx_0(X2,X1),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers))).#trainneg
cnf(i_0_52823, plain, (k4_nat_d(k3_xcmplx_0(np__1,k2_xcmplx_0(X1,X2)),X2)=k2_nat_d(k2_xcmplx_0(X1,X2),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers))).#trainneg
cnf(i_0_52911, negated_conjecture, (k13_newton(np__3,np__8)!=np__4|np__136!=np__4|~v7_ordinal1(k13_newton(np__3,np__8)))).#trainneg
cnf(i_0_50081, plain, (np__6305=np__0|np__6305!=np__136|np__6305!=np__257)).#trainneg
cnf(i_0_1860, plain, (v1_xreal_0(k4_xcmplx_0(np__256))|~v1_xreal_0(np__257))).#trainneg
cnf(i_0_50084, plain, (np__7=np__0|np__7!=np__136|np__7!=np__257)).#trainneg
cnf(i_0_50087, plain, (np__5=np__0|np__5!=np__136|np__5!=np__257)).#trainneg
cnf(i_0_27476, plain, (np__6536!=np__4|~r1_xxreal_0(np__257,np__6536))).#trainneg
cnf(i_0_27477, plain, (np__56!=np__4|~r1_xxreal_0(np__257,np__56))).#trainneg
cnf(i_0_27478, plain, (np__23!=np__4|~r1_xxreal_0(np__257,np__23))).#trainneg
cnf(i_0_27479, plain, (np__6289!=np__4|~r1_xxreal_0(np__257,np__6289))).#trainneg
cnf(i_0_27480, plain, (np__231!=np__4|~r1_xxreal_0(np__257,np__231))).#trainneg
cnf(i_0_52045, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X3),k2_xcmplx_0(X1,X3)),X3)=k2_nat_d(k3_xcmplx_0(X1,X2),X3)|r1_xxreal_0(X3,X1)|r1_xxreal_0(X3,X2)|~v7_ordinal1(k3_xcmplx_0(X1,X2))|~v7_ordinal1(k2_xcmplx_0(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_51872, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),k2_xcmplx_0(X3,X2)),X2)=k2_nat_d(k4_nat_1(X3,X1),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k4_nat_1(X3,X1))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_54025, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__1),X2)=k4_nat_d(k2_xcmplx_0(X2,X1),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_54270, plain, (k2_nat_d(k2_xcmplx_0(X1,np__257),np__257)=np__136|r1_xxreal_0(np__257,X1)|X1!=np__136|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50060, plain, (k2_nat_d(X1,np__257)=np__0|k2_nat_d(X1,np__257)!=np__257|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X1,np__257),k5_numbers))).#trainneg
cnf(i_0_8840, plain, (r1_xxreal_0(X1,np__6305)|r1_xxreal_0(np__6305,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50090, plain, (np__252=np__0|np__252!=np__136|np__252!=np__257)).#trainneg
cnf(i_0_3049, plain, (k3_xcmplx_0(np__2,np__1)=np__2)).#trainneg
cnf(i_0_50093, plain, (np__6558=np__0|np__6558!=np__136|np__6558!=np__257)).#trainneg
cnf(i_0_50096, plain, (np__6424=np__0|np__6424!=np__136|np__6424!=np__257)).#trainneg
cnf(i_0_50099, plain, (np__6304=np__0|np__6304!=np__136|np__6304!=np__257)).#trainneg
cnf(i_0_27482, plain, (np__79!=np__4|~r1_xxreal_0(np__257,np__79))).#trainneg
cnf(i_0_27491, plain, (np__8!=np__4|~r1_xxreal_0(np__257,np__8))).#trainneg
cnf(i_0_28300, plain, (np__6480!=np__3|~r1_xxreal_0(np__257,np__6480))).#trainneg
cnf(i_0_28301, plain, (np__254!=np__3|~r1_xxreal_0(np__257,np__254))).#trainneg
cnf(i_0_17216, plain, (k2_nat_d(k2_xcmplx_0(X1,np__136),np__1)=X1|r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__136))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14671, plain, (k4_nat_d(k3_xcmplx_0(k4_nat_d(X2,X3),X1),X3)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X3),X2),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X2,X3),k5_numbers))).#trainneg
cnf(i_0_8992, plain, (r1_xxreal_0(X1,np__6536)|r1_xxreal_0(np__6536,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50102, plain, (np__80=np__0|np__80!=np__136|np__80!=np__257)).#trainneg
cnf(i_0_2007, plain, (k3_xcmplx_0(np__4,k3_xcmplx_0(np__2,X1))=k3_xcmplx_0(np__8,X1)|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_50105, plain, (np__134=np__0|np__134!=np__136|np__134!=np__257)).#trainneg
cnf(i_0_50107, plain, (np__232=np__0|np__232!=np__136|np__232!=np__257)).#trainneg
cnf(i_0_50109, plain, (np__6400=np__0|np__6400!=np__136|np__6400!=np__257)).#trainneg
cnf(i_0_28302, plain, (np__253!=np__3|~r1_xxreal_0(np__257,np__253))).#trainneg
cnf(i_0_28303, plain, (np__255!=np__3|~r1_xxreal_0(np__257,np__255))).#trainneg
cnf(i_0_28304, plain, (np__120!=np__3|~r1_xxreal_0(np__257,np__120))).#trainneg
cnf(i_0_28305, plain, (np__55!=np__3|~r1_xxreal_0(np__257,np__55))).#trainneg
cnf(i_0_28306, plain, (np__6305!=np__3|~r1_xxreal_0(np__257,np__6305))).#trainneg
cnf(i_0_15117, plain, (k2_nat_d(k2_xcmplx_0(X1,np__136),X1)=np__136|r1_xxreal_0(X1,np__136)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_54674, plain, (k4_nat_d(k2_xcmplx_0(X2,X1),X2)=k2_nat_d(k2_xcmplx_0(X2,X1),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_54653, plain, (k4_nat_d(np__6561,np__6425)=np__136)).#trainneg
cnf(i_0_56583, plain, (k2_nat_d(np__136,np__1)=np__0)).#trainneg
cnf(i_0_54596, plain, (k4_nat_d(np__8,np__3)=k2_nat_d(np__8,np__3)|r1_xxreal_0(np__3,np__5))).#trainneg
cnf(i_0_56480, plain, (np__8=np__0|np__8!=np__4|np__8!=np__257)).#trainneg
cnf(i_0_1998, plain, (k3_xcmplx_0(np__2,k3_xcmplx_0(np__4,X1))=k3_xcmplx_0(np__8,X1)|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_50111, plain, (np__6557=np__0|np__6557!=np__136|np__6557!=np__257)).#trainneg
cnf(i_0_50113, plain, (np__6344=np__0|np__6344!=np__136|np__6344!=np__257)).#trainneg
cnf(i_0_50115, plain, (np__22=np__0|np__22!=np__136|np__22!=np__257)).#trainneg
cnf(i_0_28307, plain, (np__7!=np__3|~r1_xxreal_0(np__257,np__7))).#trainneg
cnf(i_0_28309, plain, (np__5!=np__3|~r1_xxreal_0(np__257,np__5))).#trainneg
cnf(i_0_28310, plain, (np__252!=np__3|~r1_xxreal_0(np__257,np__252))).#trainneg
cnf(i_0_28312, plain, (np__6558!=np__3|~r1_xxreal_0(np__257,np__6558))).#trainneg
cnf(i_0_28313, plain, (np__6424!=np__3|~r1_xxreal_0(np__257,np__6424))).#trainneg
cnf(i_0_56992, plain, (k3_xcmplx_0(np__8,np__1)=np__8)).#trainneg
cnf(i_0_56148, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__1),X1)=k4_nat_d(k2_xcmplx_0(X2,X1),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_54639, plain, (k4_nat_d(np__4,np__3)=np__1)).#trainneg
cnf(i_0_50117, plain, (np__111=np__0|np__111!=np__136|np__111!=np__257)).#trainneg
cnf(i_0_1997, plain, (k3_xcmplx_0(np__2,k3_xcmplx_0(np__2,X1))=k3_xcmplx_0(np__4,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_50119, plain, (np__121=np__0|np__121!=np__136|np__121!=np__257)).#trainneg
cnf(i_0_50121, plain, (np__132=np__0|np__132!=np__136|np__132!=np__257)).#trainneg
cnf(i_0_50123, plain, (np__133=np__0|np__133!=np__136|np__133!=np__257)).#trainneg
cnf(i_0_28314, plain, (np__6304!=np__3|~r1_xxreal_0(np__257,np__6304))).#trainneg
cnf(i_0_28315, plain, (np__80!=np__3|~r1_xxreal_0(np__257,np__80))).#trainneg
cnf(i_0_28316, plain, (np__134!=np__3|~r1_xxreal_0(np__257,np__134))).#trainneg
cnf(i_0_28317, plain, (np__232!=np__3|~r1_xxreal_0(np__257,np__232))).#trainneg
cnf(i_0_28318, plain, (np__6400!=np__3|~r1_xxreal_0(np__257,np__6400))).#trainneg
cnf(i_0_54719, plain, (k4_nat_d(k2_xcmplx_0(X2,X1),X1)=k2_nat_d(k2_xcmplx_0(X1,X2),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_54606, plain, (k4_nat_d(X1,X1)=k2_nat_d(X1,X1)|r1_xxreal_0(X1,np__0)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_54619, plain, (k4_nat_d(np__6,np__4)=np__2)).#trainneg
cnf(i_0_50125, plain, (np__175=np__0|np__175!=np__136|np__175!=np__257)).#trainneg
cnf(i_0_1951, plain, (k3_xcmplx_0(np__4,k7_xcmplx_0(np__2,X1))=k7_xcmplx_0(np__8,X1)|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_50127, plain, (np__176=np__0|np__176!=np__136|np__176!=np__257)).#trainneg
cnf(i_0_50129, plain, (np__21=np__0|np__21!=np__136|np__21!=np__257)).#trainneg
cnf(i_0_50131, plain, (np__6169=np__0|np__6169!=np__136|np__6169!=np__257)).#trainneg
cnf(i_0_28319, plain, (np__6557!=np__3|~r1_xxreal_0(np__257,np__6557))).#trainneg
cnf(i_0_28320, plain, (np__6344!=np__3|~r1_xxreal_0(np__257,np__6344))).#trainneg
cnf(i_0_28322, plain, (np__22!=np__3|~r1_xxreal_0(np__257,np__22))).#trainneg
cnf(i_0_28323, plain, (np__111!=np__3|~r1_xxreal_0(np__257,np__111))).#trainneg
cnf(i_0_28324, plain, (np__121!=np__3|~r1_xxreal_0(np__257,np__121))).#trainneg
cnf(i_0_54620, plain, (k4_nat_d(np__5,np__3)=np__2)).#trainneg
cnf(i_0_54977, plain, (k2_nat_d(k2_xcmplx_0(X2,X1),X1)=k2_nat_d(k2_xcmplx_0(X1,X2),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_56288, plain, (k2_nat_d(k2_xcmplx_0(np__257,X1),np__257)=np__136|X1!=np__136|~v7_ordinal1(k2_xcmplx_0(np__257,X1))|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_56898, plain, (k4_nat_d(k3_xcmplx_0(X3,k4_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X3,X2),X1),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k4_nat_d(X1,X2),k5_numbers)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_56893, plain, (k4_nat_d(k4_nat_d(X1,X2),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__1,X2),X1),X2)|r1_xxreal_0(X2,np__1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X1,X2),k1_numbers))).#trainneg
cnf(i_0_50133, plain, (np__6421=np__0|np__6421!=np__136|np__6421!=np__257)).#trainneg
cnf(i_0_1942, plain, (k3_xcmplx_0(np__2,k7_xcmplx_0(np__4,X1))=k7_xcmplx_0(np__8,X1)|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_50135, plain, (np__6422=np__0|np__6422!=np__136|np__6422!=np__257)).#trainneg
cnf(i_0_50137, plain, (np__6423=np__0|np__6423!=np__136|np__6423!=np__257)).#trainneg
cnf(i_0_50139, plain, (np__6559=np__0|np__6559!=np__136|np__6559!=np__257)).#trainneg
cnf(i_0_28325, plain, (np__132!=np__3|~r1_xxreal_0(np__257,np__132))).#trainneg
cnf(i_0_28326, plain, (np__133!=np__3|~r1_xxreal_0(np__257,np__133))).#trainneg
cnf(i_0_28327, plain, (np__175!=np__3|~r1_xxreal_0(np__257,np__175))).#trainneg
cnf(i_0_28328, plain, (np__176!=np__3|~r1_xxreal_0(np__257,np__176))).#trainneg
cnf(i_0_28329, plain, (np__21!=np__3|~r1_xxreal_0(np__257,np__21))).#trainneg
cnf(i_0_57773, plain, (k4_nat_d(k3_xcmplx_0(np__6,np__1),np__4)=k4_nat_d(np__6,np__4))).#trainneg
cnf(i_0_57793, plain, (k4_nat_d(k3_xcmplx_0(np__4,np__1),np__3)=k4_nat_d(np__4,np__3))).#trainneg
cnf(i_0_8996, plain, (r1_xxreal_0(X1,np__6557)|r1_xxreal_0(np__6557,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50141, plain, (np__6560=np__0|np__6560!=np__136|np__6560!=np__257)).#trainneg
cnf(i_0_1941, plain, (k3_xcmplx_0(np__2,k7_xcmplx_0(np__2,X1))=k7_xcmplx_0(np__4,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_50143, plain, (np__77=np__0|np__77!=np__136|np__77!=np__257)).#trainneg
cnf(i_0_50145, plain, (np__78=np__0|np__78!=np__136|np__78!=np__257)).#trainneg
cnf(i_0_50147, plain, (np__24=np__0|np__24!=np__136|np__24!=np__257)).#trainneg
cnf(i_0_28330, plain, (np__6169!=np__3|~r1_xxreal_0(np__257,np__6169))).#trainneg
cnf(i_0_28331, plain, (np__6421!=np__3|~r1_xxreal_0(np__257,np__6421))).#trainneg
cnf(i_0_28332, plain, (np__6422!=np__3|~r1_xxreal_0(np__257,np__6422))).#trainneg
cnf(i_0_28333, plain, (np__6423!=np__3|~r1_xxreal_0(np__257,np__6423))).#trainneg
cnf(i_0_28334, plain, (np__6559!=np__3|~r1_xxreal_0(np__257,np__6559))).#trainneg
cnf(i_0_57774, plain, (k4_nat_d(k3_xcmplx_0(np__5,np__1),np__3)=k4_nat_d(np__5,np__3))).#trainneg
cnf(i_0_57299, plain, (k4_nat_d(np__4,np__4)=np__0)).#trainneg
cnf(i_0_57940, plain, (k3_xcmplx_0(np__4,np__1)=np__4)).#trainneg
cnf(i_0_58723, plain, (k2_nat_d(k2_xcmplx_0(np__0,X1),X1)=k2_nat_d(X1,X1)|r1_xxreal_0(X1,np__0)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9000, plain, (r1_xxreal_0(X1,np__6558)|r1_xxreal_0(np__6558,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50152, plain, (np__135=np__0|np__135!=np__136|np__135!=np__257)).#trainneg
cnf(i_0_1932, plain, (v1_abian(k4_xcmplx_0(np__2))|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_50153, plain, (np__6168=np__0|np__6168!=np__136|np__6168!=np__257)).#trainneg
cnf(i_0_50157, plain, (np__6536=np__0|np__6536!=np__136|np__6536!=np__257)).#trainneg
cnf(i_0_50159, plain, (np__56=np__0|np__56!=np__136|np__56!=np__257)).#trainneg
cnf(i_0_28335, plain, (np__6560!=np__3|~r1_xxreal_0(np__257,np__6560))).#trainneg
cnf(i_0_28336, plain, (np__77!=np__3|~r1_xxreal_0(np__257,np__77))).#trainneg
cnf(i_0_28337, plain, (np__78!=np__3|~r1_xxreal_0(np__257,np__78))).#trainneg
cnf(i_0_28338, plain, (np__24!=np__3|~r1_xxreal_0(np__257,np__24))).#trainneg
cnf(i_0_28342, plain, (np__135!=np__3|~r1_xxreal_0(np__257,np__135))).#trainneg
cnf(i_0_57616, plain, (k2_nat_d(k2_xcmplx_0(np__8,X1),np__8)=X1|r1_xxreal_0(np__8,X1)|~v7_ordinal1(k2_xcmplx_0(np__8,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_54591, plain, (k2_nat_d(k2_xcmplx_0(k2_nat_d(X1,X2),X2),X2)=k4_nat_d(k3_xcmplx_0(X1,np__1),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_xcmplx_0(k2_nat_d(X1,X2),X2),k1_numbers))).#trainneg
cnf(i_0_59457, plain, (k4_nat_d(k2_nat_d(X1,X2),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__1,X2),X1),X2)|r1_xxreal_0(X2,np__1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X1,X2),k1_numbers))).#trainneg
cnf(i_0_51786, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X2,X3)),X2)=k2_nat_d(k4_nat_1(X3,X1),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k4_nat_1(X3,X1))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_49801, negated_conjecture, (r1_xxreal_0(np__257,k3_xcmplx_0(k13_newton(np__3,np__8),np__1))|r1_xxreal_0(np__257,k13_newton(np__3,np__8))|k3_xcmplx_0(k13_newton(np__3,np__8),np__1)!=np__136|~v7_ordinal1(k3_xcmplx_0(k13_newton(np__3,np__8),np__1))|~v7_ordinal1(k13_newton(np__3,np__8)))).#trainneg
cnf(i_0_9004, plain, (r1_xxreal_0(X1,np__6480)|r1_xxreal_0(np__6480,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50161, plain, (np__23=np__0|np__23!=np__136|np__23!=np__257)).#trainneg
cnf(i_0_2000, plain, (k3_xcmplx_0(np__6,k3_xcmplx_0(np__4,X1))=k3_xcmplx_0(np__24,X1)|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_50163, plain, (np__6289=np__0|np__6289!=np__136|np__6289!=np__257)).#trainneg
cnf(i_0_50165, plain, (np__231=np__0|np__231!=np__136|np__231!=np__257)).#trainneg
cnf(i_0_50168, plain, (np__79=np__0|np__79!=np__136|np__79!=np__257)).#trainneg
cnf(i_0_28343, plain, (np__6168!=np__3|~r1_xxreal_0(np__257,np__6168))).#trainneg
cnf(i_0_28344, plain, (np__6!=np__3|~r1_xxreal_0(np__257,np__6))).#trainneg
cnf(i_0_28345, plain, (np__6536!=np__3|~r1_xxreal_0(np__257,np__6536))).#trainneg
cnf(i_0_28346, plain, (np__56!=np__3|~r1_xxreal_0(np__257,np__56))).#trainneg
cnf(i_0_28347, plain, (np__23!=np__3|~r1_xxreal_0(np__257,np__23))).#trainneg
cnf(i_0_61086, negated_conjecture, (r1_xxreal_0(np__257,k13_newton(np__3,np__8))|k3_xcmplx_0(k13_newton(np__3,np__8),np__1)!=np__136|~v7_ordinal1(k13_newton(np__3,np__8)))).#trainneg
cnf(i_0_59919, plain, (k3_xcmplx_0(np__1,np__4)=np__4)).#trainneg
cnf(i_0_60527, plain, (k2_nat_d(k2_xcmplx_0(X2,k2_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__1),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X1,X2),k1_numbers))).#trainneg
cnf(i_0_60343, plain, (X1=np__0|r1_xxreal_0(X1,k2_nat_d(X2,X1))|~r1_xxreal_0(X1,k4_nat_d(k3_xcmplx_0(X2,np__1),X1))|~v7_ordinal1(k2_xcmplx_0(k2_nat_d(X2,X1),X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_9195, plain, (r1_xxreal_0(X1,np__77)|r1_xxreal_0(np__77,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1944, plain, (k3_xcmplx_0(np__6,k7_xcmplx_0(np__4,X1))=k7_xcmplx_0(np__24,X1)|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_28348, plain, (np__6289!=np__3|~r1_xxreal_0(np__257,np__6289))).#trainneg
cnf(i_0_28349, plain, (np__231!=np__3|~r1_xxreal_0(np__257,np__231))).#trainneg
cnf(i_0_28351, plain, (np__79!=np__3|~r1_xxreal_0(np__257,np__79))).#trainneg
cnf(i_0_28360, plain, (np__8!=np__3|~r1_xxreal_0(np__257,np__8))).#trainneg
cnf(i_0_28365, plain, (np__6480!=np__81|~r1_xxreal_0(np__257,np__6480))).#trainneg
cnf(i_0_60247, plain, (k2_nat_d(k2_xcmplx_0(X1,np__8),np__8)=X1|r1_xxreal_0(np__8,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__8))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15181, plain, (k4_nat_d(k3_xcmplx_0(k1_newton(np__3,np__4),X1),np__257)=k4_nat_d(k4_nat_1(np__81,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k1_newton(np__3,np__4))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_45393, plain, (k4_nat_d(k4_nat_1(np__136,X1),X2)=k4_nat_d(k3_xcmplx_0(X1,np__136),X2)|r1_xxreal_0(X2,np__136)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9199, plain, (r1_xxreal_0(X1,np__6304)|r1_xxreal_0(np__6304,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1850, plain, (v1_int_1(k4_xcmplx_0(np__6559))|~v1_int_1(np__6561))).#trainneg
cnf(i_0_28366, plain, (np__254!=np__81|~r1_xxreal_0(np__257,np__254))).#trainneg
cnf(i_0_28367, plain, (np__253!=np__81|~r1_xxreal_0(np__257,np__253))).#trainneg
cnf(i_0_28368, plain, (np__255!=np__81|~r1_xxreal_0(np__257,np__255))).#trainneg
cnf(i_0_28369, plain, (np__120!=np__81|~r1_xxreal_0(np__257,np__120))).#trainneg
cnf(i_0_28370, plain, (np__55!=np__81|~r1_xxreal_0(np__257,np__55))).#trainneg
cnf(i_0_18791, plain, (k2_nat_d(k2_xcmplx_0(k2_nat_d(X1,X2),X2),X2)=k4_nat_d(k3_xcmplx_0(np__1,X1),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_xcmplx_0(k2_nat_d(X1,X2),X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_24705, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),k2_xcmplx_0(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9203, plain, (r1_xxreal_0(X1,np__80)|r1_xxreal_0(np__80,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1802, plain, (v1_int_1(k4_xcmplx_0(np__3))|~v1_int_1(np__5))).#trainneg
cnf(i_0_28371, plain, (np__6305!=np__81|~r1_xxreal_0(np__257,np__6305))).#trainneg
cnf(i_0_28372, plain, (np__7!=np__81|~r1_xxreal_0(np__257,np__7))).#trainneg
cnf(i_0_28374, plain, (np__5!=np__81|~r1_xxreal_0(np__257,np__5))).#trainneg
cnf(i_0_28375, plain, (np__252!=np__81|~r1_xxreal_0(np__257,np__252))).#trainneg
cnf(i_0_28377, plain, (np__6558!=np__81|~r1_xxreal_0(np__257,np__6558))).#trainneg
cnf(i_0_14681, plain, (k2_nat_d(k3_xcmplx_0(X1,k4_nat_d(X2,X3)),X3)=k4_nat_d(k3_xcmplx_0(X2,k2_xcmplx_0(X1,X3)),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k3_xcmplx_0(X1,k4_nat_d(X2,X3)))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_24475, plain, (k4_nat_d(k3_xcmplx_0(X3,k2_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X3),X1),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_20884, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X1)),X1)=k4_nat_d(k4_nat_1(X2,X3),X1)|r1_xxreal_0(X1,X2)|r1_xxreal_0(X1,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_61514, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(X2,X1))|~v7_ordinal1(k2_xcmplx_0(k2_nat_d(X2,X1),X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_61196, negated_conjecture, (r1_xxreal_0(np__257,k1_newton(np__3,np__8))|k3_xcmplx_0(k1_newton(np__3,np__8),np__1)!=np__136|~v7_ordinal1(k1_newton(np__3,np__8)))).#trainneg
cnf(i_0_61508, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(k3_xcmplx_0(np__1,X2),X1))|~v7_ordinal1(k2_xcmplx_0(k2_nat_d(X2,X1),X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_62232, plain, (k4_nat_d(np__136,X1)=k2_nat_d(np__136,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_61643, plain, (np__8=np__0|np__8!=np__3|np__8!=np__257)).#trainneg
cnf(i_0_1871, plain, (v1_xreal_0(k4_xcmplx_0(np__6168))|~v1_xreal_0(np__257))).#trainneg
cnf(i_0_28378, plain, (np__6424!=np__81|~r1_xxreal_0(np__257,np__6424))).#trainneg
cnf(i_0_28379, plain, (np__6304!=np__81|~r1_xxreal_0(np__257,np__6304))).#trainneg
cnf(i_0_28380, plain, (np__80!=np__81|~r1_xxreal_0(np__257,np__80))).#trainneg
cnf(i_0_28381, plain, (np__134!=np__81|~r1_xxreal_0(np__257,np__134))).#trainneg
cnf(i_0_28382, plain, (np__232!=np__81|~r1_xxreal_0(np__257,np__232))).#trainneg
cnf(i_0_62233, plain, (k4_nat_d(np__8,X1)=k2_nat_d(np__8,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_61198, negated_conjecture, (r1_xxreal_0(np__257,k13_newton(np__3,np__8))|k3_xcmplx_0(np__1,k13_newton(np__3,np__8))!=np__136|~v7_ordinal1(k13_newton(np__3,np__8)))).#trainneg
cnf(i_0_62220, plain, (k2_nat_d(k3_xcmplx_0(X1,np__1),X2)=k4_nat_d(X1,X2)|~v7_ordinal1(k3_xcmplx_0(X1,np__1))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_62234, plain, (k4_nat_d(np__3,X1)=k2_nat_d(np__3,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1858, plain, (v1_xreal_0(k4_xcmplx_0(np__1))|~v1_xreal_0(np__256))).#trainneg
cnf(i_0_28383, plain, (np__6400!=np__81|~r1_xxreal_0(np__257,np__6400))).#trainneg
cnf(i_0_28384, plain, (np__6557!=np__81|~r1_xxreal_0(np__257,np__6557))).#trainneg
cnf(i_0_28385, plain, (np__6344!=np__81|~r1_xxreal_0(np__257,np__6344))).#trainneg
cnf(i_0_28387, plain, (np__22!=np__81|~r1_xxreal_0(np__257,np__22))).#trainneg
cnf(i_0_28388, plain, (np__111!=np__81|~r1_xxreal_0(np__257,np__111))).#trainneg
cnf(i_0_62231, plain, (k4_nat_d(np__4,X1)=k2_nat_d(np__4,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_61465, plain, (k4_nat_d(np__8,np__4)=np__0)).#trainneg
cnf(i_0_62732, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X1,X3)),X1)=k4_nat_d(k3_xcmplx_0(X2,X3),X1)|r1_xxreal_0(X1,X3)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9207, plain, (r1_xxreal_0(X1,np__79)|r1_xxreal_0(np__79,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_28389, plain, (np__121!=np__81|~r1_xxreal_0(np__257,np__121))).#trainneg
cnf(i_0_28390, plain, (np__132!=np__81|~r1_xxreal_0(np__257,np__132))).#trainneg
cnf(i_0_28391, plain, (np__133!=np__81|~r1_xxreal_0(np__257,np__133))).#trainneg
cnf(i_0_28392, plain, (np__175!=np__81|~r1_xxreal_0(np__257,np__175))).#trainneg
cnf(i_0_28393, plain, (np__176!=np__81|~r1_xxreal_0(np__257,np__176))).#trainneg
cnf(i_0_62228, plain, (k4_nat_d(np__2,X1)=k2_nat_d(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_63511, plain, (k4_nat_d(k2_nat_d(X1,X2),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,np__1),X1),X2)|r1_xxreal_0(X2,np__1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X1,X2),k1_numbers))).#trainneg
cnf(i_0_61998, plain, (k4_nat_d(k3_xcmplx_0(X1,np__136),X2)=k2_nat_d(k3_xcmplx_0(X1,np__136),X2)|r1_xxreal_0(X2,np__136)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_62215, plain, (k4_nat_d(np__1,X1)=k2_nat_d(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_62229, plain, (k4_nat_d(np__257,X1)=k2_nat_d(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_35225, plain, (k2_nat_d(k3_xcmplx_0(k4_nat_d(X1,X2),X3),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X3,X2),X1),X2)|r1_xxreal_0(X2,k4_nat_d(X1,X2))|r1_xxreal_0(X2,X3)|~v7_ordinal1(k4_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_63992, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(X2,X1))|~v7_ordinal1(k2_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_9211, plain, (r1_xxreal_0(X1,np__175)|r1_xxreal_0(np__175,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_28394, plain, (np__21!=np__81|~r1_xxreal_0(np__257,np__21))).#trainneg
cnf(i_0_28395, plain, (np__6169!=np__81|~r1_xxreal_0(np__257,np__6169))).#trainneg
cnf(i_0_28396, plain, (np__6421!=np__81|~r1_xxreal_0(np__257,np__6421))).#trainneg
cnf(i_0_28397, plain, (np__6422!=np__81|~r1_xxreal_0(np__257,np__6422))).#trainneg
cnf(i_0_28398, plain, (np__6423!=np__81|~r1_xxreal_0(np__257,np__6423))).#trainneg
cnf(i_0_65179, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__81,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__81),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_64250, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__8,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__8),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_65271, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__2,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__2),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_65833, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__257,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__257),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_64503, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__3),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_64130, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__136,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__136),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_64615, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__4,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__4),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_65742, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9215, plain, (r1_xxreal_0(X1,np__254)|r1_xxreal_0(np__254,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2005, plain, (k3_xcmplx_0(np__81,k3_xcmplx_0(np__1,X1))=k3_xcmplx_0(np__81,X1)|~v1_xcmplx_0(np__81)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_28399, plain, (np__6559!=np__81|~r1_xxreal_0(np__257,np__6559))).#trainneg
cnf(i_0_28400, plain, (np__6560!=np__81|~r1_xxreal_0(np__257,np__6560))).#trainneg
cnf(i_0_28401, plain, (np__77!=np__81|~r1_xxreal_0(np__257,np__77))).#trainneg
cnf(i_0_28402, plain, (np__78!=np__81|~r1_xxreal_0(np__257,np__78))).#trainneg
cnf(i_0_28403, plain, (np__24!=np__81|~r1_xxreal_0(np__257,np__24))).#trainneg
cnf(i_0_65266, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(np__2,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__2),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_66220, plain, (X1=np__0|~r1_xxreal_0(X1,k2_nat_d(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_66380, plain, (X1=np__0|~r1_xxreal_0(X1,k2_nat_d(np__136,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_66387, plain, (X1=np__0|~r1_xxreal_0(X1,k2_nat_d(np__8,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_66377, plain, (X1=np__0|~r1_xxreal_0(X1,k2_nat_d(np__4,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_66390, plain, (X1=np__0|~r1_xxreal_0(X1,k2_nat_d(np__3,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_66372, plain, (X1=np__0|~r1_xxreal_0(X1,k2_nat_d(np__257,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_66374, plain, (X1=np__0|~r1_xxreal_0(X1,k2_nat_d(np__81,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_66279, plain, (X1=np__0|~r1_xxreal_0(X1,k2_nat_d(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_66213, plain, (X1=np__0|r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,k4_nat_d(k2_xcmplx_0(X2,X1),X1))|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_66393, plain, (X1=np__0|r1_xxreal_0(X1,np__0)|~r1_xxreal_0(X1,k2_nat_d(X1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9219, plain, (r1_xxreal_0(X1,np__78)|r1_xxreal_0(np__78,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2001, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(np__81,X1))=k3_xcmplx_0(np__81,X1)|~v1_xcmplx_0(np__81)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_28407, plain, (np__135!=np__81|~r1_xxreal_0(np__257,np__135))).#trainneg
cnf(i_0_28408, plain, (np__6168!=np__81|~r1_xxreal_0(np__257,np__6168))).#trainneg
cnf(i_0_28409, plain, (np__6!=np__81|~r1_xxreal_0(np__257,np__6))).#trainneg
cnf(i_0_28410, plain, (np__6536!=np__81|~r1_xxreal_0(np__257,np__6536))).#trainneg
cnf(i_0_28411, plain, (np__56!=np__81|~r1_xxreal_0(np__257,np__56))).#trainneg
cnf(i_0_66272, plain, (X1=np__0|r1_xxreal_0(X1,X3)|r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,k4_nat_d(k3_xcmplx_0(X2,X3),X1))|~v7_ordinal1(k4_nat_1(X2,X3))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_66217, plain, (X1=np__0|r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,k4_nat_d(k2_xcmplx_0(X1,X2),X1))|~v7_ordinal1(k2_xcmplx_0(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_65247, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__2,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__2),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64245, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(np__8,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__8),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_65238, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__2,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__2,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64106, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__136,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__136),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_65718, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__1,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__1),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_65272, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__2,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__2,X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_64498, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(np__3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__3),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_65149, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__81,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__81),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64226, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__8,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__8),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_69743, plain, (np__0=np__3|r1_xxreal_0(np__3,np__5)|~r1_xxreal_0(np__3,k4_nat_d(np__8,np__3)))).#trainneg
cnf(i_0_1949, plain, (k3_xcmplx_0(np__81,k7_xcmplx_0(np__1,X1))=k7_xcmplx_0(np__81,X1)|~v1_xcmplx_0(np__81)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_69744, plain, (~r1_xxreal_0(np__6561,k4_nat_d(np__6561,np__6561)))).#trainneg
cnf(i_0_28412, plain, (np__23!=np__81|~r1_xxreal_0(np__257,np__23))).#trainneg
cnf(i_0_28413, plain, (np__6289!=np__81|~r1_xxreal_0(np__257,np__6289))).#trainneg
cnf(i_0_28414, plain, (np__231!=np__81|~r1_xxreal_0(np__257,np__231))).#trainneg
cnf(i_0_69826, plain, (X1=np__0|r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,k2_nat_d(k2_xcmplx_0(X2,X1),X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_69751, plain, (X1=np__0|r1_xxreal_0(X1,np__0)|~r1_xxreal_0(X1,k4_nat_d(X1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1945, plain, (k3_xcmplx_0(np__1,k7_xcmplx_0(np__81,X1))=k7_xcmplx_0(np__81,X1)|~v1_xcmplx_0(np__81)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_28416, plain, (np__79!=np__81|~r1_xxreal_0(np__257,np__79))).#trainneg
cnf(i_0_28425, plain, (np__8!=np__81|~r1_xxreal_0(np__257,np__8))).#trainneg
cnf(i_0_70225, plain, (X1=np__0|r1_xxreal_0(X1,k2_nat_d(X2,X1))|~r1_xxreal_0(X1,k4_nat_d(k3_xcmplx_0(X2,np__1),X1))|~v7_ordinal1(k2_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_69828, plain, (X1=np__0|r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,k2_nat_d(k2_xcmplx_0(X1,X2),X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_70224, plain, (X1=np__0|r1_xxreal_0(X1,k2_nat_d(X2,X1))|~r1_xxreal_0(X1,k4_nat_d(k3_xcmplx_0(np__1,X2),X1))|~v7_ordinal1(k2_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_70212, plain, (X1=np__0|r1_xxreal_0(X1,np__136)|~r1_xxreal_0(X1,k4_nat_d(k2_xcmplx_0(np__136,X1),X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_70995, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__2,np__257),X1),np__257)=k4_nat_d(k3_xcmplx_0(np__2,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9223, plain, (r1_xxreal_0(X1,np__24)|r1_xxreal_0(np__24,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1938, plain, (k3_xcmplx_0(np__0,k7_xcmplx_0(np__257,X1))=k7_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_56482, plain, (np__6480=np__0|np__6480!=np__3|np__6480!=np__257)).#trainneg
cnf(i_0_56483, plain, (np__254=np__0|np__254!=np__3|np__254!=np__257)).#trainneg
cnf(i_0_69796, plain, (X1=np__0|r1_xxreal_0(X1,np__136)|~r1_xxreal_0(X1,k4_nat_d(k2_xcmplx_0(X1,np__136),X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_66273, plain, (X1=np__0|r1_xxreal_0(X1,X3)|r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,k4_nat_d(k3_xcmplx_0(X3,X2),X1))|~v7_ordinal1(k4_nat_1(X2,X3))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_65155, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__81,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__81),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_3125, plain, (k2_xcmplx_0(X1,np__1)=k2_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_64220, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__8,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__8),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_65241, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__2,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64473, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__3,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__3),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64610, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(np__4,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__4),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_63995, negated_conjecture, (r1_xxreal_0(np__257,k1_newton(np__3,np__8))|k3_xcmplx_0(np__1,k1_newton(np__3,np__8))!=np__136|~v7_ordinal1(k1_newton(np__3,np__8)))).#trainneg
cnf(i_0_9446, plain, (r1_xxreal_0(X1,np__23)|r1_xxreal_0(np__23,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_57008, plain, (np__253=np__0|np__253!=np__3|np__253!=np__257)).#trainneg
cnf(i_0_1947, plain, (k3_xcmplx_0(np__257,k7_xcmplx_0(np__0,X1))=k7_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_57009, plain, (np__255=np__0|np__255!=np__3|np__255!=np__257)).#trainneg
cnf(i_0_57010, plain, (np__120=np__0|np__120!=np__3|np__120!=np__257)).#trainneg
cnf(i_0_57011, plain, (np__55=np__0|np__55!=np__3|np__55!=np__257)).#trainneg
cnf(i_0_64470, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__3,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__3,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_65737, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(np__1,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_72695, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(np__2,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_72891, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(np__3,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_72704, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(np__8,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_72702, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(np__81,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_72889, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(np__136,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_72688, plain, (np__8=np__0|np__8!=np__81|np__8!=np__257)).#trainneg
cnf(i_0_1994, plain, (k3_xcmplx_0(np__0,k3_xcmplx_0(np__257,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_57012, plain, (np__6305=np__0|np__6305!=np__3|np__6305!=np__257)).#trainneg
cnf(i_0_57606, plain, (np__7=np__0|np__7!=np__3|np__7!=np__257)).#trainneg
cnf(i_0_57607, plain, (np__5=np__0|np__5!=np__3|np__5!=np__257)).#trainneg
cnf(i_0_46520, plain, (np__136!=np__6561|~r1_xxreal_0(np__6425,np__6561))).#trainneg
cnf(i_0_51277, plain, (np__6480!=np__0|~r1_xxreal_0(np__6561,np__6480))).#trainneg
cnf(i_0_51278, plain, (np__254!=np__0|~r1_xxreal_0(np__6561,np__254))).#trainneg
cnf(i_0_72700, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(np__257,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_72703, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(np__4,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_72696, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9450, plain, (r1_xxreal_0(X1,np__6559)|r1_xxreal_0(np__6559,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_57608, plain, (np__252=np__0|np__252!=np__3|np__252!=np__257)).#trainneg
cnf(i_0_2003, plain, (k3_xcmplx_0(np__257,k3_xcmplx_0(np__0,X1))=k3_xcmplx_0(np__0,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_57609, plain, (np__6558=np__0|np__6558!=np__3|np__6558!=np__257)).#trainneg
cnf(i_0_57610, plain, (np__6424=np__0|np__6424!=np__3|np__6424!=np__257)).#trainneg
cnf(i_0_57957, plain, (np__6304=np__0|np__6304!=np__3|np__6304!=np__257)).#trainneg
cnf(i_0_51279, plain, (np__253!=np__0|~r1_xxreal_0(np__6561,np__253))).#trainneg
cnf(i_0_51280, plain, (np__255!=np__0|~r1_xxreal_0(np__6561,np__255))).#trainneg
cnf(i_0_51281, plain, (np__120!=np__0|~r1_xxreal_0(np__6561,np__120))).#trainneg
cnf(i_0_51282, plain, (np__55!=np__0|~r1_xxreal_0(np__6561,np__55))).#trainneg
cnf(i_0_51283, plain, (np__6305!=np__0|~r1_xxreal_0(np__6561,np__6305))).#trainneg
cnf(i_0_72689, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(np__0,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9454, plain, (r1_xxreal_0(X1,np__231)|r1_xxreal_0(np__231,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_57958, plain, (np__80=np__0|np__80!=np__3|np__80!=np__257)).#trainneg
cnf(i_0_2004, plain, (k3_xcmplx_0(np__257,k3_xcmplx_0(np__25,X1))=k3_xcmplx_0(np__6425,X1)|~v1_xcmplx_0(np__25)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_57959, plain, (np__134=np__0|np__134!=np__3|np__134!=np__257)).#trainneg
cnf(i_0_57960, plain, (np__232=np__0|np__232!=np__3|np__232!=np__257)).#trainneg
cnf(i_0_57961, plain, (np__6400=np__0|np__6400!=np__3|np__6400!=np__257)).#trainneg
cnf(i_0_51284, plain, (np__7!=np__0|~r1_xxreal_0(np__6561,np__7))).#trainneg
cnf(i_0_51286, plain, (np__5!=np__0|~r1_xxreal_0(np__6561,np__5))).#trainneg
cnf(i_0_51287, plain, (np__252!=np__0|~r1_xxreal_0(np__6561,np__252))).#trainneg
cnf(i_0_51289, plain, (np__6558!=np__0|~r1_xxreal_0(np__6561,np__6558))).#trainneg
cnf(i_0_51290, plain, (np__6424!=np__0|~r1_xxreal_0(np__6561,np__6424))).#trainneg
cnf(i_0_65803, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__257,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__257),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_65809, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__257,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__257),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_9458, plain, (r1_xxreal_0(X1,np__6289)|r1_xxreal_0(np__6289,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_58352, plain, (np__6557=np__0|np__6557!=np__3|np__6557!=np__257)).#trainneg
cnf(i_0_1996, plain, (k3_xcmplx_0(np__25,k3_xcmplx_0(np__257,X1))=k3_xcmplx_0(np__6425,X1)|~v1_xcmplx_0(np__25)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_58353, plain, (np__6344=np__0|np__6344!=np__3|np__6344!=np__257)).#trainneg
cnf(i_0_58354, plain, (np__22=np__0|np__22!=np__3|np__22!=np__257)).#trainneg
cnf(i_0_58355, plain, (np__111=np__0|np__111!=np__3|np__111!=np__257)).#trainneg
cnf(i_0_51291, plain, (np__6304!=np__0|~r1_xxreal_0(np__6561,np__6304))).#trainneg
cnf(i_0_51292, plain, (np__80!=np__0|~r1_xxreal_0(np__6561,np__80))).#trainneg
cnf(i_0_51293, plain, (np__134!=np__0|~r1_xxreal_0(np__6561,np__134))).#trainneg
cnf(i_0_51294, plain, (np__232!=np__0|~r1_xxreal_0(np__6561,np__232))).#trainneg
cnf(i_0_51295, plain, (np__6400!=np__0|~r1_xxreal_0(np__6561,np__6400))).#trainneg
cnf(i_0_65232, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__2,X1),k4_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__2,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_65828, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(np__257,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__257),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_64125, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(np__136,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__136),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_64251, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__8,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__8,X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_65127, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__81,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__81),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_64479, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__3,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__3),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2989, plain, (k3_xcmplx_0(np__1,np__1)=np__1)).#trainneg
cnf(i_0_64217, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__8,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__8,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_75848, plain, (X1=np__0|~r1_xxreal_0(X1,k2_nat_d(np__0,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9462, plain, (r1_xxreal_0(X1,np__56)|r1_xxreal_0(np__56,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_58356, plain, (np__121=np__0|np__121!=np__3|np__121!=np__257)).#trainneg
cnf(i_0_1948, plain, (k3_xcmplx_0(np__257,k7_xcmplx_0(np__25,X1))=k7_xcmplx_0(np__6425,X1)|~v1_xcmplx_0(np__25)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_59482, plain, (np__132=np__0|np__132!=np__3|np__132!=np__257)).#trainneg
cnf(i_0_59483, plain, (np__133=np__0|np__133!=np__3|np__133!=np__257)).#trainneg
cnf(i_0_59484, plain, (np__175=np__0|np__175!=np__3|np__175!=np__257)).#trainneg
cnf(i_0_51296, plain, (np__6557!=np__0|~r1_xxreal_0(np__6561,np__6557))).#trainneg
cnf(i_0_51297, plain, (np__6344!=np__0|~r1_xxreal_0(np__6561,np__6344))).#trainneg
cnf(i_0_51300, plain, (np__22!=np__0|~r1_xxreal_0(np__6561,np__22))).#trainneg
cnf(i_0_51301, plain, (np__111!=np__0|~r1_xxreal_0(np__6561,np__111))).#trainneg
cnf(i_0_51302, plain, (np__121!=np__0|~r1_xxreal_0(np__6561,np__121))).#trainneg
cnf(i_0_64591, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__4,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__4),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64100, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__136,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__136),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_65146, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__81,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__81,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_65519, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__257,np__1),X1),np__257)=k4_nat_d(X1,np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_65249, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__2,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__2),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_65180, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__81,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__81,X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9466, plain, (r1_xxreal_0(X1,np__135)|r1_xxreal_0(np__135,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_59485, plain, (np__176=np__0|np__176!=np__3|np__176!=np__257)).#trainneg
cnf(i_0_1940, plain, (k3_xcmplx_0(np__25,k7_xcmplx_0(np__257,X1))=k7_xcmplx_0(np__6425,X1)|~v1_xcmplx_0(np__25)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_59486, plain, (np__21=np__0|np__21!=np__3|np__21!=np__257)).#trainneg
cnf(i_0_59748, plain, (np__6169=np__0|np__6169!=np__3|np__6169!=np__257)).#trainneg
cnf(i_0_59749, plain, (np__6421=np__0|np__6421!=np__3|np__6421!=np__257)).#trainneg
cnf(i_0_51303, plain, (np__132!=np__0|~r1_xxreal_0(np__6561,np__132))).#trainneg
cnf(i_0_51304, plain, (np__133!=np__0|~r1_xxreal_0(np__6561,np__133))).#trainneg
cnf(i_0_51305, plain, (np__175!=np__0|~r1_xxreal_0(np__6561,np__175))).#trainneg
cnf(i_0_51306, plain, (np__176!=np__0|~r1_xxreal_0(np__6561,np__176))).#trainneg
cnf(i_0_51307, plain, (np__21!=np__0|~r1_xxreal_0(np__6561,np__21))).#trainneg
cnf(i_0_64464, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__3,X1),k4_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__3,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_75612, plain, (k4_nat_d(k4_nat_1(np__136,k2_nat_d(np__1,np__257)),np__257)=np__136)).#trainneg
cnf(i_0_64198, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__8,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__8),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_65242, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__2,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__2,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64187, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__8,np__257),X1),np__257)=k4_nat_d(k3_xcmplx_0(np__8,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_80938, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__8),X2)=k4_nat_d(k3_xcmplx_0(np__8,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_80768, plain, (k4_nat_d(k4_nat_1(np__136,np__1),np__257)=np__136)).#trainneg
cnf(i_0_64504, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__3,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__3,X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_65800, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__257,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__257,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64585, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__4,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__4),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_9470, plain, (r1_xxreal_0(X1,np__6423)|r1_xxreal_0(np__6423,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_59750, plain, (np__6422=np__0|np__6422!=np__3|np__6422!=np__257)).#trainneg
cnf(i_0_1849, plain, (v1_int_1(k4_xcmplx_0(np__256))|~v1_int_1(np__257))).#trainneg
cnf(i_0_59751, plain, (np__6423=np__0|np__6423!=np__3|np__6423!=np__257)).#trainneg
cnf(i_0_59752, plain, (np__6559=np__0|np__6559!=np__3|np__6559!=np__257)).#trainneg
cnf(i_0_60146, plain, (np__6560=np__0|np__6560!=np__3|np__6560!=np__257)).#trainneg
cnf(i_0_51308, plain, (np__6169!=np__0|~r1_xxreal_0(np__6561,np__6169))).#trainneg
cnf(i_0_51309, plain, (np__6421!=np__0|~r1_xxreal_0(np__6561,np__6421))).#trainneg
cnf(i_0_51310, plain, (np__6422!=np__0|~r1_xxreal_0(np__6561,np__6422))).#trainneg
cnf(i_0_51311, plain, (np__6423!=np__0|~r1_xxreal_0(np__6561,np__6423))).#trainneg
cnf(i_0_51312, plain, (np__6559!=np__0|~r1_xxreal_0(np__6561,np__6559))).#trainneg
cnf(i_0_65720, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__1,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__1),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_65246, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__2,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(np__2,k2_xcmplx_0(X1,X2)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64210, plain, (m1_subset_1(k2_nat_d(np__8,X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_64211, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__8,X1),k4_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__8,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_65219, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__2,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__2),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_65231, plain, (m1_subset_1(k2_nat_d(np__2,X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_65157, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__81,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__81),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64078, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__136,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__136),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_64097, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__136,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__136,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_84052, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__136),X2)=k4_nat_d(k3_xcmplx_0(np__136,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_83671, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__2),X2)=k4_nat_d(k3_xcmplx_0(np__2,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9474, plain, (r1_xxreal_0(X1,np__6422)|r1_xxreal_0(np__6422,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_60147, plain, (np__77=np__0|np__77!=np__3|np__77!=np__257)).#trainneg
cnf(i_0_1673, plain, (v1_int_1(np__256))).#trainneg
cnf(i_0_60148, plain, (np__78=np__0|np__78!=np__3|np__78!=np__257)).#trainneg
cnf(i_0_51314, plain, (np__77!=np__0|~r1_xxreal_0(np__6561,np__77))).#trainneg
cnf(i_0_51315, plain, (np__78!=np__0|~r1_xxreal_0(np__6561,np__78))).#trainneg
cnf(i_0_51316, plain, (np__24!=np__0|~r1_xxreal_0(np__6561,np__24))).#trainneg
cnf(i_0_51320, plain, (np__135!=np__0|~r1_xxreal_0(np__6561,np__135))).#trainneg
cnf(i_0_51321, plain, (np__6168!=np__0|~r1_xxreal_0(np__6561,np__6168))).#trainneg
cnf(i_0_81730, plain, (k2_nat_d(k4_nat_1(np__1,np__136),np__257)=np__136|~v7_ordinal1(k4_nat_1(np__1,np__136)))).#trainneg
cnf(i_0_65743, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__1,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__1,X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_64228, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__8,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__8),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64108, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__136,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__136),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_80909, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(k2_nat_d(X1,X2),X2),np__8),X2)=k4_nat_d(k3_xcmplx_0(X1,np__8),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_64474, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__3,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__3,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_65712, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__1,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__1),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_65151, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__81,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__81),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_80940, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__8),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__8),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9671, plain, (r1_xxreal_0(X1,np__6421)|r1_xxreal_0(np__6421,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_60149, plain, (np__24=np__0|np__24!=np__3|np__24!=np__257)).#trainneg
cnf(i_0_1848, plain, (v1_int_1(k4_xcmplx_0(np__6168))|~v1_int_1(np__6425))).#trainneg
cnf(i_0_1787, plain, (v1_int_1(np__175))).#trainneg
cnf(i_0_1784, plain, (v1_int_1(np__6305))).#trainneg
cnf(i_0_60150, plain, (np__135=np__0|np__135!=np__3|np__135!=np__257)).#trainneg
cnf(i_0_1442, plain, (~v1_abian(np__257)|~v1_abian(np__256))).#trainneg
cnf(i_0_51322, plain, (np__6!=np__0|~r1_xxreal_0(np__6561,np__6))).#trainneg
cnf(i_0_51323, plain, (np__6536!=np__0|~r1_xxreal_0(np__6561,np__6536))).#trainneg
cnf(i_0_51324, plain, (np__56!=np__0|~r1_xxreal_0(np__6561,np__56))).#trainneg
cnf(i_0_51325, plain, (np__23!=np__0|~r1_xxreal_0(np__6561,np__23))).#trainneg
cnf(i_0_86539, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__1,np__257),np__136),np__257)=np__136)).#trainneg
cnf(i_0_85967, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,k2_nat_d(X1,X2)),np__8),X2)=k4_nat_d(k3_xcmplx_0(X1,np__8),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X1,X2),k1_numbers))).#trainneg
cnf(i_0_84603, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__136),X2)=k4_nat_d(k3_xcmplx_0(np__136,k2_xcmplx_0(X2,X1)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_85649, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__136),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__136),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_84801, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__2),X2)=k4_nat_d(k3_xcmplx_0(np__2,k2_xcmplx_0(X2,X1)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9676, plain, (r1_xxreal_0(X1,np__6169)|r1_xxreal_0(np__6169,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_61190, plain, (np__6168=np__0|np__6168!=np__3|np__6168!=np__257)).#trainneg
cnf(i_0_1759, plain, (v1_int_1(k4_xcmplx_0(np__1))|~v1_int_1(np__256))).#trainneg
cnf(i_0_61191, plain, (np__6=np__0|np__6!=np__3|np__6!=np__257)).#trainneg
cnf(i_0_61192, plain, (np__6536=np__0|np__6536!=np__3|np__6536!=np__257)).#trainneg
cnf(i_0_61193, plain, (np__56=np__0|np__56!=np__3|np__56!=np__257)).#trainneg
cnf(i_0_51326, plain, (np__6289!=np__0|~r1_xxreal_0(np__6561,np__6289))).#trainneg
cnf(i_0_51327, plain, (np__231!=np__0|~r1_xxreal_0(np__6561,np__231))).#trainneg
cnf(i_0_51329, plain, (np__79!=np__0|~r1_xxreal_0(np__6561,np__79))).#trainneg
cnf(i_0_85223, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__1,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__1,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9846, plain, (r1_xxreal_0(X1,np__176)|r1_xxreal_0(np__176,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_61194, plain, (np__23=np__0|np__23!=np__3|np__23!=np__257)).#trainneg
cnf(i_0_2006, plain, (k3_xcmplx_0(np__81,k3_xcmplx_0(np__81,X1))=k3_xcmplx_0(np__6561,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_61639, plain, (np__6289=np__0|np__6289!=np__3|np__6289!=np__257)).#trainneg
cnf(i_0_61640, plain, (np__231=np__0|np__231!=np__3|np__231!=np__257)).#trainneg
cnf(i_0_61641, plain, (np__79=np__0|np__79!=np__3|np__79!=np__257)).#trainneg
cnf(i_0_79653, plain, (~r1_xxreal_0(np__257,k4_nat_d(np__1,np__257))|~v7_ordinal1(k2_xcmplx_0(np__257,np__1)))).#trainneg
cnf(i_0_44625, plain, (k2_xcmplx_0(np__257,np__136)!=np__136|~r1_xxreal_0(np__257,k2_xcmplx_0(np__257,np__136)))).#trainneg
cnf(i_0_44626, plain, (k2_xcmplx_0(np__257,np__136)!=np__3|~r1_xxreal_0(np__257,k2_xcmplx_0(np__257,np__136)))).#trainneg
cnf(i_0_65140, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__81,X1),k4_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__81,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64221, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__8,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__8,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_83043, plain, (k4_nat_d(k3_xcmplx_0(np__2,k2_xcmplx_0(X1,k2_nat_d(X2,X1))),X1)=k4_nat_d(k3_xcmplx_0(X2,np__2),X1)|r1_xxreal_0(X1,k2_nat_d(X2,X1))|~v7_ordinal1(k2_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_65794, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__257,X1),k4_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__257,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_65834, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__257,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__257,X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9852, plain, (r1_xxreal_0(X1,np__132)|r1_xxreal_0(np__132,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1950, plain, (k3_xcmplx_0(np__81,k7_xcmplx_0(np__81,X1))=k7_xcmplx_0(np__6561,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_88772, plain, (~r1_xxreal_0(np__257,k2_nat_d(np__1,np__257))|~v7_ordinal1(k2_xcmplx_0(np__257,np__1)))).#trainneg
cnf(i_0_44627, plain, (k2_xcmplx_0(np__257,np__136)!=np__81|~r1_xxreal_0(np__257,k2_xcmplx_0(np__257,np__136)))).#trainneg
cnf(i_0_44628, plain, (k2_xcmplx_0(np__257,np__136)!=np__0|~r1_xxreal_0(np__257,k2_xcmplx_0(np__257,np__136)))).#trainneg
cnf(i_0_44629, plain, (k2_xcmplx_0(np__257,np__136)!=np__4|~r1_xxreal_0(np__257,k2_xcmplx_0(np__257,np__136)))).#trainneg
cnf(i_0_87401, plain, (k4_nat_d(k4_nat_1(np__1,np__136),np__257)=np__136)).#trainneg
cnf(i_0_88645, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__1),X2)=k4_nat_d(k3_xcmplx_0(np__1,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_90407, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__257,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__257,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_85470, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__8),X1)=k4_nat_d(k3_xcmplx_0(np__8,k2_xcmplx_0(X1,X2)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_65709, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__1,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__1,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_65690, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__1,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__1),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_71884, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__1,X1),k2_nat_d(np__3,X1)),X1)=k4_nat_d(np__3,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_64222, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__8,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__8),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_9858, plain, (r1_xxreal_0(X1,np__121)|r1_xxreal_0(np__121,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18074, plain, (k2_xcmplx_0(np__257,np__136)=np__136|r1_xxreal_0(np__257,k2_xcmplx_0(np__257,np__136)))).#trainneg
cnf(i_0_51298, plain, (k2_xcmplx_0(np__257,np__136)!=np__0|~r1_xxreal_0(np__6561,k2_xcmplx_0(np__257,np__136)))).#trainneg
cnf(i_0_80070, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__81,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__81,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_65679, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__1,np__257),X1),np__257)=k4_nat_d(k3_xcmplx_0(np__1,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_64478, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__3,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(np__3,k2_xcmplx_0(X1,X2)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_56098, plain, (k4_nat_d(k3_xcmplx_0(np__6561,np__1),np__6425)=k4_nat_d(np__6561,np__6425))).#trainneg
cnf(i_0_91140, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__257),X2)=k4_nat_d(k3_xcmplx_0(np__257,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_91279, plain, (k4_nat_d(k3_xcmplx_0(np__8,k2_xcmplx_0(X1,k2_nat_d(X2,X1))),X1)=k4_nat_d(k3_xcmplx_0(X2,np__8),X1)|r1_xxreal_0(X1,k2_nat_d(X2,X1))|~v7_ordinal1(k2_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_92953, plain, (k4_nat_d(k3_xcmplx_0(np__3,k2_xcmplx_0(X1,k2_nat_d(X2,X1))),X1)=k4_nat_d(k3_xcmplx_0(X2,np__3),X1)|r1_xxreal_0(X1,k2_nat_d(X2,X1))|~v7_ordinal1(k2_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_9864, plain, (r1_xxreal_0(X1,np__111)|r1_xxreal_0(np__111,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50065, plain, (np__6480=np__0|np__6480!=np__4|np__6480!=np__257)).#trainneg
cnf(i_0_50068, plain, (np__254=np__0|np__254!=np__4|np__254!=np__257)).#trainneg
cnf(i_0_50071, plain, (np__253=np__0|np__253!=np__4|np__253!=np__257)).#trainneg
cnf(i_0_50074, plain, (np__255=np__0|np__255!=np__4|np__255!=np__257)).#trainneg
cnf(i_0_72899, plain, (~r1_xxreal_0(np__257,k2_nat_d(k2_xcmplx_0(np__1,np__257),np__257))|~v7_ordinal1(k2_xcmplx_0(np__1,np__257)))).#trainneg
cnf(i_0_42182, negated_conjecture, (k1_newton(np__3,np__8)!=np__0|np__136!=np__0|~v7_ordinal1(k1_newton(np__3,np__8)))).#trainneg
cnf(i_0_92605, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__81),X2)=k4_nat_d(k3_xcmplx_0(np__81,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_92653, plain, (k2_nat_d(k4_nat_1(k2_nat_d(np__1,np__257),X1),np__257)=k4_nat_d(k3_xcmplx_0(np__1,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k4_nat_1(k2_nat_d(np__1,np__257),X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_91550, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__1,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(np__1,k2_xcmplx_0(X2,X1)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_90754, plain, (k2_nat_d(k4_nat_1(np__136,np__1),np__257)=np__136|~v7_ordinal1(k4_nat_1(np__136,np__1)))).#trainneg
cnf(i_0_91548, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__1,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(np__1,k2_xcmplx_0(X1,X2)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64582, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__4,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__4,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_9870, plain, (r1_xxreal_0(X1,np__6344)|r1_xxreal_0(np__6344,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50077, plain, (np__120=np__0|np__120!=np__4|np__120!=np__257)).#trainneg
cnf(i_0_1986, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,k4_xcmplx_0(np__1)))=k4_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_50080, plain, (np__55=np__0|np__55!=np__4|np__55!=np__257)).#trainneg
cnf(i_0_50083, plain, (np__6305=np__0|np__6305!=np__4|np__6305!=np__257)).#trainneg
cnf(i_0_50086, plain, (np__7=np__0|np__7!=np__4|np__7!=np__257)).#trainneg
cnf(i_0_52850, negated_conjecture, (k1_newton(np__3,np__8)!=np__4|np__136!=np__4|~v7_ordinal1(k1_newton(np__3,np__8)))).#trainneg
cnf(i_0_64171, plain, (v1_xreal_0(np__6168)|~v1_xcmplx_0(np__6168))).#trainneg
cnf(i_0_87270, plain, (v1_int_1(np__6168)|~v1_xcmplx_0(np__6168))).#trainneg
cnf(i_0_79678, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__1,np__257),X1),np__257)=k4_nat_d(X1,np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9876, plain, (r1_xxreal_0(X1,np__6400)|r1_xxreal_0(np__6400,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50089, plain, (np__5=np__0|np__5!=np__4|np__5!=np__257)).#trainneg
cnf(i_0_1952, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(k4_xcmplx_0(np__1),X2))=k7_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_94030, negated_conjecture, (k1_newton(np__3,np__8)!=np__0|np__136!=np__0)).#trainneg
cnf(i_0_95465, negated_conjecture, (k1_newton(np__3,np__8)!=np__4|np__136!=np__4)).#trainneg
cnf(i_0_50092, plain, (np__252=np__0|np__252!=np__4|np__252!=np__257)).#trainneg
cnf(i_0_95596, plain, (~r1_xxreal_0(np__257,k4_nat_d(np__1,np__257))|~v7_ordinal1(k2_xcmplx_0(np__1,np__257)))).#trainneg
cnf(i_0_45535, negated_conjecture, (k13_newton(np__3,np__8)!=np__0|np__136!=np__0|~m1_subset_1(k13_newton(np__3,np__8),k5_numbers))).#trainneg
cnf(i_0_55324, negated_conjecture, (k13_newton(np__3,np__8)!=np__4|np__136!=np__4|~m1_subset_1(k13_newton(np__3,np__8),k5_numbers))).#trainneg
cnf(i_0_95127, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__4,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(np__4,k2_xcmplx_0(X1,X2)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_94549, plain, (k4_nat_d(k3_xcmplx_0(np__1,k2_xcmplx_0(k2_nat_d(X2,X1),X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__1),X1)|r1_xxreal_0(X1,k2_nat_d(X2,X1))|~v7_ordinal1(k2_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_93297, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__257),X2)=k4_nat_d(k3_xcmplx_0(np__257,k2_xcmplx_0(X2,X1)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_94896, plain, (k4_nat_d(k3_xcmplx_0(np__1,k2_xcmplx_0(X1,k2_nat_d(X2,X1))),X1)=k4_nat_d(k3_xcmplx_0(X2,np__1),X1)|r1_xxreal_0(X1,k2_nat_d(X2,X1))|~v7_ordinal1(k2_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_22394, plain, (k2_xcmplx_0(X1,np__136)=k2_xcmplx_0(np__136,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_51692, plain, (k2_nat_d(k4_nat_1(X3,X1),X2)=k2_nat_d(k4_nat_1(X1,X3),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k4_nat_1(X1,X3))|~v7_ordinal1(k4_nat_1(X3,X1))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_10079, plain, (r1_xxreal_0(X1,np__232)|r1_xxreal_0(np__232,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50095, plain, (np__6558=np__0|np__6558!=np__4|np__6558!=np__257)).#trainneg
cnf(i_0_1932, plain, (v1_abian(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_50098, plain, (np__6424=np__0|np__6424!=np__4|np__6424!=np__257)).#trainneg
cnf(i_0_50101, plain, (np__6304=np__0|np__6304!=np__4|np__6304!=np__257)).#trainneg
cnf(i_0_50104, plain, (np__80=np__0|np__80!=np__4|np__80!=np__257)).#trainneg
cnf(i_0_7748, plain, (r1_xxreal_0(np__256,np__256))).#trainneg
cnf(i_0_94223, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__81),X2)=k4_nat_d(k3_xcmplx_0(np__81,k2_xcmplx_0(X2,X1)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_34943, plain, (k2_nat_d(k3_xcmplx_0(X3,k4_nat_d(X1,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X3,X2),X1),X2)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k3_xcmplx_0(X3,k4_nat_d(X1,X2)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k4_nat_d(X1,X2),k5_numbers)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_15225, plain, (k4_nat_d(k3_xcmplx_0(k4_nat_d(X1,np__257),X2),np__257)=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|r1_xxreal_0(np__257,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X1,np__257),k5_numbers))).#trainneg
cnf(i_0_93091, plain, (k4_nat_d(k3_xcmplx_0(np__1,np__6561),np__6425)=np__136)).#trainneg
cnf(i_0_10085, plain, (r1_xxreal_0(X1,np__134)|r1_xxreal_0(np__134,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50283, plain, (np__134=np__0|np__134!=np__4|np__134!=np__257)).#trainneg
cnf(i_0_2237, plain, (k4_nat_d(k3_xcmplx_0(k4_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(k4_nat_d(X1,X2))|~v7_ordinal1(k4_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k4_nat_d(X1,X2),k5_numbers))).#trainneg
cnf(i_0_95818, negated_conjecture, (k13_newton(np__3,np__8)!=np__0|np__136!=np__0)).#trainneg
cnf(i_0_95821, negated_conjecture, (k13_newton(np__3,np__8)!=np__4|np__136!=np__4)).#trainneg
cnf(i_0_50284, plain, (np__232=np__0|np__232!=np__4|np__232!=np__257)).#trainneg
cnf(i_0_95812, plain, (~r1_xxreal_0(np__257,k2_nat_d(np__1,np__257))|~v7_ordinal1(k2_xcmplx_0(np__1,np__257)))).#trainneg
cnf(i_0_7822, plain, (r1_xxreal_0(np__6425,np__6425))).#trainneg
cnf(i_0_7938, plain, (r1_xxreal_0(np__25,np__25))).#trainneg
cnf(i_0_7942, plain, (r1_xxreal_0(np__6,np__6))).#trainneg
cnf(i_0_96989, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X2,np__136)),X2)=k4_nat_d(k4_nat_1(X1,np__136),X2)|r1_xxreal_0(X2,np__136)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_97120, plain, (k2_nat_d(X1,X2)=np__136|r1_xxreal_0(X2,np__136)|X1!=k2_xcmplx_0(np__136,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_96242, plain, (k4_nat_d(k3_xcmplx_0(np__1,k2_xcmplx_0(np__257,np__136)),np__257)=np__136)).#trainneg
cnf(i_0_96973, plain, (X1=np__0|r1_xxreal_0(X1,np__136)|~r1_xxreal_0(X1,k2_nat_d(k2_xcmplx_0(X1,np__136),X1))|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_96994, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X2,np__136)),X2)=k2_nat_d(k3_xcmplx_0(X1,np__136),X2)|r1_xxreal_0(X2,np__136)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_97073, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__136,X1),k2_xcmplx_0(X2,X1)),X1)=k4_nat_d(k4_nat_1(X2,np__136),X1)|r1_xxreal_0(X1,np__136)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10091, plain, (r1_xxreal_0(X1,np__6424)|r1_xxreal_0(np__6424,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50285, plain, (np__6400=np__0|np__6400!=np__4|np__6400!=np__257)).#trainneg
cnf(i_0_2265, plain, (k4_nat_d(k3_xcmplx_0(k4_nat_d(X1,X2),k4_nat_d(X3,X2)),X2)=k4_nat_d(k3_xcmplx_0(X3,X1),X2)|~v7_ordinal1(k4_nat_d(X3,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~m1_subset_1(k4_nat_d(X1,X2),k5_numbers))).#trainneg
cnf(i_0_50286, plain, (np__6557=np__0|np__6557!=np__4|np__6557!=np__257)).#trainneg
cnf(i_0_50287, plain, (np__6344=np__0|np__6344!=np__4|np__6344!=np__257)).#trainneg
cnf(i_0_50611, plain, (np__22=np__0|np__22!=np__4|np__22!=np__257)).#trainneg
cnf(i_0_8255, plain, (r1_xxreal_0(np__21,np__21))).#trainneg
cnf(i_0_8476, plain, (r1_xxreal_0(np__5,np__5))).#trainneg
cnf(i_0_64616, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__4,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__4,X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_21361, plain, (k2_nat_d(k2_xcmplx_0(X1,np__257),np__1)=X1|r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_36290, plain, (r1_xxreal_0(np__0,np__25))).#trainneg
cnf(i_0_90968, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__1),X2)=k4_nat_d(k3_xcmplx_0(np__1,k2_xcmplx_0(X2,X1)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_100788, plain, (k3_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_100787, plain, (k3_xcmplx_0(np__3,np__1)=np__3)).#trainneg
cnf(i_0_10097, plain, (r1_xxreal_0(X1,np__252)|r1_xxreal_0(np__252,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50612, plain, (np__111=np__0|np__111!=np__4|np__111!=np__257)).#trainneg
cnf(i_0_50613, plain, (np__121=np__0|np__121!=np__4|np__121!=np__257)).#trainneg
cnf(i_0_50614, plain, (np__132=np__0|np__132!=np__4|np__132!=np__257)).#trainneg
cnf(i_0_50615, plain, (np__133=np__0|np__133!=np__4|np__133!=np__257)).#trainneg
cnf(i_0_8480, plain, (r1_xxreal_0(np__22,np__22))).#trainneg
cnf(i_0_8484, plain, (r1_xxreal_0(np__7,np__7))).#trainneg
cnf(i_0_8830, plain, (r1_xxreal_0(np__255,np__255))).#trainneg
cnf(i_0_8834, plain, (r1_xxreal_0(np__6560,np__6560))).#trainneg
cnf(i_0_8838, plain, (r1_xxreal_0(np__6305,np__6305))).#trainneg
cnf(i_0_100783, plain, (k3_xcmplx_0(np__256,np__1)=np__256)).#trainneg
cnf(i_0_100781, plain, (k3_xcmplx_0(np__6,np__1)=np__6)).#trainneg
cnf(i_0_100794, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),np__1)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_64563, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__4,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__4),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9753, plain, (k2_nat_d(X1,np__2)=np__1|X1!=np__3|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10103, plain, (r1_xxreal_0(X1,np__55)|r1_xxreal_0(np__55,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_51501, plain, (np__175=np__0|np__175!=np__4|np__175!=np__257)).#trainneg
cnf(i_0_2280, plain, (k2_nat_d(k2_xcmplx_0(np__6425,X1),np__25)=X1|r1_xxreal_0(np__25,X1)|~v7_ordinal1(k2_xcmplx_0(np__6425,X1))|~v7_ordinal1(np__25)|~v7_ordinal1(np__257)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_51502, plain, (np__176=np__0|np__176!=np__4|np__176!=np__257)).#trainneg
cnf(i_0_51503, plain, (np__21=np__0|np__21!=np__4|np__21!=np__257)).#trainneg
cnf(i_0_51504, plain, (np__6169=np__0|np__6169!=np__4|np__6169!=np__257)).#trainneg
cnf(i_0_8990, plain, (r1_xxreal_0(np__6536,np__6536))).#trainneg
cnf(i_0_8994, plain, (r1_xxreal_0(np__6557,np__6557))).#trainneg
cnf(i_0_8998, plain, (r1_xxreal_0(np__6558,np__6558))).#trainneg
cnf(i_0_9002, plain, (r1_xxreal_0(np__6480,np__6480))).#trainneg
cnf(i_0_9193, plain, (r1_xxreal_0(np__77,np__77))).#trainneg
cnf(i_0_87527, plain, (k4_nat_d(k3_xcmplx_0(np__257,np__8),np__257)=np__0)).#trainneg
cnf(i_0_65811, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__257,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__257),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64451, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__3,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__3),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13853, plain, (k2_nat_d(k2_xcmplx_0(np__136,X1),np__1)=X1|r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(np__136,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10109, plain, (r1_xxreal_0(X1,np__120)|r1_xxreal_0(np__120,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_51505, plain, (np__6421=np__0|np__6421!=np__4|np__6421!=np__257)).#trainneg
cnf(i_0_2284, plain, (k2_nat_d(k2_xcmplx_0(np__24,X1),np__6)=X1|r1_xxreal_0(np__6,X1)|~v7_ordinal1(k2_xcmplx_0(np__24,X1))|~v7_ordinal1(np__6)|~v7_ordinal1(np__4)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_53211, plain, (np__6422=np__0|np__6422!=np__4|np__6422!=np__257)).#trainneg
cnf(i_0_53212, plain, (np__6423=np__0|np__6423!=np__4|np__6423!=np__257)).#trainneg
cnf(i_0_53213, plain, (np__6559=np__0|np__6559!=np__4|np__6559!=np__257)).#trainneg
cnf(i_0_9197, plain, (r1_xxreal_0(np__6304,np__6304))).#trainneg
cnf(i_0_9201, plain, (r1_xxreal_0(np__80,np__80))).#trainneg
cnf(i_0_9205, plain, (r1_xxreal_0(np__79,np__79))).#trainneg
cnf(i_0_9209, plain, (r1_xxreal_0(np__175,np__175))).#trainneg
cnf(i_0_9213, plain, (r1_xxreal_0(np__254,np__254))).#trainneg
cnf(i_0_85794, plain, (X1=np__0|r1_xxreal_0(X1,k2_nat_d(X2,X1))|~r1_xxreal_0(X1,k4_nat_d(k3_xcmplx_0(X2,np__8),X1))|~v7_ordinal1(k3_xcmplx_0(k2_xcmplx_0(k2_nat_d(X2,X1),X1),np__8))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_105150, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(k3_xcmplx_0(X2,np__8),X1))|~v7_ordinal1(k2_xcmplx_0(k2_nat_d(X2,X1),X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_102639, plain, (k3_xcmplx_0(np__136,np__1)=np__136)).#trainneg
cnf(i_0_105149, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(k3_xcmplx_0(X2,np__8),X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k3_xcmplx_0(k2_xcmplx_0(k2_nat_d(X2,X1),X1),np__8),k5_numbers))).#trainneg
cnf(i_0_102782, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(np__256,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_102633, plain, (k3_xcmplx_0(np__6425,np__1)=np__6425)).#trainneg
cnf(i_0_103259, plain, (np__0=np__1|np__8!=np__3)).#trainneg
cnf(i_0_103137, plain, (np__0=np__1|np__6!=np__3)).#trainneg
cnf(i_0_2286, plain, (k2_nat_d(k2_xcmplx_0(np__0,X1),np__6561)=X1|r1_xxreal_0(np__6561,X1)|~v7_ordinal1(k2_xcmplx_0(np__0,X1))|~v7_ordinal1(np__6561)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_103402, plain, (np__0=np__1|np__3!=np__2)).#trainneg
cnf(i_0_103237, plain, (np__0=np__1|np__4!=np__3)).#trainneg
cnf(i_0_53214, plain, (np__6560=np__0|np__6560!=np__4|np__6560!=np__257)).#trainneg
cnf(i_0_105049, plain, (~r1_xxreal_0(np__257,k4_nat_d(k3_xcmplx_0(np__6425,np__8),np__257))|~v7_ordinal1(k3_xcmplx_0(np__257,np__8)))).#trainneg
cnf(i_0_105050, plain, (~r1_xxreal_0(np__6425,k4_nat_d(k3_xcmplx_0(np__6561,np__8),np__6425))|~v7_ordinal1(k3_xcmplx_0(np__6561,np__8)))).#trainneg
cnf(i_0_105051, plain, (~r1_xxreal_0(np__6561,k4_nat_d(k3_xcmplx_0(np__6561,np__8),np__6561))|~v7_ordinal1(k3_xcmplx_0(np__6561,np__8)))).#trainneg
cnf(i_0_105052, plain, (~r1_xxreal_0(np__257,k4_nat_d(k3_xcmplx_0(np__6561,np__8),np__257))|~v7_ordinal1(k3_xcmplx_0(k2_xcmplx_0(np__257,np__136),np__8)))).#trainneg
cnf(i_0_105090, plain, (~r1_xxreal_0(np__257,k4_nat_d(k3_xcmplx_0(np__81,np__8),np__257))|~v7_ordinal1(k3_xcmplx_0(k2_xcmplx_0(np__81,np__257),np__8)))).#trainneg
cnf(i_0_102608, plain, (k3_xcmplx_0(np__5,np__1)=np__5)).#trainneg
cnf(i_0_102607, plain, (k3_xcmplx_0(np__7,np__1)=np__7)).#trainneg
cnf(i_0_102818, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(np__6,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_102790, plain, (k3_xcmplx_0(np__1,np__6)=np__6)).#trainneg
cnf(i_0_10115, plain, (r1_xxreal_0(X1,np__253)|r1_xxreal_0(np__253,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_53215, plain, (np__77=np__0|np__77!=np__4|np__77!=np__257)).#trainneg
cnf(i_0_54145, plain, (np__78=np__0|np__78!=np__4|np__78!=np__257)).#trainneg
cnf(i_0_54146, plain, (np__24=np__0|np__24!=np__4|np__24!=np__257)).#trainneg
cnf(i_0_54147, plain, (np__135=np__0|np__135!=np__4|np__135!=np__257)).#trainneg
cnf(i_0_105093, plain, (~r1_xxreal_0(np__257,k4_nat_d(k3_xcmplx_0(np__4,np__8),np__257))|~v7_ordinal1(k3_xcmplx_0(k2_xcmplx_0(np__4,np__257),np__8)))).#trainneg
cnf(i_0_105097, plain, (~r1_xxreal_0(np__257,k4_nat_d(k3_xcmplx_0(np__136,np__8),np__257))|~v7_ordinal1(k3_xcmplx_0(k2_xcmplx_0(np__257,np__136),np__8)))).#trainneg
cnf(i_0_105105, plain, (~r1_xxreal_0(np__257,k4_nat_d(k3_xcmplx_0(np__3,np__8),np__257))|~v7_ordinal1(k3_xcmplx_0(k2_xcmplx_0(np__3,np__257),np__8)))).#trainneg
cnf(i_0_9217, plain, (r1_xxreal_0(np__78,np__78))).#trainneg
cnf(i_0_9221, plain, (r1_xxreal_0(np__24,np__24))).#trainneg
cnf(i_0_102754, plain, (k3_xcmplx_0(np__1,np__256)=np__256)).#trainneg
cnf(i_0_102895, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(np__6561,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10121, plain, (r1_xxreal_0(X1,np__6168)|r1_xxreal_0(np__6168,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_54148, plain, (np__6168=np__0|np__6168!=np__4|np__6168!=np__257)).#trainneg
cnf(i_0_54149, plain, (np__6=np__0|np__6!=np__4|np__6!=np__257)).#trainneg
cnf(i_0_55339, plain, (np__6536=np__0|np__6536!=np__4|np__6536!=np__257)).#trainneg
cnf(i_0_55340, plain, (np__56=np__0|np__56!=np__4|np__56!=np__257)).#trainneg
cnf(i_0_9444, plain, (r1_xxreal_0(np__23,np__23))).#trainneg
cnf(i_0_9448, plain, (r1_xxreal_0(np__6559,np__6559))).#trainneg
cnf(i_0_9452, plain, (r1_xxreal_0(np__231,np__231))).#trainneg
cnf(i_0_9456, plain, (r1_xxreal_0(np__6289,np__6289))).#trainneg
cnf(i_0_9460, plain, (r1_xxreal_0(np__56,np__56))).#trainneg
cnf(i_0_104553, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__3),X2)=k4_nat_d(k3_xcmplx_0(np__3,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_103603, plain, (k2_nat_d(np__6425,np__25)=np__0)).#trainneg
cnf(i_0_102781, plain, (k4_nat_d(np__256,X1)=k2_nat_d(np__256,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_104555, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__3),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__3),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_104517, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(k2_nat_d(X1,X2),X2),np__3),X2)=k4_nat_d(k3_xcmplx_0(X1,np__3),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_103069, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__4),X2)=k4_nat_d(k3_xcmplx_0(np__4,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_105635, plain, (k2_nat_d(np__257,np__6561)=np__257)).#trainneg
cnf(i_0_105638, plain, (k2_nat_d(np__136,np__6561)=np__136)).#trainneg
cnf(i_0_105705, plain, (np__8=np__1|np__3!=np__2)).#trainneg
cnf(i_0_2403, plain, (m2_subset_1(k4_pepin(X1),k1_numbers,k5_numbers)|~v7_ordinal1(k2_newton(np__2,X1))|~v7_ordinal1(np__1)|~v7_ordinal1(X1)|~m1_subset_1(k1_newton(np__2,k2_newton(np__2,X1)),k5_numbers))).#trainneg
cnf(i_0_105666, plain, (np__3!=np__2)).#trainneg
cnf(i_0_105772, plain, (np__4!=np__3)).#trainneg
cnf(i_0_105639, plain, (k2_nat_d(np__8,np__6561)=np__8)).#trainneg
cnf(i_0_105640, plain, (k2_nat_d(np__3,np__6561)=np__3)).#trainneg
cnf(i_0_105440, plain, (X1=np__0|~r1_xxreal_0(X1,k2_nat_d(np__256,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_105637, plain, (k2_nat_d(np__4,np__6561)=np__4)).#trainneg
cnf(i_0_105437, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(k3_xcmplx_0(X2,np__8),X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k2_xcmplx_0(k2_nat_d(X2,X1),X1),k5_numbers))).#trainneg
cnf(i_0_106630, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__256,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__256),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_105964, plain, (X1=np__0|~r1_xxreal_0(X1,k2_nat_d(np__6,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_105452, plain, (k4_nat_d(k3_xcmplx_0(np__1,np__6425),X1)=k4_nat_d(np__6425,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10383, plain, (r1_xxreal_0(X1,np__133)|r1_xxreal_0(np__133,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_105636, plain, (k2_nat_d(np__81,np__6561)=np__81)).#trainneg
cnf(i_0_108927, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(np__6425,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_107690, plain, (X1=np__0|X1!=np__136|~r1_xxreal_0(np__6561,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_107547, plain, (X1=np__0|X1!=np__257|~r1_xxreal_0(np__6561,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_107982, plain, (X1=np__0|X1!=np__8|~r1_xxreal_0(np__6561,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_105934, plain, (k3_xcmplx_0(np__1,np__7)=np__7)).#trainneg
cnf(i_0_107508, plain, (k2_xcmplx_0(np__6561,np__136)!=np__257)).#trainneg
cnf(i_0_107544, plain, (k2_xcmplx_0(np__136,np__6561)!=np__257)).#trainneg
cnf(i_0_106625, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(np__256,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__256),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_105926, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(np__5,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_105641, plain, (k2_nat_d(X1,np__6561)=X1|r1_xxreal_0(np__6561,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_105962, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(np__7,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_105898, plain, (k3_xcmplx_0(np__1,np__5)=np__5)).#trainneg
cnf(i_0_107268, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,k2_nat_d(X1,X2)),np__3),X2)=k4_nat_d(k3_xcmplx_0(X1,np__3),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_106568, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__256,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__256),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_106581, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__256,X1),k4_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__256,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_105988, plain, (k4_nat_d(np__6,X1)=k2_nat_d(np__6,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10400, plain, (r1_xxreal_0(X1,np__6425)|r1_xxreal_0(np__6425,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_55341, plain, (np__23=np__0|np__23!=np__4|np__23!=np__257)).#trainneg
cnf(i_0_106593, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__256,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__256),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_106592, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__256,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__256,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_109143, plain, (X1=np__0|~r1_xxreal_0(X1,k2_nat_d(np__6425,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_109277, plain, (X1!=np__257|~r1_xxreal_0(np__6561,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_109343, plain, (X1!=np__8|~r1_xxreal_0(np__6561,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_109211, plain, (X1!=np__136|~r1_xxreal_0(np__6561,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_109867, plain, (X1=np__0|~r1_xxreal_0(X1,k2_nat_d(np__5,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_110185, plain, (X1=np__0|~r1_xxreal_0(X1,k2_nat_d(np__7,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_109203, plain, (np__136!=np__6561)).#trainneg
cnf(i_0_55342, plain, (np__6289=np__0|np__6289!=np__4|np__6289!=np__257)).#trainneg
cnf(i_0_109335, plain, (np__8!=np__6561)).#trainneg
cnf(i_0_109144, plain, (~r1_xxreal_0(np__6425,k2_nat_d(np__6425,np__6425)))).#trainneg
cnf(i_0_109148, plain, (np__6480!=np__136|~r1_xxreal_0(np__6561,np__6480))).#trainneg
cnf(i_0_109149, plain, (np__254!=np__136|~r1_xxreal_0(np__6561,np__254))).#trainneg
cnf(i_0_109150, plain, (np__253!=np__136|~r1_xxreal_0(np__6561,np__253))).#trainneg
cnf(i_0_109151, plain, (np__255!=np__136|~r1_xxreal_0(np__6561,np__255))).#trainneg
cnf(i_0_110208, plain, (k4_nat_d(np__5,X1)=k2_nat_d(np__5,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_109364, plain, (k4_nat_d(np__7,X1)=k2_nat_d(np__7,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_111199, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__6,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__6),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_111262, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__6,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__6),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_109060, plain, (X1=np__0|X1!=np__81|~r1_xxreal_0(np__6561,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_106590, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__256,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__256),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_106587, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__256,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__256,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_106631, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__256,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__256,X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_105450, plain, (k3_xcmplx_0(np__1,np__6425)=np__6425)).#trainneg
cnf(i_0_55343, plain, (np__231=np__0|np__231!=np__4|np__231!=np__257)).#trainneg
cnf(i_0_2481, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k1_nat_1(X1,X2),k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_56478, plain, (np__79=np__0|np__79!=np__4|np__79!=np__257)).#trainneg
cnf(i_0_109152, plain, (np__120!=np__136|~r1_xxreal_0(np__6561,np__120))).#trainneg
cnf(i_0_109153, plain, (np__55!=np__136|~r1_xxreal_0(np__6561,np__55))).#trainneg
cnf(i_0_109154, plain, (np__6305!=np__136|~r1_xxreal_0(np__6561,np__6305))).#trainneg
cnf(i_0_109155, plain, (np__7!=np__136|~r1_xxreal_0(np__6561,np__7))).#trainneg
cnf(i_0_109157, plain, (np__5!=np__136|~r1_xxreal_0(np__6561,np__5))).#trainneg
cnf(i_0_106638, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(X1,X2),k2_nat_d(np__256,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__256,X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_106597, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__256,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__256),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_108300, plain, (X1=np__0|X1!=np__4|~r1_xxreal_0(np__6561,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_111228, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__6,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__6),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_61645, plain, (np__6480=np__0|np__6480!=np__81|np__6480!=np__257)).#trainneg
cnf(i_0_2482, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k1_nat_1(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_109158, plain, (np__252!=np__136|~r1_xxreal_0(np__6561,np__252))).#trainneg
cnf(i_0_109160, plain, (np__6558!=np__136|~r1_xxreal_0(np__6561,np__6558))).#trainneg
cnf(i_0_109161, plain, (np__6424!=np__136|~r1_xxreal_0(np__6561,np__6424))).#trainneg
cnf(i_0_112172, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__5,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__5),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_112286, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__7,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__7),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_112235, plain, (m1_subset_1(k2_nat_d(np__7,X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_112121, plain, (m1_subset_1(k2_nat_d(np__5,X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_112109, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__5,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__5),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_108132, plain, (X1=np__0|X1!=np__3|~r1_xxreal_0(np__6561,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_106579, plain, (k4_nat_d(k3_xcmplx_0(X1,k2_nat_d(np__256,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__256),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_112223, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__7,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__7),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_105164, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(k3_xcmplx_0(np__8,X2),X1))|~v7_ordinal1(k2_xcmplx_0(k2_nat_d(X2,X1),X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10489, plain, (r1_xxreal_0(X1,np__257)|r1_xxreal_0(np__257,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_62079, plain, (np__254=np__0|np__254!=np__81|np__254!=np__257)).#trainneg
cnf(i_0_2495, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k2_nat_1(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_62080, plain, (np__253=np__0|np__253!=np__81|np__253!=np__257)).#trainneg
cnf(i_0_62081, plain, (np__255=np__0|np__255!=np__81|np__255!=np__257)).#trainneg
cnf(i_0_62082, plain, (np__120=np__0|np__120!=np__81|np__120!=np__257)).#trainneg
cnf(i_0_109162, plain, (np__6304!=np__136|~r1_xxreal_0(np__6561,np__6304))).#trainneg
cnf(i_0_109163, plain, (np__80!=np__136|~r1_xxreal_0(np__6561,np__80))).#trainneg
cnf(i_0_109164, plain, (np__134!=np__136|~r1_xxreal_0(np__6561,np__134))).#trainneg
cnf(i_0_109165, plain, (np__232!=np__136|~r1_xxreal_0(np__6561,np__232))).#trainneg
cnf(i_0_109166, plain, (np__6400!=np__136|~r1_xxreal_0(np__6561,np__6400))).#trainneg
cnf(i_0_105477, plain, (k4_nat_d(np__6425,X1)=k2_nat_d(np__6425,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_111211, plain, (m1_subset_1(k2_nat_d(np__6,X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_103871, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,k4_nat_d(X2,X1)),np__257),X1)=k4_nat_d(k3_xcmplx_0(X2,np__257),X1)|r1_xxreal_0(X1,k4_nat_d(X2,X1))|~v7_ordinal1(k4_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_112245, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__7,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__7),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_105621, plain, (k2_nat_d(np__0,np__6561)=np__0)).#trainneg
cnf(i_0_62083, plain, (np__55=np__0|np__55!=np__81|np__55!=np__257)).#trainneg
cnf(i_0_2496, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_62899, plain, (np__6305=np__0|np__6305!=np__81|np__6305!=np__257)).#trainneg
cnf(i_0_62900, plain, (np__7=np__0|np__7!=np__81|np__7!=np__257)).#trainneg
cnf(i_0_62901, plain, (np__5=np__0|np__5!=np__81|np__5!=np__257)).#trainneg
cnf(i_0_109167, plain, (np__6557!=np__136|~r1_xxreal_0(np__6561,np__6557))).#trainneg
cnf(i_0_109168, plain, (np__6344!=np__136|~r1_xxreal_0(np__6561,np__6344))).#trainneg
cnf(i_0_109171, plain, (np__22!=np__136|~r1_xxreal_0(np__6561,np__22))).#trainneg
cnf(i_0_109172, plain, (np__111!=np__136|~r1_xxreal_0(np__6561,np__111))).#trainneg
cnf(i_0_109173, plain, (np__121!=np__136|~r1_xxreal_0(np__6561,np__121))).#trainneg
cnf(i_0_111210, plain, (k4_nat_d(k3_xcmplx_0(X1,k2_nat_d(np__6,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__6),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_103034, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(k2_nat_d(X1,X2),X2),np__4),X2)=k4_nat_d(k3_xcmplx_0(X1,np__4),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_116384, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_115880, plain, (X1!=np__3|~r1_xxreal_0(np__6561,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_116411, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__6425,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__6425),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_116474, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__6425,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__6425),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_116440, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__6425,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__6425),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_116430, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__6425,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__6425,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_112131, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__5,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__5),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10589, plain, (r1_xxreal_0(X1,np__22)|r1_xxreal_0(np__22,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_62902, plain, (np__252=np__0|np__252!=np__81|np__252!=np__257)).#trainneg
cnf(i_0_62903, plain, (np__6558=np__0|np__6558!=np__81|np__6558!=np__257)).#trainneg
cnf(i_0_64172, plain, (np__6424=np__0|np__6424!=np__81|np__6424!=np__257)).#trainneg
cnf(i_0_64173, plain, (np__6304=np__0|np__6304!=np__81|np__6304!=np__257)).#trainneg
cnf(i_0_109174, plain, (np__132!=np__136|~r1_xxreal_0(np__6561,np__132))).#trainneg
cnf(i_0_109175, plain, (np__133!=np__136|~r1_xxreal_0(np__6561,np__133))).#trainneg
cnf(i_0_109176, plain, (np__175!=np__136|~r1_xxreal_0(np__6561,np__175))).#trainneg
cnf(i_0_109177, plain, (np__176!=np__136|~r1_xxreal_0(np__6561,np__176))).#trainneg
cnf(i_0_109178, plain, (np__21!=np__136|~r1_xxreal_0(np__6561,np__21))).#trainneg
cnf(i_0_111200, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__6,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__6,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_112167, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(np__5,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__5),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_111257, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(np__6,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__6),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_104370, plain, (k4_nat_d(k4_nat_d(X1,np__257),np__257)=k4_nat_d(k3_xcmplx_0(np__1,X1),np__257)|~v7_ordinal1(k4_nat_d(X1,np__257))|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X1,np__257),k1_numbers))).#trainneg
cnf(i_0_112128, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__5,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__5,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_107475, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__4),X2)=k4_nat_d(k3_xcmplx_0(np__4,k2_xcmplx_0(X2,X1)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_108394, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(k3_xcmplx_0(np__8,X2),X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k2_xcmplx_0(k2_nat_d(X2,X1),X1),k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_64174, plain, (np__80=np__0|np__80!=np__81|np__80!=np__257)).#trainneg
cnf(i_0_2515, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k1_newton(X1,X2),k1_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_64175, plain, (np__134=np__0|np__134!=np__81|np__134!=np__257)).#trainneg
cnf(i_0_64176, plain, (np__232=np__0|np__232!=np__81|np__232!=np__257)).#trainneg
cnf(i_0_64537, plain, (np__6400=np__0|np__6400!=np__81|np__6400!=np__257)).#trainneg
cnf(i_0_109179, plain, (np__6169!=np__136|~r1_xxreal_0(np__6561,np__6169))).#trainneg
cnf(i_0_109180, plain, (np__6421!=np__136|~r1_xxreal_0(np__6561,np__6421))).#trainneg
cnf(i_0_109181, plain, (np__6422!=np__136|~r1_xxreal_0(np__6561,np__6422))).#trainneg
cnf(i_0_109182, plain, (np__6423!=np__136|~r1_xxreal_0(np__6561,np__6423))).#trainneg
cnf(i_0_109183, plain, (np__6559!=np__136|~r1_xxreal_0(np__6561,np__6559))).#trainneg
cnf(i_0_103071, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__4),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__4),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_121286, plain, (k4_nat_d(np__6425,np__257)=np__0)).#trainneg
cnf(i_0_106406, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__3),X2)=k4_nat_d(k3_xcmplx_0(np__3,k2_xcmplx_0(X2,X1)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_112120, plain, (k4_nat_d(k3_xcmplx_0(X1,k2_nat_d(np__5,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__5),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_106555, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__256,np__257),X1),np__257)=k4_nat_d(k3_xcmplx_0(np__256,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_112242, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__7,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__7,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_116423, plain, (m1_subset_1(k2_nat_d(np__6425,X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_106605, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__256,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__256),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10593, plain, (r1_xxreal_0(X1,np__23)|r1_xxreal_0(np__23,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_64538, plain, (np__6557=np__0|np__6557!=np__81|np__6557!=np__257)).#trainneg
cnf(i_0_2525, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k4_nat_1(X1,X2),k1_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_64539, plain, (np__6344=np__0|np__6344!=np__81|np__6344!=np__257)).#trainneg
cnf(i_0_64540, plain, (np__22=np__0|np__22!=np__81|np__22!=np__257)).#trainneg
cnf(i_0_64541, plain, (np__111=np__0|np__111!=np__81|np__111!=np__257)).#trainneg
cnf(i_0_109185, plain, (np__77!=np__136|~r1_xxreal_0(np__6561,np__77))).#trainneg
cnf(i_0_109186, plain, (np__78!=np__136|~r1_xxreal_0(np__6561,np__78))).#trainneg
cnf(i_0_109187, plain, (np__24!=np__136|~r1_xxreal_0(np__6561,np__24))).#trainneg
cnf(i_0_109191, plain, (np__135!=np__136|~r1_xxreal_0(np__6561,np__135))).#trainneg
cnf(i_0_109192, plain, (np__6168!=np__136|~r1_xxreal_0(np__6561,np__6168))).#trainneg
cnf(i_0_106604, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__256,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__256),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_112281, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(np__7,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__7),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_116435, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__6425,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__6425,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_106599, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__256,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__256),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_116469, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,X2),k2_nat_d(np__6425,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__6425),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_123716, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__256),X1)=k4_nat_d(k3_xcmplx_0(np__256,k2_xcmplx_0(X2,X1)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_118101, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,k2_nat_d(X1,X2)),np__4),X2)=k4_nat_d(k3_xcmplx_0(X1,np__4),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X1,X2),k1_numbers))).#trainneg
cnf(i_0_113998, plain, (X1=k2_nat_d(X2,np__1)|r1_xxreal_0(np__1,X1)|X2!=k2_xcmplx_0(X1,np__6425)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_65101, plain, (np__121=np__0|np__121!=np__81|np__121!=np__257)).#trainneg
cnf(i_0_2526, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k4_nat_1(X1,X2),k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_65102, plain, (np__132=np__0|np__132!=np__81|np__132!=np__257)).#trainneg
cnf(i_0_65103, plain, (np__133=np__0|np__133!=np__81|np__133!=np__257)).#trainneg
cnf(i_0_65104, plain, (np__175=np__0|np__175!=np__81|np__175!=np__257)).#trainneg
cnf(i_0_109193, plain, (np__136!=np__6|~r1_xxreal_0(np__6561,np__6))).#trainneg
cnf(i_0_109194, plain, (np__6536!=np__136|~r1_xxreal_0(np__6561,np__6536))).#trainneg
cnf(i_0_109195, plain, (np__56!=np__136|~r1_xxreal_0(np__6561,np__56))).#trainneg
cnf(i_0_109196, plain, (np__23!=np__136|~r1_xxreal_0(np__6561,np__23))).#trainneg
cnf(i_0_109197, plain, (np__6289!=np__136|~r1_xxreal_0(np__6561,np__6289))).#trainneg
cnf(i_0_112173, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__5,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__5,X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_122898, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__5),X2)=k4_nat_d(k3_xcmplx_0(np__5,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_111221, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__6,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__6),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_116422, plain, (k4_nat_d(k3_xcmplx_0(X1,k2_nat_d(np__6425,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__6425),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_103875, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__257),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__257),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_123132, plain, (k4_nat_d(k4_nat_1(np__256,X1),np__257)=k4_nat_d(k3_xcmplx_0(np__256,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10818, plain, (r1_xxreal_0(X1,np__6558)|r1_xxreal_0(np__6558,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_65105, plain, (np__176=np__0|np__176!=np__81|np__176!=np__257)).#trainneg
cnf(i_0_2537, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k13_newton(X1,X2),k1_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_66461, plain, (np__21=np__0|np__21!=np__81|np__21!=np__257)).#trainneg
cnf(i_0_66462, plain, (np__6169=np__0|np__6169!=np__81|np__6169!=np__257)).#trainneg
cnf(i_0_66463, plain, (np__6421=np__0|np__6421!=np__81|np__6421!=np__257)).#trainneg
cnf(i_0_109198, plain, (np__231!=np__136|~r1_xxreal_0(np__6561,np__231))).#trainneg
cnf(i_0_109200, plain, (np__79!=np__136|~r1_xxreal_0(np__6561,np__79))).#trainneg
cnf(i_0_109214, plain, (np__6480!=np__257|~r1_xxreal_0(np__6561,np__6480))).#trainneg
cnf(i_0_109215, plain, (np__254!=np__257|~r1_xxreal_0(np__6561,np__254))).#trainneg
cnf(i_0_109216, plain, (np__253!=np__257|~r1_xxreal_0(np__6561,np__253))).#trainneg
cnf(i_0_112234, plain, (k4_nat_d(k3_xcmplx_0(X1,k2_nat_d(np__7,X2)),X2)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__7),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_126610, plain, (m1_subset_1(k3_xcmplx_0(X2,X1),k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_114002, plain, (X1=k2_nat_d(X2,np__6425)|r1_xxreal_0(np__6425,X1)|X2!=k2_xcmplx_0(X1,np__6425)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_112138, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__5,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__5),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_111263, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__6,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__6,X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_128495, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__7),X2)=k4_nat_d(k3_xcmplx_0(np__7,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_112287, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__7,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__7,X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_116433, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__6425,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__6425),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_66464, plain, (np__6422=np__0|np__6422!=np__81|np__6422!=np__257)).#trainneg
cnf(i_0_66465, plain, (np__6423=np__0|np__6423!=np__81|np__6423!=np__257)).#trainneg
cnf(i_0_69214, plain, (np__6559=np__0|np__6559!=np__81|np__6559!=np__257)).#trainneg
cnf(i_0_69215, plain, (np__6560=np__0|np__6560!=np__81|np__6560!=np__257)).#trainneg
cnf(i_0_109217, plain, (np__255!=np__257|~r1_xxreal_0(np__6561,np__255))).#trainneg
cnf(i_0_109218, plain, (np__120!=np__257|~r1_xxreal_0(np__6561,np__120))).#trainneg
cnf(i_0_109219, plain, (np__55!=np__257|~r1_xxreal_0(np__6561,np__55))).#trainneg
cnf(i_0_109220, plain, (np__6305!=np__257|~r1_xxreal_0(np__6561,np__6305))).#trainneg
cnf(i_0_109221, plain, (np__7!=np__257|~r1_xxreal_0(np__6561,np__7))).#trainneg
cnf(i_0_111218, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__6,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__6,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_127732, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__6425),X2)=k4_nat_d(k3_xcmplx_0(np__6425,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_116482, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(X1,X2),k2_nat_d(np__6425,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__6425,X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_123461, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__7,np__257),X1),np__257)=k4_nat_d(k3_xcmplx_0(np__7,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_112142, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__5,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(np__5,k2_xcmplx_0(X2,X1)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_111232, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__6,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(np__6,k2_xcmplx_0(X2,X1)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_111236, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__6,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__6),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10822, plain, (r1_xxreal_0(X1,np__6557)|r1_xxreal_0(np__6557,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_69216, plain, (np__77=np__0|np__77!=np__81|np__77!=np__257)).#trainneg
cnf(i_0_2555, plain, (X1=k2_nat_d(X2,k1_newton(X3,k2_newton(np__2,k2_newton(np__2,X4))))|r1_xxreal_0(k1_newton(X3,k2_newton(np__2,k2_newton(np__2,X4))),X1)|X2!=k2_xcmplx_0(k1_newton(X3,k4_pepin(X4)),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,k2_newton(np__2,k2_newton(np__2,X4))))|~v7_ordinal1(k2_newton(np__2,k2_newton(np__2,X4)))|~v7_ordinal1(k1_newton(X3,np__1))|~v7_ordinal1(np__1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_69217, plain, (np__78=np__0|np__78!=np__81|np__78!=np__257)).#trainneg
cnf(i_0_69218, plain, (np__24=np__0|np__24!=np__81|np__24!=np__257)).#trainneg
cnf(i_0_69967, plain, (np__135=np__0|np__135!=np__81|np__135!=np__257)).#trainneg
cnf(i_0_109223, plain, (np__5!=np__257|~r1_xxreal_0(np__6561,np__5))).#trainneg
cnf(i_0_109224, plain, (np__252!=np__257|~r1_xxreal_0(np__6561,np__252))).#trainneg
cnf(i_0_109226, plain, (np__6558!=np__257|~r1_xxreal_0(np__6561,np__6558))).#trainneg
cnf(i_0_109227, plain, (np__6424!=np__257|~r1_xxreal_0(np__6561,np__6424))).#trainneg
cnf(i_0_109228, plain, (np__6304!=np__257|~r1_xxreal_0(np__6561,np__6304))).#trainneg
cnf(i_0_112256, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__7,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(np__7,k2_xcmplx_0(X2,X1)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_112146, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__5,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__5),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_111230, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__6,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__6),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_116442, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__6425,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__6425),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_69968, plain, (np__6168=np__0|np__6168!=np__81|np__6168!=np__257)).#trainneg
cnf(i_0_2558, plain, (X1=k2_nat_d(X2,k1_newton(X3,np__0))|r1_xxreal_0(k1_newton(X3,np__0),X1)|X2!=k2_xcmplx_0(k1_newton(X3,np__6561),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,np__0))|~v7_ordinal1(k1_newton(X3,np__6561))|~v7_ordinal1(np__6561)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_69969, plain, (np__6=np__0|np__6!=np__81|np__6!=np__257)).#trainneg
cnf(i_0_69970, plain, (np__6536=np__0|np__6536!=np__81|np__6536!=np__257)).#trainneg
cnf(i_0_69971, plain, (np__56=np__0|np__56!=np__81|np__56!=np__257)).#trainneg
cnf(i_0_109229, plain, (np__80!=np__257|~r1_xxreal_0(np__6561,np__80))).#trainneg
cnf(i_0_109230, plain, (np__134!=np__257|~r1_xxreal_0(np__6561,np__134))).#trainneg
cnf(i_0_109231, plain, (np__232!=np__257|~r1_xxreal_0(np__6561,np__232))).#trainneg
cnf(i_0_109232, plain, (np__6400!=np__257|~r1_xxreal_0(np__6561,np__6400))).#trainneg
cnf(i_0_109233, plain, (np__6557!=np__257|~r1_xxreal_0(np__6561,np__6557))).#trainneg
cnf(i_0_112252, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__7,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__7),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_116424, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__6425,X1),k4_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__6425,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_116436, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__6425,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__6425),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_132112, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__6),X1)=k4_nat_d(k3_xcmplx_0(np__6,k2_xcmplx_0(X2,X1)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_112122, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__5,X1),k4_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__5,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_111186, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__6,np__257),X1),np__257)=k4_nat_d(k3_xcmplx_0(np__6,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10826, plain, (r1_xxreal_0(X1,np__6304)|r1_xxreal_0(np__6304,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_72495, plain, (np__23=np__0|np__23!=np__81|np__23!=np__257)).#trainneg
cnf(i_0_2561, plain, (X1=k2_nat_d(X2,k1_newton(X3,np__0))|r1_xxreal_0(k1_newton(X3,np__0),X1)|X2!=k2_xcmplx_0(k1_newton(X3,np__3),X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k1_newton(X3,np__0))|~v7_ordinal1(k1_newton(X3,np__3))|~v7_ordinal1(np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_72496, plain, (np__6289=np__0|np__6289!=np__81|np__6289!=np__257)).#trainneg
cnf(i_0_72497, plain, (np__231=np__0|np__231!=np__81|np__231!=np__257)).#trainneg
cnf(i_0_72686, plain, (np__79=np__0|np__79!=np__81|np__79!=np__257)).#trainneg
cnf(i_0_109234, plain, (np__6344!=np__257|~r1_xxreal_0(np__6561,np__6344))).#trainneg
cnf(i_0_109237, plain, (np__22!=np__257|~r1_xxreal_0(np__6561,np__22))).#trainneg
cnf(i_0_109238, plain, (np__111!=np__257|~r1_xxreal_0(np__6561,np__111))).#trainneg
cnf(i_0_109239, plain, (np__121!=np__257|~r1_xxreal_0(np__6561,np__121))).#trainneg
cnf(i_0_109240, plain, (np__132!=np__257|~r1_xxreal_0(np__6561,np__132))).#trainneg
cnf(i_0_116475, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,X2),k2_nat_d(np__6425,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__6425,X1),X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_128756, plain, (k4_nat_d(k3_xcmplx_0(np__1,X1),np__257)=np__136|k2_nat_d(X1,np__257)!=np__136|~v7_ordinal1(k2_nat_d(X1,np__257))|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X1,np__257),k1_numbers))).#trainneg
cnf(i_0_112236, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__7,X1),k4_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__7,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_105136, plain, (X1=np__0|~r1_xxreal_0(X1,k4_nat_d(k3_xcmplx_0(X2,np__8),X1))|~v7_ordinal1(k3_xcmplx_0(k2_xcmplx_0(X1,k2_nat_d(X2,X1)),np__8))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k2_nat_d(X2,X1),k1_numbers))).#trainneg
cnf(i_0_103874, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__257),X1)=k4_nat_d(k3_xcmplx_0(np__257,k2_xcmplx_0(X2,X1)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_112248, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__7,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__7),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_136237, plain, (k4_nat_d(X1,np__257)=np__136|k2_nat_d(X1,np__257)!=np__136|~v7_ordinal1(k2_nat_d(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_22841, plain, (np__136=np__3|k2_xcmplx_0(np__257,np__136)!=np__3)).#trainneg
cnf(i_0_136016, plain, (~r1_xxreal_0(np__257,k4_nat_d(np__1,np__257)))).#trainneg
cnf(i_0_109241, plain, (np__133!=np__257|~r1_xxreal_0(np__6561,np__133))).#trainneg
cnf(i_0_109242, plain, (np__175!=np__257|~r1_xxreal_0(np__6561,np__175))).#trainneg
cnf(i_0_109243, plain, (np__176!=np__257|~r1_xxreal_0(np__6561,np__176))).#trainneg
cnf(i_0_109244, plain, (np__21!=np__257|~r1_xxreal_0(np__6561,np__21))).#trainneg
cnf(i_0_136042, plain, (k4_nat_d(k3_xcmplx_0(X1,np__1),np__257)=k4_nat_d(X1,np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_136024, plain, (k2_nat_d(k3_xcmplx_0(X1,np__1),np__257)=k4_nat_d(X1,np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_136032, plain, (r1_xxreal_0(np__257,X1)|~r1_xxreal_0(np__257,k4_nat_d(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_137121, plain, (k4_nat_d(k3_xcmplx_0(np__257,k2_xcmplx_0(k2_nat_d(X2,X1),X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__257),X1)|r1_xxreal_0(X1,k2_nat_d(X2,X1))|~v7_ordinal1(k2_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_136190, plain, (k2_nat_d(k3_xcmplx_0(np__1,X1),np__257)=np__136|r1_xxreal_0(np__257,X1)|k2_nat_d(X1,np__257)!=np__136|~v7_ordinal1(k2_nat_d(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_136189, plain, (k2_nat_d(k3_xcmplx_0(X1,np__1),np__257)=np__136|r1_xxreal_0(np__257,X1)|k2_nat_d(X1,np__257)!=np__136|~v7_ordinal1(k2_nat_d(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_137933, plain, (~r1_xxreal_0(np__257,k2_nat_d(np__1,np__257)))).#trainneg
cnf(i_0_10834, plain, (r1_xxreal_0(X1,np__6559)|r1_xxreal_0(np__6559,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_107943, plain, (np__136=np__8|k2_xcmplx_0(np__6561,np__136)!=np__8)).#trainneg
cnf(i_0_107979, plain, (np__136=np__8|k2_xcmplx_0(np__136,np__6561)!=np__8)).#trainneg
cnf(i_0_138384, plain, (~r1_xxreal_0(np__257,k2_nat_d(np__2,np__257)))).#trainneg
cnf(i_0_138568, plain, (~r1_xxreal_0(np__257,k2_nat_d(np__256,np__257)))).#trainneg
cnf(i_0_109245, plain, (np__6169!=np__257|~r1_xxreal_0(np__6561,np__6169))).#trainneg
cnf(i_0_109246, plain, (np__6421!=np__257|~r1_xxreal_0(np__6561,np__6421))).#trainneg
cnf(i_0_109247, plain, (np__6422!=np__257|~r1_xxreal_0(np__6561,np__6422))).#trainneg
cnf(i_0_137882, plain, (k4_nat_d(np__1,np__257)=k2_nat_d(np__1,np__257))).#trainneg
cnf(i_0_137924, plain, (k4_nat_d(np__3,np__257)=np__3)).#trainneg
cnf(i_0_137925, plain, (k4_nat_d(X1,np__257)=k2_nat_d(X1,np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_108093, plain, (np__136=np__3|k2_xcmplx_0(np__6561,np__136)!=np__3)).#trainneg
cnf(i_0_108129, plain, (np__136=np__3|k2_xcmplx_0(np__136,np__6561)!=np__3)).#trainneg
cnf(i_0_109248, plain, (np__6423!=np__257|~r1_xxreal_0(np__6561,np__6423))).#trainneg
cnf(i_0_109249, plain, (np__6559!=np__257|~r1_xxreal_0(np__6561,np__6559))).#trainneg
cnf(i_0_109251, plain, (np__77!=np__257|~r1_xxreal_0(np__6561,np__77))).#trainneg
cnf(i_0_109252, plain, (np__78!=np__257|~r1_xxreal_0(np__6561,np__78))).#trainneg
cnf(i_0_109253, plain, (np__24!=np__257|~r1_xxreal_0(np__6561,np__24))).#trainneg
cnf(i_0_137687, plain, (k4_nat_d(X1,np__257)=np__136|k2_nat_d(X1,np__257)!=np__136|~v7_ordinal1(X1))).#trainneg
cnf(i_0_140421, plain, (r1_xxreal_0(np__257,X1)|~r1_xxreal_0(np__257,k2_nat_d(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_137912, plain, (k4_nat_d(np__2,np__257)=k2_nat_d(np__2,np__257))).#trainneg
cnf(i_0_137914, plain, (k4_nat_d(np__256,np__257)=k2_nat_d(np__256,np__257))).#trainneg
cnf(i_0_137921, plain, (k4_nat_d(np__4,np__257)=np__4)).#trainneg
cnf(i_0_10972, plain, (r1_xxreal_0(X1,np__5)|r1_xxreal_0(np__5,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_22915, plain, (np__136=np__81|k2_xcmplx_0(np__257,np__136)!=np__81)).#trainneg
cnf(i_0_23043, plain, (np__136=np__0|k2_xcmplx_0(np__257,np__136)!=np__0)).#trainneg
cnf(i_0_109257, plain, (np__135!=np__257|~r1_xxreal_0(np__6561,np__135))).#trainneg
cnf(i_0_109258, plain, (np__6168!=np__257|~r1_xxreal_0(np__6561,np__6168))).#trainneg
cnf(i_0_109259, plain, (np__6!=np__257|~r1_xxreal_0(np__6561,np__6))).#trainneg
cnf(i_0_109260, plain, (np__6536!=np__257|~r1_xxreal_0(np__6561,np__6536))).#trainneg
cnf(i_0_109261, plain, (np__56!=np__257|~r1_xxreal_0(np__6561,np__56))).#trainneg
cnf(i_0_138108, plain, (k2_nat_d(k2_xcmplx_0(X1,np__257),np__257)=k4_nat_d(X1,np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_23818, plain, (np__136=np__0|k2_xcmplx_0(np__257,np__136)!=np__6425)).#trainneg
cnf(i_0_24852, plain, (np__136=np__4|k2_xcmplx_0(np__257,np__136)!=np__4)).#trainneg
cnf(i_0_109262, plain, (np__23!=np__257|~r1_xxreal_0(np__6561,np__23))).#trainneg
cnf(i_0_109263, plain, (np__6289!=np__257|~r1_xxreal_0(np__6561,np__6289))).#trainneg
cnf(i_0_109264, plain, (np__231!=np__257|~r1_xxreal_0(np__6561,np__231))).#trainneg
cnf(i_0_109266, plain, (np__79!=np__257|~r1_xxreal_0(np__6561,np__79))).#trainneg
cnf(i_0_109280, plain, (np__6480!=np__8|~r1_xxreal_0(np__6561,np__6480))).#trainneg
cnf(i_0_139075, plain, (k4_nat_d(k3_xcmplx_0(np__257,k2_xcmplx_0(X2,k2_nat_d(X1,X2))),X2)=k4_nat_d(k3_xcmplx_0(X1,np__257),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X1,X2),k1_numbers))).#trainneg
cnf(i_0_140384, plain, (k4_nat_d(k2_nat_d(X1,np__257),np__257)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__1,np__257),X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X1,np__257),k1_numbers))).#trainneg
cnf(i_0_137864, plain, (k4_nat_d(X1,np__257)=k3_xcmplx_0(np__1,X1)|r1_xxreal_0(np__257,k3_xcmplx_0(np__1,X1))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k3_xcmplx_0(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10976, plain, (r1_xxreal_0(X1,np__7)|r1_xxreal_0(np__7,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_50412, plain, (np__136=np__0|k2_xcmplx_0(np__6561,np__136)!=np__6561)).#trainneg
cnf(i_0_109281, plain, (np__254!=np__8|~r1_xxreal_0(np__6561,np__254))).#trainneg
cnf(i_0_109282, plain, (np__253!=np__8|~r1_xxreal_0(np__6561,np__253))).#trainneg
cnf(i_0_109283, plain, (np__255!=np__8|~r1_xxreal_0(np__6561,np__255))).#trainneg
cnf(i_0_109284, plain, (np__120!=np__8|~r1_xxreal_0(np__6561,np__120))).#trainneg
cnf(i_0_109285, plain, (np__55!=np__8|~r1_xxreal_0(np__6561,np__55))).#trainneg
cnf(i_0_139881, plain, (k2_nat_d(k2_xcmplx_0(np__257,np__3),np__257)=np__3)).#trainneg
cnf(i_0_134641, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__6),X2)=k4_nat_d(k3_xcmplx_0(np__6,k2_xcmplx_0(X2,X1)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_142529, plain, (k2_nat_d(np__2,np__257)=np__2)).#trainneg
cnf(i_0_142716, plain, (k2_nat_d(np__256,np__257)=np__256)).#trainneg
cnf(i_0_142722, plain, (k2_nat_d(np__1,np__257)=np__1)).#trainneg
cnf(i_0_109286, plain, (np__6305!=np__8|~r1_xxreal_0(np__6561,np__6305))).#trainneg
cnf(i_0_109287, plain, (np__7!=np__8|~r1_xxreal_0(np__6561,np__7))).#trainneg
cnf(i_0_109289, plain, (np__5!=np__8|~r1_xxreal_0(np__6561,np__5))).#trainneg
cnf(i_0_109290, plain, (np__252!=np__8|~r1_xxreal_0(np__6561,np__252))).#trainneg
cnf(i_0_109292, plain, (np__6558!=np__8|~r1_xxreal_0(np__6561,np__6558))).#trainneg
cnf(i_0_141749, plain, (k4_nat_d(X1,np__257)=k2_xcmplx_0(X1,np__257)|r1_xxreal_0(np__257,k2_xcmplx_0(X1,np__257))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_142369, plain, (k2_nat_d(k2_nat_d(X1,np__257),np__257)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__1,np__257),X1),np__257)|r1_xxreal_0(np__257,k2_nat_d(X1,np__257))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_nat_d(X1,np__257))|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X1,np__257),k1_numbers))).#trainneg
cnf(i_0_141894, plain, (k2_nat_d(k2_xcmplx_0(np__257,np__81),np__257)=np__81)).#trainneg
cnf(i_0_141799, plain, (k2_nat_d(k2_xcmplx_0(np__257,X1),np__257)=k4_nat_d(X1,np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(np__257,X1))|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_141774, plain, (k4_nat_d(X1,np__257)=X1|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_144018, plain, (X1!=np__1|~r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10980, plain, (r1_xxreal_0(X1,np__6480)|r1_xxreal_0(np__6480,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_143712, plain, (np__136=np__2|np__6561!=np__2)).#trainneg
cnf(i_0_137912, plain, (k4_nat_d(np__2,np__257)=np__2)).#trainneg
cnf(i_0_143914, plain, (np__136=np__256|np__256!=np__6561)).#trainneg
cnf(i_0_100035, plain, (np__136=np__0|k2_xcmplx_0(np__136,np__6561)!=np__6561)).#trainneg
cnf(i_0_109293, plain, (np__6424!=np__8|~r1_xxreal_0(np__6561,np__6424))).#trainneg
cnf(i_0_109294, plain, (np__6304!=np__8|~r1_xxreal_0(np__6561,np__6304))).#trainneg
cnf(i_0_109295, plain, (np__80!=np__8|~r1_xxreal_0(np__6561,np__80))).#trainneg
cnf(i_0_109296, plain, (np__134!=np__8|~r1_xxreal_0(np__6561,np__134))).#trainneg
cnf(i_0_109297, plain, (np__232!=np__8|~r1_xxreal_0(np__6561,np__232))).#trainneg
cnf(i_0_143611, plain, (X1!=np__2|~r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_145338, plain, (k2_nat_d(k2_xcmplx_0(np__81,np__257),np__257)=np__81)).#trainneg
cnf(i_0_70995, plain, (k4_nat_d(k4_nat_1(np__2,X1),np__257)=k4_nat_d(k3_xcmplx_0(np__2,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_108261, plain, (np__136=np__4|k2_xcmplx_0(np__6561,np__136)!=np__4)).#trainneg
cnf(i_0_108297, plain, (np__136=np__4|k2_xcmplx_0(np__136,np__6561)!=np__4)).#trainneg
cnf(i_0_109020, plain, (np__136=np__81|k2_xcmplx_0(np__6561,np__136)!=np__81)).#trainneg
cnf(i_0_109298, plain, (np__6400!=np__8|~r1_xxreal_0(np__6561,np__6400))).#trainneg
cnf(i_0_109299, plain, (np__6557!=np__8|~r1_xxreal_0(np__6561,np__6557))).#trainneg
cnf(i_0_109300, plain, (np__6344!=np__8|~r1_xxreal_0(np__6561,np__6344))).#trainneg
cnf(i_0_109303, plain, (np__22!=np__8|~r1_xxreal_0(np__6561,np__22))).#trainneg
cnf(i_0_109304, plain, (np__111!=np__8|~r1_xxreal_0(np__6561,np__111))).#trainneg
cnf(i_0_143813, plain, (X1!=np__256|~r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_143836, plain, (k2_nat_d(k2_xcmplx_0(np__257,np__256),np__257)=np__256)).#trainneg
cnf(i_0_144041, plain, (k2_nat_d(k2_xcmplx_0(np__257,np__1),np__257)=np__1)).#trainneg
cnf(i_0_143634, plain, (k2_nat_d(k2_xcmplx_0(np__257,np__2),np__257)=np__2)).#trainneg
cnf(i_0_143982, plain, (k2_nat_d(k2_xcmplx_0(np__1,np__257),np__257)=np__1)).#trainneg
cnf(i_0_146194, plain, (np__257!=np__1)).#trainneg
cnf(i_0_11583, plain, (r1_xxreal_0(X1,np__24)|r1_xxreal_0(np__24,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_146351, plain, (np__8=np__0|np__8!=np__2|np__8!=np__257|np__6561!=np__2)).#trainneg
cnf(i_0_137914, plain, (k4_nat_d(np__256,np__257)=np__256)).#trainneg
cnf(i_0_146190, plain, (np__6425!=np__1)).#trainneg
cnf(i_0_146250, plain, (np__6561!=np__2)).#trainneg
cnf(i_0_146657, plain, (np__256!=np__6561)).#trainneg
cnf(i_0_146729, plain, (np__6425!=np__2)).#trainneg
cnf(i_0_146733, plain, (np__257!=np__2)).#trainneg
cnf(i_0_146739, plain, (X1!=np__2|~r1_xxreal_0(np__257,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_143575, plain, (k2_nat_d(k2_xcmplx_0(np__2,np__257),np__257)=np__2)).#trainneg
cnf(i_0_137882, plain, (k4_nat_d(np__1,np__257)=np__1)).#trainneg
cnf(i_0_145142, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__1,np__257),X1),np__257)=k2_nat_d(X1,np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_146200, plain, (X1!=np__1|~r1_xxreal_0(np__257,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_143777, plain, (k2_nat_d(k2_xcmplx_0(np__256,np__257),np__257)=np__256)).#trainneg
cnf(i_0_144702, negated_conjecture, (r1_xxreal_0(np__257,k2_xcmplx_0(k13_newton(np__3,np__8),np__257))|r1_xxreal_0(np__257,k13_newton(np__3,np__8))|k2_xcmplx_0(k13_newton(np__3,np__8),np__257)!=np__136|~v7_ordinal1(k2_xcmplx_0(k13_newton(np__3,np__8),np__257))|~v7_ordinal1(k13_newton(np__3,np__8)))).#trainneg
cnf(i_0_92653, plain, (k2_nat_d(k4_nat_1(np__1,X1),np__257)=k4_nat_d(k3_xcmplx_0(np__1,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k4_nat_1(np__1,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_147185, plain, (np__6425!=np__256)).#trainneg
cnf(i_0_147189, plain, (np__256!=np__257)).#trainneg
cnf(i_0_109305, plain, (np__121!=np__8|~r1_xxreal_0(np__6561,np__121))).#trainneg
cnf(i_0_109306, plain, (np__132!=np__8|~r1_xxreal_0(np__6561,np__132))).#trainneg
cnf(i_0_109307, plain, (np__133!=np__8|~r1_xxreal_0(np__6561,np__133))).#trainneg
cnf(i_0_109308, plain, (np__175!=np__8|~r1_xxreal_0(np__6561,np__175))).#trainneg
cnf(i_0_109309, plain, (np__176!=np__8|~r1_xxreal_0(np__6561,np__176))).#trainneg
cnf(i_0_11595, plain, (r1_xxreal_0(X1,np__254)|r1_xxreal_0(np__254,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_65679, plain, (k4_nat_d(k4_nat_1(np__1,X1),np__257)=k4_nat_d(k3_xcmplx_0(np__1,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_109310, plain, (np__21!=np__8|~r1_xxreal_0(np__6561,np__21))).#trainneg
cnf(i_0_109311, plain, (np__6169!=np__8|~r1_xxreal_0(np__6561,np__6169))).#trainneg
cnf(i_0_109312, plain, (np__6421!=np__8|~r1_xxreal_0(np__6561,np__6421))).#trainneg
cnf(i_0_109313, plain, (np__6422!=np__8|~r1_xxreal_0(np__6561,np__6422))).#trainneg
cnf(i_0_109314, plain, (np__6423!=np__8|~r1_xxreal_0(np__6561,np__6423))).#trainneg
cnf(i_0_145790, plain, (k4_nat_1(X1,k2_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,k2_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(k4_nat_1(X1,k2_nat_d(X2,np__257)),np__257))|~v7_ordinal1(k4_nat_1(X1,k2_nat_d(X2,np__257)))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_145793, plain, (k4_nat_1(X1,k4_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(X2,X1),np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,k4_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(k4_nat_1(X1,k4_nat_d(X2,np__257)),np__257))|~v7_ordinal1(k4_nat_1(X1,k4_nat_d(X2,np__257)))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_145549, plain, (k4_nat_d(X1,np__257)=X1|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(np__257,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_148668, plain, (k4_nat_d(k2_xcmplx_0(np__1,np__257),np__257)=np__1)).#trainneg
cnf(i_0_148118, plain, (k4_nat_d(k2_xcmplx_0(np__257,np__256),np__257)=np__256)).#trainneg
cnf(i_0_148923, negated_conjecture, (r1_xxreal_0(np__257,k13_newton(np__3,np__8))|k2_xcmplx_0(k13_newton(np__3,np__8),np__257)!=np__136|~v7_ordinal1(k13_newton(np__3,np__8)))).#trainneg
cnf(i_0_148444, plain, (k4_nat_d(k2_xcmplx_0(np__257,np__1),np__257)=np__1)).#trainneg
cnf(i_0_109315, plain, (np__6559!=np__8|~r1_xxreal_0(np__6561,np__6559))).#trainneg
cnf(i_0_109317, plain, (np__77!=np__8|~r1_xxreal_0(np__6561,np__77))).#trainneg
cnf(i_0_109318, plain, (np__78!=np__8|~r1_xxreal_0(np__6561,np__78))).#trainneg
cnf(i_0_109319, plain, (np__24!=np__8|~r1_xxreal_0(np__6561,np__24))).#trainneg
cnf(i_0_109323, plain, (np__135!=np__8|~r1_xxreal_0(np__6561,np__135))).#trainneg
cnf(i_0_147195, plain, (X1!=np__256|~r1_xxreal_0(np__257,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_141476, plain, (k2_nat_d(k2_xcmplx_0(np__257,np__4),np__257)=np__4)).#trainneg
cnf(i_0_144727, plain, (k2_nat_d(X1,np__257)=k2_xcmplx_0(X1,np__257)|r1_xxreal_0(np__257,k2_xcmplx_0(X1,np__257))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_146126, plain, (X1=np__136|r1_xxreal_0(np__257,X1)|k2_nat_d(X1,np__257)!=np__136|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_145143, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__1,np__257),X1),np__257)=np__136|r1_xxreal_0(np__257,X1)|X1!=np__136|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16691, plain, (r1_xxreal_0(X1,np__78)|r1_xxreal_0(np__78,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2205, plain, (k4_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),k1_newton(np__3,np__4)),np__257)=np__136|~v7_ordinal1(k13_newton(np__3,np__4))|~v7_ordinal1(k1_newton(np__3,np__4)))).#trainneg
cnf(i_0_146198, plain, (np__8!=np__1|~r1_xxreal_0(np__257,np__8))).#trainneg
cnf(i_0_109324, plain, (np__6168!=np__8|~r1_xxreal_0(np__6561,np__6168))).#trainneg
cnf(i_0_109325, plain, (np__6!=np__8|~r1_xxreal_0(np__6561,np__6))).#trainneg
cnf(i_0_109326, plain, (np__6536!=np__8|~r1_xxreal_0(np__6561,np__6536))).#trainneg
cnf(i_0_109327, plain, (np__56!=np__8|~r1_xxreal_0(np__6561,np__56))).#trainneg
cnf(i_0_109328, plain, (np__23!=np__8|~r1_xxreal_0(np__6561,np__23))).#trainneg
cnf(i_0_144825, plain, (k4_nat_d(k3_xcmplx_0(np__1,X1),np__257)=k4_nat_d(k2_xcmplx_0(X1,np__257),np__257)|r1_xxreal_0(np__257,k2_xcmplx_0(X1,np__257))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1)|~m1_subset_1(k2_xcmplx_0(X1,np__257),k1_numbers))).#trainneg
cnf(i_0_149311, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__1,np__257),np__136),np__257)=np__136)).#trainneg
cnf(i_0_145792, plain, (k4_nat_1(X1,k4_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,k4_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(k4_nat_1(X1,k4_nat_d(X2,np__257)),np__257))|~v7_ordinal1(k4_nat_1(X1,k4_nat_d(X2,np__257)))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_150741, negated_conjecture, (r1_xxreal_0(np__257,k1_newton(np__3,np__8))|k2_xcmplx_0(k1_newton(np__3,np__8),np__257)!=np__136|~v7_ordinal1(k1_newton(np__3,np__8)))).#trainneg
cnf(i_0_151235, plain, (k2_nat_d(k2_xcmplx_0(np__4,np__257),np__257)=np__4)).#trainneg
cnf(i_0_146737, plain, (np__8!=np__2|~r1_xxreal_0(np__257,np__8))).#trainneg
cnf(i_0_147193, plain, (np__256!=np__8|~r1_xxreal_0(np__257,np__8))).#trainneg
cnf(i_0_109329, plain, (np__6289!=np__8|~r1_xxreal_0(np__6561,np__6289))).#trainneg
cnf(i_0_109330, plain, (np__231!=np__8|~r1_xxreal_0(np__6561,np__231))).#trainneg
cnf(i_0_109332, plain, (np__79!=np__8|~r1_xxreal_0(np__6561,np__79))).#trainneg
cnf(i_0_112872, plain, (np__6480!=np__81|~r1_xxreal_0(np__6561,np__6480))).#trainneg
cnf(i_0_112873, plain, (np__254!=np__81|~r1_xxreal_0(np__6561,np__254))).#trainneg
cnf(i_0_152437, plain, (k4_nat_d(k2_xcmplx_0(X1,np__257),np__257)=k4_nat_d(X1,np__257)|r1_xxreal_0(np__257,k2_xcmplx_0(X1,np__257))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_17142, plain, (r1_xxreal_0(X1,np__80)|r1_xxreal_0(np__80,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2672, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k2_nat_1(X1,X2),k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_112874, plain, (np__253!=np__81|~r1_xxreal_0(np__6561,np__253))).#trainneg
cnf(i_0_112875, plain, (np__255!=np__81|~r1_xxreal_0(np__6561,np__255))).#trainneg
cnf(i_0_112876, plain, (np__120!=np__81|~r1_xxreal_0(np__6561,np__120))).#trainneg
cnf(i_0_112877, plain, (np__55!=np__81|~r1_xxreal_0(np__6561,np__55))).#trainneg
cnf(i_0_112878, plain, (np__6305!=np__81|~r1_xxreal_0(np__6561,np__6305))).#trainneg
cnf(i_0_150913, plain, (k2_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),k13_newton(np__3,np__4)),np__257)=np__136|~v7_ordinal1(k3_xcmplx_0(k13_newton(np__3,np__4),k13_newton(np__3,np__4)))|~v7_ordinal1(k13_newton(np__3,np__4)))).#trainneg
cnf(i_0_152054, plain, (k4_nat_d(k3_xcmplx_0(np__136,k2_xcmplx_0(np__1,np__257)),np__257)=np__136)).#trainneg
cnf(i_0_152446, plain, (k4_nat_d(k2_xcmplx_0(X1,np__257),np__257)=np__136|r1_xxreal_0(np__257,k2_xcmplx_0(X1,np__257))|r1_xxreal_0(np__257,X1)|X1!=np__136|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_153069, plain, (k4_nat_1(X1,k4_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,k4_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k4_nat_1(X1,k4_nat_d(X2,np__257)))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_150585, plain, (k4_nat_d(k4_nat_1(np__1,X1),np__257)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__1,np__257),X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(np__1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_109056, plain, (np__136=np__81|k2_xcmplx_0(np__136,np__6561)!=np__81)).#trainneg
cnf(i_0_2673, plain, (v1_xboole_0(k1_numbers)|m1_subset_1(k2_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_115817, plain, (np__6480!=np__3|~r1_xxreal_0(np__6561,np__6480))).#trainneg
cnf(i_0_115818, plain, (np__254!=np__3|~r1_xxreal_0(np__6561,np__254))).#trainneg
cnf(i_0_115819, plain, (np__253!=np__3|~r1_xxreal_0(np__6561,np__253))).#trainneg
cnf(i_0_112879, plain, (np__7!=np__81|~r1_xxreal_0(np__6561,np__7))).#trainneg
cnf(i_0_112881, plain, (np__5!=np__81|~r1_xxreal_0(np__6561,np__5))).#trainneg
cnf(i_0_112882, plain, (np__252!=np__81|~r1_xxreal_0(np__6561,np__252))).#trainneg
cnf(i_0_112884, plain, (np__6558!=np__81|~r1_xxreal_0(np__6561,np__6558))).#trainneg
cnf(i_0_112885, plain, (np__6424!=np__81|~r1_xxreal_0(np__6561,np__6424))).#trainneg
cnf(i_0_145027, plain, (r1_xxreal_0(np__257,k2_nat_d(X1,np__257))|r1_xxreal_0(np__257,X1)|~r1_xxreal_0(np__257,k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__1,np__257),X1),np__257))|~v7_ordinal1(k2_nat_d(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_148451, plain, (k4_nat_d(k4_nat_1(np__81,np__1),np__257)=np__81)).#trainneg
cnf(i_0_142364, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__1,np__257),X1),np__257)=np__136|r1_xxreal_0(np__257,X1)|k2_nat_d(X1,np__257)!=np__136|~v7_ordinal1(k2_nat_d(X1,np__257))|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X1,np__257),k1_numbers))).#trainneg
cnf(i_0_145791, plain, (k4_nat_1(X1,k2_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(X2,X1),np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,k2_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(k4_nat_1(X1,k2_nat_d(X2,np__257)),np__257))|~v7_ordinal1(k4_nat_1(X1,k2_nat_d(X2,np__257)))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_18113, plain, (r1_xxreal_0(X1,np__77)|r1_xxreal_0(np__77,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_117377, plain, (np__136=np__0|k2_xcmplx_0(np__6561,np__136)!=np__0)).#trainneg
cnf(i_0_2688, plain, (k2_nat_d(k2_xcmplx_0(k1_newton(X3,X4),k3_xcmplx_0(X1,X2)),X1)=k1_newton(X3,X4)|r1_xxreal_0(X1,k1_newton(X3,X4))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v7_ordinal1(k2_xcmplx_0(k1_newton(X3,X4),k3_xcmplx_0(X1,X2)))|~v7_ordinal1(k1_newton(X3,X4))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_115820, plain, (np__255!=np__3|~r1_xxreal_0(np__6561,np__255))).#trainneg
cnf(i_0_115821, plain, (np__120!=np__3|~r1_xxreal_0(np__6561,np__120))).#trainneg
cnf(i_0_115822, plain, (np__55!=np__3|~r1_xxreal_0(np__6561,np__55))).#trainneg
cnf(i_0_112886, plain, (np__6304!=np__81|~r1_xxreal_0(np__6561,np__6304))).#trainneg
cnf(i_0_112887, plain, (np__80!=np__81|~r1_xxreal_0(np__6561,np__80))).#trainneg
cnf(i_0_112888, plain, (np__134!=np__81|~r1_xxreal_0(np__6561,np__134))).#trainneg
cnf(i_0_112889, plain, (np__232!=np__81|~r1_xxreal_0(np__6561,np__232))).#trainneg
cnf(i_0_112890, plain, (np__6400!=np__81|~r1_xxreal_0(np__6561,np__6400))).#trainneg
cnf(i_0_117413, plain, (np__136=np__0|k2_xcmplx_0(np__136,np__6561)!=np__0)).#trainneg
cnf(i_0_2716, plain, (~v1_xcmplx_0(X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_115823, plain, (np__6305!=np__3|~r1_xxreal_0(np__6561,np__6305))).#trainneg
cnf(i_0_115824, plain, (np__7!=np__3|~r1_xxreal_0(np__6561,np__7))).#trainneg
cnf(i_0_115826, plain, (np__5!=np__3|~r1_xxreal_0(np__6561,np__5))).#trainneg
cnf(i_0_112891, plain, (np__6557!=np__81|~r1_xxreal_0(np__6561,np__6557))).#trainneg
cnf(i_0_112892, plain, (np__6344!=np__81|~r1_xxreal_0(np__6561,np__6344))).#trainneg
cnf(i_0_112895, plain, (np__22!=np__81|~r1_xxreal_0(np__6561,np__22))).#trainneg
cnf(i_0_112896, plain, (np__111!=np__81|~r1_xxreal_0(np__6561,np__111))).#trainneg
cnf(i_0_112897, plain, (np__121!=np__81|~r1_xxreal_0(np__6561,np__121))).#trainneg
cnf(i_0_148416, plain, (k4_nat_d(k4_nat_1(np__1,np__81),np__257)=np__81)).#trainneg
cnf(i_0_148478, plain, (k2_xcmplx_0(np__1,np__257)=np__1|r1_xxreal_0(np__257,k2_xcmplx_0(np__1,np__257))|~v7_ordinal1(k2_xcmplx_0(np__1,np__257)))).#trainneg
cnf(i_0_145999, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__1,np__257),X1),np__257)=k4_nat_d(X1,np__257)|r1_xxreal_0(np__257,k4_nat_d(X1,np__257))|~v7_ordinal1(k2_xcmplx_0(k4_nat_d(X1,np__257),np__257))|~v7_ordinal1(k4_nat_d(X1,np__257))|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X1,np__257),k1_numbers))).#trainneg
cnf(i_0_18120, plain, (r1_xxreal_0(X1,np__79)|r1_xxreal_0(np__79,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_143559, plain, (np__136=np__2|k2_xcmplx_0(np__257,np__136)!=np__2)).#trainneg
cnf(i_0_2735, plain, (v1_xreal_0(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_115827, plain, (np__252!=np__3|~r1_xxreal_0(np__6561,np__252))).#trainneg
cnf(i_0_115829, plain, (np__6558!=np__3|~r1_xxreal_0(np__6561,np__6558))).#trainneg
cnf(i_0_115830, plain, (np__6424!=np__3|~r1_xxreal_0(np__6561,np__6424))).#trainneg
cnf(i_0_112898, plain, (np__132!=np__81|~r1_xxreal_0(np__6561,np__132))).#trainneg
cnf(i_0_112899, plain, (np__133!=np__81|~r1_xxreal_0(np__6561,np__133))).#trainneg
cnf(i_0_112900, plain, (np__175!=np__81|~r1_xxreal_0(np__6561,np__175))).#trainneg
cnf(i_0_112901, plain, (np__176!=np__81|~r1_xxreal_0(np__6561,np__176))).#trainneg
cnf(i_0_112902, plain, (np__21!=np__81|~r1_xxreal_0(np__6561,np__21))).#trainneg
cnf(i_0_141752, plain, (k4_nat_d(X1,np__257)=np__136|r1_xxreal_0(np__257,X1)|k2_xcmplx_0(X1,np__257)!=np__136|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_144107, plain, (k4_nat_d(k4_nat_1(X1,np__1),np__257)=k4_nat_d(k3_xcmplx_0(X1,np__1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_154754, plain, (k2_nat_d(k4_nat_1(np__81,np__1),np__257)=np__81|~v7_ordinal1(k4_nat_1(np__81,np__1)))).#trainneg
cnf(i_0_154871, plain, (k2_nat_d(k4_nat_1(np__1,np__81),np__257)=np__81|~v7_ordinal1(k4_nat_1(np__1,np__81)))).#trainneg
cnf(i_0_156542, plain, (k2_nat_d(X1,np__257)=np__136|r1_xxreal_0(np__257,X1)|k2_xcmplx_0(X1,np__257)!=np__136|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_153279, plain, (np__8=np__0|np__8!=np__2|np__8!=np__257)).#trainneg
cnf(i_0_153282, plain, (np__8=np__0|np__256!=np__8|np__8!=np__257)).#trainneg
cnf(i_0_115831, plain, (np__6304!=np__3|~r1_xxreal_0(np__6561,np__6304))).#trainneg
cnf(i_0_115832, plain, (np__80!=np__3|~r1_xxreal_0(np__6561,np__80))).#trainneg
cnf(i_0_112903, plain, (np__6169!=np__81|~r1_xxreal_0(np__6561,np__6169))).#trainneg
cnf(i_0_112904, plain, (np__6421!=np__81|~r1_xxreal_0(np__6561,np__6421))).#trainneg
cnf(i_0_112905, plain, (np__6422!=np__81|~r1_xxreal_0(np__6561,np__6422))).#trainneg
cnf(i_0_112906, plain, (np__6423!=np__81|~r1_xxreal_0(np__6561,np__6423))).#trainneg
cnf(i_0_112907, plain, (np__6559!=np__81|~r1_xxreal_0(np__6561,np__6559))).#trainneg
cnf(i_0_155791, plain, (k2_xcmplx_0(np__1,np__257)=np__1|r1_xxreal_0(np__257,k2_xcmplx_0(np__1,np__257)))).#trainneg
cnf(i_0_154443, plain, (k4_nat_1(X1,k4_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,k4_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_1(X1,k4_nat_d(X2,np__257)),k5_numbers))).#trainneg
cnf(i_0_154009, plain, (k4_nat_d(k2_xcmplx_0(np__257,X1),np__257)=np__136|r1_xxreal_0(np__257,k2_xcmplx_0(np__257,X1))|X1!=np__136|~v7_ordinal1(k2_xcmplx_0(np__257,X1))|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_154435, plain, (k4_nat_1(X1,k2_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,k2_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|r1_xxreal_0(np__257,X2)|~v7_ordinal1(k4_nat_1(X1,k2_nat_d(X2,np__257)))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_155233, plain, (k4_nat_1(X1,k2_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(X2,X1),np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,k2_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(np__257,k4_nat_1(X1,k2_nat_d(X2,np__257))))|~v7_ordinal1(k4_nat_1(X1,k2_nat_d(X2,np__257)))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k4_nat_1(X1,k2_nat_d(X2,np__257)),k1_numbers))).#trainneg
cnf(i_0_155243, plain, (k4_nat_1(X1,k2_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(X2,X1),np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,k2_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k4_nat_1(X1,k2_nat_d(X2,np__257)))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_154430, plain, (k4_nat_1(X1,k2_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,k2_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k4_nat_1(X1,k2_nat_d(X2,np__257)))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_152427, plain, (k2_nat_d(k3_xcmplx_0(X1,np__1),np__257)=k4_nat_d(k2_xcmplx_0(X1,np__257),np__257)|r1_xxreal_0(np__257,k2_xcmplx_0(X1,np__257))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18127, plain, (r1_xxreal_0(X1,np__6536)|r1_xxreal_0(np__6536,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2750, plain, (k7_xcmplx_0(k3_xcmplx_0(X3,k1_newton(X1,X2)),X4)=k3_xcmplx_0(k1_newton(X1,X2),k7_xcmplx_0(X3,X4))|~v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_115833, plain, (np__134!=np__3|~r1_xxreal_0(np__6561,np__134))).#trainneg
cnf(i_0_115834, plain, (np__232!=np__3|~r1_xxreal_0(np__6561,np__232))).#trainneg
cnf(i_0_115835, plain, (np__6400!=np__3|~r1_xxreal_0(np__6561,np__6400))).#trainneg
cnf(i_0_112909, plain, (np__77!=np__81|~r1_xxreal_0(np__6561,np__77))).#trainneg
cnf(i_0_112910, plain, (np__78!=np__81|~r1_xxreal_0(np__6561,np__78))).#trainneg
cnf(i_0_112911, plain, (np__24!=np__81|~r1_xxreal_0(np__6561,np__24))).#trainneg
cnf(i_0_112915, plain, (np__135!=np__81|~r1_xxreal_0(np__6561,np__135))).#trainneg
cnf(i_0_112916, plain, (np__6168!=np__81|~r1_xxreal_0(np__6561,np__6168))).#trainneg
cnf(i_0_135234, plain, (k2_nat_d(k4_nat_1(k2_nat_d(np__6,np__257),X1),np__257)=k4_nat_d(k3_xcmplx_0(np__6,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k4_nat_1(k2_nat_d(np__6,np__257),X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_146509, plain, (k2_xcmplx_0(np__2,np__257)=np__2|r1_xxreal_0(np__257,k2_xcmplx_0(np__2,np__257))|~v7_ordinal1(k2_xcmplx_0(np__2,np__257)))).#trainneg
cnf(i_0_135757, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__6425,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__6425,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_152428, plain, (k2_nat_d(k3_xcmplx_0(np__1,X1),np__257)=k4_nat_d(k2_xcmplx_0(X1,np__257),np__257)|r1_xxreal_0(np__257,k2_xcmplx_0(X1,np__257))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_146121, plain, (k4_nat_d(k4_nat_1(X1,X2),np__257)=k4_nat_d(k3_xcmplx_0(X2,X1),np__257)|r1_xxreal_0(np__257,X2)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_148675, plain, (k4_nat_d(k3_xcmplx_0(np__3,k2_xcmplx_0(np__1,np__257)),np__257)=np__3)).#trainneg
cnf(i_0_143700, plain, (k4_nat_d(k4_nat_1(X1,np__2),np__257)=k4_nat_d(k3_xcmplx_0(X1,np__2),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_111212, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__6,X1),k4_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__6,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2751, plain, (k3_xcmplx_0(k3_xcmplx_0(X3,k1_newton(X1,X2)),X4)=k3_xcmplx_0(k1_newton(X1,X2),k3_xcmplx_0(X3,X4))|~v1_xcmplx_0(k1_newton(X1,X2))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_115836, plain, (np__6557!=np__3|~r1_xxreal_0(np__6561,np__6557))).#trainneg
cnf(i_0_115837, plain, (np__6344!=np__3|~r1_xxreal_0(np__6561,np__6344))).#trainneg
cnf(i_0_115840, plain, (np__22!=np__3|~r1_xxreal_0(np__6561,np__22))).#trainneg
cnf(i_0_112917, plain, (np__6!=np__81|~r1_xxreal_0(np__6561,np__6))).#trainneg
cnf(i_0_112918, plain, (np__6536!=np__81|~r1_xxreal_0(np__6561,np__6536))).#trainneg
cnf(i_0_112919, plain, (np__56!=np__81|~r1_xxreal_0(np__6561,np__56))).#trainneg
cnf(i_0_112920, plain, (np__23!=np__81|~r1_xxreal_0(np__6561,np__23))).#trainneg
cnf(i_0_112921, plain, (np__6289!=np__81|~r1_xxreal_0(np__6561,np__6289))).#trainneg
cnf(i_0_112134, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__5,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__5),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_106116, plain, (X1=np__0|~r1_xxreal_0(X1,k2_nat_d(np__6561,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_106559, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__256,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(X1,np__256),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_112260, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__7,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__7),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_151023, plain, (k2_nat_d(k3_xcmplx_0(k13_newton(np__3,np__4),k13_newton(np__3,np__4)),np__257)=np__136|r1_xxreal_0(np__257,k13_newton(np__3,np__4))|~v7_ordinal1(k13_newton(np__3,np__4)))).#trainneg
cnf(i_0_128665, plain, (k2_nat_d(k2_nat_d(X1,np__257),np__257)=k4_nat_d(k3_xcmplx_0(np__1,X1),np__257)|~v7_ordinal1(k2_nat_d(X1,np__257))|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X1,np__257),k1_numbers))).#trainneg
cnf(i_0_111227, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__6,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(np__6,k2_xcmplx_0(X1,X2)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_106569, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__256,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__256,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18134, plain, (r1_xxreal_0(X1,np__6305)|r1_xxreal_0(np__6305,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_143761, plain, (np__136=np__256|k2_xcmplx_0(np__257,np__136)!=np__256)).#trainneg
cnf(i_0_2757, plain, (k3_xcmplx_0(k1_newton(X3,X4),k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X1,k3_xcmplx_0(X2,k1_newton(X3,X4)))|~v1_xcmplx_0(k1_newton(X3,X4))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_115841, plain, (np__111!=np__3|~r1_xxreal_0(np__6561,np__111))).#trainneg
cnf(i_0_115842, plain, (np__121!=np__3|~r1_xxreal_0(np__6561,np__121))).#trainneg
cnf(i_0_115843, plain, (np__132!=np__3|~r1_xxreal_0(np__6561,np__132))).#trainneg
cnf(i_0_112922, plain, (np__231!=np__81|~r1_xxreal_0(np__6561,np__231))).#trainneg
cnf(i_0_112924, plain, (np__79!=np__81|~r1_xxreal_0(np__6561,np__79))).#trainneg
cnf(i_0_114563, plain, (np__6480!=np__4|~r1_xxreal_0(np__6561,np__6480))).#trainneg
cnf(i_0_114564, plain, (np__254!=np__4|~r1_xxreal_0(np__6561,np__254))).#trainneg
cnf(i_0_114565, plain, (np__253!=np__4|~r1_xxreal_0(np__6561,np__253))).#trainneg
cnf(i_0_145528, plain, (k4_nat_d(X1,np__257)=k2_xcmplx_0(np__257,X1)|r1_xxreal_0(np__257,k2_xcmplx_0(np__257,X1))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(np__257,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_96929, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,np__136),k2_xcmplx_0(X2,X1)),X1)=k4_nat_d(k4_nat_1(X2,np__136),X1)|r1_xxreal_0(X1,np__136)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_162761, plain, (k4_nat_d(k3_xcmplx_0(np__1,X1),np__257)=k2_nat_d(X1,np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_104518, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(k2_nat_d(X1,X2),X2),np__3),X2)=k4_nat_d(k3_xcmplx_0(np__3,X1),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_50880, plain, (k2_nat_d(X1,np__3)=np__0|X1!=np__3|~v7_ordinal1(X1))).#trainneg
cnf(i_0_49144, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__1),X1)=k4_nat_d(k2_xcmplx_0(X2,X1),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(X2,X1))|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64091, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__136,X1),k4_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__136,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2761, plain, (k7_xcmplx_0(k3_xcmplx_0(k1_newton(X2,X3),X1),X4)=k3_xcmplx_0(X1,k7_xcmplx_0(k1_newton(X2,X3),X4))|~v1_xcmplx_0(k1_newton(X2,X3))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_115844, plain, (np__133!=np__3|~r1_xxreal_0(np__6561,np__133))).#trainneg
cnf(i_0_115845, plain, (np__175!=np__3|~r1_xxreal_0(np__6561,np__175))).#trainneg
cnf(i_0_115846, plain, (np__176!=np__3|~r1_xxreal_0(np__6561,np__176))).#trainneg
cnf(i_0_114566, plain, (np__255!=np__4|~r1_xxreal_0(np__6561,np__255))).#trainneg
cnf(i_0_114567, plain, (np__120!=np__4|~r1_xxreal_0(np__6561,np__120))).#trainneg
cnf(i_0_114568, plain, (np__55!=np__4|~r1_xxreal_0(np__6561,np__55))).#trainneg
cnf(i_0_114569, plain, (np__6305!=np__4|~r1_xxreal_0(np__6561,np__6305))).#trainneg
cnf(i_0_114570, plain, (np__7!=np__4|~r1_xxreal_0(np__6561,np__7))).#trainneg
cnf(i_0_111201, plain, (k4_nat_d(k3_xcmplx_0(X1,k2_nat_d(np__6,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__6,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_64475, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__3,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__3),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_111235, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__6,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__6),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_164997, plain, (k4_nat_d(k3_xcmplx_0(np__257,np__3),np__257)=k4_nat_d(k3_xcmplx_0(np__3,np__257),np__257))).#trainneg
cnf(i_0_166903, plain, (k4_nat_d(k3_xcmplx_0(np__3,np__2),np__257)=k4_nat_d(np__6,np__257))).#trainneg
cnf(i_0_112133, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__5,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__5,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_66394, plain, (k4_nat_d(X1,np__257)=np__0|k4_nat_d(X1,np__257)!=np__257|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X1,np__257),k5_numbers))).#trainneg
cnf(i_0_143298, plain, (k2_nat_d(k2_xcmplx_0(np__3,np__257),np__257)=np__3)).#trainneg
cnf(i_0_22889, plain, (r1_xxreal_0(X1,np__81)|r1_xxreal_0(np__81,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_165099, plain, (np__0=np__2|np__5!=np__3)).#trainneg
cnf(i_0_2762, plain, (k3_xcmplx_0(k3_xcmplx_0(k1_newton(X2,X3),X1),X4)=k3_xcmplx_0(X1,k3_xcmplx_0(k1_newton(X2,X3),X4))|~v1_xcmplx_0(k1_newton(X2,X3))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_115847, plain, (np__21!=np__3|~r1_xxreal_0(np__6561,np__21))).#trainneg
cnf(i_0_115848, plain, (np__6169!=np__3|~r1_xxreal_0(np__6561,np__6169))).#trainneg
cnf(i_0_115849, plain, (np__6421!=np__3|~r1_xxreal_0(np__6561,np__6421))).#trainneg
cnf(i_0_114572, plain, (np__5!=np__4|~r1_xxreal_0(np__6561,np__5))).#trainneg
cnf(i_0_114573, plain, (np__252!=np__4|~r1_xxreal_0(np__6561,np__252))).#trainneg
cnf(i_0_114575, plain, (np__6558!=np__4|~r1_xxreal_0(np__6561,np__6558))).#trainneg
cnf(i_0_114576, plain, (np__6424!=np__4|~r1_xxreal_0(np__6561,np__6424))).#trainneg
cnf(i_0_114577, plain, (np__6304!=np__4|~r1_xxreal_0(np__6561,np__6304))).#trainneg
cnf(i_0_17530, plain, (X1=k2_nat_d(X2,np__1)|r1_xxreal_0(np__1,X1)|X2!=k2_xcmplx_0(X1,np__81)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_115850, plain, (np__6422!=np__3|~r1_xxreal_0(np__6561,np__6422))).#trainneg
cnf(i_0_115851, plain, (np__6423!=np__3|~r1_xxreal_0(np__6561,np__6423))).#trainneg
cnf(i_0_115852, plain, (np__6559!=np__3|~r1_xxreal_0(np__6561,np__6559))).#trainneg
cnf(i_0_114578, plain, (np__80!=np__4|~r1_xxreal_0(np__6561,np__80))).#trainneg
cnf(i_0_114579, plain, (np__134!=np__4|~r1_xxreal_0(np__6561,np__134))).#trainneg
cnf(i_0_114580, plain, (np__232!=np__4|~r1_xxreal_0(np__6561,np__232))).#trainneg
cnf(i_0_114581, plain, (np__6400!=np__4|~r1_xxreal_0(np__6561,np__6400))).#trainneg
cnf(i_0_114582, plain, (np__6557!=np__4|~r1_xxreal_0(np__6561,np__6557))).#trainneg
cnf(i_0_125808, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__256),X2)=k4_nat_d(k3_xcmplx_0(np__256,k2_xcmplx_0(X2,X1)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_63225, plain, (k2_nat_d(k4_nat_d(X1,X2),X2)=k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(np__1,X2)),X2)|r1_xxreal_0(X2,np__1)|~v7_ordinal1(k4_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X1,X2),k1_numbers))).#trainneg
cnf(i_0_65243, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__2,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_167578, plain, (k4_nat_d(np__6,np__257)=k2_nat_d(np__6,np__257))).#trainneg
cnf(i_0_96950, plain, (k4_nat_d(k2_xcmplx_0(X1,np__136),X1)=k2_nat_d(k2_xcmplx_0(X1,np__136),X1)|r1_xxreal_0(X1,np__136)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_50594, plain, (r1_xxreal_0(X1,np__175)|r1_xxreal_0(np__175,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_143966, plain, (np__136=np__1|k2_xcmplx_0(np__257,np__136)!=np__1)).#trainneg
cnf(i_0_2791, plain, (k2_xcmplx_0(X1,np__6)=k2_xcmplx_0(np__6,X1)|~v1_xcmplx_0(k4_xcmplx_0(np__6))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_115854, plain, (np__77!=np__3|~r1_xxreal_0(np__6561,np__77))).#trainneg
cnf(i_0_115855, plain, (np__78!=np__3|~r1_xxreal_0(np__6561,np__78))).#trainneg
cnf(i_0_115856, plain, (np__24!=np__3|~r1_xxreal_0(np__6561,np__24))).#trainneg
cnf(i_0_114583, plain, (np__6344!=np__4|~r1_xxreal_0(np__6561,np__6344))).#trainneg
cnf(i_0_114586, plain, (np__22!=np__4|~r1_xxreal_0(np__6561,np__22))).#trainneg
cnf(i_0_114587, plain, (np__111!=np__4|~r1_xxreal_0(np__6561,np__111))).#trainneg
cnf(i_0_114588, plain, (np__121!=np__4|~r1_xxreal_0(np__6561,np__121))).#trainneg
cnf(i_0_114589, plain, (np__132!=np__4|~r1_xxreal_0(np__6561,np__132))).#trainneg
cnf(i_0_65150, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__81,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__81,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_112110, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__5,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__5,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_11017, plain, (X1=k2_nat_d(X2,np__1)|r1_xxreal_0(np__1,X1)|X2!=k2_xcmplx_0(X3,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_97113, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(np__136,X2)),X2)=k4_nat_d(k4_nat_1(X1,np__136),X2)|r1_xxreal_0(X2,np__136)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_103267, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__1,np__2),np__1),np__2)=np__1)).#trainneg
cnf(i_0_148420, plain, (k4_nat_d(k4_nat_1(np__1,np__256),np__257)=np__256)).#trainneg
cnf(i_0_159476, plain, (k2_xcmplx_0(np__2,np__257)=np__2|r1_xxreal_0(np__257,k2_xcmplx_0(np__2,np__257)))).#trainneg
cnf(i_0_53194, plain, (r1_xxreal_0(X1,np__256)|r1_xxreal_0(np__256,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2792, plain, (k2_xcmplx_0(X1,np__23)=k2_xcmplx_0(np__23,X1)|~v1_xcmplx_0(k4_xcmplx_0(np__23))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_115860, plain, (np__135!=np__3|~r1_xxreal_0(np__6561,np__135))).#trainneg
cnf(i_0_115861, plain, (np__6168!=np__3|~r1_xxreal_0(np__6561,np__6168))).#trainneg
cnf(i_0_115862, plain, (np__6!=np__3|~r1_xxreal_0(np__6561,np__6))).#trainneg
cnf(i_0_114590, plain, (np__133!=np__4|~r1_xxreal_0(np__6561,np__133))).#trainneg
cnf(i_0_114591, plain, (np__175!=np__4|~r1_xxreal_0(np__6561,np__175))).#trainneg
cnf(i_0_114592, plain, (np__176!=np__4|~r1_xxreal_0(np__6561,np__176))).#trainneg
cnf(i_0_114593, plain, (np__21!=np__4|~r1_xxreal_0(np__6561,np__21))).#trainneg
cnf(i_0_114594, plain, (np__6169!=np__4|~r1_xxreal_0(np__6561,np__6169))).#trainneg
cnf(i_0_148664, plain, (k4_nat_d(k3_xcmplx_0(np__256,k2_xcmplx_0(np__1,np__257)),np__257)=np__256)).#trainneg
cnf(i_0_105448, plain, (k2_nat_d(k2_xcmplx_0(np__6425,X1),np__6425)=X1|r1_xxreal_0(np__6425,X1)|~v7_ordinal1(k2_xcmplx_0(np__6425,X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_97114, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(np__136,X2)),X2)=k4_nat_d(k3_xcmplx_0(X1,np__136),X2)|r1_xxreal_0(X2,np__136)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_65139, plain, (m1_subset_1(k2_nat_d(np__81,X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_37378, plain, (k4_nat_d(k1_newton(X1,k2_xcmplx_0(X3,X2)),X4)=k2_nat_d(k1_newton(X1,k2_xcmplx_0(X3,X2)),X4)|r1_xxreal_0(X4,k1_newton(X1,X3))|r1_xxreal_0(X4,k1_newton(X1,X2))|~v1_xcmplx_0(X1)|~v7_ordinal1(k1_newton(X1,X3))|~v7_ordinal1(k1_newton(X1,X2))|~v7_ordinal1(X4)|~v7_ordinal1(X2)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_7758, plain, (k3_xcmplx_0(X1,np__8)=k3_xcmplx_0(np__8,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2793, plain, (k2_xcmplx_0(X1,np__256)=k2_xcmplx_0(np__256,X1)|~v1_xcmplx_0(k4_xcmplx_0(np__256))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_115863, plain, (np__6536!=np__3|~r1_xxreal_0(np__6561,np__6536))).#trainneg
cnf(i_0_115864, plain, (np__56!=np__3|~r1_xxreal_0(np__6561,np__56))).#trainneg
cnf(i_0_115865, plain, (np__23!=np__3|~r1_xxreal_0(np__6561,np__23))).#trainneg
cnf(i_0_170240, plain, (~r1_xxreal_0(np__257,k2_nat_d(np__6,np__257)))).#trainneg
cnf(i_0_114595, plain, (np__6421!=np__4|~r1_xxreal_0(np__6561,np__6421))).#trainneg
cnf(i_0_114596, plain, (np__6422!=np__4|~r1_xxreal_0(np__6561,np__6422))).#trainneg
cnf(i_0_114597, plain, (np__6423!=np__4|~r1_xxreal_0(np__6561,np__6423))).#trainneg
cnf(i_0_171194, plain, (k2_nat_d(k2_xcmplx_0(X2,X1),np__1)=X1|r1_xxreal_0(np__1,X1)|~v7_ordinal1(k2_xcmplx_0(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_172760, plain, (k4_nat_d(k1_newton(X1,np__3),X2)=k2_nat_d(k1_newton(X1,np__3),X2)|r1_xxreal_0(X2,k1_newton(X1,np__1))|r1_xxreal_0(X2,k1_newton(X1,np__2))|~v1_xcmplx_0(X1)|~v7_ordinal1(k1_newton(X1,np__2))|~v7_ordinal1(k1_newton(X1,np__1))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_172246, plain, (k2_nat_d(np__6425,np__6425)=np__0)).#trainneg
cnf(i_0_171843, plain, (k4_nat_1(np__1,np__256)=np__256|r1_xxreal_0(np__257,k4_nat_1(np__1,np__256))|~v7_ordinal1(k2_xcmplx_0(k4_nat_1(np__1,np__256),np__257))|~v7_ordinal1(k4_nat_1(np__1,np__256)))).#trainneg
cnf(i_0_172779, plain, (k4_nat_d(k1_newton(X1,np__2),X2)=k2_nat_d(k1_newton(X1,np__2),X2)|r1_xxreal_0(X2,k1_newton(X1,np__1))|~v1_xcmplx_0(X1)|~v7_ordinal1(k1_newton(X1,np__1))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_172766, plain, (k4_nat_d(k1_newton(X1,np__6561),X2)=k2_nat_d(k1_newton(X1,np__6561),X2)|r1_xxreal_0(X2,k1_newton(X1,np__136))|r1_xxreal_0(X2,k1_newton(X1,np__6425))|~v1_xcmplx_0(X1)|~v7_ordinal1(k1_newton(X1,np__6425))|~v7_ordinal1(k1_newton(X1,np__136))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_148656, plain, (k2_nat_d(k3_xcmplx_0(np__1,X1),np__257)=k2_nat_d(k2_xcmplx_0(X1,np__257),np__257)|r1_xxreal_0(np__257,k2_xcmplx_0(X1,np__257))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7759, plain, (k2_xcmplx_0(X1,np__8)=k2_xcmplx_0(np__8,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_37225, plain, (np__136=np__3|X1!=np__3|X1!=np__6561|~v7_ordinal1(X1))).#trainneg
cnf(i_0_166903, plain, (k4_nat_d(k3_xcmplx_0(np__3,np__2),np__257)=k2_nat_d(np__6,np__257))).#trainneg
cnf(i_0_115866, plain, (np__6289!=np__3|~r1_xxreal_0(np__6561,np__6289))).#trainneg
cnf(i_0_115867, plain, (np__231!=np__3|~r1_xxreal_0(np__6561,np__231))).#trainneg
cnf(i_0_115869, plain, (np__79!=np__3|~r1_xxreal_0(np__6561,np__79))).#trainneg
cnf(i_0_114598, plain, (np__6559!=np__4|~r1_xxreal_0(np__6561,np__6559))).#trainneg
cnf(i_0_114600, plain, (np__77!=np__4|~r1_xxreal_0(np__6561,np__77))).#trainneg
cnf(i_0_114601, plain, (np__78!=np__4|~r1_xxreal_0(np__6561,np__78))).#trainneg
cnf(i_0_114602, plain, (np__24!=np__4|~r1_xxreal_0(np__6561,np__24))).#trainneg
cnf(i_0_114606, plain, (np__135!=np__4|~r1_xxreal_0(np__6561,np__135))).#trainneg
cnf(i_0_148092, plain, (k4_nat_d(k4_nat_1(np__256,np__1),np__257)=np__256)).#trainneg
cnf(i_0_140327, plain, (k2_nat_d(k4_nat_d(X1,np__257),np__257)=k4_nat_d(k3_xcmplx_0(np__1,X1),np__257)|r1_xxreal_0(np__257,k4_nat_d(X1,np__257))|~v7_ordinal1(k4_nat_d(X1,np__257))|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X1,np__257),k1_numbers))).#trainneg
cnf(i_0_17051, plain, (X1=np__136|r1_xxreal_0(np__257,X1)|k2_xcmplx_0(X1,np__257)!=k2_xcmplx_0(np__257,np__136)|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_44008, plain, (k2_nat_d(k3_xcmplx_0(np__1,X1),X2)=k4_nat_d(X1,X2)|r1_xxreal_0(X2,np__1)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_173177, plain, (k2_nat_d(np__8,np__1)=np__0)).#trainneg
cnf(i_0_7771, plain, (k3_xcmplx_0(X1,np__3)=k3_xcmplx_0(np__3,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_174169, plain, (np__136=np__3|np__6!=np__3|np__6!=np__6561)).#trainneg
cnf(i_0_174124, plain, (np__136=np__3|np__6480!=np__3|np__6480!=np__6561)).#trainneg
cnf(i_0_174125, plain, (np__136=np__3|np__254!=np__3|np__254!=np__6561)).#trainneg
cnf(i_0_174126, plain, (np__136=np__3|np__253!=np__3|np__253!=np__6561)).#trainneg
cnf(i_0_114607, plain, (np__6168!=np__4|~r1_xxreal_0(np__6561,np__6168))).#trainneg
cnf(i_0_114608, plain, (np__6!=np__4|~r1_xxreal_0(np__6561,np__6))).#trainneg
cnf(i_0_114609, plain, (np__6536!=np__4|~r1_xxreal_0(np__6561,np__6536))).#trainneg
cnf(i_0_114610, plain, (np__56!=np__4|~r1_xxreal_0(np__6561,np__56))).#trainneg
cnf(i_0_114611, plain, (np__23!=np__4|~r1_xxreal_0(np__6561,np__23))).#trainneg
cnf(i_0_173179, plain, (k2_nat_d(np__3,np__1)=np__0)).#trainneg
cnf(i_0_173119, plain, (r1_xxreal_0(np__1,np__257))).#trainneg
cnf(i_0_174120, plain, (k2_xcmplx_0(X1,np__8)=k2_xcmplx_0(np__8,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_174328, plain, (k2_nat_d(k3_xcmplx_0(np__3,np__2),np__257)=k2_nat_d(np__6,np__257))).#trainneg
cnf(i_0_173168, plain, (k2_nat_d(np__4,np__1)=np__0)).#trainneg
cnf(i_0_175666, plain, (k3_xcmplx_0(np__3,np__2)=np__6)).#trainneg
cnf(i_0_176250, plain, (k2_nat_d(k2_xcmplx_0(np__8,X1),np__8)=X1|r1_xxreal_0(np__8,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_173420, plain, (X1!=np__0|~r1_xxreal_0(np__6425,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_176055, plain, (m1_subset_1(k2_xcmplx_0(X1,np__8),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_173117, plain, (k2_nat_d(np__0,np__1)=np__0)).#trainneg
cnf(i_0_7772, plain, (k2_xcmplx_0(X1,np__3)=k2_xcmplx_0(np__3,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_174127, plain, (np__136=np__3|np__255!=np__3|np__255!=np__6561)).#trainneg
cnf(i_0_2795, plain, (k2_xcmplx_0(X1,np__6561)=k2_xcmplx_0(np__6561,X1)|~v1_xcmplx_0(k4_xcmplx_0(np__6561))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_175685, plain, (np__6!=np__3|np__6!=np__6561)).#trainneg
cnf(i_0_174128, plain, (np__136=np__3|np__120!=np__3|np__120!=np__6561)).#trainneg
cnf(i_0_174129, plain, (np__136=np__3|np__55!=np__3|np__55!=np__6561)).#trainneg
cnf(i_0_114612, plain, (np__6289!=np__4|~r1_xxreal_0(np__6561,np__6289))).#trainneg
cnf(i_0_114613, plain, (np__231!=np__4|~r1_xxreal_0(np__6561,np__231))).#trainneg
cnf(i_0_114615, plain, (np__79!=np__4|~r1_xxreal_0(np__6561,np__79))).#trainneg
cnf(i_0_146137, plain, (np__6480!=np__1|~r1_xxreal_0(np__257,np__6480))).#trainneg
cnf(i_0_146138, plain, (np__254!=np__1|~r1_xxreal_0(np__257,np__254))).#trainneg
cnf(i_0_176059, plain, (k2_nat_d(k2_xcmplx_0(X1,np__8),np__8)=X1|r1_xxreal_0(np__8,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_173094, plain, (k2_xcmplx_0(X1,X2)=X2|r1_xxreal_0(np__1,k2_xcmplx_0(X1,X2))|r1_xxreal_0(np__1,X2)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_173160, plain, (k2_nat_d(np__6561,np__1)=np__0)).#trainneg
cnf(i_0_173196, plain, (k2_nat_d(k2_xcmplx_0(X2,X1),np__1)=X2|r1_xxreal_0(np__1,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_175184, plain, (X1=np__136|r1_xxreal_0(np__257,X1)|k2_xcmplx_0(X1,np__257)!=k2_xcmplx_0(np__257,np__136)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_170523, plain, (k2_xcmplx_0(X1,np__6)=k2_xcmplx_0(np__6,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11878, plain, (k3_xcmplx_0(X1,np__136)=k3_xcmplx_0(np__136,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_174130, plain, (np__136=np__3|np__6305!=np__3|np__6305!=np__6561)).#trainneg
cnf(i_0_2796, plain, (k2_xcmplx_0(X1,np__4)=k2_xcmplx_0(np__4,X1)|~v1_xcmplx_0(k4_xcmplx_0(np__4))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_174131, plain, (np__136=np__3|np__7!=np__3|np__7!=np__6561)).#trainneg
cnf(i_0_174133, plain, (np__136=np__3|np__5!=np__3|np__5!=np__6561)).#trainneg
cnf(i_0_174134, plain, (np__136=np__3|np__252!=np__3|np__252!=np__6561)).#trainneg
cnf(i_0_146139, plain, (np__253!=np__1|~r1_xxreal_0(np__257,np__253))).#trainneg
cnf(i_0_146140, plain, (np__255!=np__1|~r1_xxreal_0(np__257,np__255))).#trainneg
cnf(i_0_146141, plain, (np__120!=np__1|~r1_xxreal_0(np__257,np__120))).#trainneg
cnf(i_0_146142, plain, (np__55!=np__1|~r1_xxreal_0(np__257,np__55))).#trainneg
cnf(i_0_146143, plain, (np__6305!=np__1|~r1_xxreal_0(np__257,np__6305))).#trainneg
cnf(i_0_177947, plain, (k2_nat_d(np__1,np__1)=np__0)).#trainneg
cnf(i_0_177966, plain, (X1=X2|r1_xxreal_0(np__1,X2)|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_178048, plain, (k2_nat_d(k2_xcmplx_0(np__257,X1),np__1)=X1|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_178050, plain, (k2_nat_d(k2_xcmplx_0(np__136,X1),np__1)=X1|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_177962, plain, (k2_nat_d(k2_xcmplx_0(np__1,X1),np__1)=X1|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_178052, plain, (k2_nat_d(k2_xcmplx_0(np__8,X1),np__1)=X1|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_178042, plain, (k2_nat_d(k2_xcmplx_0(np__0,X1),np__1)=X1|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_178428, plain, (X1=X2|r1_xxreal_0(np__1,X1)|r1_xxreal_0(np__1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_165059, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,k2_nat_d(X1,X2)),np__3),X2)=k4_nat_d(k3_xcmplx_0(np__3,X1),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_nat_d(X1,X2),k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_66275, plain, (X1=np__0|r1_xxreal_0(X1,X3)|r1_xxreal_0(X1,X2)|~r1_xxreal_0(X1,k2_nat_d(k3_xcmplx_0(X2,X3),X1))|~v7_ordinal1(k4_nat_1(X2,X3))|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_12081, plain, (r1_xxreal_0(X1,k4_xcmplx_0(np__3))|r1_xxreal_0(k4_xcmplx_0(np__3),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_174136, plain, (np__136=np__3|np__6558!=np__3|np__6558!=np__6561)).#trainneg
cnf(i_0_174137, plain, (np__136=np__3|np__6424!=np__3|np__6424!=np__6561)).#trainneg
cnf(i_0_174138, plain, (np__136=np__3|np__6304!=np__3|np__6304!=np__6561)).#trainneg
cnf(i_0_174139, plain, (np__136=np__3|np__80!=np__3|np__80!=np__6561)).#trainneg
cnf(i_0_146144, plain, (np__7!=np__1|~r1_xxreal_0(np__257,np__7))).#trainneg
cnf(i_0_146146, plain, (np__5!=np__1|~r1_xxreal_0(np__257,np__5))).#trainneg
cnf(i_0_146147, plain, (np__252!=np__1|~r1_xxreal_0(np__257,np__252))).#trainneg
cnf(i_0_146149, plain, (np__6558!=np__1|~r1_xxreal_0(np__257,np__6558))).#trainneg
cnf(i_0_146150, plain, (np__6424!=np__1|~r1_xxreal_0(np__257,np__6424))).#trainneg
cnf(i_0_106580, plain, (m1_subset_1(k2_nat_d(np__256,X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_177298, plain, (k2_xcmplx_0(X1,np__3)=k2_xcmplx_0(np__3,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_164949, plain, (k4_nat_d(k3_xcmplx_0(np__257,np__3),np__257)=k4_nat_d(k3_xcmplx_0(np__3,np__6425),np__257))).#trainneg
cnf(i_0_21356, plain, (m1_subset_1(k2_xcmplx_0(np__257,X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_172765, plain, (k4_nat_d(k1_newton(X1,np__257),X2)=k2_nat_d(k1_newton(X1,np__257),X2)|r1_xxreal_0(X2,k1_newton(X1,np__1))|r1_xxreal_0(X2,k1_newton(X1,np__256))|~v1_xcmplx_0(X1)|~v7_ordinal1(k1_newton(X1,np__256))|~v7_ordinal1(k1_newton(X1,np__1))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14418, plain, (r1_xxreal_0(np__136,k4_xcmplx_0(X1))|r1_xxreal_0(k4_xcmplx_0(X1),np__136)|~v1_int_1(X1))).#trainneg
cnf(i_0_174140, plain, (np__136=np__3|np__134!=np__3|np__134!=np__6561)).#trainneg
cnf(i_0_174141, plain, (np__136=np__3|np__232!=np__3|np__232!=np__6561)).#trainneg
cnf(i_0_174142, plain, (np__136=np__3|np__6400!=np__3|np__6400!=np__6561)).#trainneg
cnf(i_0_174143, plain, (np__136=np__3|np__6557!=np__3|np__6557!=np__6561)).#trainneg
cnf(i_0_146151, plain, (np__6304!=np__1|~r1_xxreal_0(np__257,np__6304))).#trainneg
cnf(i_0_146152, plain, (np__80!=np__1|~r1_xxreal_0(np__257,np__80))).#trainneg
cnf(i_0_146153, plain, (np__134!=np__1|~r1_xxreal_0(np__257,np__134))).#trainneg
cnf(i_0_146154, plain, (np__232!=np__1|~r1_xxreal_0(np__257,np__232))).#trainneg
cnf(i_0_146155, plain, (np__6400!=np__1|~r1_xxreal_0(np__257,np__6400))).#trainneg
cnf(i_0_112224, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__7,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__7,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_96947, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,np__136),k2_xcmplx_0(X2,X1)),X1)=k2_nat_d(k3_xcmplx_0(X2,np__136),X1)|r1_xxreal_0(X1,np__136)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_16111, plain, (r1_xxreal_0(np__8,k4_xcmplx_0(X1))|r1_xxreal_0(k4_xcmplx_0(X1),np__8)|~v1_int_1(X1))).#trainneg
cnf(i_0_174144, plain, (np__136=np__3|np__6344!=np__3|np__6344!=np__6561)).#trainneg
cnf(i_0_174147, plain, (np__136=np__3|np__22!=np__3|np__22!=np__6561)).#trainneg
cnf(i_0_174148, plain, (np__136=np__3|np__111!=np__3|np__111!=np__6561)).#trainneg
cnf(i_0_174149, plain, (np__136=np__3|np__121!=np__3|np__121!=np__6561)).#trainneg
cnf(i_0_146156, plain, (np__6557!=np__1|~r1_xxreal_0(np__257,np__6557))).#trainneg
cnf(i_0_146157, plain, (np__6344!=np__1|~r1_xxreal_0(np__257,np__6344))).#trainneg
cnf(i_0_146160, plain, (np__22!=np__1|~r1_xxreal_0(np__257,np__22))).#trainneg
cnf(i_0_146161, plain, (np__111!=np__1|~r1_xxreal_0(np__257,np__111))).#trainneg
cnf(i_0_146162, plain, (np__121!=np__1|~r1_xxreal_0(np__257,np__121))).#trainneg
cnf(i_0_148662, plain, (k4_nat_d(k3_xcmplx_0(np__2,k2_xcmplx_0(np__1,np__257)),np__257)=np__2)).#trainneg
cnf(i_0_48609, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,np__257),X2),np__257)=k3_xcmplx_0(X1,k2_nat_d(X2,np__257))|r1_xxreal_0(np__257,k3_xcmplx_0(X1,k2_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k3_xcmplx_0(X1,k2_nat_d(X2,np__257)))|~v7_ordinal1(k2_nat_d(X2,np__257))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_71809, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__2,X1),k2_nat_d(np__3,X1)),X1)=k4_nat_d(np__6,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_179685, plain, (k4_nat_d(k3_xcmplx_0(np__3,np__257),np__257)=np__0)).#trainneg
cnf(i_0_179530, plain, (X1=X2|r1_xxreal_0(np__1,X2)|r1_xxreal_0(np__1,X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_180404, plain, (k2_nat_d(k2_xcmplx_0(np__3,X1),np__3)=X1|r1_xxreal_0(np__3,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16642, plain, (r1_xxreal_0(np__3,k4_xcmplx_0(X1))|r1_xxreal_0(k4_xcmplx_0(X1),np__3)|~v1_int_1(X1))).#trainneg
cnf(i_0_174150, plain, (np__136=np__3|np__132!=np__3|np__132!=np__6561)).#trainneg
cnf(i_0_174151, plain, (np__136=np__3|np__133!=np__3|np__133!=np__6561)).#trainneg
cnf(i_0_174152, plain, (np__136=np__3|np__175!=np__3|np__175!=np__6561)).#trainneg
cnf(i_0_174153, plain, (np__136=np__3|np__176!=np__3|np__176!=np__6561)).#trainneg
cnf(i_0_146163, plain, (np__132!=np__1|~r1_xxreal_0(np__257,np__132))).#trainneg
cnf(i_0_146164, plain, (np__133!=np__1|~r1_xxreal_0(np__257,np__133))).#trainneg
cnf(i_0_146165, plain, (np__175!=np__1|~r1_xxreal_0(np__257,np__175))).#trainneg
cnf(i_0_146166, plain, (np__176!=np__1|~r1_xxreal_0(np__257,np__176))).#trainneg
cnf(i_0_146167, plain, (np__21!=np__1|~r1_xxreal_0(np__257,np__21))).#trainneg
cnf(i_0_179795, plain, (k4_nat_d(k3_xcmplx_0(np__3,np__6425),np__257)=k4_nat_d(k3_xcmplx_0(np__3,np__257),np__257))).#trainneg
cnf(i_0_180476, plain, (k2_nat_d(k2_xcmplx_0(np__3,X1),np__1)=X1|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_180397, plain, (m1_subset_1(k2_xcmplx_0(np__3,X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_106596, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__256,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(np__256,k2_xcmplx_0(X1,X2)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_112247, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__7,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__7,X2),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64593, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__4,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__4),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_182324, plain, (X1=np__0|r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_164949, plain, (k4_nat_d(k3_xcmplx_0(np__3,np__6425),np__257)=np__0)).#trainneg
cnf(i_0_183821, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__257,np__1),np__4),np__257)=np__4)).#trainneg
cnf(i_0_182387, plain, (X1=k4_pepin(X2)|r1_xxreal_0(np__1,k4_pepin(X2))|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_180135, plain, (r1_xxreal_0(k4_xcmplx_0(np__3),X1)|r1_xxreal_0(X1,k4_xcmplx_0(np__3))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_174154, plain, (np__136=np__3|np__21!=np__3|np__21!=np__6561)).#trainneg
cnf(i_0_164997, plain, (k4_nat_d(k3_xcmplx_0(np__257,np__3),np__257)=np__0)).#trainneg
cnf(i_0_174155, plain, (np__136=np__3|np__6169!=np__3|np__6169!=np__6561)).#trainneg
cnf(i_0_174156, plain, (np__136=np__3|np__6421!=np__3|np__6421!=np__6561)).#trainneg
cnf(i_0_174157, plain, (np__136=np__3|np__6422!=np__3|np__6422!=np__6561)).#trainneg
cnf(i_0_146168, plain, (np__6169!=np__1|~r1_xxreal_0(np__257,np__6169))).#trainneg
cnf(i_0_146169, plain, (np__6421!=np__1|~r1_xxreal_0(np__257,np__6421))).#trainneg
cnf(i_0_146170, plain, (np__6422!=np__1|~r1_xxreal_0(np__257,np__6422))).#trainneg
cnf(i_0_146171, plain, (np__6423!=np__1|~r1_xxreal_0(np__257,np__6423))).#trainneg
cnf(i_0_146172, plain, (np__6559!=np__1|~r1_xxreal_0(np__257,np__6559))).#trainneg
cnf(i_0_182180, plain, (k2_nat_d(k3_xcmplx_0(np__3,np__257),np__257)=np__0|~v7_ordinal1(k3_xcmplx_0(np__3,np__257)))).#trainneg
cnf(i_0_123352, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__7,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(np__7,k2_xcmplx_0(X1,X2)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_110070, plain, (X1=np__136|r1_xxreal_0(np__6561,X1)|X1!=k2_xcmplx_0(np__6561,np__136)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_154439, plain, (k4_nat_d(k3_xcmplx_0(X1,X2),np__257)=k3_xcmplx_0(X1,k4_nat_d(X2,np__257))|r1_xxreal_0(np__257,k3_xcmplx_0(X1,k4_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k3_xcmplx_0(X1,k4_nat_d(X2,np__257)))|~v7_ordinal1(k4_nat_d(X2,np__257))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_17074, plain, (X1=np__3|r1_xxreal_0(np__257,X1)|k2_xcmplx_0(X1,np__257)!=np__3|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_158163, plain, (k4_nat_1(X1,k2_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,k2_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|r1_xxreal_0(np__257,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_1(X1,k2_nat_d(X2,np__257)),k5_numbers))).#trainneg
cnf(i_0_79502, plain, (k4_nat_d(k4_nat_1(np__81,X1),np__257)=k4_nat_d(k3_xcmplx_0(np__81,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_174158, plain, (np__136=np__3|np__6423!=np__3|np__6423!=np__6561)).#trainneg
cnf(i_0_174159, plain, (np__136=np__3|np__6559!=np__3|np__6559!=np__6561)).#trainneg
cnf(i_0_174160, plain, (np__136=np__3|np__6560!=np__3|np__6560!=np__6561)).#trainneg
cnf(i_0_174161, plain, (np__136=np__3|np__77!=np__3|np__77!=np__6561)).#trainneg
cnf(i_0_146173, plain, (np__6560!=np__1|~r1_xxreal_0(np__257,np__6560))).#trainneg
cnf(i_0_146174, plain, (np__77!=np__1|~r1_xxreal_0(np__257,np__77))).#trainneg
cnf(i_0_146175, plain, (np__78!=np__1|~r1_xxreal_0(np__257,np__78))).#trainneg
cnf(i_0_146176, plain, (np__24!=np__1|~r1_xxreal_0(np__257,np__24))).#trainneg
cnf(i_0_146180, plain, (np__135!=np__1|~r1_xxreal_0(np__257,np__135))).#trainneg
cnf(i_0_104554, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__3),X2)=k4_nat_d(k3_xcmplx_0(np__3,k2_xcmplx_0(X2,X1)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_185583, plain, (X1=np__3|r1_xxreal_0(np__257,X1)|k2_xcmplx_0(X1,np__257)!=np__3|~v7_ordinal1(X1))).#trainneg
cnf(i_0_184445, plain, (k2_nat_d(k3_xcmplx_0(np__257,np__3),np__257)=np__0|~v7_ordinal1(k3_xcmplx_0(np__257,np__3)))).#trainneg
cnf(i_0_4259, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_174162, plain, (np__136=np__3|np__78!=np__3|np__78!=np__6561)).#trainneg
cnf(i_0_174163, plain, (np__136=np__3|np__24!=np__3|np__24!=np__6561)).#trainneg
cnf(i_0_174165, plain, (np__136=np__3|np__25!=np__3|np__25!=np__6561)).#trainneg
cnf(i_0_174167, plain, (np__136=np__3|np__135!=np__3|np__135!=np__6561)).#trainneg
cnf(i_0_146181, plain, (np__6168!=np__1|~r1_xxreal_0(np__257,np__6168))).#trainneg
cnf(i_0_146182, plain, (np__6!=np__1|~r1_xxreal_0(np__257,np__6))).#trainneg
cnf(i_0_146183, plain, (np__6536!=np__1|~r1_xxreal_0(np__257,np__6536))).#trainneg
cnf(i_0_146184, plain, (np__56!=np__1|~r1_xxreal_0(np__257,np__56))).#trainneg
cnf(i_0_146185, plain, (np__23!=np__1|~r1_xxreal_0(np__257,np__23))).#trainneg
cnf(i_0_185733, plain, (k4_nat_d(k3_xcmplx_0(X1,X2),np__257)=k4_nat_1(X1,X2)|r1_xxreal_0(np__257,k4_nat_1(X1,X2))|r1_xxreal_0(np__257,X2)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_1(X1,X2),k5_numbers))).#trainneg
cnf(i_0_184088, plain, (k2_xcmplx_0(k3_xcmplx_0(np__3,np__6425),np__257)=np__0|r1_xxreal_0(np__257,k2_xcmplx_0(k3_xcmplx_0(np__3,np__6425),np__257))|r1_xxreal_0(np__257,k3_xcmplx_0(np__3,np__6425))|~v7_ordinal1(k2_xcmplx_0(k3_xcmplx_0(np__3,np__6425),np__257))|~v7_ordinal1(k3_xcmplx_0(np__3,np__6425)))).#trainneg
cnf(i_0_183984, plain, (k2_nat_d(k3_xcmplx_0(np__3,np__6425),np__257)=np__0|~v7_ordinal1(k3_xcmplx_0(np__3,np__6425)))).#trainneg
cnf(i_0_184425, plain, (r1_xxreal_0(X1,k4_xcmplx_0(np__3))|r1_xxreal_0(k4_xcmplx_0(np__3),X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_185674, plain, (k4_nat_d(k3_xcmplx_0(X1,np__1),np__257)=k4_nat_1(X1,np__1)|r1_xxreal_0(np__257,k4_nat_1(X1,np__1))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_1(X1,np__1),k5_numbers))).#trainneg
cnf(i_0_185720, plain, (k4_nat_d(k3_xcmplx_0(X1,np__81),np__257)=k4_nat_1(X1,np__81)|r1_xxreal_0(np__257,k4_nat_1(X1,np__81))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_1(X1,np__81),k5_numbers))).#trainneg
cnf(i_0_185546, plain, (k4_nat_d(k3_xcmplx_0(np__1,X1),np__257)=k4_nat_d(X1,np__257)|r1_xxreal_0(np__257,k4_nat_d(X1,np__257))|~v7_ordinal1(k4_nat_d(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_185722, plain, (k4_nat_d(k3_xcmplx_0(X1,np__136),np__257)=k4_nat_1(X1,np__136)|r1_xxreal_0(np__257,k4_nat_1(X1,np__136))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_1(X1,np__136),k5_numbers))).#trainneg
cnf(i_0_4381, plain, (r1_xxreal_0(X1,k4_xcmplx_0(X2))|r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_174168, plain, (np__136=np__3|np__6168!=np__3|np__6168!=np__6561)).#trainneg
cnf(i_0_174170, plain, (np__136=np__3|np__6536!=np__3|np__6536!=np__6561)).#trainneg
cnf(i_0_174171, plain, (np__136=np__3|np__56!=np__3|np__56!=np__6561)).#trainneg
cnf(i_0_174172, plain, (np__136=np__3|np__23!=np__3|np__23!=np__6561)).#trainneg
cnf(i_0_146186, plain, (np__6289!=np__1|~r1_xxreal_0(np__257,np__6289))).#trainneg
cnf(i_0_146187, plain, (np__231!=np__1|~r1_xxreal_0(np__257,np__231))).#trainneg
cnf(i_0_146189, plain, (np__79!=np__1|~r1_xxreal_0(np__257,np__79))).#trainneg
cnf(i_0_158771, plain, (k4_nat_1(X1,k2_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(X2,X1),np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,k2_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k4_nat_1(X1,k2_nat_d(X2,np__257)),k5_numbers))).#trainneg
cnf(i_0_148452, plain, (k4_nat_d(k4_nat_1(np__2,np__1),np__257)=np__2)).#trainneg
cnf(i_0_85645, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,k4_nat_d(X2,X1)),np__136),X1)=k4_nat_d(k3_xcmplx_0(X2,np__136),X1)|r1_xxreal_0(X1,k4_nat_d(X2,X1))|~v7_ordinal1(k4_nat_d(X2,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_148695, plain, (k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(np__1,np__257)),np__257)=k2_nat_d(X1,np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)|~m1_subset_1(k2_xcmplx_0(np__1,np__257),k1_numbers))).#trainneg
cnf(i_0_64440, plain, (k4_nat_d(k4_nat_1(np__3,X1),np__257)=k4_nat_d(k3_xcmplx_0(np__3,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4382, plain, (r1_xxreal_0(X1,k4_xcmplx_0(X2))|r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_int_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_174173, plain, (np__136=np__3|np__6289!=np__3|np__6289!=np__6561)).#trainneg
cnf(i_0_174174, plain, (np__136=np__3|np__231!=np__3|np__231!=np__6561)).#trainneg
cnf(i_0_174176, plain, (np__136=np__3|np__79!=np__3|np__79!=np__6561)).#trainneg
cnf(i_0_177055, plain, (np__8!=np__0|~r1_xxreal_0(np__6425,np__8))).#trainneg
cnf(i_0_146676, plain, (np__6480!=np__2|~r1_xxreal_0(np__257,np__6480))).#trainneg
cnf(i_0_146677, plain, (np__254!=np__2|~r1_xxreal_0(np__257,np__254))).#trainneg
cnf(i_0_146678, plain, (np__253!=np__2|~r1_xxreal_0(np__257,np__253))).#trainneg
cnf(i_0_186510, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_188415, plain, (k2_nat_d(k4_nat_1(np__2,np__1),np__257)=np__2|~v7_ordinal1(k4_nat_1(np__2,np__1)))).#trainneg
cnf(i_0_187506, plain, (k4_nat_1(X1,np__1)=np__136|r1_xxreal_0(np__257,k4_nat_1(X1,np__1))|r1_xxreal_0(np__257,X1)|X1!=np__136|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_1(X1,np__1),k5_numbers))).#trainneg
cnf(i_0_186504, plain, (r1_xxreal_0(X1,np__257)|r1_xxreal_0(np__257,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_186853, plain, (k4_nat_1(np__1,X1)=np__136|r1_xxreal_0(np__257,k4_nat_1(np__1,X1))|r1_xxreal_0(np__257,X1)|X1!=np__136|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_1(np__1,X1),k5_numbers))).#trainneg
cnf(i_0_186790, plain, (k4_nat_1(np__1,X1)=k4_nat_d(X1,np__257)|r1_xxreal_0(np__257,k4_nat_1(np__1,X1))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_1(np__1,X1),k5_numbers))).#trainneg
cnf(i_0_186857, plain, (k4_nat_1(np__1,X1)=k2_nat_d(X1,np__257)|r1_xxreal_0(np__257,k4_nat_1(np__1,X1))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_1(np__1,X1),k5_numbers))).#trainneg
cnf(i_0_186822, plain, (k4_nat_1(X1,np__1)=k4_nat_d(X1,np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,np__1))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_1(X1,np__1),k5_numbers))).#trainneg
cnf(i_0_4399, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_177306, plain, (np__6480=np__0|np__6480!=np__1|np__6480!=np__257)).#trainneg
cnf(i_0_2810, plain, (k2_nat_d(k2_xcmplx_0(k3_xcmplx_0(X2,k3_xcmplx_0(X3,X4)),X1),k3_xcmplx_0(X2,X3))=X1|r1_xxreal_0(k3_xcmplx_0(X2,X3),X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v7_ordinal1(k2_xcmplx_0(k3_xcmplx_0(X2,k3_xcmplx_0(X3,X4)),X1))|~v7_ordinal1(k3_xcmplx_0(X2,X3))|~v7_ordinal1(X1)|~v7_ordinal1(X4))).#trainneg
cnf(i_0_177307, plain, (np__254=np__0|np__254!=np__1|np__254!=np__257)).#trainneg
cnf(i_0_178172, plain, (np__253=np__0|np__253!=np__1|np__253!=np__257)).#trainneg
cnf(i_0_178173, plain, (np__255=np__0|np__255!=np__1|np__255!=np__257)).#trainneg
cnf(i_0_146679, plain, (np__255!=np__2|~r1_xxreal_0(np__257,np__255))).#trainneg
cnf(i_0_146680, plain, (np__120!=np__2|~r1_xxreal_0(np__257,np__120))).#trainneg
cnf(i_0_146681, plain, (np__55!=np__2|~r1_xxreal_0(np__257,np__55))).#trainneg
cnf(i_0_146682, plain, (np__6305!=np__2|~r1_xxreal_0(np__257,np__6305))).#trainneg
cnf(i_0_146683, plain, (np__7!=np__2|~r1_xxreal_0(np__257,np__7))).#trainneg
cnf(i_0_186990, plain, (k4_nat_d(k3_xcmplx_0(X2,X1),np__257)=k4_nat_1(X1,X2)|r1_xxreal_0(np__257,k4_nat_1(X1,X2))|r1_xxreal_0(np__257,X1)|r1_xxreal_0(np__257,X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_1(X1,X2),k5_numbers))).#trainneg
cnf(i_0_65808, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__257,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(np__257,k2_xcmplx_0(X1,X2)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_27958, plain, (m1_subset_1(k3_xcmplx_0(X1,np__257),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4933, plain, (r1_xxreal_0(X1,k4_xcmplx_0(X2))|r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_178174, plain, (np__120=np__0|np__120!=np__1|np__120!=np__257)).#trainneg
cnf(i_0_178175, plain, (np__55=np__0|np__55!=np__1|np__55!=np__257)).#trainneg
cnf(i_0_178176, plain, (np__6305=np__0|np__6305!=np__1|np__6305!=np__257)).#trainneg
cnf(i_0_180138, plain, (np__7=np__0|np__7!=np__1|np__7!=np__257)).#trainneg
cnf(i_0_146685, plain, (np__5!=np__2|~r1_xxreal_0(np__257,np__5))).#trainneg
cnf(i_0_146686, plain, (np__252!=np__2|~r1_xxreal_0(np__257,np__252))).#trainneg
cnf(i_0_146688, plain, (np__6558!=np__2|~r1_xxreal_0(np__257,np__6558))).#trainneg
cnf(i_0_146689, plain, (np__6424!=np__2|~r1_xxreal_0(np__257,np__6424))).#trainneg
cnf(i_0_146690, plain, (np__6304!=np__2|~r1_xxreal_0(np__257,np__6304))).#trainneg
cnf(i_0_26595, plain, (k2_xcmplx_0(X1,np__6425)=X1|r1_xxreal_0(np__257,k2_xcmplx_0(X1,np__6425))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_111224, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__6,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__6),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_48367, plain, (k4_nat_d(k3_xcmplx_0(X1,np__1),np__257)=k4_nat_d(k2_xcmplx_0(X1,np__257),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)|~m1_subset_1(k2_xcmplx_0(X1,np__257),k1_numbers))).#trainneg
cnf(i_0_186507, plain, (r1_xxreal_0(X1,np__136)|r1_xxreal_0(np__136,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_11720, plain, (X1=k2_nat_d(X2,np__2)|r1_xxreal_0(np__2,X1)|X2!=k2_xcmplx_0(X1,np__2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4934, plain, (r1_xxreal_0(X1,k4_xcmplx_0(X2))|r1_xxreal_0(k4_xcmplx_0(X2),X1)|~v1_int_1(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_180140, plain, (np__252=np__0|np__252!=np__1|np__252!=np__257)).#trainneg
cnf(i_0_180141, plain, (np__6558=np__0|np__6558!=np__1|np__6558!=np__257)).#trainneg
cnf(i_0_180142, plain, (np__6424=np__0|np__6424!=np__1|np__6424!=np__257)).#trainneg
cnf(i_0_180864, plain, (np__6304=np__0|np__6304!=np__1|np__6304!=np__257)).#trainneg
cnf(i_0_146691, plain, (np__80!=np__2|~r1_xxreal_0(np__257,np__80))).#trainneg
cnf(i_0_146692, plain, (np__134!=np__2|~r1_xxreal_0(np__257,np__134))).#trainneg
cnf(i_0_146693, plain, (np__232!=np__2|~r1_xxreal_0(np__257,np__232))).#trainneg
cnf(i_0_146694, plain, (np__6400!=np__2|~r1_xxreal_0(np__257,np__6400))).#trainneg
cnf(i_0_146695, plain, (np__6557!=np__2|~r1_xxreal_0(np__257,np__6557))).#trainneg
cnf(i_0_180181, plain, (m1_subset_1(k2_xcmplx_0(X1,np__3),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_97088, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__136,X1),np__1),X1)=k4_nat_d(k2_xcmplx_0(np__136,X1),X1)|r1_xxreal_0(X1,np__136)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_190067, plain, (r1_xxreal_0(np__2,np__257))).#trainneg
cnf(i_0_190082, plain, (r1_xxreal_0(np__257,np__6561))).#trainneg
cnf(i_0_4951, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_180865, plain, (np__80=np__0|np__80!=np__1|np__80!=np__257)).#trainneg
cnf(i_0_2815, plain, (X1=k2_nat_d(X2,k3_xcmplx_0(X3,np__25))|r1_xxreal_0(k3_xcmplx_0(X3,np__25),X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(X3,np__6425),X1)|~v1_xcmplx_0(np__25)|~v1_xcmplx_0(X3)|~v7_ordinal1(k3_xcmplx_0(X3,np__25))|~v7_ordinal1(np__257)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_180866, plain, (np__134=np__0|np__134!=np__1|np__134!=np__257)).#trainneg
cnf(i_0_180867, plain, (np__232=np__0|np__232!=np__1|np__232!=np__257)).#trainneg
cnf(i_0_180868, plain, (np__6400=np__0|np__6400!=np__1|np__6400!=np__257)).#trainneg
cnf(i_0_146696, plain, (np__6344!=np__2|~r1_xxreal_0(np__257,np__6344))).#trainneg
cnf(i_0_146699, plain, (np__22!=np__2|~r1_xxreal_0(np__257,np__22))).#trainneg
cnf(i_0_146700, plain, (np__111!=np__2|~r1_xxreal_0(np__257,np__111))).#trainneg
cnf(i_0_146701, plain, (np__121!=np__2|~r1_xxreal_0(np__257,np__121))).#trainneg
cnf(i_0_146702, plain, (np__132!=np__2|~r1_xxreal_0(np__257,np__132))).#trainneg
cnf(i_0_191450, plain, (m1_subset_1(k3_xcmplx_0(np__257,X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_185746, plain, (k4_nat_1(k2_nat_d(X2,np__257),X1)=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|r1_xxreal_0(np__257,k4_nat_1(k2_nat_d(X2,np__257),X1))|r1_xxreal_0(np__257,X2)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_1(k2_nat_d(X2,np__257),X1),k5_numbers)|~m1_subset_1(k2_nat_d(X2,np__257),k5_numbers))).#trainneg
cnf(i_0_97089, plain, (k4_nat_d(k2_xcmplx_0(np__136,X1),X1)=k2_nat_d(k2_xcmplx_0(np__136,X1),X1)|r1_xxreal_0(X1,np__136)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_65154, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__81,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(np__81,k2_xcmplx_0(X1,X2)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_184356, plain, (k4_pepin(X1)=k4_pepin(X2)|r1_xxreal_0(np__1,k4_pepin(X2))|r1_xxreal_0(np__1,k4_pepin(X1))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_181489, plain, (np__6557=np__0|np__6557!=np__1|np__6557!=np__257)).#trainneg
cnf(i_0_2816, plain, (X1=k2_nat_d(X2,k3_xcmplx_0(X3,np__2))|r1_xxreal_0(k3_xcmplx_0(X3,np__2),X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(X3,np__4),X1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X3)|~v7_ordinal1(k3_xcmplx_0(X3,np__2))|~v7_ordinal1(np__2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_181490, plain, (np__6344=np__0|np__6344!=np__1|np__6344!=np__257)).#trainneg
cnf(i_0_181491, plain, (np__22=np__0|np__22!=np__1|np__22!=np__257)).#trainneg
cnf(i_0_181492, plain, (np__111=np__0|np__111!=np__1|np__111!=np__257)).#trainneg
cnf(i_0_146703, plain, (np__133!=np__2|~r1_xxreal_0(np__257,np__133))).#trainneg
cnf(i_0_146704, plain, (np__175!=np__2|~r1_xxreal_0(np__257,np__175))).#trainneg
cnf(i_0_146705, plain, (np__176!=np__2|~r1_xxreal_0(np__257,np__176))).#trainneg
cnf(i_0_146706, plain, (np__21!=np__2|~r1_xxreal_0(np__257,np__21))).#trainneg
cnf(i_0_146707, plain, (np__6169!=np__2|~r1_xxreal_0(np__257,np__6169))).#trainneg
cnf(i_0_47464, plain, (k4_nat_d(k3_xcmplx_0(X1,k2_xcmplx_0(X2,np__257)),np__257)=k4_nat_d(k3_xcmplx_0(X2,X1),np__257)|r1_xxreal_0(np__257,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_170268, plain, (k4_nat_d(k2_nat_d(np__6,np__257),np__257)=k2_nat_d(np__6,np__257)|~v7_ordinal1(k2_nat_d(np__6,np__257)))).#trainneg
cnf(i_0_36092, plain, (X1=k2_nat_d(X2,np__257)|r1_xxreal_0(np__257,X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(X3,np__257),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_112096, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__5,np__257),X1),np__257)=k4_nat_d(k3_xcmplx_0(np__5,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_91811, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(k2_nat_d(X1,X2),X2),np__1),X2)=k4_nat_d(k3_xcmplx_0(X1,np__1),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_184355, plain, (k4_nat_d(X1,X2)=k4_pepin(X3)|r1_xxreal_0(np__1,k4_nat_d(X1,X2))|r1_xxreal_0(np__1,k4_pepin(X3))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_181493, plain, (np__121=np__0|np__121!=np__1|np__121!=np__257)).#trainneg
cnf(i_0_2817, plain, (X1=k2_nat_d(X2,k3_xcmplx_0(X3,np__2))|r1_xxreal_0(k3_xcmplx_0(X3,np__2),X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(X3,np__8),X1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X3)|~v7_ordinal1(k3_xcmplx_0(X3,np__2))|~v7_ordinal1(np__4)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_182634, plain, (np__132=np__0|np__132!=np__1|np__132!=np__257)).#trainneg
cnf(i_0_182635, plain, (np__133=np__0|np__133!=np__1|np__133!=np__257)).#trainneg
cnf(i_0_182636, plain, (np__175=np__0|np__175!=np__1|np__175!=np__257)).#trainneg
cnf(i_0_146708, plain, (np__6421!=np__2|~r1_xxreal_0(np__257,np__6421))).#trainneg
cnf(i_0_146709, plain, (np__6422!=np__2|~r1_xxreal_0(np__257,np__6422))).#trainneg
cnf(i_0_146710, plain, (np__6423!=np__2|~r1_xxreal_0(np__257,np__6423))).#trainneg
cnf(i_0_146711, plain, (np__6559!=np__2|~r1_xxreal_0(np__257,np__6559))).#trainneg
cnf(i_0_146712, plain, (np__6560!=np__2|~r1_xxreal_0(np__257,np__6560))).#trainneg
cnf(i_0_152501, plain, (k2_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__1,np__257),np__136),np__257)=np__136|~v7_ordinal1(k3_xcmplx_0(k2_xcmplx_0(np__1,np__257),np__136)))).#trainneg
cnf(i_0_65805, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__257,X1),k2_nat_d(X2,X1)),X1)=k4_nat_d(k3_xcmplx_0(X2,np__257),X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_57291, plain, (k4_nat_d(np__6561,np__6561)=np__0)).#trainneg
cnf(i_0_194440, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__2,np__257),np__1),np__257)=np__2)).#trainneg
cnf(i_0_194482, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__1,np__257),np__1),np__257)=np__1)).#trainneg
cnf(i_0_194538, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__81,np__257),np__1),np__257)=np__81)).#trainneg
cnf(i_0_185726, plain, (k4_nat_d(k3_xcmplx_0(X1,np__3),np__257)=k4_nat_1(X1,np__3)|r1_xxreal_0(np__257,k4_nat_1(X1,np__3))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_1(X1,np__3),k5_numbers))).#trainneg
cnf(i_0_154603, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__1,np__257),X1),np__257)=k2_nat_d(k3_xcmplx_0(np__1,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(np__1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_193422, plain, (k4_pepin(X1)=k4_pepin(np__136)|r1_xxreal_0(np__1,k4_pepin(np__136))|r1_xxreal_0(np__1,k4_pepin(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_182637, plain, (np__176=np__0|np__176!=np__1|np__176!=np__257)).#trainneg
cnf(i_0_2818, plain, (X1=k2_nat_d(X2,k3_xcmplx_0(X3,np__2))|r1_xxreal_0(k3_xcmplx_0(X3,np__2),X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(X3,np__6),X1)|~v1_xcmplx_0(np__2)|~v1_xcmplx_0(X3)|~v7_ordinal1(k3_xcmplx_0(X3,np__2))|~v7_ordinal1(np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_182638, plain, (np__21=np__0|np__21!=np__1|np__21!=np__257)).#trainneg
cnf(i_0_184577, plain, (np__6169=np__0|np__6169!=np__1|np__6169!=np__257)).#trainneg
cnf(i_0_184578, plain, (np__6421=np__0|np__6421!=np__1|np__6421!=np__257)).#trainneg
cnf(i_0_146713, plain, (np__77!=np__2|~r1_xxreal_0(np__257,np__77))).#trainneg
cnf(i_0_146714, plain, (np__78!=np__2|~r1_xxreal_0(np__257,np__78))).#trainneg
cnf(i_0_146715, plain, (np__24!=np__2|~r1_xxreal_0(np__257,np__24))).#trainneg
cnf(i_0_146719, plain, (np__135!=np__2|~r1_xxreal_0(np__257,np__135))).#trainneg
cnf(i_0_146720, plain, (np__6168!=np__2|~r1_xxreal_0(np__257,np__6168))).#trainneg
cnf(i_0_68013, plain, (k4_nat_d(k3_xcmplx_0(np__136,k2_xcmplx_0(X2,k2_nat_d(X1,X2))),X2)=k4_nat_d(k3_xcmplx_0(X1,np__136),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_145754, plain, (k4_nat_1(X1,k4_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__257,X1),X2),np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,k4_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(k4_nat_1(X1,k4_nat_d(X2,np__257)),np__257))|~v7_ordinal1(k4_nat_1(X1,k4_nat_d(X2,np__257)))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_185006, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__7),X1)=k4_nat_d(k3_xcmplx_0(np__7,k2_xcmplx_0(X1,X2)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64481, plain, (k4_nat_d(k3_xcmplx_0(k2_nat_d(np__3,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__3),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_64105, plain, (k4_nat_d(k4_nat_1(k2_nat_d(np__136,X1),X2),X1)=k4_nat_d(k3_xcmplx_0(np__136,k2_xcmplx_0(X1,X2)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_145819, plain, (k3_xcmplx_0(k2_nat_d(X1,np__257),k4_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|r1_xxreal_0(np__257,k3_xcmplx_0(k2_nat_d(X1,np__257),k4_nat_d(X2,np__257)))|~v7_ordinal1(k2_xcmplx_0(k3_xcmplx_0(k2_nat_d(X1,np__257),k4_nat_d(X2,np__257)),np__257))|~v7_ordinal1(k3_xcmplx_0(k2_nat_d(X1,np__257),k4_nat_d(X2,np__257)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X2,np__257),k1_numbers))).#trainneg
cnf(i_0_195648, plain, (k4_nat_d(k2_xcmplx_0(np__2,np__257),np__257)=np__2)).#trainneg
cnf(i_0_195994, plain, (k4_nat_d(k2_xcmplx_0(np__81,np__257),np__257)=np__81)).#trainneg
cnf(i_0_193423, plain, (k4_pepin(X1)=k4_pepin(np__8)|r1_xxreal_0(np__1,k4_pepin(np__8))|r1_xxreal_0(np__1,k4_pepin(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_184579, plain, (np__6422=np__0|np__6422!=np__1|np__6422!=np__257)).#trainneg
cnf(i_0_2819, plain, (X1=k2_nat_d(X2,k3_xcmplx_0(X3,np__6))|r1_xxreal_0(k3_xcmplx_0(X3,np__6),X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(X3,np__24),X1)|~v1_xcmplx_0(np__6)|~v1_xcmplx_0(X3)|~v7_ordinal1(k3_xcmplx_0(X3,np__6))|~v7_ordinal1(np__4)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_184580, plain, (np__6423=np__0|np__6423!=np__1|np__6423!=np__257)).#trainneg
cnf(i_0_184581, plain, (np__6559=np__0|np__6559!=np__1|np__6559!=np__257)).#trainneg
cnf(i_0_185938, plain, (np__6560=np__0|np__6560!=np__1|np__6560!=np__257)).#trainneg
cnf(i_0_146721, plain, (np__6!=np__2|~r1_xxreal_0(np__257,np__6))).#trainneg
cnf(i_0_146722, plain, (np__6536!=np__2|~r1_xxreal_0(np__257,np__6536))).#trainneg
cnf(i_0_146723, plain, (np__56!=np__2|~r1_xxreal_0(np__257,np__56))).#trainneg
cnf(i_0_146724, plain, (np__23!=np__2|~r1_xxreal_0(np__257,np__23))).#trainneg
cnf(i_0_146725, plain, (np__6289!=np__2|~r1_xxreal_0(np__257,np__6289))).#trainneg
cnf(i_0_195809, plain, (k2_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__1,np__257),np__1),np__257)=np__1|r1_xxreal_0(np__257,k2_xcmplx_0(np__1,np__257))|~v7_ordinal1(k2_xcmplx_0(np__1,np__257)))).#trainneg
cnf(i_0_195861, plain, (k2_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__81,np__257),np__1),np__257)=np__81|~v7_ordinal1(k3_xcmplx_0(k2_xcmplx_0(np__81,np__257),np__1)))).#trainneg
cnf(i_0_195993, plain, (k4_nat_d(k2_xcmplx_0(np__257,np__81),np__257)=np__81)).#trainneg
cnf(i_0_195515, plain, (k2_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__2,np__257),np__1),np__257)=np__2|~v7_ordinal1(k3_xcmplx_0(k2_xcmplx_0(np__2,np__257),np__1)))).#trainneg
cnf(i_0_195647, plain, (k4_nat_d(k2_xcmplx_0(np__257,np__2),np__257)=np__2)).#trainneg
cnf(i_0_195685, plain, (k2_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__1,np__257),np__1),np__257)=np__1|~v7_ordinal1(k3_xcmplx_0(k2_xcmplx_0(np__1,np__257),np__1)))).#trainneg
cnf(i_0_197511, plain, (k4_nat_1(X1,k4_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__257,X1),X2),np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,k4_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k4_nat_1(X1,k4_nat_d(X2,np__257)))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k2_xcmplx_0(k4_nat_1(X1,k4_nat_d(X2,np__257)),np__257),k5_numbers))).#trainneg
cnf(i_0_194681, plain, (k4_nat_d(X1,X2)=k4_pepin(np__136)|r1_xxreal_0(np__1,k4_nat_d(X1,X2))|r1_xxreal_0(np__1,k4_pepin(np__136))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_185939, plain, (np__77=np__0|np__77!=np__1|np__77!=np__257)).#trainneg
cnf(i_0_2821, plain, (X1=k2_nat_d(X2,k3_xcmplx_0(X3,np__1))|r1_xxreal_0(k3_xcmplx_0(X3,np__1),X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(X3,np__81),X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X3)|~v7_ordinal1(k3_xcmplx_0(X3,np__1))|~v7_ordinal1(np__81)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_185940, plain, (np__78=np__0|np__78!=np__1|np__78!=np__257)).#trainneg
cnf(i_0_185941, plain, (np__24=np__0|np__24!=np__1|np__24!=np__257)).#trainneg
cnf(i_0_185942, plain, (np__135=np__0|np__135!=np__1|np__135!=np__257)).#trainneg
cnf(i_0_146726, plain, (np__231!=np__2|~r1_xxreal_0(np__257,np__231))).#trainneg
cnf(i_0_146728, plain, (np__79!=np__2|~r1_xxreal_0(np__257,np__79))).#trainneg
cnf(i_0_147132, plain, (np__6480!=np__256|~r1_xxreal_0(np__257,np__6480))).#trainneg
cnf(i_0_147133, plain, (np__254!=np__256|~r1_xxreal_0(np__257,np__254))).#trainneg
cnf(i_0_195639, plain, (k2_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__2,np__257),np__1),np__257)=np__2|r1_xxreal_0(np__257,k2_xcmplx_0(np__2,np__257))|~v7_ordinal1(k2_xcmplx_0(np__2,np__257)))).#trainneg
cnf(i_0_195584, plain, (k4_nat_d(k3_xcmplx_0(k3_xcmplx_0(k2_xcmplx_0(np__2,np__257),np__1),np__1),np__257)=k4_nat_d(k2_xcmplx_0(np__257,np__2),np__257)|~v7_ordinal1(k3_xcmplx_0(k2_xcmplx_0(np__2,np__257),np__1)))).#trainneg
cnf(i_0_201501, plain, (k2_nat_d(k3_xcmplx_0(k3_xcmplx_0(k2_xcmplx_0(np__2,np__257),np__1),np__1),np__257)=np__2|~v7_ordinal1(k3_xcmplx_0(k3_xcmplx_0(k2_xcmplx_0(np__2,np__257),np__1),np__1))|~v7_ordinal1(k3_xcmplx_0(k2_xcmplx_0(np__2,np__257),np__1)))).#trainneg
cnf(i_0_198232, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__136),X1)=k4_nat_d(k3_xcmplx_0(np__136,k2_xcmplx_0(X1,X2)),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_195968, plain, (k2_xcmplx_0(k3_xcmplx_0(k2_xcmplx_0(np__81,np__257),np__1),np__257)=np__81|r1_xxreal_0(np__257,k2_xcmplx_0(k3_xcmplx_0(k2_xcmplx_0(np__81,np__257),np__1),np__257))|r1_xxreal_0(np__257,k3_xcmplx_0(k2_xcmplx_0(np__81,np__257),np__1))|~v7_ordinal1(k2_xcmplx_0(k3_xcmplx_0(k2_xcmplx_0(np__81,np__257),np__1),np__257))|~v7_ordinal1(k3_xcmplx_0(k2_xcmplx_0(np__81,np__257),np__1)))).#trainneg
cnf(i_0_97095, plain, (X1=np__0|r1_xxreal_0(X1,np__136)|~r1_xxreal_0(X1,k2_nat_d(k2_xcmplx_0(np__136,X1),X1))|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_201195, plain, (k4_nat_d(X1,np__136)=k4_pepin(np__136)|r1_xxreal_0(np__1,k4_nat_d(X1,np__136))|r1_xxreal_0(np__1,k4_pepin(np__136))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_186511, plain, (np__6168=np__0|np__6168!=np__1|np__6168!=np__257)).#trainneg
cnf(i_0_2822, plain, (X1=k2_nat_d(X2,k3_xcmplx_0(X3,np__6561))|r1_xxreal_0(k3_xcmplx_0(X3,np__6561),X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(X3,np__0),X1)|~v1_xcmplx_0(np__6561)|~v1_xcmplx_0(X3)|~v7_ordinal1(k3_xcmplx_0(X3,np__6561))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_186513, plain, (np__6536=np__0|np__6536!=np__1|np__6536!=np__257)).#trainneg
cnf(i_0_186514, plain, (np__56=np__0|np__56!=np__1|np__56!=np__257)).#trainneg
cnf(i_0_186515, plain, (np__23=np__0|np__23!=np__1|np__23!=np__257)).#trainneg
cnf(i_0_147134, plain, (np__253!=np__256|~r1_xxreal_0(np__257,np__253))).#trainneg
cnf(i_0_147135, plain, (np__255!=np__256|~r1_xxreal_0(np__257,np__255))).#trainneg
cnf(i_0_147136, plain, (np__120!=np__256|~r1_xxreal_0(np__257,np__120))).#trainneg
cnf(i_0_147137, plain, (np__55!=np__256|~r1_xxreal_0(np__257,np__55))).#trainneg
cnf(i_0_147138, plain, (np__6305!=np__256|~r1_xxreal_0(np__257,np__6305))).#trainneg
cnf(i_0_17205, plain, (k2_nat_d(k2_xcmplx_0(X1,np__136),np__136)=X1|r1_xxreal_0(np__136,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_40788, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X3),k2_xcmplx_0(X1,X2)),X2)=k2_nat_d(k3_xcmplx_0(X1,X3),X2)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X2,X3)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_201196, plain, (k4_nat_d(X1,np__8)=k4_pepin(np__136)|r1_xxreal_0(np__1,k4_nat_d(X1,np__8))|r1_xxreal_0(np__1,k4_pepin(np__136))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_188228, plain, (np__6289=np__0|np__6289!=np__1|np__6289!=np__257)).#trainneg
cnf(i_0_2824, plain, (X1=k2_nat_d(X2,k3_xcmplx_0(X3,np__257))|r1_xxreal_0(k3_xcmplx_0(X3,np__257),X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(X3,np__0),X1)|~v1_xcmplx_0(np__257)|~v1_xcmplx_0(X3)|~v7_ordinal1(k3_xcmplx_0(X3,np__257))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_188229, plain, (np__231=np__0|np__231!=np__1|np__231!=np__257)).#trainneg
cnf(i_0_188230, plain, (np__79=np__0|np__79!=np__1|np__79!=np__257)).#trainneg
cnf(i_0_189684, plain, (np__6480=np__0|np__6480!=np__2|np__6480!=np__257)).#trainneg
cnf(i_0_147139, plain, (np__7!=np__256|~r1_xxreal_0(np__257,np__7))).#trainneg
cnf(i_0_147141, plain, (np__5!=np__256|~r1_xxreal_0(np__257,np__5))).#trainneg
cnf(i_0_147142, plain, (np__252!=np__256|~r1_xxreal_0(np__257,np__252))).#trainneg
cnf(i_0_147144, plain, (np__6558!=np__256|~r1_xxreal_0(np__257,np__6558))).#trainneg
cnf(i_0_147145, plain, (np__6424!=np__256|~r1_xxreal_0(np__257,np__6424))).#trainneg
cnf(i_0_12518, plain, (k2_xcmplx_0(X1,np__257)=X1|r1_xxreal_0(np__257,k2_xcmplx_0(X1,np__257))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(X1,np__257))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_81946, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__3,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__3,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_185638, plain, (k4_nat_d(k3_xcmplx_0(X1,np__2),np__257)=k4_nat_1(X1,np__2)|r1_xxreal_0(np__257,k4_nat_1(X1,np__2))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_1(X1,np__2),k5_numbers))).#trainneg
cnf(i_0_131910, plain, (k4_nat_d(k4_nat_1(np__136,X1),np__257)=k4_nat_d(k3_xcmplx_0(np__6561,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_140970, plain, (k4_nat_d(k4_nat_1(k2_nat_d(X1,np__257),np__136),np__257)=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|k2_nat_d(X2,np__257)!=np__136|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_203155, plain, (k2_xcmplx_0(X1,np__257)=X1|r1_xxreal_0(np__257,k2_xcmplx_0(X1,np__257))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)|~m1_subset_1(k2_xcmplx_0(X1,np__257),k5_numbers))).#trainneg
cnf(i_0_203837, plain, (k4_nat_d(k3_xcmplx_0(np__136,X1),np__257)=k4_nat_d(k3_xcmplx_0(np__6561,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_201197, plain, (k4_nat_d(X1,np__3)=k4_pepin(np__136)|r1_xxreal_0(np__1,k4_nat_d(X1,np__3))|r1_xxreal_0(np__1,k4_pepin(np__136))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_189685, plain, (np__254=np__0|np__254!=np__2|np__254!=np__257)).#trainneg
cnf(i_0_2825, plain, (X1=k2_nat_d(X2,k3_xcmplx_0(X3,np__257))|r1_xxreal_0(k3_xcmplx_0(X3,np__257),X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(X3,np__6425),X1)|~v1_xcmplx_0(np__257)|~v1_xcmplx_0(X3)|~v7_ordinal1(k3_xcmplx_0(X3,np__257))|~v7_ordinal1(np__25)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_189686, plain, (np__253=np__0|np__253!=np__2|np__253!=np__257)).#trainneg
cnf(i_0_190672, plain, (np__255=np__0|np__255!=np__2|np__255!=np__257)).#trainneg
cnf(i_0_190673, plain, (np__120=np__0|np__120!=np__2|np__120!=np__257)).#trainneg
cnf(i_0_147146, plain, (np__6304!=np__256|~r1_xxreal_0(np__257,np__6304))).#trainneg
cnf(i_0_147147, plain, (np__80!=np__256|~r1_xxreal_0(np__257,np__80))).#trainneg
cnf(i_0_147148, plain, (np__134!=np__256|~r1_xxreal_0(np__257,np__134))).#trainneg
cnf(i_0_147149, plain, (np__232!=np__256|~r1_xxreal_0(np__257,np__232))).#trainneg
cnf(i_0_147150, plain, (np__6400!=np__256|~r1_xxreal_0(np__257,np__6400))).#trainneg
cnf(i_0_57638, plain, (X1=k2_nat_d(X2,np__1)|r1_xxreal_0(np__1,X1)|X2!=k2_xcmplx_0(X1,np__8)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_133036, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X1,X2),np__6),X1)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,X1),np__6),X1)|r1_xxreal_0(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_68309, plain, (k4_nat_d(k3_xcmplx_0(np__4,k2_xcmplx_0(X2,k2_nat_d(X1,X2))),X2)=k4_nat_d(k3_xcmplx_0(X1,np__4),X2)|r1_xxreal_0(X2,k2_nat_d(X1,X2))|~v7_ordinal1(k2_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_194682, plain, (k4_nat_d(X1,X2)=k4_pepin(np__8)|r1_xxreal_0(np__1,k4_nat_d(X1,X2))|r1_xxreal_0(np__1,k4_pepin(np__8))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_190674, plain, (np__55=np__0|np__55!=np__2|np__55!=np__257)).#trainneg
cnf(i_0_190675, plain, (np__6305=np__0|np__6305!=np__2|np__6305!=np__257)).#trainneg
cnf(i_0_190676, plain, (np__7=np__0|np__7!=np__2|np__7!=np__257)).#trainneg
cnf(i_0_191501, plain, (np__5=np__0|np__5!=np__2|np__5!=np__257)).#trainneg
cnf(i_0_147151, plain, (np__6557!=np__256|~r1_xxreal_0(np__257,np__6557))).#trainneg
cnf(i_0_147152, plain, (np__6344!=np__256|~r1_xxreal_0(np__257,np__6344))).#trainneg
cnf(i_0_147155, plain, (np__22!=np__256|~r1_xxreal_0(np__257,np__22))).#trainneg
cnf(i_0_147156, plain, (np__111!=np__256|~r1_xxreal_0(np__257,np__111))).#trainneg
cnf(i_0_147157, plain, (np__121!=np__256|~r1_xxreal_0(np__257,np__121))).#trainneg
cnf(i_0_45529, plain, (k2_nat_d(k3_xcmplx_0(np__1,X1),X2)=k2_nat_d(X1,X2)|r1_xxreal_0(X2,np__1)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_145736, plain, (k4_nat_1(X1,k2_nat_d(X2,np__257))=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__257,X1),X2),np__257)|r1_xxreal_0(np__257,k4_nat_1(X1,k2_nat_d(X2,np__257)))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(k2_xcmplx_0(k4_nat_1(X1,k2_nat_d(X2,np__257)),np__257))|~v7_ordinal1(k4_nat_1(X1,k2_nat_d(X2,np__257)))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_65220, plain, (k4_nat_d(k4_nat_1(X1,k2_nat_d(np__2,X2)),X2)=k4_nat_d(k3_xcmplx_0(np__2,k2_xcmplx_0(X1,X2)),X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_140325, plain, (k2_nat_d(k4_nat_d(X1,np__257),np__257)=k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(np__1,np__257),X1),np__257)|r1_xxreal_0(np__257,k4_nat_d(X1,np__257))|~v7_ordinal1(k4_nat_d(X1,np__257))|~v7_ordinal1(X1)|~m1_subset_1(k4_nat_d(X1,np__257),k1_numbers))).#trainneg
cnf(i_0_204604, plain, (k2_xcmplx_0(np__257,X1)=X1|r1_xxreal_0(np__257,k2_xcmplx_0(np__257,X1))|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_205807, plain, (k4_nat_d(X1,np__136)=k4_pepin(np__8)|r1_xxreal_0(np__1,k4_nat_d(X1,np__136))|r1_xxreal_0(np__1,k4_pepin(np__8))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_191502, plain, (np__252=np__0|np__252!=np__2|np__252!=np__257)).#trainneg
cnf(i_0_2827, plain, (X1=k2_nat_d(X2,k3_xcmplx_0(X3,np__81))|r1_xxreal_0(k3_xcmplx_0(X3,np__81),X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(X3,np__6561),X1)|~v1_xcmplx_0(np__81)|~v1_xcmplx_0(X3)|~v7_ordinal1(k3_xcmplx_0(X3,np__81))|~v7_ordinal1(np__81)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_191503, plain, (np__6558=np__0|np__6558!=np__2|np__6558!=np__257)).#trainneg
cnf(i_0_191504, plain, (np__6424=np__0|np__6424!=np__2|np__6424!=np__257)).#trainneg
cnf(i_0_191505, plain, (np__6304=np__0|np__6304!=np__2|np__6304!=np__257)).#trainneg
cnf(i_0_147158, plain, (np__132!=np__256|~r1_xxreal_0(np__257,np__132))).#trainneg
cnf(i_0_147159, plain, (np__133!=np__256|~r1_xxreal_0(np__257,np__133))).#trainneg
cnf(i_0_147160, plain, (np__175!=np__256|~r1_xxreal_0(np__257,np__175))).#trainneg
cnf(i_0_147161, plain, (np__176!=np__256|~r1_xxreal_0(np__257,np__176))).#trainneg
cnf(i_0_147162, plain, (np__21!=np__256|~r1_xxreal_0(np__257,np__21))).#trainneg
cnf(i_0_204779, plain, (k4_nat_d(k3_xcmplx_0(np__6561,X1),np__257)=k2_nat_d(k3_xcmplx_0(np__136,X1),np__257)|r1_xxreal_0(np__257,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_205808, plain, (k4_nat_d(X1,np__8)=k4_pepin(np__8)|r1_xxreal_0(np__1,k4_nat_d(X1,np__8))|r1_xxreal_0(np__1,k4_pepin(np__8))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_192378, plain, (np__80=np__0|np__80!=np__2|np__80!=np__257)).#trainneg
cnf(i_0_2828, plain, (X1=k2_nat_d(X2,k3_xcmplx_0(X3,np__4))|r1_xxreal_0(k3_xcmplx_0(X3,np__4),X1)|X2!=k2_xcmplx_0(k3_xcmplx_0(X3,np__8),X1)|~v1_xcmplx_0(np__4)|~v1_xcmplx_0(X3)|~v7_ordinal1(k3_xcmplx_0(X3,np__4))|~v7_ordinal1(np__2)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_192379, plain, (np__134=np__0|np__134!=np__2|np__134!=np__257)).#trainneg
cnf(i_0_192380, plain, (np__232=np__0|np__232!=np__2|np__232!=np__257)).#trainneg
cnf(i_0_192381, plain, (np__6400=np__0|np__6400!=np__2|np__6400!=np__257)).#trainneg
cnf(i_0_147163, plain, (np__6169!=np__256|~r1_xxreal_0(np__257,np__6169))).#trainneg
cnf(i_0_147164, plain, (np__6421!=np__256|~r1_xxreal_0(np__257,np__6421))).#trainneg
cnf(i_0_147165, plain, (np__6422!=np__256|~r1_xxreal_0(np__257,np__6422))).#trainneg
cnf(i_0_147166, plain, (np__6423!=np__256|~r1_xxreal_0(np__257,np__6423))).#trainneg
cnf(i_0_147167, plain, (np__6559!=np__256|~r1_xxreal_0(np__257,np__6559))).#trainneg
cnf(i_0_70502, plain, (k4_nat_d(k3_xcmplx_0(k2_xcmplx_0(X2,k4_nat_d(X1,X2)),np__8),X2)=k4_nat_d(k3_xcmplx_0(X1,np__8),X2)|r1_xxreal_0(X2,k4_nat_d(X1,X2))|~v7_ordinal1(k4_nat_d(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_52150, plain, (k2_nat_d(k3_xcmplx_0(k2_nat_d(X2,X3),X1),X3)=k4_nat_d(k3_xcmplx_0(X2,k2_xcmplx_0(X1,X3)),X3)|r1_xxreal_0(X3,X1)|~v7_ordinal1(k3_xcmplx_0(k2_nat_d(X2,X3),X1))|~v7_ordinal1(k2_nat_d(X2,X3))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_58623, plain, (X1=np__0|r1_xxreal_0(X1,X3)|X2!=k2_nat_d(k2_xcmplx_0(X1,X3),X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_xcmplx_0(X3,X1))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_205809, plain, (k4_nat_d(X1,np__3)=k4_pepin(np__8)|r1_xxreal_0(np__1,k4_nat_d(X1,np__3))|r1_xxreal_0(np__1,k4_pepin(np__8))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_192382, plain, (np__6557=np__0|np__6557!=np__2|np__6557!=np__257)).#trainneg
cnf(i_0_192719, plain, (np__6344=np__0|np__6344!=np__2|np__6344!=np__257)).#trainneg
cnf(i_0_192720, plain, (np__22=np__0|np__22!=np__2|np__22!=np__257)).#trainneg
cnf(i_0_192721, plain, (np__111=np__0|np__111!=np__2|np__111!=np__257)).#trainneg
cnf(i_0_147168, plain, (np__6560!=np__256|~r1_xxreal_0(np__257,np__6560))).#trainneg
cnf(i_0_147169, plain, (np__77!=np__256|~r1_xxreal_0(np__257,np__77))).#trainneg
cnf(i_0_147170, plain, (np__78!=np__256|~r1_xxreal_0(np__257,np__78))).#trainneg
cnf(i_0_147171, plain, (np__24!=np__256|~r1_xxreal_0(np__257,np__24))).#trainneg
cnf(i_0_147175, plain, (np__135!=np__256|~r1_xxreal_0(np__257,np__135))).#trainneg
cnf(i_0_94542, plain, (k4_nat_d(k3_xcmplx_0(X2,k2_nat_d(np__1,X1)),X1)=k4_nat_d(k3_xcmplx_0(np__1,k2_xcmplx_0(X2,X1)),X1)|r1_xxreal_0(X1,k2_nat_d(np__1,X1))|r1_xxreal_0(X1,X2)|~v7_ordinal1(k2_nat_d(np__1,X1))|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_140965, plain, (k4_nat_d(k4_nat_1(k4_nat_d(X1,np__257),np__136),np__257)=k4_nat_d(k3_xcmplx_0(X1,X2),np__257)|k2_nat_d(X2,np__257)!=np__136|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 397
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 591
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 588
# Processed clauses                    : 6017
# ...of these trivial                  : 400
# ...subsumed                          : 1490
# ...remaining for further processing  : 4127
# Other redundant clauses eliminated   : 84
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 142
# Backward-rewritten                   : 496
# Generated clauses                    : 209308
# ...of the previous two non-redundant : 83219
# ...frozen by parental guidance       : 108788
# ...of these subsequently unfrozen    : 0
# ...aggressively subsumed             : 7944
# Contextual simplify-reflections      : 315
# Paramodulations                      : 209118
# Factorizations                       : 0
# NegExts                              : 0
# Equation resolutions                 : 190
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 3489
#    Positive orientable unit clauses  : 803
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 230
#    Non-unit-clauses                  : 2456
# Current number of unprocessed clauses: 77280
# ...number of literals in the above   : 329890
# Current number of archived formulas  : 0
# Current number of archived clauses   : 4765
# Proof object given clauses           : 259
# Proof search given clauses           : 4127
# Clause-clause subsumption calls (NU) : 24737167
# Rec. Clause-clause subsumption calls : 6082442
# Non-unit clause-clause subsumptions  : 6112
# Unit Clause-clause subsumption calls : 96456
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 171
# BW rewrite match successes           : 160
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 5080608

# -------------------------------------------------
# User time                : 78.882 s
# System time              : 0.989 s
# Total time               : 79.871 s
# Maximum resident set size: 306904 pages
78.88user 1.02system 2:48.71elapsed 47%CPU (0avgtext+0avgdata 307164maxresident)k
80inputs+1128outputs (0major+221279minor)pagefaults 0swaps
