2020-06-03 22:11:04.116673: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_8_bb1/premsel_enigma_01_2020_T30_loop02_epoch_243
2020-06-03 22:11:04.212377: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-06-03 22:11:04.345396: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-06-03 22:11:04.869240: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-06-03 22:11:06.315570: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 2198919 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_8_bb1/premsel_enigma_01_2020_T30_loop02_epoch_243' loaded (query=128; context=512; weigths=real; len_mult=0.000000).
# Preprocessing time       : 2.919 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t21_extreal2, axiom, ![X1]:(m1_subset_1(X1,k7_numbers)=>![X2]:(m1_subset_1(X2,k7_numbers)=>r1_xxreal_0(k3_extreal1(k4_supinf_2(X1,X2)),k3_supinf_2(k3_extreal1(X1),k3_extreal1(X2))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', t21_extreal2)).
fof(redefinition_k4_supinf_2, axiom, ![X1, X2]:((m1_subset_1(X1,k7_numbers)&m1_subset_1(X2,k7_numbers))=>k4_supinf_2(X1,X2)=k3_xxreal_3(X1,X2)), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', redefinition_k4_supinf_2)).
fof(commutativity_k3_supinf_2, axiom, ![X1, X2]:((m1_subset_1(X1,k7_numbers)&m1_subset_1(X2,k7_numbers))=>k3_supinf_2(X1,X2)=k3_supinf_2(X2,X1)), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', commutativity_k3_supinf_2)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', t6_boole)).
fof(redefinition_k1_supinf_2, axiom, k1_supinf_2=k1_xboole_0, file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', redefinition_k1_supinf_2)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', redefinition_k6_numbers)).
fof(dt_k3_extreal1, axiom, ![X1]:(m1_subset_1(X1,k7_numbers)=>m1_subset_1(k3_extreal1(X1),k7_numbers)), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', dt_k3_extreal1)).
fof(d16_mesfunc5, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((((~(v1_xboole_0(X2))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))=>![X3]:((((((v1_funct_1(X3)&v1_funct_2(X3,X2,k7_numbers))&v10_valued_0(X3))&v6_supinf_2(X3))&v4_measure1(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))=>![X4]:((v1_funct_1(X4)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>k7_mesfunc5(X1,X2,X3,X4)=k4_supinf_2(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', d16_mesfunc5)).
fof(d1_extreal1, axiom, ![X1]:(m1_subset_1(X1,k7_numbers)=>((r1_xxreal_0(k6_numbers,X1)=>k3_extreal1(X1)=X1)&(~(r1_xxreal_0(k6_numbers,X1))=>k3_extreal1(X1)=k2_supinf_2(X1)))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', d1_extreal1)).
fof(spc0_boole, axiom, v1_xboole_0(np__0), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', spc0_boole)).
fof(d17_mesfunc5, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((((~(v1_xboole_0(X2))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))=>![X3]:((((((v1_funct_1(X3)&v1_funct_2(X3,X2,k7_numbers))&v10_valued_0(X3))&v6_supinf_2(X3))&v4_measure1(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))=>![X4]:((v1_funct_1(X4)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>(r1_mesfunc5(X1,X2,X3,X4)<=>((?[X5]:((m2_subset_1(X5,k1_zfmisc_1(X1),X2)&X5=k1_relset_1(X1,X4))&r1_mesfunc1(X1,X2,X4,X5))&~(r1_xxreal_0(k1_supinf_1,k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)))))&~(r1_xxreal_0(k1_supinf_1,k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)))))))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', d17_mesfunc5)).
fof(t101_mesfunc5, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((((~(v1_xboole_0(X2))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))=>![X3]:((((((v1_funct_1(X3)&v1_funct_2(X3,X2,k7_numbers))&v10_valued_0(X3))&v6_supinf_2(X3))&v4_measure1(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))=>![X4]:((v1_funct_1(X4)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>(r1_mesfunc5(X1,X2,X3,X4)=>r1_xxreal_0(k3_extreal1(k7_mesfunc5(X1,X2,X3,X4)),k7_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4)))))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', t101_mesfunc5)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', redefinition_m2_subset_1)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', cc1_subset_1)).
fof(l29_mesfunc5, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((((~(v1_xboole_0(X2))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))=>![X3]:((((((v1_funct_1(X3)&v1_funct_2(X3,X2,k7_numbers))&v10_valued_0(X3))&v6_supinf_2(X3))&v4_measure1(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))=>![X4]:((v1_funct_1(X4)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>((v6_supinf_2(k1_mesfunc2(X1,X4))&v6_supinf_2(k2_mesfunc2(X1,X4)))&v6_supinf_2(k10_mesfunc1(X1,X4))))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', l29_mesfunc5)).
fof(rc6_measure1, axiom, ![X1, X2]:((((~(v1_xboole_0(X2))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))=>?[X3]:(((((((((((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))&v1_relat_1(X3))&v4_relat_1(X3,X2))&v5_relat_1(X3,k7_numbers))&v1_funct_1(X3))&~(v1_xboole_0(X3)))&v1_partfun1(X3,X2))&v1_funct_2(X3,X2,k7_numbers))&v2_valued_0(X3))&v10_valued_0(X3))&v6_supinf_2(X3))&v4_measure1(X3,X1,X2))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', rc6_measure1)).
fof(l102_mesfunc5, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((((~(v1_xboole_0(X2))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))=>![X3]:((((((v1_funct_1(X3)&v1_funct_2(X3,X2,k7_numbers))&v10_valued_0(X3))&v6_supinf_2(X3))&v4_measure1(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))=>![X4]:((v1_funct_1(X4)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>![X5]:((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>((v6_supinf_2(X4)&v6_supinf_2(X5))=>(![X6]:(m2_subset_1(X6,k1_zfmisc_1(X1),X2)=>~((((X6=k1_relset_1(X1,X4)&X6=k1_relset_1(X1,X5))&r1_mesfunc1(X1,X2,X4,X6))&r1_mesfunc1(X1,X2,X5,X6))))|k6_mesfunc5(X1,X2,X3,k3_mesfunc1(X1,X4,X5))=k3_supinf_2(k6_mesfunc5(X1,X2,X3,X4),k6_mesfunc5(X1,X2,X3,X5))))))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', l102_mesfunc5)).
fof(redefinition_r2_relset_1, axiom, ![X1, X2, X3, X4]:((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>(r2_relset_1(X1,X2,X3,X4)<=>X3=X4)), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', redefinition_r2_relset_1)).
fof(t24_mesfunc2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>r2_relset_1(X1,k7_numbers,k10_mesfunc1(X1,X2),k3_mesfunc1(X1,k1_mesfunc2(X1,X2),k2_mesfunc2(X1,X2))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', t24_mesfunc2)).
fof(t88_mesfunc5, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((((~(v1_xboole_0(X2))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))=>![X3]:((((((v1_funct_1(X3)&v1_funct_2(X3,X2,k7_numbers))&v10_valued_0(X3))&v6_supinf_2(X3))&v4_measure1(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))=>![X4]:((v1_funct_1(X4)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>(v6_supinf_2(X4)=>(![X5]:(m2_subset_1(X5,k1_zfmisc_1(X1),X2)=>~((X5=k1_relset_1(X1,X4)&r1_mesfunc1(X1,X2,X4,X5))))|k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4))))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', t88_mesfunc5)).
fof(dt_k3_mesfunc1, axiom, ![X1, X2, X3]:(((((((((~(v1_xboole_0(X1))&v1_relat_1(X2))&v4_relat_1(X2,X1))&v5_relat_1(X2,k7_numbers))&v1_funct_1(X2))&v1_relat_1(X3))&v4_relat_1(X3,X1))&v5_relat_1(X3,k7_numbers))&v1_funct_1(X3))=>(v1_funct_1(k3_mesfunc1(X1,X2,X3))&m1_subset_1(k3_mesfunc1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', dt_k3_mesfunc1)).
fof(t27_mesfunc2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((((~(v1_xboole_0(X2))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))=>![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>![X4]:(m1_subset_1(X4,X2)=>((r1_mesfunc1(X1,X2,X3,X4)&r1_tarski(X4,k1_relset_1(X1,X3)))=>r1_mesfunc1(X1,X2,k10_mesfunc1(X1,X3),X4)))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', t27_mesfunc2)).
fof(dt_k1_mesfunc2, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&v1_funct_1(X2))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>(v1_funct_1(k1_mesfunc2(X1,X2))&m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', dt_k1_mesfunc2)).
fof(dt_k2_mesfunc2, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&v1_funct_1(X2))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>(v1_funct_1(k2_mesfunc2(X1,X2))&m1_subset_1(k2_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', dt_k2_mesfunc2)).
fof(t79_mesfunc5, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((((~(v1_xboole_0(X2))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))=>![X3]:((((((v1_funct_1(X3)&v1_funct_2(X3,X2,k7_numbers))&v10_valued_0(X3))&v6_supinf_2(X3))&v4_measure1(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))=>![X4]:((v1_funct_1(X4)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>(v6_supinf_2(X4)=>(![X5]:(m2_subset_1(X5,k1_zfmisc_1(X1),X2)=>~((X5=k1_relset_1(X1,X4)&r1_mesfunc1(X1,X2,X4,X5))))|r1_xxreal_0(k6_numbers,k6_mesfunc5(X1,X2,X3,X4)))))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', t79_mesfunc5)).
fof(reflexivity_r1_tarski, axiom, ![X1, X2]:r1_tarski(X1,X1), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', reflexivity_r1_tarski)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', cc1_relset_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', cc2_relset_1)).
fof(t26_mesfunc2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>![X3]:((((~(v1_xboole_0(X3))&v1_prob_1(X3,X1))&v4_prob_1(X3,X1))&m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1))))=>![X4]:(m1_subset_1(X4,X3)=>((r1_mesfunc1(X1,X3,X2,X4)&r1_tarski(X4,k1_relset_1(X1,X2)))=>r1_mesfunc1(X1,X3,k2_mesfunc2(X1,X2),X4)))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', t26_mesfunc2)).
fof(t25_mesfunc2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>![X3]:((((~(v1_xboole_0(X3))&v1_prob_1(X3,X1))&v4_prob_1(X3,X1))&m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1))))=>![X4]:(m1_subset_1(X4,X3)=>(r1_mesfunc1(X1,X3,X2,X4)=>r1_mesfunc1(X1,X3,k1_mesfunc2(X1,X2),X4)))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', t25_mesfunc2)).
fof(dt_k6_mesfunc5, axiom, ![X1, X2, X3, X4]:(((((((((((((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))&v1_funct_1(X3))&v1_funct_2(X3,X2,k7_numbers))&v10_valued_0(X3))&v6_supinf_2(X3))&v4_measure1(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))&v1_funct_1(X4))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>m1_subset_1(k6_mesfunc5(X1,X2,X3,X4),k7_numbers)), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', dt_k6_mesfunc5)).
fof(cc5_measure1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>((v1_prob_1(X2,X1)&v3_measure1(X2,X1))=>v4_prob_1(X2,X1))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', cc5_measure1)).
fof(rc4_measure1, axiom, ![X1]:?[X2]:(((m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))&~(v1_xboole_0(X2)))&v1_prob_1(X2,X1))&v3_measure1(X2,X1)), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', rc4_measure1)).
fof(dt_k10_mesfunc1, axiom, ![X1, X2]:(((((~(v1_xboole_0(X1))&v1_relat_1(X2))&v4_relat_1(X2,X1))&v5_relat_1(X2,k7_numbers))&v1_funct_1(X2))=>(v1_funct_1(k10_mesfunc1(X1,X2))&m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', dt_k10_mesfunc1)).
fof(d10_mesfunc1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((((v1_relat_1(X2)&v4_relat_1(X2,X1))&v5_relat_1(X2,k7_numbers))&v1_funct_1(X2))=>![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>(X3=k10_mesfunc1(X1,X2)<=>(k9_xtuple_0(X3)=k9_xtuple_0(X2)&![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,k9_xtuple_0(X3))=>k12_supinf_2(X3,X4)=k3_extreal1(k12_supinf_2(X2,X4))))))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', d10_mesfunc1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>k1_relset_1(X1,X2)=k9_xtuple_0(X2)), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', redefinition_k1_relset_1)).
fof(d3_mesfunc2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>(X3=k2_mesfunc2(X1,X2)<=>(k1_relset_1(X1,X3)=k1_relset_1(X1,X2)&![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,k1_relset_1(X1,X3))=>k12_supinf_2(X3,X4)=k1_extreal2(k2_supinf_2(k12_supinf_2(X2,X4)),k1_supinf_2)))))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', d3_mesfunc2)).
fof(d2_mesfunc2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>(X3=k1_mesfunc2(X1,X2)<=>(k1_relset_1(X1,X3)=k1_relset_1(X1,X2)&![X4]:(m1_subset_1(X4,X1)=>(r2_hidden(X4,k1_relset_1(X1,X3))=>k12_supinf_2(X3,X4)=k1_extreal2(k12_supinf_2(X2,X4),k1_supinf_2)))))))), file('/local1/mptp/loop2/hardpred/preds__0.05/t101_mesfunc5', d2_mesfunc2)).
fof(c_0_38, plain, ![X84, X85]:(~m1_subset_1(X84,k7_numbers)|(~m1_subset_1(X85,k7_numbers)|r1_xxreal_0(k3_extreal1(k4_supinf_2(X84,X85)),k3_supinf_2(k3_extreal1(X84),k3_extreal1(X85))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t21_extreal2])])])).
fof(c_0_39, plain, ![X100, X101]:(~m1_subset_1(X100,k7_numbers)|~m1_subset_1(X101,k7_numbers)|k4_supinf_2(X100,X101)=k3_xxreal_3(X100,X101)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_supinf_2])])).
cnf(c_0_40, plain, (r1_xxreal_0(k3_extreal1(k4_supinf_2(X1,X2)),k3_supinf_2(k3_extreal1(X1),k3_extreal1(X2)))|~m1_subset_1(X1,k7_numbers)|~m1_subset_1(X2,k7_numbers)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_41, plain, (k4_supinf_2(X1,X2)=k3_xxreal_3(X1,X2)|~m1_subset_1(X1,k7_numbers)|~m1_subset_1(X2,k7_numbers)), inference(split_conjunct,[status(thm)],[c_0_39])).
fof(c_0_42, plain, ![X96, X97]:(~m1_subset_1(X96,k7_numbers)|~m1_subset_1(X97,k7_numbers)|k3_supinf_2(X96,X97)=k3_supinf_2(X97,X96)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k3_supinf_2])])).
fof(c_0_43, plain, ![X133]:(~v1_xboole_0(X133)|X133=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_44, plain, (k1_supinf_2=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k1_supinf_2])).
cnf(c_0_45, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
cnf(c_0_46, plain, (r1_xxreal_0(k3_extreal1(k4_supinf_2(X1,X2)),k3_supinf_2(k3_extreal1(X1),k3_extreal1(X2)))|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)), inference(evalgc,[status(thm)],[c_0_40])).
cnf(c_0_47, plain, (k4_supinf_2(X1,X2)=k3_xxreal_3(X1,X2)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)), inference(evalgc,[status(thm)],[c_0_41])).
cnf(c_0_48, plain, (k3_supinf_2(X1,X2)=k3_supinf_2(X2,X1)|~m1_subset_1(X1,k7_numbers)|~m1_subset_1(X2,k7_numbers)), inference(split_conjunct,[status(thm)],[c_0_42])).
fof(c_0_49, plain, ![X171]:(~m1_subset_1(X171,k7_numbers)|m1_subset_1(k3_extreal1(X171),k7_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_extreal1])])).
fof(c_0_50, plain, ![X152, X153, X154, X155]:(v1_xboole_0(X152)|(v1_xboole_0(X153)|~v1_prob_1(X153,X152)|~v4_prob_1(X153,X152)|~m1_subset_1(X153,k1_zfmisc_1(k1_zfmisc_1(X152)))|(~v1_funct_1(X154)|~v1_funct_2(X154,X153,k7_numbers)|~v10_valued_0(X154)|~v6_supinf_2(X154)|~v4_measure1(X154,X152,X153)|~m1_subset_1(X154,k1_zfmisc_1(k2_zfmisc_1(X153,k7_numbers)))|(~v1_funct_1(X155)|~m1_subset_1(X155,k1_zfmisc_1(k2_zfmisc_1(X152,k7_numbers)))|k7_mesfunc5(X152,X153,X154,X155)=k4_supinf_2(k6_mesfunc5(X152,X153,X154,k1_mesfunc2(X152,X155)),k6_mesfunc5(X152,X153,X154,k2_mesfunc2(X152,X155))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d16_mesfunc5])])])])).
fof(c_0_51, plain, ![X27]:((~r1_xxreal_0(k6_numbers,X27)|k3_extreal1(X27)=X27|~m1_subset_1(X27,k7_numbers))&(r1_xxreal_0(k6_numbers,X27)|k3_extreal1(X27)=k2_supinf_2(X27)|~m1_subset_1(X27,k7_numbers))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d1_extreal1])])])])).
cnf(c_0_52, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_53, plain, (k1_xboole_0=k1_supinf_2), inference(evalgc,[status(thm)],[c_0_44])).
cnf(c_0_54, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_45])).
cnf(c_0_55, plain, (r1_xxreal_0(k3_extreal1(k3_xxreal_3(X1,X2)),k3_supinf_2(k3_extreal1(X1),k3_extreal1(X2)))|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_46, c_0_47])])).
cnf(c_0_56, plain, (k3_supinf_2(X1,X2)=k3_supinf_2(X2,X1)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)), inference(evalgc,[status(thm)],[c_0_48])).
cnf(c_0_57, plain, (m1_subset_1(k3_extreal1(X1),k7_numbers)|~m1_subset_1(X1,k7_numbers)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_58, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|k7_mesfunc5(X1,X2,X3,X4)=k4_supinf_2(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v1_funct_1(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v4_measure1(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_59, plain, (k3_extreal1(X1)=X1|~r1_xxreal_0(k6_numbers,X1)|~m1_subset_1(X1,k7_numbers)), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_60, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_52])).
cnf(c_0_61, plain, (k6_numbers=k1_supinf_2), inference(rw,[status(thm)],[c_0_53, c_0_54])).
cnf(c_0_62, plain, (v1_xboole_0(np__0)), inference(split_conjunct,[status(thm)],[spc0_boole])).
cnf(c_0_63, plain, (r1_xxreal_0(k3_extreal1(k3_xxreal_3(X1,X2)),k3_supinf_2(k3_extreal1(X2),k3_extreal1(X1)))|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)|~m1_subset_1(k3_extreal1(X2),k7_numbers)|~m1_subset_1(k3_extreal1(X1),k7_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56])])).
cnf(c_0_64, plain, (m1_subset_1(k3_extreal1(X1),k7_numbers)|~m1_subset_1(X1,k7_numbers)), inference(evalgc,[status(thm)],[c_0_57])).
cnf(c_0_65, plain, (k4_supinf_2(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)))=k7_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~v4_measure1(X3,X1,X2)|~v1_funct_2(X3,X2,k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_58])).
cnf(c_0_66, plain, (k3_extreal1(X1)=X1|~m1_subset_1(X1,k7_numbers)|~r1_xxreal_0(k6_numbers,X1)), inference(evalgc,[status(thm)],[c_0_59])).
cnf(c_0_67, plain, (X1=k1_supinf_2|~v1_xboole_0(X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_60, c_0_54]), c_0_61])).
cnf(c_0_68, plain, (v1_xboole_0(np__0)), inference(evalgc,[status(thm)],[c_0_62])).
fof(c_0_69, plain, ![X11, X12, X13, X14, X16]:((((((m2_subset_1(esk5_4(X11,X12,X13,X14),k1_zfmisc_1(X11),X12)|~r1_mesfunc5(X11,X12,X13,X14)|(~v1_funct_1(X14)|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(X11,k7_numbers))))|(~v1_funct_1(X13)|~v1_funct_2(X13,X12,k7_numbers)|~v10_valued_0(X13)|~v6_supinf_2(X13)|~v4_measure1(X13,X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X12,k7_numbers))))|(v1_xboole_0(X12)|~v1_prob_1(X12,X11)|~v4_prob_1(X12,X11)|~m1_subset_1(X12,k1_zfmisc_1(k1_zfmisc_1(X11))))|v1_xboole_0(X11))&(esk5_4(X11,X12,X13,X14)=k1_relset_1(X11,X14)|~r1_mesfunc5(X11,X12,X13,X14)|(~v1_funct_1(X14)|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(X11,k7_numbers))))|(~v1_funct_1(X13)|~v1_funct_2(X13,X12,k7_numbers)|~v10_valued_0(X13)|~v6_supinf_2(X13)|~v4_measure1(X13,X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X12,k7_numbers))))|(v1_xboole_0(X12)|~v1_prob_1(X12,X11)|~v4_prob_1(X12,X11)|~m1_subset_1(X12,k1_zfmisc_1(k1_zfmisc_1(X11))))|v1_xboole_0(X11)))&(r1_mesfunc1(X11,X12,X14,esk5_4(X11,X12,X13,X14))|~r1_mesfunc5(X11,X12,X13,X14)|(~v1_funct_1(X14)|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(X11,k7_numbers))))|(~v1_funct_1(X13)|~v1_funct_2(X13,X12,k7_numbers)|~v10_valued_0(X13)|~v6_supinf_2(X13)|~v4_measure1(X13,X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X12,k7_numbers))))|(v1_xboole_0(X12)|~v1_prob_1(X12,X11)|~v4_prob_1(X12,X11)|~m1_subset_1(X12,k1_zfmisc_1(k1_zfmisc_1(X11))))|v1_xboole_0(X11)))&(~r1_xxreal_0(k1_supinf_1,k6_mesfunc5(X11,X12,X13,k1_mesfunc2(X11,X14)))|~r1_mesfunc5(X11,X12,X13,X14)|(~v1_funct_1(X14)|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(X11,k7_numbers))))|(~v1_funct_1(X13)|~v1_funct_2(X13,X12,k7_numbers)|~v10_valued_0(X13)|~v6_supinf_2(X13)|~v4_measure1(X13,X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X12,k7_numbers))))|(v1_xboole_0(X12)|~v1_prob_1(X12,X11)|~v4_prob_1(X12,X11)|~m1_subset_1(X12,k1_zfmisc_1(k1_zfmisc_1(X11))))|v1_xboole_0(X11)))&(~r1_xxreal_0(k1_supinf_1,k6_mesfunc5(X11,X12,X13,k2_mesfunc2(X11,X14)))|~r1_mesfunc5(X11,X12,X13,X14)|(~v1_funct_1(X14)|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(X11,k7_numbers))))|(~v1_funct_1(X13)|~v1_funct_2(X13,X12,k7_numbers)|~v10_valued_0(X13)|~v6_supinf_2(X13)|~v4_measure1(X13,X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X12,k7_numbers))))|(v1_xboole_0(X12)|~v1_prob_1(X12,X11)|~v4_prob_1(X12,X11)|~m1_subset_1(X12,k1_zfmisc_1(k1_zfmisc_1(X11))))|v1_xboole_0(X11)))&(~m2_subset_1(X16,k1_zfmisc_1(X11),X12)|X16!=k1_relset_1(X11,X14)|~r1_mesfunc1(X11,X12,X14,X16)|r1_xxreal_0(k1_supinf_1,k6_mesfunc5(X11,X12,X13,k1_mesfunc2(X11,X14)))|r1_xxreal_0(k1_supinf_1,k6_mesfunc5(X11,X12,X13,k2_mesfunc2(X11,X14)))|r1_mesfunc5(X11,X12,X13,X14)|(~v1_funct_1(X14)|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(X11,k7_numbers))))|(~v1_funct_1(X13)|~v1_funct_2(X13,X12,k7_numbers)|~v10_valued_0(X13)|~v6_supinf_2(X13)|~v4_measure1(X13,X11,X12)|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(X12,k7_numbers))))|(v1_xboole_0(X12)|~v1_prob_1(X12,X11)|~v4_prob_1(X12,X11)|~m1_subset_1(X12,k1_zfmisc_1(k1_zfmisc_1(X11))))|v1_xboole_0(X11))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d17_mesfunc5])])])])])])).
fof(c_0_70, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:((((~(v1_xboole_0(X2))&v1_prob_1(X2,X1))&v4_prob_1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))=>![X3]:((((((v1_funct_1(X3)&v1_funct_2(X3,X2,k7_numbers))&v10_valued_0(X3))&v6_supinf_2(X3))&v4_measure1(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))=>![X4]:((v1_funct_1(X4)&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=>(r1_mesfunc5(X1,X2,X3,X4)=>r1_xxreal_0(k3_extreal1(k7_mesfunc5(X1,X2,X3,X4)),k7_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4))))))))), inference(assume_negation,[status(cth)],[t101_mesfunc5])).
cnf(c_0_71, plain, (r1_xxreal_0(k3_extreal1(k3_xxreal_3(X1,X2)),k3_supinf_2(k3_extreal1(X2),k3_extreal1(X1)))|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_63, c_0_64]), c_0_64])).
cnf(c_0_72, plain, (k3_xxreal_3(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)))=k7_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_65])])).
cnf(c_0_73, plain, (k3_extreal1(X1)=X1|~r1_xxreal_0(k1_supinf_2,X1)|~m1_subset_1(X1,k7_numbers)), inference(rw,[status(thm)],[c_0_66, c_0_61])).
cnf(c_0_74, plain, (k1_supinf_2=np__0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68])])).
cnf(c_0_75, plain, (m2_subset_1(esk5_4(X1,X2,X3,X4),k1_zfmisc_1(X1),X2)|v1_xboole_0(X2)|v1_xboole_0(X1)|~r1_mesfunc5(X1,X2,X3,X4)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v4_measure1(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_76, plain, (esk5_4(X1,X2,X3,X4)=k1_relset_1(X1,X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|~r1_mesfunc5(X1,X2,X3,X4)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v4_measure1(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_69])).
fof(c_0_77, negated_conjecture, (~v1_xboole_0(esk1_0)&((((~v1_xboole_0(esk2_0)&v1_prob_1(esk2_0,esk1_0))&v4_prob_1(esk2_0,esk1_0))&m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0))))&((((((v1_funct_1(esk3_0)&v1_funct_2(esk3_0,esk2_0,k7_numbers))&v10_valued_0(esk3_0))&v6_supinf_2(esk3_0))&v4_measure1(esk3_0,esk1_0,esk2_0))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))&((v1_funct_1(esk4_0)&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))&(r1_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0)&~r1_xxreal_0(k3_extreal1(k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0)),k7_mesfunc5(esk1_0,esk2_0,esk3_0,k10_mesfunc1(esk1_0,esk4_0)))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_70])])])])).
fof(c_0_78, plain, ![X93, X94, X95]:((~m2_subset_1(X95,X93,X94)|m1_subset_1(X95,X94)|(v1_xboole_0(X93)|v1_xboole_0(X94)|~m1_subset_1(X94,k1_zfmisc_1(X93))))&(~m1_subset_1(X95,X94)|m2_subset_1(X95,X93,X94)|(v1_xboole_0(X93)|v1_xboole_0(X94)|~m1_subset_1(X94,k1_zfmisc_1(X93))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])).
fof(c_0_79, plain, ![X102, X103]:(~v1_xboole_0(X102)|(~m1_subset_1(X103,k1_zfmisc_1(X102))|v1_xboole_0(X103))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
fof(c_0_80, plain, ![X63, X64, X65, X66]:(((v6_supinf_2(k1_mesfunc2(X63,X66))|(~v1_funct_1(X66)|~m1_subset_1(X66,k1_zfmisc_1(k2_zfmisc_1(X63,k7_numbers))))|(~v1_funct_1(X65)|~v1_funct_2(X65,X64,k7_numbers)|~v10_valued_0(X65)|~v6_supinf_2(X65)|~v4_measure1(X65,X63,X64)|~m1_subset_1(X65,k1_zfmisc_1(k2_zfmisc_1(X64,k7_numbers))))|(v1_xboole_0(X64)|~v1_prob_1(X64,X63)|~v4_prob_1(X64,X63)|~m1_subset_1(X64,k1_zfmisc_1(k1_zfmisc_1(X63))))|v1_xboole_0(X63))&(v6_supinf_2(k2_mesfunc2(X63,X66))|(~v1_funct_1(X66)|~m1_subset_1(X66,k1_zfmisc_1(k2_zfmisc_1(X63,k7_numbers))))|(~v1_funct_1(X65)|~v1_funct_2(X65,X64,k7_numbers)|~v10_valued_0(X65)|~v6_supinf_2(X65)|~v4_measure1(X65,X63,X64)|~m1_subset_1(X65,k1_zfmisc_1(k2_zfmisc_1(X64,k7_numbers))))|(v1_xboole_0(X64)|~v1_prob_1(X64,X63)|~v4_prob_1(X64,X63)|~m1_subset_1(X64,k1_zfmisc_1(k1_zfmisc_1(X63))))|v1_xboole_0(X63)))&(v6_supinf_2(k10_mesfunc1(X63,X66))|(~v1_funct_1(X66)|~m1_subset_1(X66,k1_zfmisc_1(k2_zfmisc_1(X63,k7_numbers))))|(~v1_funct_1(X65)|~v1_funct_2(X65,X64,k7_numbers)|~v10_valued_0(X65)|~v6_supinf_2(X65)|~v4_measure1(X65,X63,X64)|~m1_subset_1(X65,k1_zfmisc_1(k2_zfmisc_1(X64,k7_numbers))))|(v1_xboole_0(X64)|~v1_prob_1(X64,X63)|~v4_prob_1(X64,X63)|~m1_subset_1(X64,k1_zfmisc_1(k1_zfmisc_1(X63))))|v1_xboole_0(X63))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l29_mesfunc5])])])])])).
fof(c_0_81, plain, ![X141, X142]:((((((((((((m1_subset_1(esk10_2(X141,X142),k1_zfmisc_1(k2_zfmisc_1(X142,k7_numbers)))|(v1_xboole_0(X142)|~v1_prob_1(X142,X141)|~v4_prob_1(X142,X141)|~m1_subset_1(X142,k1_zfmisc_1(k1_zfmisc_1(X141)))))&(v1_relat_1(esk10_2(X141,X142))|(v1_xboole_0(X142)|~v1_prob_1(X142,X141)|~v4_prob_1(X142,X141)|~m1_subset_1(X142,k1_zfmisc_1(k1_zfmisc_1(X141))))))&(v4_relat_1(esk10_2(X141,X142),X142)|(v1_xboole_0(X142)|~v1_prob_1(X142,X141)|~v4_prob_1(X142,X141)|~m1_subset_1(X142,k1_zfmisc_1(k1_zfmisc_1(X141))))))&(v5_relat_1(esk10_2(X141,X142),k7_numbers)|(v1_xboole_0(X142)|~v1_prob_1(X142,X141)|~v4_prob_1(X142,X141)|~m1_subset_1(X142,k1_zfmisc_1(k1_zfmisc_1(X141))))))&(v1_funct_1(esk10_2(X141,X142))|(v1_xboole_0(X142)|~v1_prob_1(X142,X141)|~v4_prob_1(X142,X141)|~m1_subset_1(X142,k1_zfmisc_1(k1_zfmisc_1(X141))))))&(~v1_xboole_0(esk10_2(X141,X142))|(v1_xboole_0(X142)|~v1_prob_1(X142,X141)|~v4_prob_1(X142,X141)|~m1_subset_1(X142,k1_zfmisc_1(k1_zfmisc_1(X141))))))&(v1_partfun1(esk10_2(X141,X142),X142)|(v1_xboole_0(X142)|~v1_prob_1(X142,X141)|~v4_prob_1(X142,X141)|~m1_subset_1(X142,k1_zfmisc_1(k1_zfmisc_1(X141))))))&(v1_funct_2(esk10_2(X141,X142),X142,k7_numbers)|(v1_xboole_0(X142)|~v1_prob_1(X142,X141)|~v4_prob_1(X142,X141)|~m1_subset_1(X142,k1_zfmisc_1(k1_zfmisc_1(X141))))))&(v2_valued_0(esk10_2(X141,X142))|(v1_xboole_0(X142)|~v1_prob_1(X142,X141)|~v4_prob_1(X142,X141)|~m1_subset_1(X142,k1_zfmisc_1(k1_zfmisc_1(X141))))))&(v10_valued_0(esk10_2(X141,X142))|(v1_xboole_0(X142)|~v1_prob_1(X142,X141)|~v4_prob_1(X142,X141)|~m1_subset_1(X142,k1_zfmisc_1(k1_zfmisc_1(X141))))))&(v6_supinf_2(esk10_2(X141,X142))|(v1_xboole_0(X142)|~v1_prob_1(X142,X141)|~v4_prob_1(X142,X141)|~m1_subset_1(X142,k1_zfmisc_1(k1_zfmisc_1(X141))))))&(v4_measure1(esk10_2(X141,X142),X141,X142)|(v1_xboole_0(X142)|~v1_prob_1(X142,X141)|~v4_prob_1(X142,X141)|~m1_subset_1(X142,k1_zfmisc_1(k1_zfmisc_1(X141)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[rc6_measure1])])])])])).
cnf(c_0_82, plain, (r1_xxreal_0(k3_extreal1(k7_mesfunc5(X1,X2,X3,X4)),k3_supinf_2(k3_extreal1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))),k3_extreal1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)))))|v1_xboole_0(X2)|v1_xboole_0(X1)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72])])).
cnf(c_0_83, plain, (k3_extreal1(X1)=X1|~r1_xxreal_0(np__0,X1)|~m1_subset_1(X1,k7_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_74])])).
fof(c_0_84, plain, ![X28, X29, X30, X31, X32, X33]:(v1_xboole_0(X28)|(v1_xboole_0(X29)|~v1_prob_1(X29,X28)|~v4_prob_1(X29,X28)|~m1_subset_1(X29,k1_zfmisc_1(k1_zfmisc_1(X28)))|(~v1_funct_1(X30)|~v1_funct_2(X30,X29,k7_numbers)|~v10_valued_0(X30)|~v6_supinf_2(X30)|~v4_measure1(X30,X28,X29)|~m1_subset_1(X30,k1_zfmisc_1(k2_zfmisc_1(X29,k7_numbers)))|(~v1_funct_1(X31)|~m1_subset_1(X31,k1_zfmisc_1(k2_zfmisc_1(X28,k7_numbers)))|(~v1_funct_1(X32)|~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(X28,k7_numbers)))|(~v6_supinf_2(X31)|~v6_supinf_2(X32)|(~m2_subset_1(X33,k1_zfmisc_1(X28),X29)|(X33!=k1_relset_1(X28,X31)|X33!=k1_relset_1(X28,X32)|~r1_mesfunc1(X28,X29,X31,X33)|~r1_mesfunc1(X28,X29,X32,X33))|k6_mesfunc5(X28,X29,X30,k3_mesfunc1(X28,X31,X32))=k3_supinf_2(k6_mesfunc5(X28,X29,X30,X31),k6_mesfunc5(X28,X29,X30,X32))))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[l102_mesfunc5])])])])).
cnf(c_0_85, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m2_subset_1(esk5_4(X1,X2,X3,X4),k1_zfmisc_1(X1),X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~v4_measure1(X3,X1,X2)|~v1_funct_2(X3,X2,k7_numbers)|~r1_mesfunc5(X1,X2,X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_75])).
cnf(c_0_86, plain, (esk5_4(X1,X2,X3,X4)=k1_relset_1(X1,X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~v4_measure1(X3,X1,X2)|~v1_funct_2(X3,X2,k7_numbers)|~r1_mesfunc5(X1,X2,X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_87, negated_conjecture, (r1_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_88, negated_conjecture, (v4_measure1(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_89, negated_conjecture, (v6_supinf_2(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_90, negated_conjecture, (v10_valued_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_91, negated_conjecture, (v1_funct_2(esk3_0,esk2_0,k7_numbers)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_92, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_93, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_94, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_95, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_96, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_97, negated_conjecture, (v4_prob_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_98, negated_conjecture, (v1_prob_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_99, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_100, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_77])).
fof(c_0_101, plain, ![X67, X68, X69, X70]:((~r2_relset_1(X67,X68,X69,X70)|X69=X70|(~m1_subset_1(X69,k1_zfmisc_1(k2_zfmisc_1(X67,X68)))|~m1_subset_1(X70,k1_zfmisc_1(k2_zfmisc_1(X67,X68)))))&(X69!=X70|r2_relset_1(X67,X68,X69,X70)|(~m1_subset_1(X69,k1_zfmisc_1(k2_zfmisc_1(X67,X68)))|~m1_subset_1(X70,k1_zfmisc_1(k2_zfmisc_1(X67,X68)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_r2_relset_1])])])).
fof(c_0_102, plain, ![X50, X51]:(v1_xboole_0(X50)|(~v1_funct_1(X51)|~m1_subset_1(X51,k1_zfmisc_1(k2_zfmisc_1(X50,k7_numbers)))|r2_relset_1(X50,k7_numbers,k10_mesfunc1(X50,X51),k3_mesfunc1(X50,k1_mesfunc2(X50,X51),k2_mesfunc2(X50,X51))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t24_mesfunc2])])])])).
fof(c_0_103, plain, ![X45, X46, X47, X48, X49]:(v1_xboole_0(X45)|(v1_xboole_0(X46)|~v1_prob_1(X46,X45)|~v4_prob_1(X46,X45)|~m1_subset_1(X46,k1_zfmisc_1(k1_zfmisc_1(X45)))|(~v1_funct_1(X47)|~v1_funct_2(X47,X46,k7_numbers)|~v10_valued_0(X47)|~v6_supinf_2(X47)|~v4_measure1(X47,X45,X46)|~m1_subset_1(X47,k1_zfmisc_1(k2_zfmisc_1(X46,k7_numbers)))|(~v1_funct_1(X48)|~m1_subset_1(X48,k1_zfmisc_1(k2_zfmisc_1(X45,k7_numbers)))|(~v6_supinf_2(X48)|(~m2_subset_1(X49,k1_zfmisc_1(X45),X46)|(X49!=k1_relset_1(X45,X48)|~r1_mesfunc1(X45,X46,X48,X49))|k7_mesfunc5(X45,X46,X47,X48)=k6_mesfunc5(X45,X46,X47,X48))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t88_mesfunc5])])])])).
cnf(c_0_104, plain, (m1_subset_1(X1,X3)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_78])).
cnf(c_0_105, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_79])).
cnf(c_0_106, plain, (v6_supinf_2(k10_mesfunc1(X1,X2))|v1_xboole_0(X4)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X3)|~v1_funct_2(X3,X4,k7_numbers)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v4_measure1(X3,X1,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~v1_prob_1(X4,X1)|~v4_prob_1(X4,X1)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_107, plain, (v4_measure1(esk10_2(X1,X2),X1,X2)|v1_xboole_0(X2)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_108, plain, (r1_xxreal_0(k3_extreal1(k7_mesfunc5(X1,X2,X3,X4)),k3_supinf_2(k3_extreal1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))),k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83])])).
cnf(c_0_109, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|k6_mesfunc5(X1,X2,X3,k3_mesfunc1(X1,X4,X5))=k3_supinf_2(k6_mesfunc5(X1,X2,X3,X4),k6_mesfunc5(X1,X2,X3,X5))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v1_funct_1(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v4_measure1(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v6_supinf_2(X4)|~v6_supinf_2(X5)|~m2_subset_1(X6,k1_zfmisc_1(X1),X2)|X6!=k1_relset_1(X1,X4)|X6!=k1_relset_1(X1,X5)|~r1_mesfunc1(X1,X2,X4,X6)|~r1_mesfunc1(X1,X2,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_110, plain, (m2_subset_1(k1_relset_1(X1,X4),k1_zfmisc_1(X1),X2)|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_mesfunc5(X1,X2,X3,X4)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86])])).
cnf(c_0_111, negated_conjecture, (r1_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0)), inference(evalgc,[status(thm)],[c_0_87])).
cnf(c_0_112, negated_conjecture, (v4_measure1(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_88])).
cnf(c_0_113, negated_conjecture, (v6_supinf_2(esk3_0)), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_114, negated_conjecture, (v10_valued_0(esk3_0)), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_115, negated_conjecture, (v1_funct_2(esk3_0,esk2_0,k7_numbers)), inference(evalgc,[status(thm)],[c_0_91])).
cnf(c_0_116, negated_conjecture, (v1_funct_1(esk4_0)), inference(evalgc,[status(thm)],[c_0_92])).
cnf(c_0_117, negated_conjecture, (v1_funct_1(esk3_0)), inference(evalgc,[status(thm)],[c_0_93])).
cnf(c_0_118, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_94])).
cnf(c_0_119, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_95])).
cnf(c_0_120, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0)))), inference(evalgc,[status(thm)],[c_0_96])).
cnf(c_0_121, negated_conjecture, (v4_prob_1(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_97])).
cnf(c_0_122, negated_conjecture, (v1_prob_1(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_123, negated_conjecture, (~v1_xboole_0(esk2_0)), inference(evalgc,[status(thm)],[c_0_99])).
cnf(c_0_124, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_100])).
cnf(c_0_125, plain, (X3=X4|~r2_relset_1(X1,X2,X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_126, plain, (v1_xboole_0(X1)|r2_relset_1(X1,k7_numbers,k10_mesfunc1(X1,X2),k3_mesfunc1(X1,k1_mesfunc2(X1,X2),k2_mesfunc2(X1,X2)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_102])).
fof(c_0_127, plain, ![X86, X87, X88]:((v1_funct_1(k3_mesfunc1(X86,X87,X88))|(v1_xboole_0(X86)|~v1_relat_1(X87)|~v4_relat_1(X87,X86)|~v5_relat_1(X87,k7_numbers)|~v1_funct_1(X87)|~v1_relat_1(X88)|~v4_relat_1(X88,X86)|~v5_relat_1(X88,k7_numbers)|~v1_funct_1(X88)))&(m1_subset_1(k3_mesfunc1(X86,X87,X88),k1_zfmisc_1(k2_zfmisc_1(X86,k7_numbers)))|(v1_xboole_0(X86)|~v1_relat_1(X87)|~v4_relat_1(X87,X86)|~v5_relat_1(X87,k7_numbers)|~v1_funct_1(X87)|~v1_relat_1(X88)|~v4_relat_1(X88,X86)|~v5_relat_1(X88,k7_numbers)|~v1_funct_1(X88)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_mesfunc1])])])])).
cnf(c_0_128, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v1_funct_1(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v4_measure1(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v6_supinf_2(X4)|~m2_subset_1(X5,k1_zfmisc_1(X1),X2)|X5!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_103])).
fof(c_0_129, plain, ![X34, X35, X36, X37]:(v1_xboole_0(X34)|(v1_xboole_0(X35)|~v1_prob_1(X35,X34)|~v4_prob_1(X35,X34)|~m1_subset_1(X35,k1_zfmisc_1(k1_zfmisc_1(X34)))|(~v1_funct_1(X36)|~m1_subset_1(X36,k1_zfmisc_1(k2_zfmisc_1(X34,k7_numbers)))|(~m1_subset_1(X37,X35)|(~r1_mesfunc1(X34,X35,X36,X37)|~r1_tarski(X37,k1_relset_1(X34,X36))|r1_mesfunc1(X34,X35,k10_mesfunc1(X34,X36),X37)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t27_mesfunc2])])])])).
cnf(c_0_130, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_104])).
cnf(c_0_131, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_132, plain, (v1_xboole_0(X4)|v1_xboole_0(X1)|v6_supinf_2(k10_mesfunc1(X1,X2))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v1_prob_1(X4,X1)|~v4_prob_1(X4,X1)|~v4_measure1(X3,X1,X4)|~v1_funct_2(X3,X4,k7_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_133, plain, (v1_xboole_0(X2)|v4_measure1(esk10_2(X1,X2),X1,X2)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(evalgc,[status(thm)],[c_0_107])).
cnf(c_0_134, plain, (m1_subset_1(esk10_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|v1_xboole_0(X2)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_135, plain, (v1_funct_1(esk10_2(X1,X2))|v1_xboole_0(X2)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_136, plain, (v1_funct_2(esk10_2(X1,X2),X2,k7_numbers)|v1_xboole_0(X2)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_137, plain, (v10_valued_0(esk10_2(X1,X2))|v1_xboole_0(X2)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_138, plain, (v6_supinf_2(esk10_2(X1,X2))|v1_xboole_0(X2)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_139, plain, (r1_xxreal_0(k3_extreal1(k7_mesfunc5(X1,X2,X3,X4)),k3_supinf_2(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))))|v1_xboole_0(X2)|v1_xboole_0(X1)|~r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_83])])).
cnf(c_0_140, plain, (k3_supinf_2(k6_mesfunc5(X1,X2,X3,X4),k6_mesfunc5(X1,X2,X3,X5))=k6_mesfunc5(X1,X2,X3,k3_mesfunc1(X1,X4,X5))|v1_xboole_0(X2)|v1_xboole_0(X1)|X6!=k1_relset_1(X1,X5)|X6!=k1_relset_1(X1,X4)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v10_valued_0(X3)|~v6_supinf_2(X5)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~v4_measure1(X3,X1,X2)|~v1_funct_2(X3,X2,k7_numbers)|~r1_mesfunc1(X1,X2,X5,X6)|~r1_mesfunc1(X1,X2,X4,X6)|~m2_subset_1(X6,k1_zfmisc_1(X1),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_109])).
cnf(c_0_141, negated_conjecture, (m2_subset_1(k1_relset_1(esk1_0,esk4_0),k1_zfmisc_1(esk1_0),esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111]), c_0_112]), c_0_113]), c_0_114]), c_0_115]), c_0_116]), c_0_117]), c_0_118]), c_0_119]), c_0_120]), c_0_121]), c_0_122])]), c_0_123]), c_0_124])])).
cnf(c_0_142, plain, (v6_supinf_2(k1_mesfunc2(X1,X2))|v1_xboole_0(X4)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X3)|~v1_funct_2(X3,X4,k7_numbers)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v4_measure1(X3,X1,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~v1_prob_1(X4,X1)|~v4_prob_1(X4,X1)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_143, plain, (v6_supinf_2(k2_mesfunc2(X1,X2))|v1_xboole_0(X4)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X3)|~v1_funct_2(X3,X4,k7_numbers)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v4_measure1(X3,X1,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~v1_prob_1(X4,X1)|~v4_prob_1(X4,X1)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_80])).
cnf(c_0_144, plain, (X3=X4|~r2_relset_1(X1,X2,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_125])).
cnf(c_0_145, plain, (v1_xboole_0(X1)|r2_relset_1(X1,k7_numbers,k10_mesfunc1(X1,X2),k3_mesfunc1(X1,k1_mesfunc2(X1,X2),k2_mesfunc2(X1,X2)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_126])).
cnf(c_0_146, plain, (m1_subset_1(k3_mesfunc1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|v1_xboole_0(X1)|~v1_relat_1(X2)|~v4_relat_1(X2,X1)|~v5_relat_1(X2,k7_numbers)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v4_relat_1(X3,X1)|~v5_relat_1(X3,k7_numbers)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_127])).
fof(c_0_147, plain, ![X126, X127]:((v1_funct_1(k1_mesfunc2(X126,X127))|(v1_xboole_0(X126)|~v1_funct_1(X127)|~m1_subset_1(X127,k1_zfmisc_1(k2_zfmisc_1(X126,k7_numbers)))))&(m1_subset_1(k1_mesfunc2(X126,X127),k1_zfmisc_1(k2_zfmisc_1(X126,k7_numbers)))|(v1_xboole_0(X126)|~v1_funct_1(X127)|~m1_subset_1(X127,k1_zfmisc_1(k2_zfmisc_1(X126,k7_numbers)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k1_mesfunc2])])])])).
fof(c_0_148, plain, ![X124, X125]:((v1_funct_1(k2_mesfunc2(X124,X125))|(v1_xboole_0(X124)|~v1_funct_1(X125)|~m1_subset_1(X125,k1_zfmisc_1(k2_zfmisc_1(X124,k7_numbers)))))&(m1_subset_1(k2_mesfunc2(X124,X125),k1_zfmisc_1(k2_zfmisc_1(X124,k7_numbers)))|(v1_xboole_0(X124)|~v1_funct_1(X125)|~m1_subset_1(X125,k1_zfmisc_1(k2_zfmisc_1(X124,k7_numbers)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_mesfunc2])])])])).
cnf(c_0_149, plain, (k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|X5!=k1_relset_1(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v10_valued_0(X3)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~v4_measure1(X3,X1,X2)|~v1_funct_2(X3,X2,k7_numbers)|~r1_mesfunc1(X1,X2,X4,X5)|~m2_subset_1(X5,k1_zfmisc_1(X1),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_128])).
cnf(c_0_150, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|r1_mesfunc1(X1,X2,k10_mesfunc1(X1,X3),X4)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,X2)|~r1_mesfunc1(X1,X2,X3,X4)|~r1_tarski(X4,k1_relset_1(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_151, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~m2_subset_1(X1,X3,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(csr,[status(thm)],[c_0_130, c_0_131])).
cnf(c_0_152, plain, (v6_supinf_2(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|v1_xboole_0(X3)|~v6_supinf_2(esk10_2(X1,X3))|~v10_valued_0(esk10_2(X1,X3))|~v1_funct_2(esk10_2(X1,X3),X3,k7_numbers)|~v1_funct_1(esk10_2(X1,X3))|~v1_funct_1(X2)|~m1_subset_1(esk10_2(X1,X3),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X3,X1)|~v1_prob_1(X3,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_133])])).
cnf(c_0_153, plain, (v1_xboole_0(X2)|m1_subset_1(esk10_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(evalgc,[status(thm)],[c_0_134])).
cnf(c_0_154, plain, (v1_xboole_0(X2)|v1_funct_1(esk10_2(X1,X2))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(evalgc,[status(thm)],[c_0_135])).
cnf(c_0_155, plain, (v1_xboole_0(X2)|v1_funct_2(esk10_2(X1,X2),X2,k7_numbers)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(evalgc,[status(thm)],[c_0_136])).
cnf(c_0_156, plain, (v1_xboole_0(X2)|v10_valued_0(esk10_2(X1,X2))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_157, plain, (v1_xboole_0(X2)|v6_supinf_2(esk10_2(X1,X2))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(evalgc,[status(thm)],[c_0_138])).
cnf(c_0_158, plain, (r1_mesfunc1(X1,X2,X3,esk5_4(X1,X2,X4,X3))|v1_xboole_0(X2)|v1_xboole_0(X1)|~r1_mesfunc5(X1,X2,X4,X3)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X4)|~v1_funct_2(X4,X2,k7_numbers)|~v10_valued_0(X4)|~v6_supinf_2(X4)|~v4_measure1(X4,X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_159, plain, (r1_xxreal_0(k3_extreal1(k7_mesfunc5(X1,X2,X3,X4)),k3_supinf_2(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)))|~r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_56])])).
cnf(c_0_160, negated_conjecture, (k3_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,X3))=k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X2,X3))|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X3)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X3,k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X3)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_141]), c_0_120]), c_0_121]), c_0_122])]), c_0_124]), c_0_123])])).
cnf(c_0_161, plain, (v1_xboole_0(X4)|v1_xboole_0(X1)|v6_supinf_2(k1_mesfunc2(X1,X2))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v1_prob_1(X4,X1)|~v4_prob_1(X4,X1)|~v4_measure1(X3,X1,X4)|~v1_funct_2(X3,X4,k7_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_142])).
cnf(c_0_162, plain, (v1_xboole_0(X4)|v1_xboole_0(X1)|v6_supinf_2(k2_mesfunc2(X1,X2))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v1_prob_1(X4,X1)|~v4_prob_1(X4,X1)|~v4_measure1(X3,X1,X4)|~v1_funct_2(X3,X4,k7_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_143])).
cnf(c_0_163, plain, (k3_mesfunc1(X1,k1_mesfunc2(X1,X2),k2_mesfunc2(X1,X2))=k10_mesfunc1(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(k3_mesfunc1(X1,k1_mesfunc2(X1,X2),k2_mesfunc2(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_145])])).
cnf(c_0_164, plain, (v1_xboole_0(X1)|m1_subset_1(k3_mesfunc1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)), inference(evalgc,[status(thm)],[c_0_146])).
cnf(c_0_165, plain, (v1_funct_1(k1_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_147])).
cnf(c_0_166, plain, (v1_funct_1(k2_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_148])).
fof(c_0_167, plain, ![X22, X23, X24, X25, X26]:(v1_xboole_0(X22)|(v1_xboole_0(X23)|~v1_prob_1(X23,X22)|~v4_prob_1(X23,X22)|~m1_subset_1(X23,k1_zfmisc_1(k1_zfmisc_1(X22)))|(~v1_funct_1(X24)|~v1_funct_2(X24,X23,k7_numbers)|~v10_valued_0(X24)|~v6_supinf_2(X24)|~v4_measure1(X24,X22,X23)|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(X23,k7_numbers)))|(~v1_funct_1(X25)|~m1_subset_1(X25,k1_zfmisc_1(k2_zfmisc_1(X22,k7_numbers)))|(~v6_supinf_2(X25)|(~m2_subset_1(X26,k1_zfmisc_1(X22),X23)|(X26!=k1_relset_1(X22,X25)|~r1_mesfunc1(X22,X23,X25,X26))|r1_xxreal_0(k6_numbers,k6_mesfunc5(X22,X23,X24,X25)))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t79_mesfunc5])])])])).
cnf(c_0_168, negated_conjecture, (k7_mesfunc5(esk1_0,esk2_0,X1,X2)=k6_mesfunc5(esk1_0,esk2_0,X1,X2)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_141]), c_0_120]), c_0_121]), c_0_122])]), c_0_124]), c_0_123])])).
cnf(c_0_169, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|r1_mesfunc1(X1,X2,k10_mesfunc1(X1,X3),X4)|~v1_funct_1(X3)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X4,X2)|~r1_mesfunc1(X1,X2,X3,X4)|~r1_tarski(X4,k1_relset_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_150])).
cnf(c_0_170, negated_conjecture, (m1_subset_1(k1_relset_1(esk1_0,esk4_0),esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_141]), c_0_120])]), c_0_123])])).
cnf(c_0_171, plain, (v6_supinf_2(k10_mesfunc1(X1,X2))|v1_xboole_0(X3)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X3,X1)|~v1_prob_1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_152, c_0_153]), c_0_154]), c_0_155]), c_0_156]), c_0_157])).
fof(c_0_172, plain, ![X128]:r1_tarski(X128,X128), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r1_tarski])])).
cnf(c_0_173, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|r1_mesfunc1(X1,X2,X3,esk5_4(X1,X2,X4,X3))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v10_valued_0(X4)|~v6_supinf_2(X4)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~v4_measure1(X4,X1,X2)|~v1_funct_2(X4,X2,k7_numbers)|~r1_mesfunc5(X1,X2,X4,X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_158])).
cnf(c_0_174, negated_conjecture, (r1_xxreal_0(k3_extreal1(k7_mesfunc5(esk1_0,esk2_0,X1,X2)),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,k1_mesfunc2(esk1_0,X2),k2_mesfunc2(esk1_0,X2))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~v6_supinf_2(k2_mesfunc2(esk1_0,X2))|~v6_supinf_2(k1_mesfunc2(esk1_0,X2))|~v1_funct_1(k2_mesfunc2(esk1_0,X2))|~v1_funct_1(k1_mesfunc2(esk1_0,X2))|~m1_subset_1(k2_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_160]), c_0_120]), c_0_121]), c_0_122])]), c_0_123]), c_0_124])])).
cnf(c_0_175, negated_conjecture, (v6_supinf_2(k1_mesfunc2(esk1_0,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_112]), c_0_113]), c_0_114]), c_0_115]), c_0_117]), c_0_119]), c_0_120]), c_0_121]), c_0_122])]), c_0_124]), c_0_123])])).
cnf(c_0_176, negated_conjecture, (v6_supinf_2(k2_mesfunc2(esk1_0,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_112]), c_0_113]), c_0_114]), c_0_115]), c_0_117]), c_0_119]), c_0_120]), c_0_121]), c_0_122])]), c_0_124]), c_0_123])])).
cnf(c_0_177, plain, (k3_mesfunc1(X1,k1_mesfunc2(X1,X2),k2_mesfunc2(X1,X2))=k10_mesfunc1(X1,X2)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v5_relat_1(k2_mesfunc2(X1,X2),k7_numbers)|~v5_relat_1(k1_mesfunc2(X1,X2),k7_numbers)|~v4_relat_1(k2_mesfunc2(X1,X2),X1)|~v4_relat_1(k1_mesfunc2(X1,X2),X1)|~v1_relat_1(k2_mesfunc2(X1,X2))|~v1_relat_1(k1_mesfunc2(X1,X2))|~v1_funct_1(k2_mesfunc2(X1,X2))|~v1_funct_1(k1_mesfunc2(X1,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_164])])).
cnf(c_0_178, plain, (v1_xboole_0(X1)|v1_funct_1(k1_mesfunc2(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_165])).
cnf(c_0_179, plain, (v1_xboole_0(X1)|v1_funct_1(k2_mesfunc2(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_166])).
fof(c_0_180, plain, ![X71, X72, X73]:(~m1_subset_1(X73,k1_zfmisc_1(k2_zfmisc_1(X71,X72)))|v1_relat_1(X73)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_181, plain, ![X74, X75, X76]:((v4_relat_1(X76,X74)|~m1_subset_1(X76,k1_zfmisc_1(k2_zfmisc_1(X74,X75))))&(v5_relat_1(X76,X75)|~m1_subset_1(X76,k1_zfmisc_1(k2_zfmisc_1(X74,X75))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
cnf(c_0_182, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|r1_xxreal_0(k6_numbers,k6_mesfunc5(X1,X2,X3,X4))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v1_funct_1(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v4_measure1(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v6_supinf_2(X4)|~m2_subset_1(X5,k1_zfmisc_1(X1),X2)|X5!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,X5)), inference(split_conjunct,[status(thm)],[c_0_167])).
cnf(c_0_183, negated_conjecture, (~r1_xxreal_0(k3_extreal1(k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0)),k7_mesfunc5(esk1_0,esk2_0,esk3_0,k10_mesfunc1(esk1_0,esk4_0)))), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_184, negated_conjecture, (k7_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,X2))=k6_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,X2))|k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(k10_mesfunc1(esk1_0,X2))|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k10_mesfunc1(esk1_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k10_mesfunc1(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|~r1_tarski(k1_relset_1(esk1_0,esk4_0),k1_relset_1(esk1_0,X2))|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_169]), c_0_120]), c_0_170]), c_0_121]), c_0_122])]), c_0_124]), c_0_123])])).
cnf(c_0_185, negated_conjecture, (v6_supinf_2(k10_mesfunc1(esk1_0,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_121]), c_0_120]), c_0_122])]), c_0_123]), c_0_124])])).
cnf(c_0_186, plain, (r1_tarski(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_172])).
cnf(c_0_187, plain, (r1_mesfunc1(X1,X2,X3,k1_relset_1(X1,X3))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_mesfunc5(X1,X2,X4,X3)|~v4_measure1(X4,X1,X2)|~v6_supinf_2(X4)|~v10_valued_0(X4)|~v1_funct_2(X4,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_86])])).
cnf(c_0_188, negated_conjecture, (r1_xxreal_0(k3_extreal1(k7_mesfunc5(esk1_0,esk2_0,X1,X2)),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,k1_mesfunc2(esk1_0,X2),k2_mesfunc2(esk1_0,X2))))|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k2_mesfunc2(esk1_0,X2))|~v1_funct_1(k1_mesfunc2(esk1_0,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_174, c_0_175]), c_0_176])).
cnf(c_0_189, plain, (k3_mesfunc1(X1,k1_mesfunc2(X1,X2),k2_mesfunc2(X1,X2))=k10_mesfunc1(X1,X2)|v1_xboole_0(X1)|~v5_relat_1(k2_mesfunc2(X1,X2),k7_numbers)|~v5_relat_1(k1_mesfunc2(X1,X2),k7_numbers)|~v4_relat_1(k2_mesfunc2(X1,X2),X1)|~v4_relat_1(k1_mesfunc2(X1,X2),X1)|~v1_relat_1(k2_mesfunc2(X1,X2))|~v1_relat_1(k1_mesfunc2(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_177, c_0_178]), c_0_179])).
cnf(c_0_190, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_180])).
cnf(c_0_191, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_181])).
cnf(c_0_192, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_181])).
cnf(c_0_193, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|r1_xxreal_0(k6_numbers,k6_mesfunc5(X1,X2,X3,X4))|X5!=k1_relset_1(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v10_valued_0(X3)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~v4_measure1(X3,X1,X2)|~v1_funct_2(X3,X2,k7_numbers)|~r1_mesfunc1(X1,X2,X4,X5)|~m2_subset_1(X5,k1_zfmisc_1(X1),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_182])).
cnf(c_0_194, negated_conjecture, (~r1_xxreal_0(k3_extreal1(k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0)),k7_mesfunc5(esk1_0,esk2_0,esk3_0,k10_mesfunc1(esk1_0,esk4_0)))), inference(evalgc,[status(thm)],[c_0_183])).
cnf(c_0_195, negated_conjecture, (k7_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,X2))=k6_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,X2))|k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~r1_tarski(k1_relset_1(esk1_0,esk4_0),k1_relset_1(esk1_0,X2))|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k10_mesfunc1(esk1_0,X2))|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(k10_mesfunc1(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(csr,[status(thm)],[c_0_184, c_0_185])).
cnf(c_0_196, plain, (r1_tarski(X1,X1)), inference(evalgc,[status(thm)],[c_0_186])).
cnf(c_0_197, negated_conjecture, (r1_mesfunc1(esk1_0,esk2_0,esk4_0,k1_relset_1(esk1_0,esk4_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_111]), c_0_112]), c_0_113]), c_0_114]), c_0_115]), c_0_117]), c_0_116]), c_0_119]), c_0_118]), c_0_120]), c_0_121]), c_0_122])]), c_0_123]), c_0_124])])).
cnf(c_0_198, negated_conjecture, (r1_xxreal_0(k3_extreal1(k7_mesfunc5(esk1_0,esk2_0,X1,X2)),k6_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,X2)))|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k2_mesfunc2(esk1_0,X2))|~v1_funct_1(k1_mesfunc2(esk1_0,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|~v5_relat_1(k2_mesfunc2(esk1_0,X2),k7_numbers)|~v5_relat_1(k1_mesfunc2(esk1_0,X2),k7_numbers)|~v4_relat_1(k2_mesfunc2(esk1_0,X2),esk1_0)|~v4_relat_1(k1_mesfunc2(esk1_0,X2),esk1_0)|~v1_relat_1(k2_mesfunc2(esk1_0,X2))|~v1_relat_1(k1_mesfunc2(esk1_0,X2))|~m1_subset_1(k10_mesfunc1(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_189]), c_0_124])])).
cnf(c_0_199, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_190])).
cnf(c_0_200, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_191])).
cnf(c_0_201, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(evalgc,[status(thm)],[c_0_192])).
cnf(c_0_202, plain, (r1_xxreal_0(k1_supinf_2,k6_mesfunc5(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|X5!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,X5)|~m2_subset_1(X5,k1_zfmisc_1(X1),X2)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)), inference(rw,[status(thm)],[c_0_193, c_0_61])).
cnf(c_0_203, negated_conjecture, (~r1_xxreal_0(k3_extreal1(k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0)),k6_mesfunc5(esk1_0,esk2_0,esk3_0,k10_mesfunc1(esk1_0,esk4_0)))|k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_195]), c_0_196]), c_0_197]), c_0_112]), c_0_113]), c_0_114]), c_0_115]), c_0_117]), c_0_116]), c_0_119]), c_0_118])])])).
cnf(c_0_204, negated_conjecture, (r1_xxreal_0(k3_extreal1(k7_mesfunc5(esk1_0,esk2_0,X1,X2)),k6_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,X2)))|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k2_mesfunc2(esk1_0,X2))|~v1_funct_1(k1_mesfunc2(esk1_0,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k10_mesfunc1(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_198, c_0_199]), c_0_199]), c_0_200]), c_0_200]), c_0_201]), c_0_201])).
cnf(c_0_205, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|X5!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,X5)|~m2_subset_1(X5,k1_zfmisc_1(X1),X2)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_202, c_0_74])])).
cnf(c_0_206, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_204]), c_0_112]), c_0_113]), c_0_114]), c_0_115]), c_0_116]), c_0_117]), c_0_118]), c_0_119])])])).
cnf(c_0_207, negated_conjecture, (r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X2))|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_141]), c_0_120]), c_0_121]), c_0_122])]), c_0_123]), c_0_124])])).
fof(c_0_208, plain, ![X38, X39, X40, X41]:(v1_xboole_0(X38)|(~v1_funct_1(X39)|~m1_subset_1(X39,k1_zfmisc_1(k2_zfmisc_1(X38,k7_numbers)))|(v1_xboole_0(X40)|~v1_prob_1(X40,X38)|~v4_prob_1(X40,X38)|~m1_subset_1(X40,k1_zfmisc_1(k1_zfmisc_1(X38)))|(~m1_subset_1(X41,X40)|(~r1_mesfunc1(X38,X40,X39,X41)|~r1_tarski(X41,k1_relset_1(X38,X39))|r1_mesfunc1(X38,X40,k2_mesfunc2(X38,X39),X41)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t26_mesfunc2])])])])).
cnf(c_0_209, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)))|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~v6_supinf_2(k2_mesfunc2(esk1_0,esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_207]), c_0_112]), c_0_113]), c_0_114]), c_0_115]), c_0_117]), c_0_119])])])).
cnf(c_0_210, plain, (v1_xboole_0(X1)|v1_xboole_0(X3)|r1_mesfunc1(X1,X3,k2_mesfunc2(X1,X2),X4)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_prob_1(X3,X1)|~v4_prob_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X3)|~r1_mesfunc1(X1,X3,X2,X4)|~r1_tarski(X4,k1_relset_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_208])).
fof(c_0_211, plain, ![X80, X81, X82, X83]:(v1_xboole_0(X80)|(~v1_funct_1(X81)|~m1_subset_1(X81,k1_zfmisc_1(k2_zfmisc_1(X80,k7_numbers)))|(v1_xboole_0(X82)|~v1_prob_1(X82,X80)|~v4_prob_1(X82,X80)|~m1_subset_1(X82,k1_zfmisc_1(k1_zfmisc_1(X80)))|(~m1_subset_1(X83,X82)|(~r1_mesfunc1(X80,X82,X81,X83)|r1_mesfunc1(X80,X82,k1_mesfunc2(X80,X81),X83)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t25_mesfunc2])])])])).
cnf(c_0_212, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~v6_supinf_2(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~v6_supinf_2(k1_mesfunc2(esk1_0,esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_207]), c_0_112]), c_0_113]), c_0_114]), c_0_115]), c_0_117]), c_0_119])])])).
cnf(c_0_213, plain, (v1_xboole_0(X3)|v1_xboole_0(X1)|r1_mesfunc1(X1,X3,k2_mesfunc2(X1,X2),X4)|~v1_funct_1(X2)|~v1_prob_1(X3,X1)|~v4_prob_1(X3,X1)|~m1_subset_1(X4,X3)|~r1_mesfunc1(X1,X3,X2,X4)|~r1_tarski(X4,k1_relset_1(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_210])).
cnf(c_0_214, plain, (v1_xboole_0(X1)|v1_xboole_0(X3)|r1_mesfunc1(X1,X3,k1_mesfunc2(X1,X2),X4)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_prob_1(X3,X1)|~v4_prob_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,X3)|~r1_mesfunc1(X1,X3,X2,X4)), inference(split_conjunct,[status(thm)],[c_0_211])).
fof(c_0_215, plain, ![X167, X168, X169, X170]:(v1_xboole_0(X167)|v1_xboole_0(X168)|~v1_prob_1(X168,X167)|~v4_prob_1(X168,X167)|~m1_subset_1(X168,k1_zfmisc_1(k1_zfmisc_1(X167)))|~v1_funct_1(X169)|~v1_funct_2(X169,X168,k7_numbers)|~v10_valued_0(X169)|~v6_supinf_2(X169)|~v4_measure1(X169,X167,X168)|~m1_subset_1(X169,k1_zfmisc_1(k2_zfmisc_1(X168,k7_numbers)))|~v1_funct_1(X170)|~m1_subset_1(X170,k1_zfmisc_1(k2_zfmisc_1(X167,k7_numbers)))|m1_subset_1(k6_mesfunc5(X167,X168,X169,X170),k7_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_mesfunc5])])])).
cnf(c_0_216, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~v6_supinf_2(k2_mesfunc2(esk1_0,esk4_0))|~v6_supinf_2(k1_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_213]), c_0_196]), c_0_197]), c_0_116]), c_0_118]), c_0_120]), c_0_170]), c_0_121]), c_0_122])]), c_0_124]), c_0_123])])).
cnf(c_0_217, plain, (v1_xboole_0(X3)|v1_xboole_0(X1)|r1_mesfunc1(X1,X3,k1_mesfunc2(X1,X2),X4)|~v1_funct_1(X2)|~v1_prob_1(X3,X1)|~v4_prob_1(X3,X1)|~m1_subset_1(X4,X3)|~r1_mesfunc1(X1,X3,X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_214])).
cnf(c_0_218, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(k6_mesfunc5(X1,X2,X3,X4),k7_numbers)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v1_funct_1(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v4_measure1(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_215])).
fof(c_0_219, plain, ![X192, X193]:(~m1_subset_1(X193,k1_zfmisc_1(k1_zfmisc_1(X192)))|(~v1_prob_1(X193,X192)|~v3_measure1(X193,X192)|v4_prob_1(X193,X192))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_measure1])])).
fof(c_0_220, plain, ![X243]:(((m1_subset_1(esk14_1(X243),k1_zfmisc_1(k1_zfmisc_1(X243)))&~v1_xboole_0(esk14_1(X243)))&v1_prob_1(esk14_1(X243),X243))&v3_measure1(esk14_1(X243),X243)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[rc4_measure1])])])).
fof(c_0_221, plain, ![X89, X90]:((v1_funct_1(k10_mesfunc1(X89,X90))|(v1_xboole_0(X89)|~v1_relat_1(X90)|~v4_relat_1(X90,X89)|~v5_relat_1(X90,k7_numbers)|~v1_funct_1(X90)))&(m1_subset_1(k10_mesfunc1(X89,X90),k1_zfmisc_1(k2_zfmisc_1(X89,k7_numbers)))|(v1_xboole_0(X89)|~v1_relat_1(X90)|~v4_relat_1(X90,X89)|~v5_relat_1(X90,k7_numbers)|~v1_funct_1(X90)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k10_mesfunc1])])])])).
fof(c_0_222, plain, ![X17, X18, X19, X20]:(((k9_xtuple_0(X19)=k9_xtuple_0(X18)|X19!=k10_mesfunc1(X17,X18)|(~v1_funct_1(X19)|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(X17,k7_numbers))))|(~v1_relat_1(X18)|~v4_relat_1(X18,X17)|~v5_relat_1(X18,k7_numbers)|~v1_funct_1(X18))|v1_xboole_0(X17))&(~m1_subset_1(X20,X17)|(~r2_hidden(X20,k9_xtuple_0(X19))|k12_supinf_2(X19,X20)=k3_extreal1(k12_supinf_2(X18,X20)))|X19!=k10_mesfunc1(X17,X18)|(~v1_funct_1(X19)|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(X17,k7_numbers))))|(~v1_relat_1(X18)|~v4_relat_1(X18,X17)|~v5_relat_1(X18,k7_numbers)|~v1_funct_1(X18))|v1_xboole_0(X17)))&((m1_subset_1(esk6_3(X17,X18,X19),X17)|k9_xtuple_0(X19)!=k9_xtuple_0(X18)|X19=k10_mesfunc1(X17,X18)|(~v1_funct_1(X19)|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(X17,k7_numbers))))|(~v1_relat_1(X18)|~v4_relat_1(X18,X17)|~v5_relat_1(X18,k7_numbers)|~v1_funct_1(X18))|v1_xboole_0(X17))&((r2_hidden(esk6_3(X17,X18,X19),k9_xtuple_0(X19))|k9_xtuple_0(X19)!=k9_xtuple_0(X18)|X19=k10_mesfunc1(X17,X18)|(~v1_funct_1(X19)|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(X17,k7_numbers))))|(~v1_relat_1(X18)|~v4_relat_1(X18,X17)|~v5_relat_1(X18,k7_numbers)|~v1_funct_1(X18))|v1_xboole_0(X17))&(k12_supinf_2(X19,esk6_3(X17,X18,X19))!=k3_extreal1(k12_supinf_2(X18,esk6_3(X17,X18,X19)))|k9_xtuple_0(X19)!=k9_xtuple_0(X18)|X19=k10_mesfunc1(X17,X18)|(~v1_funct_1(X19)|~m1_subset_1(X19,k1_zfmisc_1(k2_zfmisc_1(X17,k7_numbers))))|(~v1_relat_1(X18)|~v4_relat_1(X18,X17)|~v5_relat_1(X18,k7_numbers)|~v1_funct_1(X18))|v1_xboole_0(X17))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_mesfunc1])])])])])])).
cnf(c_0_223, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v6_supinf_2(k2_mesfunc2(esk1_0,esk4_0))|~v6_supinf_2(k1_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_217]), c_0_197]), c_0_116]), c_0_118]), c_0_120]), c_0_170]), c_0_121]), c_0_122])]), c_0_124]), c_0_123])])).
cnf(c_0_224, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(k6_mesfunc5(X1,X2,X3,X4),k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~v4_measure1(X3,X1,X2)|~v1_funct_2(X3,X2,k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_218])).
cnf(c_0_225, plain, (v4_prob_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v1_prob_1(X1,X2)|~v3_measure1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_219])).
cnf(c_0_226, plain, (v3_measure1(esk14_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_220])).
cnf(c_0_227, plain, (m1_subset_1(esk14_1(X1),k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_220])).
cnf(c_0_228, plain, (v1_prob_1(esk14_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_220])).
fof(c_0_229, plain, ![X110, X111]:(~v1_relat_1(X111)|~v4_relat_1(X111,X110)|k1_relset_1(X110,X111)=k9_xtuple_0(X111)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_230, plain, (m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|v1_xboole_0(X1)|~v1_relat_1(X2)|~v4_relat_1(X2,X1)|~v5_relat_1(X2,k7_numbers)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_221])).
cnf(c_0_231, plain, (k9_xtuple_0(X1)=k9_xtuple_0(X2)|v1_xboole_0(X3)|X1!=k10_mesfunc1(X3,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~v1_relat_1(X2)|~v4_relat_1(X2,X3)|~v5_relat_1(X2,k7_numbers)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_222])).
cnf(c_0_232, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v6_supinf_2(k2_mesfunc2(esk1_0,esk4_0))|~v6_supinf_2(k1_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_224]), c_0_112]), c_0_113]), c_0_114]), c_0_115]), c_0_117]), c_0_119]), c_0_120]), c_0_121]), c_0_122])]), c_0_124]), c_0_123])])).
cnf(c_0_233, plain, (v6_supinf_2(k2_mesfunc2(X1,X2))|v1_xboole_0(X1)|v1_xboole_0(X3)|~v6_supinf_2(esk10_2(X1,X3))|~v10_valued_0(esk10_2(X1,X3))|~v1_funct_2(esk10_2(X1,X3),X3,k7_numbers)|~v1_funct_1(esk10_2(X1,X3))|~v1_funct_1(X2)|~m1_subset_1(esk10_2(X1,X3),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X3,X1)|~v1_prob_1(X3,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_133])])).
cnf(c_0_234, plain, (v4_prob_1(X1,X2)|~v1_prob_1(X1,X2)|~v3_measure1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(evalgc,[status(thm)],[c_0_225])).
cnf(c_0_235, plain, (v3_measure1(esk14_1(X1),X1)), inference(evalgc,[status(thm)],[c_0_226])).
cnf(c_0_236, plain, (m1_subset_1(esk14_1(X1),k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(evalgc,[status(thm)],[c_0_227])).
cnf(c_0_237, plain, (v1_prob_1(esk14_1(X1),X1)), inference(evalgc,[status(thm)],[c_0_228])).
cnf(c_0_238, plain, (~v1_xboole_0(esk14_1(X1))), inference(split_conjunct,[status(thm)],[c_0_220])).
cnf(c_0_239, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_229])).
cnf(c_0_240, plain, (v1_xboole_0(X1)|m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~v1_relat_1(X2)|~v4_relat_1(X2,X1)|~v5_relat_1(X2,k7_numbers)), inference(evalgc,[status(thm)],[c_0_230])).
cnf(c_0_241, plain, (k9_xtuple_0(X1)=k9_xtuple_0(X2)|v1_xboole_0(X3)|X1!=k10_mesfunc1(X3,X2)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v4_relat_1(X2,X3)|~v5_relat_1(X2,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_231])).
cnf(c_0_242, plain, (v1_funct_1(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v1_relat_1(X2)|~v4_relat_1(X2,X1)|~v5_relat_1(X2,k7_numbers)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_221])).
cnf(c_0_243, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v6_supinf_2(k2_mesfunc2(esk1_0,esk4_0))|~v6_supinf_2(k1_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_224]), c_0_112]), c_0_113]), c_0_114]), c_0_115]), c_0_117]), c_0_119]), c_0_120]), c_0_121]), c_0_122])]), c_0_124]), c_0_123])])).
cnf(c_0_244, negated_conjecture, (v5_relat_1(esk4_0,k7_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_118])])).
cnf(c_0_245, negated_conjecture, (v4_relat_1(esk4_0,esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_118])])).
cnf(c_0_246, negated_conjecture, (v1_relat_1(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_118])])).
cnf(c_0_247, plain, (v6_supinf_2(k2_mesfunc2(X1,X2))|v1_xboole_0(X3)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X3,X1)|~v1_prob_1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_233, c_0_153]), c_0_154]), c_0_155]), c_0_156]), c_0_157])).
cnf(c_0_248, plain, (v4_prob_1(esk14_1(X1),X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_235]), c_0_236]), c_0_237])])])).
cnf(c_0_249, plain, (~v1_xboole_0(esk14_1(X1))), inference(evalgc,[status(thm)],[c_0_238])).
cnf(c_0_250, plain, (v6_supinf_2(k1_mesfunc2(X1,X2))|v1_xboole_0(X1)|v1_xboole_0(X3)|~v6_supinf_2(esk10_2(X1,X3))|~v10_valued_0(esk10_2(X1,X3))|~v1_funct_2(esk10_2(X1,X3),X3,k7_numbers)|~v1_funct_1(esk10_2(X1,X3))|~v1_funct_1(X2)|~m1_subset_1(esk10_2(X1,X3),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X3,X1)|~v1_prob_1(X3,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_133])])).
cnf(c_0_251, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_239])).
cnf(c_0_252, plain, (v4_relat_1(k10_mesfunc1(X1,X2),X1)|v1_xboole_0(X1)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_240])])).
cnf(c_0_253, plain, (k9_xtuple_0(k10_mesfunc1(X1,X2))=k9_xtuple_0(X3)|v1_xboole_0(X1)|k10_mesfunc1(X1,X2)!=k10_mesfunc1(X1,X3)|~v5_relat_1(X3,k7_numbers)|~v4_relat_1(X3,X1)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v1_funct_1(k10_mesfunc1(X1,X2))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_240])])).
cnf(c_0_254, plain, (v1_xboole_0(X1)|v1_funct_1(k10_mesfunc1(X1,X2))|~v1_funct_1(X2)|~v1_relat_1(X2)|~v4_relat_1(X2,X1)|~v5_relat_1(X2,k7_numbers)), inference(evalgc,[status(thm)],[c_0_242])).
cnf(c_0_255, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v6_supinf_2(k2_mesfunc2(esk1_0,esk4_0))|~v6_supinf_2(k1_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_240]), c_0_244]), c_0_245]), c_0_246]), c_0_116])]), c_0_124])])).
cnf(c_0_256, plain, (v6_supinf_2(k2_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_248]), c_0_236]), c_0_237])]), c_0_249])])).
cnf(c_0_257, plain, (v6_supinf_2(k1_mesfunc2(X1,X2))|v1_xboole_0(X3)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X3,X1)|~v1_prob_1(X3,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_250, c_0_153]), c_0_154]), c_0_155]), c_0_156]), c_0_157])).
cnf(c_0_258, plain, (k1_relset_1(X1,k10_mesfunc1(X1,X2))=k9_xtuple_0(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v1_relat_1(k10_mesfunc1(X1,X2))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_252])])).
cnf(c_0_259, plain, (v1_relat_1(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_240])])).
cnf(c_0_260, plain, (k9_xtuple_0(k10_mesfunc1(X1,X2))=k9_xtuple_0(X3)|v1_xboole_0(X1)|k10_mesfunc1(X1,X2)!=k10_mesfunc1(X1,X3)|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2)), inference(csr,[status(thm)],[c_0_253, c_0_254])).
fof(c_0_261, plain, ![X57, X58, X59, X60]:(((k1_relset_1(X57,X59)=k1_relset_1(X57,X58)|X59!=k2_mesfunc2(X57,X58)|(~v1_funct_1(X59)|~m1_subset_1(X59,k1_zfmisc_1(k2_zfmisc_1(X57,k7_numbers))))|(~v1_funct_1(X58)|~m1_subset_1(X58,k1_zfmisc_1(k2_zfmisc_1(X57,k7_numbers))))|v1_xboole_0(X57))&(~m1_subset_1(X60,X57)|(~r2_hidden(X60,k1_relset_1(X57,X59))|k12_supinf_2(X59,X60)=k1_extreal2(k2_supinf_2(k12_supinf_2(X58,X60)),k1_supinf_2))|X59!=k2_mesfunc2(X57,X58)|(~v1_funct_1(X59)|~m1_subset_1(X59,k1_zfmisc_1(k2_zfmisc_1(X57,k7_numbers))))|(~v1_funct_1(X58)|~m1_subset_1(X58,k1_zfmisc_1(k2_zfmisc_1(X57,k7_numbers))))|v1_xboole_0(X57)))&((m1_subset_1(esk9_3(X57,X58,X59),X57)|k1_relset_1(X57,X59)!=k1_relset_1(X57,X58)|X59=k2_mesfunc2(X57,X58)|(~v1_funct_1(X59)|~m1_subset_1(X59,k1_zfmisc_1(k2_zfmisc_1(X57,k7_numbers))))|(~v1_funct_1(X58)|~m1_subset_1(X58,k1_zfmisc_1(k2_zfmisc_1(X57,k7_numbers))))|v1_xboole_0(X57))&((r2_hidden(esk9_3(X57,X58,X59),k1_relset_1(X57,X59))|k1_relset_1(X57,X59)!=k1_relset_1(X57,X58)|X59=k2_mesfunc2(X57,X58)|(~v1_funct_1(X59)|~m1_subset_1(X59,k1_zfmisc_1(k2_zfmisc_1(X57,k7_numbers))))|(~v1_funct_1(X58)|~m1_subset_1(X58,k1_zfmisc_1(k2_zfmisc_1(X57,k7_numbers))))|v1_xboole_0(X57))&(k12_supinf_2(X59,esk9_3(X57,X58,X59))!=k1_extreal2(k2_supinf_2(k12_supinf_2(X58,esk9_3(X57,X58,X59))),k1_supinf_2)|k1_relset_1(X57,X59)!=k1_relset_1(X57,X58)|X59=k2_mesfunc2(X57,X58)|(~v1_funct_1(X59)|~m1_subset_1(X59,k1_zfmisc_1(k2_zfmisc_1(X57,k7_numbers))))|(~v1_funct_1(X58)|~m1_subset_1(X58,k1_zfmisc_1(k2_zfmisc_1(X57,k7_numbers))))|v1_xboole_0(X57))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d3_mesfunc2])])])])])])).
cnf(c_0_262, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v6_supinf_2(k1_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_256]), c_0_116]), c_0_118])]), c_0_124])])).
cnf(c_0_263, plain, (v6_supinf_2(k1_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_248]), c_0_236]), c_0_237])]), c_0_249])])).
cnf(c_0_264, plain, (k1_relset_1(X1,k10_mesfunc1(X1,X2))=k9_xtuple_0(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(csr,[status(thm)],[c_0_258, c_0_259])).
cnf(c_0_265, plain, (k9_xtuple_0(k10_mesfunc1(X1,X2))=k9_xtuple_0(X2)|v1_xboole_0(X1)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_260])])).
cnf(c_0_266, plain, (k1_relset_1(X1,X2)=k1_relset_1(X1,X3)|v1_xboole_0(X1)|X2!=k2_mesfunc2(X1,X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_261])).
cnf(c_0_267, plain, (m1_subset_1(k2_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_148])).
fof(c_0_268, plain, ![X52, X53, X54, X55]:(((k1_relset_1(X52,X54)=k1_relset_1(X52,X53)|X54!=k1_mesfunc2(X52,X53)|(~v1_funct_1(X54)|~m1_subset_1(X54,k1_zfmisc_1(k2_zfmisc_1(X52,k7_numbers))))|(~v1_funct_1(X53)|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(X52,k7_numbers))))|v1_xboole_0(X52))&(~m1_subset_1(X55,X52)|(~r2_hidden(X55,k1_relset_1(X52,X54))|k12_supinf_2(X54,X55)=k1_extreal2(k12_supinf_2(X53,X55),k1_supinf_2))|X54!=k1_mesfunc2(X52,X53)|(~v1_funct_1(X54)|~m1_subset_1(X54,k1_zfmisc_1(k2_zfmisc_1(X52,k7_numbers))))|(~v1_funct_1(X53)|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(X52,k7_numbers))))|v1_xboole_0(X52)))&((m1_subset_1(esk8_3(X52,X53,X54),X52)|k1_relset_1(X52,X54)!=k1_relset_1(X52,X53)|X54=k1_mesfunc2(X52,X53)|(~v1_funct_1(X54)|~m1_subset_1(X54,k1_zfmisc_1(k2_zfmisc_1(X52,k7_numbers))))|(~v1_funct_1(X53)|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(X52,k7_numbers))))|v1_xboole_0(X52))&((r2_hidden(esk8_3(X52,X53,X54),k1_relset_1(X52,X54))|k1_relset_1(X52,X54)!=k1_relset_1(X52,X53)|X54=k1_mesfunc2(X52,X53)|(~v1_funct_1(X54)|~m1_subset_1(X54,k1_zfmisc_1(k2_zfmisc_1(X52,k7_numbers))))|(~v1_funct_1(X53)|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(X52,k7_numbers))))|v1_xboole_0(X52))&(k12_supinf_2(X54,esk8_3(X52,X53,X54))!=k1_extreal2(k12_supinf_2(X53,esk8_3(X52,X53,X54)),k1_supinf_2)|k1_relset_1(X52,X54)!=k1_relset_1(X52,X53)|X54=k1_mesfunc2(X52,X53)|(~v1_funct_1(X54)|~m1_subset_1(X54,k1_zfmisc_1(k2_zfmisc_1(X52,k7_numbers))))|(~v1_funct_1(X53)|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(X52,k7_numbers))))|v1_xboole_0(X52))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_mesfunc2])])])])])])).
cnf(c_0_269, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_263]), c_0_116]), c_0_118])]), c_0_124])])).
cnf(c_0_270, plain, (k1_relset_1(X1,k10_mesfunc1(X1,X2))=k9_xtuple_0(X2)|v1_xboole_0(X1)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_265])])).
cnf(c_0_271, negated_conjecture, (k9_xtuple_0(esk4_0)=k1_relset_1(esk1_0,esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_245]), c_0_246])])])).
cnf(c_0_272, plain, (k1_relset_1(X1,X2)=k1_relset_1(X1,X3)|v1_xboole_0(X1)|X2!=k2_mesfunc2(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_266])).
cnf(c_0_273, plain, (v1_xboole_0(X1)|m1_subset_1(k2_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_267])).
cnf(c_0_274, plain, (k1_relset_1(X1,X2)=k1_relset_1(X1,X3)|v1_xboole_0(X1)|X2!=k1_mesfunc2(X1,X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_268])).
cnf(c_0_275, plain, (m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(split_conjunct,[status(thm)],[c_0_147])).
cnf(c_0_276, negated_conjecture, (k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_269, c_0_270]), c_0_271]), c_0_244]), c_0_245]), c_0_246]), c_0_116])]), c_0_124])])).
cnf(c_0_277, plain, (k1_relset_1(X1,k2_mesfunc2(X1,X2))=k1_relset_1(X1,X3)|v1_xboole_0(X1)|k2_mesfunc2(X1,X2)!=k2_mesfunc2(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(k2_mesfunc2(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_273])])).
cnf(c_0_278, plain, (k1_relset_1(X1,X2)=k1_relset_1(X1,X3)|v1_xboole_0(X1)|X2!=k1_mesfunc2(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_274])).
cnf(c_0_279, plain, (v1_xboole_0(X1)|m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[c_0_275])).
cnf(c_0_280, negated_conjecture, (k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_276, c_0_273]), c_0_116]), c_0_118])]), c_0_124])])).
cnf(c_0_281, plain, (k1_relset_1(X1,k2_mesfunc2(X1,X2))=k1_relset_1(X1,X3)|v1_xboole_0(X1)|k2_mesfunc2(X1,X2)!=k2_mesfunc2(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(csr,[status(thm)],[c_0_277, c_0_179])).
cnf(c_0_282, plain, (k1_relset_1(X1,k1_mesfunc2(X1,X2))=k1_relset_1(X1,X3)|v1_xboole_0(X1)|k1_mesfunc2(X1,X2)!=k1_mesfunc2(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_279])])).
cnf(c_0_283, negated_conjecture, (k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_279]), c_0_116]), c_0_118])]), c_0_124])])).
cnf(c_0_284, plain, (k1_relset_1(X1,k2_mesfunc2(X1,X2))=k1_relset_1(X1,X2)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_281])])).
cnf(c_0_285, plain, (k1_relset_1(X1,k1_mesfunc2(X1,X2))=k1_relset_1(X1,X3)|v1_xboole_0(X1)|k1_mesfunc2(X1,X2)!=k1_mesfunc2(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(csr,[status(thm)],[c_0_282, c_0_178])).
cnf(c_0_286, negated_conjecture, (k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_284]), c_0_116]), c_0_118])]), c_0_124])])).
cnf(c_0_287, plain, (k1_relset_1(X1,k1_mesfunc2(X1,X2))=k1_relset_1(X1,X2)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_285])])).
cnf(c_0_288, negated_conjecture, (~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_287]), c_0_116]), c_0_118])]), c_0_124])])).
cnf(c_0_289, negated_conjecture, (~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_254]), c_0_244]), c_0_245]), c_0_246]), c_0_116])]), c_0_124])])).
cnf(c_0_290, negated_conjecture, (~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_179]), c_0_116]), c_0_118])]), c_0_124])])).
cnf(c_0_291, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_178]), c_0_116]), c_0_118])]), c_0_124]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 292
# Proof object clause steps            : 218
# Proof object formula steps           : 74
# Proof object conjectures             : 71
# Proof object clause conjectures      : 68
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 68
# Proof object initial formulas used   : 38
# Proof object generating inferences   : 63
# Proof object simplifying inferences  : 273
# Training examples: 132 positive, 995 negative
# Training: Positive examples begin
cnf(c_0_54, plain, (k1_xboole_0=k6_numbers)).# trainpos
cnf(c_0_53, plain, (k1_xboole_0=k1_supinf_2)).# trainpos
cnf(c_0_68, plain, (v1_xboole_0(np__0))).# trainpos
cnf(c_0_117, negated_conjecture, (v1_funct_1(esk3_0))).# trainpos
cnf(c_0_116, negated_conjecture, (v1_funct_1(esk4_0))).# trainpos
cnf(c_0_124, negated_conjecture, (~v1_xboole_0(esk1_0))).# trainpos
cnf(c_0_123, negated_conjecture, (~v1_xboole_0(esk2_0))).# trainpos
cnf(c_0_114, negated_conjecture, (v10_valued_0(esk3_0))).# trainpos
cnf(c_0_113, negated_conjecture, (v6_supinf_2(esk3_0))).# trainpos
cnf(c_0_111, negated_conjecture, (r1_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0))).# trainpos
cnf(c_0_122, negated_conjecture, (v1_prob_1(esk2_0,esk1_0))).# trainpos
cnf(c_0_121, negated_conjecture, (v4_prob_1(esk2_0,esk1_0))).# trainpos
cnf(c_0_115, negated_conjecture, (v1_funct_2(esk3_0,esk2_0,k7_numbers))).# trainpos
cnf(c_0_112, negated_conjecture, (v4_measure1(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_194, negated_conjecture, (~r1_xxreal_0(k3_extreal1(k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0)),k7_mesfunc5(esk1_0,esk2_0,esk3_0,k10_mesfunc1(esk1_0,esk4_0))))).# trainpos
cnf(c_0_60, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_196, plain, (r1_tarski(X1,X1))).# trainpos
cnf(c_0_249, plain, (~v1_xboole_0(esk14_1(X1)))).# trainpos
cnf(c_0_237, plain, (v1_prob_1(esk14_1(X1),X1))).# trainpos
cnf(c_0_235, plain, (v3_measure1(esk14_1(X1),X1))).# trainpos
cnf(c_0_120, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).# trainpos
cnf(c_0_119, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).# trainpos
cnf(c_0_118, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).# trainpos
cnf(c_0_236, plain, (m1_subset_1(esk14_1(X1),k1_zfmisc_1(k1_zfmisc_1(X1))))).# trainpos
cnf(c_0_234, plain, (v4_prob_1(X1,X2)|~v1_prob_1(X1,X2)|~v3_measure1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).# trainpos
cnf(c_0_64, plain, (m1_subset_1(k3_extreal1(X1),k7_numbers)|~m1_subset_1(X1,k7_numbers))).# trainpos
cnf(c_0_199, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_131, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_66, plain, (k3_extreal1(X1)=X1|~m1_subset_1(X1,k7_numbers)|~r1_xxreal_0(k6_numbers,X1))).# trainpos
cnf(c_0_200, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_201, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).# trainpos
cnf(c_0_251, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).# trainpos
cnf(c_0_254, plain, (v1_xboole_0(X1)|v1_funct_1(k10_mesfunc1(X1,X2))|~v1_funct_1(X2)|~v1_relat_1(X2)|~v4_relat_1(X2,X1)|~v5_relat_1(X2,k7_numbers))).# trainpos
cnf(c_0_56, plain, (k3_supinf_2(X1,X2)=k3_supinf_2(X2,X1)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).# trainpos
cnf(c_0_47, plain, (k4_supinf_2(X1,X2)=k3_xxreal_3(X1,X2)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).# trainpos
cnf(c_0_178, plain, (v1_xboole_0(X1)|v1_funct_1(k1_mesfunc2(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_179, plain, (v1_xboole_0(X1)|v1_funct_1(k2_mesfunc2(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_46, plain, (r1_xxreal_0(k3_extreal1(k4_supinf_2(X1,X2)),k3_supinf_2(k3_extreal1(X1),k3_extreal1(X2)))|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).# trainpos
cnf(c_0_130, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_144, plain, (X3=X4|~r2_relset_1(X1,X2,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_279, plain, (v1_xboole_0(X1)|m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_154, plain, (v1_xboole_0(X2)|v1_funct_1(esk10_2(X1,X2))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).# trainpos
cnf(c_0_156, plain, (v1_xboole_0(X2)|v10_valued_0(esk10_2(X1,X2))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).# trainpos
cnf(c_0_240, plain, (v1_xboole_0(X1)|m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~v1_relat_1(X2)|~v4_relat_1(X2,X1)|~v5_relat_1(X2,k7_numbers))).# trainpos
cnf(c_0_157, plain, (v1_xboole_0(X2)|v6_supinf_2(esk10_2(X1,X2))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).# trainpos
cnf(c_0_273, plain, (v1_xboole_0(X1)|m1_subset_1(k2_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_155, plain, (v1_xboole_0(X2)|v1_funct_2(esk10_2(X1,X2),X2,k7_numbers)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).# trainpos
cnf(c_0_133, plain, (v1_xboole_0(X2)|v4_measure1(esk10_2(X1,X2),X1,X2)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).# trainpos
cnf(c_0_153, plain, (v1_xboole_0(X2)|m1_subset_1(esk10_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).# trainpos
cnf(c_0_145, plain, (v1_xboole_0(X1)|r2_relset_1(X1,k7_numbers,k10_mesfunc1(X1,X2),k3_mesfunc1(X1,k1_mesfunc2(X1,X2),k2_mesfunc2(X1,X2)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_278, plain, (k1_relset_1(X1,X2)=k1_relset_1(X1,X3)|v1_xboole_0(X1)|X2!=k1_mesfunc2(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_272, plain, (k1_relset_1(X1,X2)=k1_relset_1(X1,X3)|v1_xboole_0(X1)|X2!=k2_mesfunc2(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_241, plain, (k9_xtuple_0(X1)=k9_xtuple_0(X2)|v1_xboole_0(X3)|X1!=k10_mesfunc1(X3,X2)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v4_relat_1(X2,X3)|~v5_relat_1(X2,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers))))).# trainpos
cnf(c_0_164, plain, (v1_xboole_0(X1)|m1_subset_1(k3_mesfunc1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers))).# trainpos
cnf(c_0_132, plain, (v1_xboole_0(X4)|v1_xboole_0(X1)|v6_supinf_2(k10_mesfunc1(X1,X2))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v1_prob_1(X4,X1)|~v4_prob_1(X4,X1)|~v4_measure1(X3,X1,X4)|~v1_funct_2(X3,X4,k7_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_169, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|r1_mesfunc1(X1,X2,k10_mesfunc1(X1,X3),X4)|~v1_funct_1(X3)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X4,X2)|~r1_mesfunc1(X1,X2,X3,X4)|~r1_tarski(X4,k1_relset_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_217, plain, (v1_xboole_0(X3)|v1_xboole_0(X1)|r1_mesfunc1(X1,X3,k1_mesfunc2(X1,X2),X4)|~v1_funct_1(X2)|~v1_prob_1(X3,X1)|~v4_prob_1(X3,X1)|~m1_subset_1(X4,X3)|~r1_mesfunc1(X1,X3,X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_213, plain, (v1_xboole_0(X3)|v1_xboole_0(X1)|r1_mesfunc1(X1,X3,k2_mesfunc2(X1,X2),X4)|~v1_funct_1(X2)|~v1_prob_1(X3,X1)|~v4_prob_1(X3,X1)|~m1_subset_1(X4,X3)|~r1_mesfunc1(X1,X3,X2,X4)|~r1_tarski(X4,k1_relset_1(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_161, plain, (v1_xboole_0(X4)|v1_xboole_0(X1)|v6_supinf_2(k1_mesfunc2(X1,X2))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v1_prob_1(X4,X1)|~v4_prob_1(X4,X1)|~v4_measure1(X3,X1,X4)|~v1_funct_2(X3,X4,k7_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_162, plain, (v1_xboole_0(X4)|v1_xboole_0(X1)|v6_supinf_2(k2_mesfunc2(X1,X2))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v1_prob_1(X4,X1)|~v4_prob_1(X4,X1)|~v4_measure1(X3,X1,X4)|~v1_funct_2(X3,X4,k7_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_224, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(k6_mesfunc5(X1,X2,X3,X4),k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~v4_measure1(X3,X1,X2)|~v1_funct_2(X3,X2,k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).# trainpos
cnf(c_0_86, plain, (esk5_4(X1,X2,X3,X4)=k1_relset_1(X1,X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~v4_measure1(X3,X1,X2)|~v1_funct_2(X3,X2,k7_numbers)|~r1_mesfunc5(X1,X2,X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).# trainpos
cnf(c_0_85, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m2_subset_1(esk5_4(X1,X2,X3,X4),k1_zfmisc_1(X1),X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~v4_measure1(X3,X1,X2)|~v1_funct_2(X3,X2,k7_numbers)|~r1_mesfunc5(X1,X2,X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).# trainpos
cnf(c_0_173, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|r1_mesfunc1(X1,X2,X3,esk5_4(X1,X2,X4,X3))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v10_valued_0(X4)|~v6_supinf_2(X4)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~v4_measure1(X4,X1,X2)|~v1_funct_2(X4,X2,k7_numbers)|~r1_mesfunc5(X1,X2,X4,X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_193, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|r1_xxreal_0(k6_numbers,k6_mesfunc5(X1,X2,X3,X4))|X5!=k1_relset_1(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v10_valued_0(X3)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~v4_measure1(X3,X1,X2)|~v1_funct_2(X3,X2,k7_numbers)|~r1_mesfunc1(X1,X2,X4,X5)|~m2_subset_1(X5,k1_zfmisc_1(X1),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).# trainpos
cnf(c_0_65, plain, (k4_supinf_2(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)))=k7_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~v4_measure1(X3,X1,X2)|~v1_funct_2(X3,X2,k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).# trainpos
cnf(c_0_149, plain, (k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|X5!=k1_relset_1(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v10_valued_0(X3)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~v4_measure1(X3,X1,X2)|~v1_funct_2(X3,X2,k7_numbers)|~r1_mesfunc1(X1,X2,X4,X5)|~m2_subset_1(X5,k1_zfmisc_1(X1),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).# trainpos
cnf(c_0_140, plain, (k3_supinf_2(k6_mesfunc5(X1,X2,X3,X4),k6_mesfunc5(X1,X2,X3,X5))=k6_mesfunc5(X1,X2,X3,k3_mesfunc1(X1,X4,X5))|v1_xboole_0(X2)|v1_xboole_0(X1)|X6!=k1_relset_1(X1,X5)|X6!=k1_relset_1(X1,X4)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v10_valued_0(X3)|~v6_supinf_2(X5)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~v4_measure1(X3,X1,X2)|~v1_funct_2(X3,X2,k7_numbers)|~r1_mesfunc1(X1,X2,X5,X6)|~r1_mesfunc1(X1,X2,X4,X6)|~m2_subset_1(X6,k1_zfmisc_1(X1),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).# trainpos
cnf(c_0_246, negated_conjecture, (v1_relat_1(esk4_0))).# trainpos
cnf(c_0_163, plain, (k3_mesfunc1(X1,k1_mesfunc2(X1,X2),k2_mesfunc2(X1,X2))=k10_mesfunc1(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(k3_mesfunc1(X1,k1_mesfunc2(X1,X2),k2_mesfunc2(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_282, plain, (k1_relset_1(X1,k1_mesfunc2(X1,X2))=k1_relset_1(X1,X3)|v1_xboole_0(X1)|k1_mesfunc2(X1,X2)!=k1_mesfunc2(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_277, plain, (k1_relset_1(X1,k2_mesfunc2(X1,X2))=k1_relset_1(X1,X3)|v1_xboole_0(X1)|k2_mesfunc2(X1,X2)!=k2_mesfunc2(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(k2_mesfunc2(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_55, plain, (r1_xxreal_0(k3_extreal1(k3_xxreal_3(X1,X2)),k3_supinf_2(k3_extreal1(X1),k3_extreal1(X2)))|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).# trainpos
cnf(c_0_74, plain, (k1_supinf_2=np__0)).# trainpos
cnf(c_0_245, negated_conjecture, (v4_relat_1(esk4_0,esk1_0))).# trainpos
cnf(c_0_187, plain, (r1_mesfunc1(X1,X2,X3,k1_relset_1(X1,X3))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_mesfunc5(X1,X2,X4,X3)|~v4_measure1(X4,X1,X2)|~v6_supinf_2(X4)|~v10_valued_0(X4)|~v1_funct_2(X4,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).# trainpos
cnf(c_0_177, plain, (k3_mesfunc1(X1,k1_mesfunc2(X1,X2),k2_mesfunc2(X1,X2))=k10_mesfunc1(X1,X2)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v5_relat_1(k2_mesfunc2(X1,X2),k7_numbers)|~v5_relat_1(k1_mesfunc2(X1,X2),k7_numbers)|~v4_relat_1(k2_mesfunc2(X1,X2),X1)|~v4_relat_1(k1_mesfunc2(X1,X2),X1)|~v1_relat_1(k2_mesfunc2(X1,X2))|~v1_relat_1(k1_mesfunc2(X1,X2))|~v1_funct_1(k2_mesfunc2(X1,X2))|~v1_funct_1(k1_mesfunc2(X1,X2)))).# trainpos
cnf(c_0_110, plain, (m2_subset_1(k1_relset_1(X1,X4),k1_zfmisc_1(X1),X2)|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_mesfunc5(X1,X2,X3,X4)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).# trainpos
cnf(c_0_284, plain, (k1_relset_1(X1,k2_mesfunc2(X1,X2))=k1_relset_1(X1,X2)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_287, plain, (k1_relset_1(X1,k1_mesfunc2(X1,X2))=k1_relset_1(X1,X2)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_72, plain, (k3_xxreal_3(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)))=k7_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).# trainpos
cnf(c_0_152, plain, (v6_supinf_2(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|v1_xboole_0(X3)|~v6_supinf_2(esk10_2(X1,X3))|~v10_valued_0(esk10_2(X1,X3))|~v1_funct_2(esk10_2(X1,X3),X3,k7_numbers)|~v1_funct_1(esk10_2(X1,X3))|~v1_funct_1(X2)|~m1_subset_1(esk10_2(X1,X3),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X3,X1)|~v1_prob_1(X3,X1))).# trainpos
cnf(c_0_176, negated_conjecture, (v6_supinf_2(k2_mesfunc2(esk1_0,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).# trainpos
cnf(c_0_175, negated_conjecture, (v6_supinf_2(k1_mesfunc2(esk1_0,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).# trainpos
cnf(c_0_233, plain, (v6_supinf_2(k2_mesfunc2(X1,X2))|v1_xboole_0(X1)|v1_xboole_0(X3)|~v6_supinf_2(esk10_2(X1,X3))|~v10_valued_0(esk10_2(X1,X3))|~v1_funct_2(esk10_2(X1,X3),X3,k7_numbers)|~v1_funct_1(esk10_2(X1,X3))|~v1_funct_1(X2)|~m1_subset_1(esk10_2(X1,X3),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X3,X1)|~v1_prob_1(X3,X1))).# trainpos
cnf(c_0_250, plain, (v6_supinf_2(k1_mesfunc2(X1,X2))|v1_xboole_0(X1)|v1_xboole_0(X3)|~v6_supinf_2(esk10_2(X1,X3))|~v10_valued_0(esk10_2(X1,X3))|~v1_funct_2(esk10_2(X1,X3),X3,k7_numbers)|~v1_funct_1(esk10_2(X1,X3))|~v1_funct_1(X2)|~m1_subset_1(esk10_2(X1,X3),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X3,X1)|~v1_prob_1(X3,X1))).# trainpos
cnf(c_0_63, plain, (r1_xxreal_0(k3_extreal1(k3_xxreal_3(X1,X2)),k3_supinf_2(k3_extreal1(X2),k3_extreal1(X1)))|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)|~m1_subset_1(k3_extreal1(X2),k7_numbers)|~m1_subset_1(k3_extreal1(X1),k7_numbers))).# trainpos
cnf(c_0_244, negated_conjecture, (v5_relat_1(esk4_0,k7_numbers))).# trainpos
cnf(c_0_197, negated_conjecture, (r1_mesfunc1(esk1_0,esk2_0,esk4_0,k1_relset_1(esk1_0,esk4_0)))).# trainpos
cnf(c_0_82, plain, (r1_xxreal_0(k3_extreal1(k7_mesfunc5(X1,X2,X3,X4)),k3_supinf_2(k3_extreal1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))),k3_extreal1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)))))|v1_xboole_0(X2)|v1_xboole_0(X1)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).# trainpos
cnf(c_0_83, plain, (k3_extreal1(X1)=X1|~r1_xxreal_0(np__0,X1)|~m1_subset_1(X1,k7_numbers))).# trainpos
cnf(c_0_205, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|X5!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,X5)|~m2_subset_1(X5,k1_zfmisc_1(X1),X2)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).# trainpos
cnf(c_0_271, negated_conjecture, (k9_xtuple_0(esk4_0)=k1_relset_1(esk1_0,esk4_0))).# trainpos
cnf(c_0_248, plain, (v4_prob_1(esk14_1(X1),X1))).# trainpos
cnf(c_0_141, negated_conjecture, (m2_subset_1(k1_relset_1(esk1_0,esk4_0),k1_zfmisc_1(esk1_0),esk2_0))).# trainpos
cnf(c_0_185, negated_conjecture, (v6_supinf_2(k10_mesfunc1(esk1_0,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).# trainpos
cnf(c_0_263, plain, (v6_supinf_2(k1_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_256, plain, (v6_supinf_2(k2_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).# trainpos
cnf(c_0_170, negated_conjecture, (m1_subset_1(k1_relset_1(esk1_0,esk4_0),esk2_0))).# trainpos
cnf(c_0_168, negated_conjecture, (k7_mesfunc5(esk1_0,esk2_0,X1,X2)=k6_mesfunc5(esk1_0,esk2_0,X1,X2)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).# trainpos
cnf(c_0_108, plain, (r1_xxreal_0(k3_extreal1(k7_mesfunc5(X1,X2,X3,X4)),k3_supinf_2(k3_extreal1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))),k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))))).# trainpos
cnf(c_0_207, negated_conjecture, (r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X2))|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).# trainpos
cnf(c_0_139, plain, (r1_xxreal_0(k3_extreal1(k7_mesfunc5(X1,X2,X3,X4)),k3_supinf_2(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))))|v1_xboole_0(X2)|v1_xboole_0(X1)|~r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))))).# trainpos
cnf(c_0_159, plain, (r1_xxreal_0(k3_extreal1(k7_mesfunc5(X1,X2,X3,X4)),k3_supinf_2(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)))|~r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).# trainpos
cnf(c_0_160, negated_conjecture, (k3_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,X3))=k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X2,X3))|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X3)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X3,k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X3)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).# trainpos
cnf(c_0_174, negated_conjecture, (r1_xxreal_0(k3_extreal1(k7_mesfunc5(esk1_0,esk2_0,X1,X2)),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,k1_mesfunc2(esk1_0,X2),k2_mesfunc2(esk1_0,X2))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~v6_supinf_2(k2_mesfunc2(esk1_0,X2))|~v6_supinf_2(k1_mesfunc2(esk1_0,X2))|~v1_funct_1(k2_mesfunc2(esk1_0,X2))|~v1_funct_1(k1_mesfunc2(esk1_0,X2))|~m1_subset_1(k2_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).# trainpos
cnf(c_0_198, negated_conjecture, (r1_xxreal_0(k3_extreal1(k7_mesfunc5(esk1_0,esk2_0,X1,X2)),k6_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,X2)))|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k2_mesfunc2(esk1_0,X2))|~v1_funct_1(k1_mesfunc2(esk1_0,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|~v5_relat_1(k2_mesfunc2(esk1_0,X2),k7_numbers)|~v5_relat_1(k1_mesfunc2(esk1_0,X2),k7_numbers)|~v4_relat_1(k2_mesfunc2(esk1_0,X2),esk1_0)|~v4_relat_1(k1_mesfunc2(esk1_0,X2),esk1_0)|~v1_relat_1(k2_mesfunc2(esk1_0,X2))|~v1_relat_1(k1_mesfunc2(esk1_0,X2))|~m1_subset_1(k10_mesfunc1(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).# trainpos
cnf(c_0_259, plain, (v1_relat_1(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).# trainpos
cnf(c_0_184, negated_conjecture, (k7_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,X2))=k6_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,X2))|k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(k10_mesfunc1(esk1_0,X2))|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k10_mesfunc1(esk1_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k10_mesfunc1(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|~r1_tarski(k1_relset_1(esk1_0,esk4_0),k1_relset_1(esk1_0,X2))|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).# trainpos
cnf(c_0_203, negated_conjecture, (~r1_xxreal_0(k3_extreal1(k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0)),k6_mesfunc5(esk1_0,esk2_0,esk3_0,k10_mesfunc1(esk1_0,esk4_0)))|k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).# trainpos
cnf(c_0_206, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).# trainpos
cnf(c_0_209, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)))|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~v6_supinf_2(k2_mesfunc2(esk1_0,esk4_0)))).# trainpos
cnf(c_0_252, plain, (v4_relat_1(k10_mesfunc1(X1,X2),X1)|v1_xboole_0(X1)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).# trainpos
cnf(c_0_212, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~v6_supinf_2(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~v6_supinf_2(k1_mesfunc2(esk1_0,esk4_0)))).# trainpos
cnf(c_0_216, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~v6_supinf_2(k2_mesfunc2(esk1_0,esk4_0))|~v6_supinf_2(k1_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).# trainpos
cnf(c_0_258, plain, (k1_relset_1(X1,k10_mesfunc1(X1,X2))=k9_xtuple_0(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v1_relat_1(k10_mesfunc1(X1,X2))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).# trainpos
cnf(c_0_223, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v6_supinf_2(k2_mesfunc2(esk1_0,esk4_0))|~v6_supinf_2(k1_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).# trainpos
cnf(c_0_232, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v6_supinf_2(k2_mesfunc2(esk1_0,esk4_0))|~v6_supinf_2(k1_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).# trainpos
cnf(c_0_243, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v6_supinf_2(k2_mesfunc2(esk1_0,esk4_0))|~v6_supinf_2(k1_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).# trainpos
cnf(c_0_255, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v6_supinf_2(k2_mesfunc2(esk1_0,esk4_0))|~v6_supinf_2(k1_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).# trainpos
cnf(c_0_262, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v6_supinf_2(k1_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).# trainpos
cnf(c_0_269, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).# trainpos
cnf(c_0_253, plain, (k9_xtuple_0(k10_mesfunc1(X1,X2))=k9_xtuple_0(X3)|v1_xboole_0(X1)|k10_mesfunc1(X1,X2)!=k10_mesfunc1(X1,X3)|~v5_relat_1(X3,k7_numbers)|~v4_relat_1(X3,X1)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v1_funct_1(k10_mesfunc1(X1,X2))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).# trainpos
cnf(c_0_265, plain, (k9_xtuple_0(k10_mesfunc1(X1,X2))=k9_xtuple_0(X2)|v1_xboole_0(X1)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).# trainpos
cnf(c_0_270, plain, (k1_relset_1(X1,k10_mesfunc1(X1,X2))=k9_xtuple_0(X2)|v1_xboole_0(X1)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).# trainpos
cnf(c_0_276, negated_conjecture, (k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).# trainpos
cnf(c_0_280, negated_conjecture, (k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).# trainpos
cnf(c_0_283, negated_conjecture, (k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0)))).# trainpos
cnf(c_0_286, negated_conjecture, (k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0)))).# trainpos
cnf(c_0_288, negated_conjecture, (~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0)))).# trainpos
cnf(c_0_289, negated_conjecture, (~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0)))).# trainpos
cnf(c_0_290, negated_conjecture, (~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0)))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_61, plain, (k1_xxreal_0=k1_supinf_1)).#trainneg
cnf(i_0_106, plain, (k5_numbers=k4_ordinal1)).#trainneg
cnf(i_0_132, plain, (o_0_0_xboole_0=k1_xboole_0)).#trainneg
cnf(i_0_218, plain, (k4_xcmplx_0(np__0)=np__0)).#trainneg
cnf(i_0_97, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_180, plain, (v1_xboole_0(esk13_0))).#trainneg
cnf(i_0_121, plain, (~v1_xboole_0(k7_numbers))).#trainneg
cnf(i_0_204, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_129, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_98, plain, (v2_membered(k7_numbers))).#trainneg
cnf(i_0_134, plain, (v1_xxreal_0(k1_xxreal_0))).#trainneg
cnf(i_0_178, plain, (v1_xxreal_0(esk13_0))).#trainneg
cnf(i_0_93, plain, (v3_membered(k1_numbers))).#trainneg
cnf(i_0_179, plain, (v1_xcmplx_0(esk13_0))).#trainneg
cnf(i_0_177, plain, (v1_xreal_0(esk13_0))).#trainneg
cnf(i_0_128, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_164, plain, (v2_xxreal_0(k1_xxreal_0))).#trainneg
cnf(i_0_161, plain, (v6_membered(k4_ordinal1))).#trainneg
cnf(i_0_174, plain, (~v1_xreal_0(k1_xxreal_0))).#trainneg
cnf(i_0_159, plain, (v1_xboole_0(esk12_1(X1)))).#trainneg
cnf(i_0_84, plain, (m1_subset_1(k1_supinf_1,k7_numbers))).#trainneg
cnf(i_0_75, plain, (m1_subset_1(k1_supinf_2,k7_numbers))).#trainneg
cnf(i_0_206, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_103, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_105, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_104, plain, (m1_subset_1(np__0,k5_numbers))).#trainneg
cnf(i_0_158, plain, (r1_xxreal_0(np__0,np__0))).#trainneg
cnf(i_0_211, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_220, plain, (v2_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_127, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_167, plain, (v1_xcmplx_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_94, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_138, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_120, plain, (m1_subset_1(esk11_1(X1),X1))).#trainneg
cnf(i_0_221, plain, (~v1_subset_1(esk15_1(X1),X1))).#trainneg
cnf(i_0_168, plain, (v1_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_160, plain, (m1_subset_1(esk12_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_222, plain, (m1_subset_1(esk15_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_173, plain, (v1_xxreal_0(X1)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_156, plain, (v1_xboole_0(X1)|~v1_xcmplx_0(X1)|~v1_xboole_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_228, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_144, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_86, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_157, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_181, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_170, plain, (v2_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k7_numbers))).#trainneg
cnf(i_0_149, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_123, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_122, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_107, plain, (k2_xxreal_3(k2_xxreal_3(X1))=X1|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_100, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_143, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_124, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_152, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_85, plain, (k2_xxreal_3(X1)=k2_supinf_2(X1)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_81, plain, (k2_supinf_2(k2_supinf_2(X1))=X1|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_87, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_102, plain, (k3_extreal1(k3_extreal1(X1))=k3_extreal1(X1)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_48, plain, (r1_xxreal_0(k6_numbers,k3_extreal1(X1))|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_76, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_114, plain, (v1_xboole_0(X2)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~v1_xboole_0(esk10_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_165, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_140, plain, (m1_subset_1(k2_supinf_2(X1),k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_147, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_80, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_77, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_216, plain, (v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v3_xxreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_175, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_205, plain, (k4_xxreal_0(X1,X1)=X1|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_150, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_142, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_146, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_176, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_145, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_230, plain, (v1_xxreal_0(k3_xxreal_3(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_225, plain, (v1_xcmplx_0(k6_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_33, plain, (v6_supinf_2(X1)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~r1_xxreal_0(k1_supinf_2,k12_supinf_2(X1,esk7_2(X2,X1))))).#trainneg
cnf(i_0_131, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_28, plain, (k2_supinf_2(X1)=k3_extreal1(X1)|r1_xxreal_0(k6_numbers,X1)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_22, plain, (X1=k10_mesfunc1(X2,X3)|v1_xboole_0(X2)|k9_xtuple_0(X1)!=k9_xtuple_0(X3)|k12_supinf_2(X1,esk6_3(X2,X3,X1))!=k3_extreal1(k12_supinf_2(X3,esk6_3(X2,X3,X1)))|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v4_relat_1(X3,X2)|~v5_relat_1(X3,k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).#trainneg
cnf(i_0_207, plain, (k4_xcmplx_0(X2)=k2_xxreal_3(X1)|X1!=X2|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_226, plain, (v1_valued_0(X2)|~v1_relat_1(X1)|~v1_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_183, plain, (v3_valued_0(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_208, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_68, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_71, plain, (k4_xxreal_0(X1,X2)=k4_xxreal_0(X2,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_83, plain, (k1_xxreal_3(X1,X2)=k1_xxreal_3(X2,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_229, plain, (v1_xreal_0(k3_xxreal_3(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_141, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_151, plain, (k1_extreal2(X1,X1)=X1|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_154, plain, (k1_xxreal_3(X1,k2_xxreal_3(X2))=k3_xxreal_3(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_130, plain, (v2_valued_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_162, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_219, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_163, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_valued_0(X1))).#trainneg
cnf(i_0_203, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_186, plain, (v2_finsub_1(X1)|~v1_finsub_1(X1)|~v1_prob_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_209, plain, (v1_xboole_0(X1)|v1_finsub_1(X1)|~v1_prob_1(X1,X2)|~v3_measure1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_169, plain, (m1_subset_1(k3_supinf_2(X1,X2),k7_numbers)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_37, plain, (X1=k1_mesfunc2(X2,X3)|v1_xboole_0(X2)|k1_relset_1(X2,X1)!=k1_relset_1(X2,X3)|k12_supinf_2(X1,esk8_3(X2,X3,X1))!=k1_extreal2(k12_supinf_2(X3,esk8_3(X2,X3,X1)),k1_supinf_2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).#trainneg
cnf(i_0_66, plain, (k1_funct_1(X1,X2)=k12_supinf_2(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_valued_0(X1))).#trainneg
cnf(i_0_74, plain, (k4_xxreal_0(X1,X2)=k1_extreal2(X1,X2)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_224, plain, (v3_finsub_1(X1)|~v2_finsub_1(X1)|~v1_prob_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_185, plain, (m1_subset_1(k1_extreal2(X1,X2),k7_numbers)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_137, plain, (v3_measure1(X1,X2)|~v1_prob_1(X1,X2)|~v4_prob_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_227, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_101, plain, (m1_subset_1(k12_supinf_2(X1,X2),k7_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_valued_0(X1))).#trainneg
cnf(i_0_82, plain, (k1_extreal2(X1,X2)=k1_extreal2(X2,X1)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_78, plain, (k1_xxreal_3(X1,X2)=k3_supinf_2(X1,X2)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_99, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_184, plain, (m1_subset_1(k4_supinf_2(X1,X2),k7_numbers)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_95, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_133, plain, (r2_relset_1(X2,X3,X1,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_125, plain, (r2_relset_1(X2,X3,X4,X1)|~r2_relset_1(X2,X3,X1,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_42, plain, (X1=k2_mesfunc2(X2,X3)|v1_xboole_0(X2)|k1_relset_1(X2,X1)!=k1_relset_1(X2,X3)|k12_supinf_2(X1,esk9_3(X2,X3,X1))!=k1_extreal2(k2_supinf_2(k12_supinf_2(X3,esk9_3(X2,X3,X1))),k1_supinf_2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).#trainneg
cnf(i_0_223, plain, (k2_xcmplx_0(X3,X4)=k1_xxreal_3(X1,X2)|X2!=X4|X1!=X3|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_116, plain, (v1_xboole_0(X2)|v5_relat_1(esk10_2(X1,X2),k7_numbers)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_52, plain, (r2_relset_1(X3,X4,X1,X2)|X1!=X2|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_172, plain, (k6_xcmplx_0(X3,X4)=k3_xxreal_3(X1,X2)|X2!=X4|X1!=X3|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_34, plain, (v6_supinf_2(X2)|r2_hidden(esk7_2(X1,X2),k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_118, plain, (v1_xboole_0(X2)|v1_relat_1(esk10_2(X1,X2))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_111, plain, (v1_xboole_0(X2)|v2_valued_0(esk10_2(X1,X2))|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_117, plain, (v1_xboole_0(X2)|v4_relat_1(esk10_2(X1,X2),X2)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_113, plain, (v1_xboole_0(X2)|v1_partfun1(esk10_2(X1,X2),X2)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_25, plain, (k12_supinf_2(X3,X1)=k3_extreal1(k12_supinf_2(X4,X1))|v1_xboole_0(X2)|X3!=k10_mesfunc1(X2,X4)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_relat_1(X4)|~m1_subset_1(X1,X2)|~v4_relat_1(X4,X2)|~v5_relat_1(X4,k7_numbers)|~r2_hidden(X1,k9_xtuple_0(X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).#trainneg
cnf(i_0_63, plain, (v1_xboole_0(X1)|v1_funct_1(k3_mesfunc1(X1,X2,X3))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers))).#trainneg
cnf(i_0_18, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~v4_measure1(X3,X1,X2)|~v1_funct_2(X3,X2,k7_numbers)|~r1_mesfunc5(X1,X2,X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~r1_xxreal_0(k1_supinf_1,k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))))).#trainneg
cnf(i_0_58, plain, (k3_mesfunc1(X1,X2,X3)=k3_mesfunc1(X1,X3,X2)|v1_xboole_0(X1)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers))).#trainneg
cnf(i_0_40, plain, (k12_supinf_2(X3,X1)=k1_extreal2(k12_supinf_2(X4,X1),k1_supinf_2)|v1_xboole_0(X2)|X3!=k1_mesfunc2(X2,X4)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).#trainneg
cnf(i_0_45, plain, (k12_supinf_2(X3,X1)=k1_extreal2(k2_supinf_2(k12_supinf_2(X4,X1)),k1_supinf_2)|v1_xboole_0(X2)|X3!=k2_mesfunc2(X2,X4)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).#trainneg
cnf(i_0_24, plain, (X3=k10_mesfunc1(X1,X2)|v1_xboole_0(X1)|m1_subset_1(esk6_3(X1,X2,X3),X1)|k9_xtuple_0(X2)!=k9_xtuple_0(X3)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v4_relat_1(X2,X1)|~v5_relat_1(X2,k7_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_23, plain, (X3=k10_mesfunc1(X1,X2)|v1_xboole_0(X1)|r2_hidden(esk6_3(X1,X2,X3),k9_xtuple_0(X3))|k9_xtuple_0(X2)!=k9_xtuple_0(X3)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v4_relat_1(X2,X1)|~v5_relat_1(X2,k7_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_39, plain, (X3=k1_mesfunc2(X1,X2)|v1_xboole_0(X1)|m1_subset_1(esk8_3(X1,X2,X3),X1)|k1_relset_1(X1,X2)!=k1_relset_1(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_44, plain, (X3=k2_mesfunc2(X1,X2)|v1_xboole_0(X1)|m1_subset_1(esk9_3(X1,X2,X3),X1)|k1_relset_1(X1,X2)!=k1_relset_1(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_38, plain, (X3=k1_mesfunc2(X1,X2)|v1_xboole_0(X1)|r2_hidden(esk8_3(X1,X2,X3),k1_relset_1(X1,X3))|k1_relset_1(X1,X2)!=k1_relset_1(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_43, plain, (X3=k2_mesfunc2(X1,X2)|v1_xboole_0(X1)|r2_hidden(esk9_3(X1,X2,X3),k1_relset_1(X1,X3))|k1_relset_1(X1,X2)!=k1_relset_1(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_17, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v10_valued_0(X3)|~v6_supinf_2(X3)|~v1_prob_1(X2,X1)|~v4_prob_1(X2,X1)|~v4_measure1(X3,X1,X2)|~v1_funct_2(X3,X2,k7_numbers)|~r1_mesfunc5(X1,X2,X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~r1_xxreal_0(k1_supinf_1,k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))))).#trainneg
cnf(i_0_16, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|r1_mesfunc5(X2,X3,X5,X4)|r1_xxreal_0(k1_supinf_1,k6_mesfunc5(X2,X3,X5,k1_mesfunc2(X2,X4)))|r1_xxreal_0(k1_supinf_1,k6_mesfunc5(X2,X3,X5,k2_mesfunc2(X2,X4)))|X1!=k1_relset_1(X2,X4)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v10_valued_0(X5)|~v6_supinf_2(X5)|~v1_prob_1(X3,X2)|~v4_prob_1(X3,X2)|~v4_measure1(X5,X2,X3)|~v1_funct_2(X5,X3,k7_numbers)|~r1_mesfunc1(X2,X3,X4,X1)|~m2_subset_1(X1,k1_zfmisc_1(X2),X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).#trainneg
cnf(i_0_326, plain, (~v3_membered(k7_numbers))).#trainneg
cnf(i_0_347, negated_conjecture, (v1_relat_1(esk3_0))).#trainneg
cnf(i_0_67, plain, (k1_xboole_0=k1_supinf_2)).#trainneg
cnf(i_0_231, plain, (esk13_0=k1_supinf_2)).#trainneg
cnf(i_0_276, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_463, plain, (v1_finsub_1(X1)|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v1_prob_1(X1,X2)|~v4_prob_1(X1,X2))).#trainneg
cnf(i_0_484, plain, (r1_xxreal_0(k3_extreal1(k4_supinf_2(X1,X2)),k3_supinf_2(k3_extreal1(X1),X2))|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)|~r1_xxreal_0(k1_supinf_2,X2))).#trainneg
cnf(i_0_599, plain, (k1_relset_1(X1,k1_mesfunc2(X1,X2))=k1_relset_1(X1,X3)|v1_xboole_0(X1)|k1_mesfunc2(X1,X2)!=k2_mesfunc2(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_379, negated_conjecture, (~r1_xxreal_0(k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0),k7_mesfunc5(esk1_0,esk2_0,esk3_0,k10_mesfunc1(esk1_0,esk4_0)))|~r1_xxreal_0(k1_supinf_2,k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_subset_1(k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0),k7_numbers))).#trainneg
cnf(i_0_589, plain, (k1_relset_1(X1,k2_mesfunc2(X1,X2))=k1_relset_1(X1,X3)|v1_xboole_0(X1)|k2_mesfunc2(X1,X2)!=k1_mesfunc2(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(k2_mesfunc2(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_264, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_207, plain, (k2_xxreal_3(X1)=k4_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_237, plain, (X1=esk12_1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_256, plain, (~m1_subset_1(k1_supinf_1,k1_numbers))).#trainneg
cnf(i_0_429, negated_conjecture, (v2_valued_0(esk4_0))).#trainneg
cnf(i_0_430, negated_conjecture, (v2_valued_0(esk3_0))).#trainneg
cnf(i_0_238, plain, (X1=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_744, plain, (k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X1)|v1_xboole_0(X2)|esk5_4(X1,X2,X5,X6)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,esk5_4(X1,X2,X5,X6))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_mesfunc5(X1,X2,X5,X6)|~v4_measure1(X5,X1,X2)|~v6_supinf_2(X5)|~v10_valued_0(X5)|~v1_funct_2(X5,X2,k7_numbers)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).#trainneg
cnf(i_0_736, plain, (m1_subset_1(esk5_4(X1,X2,X3,X4),k1_zfmisc_1(X1))|v1_xboole_0(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~r1_mesfunc5(X1,X2,X3,X4)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_745, plain, (k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X1)|v1_xboole_0(X2)|X5!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,X5)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~m1_subset_1(X5,X2))).#trainneg
cnf(i_0_748, plain, (k3_supinf_2(k6_mesfunc5(X1,X2,X3,X4),k6_mesfunc5(X1,X2,X3,X5))=k6_mesfunc5(X1,X2,X3,k3_mesfunc1(X1,X4,X5))|v1_xboole_0(X1)|v1_xboole_0(X2)|esk5_4(X1,X2,X6,X7)!=k1_relset_1(X1,X5)|esk5_4(X1,X2,X6,X7)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X5,esk5_4(X1,X2,X6,X7))|~r1_mesfunc1(X1,X2,X4,esk5_4(X1,X2,X6,X7))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X5)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_mesfunc5(X1,X2,X6,X7)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X6)|~v10_valued_0(X6)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).#trainneg
cnf(i_0_735, plain, (m1_subset_1(esk5_4(X1,X2,X3,X4),X2)|v1_xboole_0(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~r1_mesfunc5(X1,X2,X3,X4)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_297, plain, (v1_xxreal_0(k3_extreal1(X1))|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_366, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk15_1(X2)))).#trainneg
cnf(i_0_368, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk12_1(X2)))).#trainneg
cnf(i_0_322, plain, (v1_xreal_0(np__0))).#trainneg
cnf(i_0_223, plain, (k2_xcmplx_0(X1,X2)=k1_xxreal_3(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_409, negated_conjecture, (v4_relat_1(esk3_0,esk2_0))).#trainneg
cnf(i_0_336, plain, (v7_ordinal1(np__0))).#trainneg
cnf(i_0_363, plain, (v3_membered(k4_ordinal1))).#trainneg
cnf(i_0_253, plain, (v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_254, plain, (v1_xboole_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xboole_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_743, plain, (r1_xxreal_0(k3_extreal1(k7_mesfunc5(X1,X2,X3,X4)),k3_supinf_2(k3_extreal1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))),k3_extreal1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)))))|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_742, plain, (m1_subset_1(k7_mesfunc5(X1,X2,X3,X4),k7_numbers)|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_414, negated_conjecture, (v5_relat_1(esk3_0,k7_numbers))).#trainneg
cnf(i_0_487, plain, (r1_xxreal_0(k3_extreal1(k4_supinf_2(X1,X2)),k3_supinf_2(k3_extreal1(X2),k3_extreal1(X1)))|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)|~m1_subset_1(k3_extreal1(X2),k7_numbers)|~m1_subset_1(k3_extreal1(X1),k7_numbers))).#trainneg
cnf(i_0_367, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(X1,esk14_1(X2)))).#trainneg
cnf(i_0_369, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk11_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_67, plain, (k1_xboole_0=np__0)).#trainneg
cnf(i_0_379, negated_conjecture, (~r1_xxreal_0(k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0),k7_mesfunc5(esk1_0,esk2_0,esk3_0,k10_mesfunc1(esk1_0,esk4_0)))|~r1_xxreal_0(np__0,k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_subset_1(k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0),k7_numbers))).#trainneg
cnf(i_0_47, plain, (k6_numbers=np__0)).#trainneg
cnf(i_0_179, plain, (v1_xcmplx_0(k1_supinf_2))).#trainneg
cnf(i_0_457, negated_conjecture, (v3_finsub_1(esk2_0)|~v2_finsub_1(esk2_0))).#trainneg
cnf(i_0_434, negated_conjecture, (v1_valued_0(esk4_0)|~v1_membered(k7_numbers))).#trainneg
cnf(i_0_258, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_260, plain, (k2_supinf_2(k2_xxreal_3(X1))=X1|~v1_xxreal_0(X1)|~m1_subset_1(k2_xxreal_3(X1),k7_numbers))).#trainneg
cnf(i_0_801, plain, (k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|k1_relset_1(X1,X6)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X6))|~r1_mesfunc5(X1,X2,X5,X6)|~v4_measure1(X3,X1,X2)|~v4_measure1(X5,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X5)|~v10_valued_0(X3)|~v10_valued_0(X5)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X5,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_48, plain, (r1_xxreal_0(np__0,k3_extreal1(X1))|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_802, plain, (k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|esk5_4(X1,X2,X5,X4)!=k1_relset_1(X1,X4)|~r1_mesfunc5(X1,X2,X5,X4)|~v4_measure1(X3,X1,X2)|~v4_measure1(X5,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X5)|~v10_valued_0(X3)|~v10_valued_0(X5)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X5,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X5)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_846, plain, (k3_supinf_2(k6_mesfunc5(X1,X2,X3,X4),k6_mesfunc5(X1,X2,X3,X5))=k6_mesfunc5(X1,X2,X3,k3_mesfunc1(X1,X4,X5))|v1_xboole_0(X2)|v1_xboole_0(X1)|esk5_4(X1,X2,X6,X5)!=k1_relset_1(X1,X5)|esk5_4(X1,X2,X6,X5)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,esk5_4(X1,X2,X6,X5))|~r1_mesfunc5(X1,X2,X6,X5)|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X5)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_881, negated_conjecture, (k9_xtuple_0(esk3_0)=k1_relset_1(esk2_0,esk3_0))).#trainneg
cnf(i_0_33, plain, (v6_supinf_2(X1)|~r1_xxreal_0(np__0,k12_supinf_2(X1,esk7_2(X2,X1)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).#trainneg
cnf(i_0_178, plain, (v1_xxreal_0(k1_supinf_2))).#trainneg
cnf(i_0_132, plain, (o_0_0_xboole_0=np__0)).#trainneg
cnf(i_0_261, plain, (k2_xxreal_3(k2_supinf_2(X1))=X1|~v1_xxreal_0(X1)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_231, plain, (esk13_0=np__0)).#trainneg
cnf(i_0_750, negated_conjecture, (v1_finsub_1(esk2_0))).#trainneg
cnf(i_0_435, negated_conjecture, (v1_valued_0(esk3_0)|~v1_membered(k7_numbers))).#trainneg
cnf(i_0_445, negated_conjecture, (v2_finsub_1(esk2_0)|~v1_finsub_1(esk2_0))).#trainneg
cnf(i_0_268, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|~v2_xxreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_484, plain, (r1_xxreal_0(k3_extreal1(k4_supinf_2(X1,X2)),k3_supinf_2(k3_extreal1(X1),X2))|~r1_xxreal_0(np__0,X2)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_895, plain, (m1_subset_1(k7_mesfunc5(X1,X2,X3,X4),k7_numbers)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(k2_mesfunc2(X1,X4))|~m1_subset_1(k2_mesfunc2(X1,X4),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_845, plain, (k3_supinf_2(k6_mesfunc5(X1,X2,X3,X4),k6_mesfunc5(X1,X2,X3,X5))=k6_mesfunc5(X1,X2,X3,k3_mesfunc1(X1,X4,X5))|v1_xboole_0(X2)|v1_xboole_0(X1)|k1_relset_1(X1,X7)!=k1_relset_1(X1,X5)|k1_relset_1(X1,X7)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X5,k1_relset_1(X1,X7))|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X7))|~r1_mesfunc5(X1,X2,X6,X7)|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X5)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_844, plain, (k7_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))=k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|X5!=k1_relset_1(X1,k2_mesfunc2(X1,X4))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(k2_mesfunc2(X1,X4))|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(k2_mesfunc2(X1,X4))|~v1_funct_1(X3)|~m1_subset_1(k2_mesfunc2(X1,X4),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X5,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_tarski(X5,k1_relset_1(X1,X4))|~r1_mesfunc1(X1,X2,X4,X5)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_28, plain, (k2_supinf_2(X1)=k3_extreal1(X1)|r1_xxreal_0(np__0,X1)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_842, plain, (k7_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4))=k6_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|X5!=k1_relset_1(X1,k10_mesfunc1(X1,X4))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(k10_mesfunc1(X1,X4))|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(k10_mesfunc1(X1,X4))|~v1_funct_1(X3)|~m1_subset_1(k10_mesfunc1(X1,X4),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X5,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_tarski(X5,k1_relset_1(X1,X4))|~r1_mesfunc1(X1,X2,X4,X5)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_352, negated_conjecture, (v1_xboole_0(esk4_0)|~v1_xboole_0(k2_zfmisc_1(esk1_0,k7_numbers)))).#trainneg
cnf(i_0_75, plain, (m1_subset_1(np__0,k7_numbers))).#trainneg
cnf(i_0_37, plain, (X1=k1_mesfunc2(X2,X3)|v1_xboole_0(X2)|k12_supinf_2(X1,esk8_3(X2,X3,X1))!=k1_extreal2(k12_supinf_2(X3,esk8_3(X2,X3,X1)),np__0)|k1_relset_1(X2,X1)!=k1_relset_1(X2,X3)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).#trainneg
cnf(i_0_269, plain, (v3_xxreal_0(np__0)|~v2_xxreal_0(np__0)|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_780, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_776, plain, (v2_valued_0(k2_zfmisc_1(X1,X2))|~v2_membered(X2))).#trainneg
cnf(i_0_273, plain, (v2_xxreal_0(k4_xcmplx_0(X1))|~v3_xxreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_274, plain, (v2_xxreal_0(np__0)|~v3_xxreal_0(np__0)|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_40, plain, (k12_supinf_2(X1,X2)=k1_extreal2(k12_supinf_2(X3,X2),np__0)|v1_xboole_0(X4)|X1!=k1_mesfunc2(X4,X3)|~r2_hidden(X2,k1_relset_1(X4,X1))|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,X4))).#trainneg
cnf(i_0_843, plain, (k7_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))=k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|X5!=k1_relset_1(X1,k1_mesfunc2(X1,X4))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(k1_mesfunc2(X1,X4))|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(k1_mesfunc2(X1,X4))|~v1_funct_1(X3)|~m1_subset_1(k1_mesfunc2(X1,X4),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X5,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_mesfunc1(X1,X2,X4,X5)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_353, negated_conjecture, (v1_xboole_0(esk3_0)|~v1_xboole_0(k2_zfmisc_1(esk2_0,k7_numbers)))).#trainneg
cnf(i_0_348, plain, (v1_xboole_0(esk15_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_42, plain, (X1=k2_mesfunc2(X2,X3)|v1_xboole_0(X2)|k12_supinf_2(X1,esk9_3(X2,X3,X1))!=k1_extreal2(k2_supinf_2(k12_supinf_2(X3,esk9_3(X2,X3,X1))),np__0)|k1_relset_1(X2,X1)!=k1_relset_1(X2,X3)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).#trainneg
cnf(i_0_280, plain, (k4_xxreal_0(X1,X1)=X1|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_311, plain, (v1_xcmplx_0(k1_supinf_1)|~v1_membered(k7_numbers))).#trainneg
cnf(i_0_777, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_778, plain, (v3_valued_0(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_283, plain, (v1_xxreal_0(esk15_1(X1))|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_284, plain, (v1_xxreal_0(esk14_1(X1))|~v2_membered(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_591, plain, (k1_relset_1(X1,k10_mesfunc1(X1,X2))=k1_relset_1(X1,X3)|v1_xboole_0(X1)|k10_mesfunc1(X1,X2)!=k1_mesfunc2(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(k10_mesfunc1(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_806, plain, (v1_funct_1(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v5_relat_1(k2_mesfunc2(X1,X2),k7_numbers)|~v5_relat_1(k1_mesfunc2(X1,X2),k7_numbers)|~v4_relat_1(k2_mesfunc2(X1,X2),X1)|~v4_relat_1(k1_mesfunc2(X1,X2),X1)|~v1_relat_1(k2_mesfunc2(X1,X2))|~v1_relat_1(k1_mesfunc2(X1,X2))|~v1_funct_1(k2_mesfunc2(X1,X2))|~v1_funct_1(k1_mesfunc2(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_596, negated_conjecture, (k1_relset_1(esk1_0,esk4_0)=k1_relset_1(esk1_0,X1)|k2_mesfunc2(esk1_0,X1)!=esk4_0|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_597, negated_conjecture, (k1_relset_1(esk2_0,esk3_0)=k1_relset_1(esk2_0,X1)|k2_mesfunc2(esk2_0,X1)!=esk3_0|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_600, plain, (k1_relset_1(X1,k10_mesfunc1(X1,X2))=k1_relset_1(X1,X3)|v1_xboole_0(X1)|k10_mesfunc1(X1,X2)!=k2_mesfunc2(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(k10_mesfunc1(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_561, plain, (v4_relat_1(k2_mesfunc2(X1,X2),X1)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_581, plain, (r2_relset_1(X1,k7_numbers,k3_mesfunc1(X1,k1_mesfunc2(X1,X2),k2_mesfunc2(X1,X2)),k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~m1_subset_1(k3_mesfunc1(X1,k1_mesfunc2(X1,X2),k2_mesfunc2(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_798, negated_conjecture, (m1_subset_1(k12_supinf_2(esk3_0,X1),k7_numbers))).#trainneg
cnf(i_0_373, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k4_ordinal1))).#trainneg
cnf(i_0_351, plain, (v1_xboole_0(esk11_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_457, negated_conjecture, (v3_finsub_1(esk2_0))).#trainneg
cnf(i_0_45, plain, (k12_supinf_2(X1,X2)=k1_extreal2(k2_supinf_2(k12_supinf_2(X3,X2)),np__0)|v1_xboole_0(X4)|X1!=k2_mesfunc2(X4,X3)|~r2_hidden(X2,k1_relset_1(X4,X1))|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,X4))).#trainneg
cnf(i_0_285, plain, (v1_xxreal_0(esk12_1(X1))|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_341, plain, (v7_ordinal1(k1_supinf_1)|~v6_membered(k7_numbers))).#trainneg
cnf(i_0_779, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_781, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_286, plain, (v1_xxreal_0(esk11_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_287, negated_conjecture, (v1_xxreal_0(esk4_0)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_925, negated_conjecture, (k7_mesfunc5(esk1_0,esk2_0,X1,esk4_0)=k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0)|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(esk4_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_989, plain, (m1_subset_1(k7_mesfunc5(X1,X2,X3,X4),k7_numbers)|v1_xboole_0(X1)|v1_xboole_0(X2)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(k1_mesfunc2(X1,X4))|~m1_subset_1(k1_mesfunc2(X1,X4),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_969, plain, (v6_supinf_2(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1008, plain, (k1_extreal2(X1,X1)=X1|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_940, plain, (r1_xxreal_0(k3_extreal1(k4_supinf_2(X1,X2)),k3_supinf_2(X2,k3_extreal1(X1)))|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)|~r1_xxreal_0(np__0,X2))).#trainneg
cnf(i_0_1036, plain, (k1_relset_1(X1,k2_mesfunc2(X1,X2))=k9_xtuple_0(k2_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_relat_1(k2_mesfunc2(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_951, plain, (r1_xxreal_0(k3_extreal1(k4_supinf_2(X1,X2)),k3_supinf_2(X1,k3_extreal1(X2)))|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)|~r1_xxreal_0(np__0,X1))).#trainneg
cnf(i_0_372, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(esk1_0))|~r2_hidden(X1,esk2_0))).#trainneg
cnf(i_0_793, plain, (esk12_1(X1)=esk12_1(X2))).#trainneg
cnf(i_0_143, plain, (k2_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_288, negated_conjecture, (v1_xxreal_0(esk3_0)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_290, plain, (v1_xxreal_0(k4_ordinal1)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_312, plain, (v1_xcmplx_0(k3_extreal1(X1))|~v1_membered(k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_289, negated_conjecture, (v1_xxreal_0(esk2_0)|~v2_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_295, plain, (v1_xxreal_0(k2_supinf_2(X1))|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_1023, plain, (k7_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))=k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|X5!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,X5)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X5,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_942, plain, (r1_xxreal_0(k3_extreal1(k3_xxreal_3(X1,X2)),k3_supinf_2(X2,k3_extreal1(X1)))|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)|~r1_xxreal_0(np__0,X2))).#trainneg
cnf(i_0_954, plain, (k3_supinf_2(k6_mesfunc5(X1,X2,X3,X4),k6_mesfunc5(X1,X2,X3,X5))=k6_mesfunc5(X1,X2,X3,k3_mesfunc1(X1,X4,X5))|v1_xboole_0(X1)|v1_xboole_0(X2)|k1_relset_1(X1,X5)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X5))|~r1_mesfunc5(X1,X2,X6,X5)|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X5)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_941, plain, (r1_xxreal_0(k3_extreal1(k3_xxreal_3(X1,X2)),k3_supinf_2(k3_extreal1(X1),X2))|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)|~r1_xxreal_0(np__0,X2))).#trainneg
cnf(i_0_952, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|esk5_4(X1,X2,X5,X6)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,esk5_4(X1,X2,X5,X6))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_mesfunc5(X1,X2,X5,X6)|~v4_measure1(X5,X1,X2)|~v6_supinf_2(X5)|~v10_valued_0(X5)|~v1_funct_2(X5,X2,k7_numbers)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).#trainneg
cnf(i_0_946, plain, (r1_xxreal_0(k3_extreal1(k4_supinf_2(X1,X2)),k3_supinf_2(k3_extreal1(X2),X1))|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)|~r1_xxreal_0(np__0,X1))).#trainneg
cnf(i_0_1041, negated_conjecture, (v1_xxreal_0(k12_supinf_2(esk3_0,X1)))).#trainneg
cnf(i_0_390, plain, (v1_valued_0(k4_ordinal1)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_172, plain, (k6_xcmplx_0(X1,X2)=k3_xxreal_3(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_298, plain, (v1_xcmplx_0(esk15_1(X1))|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_398, plain, (v3_valued_0(k4_ordinal1)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_305, plain, (v1_xcmplx_0(k4_ordinal1)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_320, plain, (v1_xreal_0(k4_ordinal1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_342, plain, (v7_ordinal1(k3_extreal1(X1))|~v6_membered(k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_402, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_299, plain, (v1_xcmplx_0(esk14_1(X1))|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_300, plain, (v1_xcmplx_0(esk12_1(X1))|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_953, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|X5!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,X5)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~m1_subset_1(X5,X2))).#trainneg
cnf(i_0_1090, plain, (k9_xtuple_0(k2_mesfunc2(X1,X2))=k1_relset_1(X1,X2)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_relat_1(k2_mesfunc2(X1,X2)))).#trainneg
cnf(i_0_1174, negated_conjecture, (r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(esk4_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_1064, negated_conjecture, (~r1_xxreal_0(k3_extreal1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0)),k7_mesfunc5(esk1_0,esk2_0,esk3_0,k10_mesfunc1(esk1_0,esk4_0)))|~v6_supinf_2(esk4_0))).#trainneg
cnf(i_0_1149, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|k1_relset_1(X1,X6)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X6))|~r1_mesfunc5(X1,X2,X5,X6)|~v4_measure1(X3,X1,X2)|~v4_measure1(X5,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X5)|~v10_valued_0(X3)|~v10_valued_0(X5)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X5,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1150, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|esk5_4(X1,X2,X5,X4)!=k1_relset_1(X1,X4)|~r1_mesfunc5(X1,X2,X5,X4)|~v4_measure1(X3,X1,X2)|~v4_measure1(X5,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X5)|~v10_valued_0(X3)|~v10_valued_0(X5)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X5,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X5)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1103, negated_conjecture, (k7_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2))=k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2))|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(k1_mesfunc2(esk1_0,X2))|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k1_mesfunc2(esk1_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k1_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_1131, negated_conjecture, (k7_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,esk4_0))=k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_1144, plain, (r1_xxreal_0(k3_extreal1(k3_xxreal_3(X1,X2)),k3_supinf_2(X1,X2))|~r1_xxreal_0(np__0,X2)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)|~r1_xxreal_0(np__0,X1))).#trainneg
cnf(i_0_1111, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk12_1(X3)))).#trainneg
cnf(i_0_1098, negated_conjecture, (v1_xxreal_0(k1_relset_1(esk1_0,esk4_0))|~v2_membered(esk2_0))).#trainneg
cnf(i_0_52, plain, (r2_relset_1(X1,X2,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_301, plain, (v1_xcmplx_0(esk11_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_360, negated_conjecture, (v3_membered(esk4_0)|~v3_membered(k2_zfmisc_1(esk1_0,k7_numbers)))).#trainneg
cnf(i_0_335, plain, (v7_ordinal1(k4_ordinal1)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_362, negated_conjecture, (v3_membered(esk2_0)|~v3_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1164, plain, (v1_xcmplx_0(k3_extreal1(X1))|~v6_membered(k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_768, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_302, negated_conjecture, (v1_xcmplx_0(esk4_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_303, negated_conjecture, (v1_xcmplx_0(esk3_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_1082, plain, (k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X1)|v1_xboole_0(X2)|k9_xtuple_0(k2_mesfunc2(X1,X5))!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k9_xtuple_0(k2_mesfunc2(X1,X5)))|~r1_mesfunc5(X1,X2,X6,k2_mesfunc2(X1,X5))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(k2_mesfunc2(X1,X5))|~v1_funct_1(X6)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X1,X5),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_relat_1(k2_mesfunc2(X1,X5))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1063, negated_conjecture, (~r1_xxreal_0(k6_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0),k7_mesfunc5(esk1_0,esk2_0,esk3_0,k10_mesfunc1(esk1_0,esk4_0)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0),k7_numbers)|~v6_supinf_2(esk4_0))).#trainneg
cnf(i_0_1092, plain, (r1_xxreal_0(k3_extreal1(k4_supinf_2(X1,X2)),k3_supinf_2(X1,X2))|~r1_xxreal_0(np__0,X1)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)|~r1_xxreal_0(np__0,X2))).#trainneg
cnf(i_0_1178, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)))|v1_xboole_0(X1)|v1_xboole_0(X2)|X5!=k1_relset_1(X1,k2_mesfunc2(X1,X4))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(k2_mesfunc2(X1,X4))|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(k2_mesfunc2(X1,X4))|~v1_funct_1(X3)|~m1_subset_1(k2_mesfunc2(X1,X4),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X5,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_tarski(X5,k1_relset_1(X1,X4))|~r1_mesfunc1(X1,X2,X4,X5)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_919, plain, (k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X1)|v1_xboole_0(X2)|k1_relset_1(X1,X5)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X5))|~r1_mesfunc5(X1,X2,X6,k1_mesfunc2(X1,X5))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(k1_mesfunc2(X1,X5))|~v1_funct_1(X6)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X1,X5),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_980, plain, (r1_xxreal_0(k3_extreal1(k7_mesfunc5(X1,X2,X3,X4)),k3_supinf_2(k3_extreal1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))),k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)))|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1089, plain, (k7_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))=k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|X5!=k9_xtuple_0(k2_mesfunc2(X1,X4))|~r1_tarski(X5,k1_relset_1(X1,X4))|~r1_mesfunc1(X1,X2,X4,X5)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X5,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_relat_1(k2_mesfunc2(X1,X4)))).#trainneg
cnf(i_0_1104, negated_conjecture, (k7_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2))=k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2))|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(k2_mesfunc2(esk1_0,X2))|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k2_mesfunc2(esk1_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k2_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|~r1_tarski(k1_relset_1(esk1_0,esk4_0),k1_relset_1(esk1_0,X2))|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_494, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_1099, negated_conjecture, (v1_xcmplx_0(k1_relset_1(esk1_0,esk4_0))|~v1_membered(esk2_0))).#trainneg
cnf(i_0_605, negated_conjecture, (k9_xtuple_0(esk4_0)=k9_xtuple_0(X1)|k10_mesfunc1(esk1_0,X1)!=esk4_0|~v5_relat_1(X1,k7_numbers)|~v4_relat_1(X1,esk1_0)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_304, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_361, negated_conjecture, (v3_membered(esk3_0)|~v3_membered(k2_zfmisc_1(esk2_0,k7_numbers)))).#trainneg
cnf(i_0_1123, plain, (v1_xxreal_0(k4_ordinal1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1163, plain, (v1_xcmplx_0(k4_ordinal1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_769, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_770, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_310, plain, (v1_xcmplx_0(k2_supinf_2(X1))|~v1_membered(k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_313, plain, (v1_xreal_0(esk15_1(X1))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1086, plain, (k3_supinf_2(k6_mesfunc5(X1,X2,X3,X4),k6_mesfunc5(X1,X2,X3,X5))=k6_mesfunc5(X1,X2,X3,k3_mesfunc1(X1,X4,X5))|v1_xboole_0(X1)|v1_xboole_0(X2)|k9_xtuple_0(k2_mesfunc2(X1,X6))!=k1_relset_1(X1,X5)|k9_xtuple_0(k2_mesfunc2(X1,X6))!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X5,k9_xtuple_0(k2_mesfunc2(X1,X6)))|~r1_mesfunc1(X1,X2,X4,k9_xtuple_0(k2_mesfunc2(X1,X6)))|~r1_mesfunc5(X1,X2,X7,k2_mesfunc2(X1,X6))|~v4_measure1(X3,X1,X2)|~v4_measure1(X7,X1,X2)|~v6_supinf_2(X5)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X7)|~v10_valued_0(X3)|~v10_valued_0(X7)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X7,X2,k7_numbers)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(k2_mesfunc2(X1,X6))|~v1_funct_1(X7)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X1,X6),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_relat_1(k2_mesfunc2(X1,X6))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_918, plain, (k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X1)|v1_xboole_0(X2)|k1_relset_1(X1,X5)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X5))|~r1_mesfunc5(X1,X2,X6,k2_mesfunc2(X1,X5))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(k2_mesfunc2(X1,X5))|~v1_funct_1(X6)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X1,X5),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1132, plain, (k7_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))=k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|esk5_4(X1,X2,X5,X4)!=k1_relset_1(X1,X4)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(esk5_4(X1,X2,X5,X4),X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_mesfunc5(X1,X2,X5,X4)|~v4_measure1(X5,X1,X2)|~v6_supinf_2(X5)|~v10_valued_0(X5)|~v1_funct_2(X5,X2,k7_numbers)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).#trainneg
cnf(i_0_934, plain, (r1_xxreal_0(k3_extreal1(k7_mesfunc5(X1,X2,X3,X4)),k3_supinf_2(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k3_extreal1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)))))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))))).#trainneg
cnf(i_0_935, plain, (r1_xxreal_0(k3_extreal1(k7_mesfunc5(X1,X2,X3,X4)),k3_supinf_2(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k3_extreal1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)))))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))))).#trainneg
cnf(i_0_997, plain, (k7_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))=k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|X5!=k1_relset_1(X1,X4)|~r1_tarski(X5,k1_relset_1(X1,X4))|~r1_mesfunc1(X1,X2,X4,X5)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X5,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1155, plain, (r1_xxreal_0(k3_extreal1(k4_supinf_2(X1,X2)),k3_supinf_2(X2,X1))|~r1_xxreal_0(np__0,X1)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)|~r1_xxreal_0(np__0,X2))).#trainneg
cnf(i_0_1100, negated_conjecture, (v1_xreal_0(k1_relset_1(esk1_0,esk4_0))|~v3_membered(esk2_0))).#trainneg
cnf(i_0_606, negated_conjecture, (k9_xtuple_0(esk3_0)=k9_xtuple_0(X1)|k10_mesfunc1(esk2_0,X1)!=esk3_0|~v5_relat_1(X1,k7_numbers)|~v4_relat_1(X1,esk2_0)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_314, plain, (v1_xreal_0(esk14_1(X1))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_319, negated_conjecture, (v1_xreal_0(esk2_0)|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_334, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_771, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1205, plain, (v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_315, plain, (v1_xreal_0(esk12_1(X1))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_316, plain, (v1_xreal_0(esk11_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_1286, plain, (k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|k1_relset_1(X1,X5)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X5))|~r1_mesfunc5(X1,X2,X6,k2_mesfunc2(X1,k1_mesfunc2(X1,X5)))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X6)|~v1_funct_1(k1_mesfunc2(X1,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X1,X5),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1223, plain, (k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|k1_relset_1(X1,X5)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X5))|~r1_mesfunc5(X1,X2,X6,k1_mesfunc2(X1,k1_mesfunc2(X1,X5)))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X6)|~v1_funct_1(k1_mesfunc2(X1,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X1,X5),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1183, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|k1_relset_1(X1,X5)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X5))|~r1_mesfunc5(X1,X2,X6,k2_mesfunc2(X1,X5))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(k2_mesfunc2(X1,X5))|~v1_funct_1(X6)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X1,X5),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1222, plain, (k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|k1_relset_1(X1,X5)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X5))|~r1_mesfunc5(X1,X2,X6,k1_mesfunc2(X1,k2_mesfunc2(X1,X5)))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X6)|~v1_funct_1(k2_mesfunc2(X1,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X1,X5),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1285, plain, (k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|k1_relset_1(X1,X5)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X5))|~r1_mesfunc5(X1,X2,X6,k2_mesfunc2(X1,k2_mesfunc2(X1,X5)))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X6)|~v1_funct_1(k2_mesfunc2(X1,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X1,X5),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1184, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|k1_relset_1(X1,X5)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X5))|~r1_mesfunc5(X1,X2,X6,k1_mesfunc2(X1,X5))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(k1_mesfunc2(X1,X5))|~v1_funct_1(X6)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X1,X5),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_913, plain, (m1_subset_1(X1,k7_numbers)|~m1_subset_1(k2_xxreal_3(X1),k7_numbers)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1203, negated_conjecture, (v1_xxreal_0(k1_relset_1(esk1_0,esk4_0))|~v3_membered(esk2_0))).#trainneg
cnf(i_0_587, negated_conjecture, (k1_relset_1(esk1_0,esk4_0)=k1_relset_1(esk1_0,X1)|k1_mesfunc2(esk1_0,X1)!=esk4_0|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_317, negated_conjecture, (v1_xreal_0(esk4_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_1126, negated_conjecture, (v1_xxreal_0(esk2_0)|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_1264, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(esk1_0))))).#trainneg
cnf(i_0_1265, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1254, negated_conjecture, (m1_subset_1(X1,k2_zfmisc_1(esk1_0,k7_numbers))|v1_xboole_0(esk4_0)|~m1_subset_1(X1,esk4_0))).#trainneg
cnf(i_0_318, negated_conjecture, (v1_xreal_0(esk3_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_1199, plain, (r1_xxreal_0(k3_xxreal_3(X1,X2),k3_supinf_2(X1,X2))|~r1_xxreal_0(np__0,X2)|~r1_xxreal_0(np__0,X1)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)|~r1_xxreal_0(np__0,k3_xxreal_3(X1,X2))|~m1_subset_1(k3_xxreal_3(X1,X2),k7_numbers))).#trainneg
cnf(i_0_1217, plain, (r1_xxreal_0(k4_supinf_2(X1,X2),k3_supinf_2(X1,X2))|~r1_xxreal_0(np__0,X1)|~r1_xxreal_0(np__0,X2)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)|~r1_xxreal_0(np__0,k4_supinf_2(X1,X2))|~m1_subset_1(k4_supinf_2(X1,X2),k7_numbers))).#trainneg
cnf(i_0_1003, plain, (k7_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4))=k6_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_tarski(k1_relset_1(X1,k10_mesfunc1(X1,X4)),k1_relset_1(X1,X4))|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,k10_mesfunc1(X1,X4)))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(k10_mesfunc1(X1,X4))|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(k10_mesfunc1(X1,X4),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k1_relset_1(X1,k10_mesfunc1(X1,X4)),X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1070, negated_conjecture, (m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k7_numbers)|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|~v6_supinf_2(esk4_0))).#trainneg
cnf(i_0_1088, plain, (k12_supinf_2(k2_mesfunc2(X1,X2),X3)=k1_extreal2(k12_supinf_2(X4,X3),np__0)|v1_xboole_0(X1)|k2_mesfunc2(X1,X2)!=k1_mesfunc2(X1,X4)|~r2_hidden(X3,k9_xtuple_0(k2_mesfunc2(X1,X2)))|~v1_funct_1(X4)|~v1_funct_1(k2_mesfunc2(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1)|~v1_relat_1(k2_mesfunc2(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_876, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk15_1(X2))|~m1_subset_1(X1,esk15_1(X2)))).#trainneg
cnf(i_0_1263, negated_conjecture, (v1_xcmplx_0(k1_relset_1(esk1_0,esk4_0))|~v3_membered(esk2_0))).#trainneg
cnf(i_0_588, negated_conjecture, (k1_relset_1(esk2_0,esk3_0)=k1_relset_1(esk2_0,X1)|k1_mesfunc2(esk2_0,X1)!=esk3_0|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_328, plain, (v7_ordinal1(esk15_1(X1))|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_332, negated_conjecture, (v7_ordinal1(esk4_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_333, negated_conjecture, (v7_ordinal1(esk3_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_1255, negated_conjecture, (m1_subset_1(X1,k2_zfmisc_1(esk2_0,k7_numbers))|v1_xboole_0(esk3_0)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_329, plain, (v7_ordinal1(esk14_1(X1))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_330, plain, (v7_ordinal1(esk12_1(X1))|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1159, plain, (r1_xxreal_0(k3_extreal1(k3_xxreal_3(X1,X2)),k3_supinf_2(k3_extreal1(X2),X1))|~r1_xxreal_0(np__0,X1)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_1016, plain, (k12_supinf_2(k2_mesfunc2(X1,X2),X3)=k1_extreal2(k12_supinf_2(X4,X3),np__0)|v1_xboole_0(X1)|k2_mesfunc2(X1,X2)!=k1_mesfunc2(X1,X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X4)|~v1_funct_1(k2_mesfunc2(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1017, plain, (k12_supinf_2(k1_mesfunc2(X1,X2),X3)=k1_extreal2(k12_supinf_2(X4,X3),np__0)|v1_xboole_0(X1)|k1_mesfunc2(X1,X2)!=k1_mesfunc2(X1,X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X4)|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1096, plain, (r1_xxreal_0(k3_extreal1(k3_xxreal_3(X1,X2)),k3_supinf_2(X1,k3_extreal1(X2)))|~r1_xxreal_0(np__0,X1)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_976, plain, (r1_xxreal_0(k3_extreal1(k4_supinf_2(k3_extreal1(X1),X2)),k3_supinf_2(k3_extreal1(X1),X2))|~r1_xxreal_0(np__0,X2)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(k3_extreal1(X1),k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_1037, plain, (m1_subset_1(k1_relset_1(X1,k2_mesfunc2(X1,X2)),k1_zfmisc_1(X1))|v1_xboole_0(X1)|~v1_relat_1(k2_mesfunc2(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_797, negated_conjecture, (m1_subset_1(k12_supinf_2(esk4_0,X1),k7_numbers))).#trainneg
cnf(i_0_1101, negated_conjecture, (v7_ordinal1(k1_relset_1(esk1_0,esk4_0))|~v6_membered(esk2_0))).#trainneg
cnf(i_0_331, plain, (v7_ordinal1(esk11_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_401, plain, (r1_xxreal_0(X1,k1_supinf_1)|r1_xxreal_0(k1_supinf_1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1443, plain, (r1_xxreal_0(k1_supinf_1,k1_supinf_1))).#trainneg
cnf(i_0_1060, negated_conjecture, (v1_xxreal_0(esk4_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_370, negated_conjecture, (m1_subset_1(X1,k2_zfmisc_1(esk1_0,k7_numbers))|~r2_hidden(X1,esk4_0))).#trainneg
cnf(i_0_371, negated_conjecture, (m1_subset_1(X1,k2_zfmisc_1(esk2_0,k7_numbers))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_340, plain, (v7_ordinal1(k2_supinf_2(X1))|~v6_membered(k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_343, plain, (v1_relat_1(esk15_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1383, plain, (k12_supinf_2(k1_mesfunc2(X1,X2),X3)=k1_extreal2(k12_supinf_2(X2,X3),np__0)|v1_xboole_0(X1)|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_1331, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|k1_relset_1(X1,X5)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X5))|~r1_mesfunc5(X1,X2,X6,k1_mesfunc2(X1,k1_mesfunc2(X1,X5)))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X6)|~v1_funct_1(k1_mesfunc2(X1,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X1,X5),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1324, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|k1_relset_1(X1,X5)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X5))|~r1_mesfunc5(X1,X2,X6,k2_mesfunc2(X1,k1_mesfunc2(X1,X5)))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X6)|~v1_funct_1(k1_mesfunc2(X1,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X1,X5),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1350, plain, (r1_xxreal_0(k3_xxreal_3(X1,X2),k3_supinf_2(X1,X2))|~r1_xxreal_0(np__0,k3_xxreal_3(X1,X2))|~r1_xxreal_0(np__0,X1)|~r1_xxreal_0(np__0,X2)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_1338, plain, (m1_subset_1(X1,k7_numbers)|~v1_xxreal_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k7_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1428, negated_conjecture, (v1_xxreal_0(k12_supinf_2(esk4_0,X1)))).#trainneg
cnf(i_0_1440, plain, (r1_xxreal_0(k1_supinf_1,np__0)|r1_xxreal_0(np__0,k1_supinf_1))).#trainneg
cnf(i_0_912, plain, (k2_xxreal_3(X1)=k2_supinf_2(X1)|~m1_subset_1(k2_xxreal_3(X1),k7_numbers)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_344, plain, (v1_relat_1(esk12_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1446, plain, (r1_xxreal_0(k1_supinf_1,X1)|r1_xxreal_0(X1,k1_supinf_1)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_1122, negated_conjecture, (v1_xxreal_0(esk3_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_1206, negated_conjecture, (v1_xcmplx_0(esk4_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_403, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k7_numbers))).#trainneg
cnf(i_0_345, plain, (v1_relat_1(esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_356, plain, (v3_membered(esk15_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_1330, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|k1_relset_1(X1,X5)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X5))|~r1_mesfunc5(X1,X2,X6,k1_mesfunc2(X1,k2_mesfunc2(X1,X5)))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X6)|~v1_funct_1(k2_mesfunc2(X1,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X1,X5),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1323, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|k1_relset_1(X1,X5)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X5))|~r1_mesfunc5(X1,X2,X6,k2_mesfunc2(X1,k2_mesfunc2(X1,X5)))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X6)|~v1_funct_1(k2_mesfunc2(X1,X5))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X1,X5),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1378, plain, (r1_xxreal_0(k3_extreal1(k3_xxreal_3(X1,X2)),k3_supinf_2(X2,X1))|~r1_xxreal_0(np__0,X1)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)|~r1_xxreal_0(np__0,X2))).#trainneg
cnf(i_0_1022, plain, (k7_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))=k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,k1_mesfunc2(X1,X4)))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k1_relset_1(X1,k1_mesfunc2(X1,X4)),X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1176, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4)))|v1_xboole_0(X1)|v1_xboole_0(X2)|X5!=k1_relset_1(X1,k10_mesfunc1(X1,X4))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(k10_mesfunc1(X1,X4))|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(k10_mesfunc1(X1,X4))|~v1_funct_1(X3)|~m1_subset_1(k10_mesfunc1(X1,X4),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X5,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_tarski(X5,k1_relset_1(X1,X4))|~r1_mesfunc1(X1,X2,X4,X5)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1109, plain, (m1_subset_1(esk12_1(X2),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1432, negated_conjecture, (v1_xcmplx_0(k1_relset_1(esk1_0,esk4_0))|~v6_membered(esk2_0))).#trainneg
cnf(i_0_1024, plain, (X1=esk15_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_357, plain, (v3_membered(esk14_1(X1))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_962, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1207, negated_conjecture, (v1_xcmplx_0(esk3_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_389, negated_conjecture, (v1_valued_0(esk2_0)|~v1_valued_0(k1_zfmisc_1(esk1_0))|~v1_relat_1(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_397, negated_conjecture, (v3_valued_0(esk2_0)|~v3_valued_0(k1_zfmisc_1(esk1_0))|~v1_relat_1(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_359, plain, (v3_membered(esk11_1(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_383, plain, (v1_valued_0(esk15_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1420, plain, (m1_subset_1(k1_relset_1(X1,X2),k1_zfmisc_1(X1))|v1_xboole_0(X1)|~v1_relat_1(k2_mesfunc2(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1419, plain, (m1_subset_1(k9_xtuple_0(k2_mesfunc2(X1,X2)),k1_zfmisc_1(X1))|v1_xboole_0(X1)|~v1_relat_1(k2_mesfunc2(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1481, plain, (v1_relat_1(esk12_1(X3)))).#trainneg
cnf(i_0_888, plain, (X1=k4_xcmplx_0(X2)|~v1_xboole_0(X1)|~v1_xcmplx_0(X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_384, plain, (v1_valued_0(esk14_1(X1))|~v1_valued_0(k1_zfmisc_1(X1))|~v1_relat_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1579, negated_conjecture, (v1_valued_0(esk2_0)|~v1_relat_1(k1_zfmisc_1(esk1_0))|~v3_valued_0(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1580, negated_conjecture, (v3_valued_0(esk2_0)|~v1_relat_1(k1_zfmisc_1(esk1_0))|~v5_relat_1(k1_zfmisc_1(esk1_0),k1_numbers))).#trainneg
cnf(i_0_386, plain, (v1_valued_0(esk11_1(k1_zfmisc_1(X1)))|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_387, negated_conjecture, (v1_valued_0(esk4_0)|~v1_valued_0(k2_zfmisc_1(esk1_0,k7_numbers))|~v1_relat_1(k2_zfmisc_1(esk1_0,k7_numbers)))).#trainneg
cnf(i_0_1531, negated_conjecture, (r1_xxreal_0(k3_extreal1(k7_mesfunc5(esk1_0,esk2_0,X1,X2)),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,k2_mesfunc2(esk1_0,X2),k1_mesfunc2(esk1_0,X2))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~v6_supinf_2(k1_mesfunc2(esk1_0,X2))|~v6_supinf_2(k2_mesfunc2(esk1_0,X2))|~v1_funct_1(k1_mesfunc2(esk1_0,X2))|~v1_funct_1(k2_mesfunc2(esk1_0,X2))|~m1_subset_1(k1_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k2_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_1527, negated_conjecture, (r1_xxreal_0(k3_extreal1(k4_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,X3))),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X3,X2)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X2))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X3))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X3),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X3)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,X3,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X2)|~v6_supinf_2(X3)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_1521, negated_conjecture, (r1_xxreal_0(k3_extreal1(k4_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,X3))),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X2,X3)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X2))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X3))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X3),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X3)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X3,k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X3)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_1524, negated_conjecture, (r1_xxreal_0(k3_xxreal_3(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,X3)),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X2,X3)))|~r1_xxreal_0(np__0,k3_xxreal_3(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,X3)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X2))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X3))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X3),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X3)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X3,k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X3)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_1522, negated_conjecture, (r1_xxreal_0(k3_extreal1(k3_xxreal_3(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,X3))),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X2,X3)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X3))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X2))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X3),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X3)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X3,k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X3)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_1528, negated_conjecture, (r1_xxreal_0(k3_extreal1(k3_xxreal_3(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,X3))),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X3,X2)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X2))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X3))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X3),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X3)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,X3,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X2)|~v6_supinf_2(X3)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_1457, plain, (k12_supinf_2(k1_mesfunc2(X1,X2),X3)=k12_supinf_2(k1_mesfunc2(X4,X2),X3)|v1_xboole_0(X1)|v1_xboole_0(X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1)|~r2_hidden(X3,k1_relset_1(X4,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X3,X4))).#trainneg
cnf(i_0_1465, plain, (k12_supinf_2(k1_mesfunc2(X1,k1_mesfunc2(X2,X3)),X4)=k1_extreal2(k1_extreal2(k12_supinf_2(X3,X4),np__0),np__0)|v1_xboole_0(X1)|v1_xboole_0(X2)|~r2_hidden(X4,k1_relset_1(X1,k1_mesfunc2(X2,X3)))|~v1_funct_1(k1_mesfunc2(X2,X3))|~m1_subset_1(k1_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,X1)|~r2_hidden(X4,k1_relset_1(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_1523, negated_conjecture, (r1_xxreal_0(k4_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,X3)),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X2,X3)))|~r1_xxreal_0(np__0,k4_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,X3)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X2))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X3))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X3),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X3)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X3,k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X3)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_1547, plain, (v5_relat_1(esk12_1(X1),X2))).#trainneg
cnf(i_0_1483, plain, (r1_xxreal_0(esk11_1(k7_numbers),k1_supinf_1)|r1_xxreal_0(k1_supinf_1,esk11_1(k7_numbers)))).#trainneg
cnf(i_0_1047, plain, (X1=esk11_1(k1_zfmisc_1(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_388, negated_conjecture, (v1_valued_0(esk3_0)|~v1_valued_0(k2_zfmisc_1(esk2_0,k7_numbers))|~v1_relat_1(k2_zfmisc_1(esk2_0,k7_numbers)))).#trainneg
cnf(i_0_1578, plain, (r1_xxreal_0(np__0,X1)|r1_xxreal_0(X1,np__0)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_395, negated_conjecture, (v3_valued_0(esk4_0)|~v3_valued_0(k2_zfmisc_1(esk1_0,k7_numbers))|~v1_relat_1(k2_zfmisc_1(esk1_0,k7_numbers)))).#trainneg
cnf(i_0_396, negated_conjecture, (v3_valued_0(esk3_0)|~v3_valued_0(k2_zfmisc_1(esk2_0,k7_numbers))|~v1_relat_1(k2_zfmisc_1(esk2_0,k7_numbers)))).#trainneg
cnf(i_0_874, plain, (r1_xxreal_0(X1,k3_extreal1(X2))|r1_xxreal_0(k3_extreal1(X2),X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k7_numbers))).#trainneg
cnf(i_0_391, plain, (v3_valued_0(esk15_1(X1))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_392, plain, (v3_valued_0(esk14_1(X1))|~v3_valued_0(k1_zfmisc_1(X1))|~v1_relat_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1727, plain, (r1_xxreal_0(k7_mesfunc5(X1,X2,X3,X4),np__0)|r1_xxreal_0(np__0,k7_mesfunc5(X1,X2,X3,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1724, plain, (r1_xxreal_0(k6_mesfunc5(X1,X2,X3,X4),np__0)|r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1652, negated_conjecture, (r1_xxreal_0(k3_xxreal_3(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,X3)),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X3,X2)))|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X3)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,X3,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X2)|~v6_supinf_2(X3)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X3),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|~r1_xxreal_0(np__0,k3_xxreal_3(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,X3)))|~m1_subset_1(k3_xxreal_3(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,X3)),k7_numbers))).#trainneg
cnf(i_0_1632, negated_conjecture, (r1_xxreal_0(k4_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,X3)),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X3,X2)))|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X3)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,X3,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X2)|~v6_supinf_2(X3)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X3),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|~r1_xxreal_0(np__0,k4_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,X3)))|~m1_subset_1(k4_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,X3)),k7_numbers))).#trainneg
cnf(i_0_1696, plain, (k12_supinf_2(k1_mesfunc2(X1,k1_mesfunc2(X2,X3)),X4)=k1_extreal2(k12_supinf_2(k1_mesfunc2(X5,X3),X4),np__0)|v1_xboole_0(X2)|v1_xboole_0(X1)|v1_xboole_0(X5)|~r2_hidden(X4,k1_relset_1(X1,k1_mesfunc2(X2,X3)))|~r2_hidden(X4,k1_relset_1(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(k1_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,X1)|~m1_subset_1(X4,X2)|~r2_hidden(X4,k1_relset_1(X5,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X4,X5))).#trainneg
cnf(i_0_1549, plain, (v4_relat_1(esk12_1(X1),X2))).#trainneg
cnf(i_0_992, negated_conjecture, (k3_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0))=k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X2,esk4_0))|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc5(esk1_0,esk2_0,X3,esk4_0)|~v4_measure1(X1,esk1_0,esk2_0)|~v4_measure1(X3,esk1_0,esk2_0)|~v6_supinf_2(esk4_0)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v6_supinf_2(X3)|~v10_valued_0(X1)|~v10_valued_0(X3)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_2(X3,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_394, plain, (v3_valued_0(esk11_1(k1_zfmisc_1(X1)))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1612, negated_conjecture, (v1_valued_0(esk4_0)|~v3_valued_0(k2_zfmisc_1(esk1_0,k7_numbers)))).#trainneg
cnf(i_0_1709, negated_conjecture, (v1_valued_0(esk3_0)|~v3_valued_0(k2_zfmisc_1(esk2_0,k7_numbers)))).#trainneg
cnf(i_0_400, plain, (r1_xxreal_0(X1,k3_xxreal_3(X2,X3))|r1_xxreal_0(k3_xxreal_3(X2,X3),X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_404, plain, (v1_xcmplx_0(k3_xxreal_3(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1640, negated_conjecture, (r1_xxreal_0(k3_xxreal_3(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,X3)),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X3,X2)))|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X3)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X3,k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k3_xxreal_3(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,X3)))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X3)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X3),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X3,esk1_0)|~v4_relat_1(X2,esk1_0)|~v1_relat_1(X3)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1689, plain, (k1_extreal2(k1_extreal2(k12_supinf_2(X3,X4),np__0),np__0)=k1_extreal2(k12_supinf_2(k1_mesfunc2(X2,X3),X4),np__0)|v1_xboole_0(X1)|v1_xboole_0(X2)|~r2_hidden(X4,k1_relset_1(X1,k1_mesfunc2(X2,X3)))|~v1_funct_1(k1_mesfunc2(X2,X3))|~m1_subset_1(k1_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,X1)|~r2_hidden(X4,k1_relset_1(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_1536, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4)))|v1_xboole_0(X2)|v1_xboole_0(X1)|~r1_tarski(k1_relset_1(X1,k10_mesfunc1(X1,X4)),k1_relset_1(X1,X4))|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,k10_mesfunc1(X1,X4)))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(k10_mesfunc1(X1,X4))|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(k10_mesfunc1(X1,X4),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k1_relset_1(X1,k10_mesfunc1(X1,X4)),X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1453, plain, (k1_extreal2(np__0,k12_supinf_2(X1,X2))=k12_supinf_2(k1_mesfunc2(X3,X1),X2)|v1_xboole_0(X3)|~m1_subset_1(k12_supinf_2(X1,X2),k7_numbers)|~r2_hidden(X2,k1_relset_1(X3,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_1470, plain, (r1_xxreal_0(k7_mesfunc5(X1,X2,X3,X4),k3_supinf_2(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))))|v1_xboole_0(X2)|v1_xboole_0(X1)|~r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)))|~r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_xxreal_0(np__0,k7_mesfunc5(X1,X2,X3,X4))|~m1_subset_1(k7_mesfunc5(X1,X2,X3,X4),k7_numbers))).#trainneg
cnf(i_0_1471, plain, (r1_xxreal_0(k3_xxreal_3(X1,X2),k3_supinf_2(X2,X1))|~r1_xxreal_0(np__0,k3_xxreal_3(X1,X2))|~r1_xxreal_0(np__0,X1)|~r1_xxreal_0(np__0,X2)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_1171, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1026, plain, (esk15_1(X1)=esk12_1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_405, plain, (v4_relat_1(esk15_1(k2_zfmisc_1(X1,X2)),X1))).#trainneg
cnf(i_0_1728, negated_conjecture, (v3_valued_0(esk4_0)|~v5_relat_1(k2_zfmisc_1(esk1_0,k7_numbers),k1_numbers))).#trainneg
cnf(i_0_1730, negated_conjecture, (v3_valued_0(esk3_0)|~v5_relat_1(k2_zfmisc_1(esk2_0,k7_numbers),k1_numbers))).#trainneg
cnf(i_0_1747, plain, (r1_xxreal_0(k3_extreal1(k3_extreal1(X1)),X2)|r1_xxreal_0(X2,k3_extreal1(k3_extreal1(X1)))|~v1_xxreal_0(X2)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_407, plain, (v4_relat_1(esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_410, plain, (v5_relat_1(esk15_1(k2_zfmisc_1(X1,X2)),X2))).#trainneg
cnf(i_0_1785, plain, (k1_extreal2(k12_supinf_2(k1_mesfunc2(X5,X3),X4),np__0)=k1_extreal2(k12_supinf_2(k1_mesfunc2(X2,X3),X4),np__0)|v1_xboole_0(X1)|v1_xboole_0(X5)|v1_xboole_0(X2)|~r2_hidden(X4,k1_relset_1(X1,k1_mesfunc2(X2,X3)))|~v1_funct_1(k1_mesfunc2(X2,X3))|~m1_subset_1(k1_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,X1)|~r2_hidden(X4,k1_relset_1(X2,X3))|~r2_hidden(X4,k1_relset_1(X5,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X4,X2)|~m1_subset_1(X4,X5))).#trainneg
cnf(i_0_1816, plain, (k1_extreal2(k1_extreal2(k12_supinf_2(X1,X2),np__0),np__0)=k1_extreal2(k12_supinf_2(k1_mesfunc2(X3,X1),X2),np__0)|v1_xboole_0(X3)|~r2_hidden(X2,k1_relset_1(X3,X1))|~v1_funct_1(X1)|~m1_subset_1(k1_mesfunc2(X3,X1),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_1766, negated_conjecture, (r1_xxreal_0(k7_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,X2)))|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k2_mesfunc2(esk1_0,X2))|~v1_funct_1(k1_mesfunc2(esk1_0,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k10_mesfunc1(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|~r1_xxreal_0(np__0,k7_mesfunc5(esk1_0,esk2_0,X1,X2))|~m1_subset_1(k7_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers))).#trainneg
cnf(i_0_1751, negated_conjecture, (r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,esk4_0)))|r1_xxreal_0(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,esk4_0)),np__0)|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_1794, plain, (m1_subset_1(k1_relset_1(X1,esk12_1(X2)),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1875, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk15_1(X4))|~v1_xboole_0(X4))).#trainneg
cnf(i_0_1798, negated_conjecture, (v1_valued_0(esk4_0)|~v5_relat_1(k2_zfmisc_1(esk1_0,k7_numbers),k1_numbers))).#trainneg
cnf(i_0_775, plain, (k9_xtuple_0(k2_zfmisc_1(X1,k7_numbers))=k9_xtuple_0(X2)|v1_xboole_0(X1)|k10_mesfunc1(X1,X2)!=k2_zfmisc_1(X1,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_funct_1(k2_zfmisc_1(X1,k7_numbers)))).#trainneg
cnf(i_0_412, plain, (v5_relat_1(esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_1800, negated_conjecture, (v1_valued_0(esk3_0)|~v5_relat_1(k2_zfmisc_1(esk2_0,k7_numbers),k1_numbers))).#trainneg
cnf(i_0_1611, negated_conjecture, (v1_valued_0(esk2_0)|~v1_relat_1(k1_zfmisc_1(esk1_0))|~v5_relat_1(k1_zfmisc_1(esk1_0),k1_numbers))).#trainneg
cnf(i_0_796, negated_conjecture, (m1_subset_1(k1_relset_1(esk1_0,esk4_0),k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_533, plain, (v2_valued_0(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_422, plain, (k1_xxreal_3(k2_xxreal_3(X2),X1)=k3_xxreal_3(X1,X2)|~v1_xxreal_0(k2_xxreal_3(X2))|~v1_xxreal_0(X1)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_996, plain, (k7_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))=k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_tarski(k1_relset_1(X1,k2_mesfunc2(X1,X4)),k1_relset_1(X1,X4))|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,k2_mesfunc2(X1,X4)))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k1_relset_1(X1,k2_mesfunc2(X1,X4)),X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1208, plain, (k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|k1_relset_1(X1,k2_mesfunc2(X1,X5))!=k1_relset_1(X1,X4)|~v1_relat_1(k2_mesfunc2(X1,X5))|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,k2_mesfunc2(X1,X5)))|~r1_mesfunc5(X1,X2,X6,k2_mesfunc2(X1,X5))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1499, plain, (r1_xxreal_0(k7_mesfunc5(X1,X2,X3,X4),k1_supinf_1)|r1_xxreal_0(k1_supinf_1,k7_mesfunc5(X1,X2,X3,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_2025, negated_conjecture, (~r1_xxreal_0(k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0),k6_mesfunc5(esk1_0,esk2_0,esk3_0,k10_mesfunc1(esk1_0,esk4_0)))|~r1_xxreal_0(np__0,k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_subset_1(k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0),k7_numbers)|k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_1970, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,esk12_1(X3))))).#trainneg
cnf(i_0_1982, negated_conjecture, (v3_membered(k1_relset_1(esk1_0,esk4_0))|~v3_membered(esk1_0))).#trainneg
cnf(i_0_2023, negated_conjecture, (~r1_xxreal_0(k6_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0),k6_mesfunc5(esk1_0,esk2_0,esk3_0,k10_mesfunc1(esk1_0,esk4_0)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0))|~v6_supinf_2(esk4_0)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0),k7_numbers)|k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_424, plain, (k3_xxreal_3(X1,k2_xxreal_3(X2))=k1_xxreal_3(X1,X2)|~v1_xxreal_0(k2_xxreal_3(X2))|~v1_xxreal_0(X1)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_1976, negated_conjecture, (v1_xxreal_0(k1_relset_1(esk1_0,esk4_0))|~v2_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1977, negated_conjecture, (v1_xcmplx_0(k1_relset_1(esk1_0,esk4_0))|~v1_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1908, plain, (r1_xxreal_0(X1,k3_extreal1(k3_extreal1(k3_extreal1(X2))))|r1_xxreal_0(k3_extreal1(k3_extreal1(k3_extreal1(X2))),X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k7_numbers))).#trainneg
cnf(i_0_1340, negated_conjecture, (v1_xxreal_0(X1)|v1_xboole_0(esk4_0)|~v2_membered(k2_zfmisc_1(esk1_0,k7_numbers))|~m1_subset_1(X1,esk4_0))).#trainneg
cnf(i_0_425, plain, (k1_xxreal_3(X1,k2_supinf_2(X2))=k3_xxreal_3(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k7_numbers))).#trainneg
cnf(i_0_426, plain, (v2_valued_0(esk15_1(k2_zfmisc_1(X1,X2)))|~v2_membered(X2))).#trainneg
cnf(i_0_2024, negated_conjecture, (~r1_xxreal_0(k3_extreal1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0)),k6_mesfunc5(esk1_0,esk2_0,esk3_0,k10_mesfunc1(esk1_0,esk4_0)))|~v6_supinf_2(esk4_0)|k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_1949, negated_conjecture, (r1_xxreal_0(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k6_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,esk4_0)))|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,esk4_0)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|~v6_supinf_2(esk4_0))).#trainneg
cnf(i_0_1913, plain, (k1_extreal2(k12_supinf_2(k1_mesfunc2(X1,X2),X3),np__0)=k1_extreal2(k12_supinf_2(k1_mesfunc2(X4,X2),X3),np__0)|v1_xboole_0(X4)|v1_xboole_0(X1)|~r2_hidden(X3,k1_relset_1(X4,X2))|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k1_mesfunc2(X4,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X4)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_2000, plain, (k7_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))=k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X4))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k1_relset_1(X1,X4),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1946, negated_conjecture, (r1_xxreal_0(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)),k6_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,k1_mesfunc2(esk1_0,X2))))|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,X2)))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,X2)))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,X2)),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,X2)),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,X2))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,X2))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k2_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,X2)))|~v1_funct_1(k1_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,X2)))|~v1_funct_1(k1_mesfunc2(esk1_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,X2))),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,X2))),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,X2)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,X2)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k10_mesfunc1(esk1_0,k1_mesfunc2(esk1_0,X2)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(k1_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_1942, plain, (k1_extreal2(k1_extreal2(np__0,k12_supinf_2(X1,X2)),np__0)=k1_extreal2(k12_supinf_2(k1_mesfunc2(X3,X1),X2),np__0)|v1_xboole_0(X3)|~r2_hidden(X2,k1_relset_1(X3,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X2,X3)|~m1_subset_1(k12_supinf_2(X1,X2),k7_numbers))).#trainneg
cnf(i_0_2128, negated_conjecture, (k7_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0))=k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_1711, plain, (r1_xxreal_0(esk11_1(k7_numbers),np__0)|r1_xxreal_0(np__0,esk11_1(k7_numbers)))).#trainneg
cnf(i_0_428, plain, (v2_valued_0(esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v2_membered(X2))).#trainneg
cnf(i_0_1445, plain, (r1_xxreal_0(k1_supinf_1,X1)|r1_xxreal_0(X1,k1_supinf_1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1978, negated_conjecture, (v1_xreal_0(k1_relset_1(esk1_0,esk4_0))|~v3_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1979, negated_conjecture, (v7_ordinal1(k1_relset_1(esk1_0,esk4_0))|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_2077, plain, (r1_xxreal_0(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X1)))),X2)|r1_xxreal_0(X2,k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X1)))))|~v1_xxreal_0(X2)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_1341, negated_conjecture, (v1_xcmplx_0(X1)|v1_xboole_0(esk4_0)|~v1_membered(k2_zfmisc_1(esk1_0,k7_numbers))|~m1_subset_1(X1,esk4_0))).#trainneg
cnf(i_0_431, plain, (v1_valued_0(esk15_1(k2_zfmisc_1(X1,X2)))|~v1_membered(X2))).#trainneg
cnf(i_0_433, plain, (v1_valued_0(esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_membered(X2))).#trainneg
cnf(i_0_2032, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~r1_xxreal_0(np__0,k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0))|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0),k7_numbers)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_2037, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,esk15_1(X4)))|~v1_xboole_0(X4))).#trainneg
cnf(i_0_2060, negated_conjecture, (v1_xxreal_0(k1_relset_1(esk1_0,esk4_0))|~v3_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1543, plain, (k9_xtuple_0(esk12_1(X1))=k9_xtuple_0(X2)|v1_xboole_0(X3)|esk12_1(X1)!=k10_mesfunc1(X3,X2)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X3)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_funct_1(esk12_1(X1)))).#trainneg
cnf(i_0_436, plain, (v3_valued_0(esk15_1(k2_zfmisc_1(X1,X2)))|~v3_membered(X2))).#trainneg
cnf(i_0_2061, negated_conjecture, (v1_xcmplx_0(k1_relset_1(esk1_0,esk4_0))|~v3_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1983, negated_conjecture, (v1_valued_0(k1_relset_1(esk1_0,esk4_0))|~v1_valued_0(esk1_0)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1342, negated_conjecture, (v1_xreal_0(X1)|v1_xboole_0(esk4_0)|~v3_membered(k2_zfmisc_1(esk1_0,k7_numbers))|~m1_subset_1(X1,esk4_0))).#trainneg
cnf(i_0_1343, negated_conjecture, (v7_ordinal1(X1)|v1_xboole_0(esk4_0)|~v6_membered(k2_zfmisc_1(esk1_0,k7_numbers))|~m1_subset_1(X1,esk4_0))).#trainneg
cnf(i_0_438, plain, (v3_valued_0(esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v3_membered(X2))).#trainneg
cnf(i_0_443, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|r1_xxreal_0(k1_funct_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v2_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1999, plain, (k7_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))=k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|~r1_tarski(k9_xtuple_0(k2_mesfunc2(X1,X4)),k1_relset_1(X1,X4))|~r1_mesfunc1(X1,X2,X4,k9_xtuple_0(k2_mesfunc2(X1,X4)))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k9_xtuple_0(k2_mesfunc2(X1,X4)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_relat_1(k2_mesfunc2(X1,X4)))).#trainneg
cnf(i_0_2095, plain, (k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X5,X2),X3),np__0),np__0)=k1_extreal2(k12_supinf_2(k1_mesfunc2(X4,k1_mesfunc2(X1,X2)),X3),np__0)|v1_xboole_0(X4)|v1_xboole_0(X5)|v1_xboole_0(X1)|~r2_hidden(X3,k1_relset_1(X4,k1_mesfunc2(X1,X2)))|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X3,X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X5,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X5))).#trainneg
cnf(i_0_2136, plain, (k1_extreal2(k1_extreal2(k1_extreal2(np__0,k12_supinf_2(X2,X3)),np__0),np__0)=k1_extreal2(k12_supinf_2(k1_mesfunc2(X4,k1_mesfunc2(X1,X2)),X3),np__0)|v1_xboole_0(X4)|v1_xboole_0(X1)|~r2_hidden(X3,k1_relset_1(X4,k1_mesfunc2(X1,X2)))|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X3,X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k12_supinf_2(X2,X3),k7_numbers)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_2161, negated_conjecture, (r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0)))|r1_xxreal_0(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0)),np__0)|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_2244, plain, (k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X1,X2),X3),np__0),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X5,X2),X3),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X1)|v1_xboole_0(X5)|~r2_hidden(X3,k1_relset_1(X4,k1_mesfunc2(X1,X2)))|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X3,X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X5,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X5))).#trainneg
cnf(i_0_2296, plain, (k1_extreal2(k1_extreal2(k1_extreal2(np__0,k12_supinf_2(X1,X2)),np__0),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X4,X1),X2),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X3)|~r2_hidden(X2,k1_relset_1(X3,k1_mesfunc2(X4,X1)))|~r2_hidden(X2,k1_relset_1(X4,X1))|~v1_funct_1(X1)|~m1_subset_1(k1_mesfunc2(X4,X1),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(k12_supinf_2(X1,X2),k7_numbers)|~m1_subset_1(X2,X3)|~m1_subset_1(X2,X4)|~v1_funct_1(k1_mesfunc2(X4,X1)))).#trainneg
cnf(i_0_1965, plain, (v1_xboole_0(k1_relset_1(X1,esk12_1(X2)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_444, plain, (v2_finsub_1(esk14_1(X1))|~v1_finsub_1(esk14_1(X1)))).#trainneg
cnf(i_0_2184, plain, (r1_xxreal_0(X1,k1_supinf_1)|r1_xxreal_0(k1_supinf_1,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2185, negated_conjecture, (v1_xcmplx_0(k1_relset_1(esk1_0,esk4_0))|~v6_membered(k1_zfmisc_1(esk1_0)))).#trainneg
cnf(i_0_1984, negated_conjecture, (v3_valued_0(k1_relset_1(esk1_0,esk4_0))|~v3_valued_0(esk1_0)|~v1_relat_1(esk1_0))).#trainneg
cnf(i_0_1369, negated_conjecture, (v1_xxreal_0(X1)|v1_xboole_0(esk3_0)|~v2_membered(k2_zfmisc_1(esk2_0,k7_numbers))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_1370, negated_conjecture, (v1_xcmplx_0(X1)|v1_xboole_0(esk3_0)|~v1_membered(k2_zfmisc_1(esk2_0,k7_numbers))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_447, plain, (v1_xxreal_0(k3_supinf_2(X1,X2))|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_448, plain, (v1_xcmplx_0(k3_supinf_2(X1,X2))|~v1_membered(k7_numbers)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_2298, plain, (k1_extreal2(k1_extreal2(k1_extreal2(np__0,k12_supinf_2(X1,X2)),np__0),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X5,X1),X2),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X3)|v1_xboole_0(X5)|~r2_hidden(X2,k1_relset_1(X3,k1_mesfunc2(X4,X1)))|~r2_hidden(X2,k1_relset_1(X4,X1))|~v1_funct_1(X1)|~m1_subset_1(k1_mesfunc2(X4,X1),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(k12_supinf_2(X1,X2),k7_numbers)|~m1_subset_1(X2,X3)|~m1_subset_1(X2,X4)|~r2_hidden(X2,k1_relset_1(X5,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X2,X5))).#trainneg
cnf(i_0_2299, plain, (k1_extreal2(k1_extreal2(k1_extreal2(k12_supinf_2(X1,X2),np__0),np__0),np__0)=k1_extreal2(k1_extreal2(k1_extreal2(np__0,k12_supinf_2(X1,X2)),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X3)|~r2_hidden(X2,k1_relset_1(X3,k1_mesfunc2(X4,X1)))|~r2_hidden(X2,k1_relset_1(X4,X1))|~v1_funct_1(X1)|~m1_subset_1(k1_mesfunc2(X4,X1),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(k12_supinf_2(X1,X2),k7_numbers)|~m1_subset_1(X2,X3)|~m1_subset_1(X2,X4))).#trainneg
cnf(i_0_2030, negated_conjecture, (r1_xxreal_0(k3_extreal1(k6_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,X2))),k6_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,k10_mesfunc1(esk1_0,X2))))|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2)))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2)))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2)),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2)),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2))))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k2_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2)))|~v1_funct_1(k1_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2)))|~v1_funct_1(k10_mesfunc1(esk1_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2))),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2))),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k10_mesfunc1(esk1_0,k10_mesfunc1(esk1_0,X2)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k10_mesfunc1(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~r1_tarski(k1_relset_1(esk1_0,esk4_0),k1_relset_1(esk1_0,X2))|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_1986, negated_conjecture, (m1_subset_1(X1,esk1_0)|~r2_hidden(X1,k1_relset_1(esk1_0,esk4_0)))).#trainneg
cnf(i_0_1863, plain, (v1_relat_1(esk15_1(X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_882, negated_conjecture, (m1_subset_1(k1_relset_1(esk2_0,esk3_0),k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_1371, negated_conjecture, (v1_xreal_0(X1)|v1_xboole_0(esk3_0)|~v3_membered(k2_zfmisc_1(esk2_0,k7_numbers))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_1372, negated_conjecture, (v7_ordinal1(X1)|v1_xboole_0(esk3_0)|~v6_membered(k2_zfmisc_1(esk2_0,k7_numbers))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_450, plain, (v7_ordinal1(k3_supinf_2(X1,X2))|~v6_membered(k7_numbers)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_452, plain, (v1_xxreal_0(k12_supinf_2(X1,X2))|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2163, negated_conjecture, (r1_xxreal_0(k3_extreal1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0))),k3_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))),k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0)))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_2162, negated_conjecture, (r1_xxreal_0(k3_extreal1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0))),k3_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))),k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0)))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_2029, negated_conjecture, (r1_xxreal_0(k6_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,X2)),k6_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,k10_mesfunc1(esk1_0,X2))))|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2)))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2)))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2)),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2)),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,X2)))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k2_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2)))|~v1_funct_1(k1_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2)))|~v1_funct_1(k10_mesfunc1(esk1_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2))),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2))),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,k10_mesfunc1(esk1_0,X2)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k10_mesfunc1(esk1_0,k10_mesfunc1(esk1_0,X2)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,X2)),k7_numbers)|~m1_subset_1(k10_mesfunc1(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~r1_tarski(k1_relset_1(esk1_0,esk4_0),k1_relset_1(esk1_0,X2))|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_2261, plain, (k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X1,X2),X3),np__0),np__0)=k1_extreal2(k1_extreal2(k1_extreal2(k12_supinf_2(X2,X3),np__0),np__0),np__0)|v1_xboole_0(X5)|v1_xboole_0(X1)|v1_xboole_0(X4)|~r2_hidden(X3,k1_relset_1(X4,k1_mesfunc2(X5,X2)))|~r2_hidden(X3,k1_relset_1(X5,X2))|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k1_mesfunc2(X5,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X4)|~m1_subset_1(X3,X5)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_1947, negated_conjecture, (r1_xxreal_0(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)),k6_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,k2_mesfunc2(esk1_0,X2))))|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,X2)))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,X2)))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,X2)),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,X2)),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,X2))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,X2))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k2_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,X2)))|~v1_funct_1(k1_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,X2)))|~v1_funct_1(k2_mesfunc2(esk1_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,X2))),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,X2))),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,X2)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,X2)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k10_mesfunc1(esk1_0,k2_mesfunc2(esk1_0,X2)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~r1_tarski(k1_relset_1(esk1_0,esk4_0),k1_relset_1(esk1_0,X2))|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_2245, plain, (k1_extreal2(k1_extreal2(np__0,k12_supinf_2(k1_mesfunc2(X1,X2),X3)),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X5,X2),X3),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X1)|v1_xboole_0(X5)|~r2_hidden(X3,k1_relset_1(X4,k1_mesfunc2(X1,X2)))|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(k12_supinf_2(k1_mesfunc2(X1,X2),X3),k7_numbers)|~m1_subset_1(X3,X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X5,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X5))).#trainneg
cnf(i_0_1532, plain, (k7_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))=k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X4))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k1_relset_1(X1,X4),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1574, negated_conjecture, (r1_xxreal_0(np__0,k12_supinf_2(esk3_0,X1))|r1_xxreal_0(k12_supinf_2(esk3_0,X1),np__0))).#trainneg
cnf(i_0_1142, negated_conjecture, (k3_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0))=k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X2,esk4_0))|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(esk4_0)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_1948, negated_conjecture, (r1_xxreal_0(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,esk4_0)),k6_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))))|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,esk4_0)))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,esk4_0)))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,esk4_0)),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,esk4_0)),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,esk4_0))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,esk4_0))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,esk4_0)))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k2_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,esk4_0)))|~v1_funct_1(k1_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,esk4_0)))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,esk4_0))),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,esk4_0))),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,esk4_0)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,k1_mesfunc2(esk1_0,esk4_0)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k10_mesfunc1(esk1_0,k1_mesfunc2(esk1_0,esk4_0)),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_1049, plain, (esk11_1(k1_zfmisc_1(X1))=esk12_1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_453, plain, (v1_xcmplx_0(k12_supinf_2(X1,X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_valued_0(X1))).#trainneg
cnf(i_0_1577, plain, (r1_xxreal_0(np__0,X1)|r1_xxreal_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2370, negated_conjecture, (v3_membered(k1_relset_1(esk2_0,esk3_0))|~v3_membered(esk2_0))).#trainneg
cnf(i_0_2079, negated_conjecture, (v1_xxreal_0(X1)|v1_xboole_0(esk4_0)|~m1_subset_1(X1,esk4_0)|~v3_membered(k2_zfmisc_1(esk1_0,k7_numbers)))).#trainneg
cnf(i_0_2203, negated_conjecture, (v1_xcmplx_0(X1)|v1_xboole_0(esk4_0)|~m1_subset_1(X1,esk4_0)|~v3_membered(k2_zfmisc_1(esk1_0,k7_numbers)))).#trainneg
cnf(i_0_454, plain, (k4_xxreal_0(X2,X1)=k1_extreal2(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_2414, negated_conjecture, (r1_xxreal_0(k4_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0)),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X2,esk4_0)))|~r1_xxreal_0(np__0,k4_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X2))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(esk4_0)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_2423, negated_conjecture, (r1_xxreal_0(k3_xxreal_3(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k6_mesfunc5(esk1_0,esk2_0,X1,X2)),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X2,esk4_0)))|~r1_xxreal_0(np__0,k3_xxreal_3(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k6_mesfunc5(esk1_0,esk2_0,X1,X2)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X2))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k7_numbers)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(esk4_0)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_2415, negated_conjecture, (r1_xxreal_0(k3_xxreal_3(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0)),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X2,esk4_0)))|~r1_xxreal_0(np__0,k3_xxreal_3(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X2))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(esk4_0)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_2422, negated_conjecture, (r1_xxreal_0(k3_extreal1(k3_xxreal_3(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k6_mesfunc5(esk1_0,esk2_0,X1,X2))),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X2,esk4_0)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X2))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k7_numbers)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(esk4_0)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_2413, negated_conjecture, (r1_xxreal_0(k3_extreal1(k3_xxreal_3(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0))),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X2,esk4_0)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X2))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(esk4_0)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_2412, negated_conjecture, (r1_xxreal_0(k3_extreal1(k4_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0))),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X2,esk4_0)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X2))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(esk4_0)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_2421, negated_conjecture, (r1_xxreal_0(k3_extreal1(k4_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k6_mesfunc5(esk1_0,esk2_0,X1,X2))),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X2,esk4_0)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X2))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k7_numbers)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(esk4_0)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_2417, negated_conjecture, (k3_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k6_mesfunc5(esk1_0,esk2_0,X1,X2))=k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X2,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(esk4_0)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_2392, negated_conjecture, (r1_xxreal_0(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0)),k3_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))),k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0)))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0)))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers))).#trainneg
cnf(i_0_2442, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk11_1(k1_zfmisc_1(X4)))|~v1_xboole_0(X4))).#trainneg
cnf(i_0_2330, plain, (r1_xxreal_0(k1_supinf_1,esk11_1(k1_numbers))|r1_xxreal_0(esk11_1(k1_numbers),k1_supinf_1))).#trainneg
cnf(i_0_2321, plain, (X1=k1_relset_1(X2,esk12_1(X3))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_456, plain, (v3_finsub_1(esk14_1(X1))|~v2_finsub_1(esk14_1(X1)))).#trainneg
cnf(i_0_2364, negated_conjecture, (v1_xxreal_0(k1_relset_1(esk2_0,esk3_0))|~v2_membered(k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_2365, negated_conjecture, (v1_xcmplx_0(k1_relset_1(esk2_0,esk3_0))|~v1_membered(k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_2336, negated_conjecture, (v1_xxreal_0(X1)|v1_xboole_0(esk3_0)|~m1_subset_1(X1,esk3_0)|~v3_membered(k2_zfmisc_1(esk2_0,k7_numbers)))).#trainneg
cnf(i_0_2337, negated_conjecture, (v1_xcmplx_0(X1)|v1_xboole_0(esk3_0)|~m1_subset_1(X1,esk3_0)|~v3_membered(k2_zfmisc_1(esk2_0,k7_numbers)))).#trainneg
cnf(i_0_458, plain, (v1_xxreal_0(k1_extreal2(X1,X2))|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_459, plain, (v1_xcmplx_0(k1_extreal2(X1,X2))|~v1_membered(k7_numbers)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_2394, negated_conjecture, (r1_xxreal_0(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0)),k3_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))),k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0)))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,k2_mesfunc2(esk1_0,esk4_0))),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0)))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers))).#trainneg
cnf(i_0_960, negated_conjecture, (k10_mesfunc1(X1,X2)=esk4_0|r2_hidden(esk6_3(X1,X2,esk4_0),k1_relset_1(esk1_0,esk4_0))|v1_xboole_0(X1)|k9_xtuple_0(X2)!=k1_relset_1(esk1_0,esk4_0)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_2316, plain, (k1_extreal2(k1_extreal2(k1_extreal2(np__0,k12_supinf_2(X1,X2)),np__0),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X3,X1),X2),np__0),np__0)|v1_xboole_0(X3)|~r2_hidden(X2,k1_relset_1(X3,X1))|~v1_funct_1(X1)|~m1_subset_1(k1_mesfunc2(X3,X1),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(k12_supinf_2(X1,X2),k7_numbers)|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_2441, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,esk11_1(k1_zfmisc_1(X4))))|~v1_xboole_0(X4))).#trainneg
cnf(i_0_2367, negated_conjecture, (v7_ordinal1(k1_relset_1(esk2_0,esk3_0))|~v6_membered(k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_2213, plain, (k9_xtuple_0(esk15_1(X4))=k9_xtuple_0(X2)|v1_xboole_0(X3)|esk15_1(X4)!=k10_mesfunc1(X3,X2)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X3)|~v1_relat_1(X2)|~v1_funct_1(esk15_1(X4))|~v1_funct_1(X2)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_2366, negated_conjecture, (v1_xreal_0(k1_relset_1(esk2_0,esk3_0))|~v3_membered(k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_2525, negated_conjecture, (v1_xxreal_0(k1_relset_1(esk2_0,esk3_0))|~v3_membered(k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_534, plain, (v1_valued_0(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v1_membered(k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_461, plain, (v7_ordinal1(k1_extreal2(X1,X2))|~v6_membered(k7_numbers)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_465, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2499, negated_conjecture, (r1_xxreal_0(k4_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k6_mesfunc5(esk1_0,esk2_0,X1,X2)),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X2,esk4_0)))|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(esk4_0)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|~r1_xxreal_0(np__0,k4_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k6_mesfunc5(esk1_0,esk2_0,X1,X2)))|~m1_subset_1(k4_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k6_mesfunc5(esk1_0,esk2_0,X1,X2)),k7_numbers))).#trainneg
cnf(i_0_2527, plain, (r1_xxreal_0(np__0,k1_extreal2(X1,X2))|r1_xxreal_0(k1_extreal2(X1,X2),np__0)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_2285, plain, (k1_extreal2(k1_extreal2(np__0,k12_supinf_2(k1_mesfunc2(X1,X2),X3)),np__0)=k1_extreal2(k1_extreal2(k1_extreal2(np__0,k12_supinf_2(X2,X3)),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X1)|~r2_hidden(X3,k1_relset_1(X4,k1_mesfunc2(X1,X2)))|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(k12_supinf_2(k1_mesfunc2(X1,X2),X3),k7_numbers)|~m1_subset_1(X3,X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k12_supinf_2(X2,X3),k7_numbers)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_2374, negated_conjecture, (m1_subset_1(X1,esk2_0)|~r2_hidden(X1,k1_relset_1(esk2_0,esk3_0)))).#trainneg
cnf(i_0_2526, negated_conjecture, (v1_xcmplx_0(k1_relset_1(esk2_0,esk3_0))|~v3_membered(k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_2426, plain, (k9_xtuple_0(esk11_1(k1_zfmisc_1(X4)))=k9_xtuple_0(X2)|v1_xboole_0(X3)|esk11_1(k1_zfmisc_1(X4))!=k10_mesfunc1(X3,X2)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X3)|~v1_relat_1(X2)|~v1_funct_1(esk11_1(k1_zfmisc_1(X4)))|~v1_funct_1(X2)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_466, plain, (v1_xreal_0(k12_supinf_2(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v2_valued_0(X1))).#trainneg
cnf(i_0_2474, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2371, negated_conjecture, (v1_valued_0(k1_relset_1(esk2_0,esk3_0))|~v1_valued_0(esk2_0)|~v1_relat_1(esk2_0))).#trainneg
cnf(i_0_2372, negated_conjecture, (v3_valued_0(k1_relset_1(esk2_0,esk3_0))|~v3_valued_0(esk2_0)|~v1_relat_1(esk2_0))).#trainneg
cnf(i_0_536, plain, (v5_relat_1(k10_mesfunc1(X1,X2),k7_numbers)|v1_xboole_0(X1)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_467, plain, (m1_subset_1(k12_supinf_2(X1,X2),k7_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X1,k7_numbers))).#trainneg
cnf(i_0_2270, plain, (k1_extreal2(k12_supinf_2(k1_mesfunc2(X6,k1_mesfunc2(X1,X2)),X3),np__0)=k1_extreal2(k12_supinf_2(k1_mesfunc2(X4,k1_mesfunc2(X5,X2)),X3),np__0)|v1_xboole_0(X5)|v1_xboole_0(X1)|v1_xboole_0(X4)|v1_xboole_0(X6)|~r2_hidden(X3,k1_relset_1(X4,k1_mesfunc2(X5,X2)))|~r2_hidden(X3,k1_relset_1(X5,X2))|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k1_mesfunc2(X5,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X4)|~m1_subset_1(X3,X5)|~m1_subset_1(X3,X1)|~r2_hidden(X3,k1_relset_1(X6,k1_mesfunc2(X1,X2)))|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X6,k7_numbers)))|~m1_subset_1(X3,X6))).#trainneg
cnf(i_0_1051, plain, (k12_supinf_2(k1_mesfunc2(X1,X2),X3)=k1_extreal2(k2_supinf_2(k12_supinf_2(X4,X3)),np__0)|v1_xboole_0(X1)|k1_mesfunc2(X1,X2)!=k2_mesfunc2(X1,X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X4)|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1359, plain, (k12_supinf_2(k2_mesfunc2(X1,X2),X3)=k1_extreal2(k12_supinf_2(X4,X3),np__0)|v1_xboole_0(X1)|k2_mesfunc2(X1,X2)!=k1_mesfunc2(X1,X4)|~r2_hidden(X3,k1_relset_1(X1,k2_mesfunc2(X1,X2)))|~v1_relat_1(k2_mesfunc2(X1,X2))|~v1_funct_1(X4)|~v1_funct_1(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_2681, plain, (r1_xxreal_0(np__0,esk11_1(k1_numbers))|r1_xxreal_0(esk11_1(k1_numbers),np__0))).#trainneg
cnf(i_0_2615, negated_conjecture, (v1_xcmplx_0(k1_relset_1(esk2_0,esk3_0))|~v6_membered(k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_2361, negated_conjecture, (m1_subset_1(esk7_2(esk1_0,esk4_0),esk1_0)|v6_supinf_2(esk4_0))).#trainneg
cnf(i_0_2201, plain, (r1_xxreal_0(X1,k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X2))))))|r1_xxreal_0(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X2))))),X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k7_numbers))).#trainneg
cnf(i_0_474, plain, (k3_supinf_2(X1,k2_xxreal_3(X2))=k3_xxreal_3(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(k2_xxreal_3(X2),k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_475, plain, (k1_xxreal_3(X2,X1)=k3_supinf_2(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_2502, negated_conjecture, (r1_xxreal_0(k4_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k6_mesfunc5(esk1_0,esk2_0,X1,X2)),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X2,esk4_0)))|~r1_xxreal_0(np__0,k4_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k6_mesfunc5(esk1_0,esk2_0,X1,X2)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,X2))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k7_numbers)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(esk4_0)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_2379, plain, (r1_xxreal_0(np__0,k12_supinf_2(X1,X2))|r1_xxreal_0(k12_supinf_2(X1,X2),np__0)|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1080, plain, (r1_xxreal_0(k4_supinf_2(X1,X2),k3_supinf_2(X2,k3_extreal1(X1)))|~r1_xxreal_0(np__0,X2)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)|~r1_xxreal_0(np__0,k4_supinf_2(X1,X2))|~m1_subset_1(k4_supinf_2(X1,X2),k7_numbers))).#trainneg
cnf(i_0_1525, negated_conjecture, (k3_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,X3),k6_mesfunc5(esk1_0,esk2_0,X1,X2))=k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,X2,X3))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X3),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,X2),k7_numbers)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X3)|k1_relset_1(esk1_0,esk4_0)!=k1_relset_1(esk1_0,X2)|~r1_mesfunc1(esk1_0,esk2_0,X3,k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X3)|~v6_supinf_2(X2)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_2586, plain, (k1_extreal2(k1_extreal2(k1_extreal2(np__0,k12_supinf_2(k1_mesfunc2(X1,X2),X3)),np__0),np__0)=k1_extreal2(k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X5,X2),X3),np__0),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X5)|v1_xboole_0(X1)|~r2_hidden(X3,k1_relset_1(X4,k1_mesfunc2(X1,X2)))|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(k12_supinf_2(k1_mesfunc2(X1,X2),X3),k7_numbers)|~m1_subset_1(X3,X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X5,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X5))).#trainneg
cnf(i_0_886, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_243, negated_conjecture, (~r2_hidden(X1,esk4_0)|~v1_xboole_0(k2_zfmisc_1(esk1_0,k7_numbers)))).#trainneg
cnf(i_0_527, plain, (v1_xboole_0(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,k7_numbers))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_528, plain, (v3_membered(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v3_membered(k2_zfmisc_1(X1,k7_numbers))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_477, plain, (v1_xxreal_0(k4_supinf_2(X1,X2))|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_478, plain, (v1_xcmplx_0(k4_supinf_2(X1,X2))|~v1_membered(k7_numbers)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_2585, plain, (k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X4,X1),X2),np__0),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X3,X1),X2),np__0),np__0)|v1_xboole_0(X3)|v1_xboole_0(X4)|~r2_hidden(X2,k1_relset_1(X3,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(k12_supinf_2(X1,X2),k7_numbers)|~m1_subset_1(X2,X3)|~r2_hidden(X2,k1_relset_1(X4,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,X4))).#trainneg
cnf(i_0_1320, plain, (v1_xcmplx_0(esk11_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_2766, negated_conjecture, (v1_xxreal_0(esk7_2(esk1_0,esk4_0))|v6_supinf_2(esk4_0)|~v2_membered(esk1_0))).#trainneg
cnf(i_0_2767, negated_conjecture, (v1_xcmplx_0(esk7_2(esk1_0,esk4_0))|v6_supinf_2(esk4_0)|~v1_membered(esk1_0))).#trainneg
cnf(i_0_2768, negated_conjecture, (v1_xreal_0(esk7_2(esk1_0,esk4_0))|v6_supinf_2(esk4_0)|~v3_membered(esk1_0))).#trainneg
cnf(i_0_2788, plain, (r1_xxreal_0(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X1)))))),X2)|r1_xxreal_0(X2,k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X1)))))))|~v1_xxreal_0(X2)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_2864, plain, (X1=k10_mesfunc1(X2,X3)|v1_xboole_0(X2)|~v1_xboole_0(X1)|~v5_relat_1(X3,k7_numbers)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v1_xboole_0(k2_zfmisc_1(X2,k7_numbers)))).#trainneg
cnf(i_0_480, plain, (v7_ordinal1(k4_supinf_2(X1,X2))|~v6_membered(k7_numbers)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_482, plain, (m1_subset_1(k3_xxreal_3(X1,X2),k7_numbers)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_2831, plain, (r1_xxreal_0(k7_mesfunc5(X1,X2,X3,X4),k3_supinf_2(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k3_extreal1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)))))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_xxreal_0(np__0,k7_mesfunc5(X1,X2,X3,X4))|~r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)))|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_1279, plain, (k3_supinf_2(k6_mesfunc5(X1,X2,X3,X4),k6_mesfunc5(X1,X2,X3,X5))=k6_mesfunc5(X1,X2,X3,k3_mesfunc1(X1,X4,X5))|v1_xboole_0(X2)|v1_xboole_0(X1)|k1_relset_1(X1,k2_mesfunc2(X1,X6))!=k1_relset_1(X1,X5)|k1_relset_1(X1,k2_mesfunc2(X1,X6))!=k1_relset_1(X1,X4)|~v1_relat_1(k2_mesfunc2(X1,X6))|~r1_mesfunc1(X1,X2,X5,k1_relset_1(X1,k2_mesfunc2(X1,X6)))|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,k2_mesfunc2(X1,X6)))|~r1_mesfunc5(X1,X2,X7,k2_mesfunc2(X1,X6))|~v4_measure1(X3,X1,X2)|~v4_measure1(X7,X1,X2)|~v6_supinf_2(X5)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X7)|~v10_valued_0(X3)|~v10_valued_0(X7)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X7,X2,k7_numbers)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_2249, plain, (k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X1,X2),X3),np__0),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X6,X2),X3),np__0),np__0)|v1_xboole_0(X5)|v1_xboole_0(X1)|v1_xboole_0(X4)|v1_xboole_0(X6)|~r2_hidden(X3,k1_relset_1(X4,k1_mesfunc2(X5,X2)))|~r2_hidden(X3,k1_relset_1(X5,X2))|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k1_mesfunc2(X5,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X4)|~m1_subset_1(X3,X5)|~m1_subset_1(X3,X1)|~r2_hidden(X3,k1_relset_1(X6,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X6,k7_numbers)))|~m1_subset_1(X3,X6))).#trainneg
cnf(i_0_1177, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)))|v1_xboole_0(X1)|v1_xboole_0(X2)|X5!=k1_relset_1(X1,k1_mesfunc2(X1,X4))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(k1_mesfunc2(X1,X4))|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(k1_mesfunc2(X1,X4))|~v1_funct_1(X3)|~m1_subset_1(k1_mesfunc2(X1,X4),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X5,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_mesfunc1(X1,X2,X4,X5)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1433, plain, (v1_xcmplx_0(esk11_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_2769, negated_conjecture, (v7_ordinal1(esk7_2(esk1_0,esk4_0))|v6_supinf_2(esk4_0)|~v6_membered(esk1_0))).#trainneg
cnf(i_0_601, plain, (k9_xtuple_0(esk15_1(k2_zfmisc_1(X1,k7_numbers)))=k9_xtuple_0(X2)|v1_xboole_0(X1)|esk15_1(k2_zfmisc_1(X1,k7_numbers))!=k10_mesfunc1(X1,X2)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_funct_1(esk15_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_483, plain, (r1_xxreal_0(k3_extreal1(k4_supinf_2(X1,k3_extreal1(X2))),k3_supinf_2(k3_extreal1(X1),k3_extreal1(X2)))|~m1_subset_1(k3_extreal1(X2),k7_numbers)|~m1_subset_1(X1,k7_numbers)|~m1_subset_1(X2,k7_numbers))).#trainneg
cnf(i_0_1266, plain, (v1_xcmplx_0(X1)|~v1_membered(k7_numbers)|~m1_subset_1(k2_xxreal_3(X1),k7_numbers)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_244, negated_conjecture, (~r2_hidden(X1,esk3_0)|~v1_xboole_0(k2_zfmisc_1(esk2_0,k7_numbers)))).#trainneg
cnf(i_0_524, plain, (v1_xreal_0(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_525, plain, (v7_ordinal1(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_485, plain, (r1_xxreal_0(k3_extreal1(k4_supinf_2(k3_extreal1(X1),X2)),k3_supinf_2(k3_extreal1(X1),k3_extreal1(X2)))|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(k3_extreal1(X1),k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_500, plain, (v1_xxreal_0(k1_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_3038, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)))|v1_xboole_0(X2)|v1_xboole_0(X1)|X5!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,X5)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X5,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3028, plain, (k3_supinf_2(k6_mesfunc5(X1,X2,X3,X4),k6_mesfunc5(X1,X2,X3,X5))=k6_mesfunc5(X1,X2,X3,k3_mesfunc1(X1,X4,X5))|v1_xboole_0(X1)|v1_xboole_0(X2)|k1_relset_1(X1,X6)!=k1_relset_1(X1,X5)|k1_relset_1(X1,X6)!=k1_relset_1(X1,X4)|~v1_relat_1(k2_mesfunc2(X1,X6))|~r1_mesfunc1(X1,X2,X5,k1_relset_1(X1,X6))|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X6))|~r1_mesfunc5(X1,X2,X7,k2_mesfunc2(X1,X6))|~v4_measure1(X3,X1,X2)|~v4_measure1(X7,X1,X2)|~v6_supinf_2(X5)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X7)|~v10_valued_0(X3)|~v10_valued_0(X7)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X7,X2,k7_numbers)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_2994, plain, (r1_xxreal_0(k7_mesfunc5(X1,X2,X3,X4),k3_supinf_2(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))))|v1_xboole_0(X2)|v1_xboole_0(X1)|~r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)))|~r1_xxreal_0(np__0,k7_mesfunc5(X1,X2,X3,X4))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))))).#trainneg
cnf(i_0_3023, plain, (k7_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4))=k6_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|~r1_tarski(k9_xtuple_0(k10_mesfunc1(X1,X4)),k1_relset_1(X1,X4))|~r1_mesfunc1(X1,X2,X4,k9_xtuple_0(k10_mesfunc1(X1,X4)))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(k10_mesfunc1(X1,X4))|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(k10_mesfunc1(X1,X4),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k9_xtuple_0(k10_mesfunc1(X1,X4)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v5_relat_1(X4,k7_numbers)|~v4_relat_1(X4,X1)|~v1_relat_1(X4))).#trainneg
cnf(i_0_3003, plain, (k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X1)|v1_xboole_0(X2)|k9_xtuple_0(k10_mesfunc1(X1,X5))!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k9_xtuple_0(k10_mesfunc1(X1,X5)))|~r1_mesfunc5(X1,X2,X6,k10_mesfunc1(X1,X5))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(k10_mesfunc1(X1,X5))|~v1_funct_1(X6)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k10_mesfunc1(X1,X5),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v5_relat_1(X5,k7_numbers)|~v4_relat_1(X5,X1)|~v1_relat_1(X5)|~v1_funct_1(X5))).#trainneg
cnf(i_0_1544, plain, (v2_valued_0(esk12_1(X1))|~v2_membered(X2))).#trainneg
cnf(i_0_1864, plain, (v5_relat_1(esk15_1(X3),X2)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_2948, negated_conjecture, (v1_xcmplx_0(esk7_2(esk1_0,esk4_0))|v6_supinf_2(esk4_0)|~v3_membered(esk1_0))).#trainneg
cnf(i_0_501, plain, (v1_xcmplx_0(k1_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1450, plain, (v7_ordinal1(X1)|~v6_membered(k7_numbers)|~m1_subset_1(k2_xxreal_3(X1),k7_numbers)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_3040, negated_conjecture, (v1_xcmplx_0(esk7_2(esk1_0,esk4_0))|v6_supinf_2(esk4_0)|~v6_membered(esk1_0))).#trainneg
cnf(i_0_2942, negated_conjecture, (r1_xxreal_0(np__0,esk7_2(esk1_0,esk4_0))|r1_xxreal_0(esk7_2(esk1_0,esk4_0),np__0)|v6_supinf_2(esk4_0)|~v2_membered(esk1_0))).#trainneg
cnf(i_0_2967, plain, (r1_xxreal_0(X1,k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X2))))))))|r1_xxreal_0(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X2))))))),X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k7_numbers))).#trainneg
cnf(i_0_529, plain, (v1_valued_0(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v1_valued_0(k2_zfmisc_1(X1,k7_numbers))|~v1_relat_1(k2_zfmisc_1(X1,k7_numbers))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_502, plain, (v1_xreal_0(k1_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_503, plain, (v7_ordinal1(k1_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_3022, plain, (k12_supinf_2(k10_mesfunc1(X1,X2),X3)=k1_extreal2(k12_supinf_2(X4,X3),np__0)|v1_xboole_0(X1)|k10_mesfunc1(X1,X2)!=k1_mesfunc2(X1,X4)|~r2_hidden(X3,k9_xtuple_0(k10_mesfunc1(X1,X2)))|~v1_funct_1(X4)|~v1_funct_1(k10_mesfunc1(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1050, plain, (k12_supinf_2(k2_mesfunc2(X1,X2),X3)=k1_extreal2(k2_supinf_2(k12_supinf_2(X4,X3)),np__0)|v1_xboole_0(X1)|k2_mesfunc2(X1,X2)!=k2_mesfunc2(X1,X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X4)|~v1_funct_1(k2_mesfunc2(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_2338, plain, (r1_xxreal_0(np__0,k3_supinf_2(X1,X2))|r1_xxreal_0(k3_supinf_2(X1,X2),np__0)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_1545, plain, (v1_valued_0(esk12_1(X1))|~v1_membered(X2))).#trainneg
cnf(i_0_1865, plain, (v4_relat_1(esk15_1(X3),X2)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_2943, negated_conjecture, (r1_xxreal_0(k1_supinf_1,esk7_2(esk1_0,esk4_0))|r1_xxreal_0(esk7_2(esk1_0,esk4_0),k1_supinf_1)|v6_supinf_2(esk4_0)|~v2_membered(esk1_0))).#trainneg
cnf(i_0_505, plain, (v1_xboole_0(k1_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,k7_numbers))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1546, plain, (v3_valued_0(esk12_1(X1))|~v3_membered(X2))).#trainneg
cnf(i_0_2946, negated_conjecture, (r1_xxreal_0(esk7_2(esk1_0,esk4_0),np__0)|r1_xxreal_0(np__0,esk7_2(esk1_0,esk4_0))|v6_supinf_2(esk4_0)|~v3_membered(esk1_0))).#trainneg
cnf(i_0_2947, negated_conjecture, (r1_xxreal_0(esk7_2(esk1_0,esk4_0),k1_supinf_1)|r1_xxreal_0(k1_supinf_1,esk7_2(esk1_0,esk4_0))|v6_supinf_2(esk4_0)|~v3_membered(esk1_0))).#trainneg
cnf(i_0_530, plain, (v3_valued_0(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v3_valued_0(k2_zfmisc_1(X1,k7_numbers))|~v1_relat_1(k2_zfmisc_1(X1,k7_numbers))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_522, plain, (v1_xxreal_0(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_506, plain, (v3_membered(k1_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v3_membered(k2_zfmisc_1(X1,k7_numbers))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_507, plain, (v1_valued_0(k1_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_valued_0(k2_zfmisc_1(X1,k7_numbers))|~v1_relat_1(k2_zfmisc_1(X1,k7_numbers))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_939, plain, (r1_xxreal_0(k7_mesfunc5(X1,X2,X3,X4),k3_supinf_2(k3_extreal1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))),k3_extreal1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)))))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_xxreal_0(np__0,k7_mesfunc5(X1,X2,X3,X4))|~m1_subset_1(k7_mesfunc5(X1,X2,X3,X4),k7_numbers))).#trainneg
cnf(i_0_2231, plain, (k7_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,k2_mesfunc2(X1,X4)))=k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,k2_mesfunc2(X1,X4)))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_tarski(k9_xtuple_0(k2_mesfunc2(X1,k2_mesfunc2(X1,X4))),k1_relset_1(X1,X4))|~v1_relat_1(k2_mesfunc2(X1,k2_mesfunc2(X1,X4)))|~r1_mesfunc1(X1,X2,k2_mesfunc2(X1,X4),k9_xtuple_0(k2_mesfunc2(X1,k2_mesfunc2(X1,X4))))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(k2_mesfunc2(X1,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X1,X4),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k9_xtuple_0(k2_mesfunc2(X1,k2_mesfunc2(X1,X4))),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1824, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4)))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,k10_mesfunc1(X1,X4)))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(k10_mesfunc1(X1,X4))|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(k10_mesfunc1(X1,X4),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k1_relset_1(X1,k10_mesfunc1(X1,X4)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~m1_subset_1(k1_relset_1(X1,k10_mesfunc1(X1,X4)),k1_zfmisc_1(k1_relset_1(X1,X4))))).#trainneg
cnf(i_0_1553, plain, (v3_membered(esk12_1(X1))|~v3_membered(X2))).#trainneg
cnf(i_0_2425, plain, (v1_relat_1(esk11_1(k1_zfmisc_1(X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2944, negated_conjecture, (r1_xxreal_0(X1,esk7_2(esk1_0,esk4_0))|r1_xxreal_0(esk7_2(esk1_0,esk4_0),X1)|v6_supinf_2(esk4_0)|~v1_xxreal_0(X1)|~v2_membered(esk1_0))).#trainneg
cnf(i_0_1046, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_508, plain, (v3_valued_0(k1_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v3_valued_0(k2_zfmisc_1(X1,k7_numbers))|~v1_relat_1(k2_zfmisc_1(X1,k7_numbers))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_794, plain, (esk12_1(X1)=np__0)).#trainneg
cnf(i_0_1481, plain, (v1_relat_1(np__0))).#trainneg
cnf(i_0_3188, plain, (m1_subset_1(esk11_1(k4_ordinal1),k1_numbers))).#trainneg
cnf(i_0_3160, plain, (v1_xxreal_0(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_3126, plain, (r1_xxreal_0(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X1)))))))),X2)|r1_xxreal_0(X2,k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X1)))))))))|~v1_xxreal_0(X2)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_509, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,k7_numbers))|v1_xboole_0(X2)|~r2_hidden(X1,k1_mesfunc2(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).#trainneg
cnf(i_0_511, plain, (v2_valued_0(k1_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_3071, negated_conjecture, (r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,esk4_0)))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_3129, plain, (k12_supinf_2(k10_mesfunc1(X1,X2),X3)=k1_extreal2(k12_supinf_2(X4,X3),np__0)|v1_xboole_0(X1)|k10_mesfunc1(X1,X2)!=k1_mesfunc2(X1,X4)|~r2_hidden(X3,k1_relset_1(X1,k10_mesfunc1(X1,X2)))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X4)|~v1_funct_1(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_3088, negated_conjecture, (r1_xxreal_0(k7_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,k2_mesfunc2(esk1_0,X2),k1_mesfunc2(esk1_0,X2))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)))|~r1_xxreal_0(np__0,k7_mesfunc5(esk1_0,esk2_0,X1,X2))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~v6_supinf_2(k1_mesfunc2(esk1_0,X2))|~v6_supinf_2(k2_mesfunc2(esk1_0,X2))|~v1_funct_1(k1_mesfunc2(esk1_0,X2))|~v1_funct_1(k2_mesfunc2(esk1_0,X2))|~m1_subset_1(k1_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k2_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_1794, plain, (m1_subset_1(k1_relset_1(X1,np__0),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3087, negated_conjecture, (r1_xxreal_0(k7_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,k3_mesfunc1(esk1_0,k1_mesfunc2(esk1_0,X2),k2_mesfunc2(esk1_0,X2))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)))|~r1_xxreal_0(np__0,k7_mesfunc5(esk1_0,esk2_0,X1,X2))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~v6_supinf_2(k2_mesfunc2(esk1_0,X2))|~v6_supinf_2(k1_mesfunc2(esk1_0,X2))|~v1_funct_1(k2_mesfunc2(esk1_0,X2))|~v1_funct_1(k1_mesfunc2(esk1_0,X2))|~m1_subset_1(k2_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_3132, plain, (k12_supinf_2(k2_mesfunc2(X1,X2),X3)=k1_extreal2(k2_supinf_2(k12_supinf_2(X2,X3)),np__0)|v1_xboole_0(X1)|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_3078, plain, (k3_supinf_2(k6_mesfunc5(X1,X2,X3,X4),k6_mesfunc5(X1,X2,X3,X5))=k6_mesfunc5(X1,X2,X3,k3_mesfunc1(X1,X4,X5))|v1_xboole_0(X2)|v1_xboole_0(X1)|k1_relset_1(X1,X6)!=k1_relset_1(X1,X5)|k1_relset_1(X1,X6)!=k1_relset_1(X1,X4)|~v1_relat_1(k2_mesfunc2(X1,k2_mesfunc2(X1,X6)))|~r1_mesfunc1(X1,X2,X5,k1_relset_1(X1,X6))|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X6))|~r1_mesfunc5(X1,X2,X7,k2_mesfunc2(X1,k2_mesfunc2(X1,X6)))|~v4_measure1(X3,X1,X2)|~v4_measure1(X7,X1,X2)|~v6_supinf_2(X5)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X7)|~v10_valued_0(X3)|~v10_valued_0(X7)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X7,X2,k7_numbers)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X7)|~v1_funct_1(k2_mesfunc2(X1,X6))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X1,X6),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_3094, negated_conjecture, (r1_xxreal_0(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k3_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0)),k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,esk4_0))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,esk4_0)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|~v6_supinf_2(esk4_0))).#trainneg
cnf(i_0_1111, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,np__0))).#trainneg
cnf(i_0_1970, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,np__0)))).#trainneg
cnf(i_0_2949, negated_conjecture, (r1_xxreal_0(esk7_2(esk1_0,esk4_0),X1)|r1_xxreal_0(X1,esk7_2(esk1_0,esk4_0))|v6_supinf_2(esk4_0)|~v1_xxreal_0(X1)|~v3_membered(esk1_0))).#trainneg
cnf(i_0_2321, plain, (X1=k1_relset_1(X2,np__0)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_512, plain, (v1_valued_0(k1_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_membered(k7_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1553, plain, (v3_membered(np__0)|~v3_membered(X2))).#trainneg
cnf(i_0_3323, plain, (v3_membered(np__0))).#trainneg
cnf(i_0_3196, plain, (v1_xreal_0(esk11_1(k4_ordinal1)))).#trainneg
cnf(i_0_3222, plain, (r1_xxreal_0(X1,k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X2))))))))))|r1_xxreal_0(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X2))))))))),X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k7_numbers))).#trainneg
cnf(i_0_523, plain, (v1_xcmplx_0(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_514, plain, (v5_relat_1(k1_mesfunc2(X1,X2),k7_numbers)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_3302, plain, (k12_supinf_2(k2_mesfunc2(X1,k1_mesfunc2(X2,X3)),X4)=k1_extreal2(k2_supinf_2(k1_extreal2(k12_supinf_2(X3,X4),np__0)),np__0)|v1_xboole_0(X1)|v1_xboole_0(X2)|~r2_hidden(X4,k1_relset_1(X1,k1_mesfunc2(X2,X3)))|~v1_funct_1(k1_mesfunc2(X2,X3))|~m1_subset_1(k1_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,X1)|~r2_hidden(X4,k1_relset_1(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_3261, negated_conjecture, (r1_xxreal_0(k7_mesfunc5(esk1_0,esk2_0,X1,X2),k6_mesfunc5(esk1_0,esk2_0,X1,k10_mesfunc1(esk1_0,X2)))|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,X2),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)))|~r1_xxreal_0(np__0,k7_mesfunc5(esk1_0,esk2_0,X1,X2))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k2_mesfunc2(esk1_0,X2))|~v1_funct_1(k1_mesfunc2(esk1_0,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|~v5_relat_1(k2_mesfunc2(esk1_0,X2),k7_numbers)|~v5_relat_1(k1_mesfunc2(esk1_0,X2),k7_numbers)|~v4_relat_1(k2_mesfunc2(esk1_0,X2),esk1_0)|~v4_relat_1(k1_mesfunc2(esk1_0,X2),esk1_0)|~v1_relat_1(k2_mesfunc2(esk1_0,X2))|~v1_relat_1(k1_mesfunc2(esk1_0,X2))|~m1_subset_1(k10_mesfunc1(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_3231, plain, (k12_supinf_2(k10_mesfunc1(X1,X2),X3)=k1_extreal2(k12_supinf_2(X4,X3),np__0)|v1_xboole_0(X1)|v1_xboole_0(k1_relset_1(X1,k10_mesfunc1(X1,X2)))|k10_mesfunc1(X1,X2)!=k1_mesfunc2(X1,X4)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X4)|~v1_funct_1(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,k1_relset_1(X1,k10_mesfunc1(X1,X2))))).#trainneg
cnf(i_0_3305, negated_conjecture, (r1_xxreal_0(k6_mesfunc5(esk1_0,esk2_0,X1,esk4_0),k3_supinf_2(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,esk4_0)),k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0))))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0)))|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(esk4_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(X1)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers))))).#trainneg
cnf(i_0_3270, plain, (r1_xxreal_0(k12_supinf_2(k2_mesfunc2(X3,X1),X2),np__0)|r1_xxreal_0(np__0,k12_supinf_2(k2_mesfunc2(X3,X1),X2))|v1_xboole_0(X3)|~m1_subset_1(k2_supinf_2(k12_supinf_2(X1,X2)),k7_numbers)|~r2_hidden(X2,k1_relset_1(X3,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_3300, plain, (k12_supinf_2(k2_mesfunc2(X1,k1_mesfunc2(X2,X3)),X4)=k1_extreal2(k2_supinf_2(k1_extreal2(np__0,k12_supinf_2(X3,X4))),np__0)|v1_xboole_0(X1)|v1_xboole_0(X2)|~r2_hidden(X4,k1_relset_1(X1,k1_mesfunc2(X2,X3)))|~v1_funct_1(k1_mesfunc2(X2,X3))|~m1_subset_1(k1_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,X1)|~r2_hidden(X4,k1_relset_1(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k12_supinf_2(X3,X4),k7_numbers)|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_1549, plain, (v4_relat_1(np__0,X2))).#trainneg
cnf(i_0_3331, plain, (v1_xcmplx_0(esk11_1(k4_ordinal1)))).#trainneg
cnf(i_0_1543, plain, (k9_xtuple_0(np__0)=k9_xtuple_0(X2)|v1_xboole_0(X3)|k10_mesfunc1(X3,X2)!=np__0|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X3)|~v1_relat_1(X2)|~v1_funct_1(np__0)|~v1_funct_1(X2))).#trainneg
cnf(i_0_515, plain, (v1_relat_1(k1_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1546, plain, (v3_valued_0(np__0)|~v3_membered(X2))).#trainneg
cnf(i_0_3194, plain, (v1_xxreal_0(esk11_1(k4_ordinal1))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_3351, plain, (r1_xxreal_0(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X1)))))))))),X2)|r1_xxreal_0(X2,k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X1)))))))))))|~v1_xxreal_0(X2)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_3056, plain, (v1_xcmplx_0(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_516, plain, (v4_relat_1(k1_mesfunc2(X1,X2),X1)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_531, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,k7_numbers))|v1_xboole_0(X2)|~r2_hidden(X1,k10_mesfunc1(X2,X3))|~v5_relat_1(X3,k7_numbers)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_3296, plain, (k12_supinf_2(k1_mesfunc2(X1,k2_mesfunc2(X2,X3)),X4)=k1_extreal2(k1_extreal2(k2_supinf_2(k12_supinf_2(X3,X4)),np__0),np__0)|v1_xboole_0(X1)|v1_xboole_0(X2)|~r2_hidden(X4,k1_relset_1(X1,k2_mesfunc2(X2,X3)))|~v1_funct_1(k2_mesfunc2(X2,X3))|~m1_subset_1(k2_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,X1)|~r2_hidden(X4,k1_relset_1(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_3079, plain, (k3_supinf_2(k6_mesfunc5(X1,X2,X3,X4),k6_mesfunc5(X1,X2,X3,X5))=k6_mesfunc5(X1,X2,X3,k3_mesfunc1(X1,X4,X5))|v1_xboole_0(X2)|v1_xboole_0(X1)|k1_relset_1(X1,X6)!=k1_relset_1(X1,X5)|k1_relset_1(X1,X6)!=k1_relset_1(X1,X4)|~v1_relat_1(k2_mesfunc2(X1,k1_mesfunc2(X1,X6)))|~r1_mesfunc1(X1,X2,X5,k1_relset_1(X1,X6))|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X6))|~r1_mesfunc5(X1,X2,X7,k2_mesfunc2(X1,k1_mesfunc2(X1,X6)))|~v4_measure1(X3,X1,X2)|~v4_measure1(X7,X1,X2)|~v6_supinf_2(X5)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X7)|~v10_valued_0(X3)|~v10_valued_0(X7)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X7,X2,k7_numbers)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X7)|~v1_funct_1(k1_mesfunc2(X1,X6))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X1,X6),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_3096, plain, (k7_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,k2_mesfunc2(X1,X4)))=k6_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,k2_mesfunc2(X1,X4)))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_tarski(k9_xtuple_0(k10_mesfunc1(X1,k2_mesfunc2(X1,X4))),k1_relset_1(X1,X4))|~r1_mesfunc1(X1,X2,k2_mesfunc2(X1,X4),k9_xtuple_0(k10_mesfunc1(X1,k2_mesfunc2(X1,X4))))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(k2_mesfunc2(X1,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X1,X4),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k9_xtuple_0(k10_mesfunc1(X1,k2_mesfunc2(X1,X4))),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_3226, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~v6_supinf_2(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_3285, plain, (k1_extreal2(k1_extreal2(k1_extreal2(k2_supinf_2(k12_supinf_2(X2,X3)),np__0),np__0),np__0)=k1_extreal2(k12_supinf_2(k1_mesfunc2(X4,k2_mesfunc2(X1,X2)),X3),np__0)|v1_xboole_0(X4)|v1_xboole_0(X1)|~r2_hidden(X3,k1_relset_1(X4,k2_mesfunc2(X1,X2)))|~v1_funct_1(k2_mesfunc2(X1,X2))|~m1_subset_1(k2_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X3,X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_3502, plain, (k1_relset_1(X1,k1_mesfunc2(X1,X2))=k9_xtuple_0(k1_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_relat_1(k1_mesfunc2(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_3546, plain, (k1_extreal2(k1_extreal2(k2_supinf_2(k12_supinf_2(X3,X4)),np__0),np__0)=k1_extreal2(k12_supinf_2(k2_mesfunc2(X2,X3),X4),np__0)|v1_xboole_0(X1)|v1_xboole_0(X2)|~r2_hidden(X4,k1_relset_1(X1,k2_mesfunc2(X2,X3)))|~v1_funct_1(k2_mesfunc2(X2,X3))|~m1_subset_1(k2_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,X1)|~r2_hidden(X4,k1_relset_1(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_3561, plain, (k12_supinf_2(k1_mesfunc2(X1,k2_mesfunc2(X2,X3)),X4)=k1_extreal2(k12_supinf_2(k2_mesfunc2(X5,X3),X4),np__0)|v1_xboole_0(X2)|v1_xboole_0(X1)|v1_xboole_0(X5)|~r2_hidden(X4,k1_relset_1(X1,k2_mesfunc2(X2,X3)))|~r2_hidden(X4,k1_relset_1(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(k2_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,X1)|~m1_subset_1(X4,X2)|~r2_hidden(X4,k1_relset_1(X5,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X4,X5))).#trainneg
cnf(i_0_3596, plain, (k1_extreal2(k1_extreal2(k1_extreal2(k2_supinf_2(k12_supinf_2(X1,X2)),np__0),np__0),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X4,X1),X2),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X3)|~r2_hidden(X2,k1_relset_1(X3,k2_mesfunc2(X4,X1)))|~r2_hidden(X2,k1_relset_1(X4,X1))|~v1_funct_1(X1)|~m1_subset_1(k2_mesfunc2(X4,X1),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,X3)|~m1_subset_1(X2,X4)|~v1_funct_1(k2_mesfunc2(X4,X1)))).#trainneg
cnf(i_0_3564, negated_conjecture, (k1_relset_1(esk1_0,k10_mesfunc1(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~v6_supinf_2(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k10_mesfunc1(esk1_0,esk4_0))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_3472, plain, (v3_valued_0(np__0))).#trainneg
cnf(i_0_1547, plain, (v5_relat_1(np__0,X2))).#trainneg
cnf(i_0_3197, plain, (v7_ordinal1(esk11_1(k4_ordinal1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1545, plain, (v1_valued_0(np__0)|~v1_membered(X2))).#trainneg
cnf(i_0_543, plain, (m1_subset_1(esk7_2(X1,X2),k1_relset_1(X1,X2))|v6_supinf_2(X2)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1544, plain, (v2_valued_0(np__0)|~v2_membered(X2))).#trainneg
cnf(i_0_3198, plain, (r1_xxreal_0(k1_supinf_1,esk11_1(k4_ordinal1))|r1_xxreal_0(esk11_1(k4_ordinal1),k1_supinf_1))).#trainneg
cnf(i_0_3199, plain, (r1_xxreal_0(np__0,esk11_1(k4_ordinal1))|r1_xxreal_0(esk11_1(k4_ordinal1),np__0))).#trainneg
cnf(i_0_3500, plain, (r1_xxreal_0(X1,k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X2))))))))))))|r1_xxreal_0(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X2))))))))))),X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k7_numbers))).#trainneg
cnf(i_0_3751, plain, (r1_xxreal_0(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X1)))))))))))),X2)|r1_xxreal_0(X2,k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X1)))))))))))))|~v1_xxreal_0(X2)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_544, plain, (m1_subset_1(k12_supinf_2(esk10_2(X1,X2),X3),k7_numbers)|v1_xboole_0(X2)|~v1_relat_1(esk10_2(X1,X2))|~v1_funct_1(esk10_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_545, plain, (v1_xxreal_0(k2_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_3648, plain, (k9_xtuple_0(k1_mesfunc2(X1,X2))=k1_relset_1(X1,X2)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_3704, plain, (k1_extreal2(k12_supinf_2(k2_mesfunc2(X5,X3),X4),np__0)=k1_extreal2(k12_supinf_2(k2_mesfunc2(X2,X3),X4),np__0)|v1_xboole_0(X1)|v1_xboole_0(X5)|v1_xboole_0(X2)|~r2_hidden(X4,k1_relset_1(X1,k2_mesfunc2(X2,X3)))|~v1_funct_1(k2_mesfunc2(X2,X3))|~m1_subset_1(k2_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,X1)|~r2_hidden(X4,k1_relset_1(X2,X3))|~r2_hidden(X4,k1_relset_1(X5,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X4,X2)|~m1_subset_1(X4,X5))).#trainneg
cnf(i_0_3651, plain, (k1_extreal2(k1_extreal2(k2_supinf_2(k12_supinf_2(X1,X2)),np__0),np__0)=k1_extreal2(k12_supinf_2(k2_mesfunc2(X3,X1),X2),np__0)|v1_xboole_0(X3)|~r2_hidden(X2,k9_xtuple_0(k2_mesfunc2(X3,X1)))|~r2_hidden(X2,k1_relset_1(X3,X1))|~v1_funct_1(X1)|~m1_subset_1(k2_mesfunc2(X3,X1),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X2,X3)|~v1_relat_1(k2_mesfunc2(X3,X1)))).#trainneg
cnf(i_0_3645, plain, (k7_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))=k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|~r1_mesfunc1(X1,X2,X4,k9_xtuple_0(k1_mesfunc2(X1,X4)))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k9_xtuple_0(k1_mesfunc2(X1,X4)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3652, plain, (k1_extreal2(k1_extreal2(k2_supinf_2(k12_supinf_2(X1,X2)),np__0),np__0)=k1_extreal2(k12_supinf_2(k2_mesfunc2(X3,X1),X2),np__0)|v1_xboole_0(X3)|~r2_hidden(X2,k1_relset_1(X3,X1))|~v1_funct_1(X1)|~m1_subset_1(k2_mesfunc2(X3,X1),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_3695, plain, (k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X1,X2),X3),np__0),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X5,X2),X3),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X5)|v1_xboole_0(X1)|~r2_hidden(X3,k1_relset_1(X4,k2_mesfunc2(X1,X2)))|~v1_funct_1(k2_mesfunc2(X1,X2))|~m1_subset_1(k2_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X3,X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X5,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X5))).#trainneg
cnf(i_0_3794, plain, (k1_extreal2(k12_supinf_2(k2_mesfunc2(X1,X2),X3),np__0)=k1_extreal2(k12_supinf_2(k2_mesfunc2(X4,X2),X3),np__0)|v1_xboole_0(X4)|v1_xboole_0(X1)|~r2_hidden(X3,k9_xtuple_0(k2_mesfunc2(X4,X2)))|~r2_hidden(X3,k1_relset_1(X4,X2))|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k2_mesfunc2(X4,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X4)|~m1_subset_1(X3,X1)|~v1_relat_1(k2_mesfunc2(X4,X2)))).#trainneg
cnf(i_0_3867, plain, (k1_extreal2(k12_supinf_2(k2_mesfunc2(X4,X1),X2),np__0)=k1_extreal2(k12_supinf_2(k2_mesfunc2(X3,X1),X2),np__0)|v1_xboole_0(X3)|v1_xboole_0(X4)|~r2_hidden(X2,k1_relset_1(X3,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X2,X3)|~r2_hidden(X2,k1_relset_1(X4,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,X4))).#trainneg
cnf(i_0_3732, plain, (v2_valued_0(np__0))).#trainneg
cnf(i_0_1109, plain, (m1_subset_1(np__0,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1026, plain, (esk15_1(X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_546, plain, (v1_xcmplx_0(k2_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_3721, plain, (v1_valued_0(np__0)|~v3_membered(X1))).#trainneg
cnf(i_0_3770, plain, (r1_xxreal_0(X1,k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X2))))))))))))))|r1_xxreal_0(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X2))))))))))))),X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k7_numbers))).#trainneg
cnf(i_0_3059, plain, (v1_xcmplx_0(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_547, plain, (v1_xreal_0(k2_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_548, plain, (v7_ordinal1(k2_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_3616, plain, (k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X1)|v1_xboole_0(X2)|k9_xtuple_0(k1_mesfunc2(X1,X5))!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k9_xtuple_0(k1_mesfunc2(X1,X5)))|~r1_mesfunc5(X1,X2,X6,k1_mesfunc2(X1,X5))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(k1_mesfunc2(X1,X5))|~v1_funct_1(X6)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X1,X5),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_3643, plain, (k1_extreal2(k1_extreal2(k1_extreal2(np__0,k12_supinf_2(k1_mesfunc2(X1,X2),X3)),np__0),np__0)=k1_extreal2(k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X4,X2),X3),np__0),np__0),np__0)|v1_xboole_0(X1)|v1_xboole_0(X4)|~r2_hidden(X3,k9_xtuple_0(k1_mesfunc2(X1,X2)))|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X4,X2))|~v1_funct_1(X2)|~m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k12_supinf_2(k1_mesfunc2(X1,X2),X3),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X4))).#trainneg
cnf(i_0_3697, plain, (k1_extreal2(k1_extreal2(k1_extreal2(np__0,k12_supinf_2(k2_mesfunc2(X1,X2),X3)),np__0),np__0)=k1_extreal2(k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X5,X2),X3),np__0),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X5)|v1_xboole_0(X1)|~r2_hidden(X3,k1_relset_1(X4,k2_mesfunc2(X1,X2)))|~v1_funct_1(k2_mesfunc2(X1,X2))|~m1_subset_1(k2_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(k12_supinf_2(k2_mesfunc2(X1,X2),X3),k7_numbers)|~m1_subset_1(X3,X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X5,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X5))).#trainneg
cnf(i_0_3558, plain, (k12_supinf_2(k1_mesfunc2(X1,k2_mesfunc2(X2,X3)),X4)=k1_extreal2(k1_extreal2(k3_extreal1(k12_supinf_2(X3,X4)),np__0),np__0)|v1_xboole_0(X2)|v1_xboole_0(X1)|r1_xxreal_0(np__0,k12_supinf_2(X3,X4))|~r2_hidden(X4,k1_relset_1(X1,k2_mesfunc2(X2,X3)))|~r2_hidden(X4,k1_relset_1(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(k2_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,X1)|~m1_subset_1(X4,X2)|~m1_subset_1(k12_supinf_2(X3,X4),k7_numbers))).#trainneg
cnf(i_0_4110, plain, (k1_extreal2(k1_extreal2(k1_extreal2(np__0,k12_supinf_2(k2_mesfunc2(X1,X2),X3)),np__0),np__0)=k1_extreal2(k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X4,X2),X3),np__0),np__0),np__0)|v1_xboole_0(X1)|v1_xboole_0(X4)|~r2_hidden(X3,k9_xtuple_0(k2_mesfunc2(X1,X2)))|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X4,X2))|~v1_funct_1(X2)|~m1_subset_1(k2_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k12_supinf_2(k2_mesfunc2(X1,X2),X3),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X4)|~v1_relat_1(k2_mesfunc2(X1,X2)))).#trainneg
cnf(i_0_4111, plain, (k1_extreal2(k1_extreal2(k1_extreal2(np__0,k12_supinf_2(k2_mesfunc2(X1,X2),X3)),np__0),np__0)=k1_extreal2(k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X4,X2),X3),np__0),np__0),np__0)|v1_xboole_0(X1)|v1_xboole_0(X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X4,X2))|~v1_funct_1(X2)|~m1_subset_1(k2_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k12_supinf_2(k2_mesfunc2(X1,X2),X3),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X4))).#trainneg
cnf(i_0_4097, plain, (k1_extreal2(k1_extreal2(k1_extreal2(np__0,k12_supinf_2(k1_mesfunc2(X1,X2),X3)),np__0),np__0)=k1_extreal2(k12_supinf_2(k1_mesfunc2(X5,k1_mesfunc2(X6,k1_mesfunc2(X4,X2))),X3),np__0)|v1_xboole_0(X4)|v1_xboole_0(X1)|v1_xboole_0(X6)|v1_xboole_0(X5)|~r2_hidden(X3,k9_xtuple_0(k1_mesfunc2(X1,X2)))|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X4,X2))|~v1_funct_1(X2)|~m1_subset_1(k12_supinf_2(k1_mesfunc2(X1,X2),X3),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X4)|~r2_hidden(X3,k1_relset_1(X5,k1_mesfunc2(X6,k1_mesfunc2(X4,X2))))|~r2_hidden(X3,k1_relset_1(X6,k1_mesfunc2(X4,X2)))|~v1_funct_1(k1_mesfunc2(X4,X2))|~m1_subset_1(k1_mesfunc2(X6,k1_mesfunc2(X4,X2)),k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X4,X2),k1_zfmisc_1(k2_zfmisc_1(X6,k7_numbers)))|~m1_subset_1(X3,X5)|~m1_subset_1(X3,X6))).#trainneg
cnf(i_0_3714, plain, (k1_extreal2(k1_extreal2(k1_extreal2(k2_supinf_2(k12_supinf_2(X1,X2)),np__0),np__0),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X3,X1),X2),np__0),np__0)|v1_xboole_0(X3)|~r2_hidden(X2,k1_relset_1(X3,X1))|~v1_funct_1(X1)|~m1_subset_1(k2_mesfunc2(X3,X1),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_3692, plain, (k12_supinf_2(k1_mesfunc2(X1,k1_mesfunc2(X2,k2_mesfunc2(X3,X4))),X5)=k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X6,X4),X5),np__0),np__0)|v1_xboole_0(X1)|v1_xboole_0(X6)|v1_xboole_0(X2)|v1_xboole_0(X3)|~r2_hidden(X5,k1_relset_1(X1,k1_mesfunc2(X2,k2_mesfunc2(X3,X4))))|~v1_funct_1(k1_mesfunc2(X2,k2_mesfunc2(X3,X4)))|~m1_subset_1(k1_mesfunc2(X2,k2_mesfunc2(X3,X4)),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X5,X1)|~r2_hidden(X5,k1_relset_1(X2,k2_mesfunc2(X3,X4)))|~r2_hidden(X5,k1_relset_1(X3,X4))|~r2_hidden(X5,k1_relset_1(X6,X4))|~v1_funct_1(X4)|~m1_subset_1(k2_mesfunc2(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X6,k7_numbers)))|~m1_subset_1(X5,X2)|~m1_subset_1(X5,X3)|~m1_subset_1(X5,X6))).#trainneg
cnf(i_0_4018, plain, (v1_valued_0(np__0))).#trainneg
cnf(i_0_1049, plain, (esk11_1(k1_zfmisc_1(X1))=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_550, plain, (v1_xboole_0(k2_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,k7_numbers))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_3982, plain, (~r2_hidden(X1,np__0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1261, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,k7_numbers))|v1_xboole_0(k10_mesfunc1(X2,X3))|v1_xboole_0(X2)|~m1_subset_1(X1,k10_mesfunc1(X2,X3))|~v5_relat_1(X3,k7_numbers)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~v1_funct_1(X3))).#trainneg
cnf(i_0_551, plain, (v3_membered(k2_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v3_membered(k2_zfmisc_1(X1,k7_numbers))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_552, plain, (v1_valued_0(k2_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_valued_0(k2_zfmisc_1(X1,k7_numbers))|~v1_relat_1(k2_zfmisc_1(X1,k7_numbers))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_4335, plain, (k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X4,X1),X2),np__0),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X3,X1),X2),np__0),np__0)|v1_xboole_0(X3)|v1_xboole_0(X4)|~r2_hidden(X2,k1_relset_1(X3,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X2,X3)|~r2_hidden(X2,k1_relset_1(X4,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,X4))).#trainneg
cnf(i_0_4362, plain, (k1_extreal2(k12_supinf_2(k1_mesfunc2(X4,k2_mesfunc2(X5,X1)),X2),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X3,X1),X2),np__0),np__0)|v1_xboole_0(X3)|v1_xboole_0(X5)|v1_xboole_0(X4)|~r2_hidden(X2,k1_relset_1(X3,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X2,X3)|~r2_hidden(X2,k1_relset_1(X4,k2_mesfunc2(X5,X1)))|~r2_hidden(X2,k1_relset_1(X5,X1))|~m1_subset_1(k2_mesfunc2(X5,X1),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X2,X4)|~m1_subset_1(X2,X5))).#trainneg
cnf(i_0_3635, plain, (k3_supinf_2(k6_mesfunc5(X1,X2,X3,X4),k6_mesfunc5(X1,X2,X3,X5))=k6_mesfunc5(X1,X2,X3,k3_mesfunc1(X1,X4,X5))|v1_xboole_0(X1)|v1_xboole_0(X2)|k9_xtuple_0(k1_mesfunc2(X1,X6))!=k1_relset_1(X1,X5)|k9_xtuple_0(k1_mesfunc2(X1,X6))!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X5,k9_xtuple_0(k1_mesfunc2(X1,X6)))|~r1_mesfunc1(X1,X2,X4,k9_xtuple_0(k1_mesfunc2(X1,X6)))|~r1_mesfunc5(X1,X2,X7,k1_mesfunc2(X1,X6))|~v4_measure1(X3,X1,X2)|~v4_measure1(X7,X1,X2)|~v6_supinf_2(X5)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X7)|~v10_valued_0(X3)|~v10_valued_0(X7)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X7,X2,k7_numbers)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(k1_mesfunc2(X1,X6))|~v1_funct_1(X7)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X1,X6),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_4046, plain, (k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X2)|v1_xboole_0(X1)|k1_relset_1(X1,k1_mesfunc2(X1,X5))!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,k1_mesfunc2(X1,X5)))|~r1_mesfunc5(X1,X2,X6,k1_mesfunc2(X1,X5))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_4242, plain, (k1_extreal2(k1_extreal2(k1_extreal2(np__0,k12_supinf_2(k2_mesfunc2(X1,X2),X3)),np__0),np__0)=k1_extreal2(k1_extreal2(k1_extreal2(np__0,k12_supinf_2(k2_mesfunc2(X4,X2),X3)),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X1)|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X4,X2))|~v1_funct_1(X2)|~m1_subset_1(k12_supinf_2(k2_mesfunc2(X1,X2),X3),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X4)|~m1_subset_1(k12_supinf_2(k2_mesfunc2(X4,X2),X3),k7_numbers))).#trainneg
cnf(i_0_3481, plain, (r1_xxreal_0(X1,esk11_1(k4_ordinal1))|r1_xxreal_0(esk11_1(k4_ordinal1),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1965, plain, (v1_xboole_0(k1_relset_1(X1,np__0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_553, plain, (v3_valued_0(k2_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v3_valued_0(k2_zfmisc_1(X1,k7_numbers))|~v1_relat_1(k2_zfmisc_1(X1,k7_numbers))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_4007, plain, (~v1_subset_1(np__0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_240, plain, (~r2_hidden(X1,esk15_1(X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4043, plain, (r1_xxreal_0(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X1)))))))))))))),X2)|r1_xxreal_0(X2,k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X1)))))))))))))))|~v1_xxreal_0(X2)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_554, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,k7_numbers))|v1_xboole_0(X2)|~r2_hidden(X1,k2_mesfunc2(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers))))).#trainneg
cnf(i_0_556, plain, (v2_valued_0(k2_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_4546, plain, (k12_supinf_2(k1_mesfunc2(X1,k1_mesfunc2(X2,k2_mesfunc2(X3,X4))),X5)=k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X7,X4),X5),np__0),np__0)|v1_xboole_0(X6)|v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X7)|~r2_hidden(X5,k1_relset_1(X1,k1_mesfunc2(X2,k2_mesfunc2(X3,X4))))|~r2_hidden(X5,k1_relset_1(X2,k2_mesfunc2(X3,X4)))|~r2_hidden(X5,k1_relset_1(X6,k2_mesfunc2(X3,X4)))|~v1_funct_1(k2_mesfunc2(X3,X4))|~m1_subset_1(k1_mesfunc2(X2,k2_mesfunc2(X3,X4)),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X6,k7_numbers)))|~m1_subset_1(X5,X1)|~m1_subset_1(X5,X2)|~m1_subset_1(X5,X6)|~r2_hidden(X5,k1_relset_1(X7,X4))|~r2_hidden(X5,k1_relset_1(X3,X4))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X7,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X5,X7)|~m1_subset_1(X5,X3))).#trainneg
cnf(i_0_4493, plain, (k12_supinf_2(k1_mesfunc2(X1,k1_mesfunc2(X2,k2_mesfunc2(X3,X4))),X5)=k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X6,X4),X5),np__0),np__0)|v1_xboole_0(X2)|v1_xboole_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X6)|~r2_hidden(X5,k1_relset_1(X1,k1_mesfunc2(X2,k2_mesfunc2(X3,X4))))|~r2_hidden(X5,k1_relset_1(X2,k2_mesfunc2(X3,X4)))|~v1_funct_1(k2_mesfunc2(X3,X4))|~m1_subset_1(k1_mesfunc2(X2,k2_mesfunc2(X3,X4)),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X5,X1)|~m1_subset_1(X5,X2)|~r2_hidden(X5,k1_relset_1(X6,X4))|~r2_hidden(X5,k1_relset_1(X3,X4))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X6,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X5,X6)|~m1_subset_1(X5,X3))).#trainneg
cnf(i_0_4527, plain, (k12_supinf_2(k1_mesfunc2(X1,k1_mesfunc2(X2,k2_mesfunc2(X3,X4))),X5)=k1_extreal2(k12_supinf_2(k1_mesfunc2(X6,k2_mesfunc2(X7,X4)),X5),np__0)|v1_xboole_0(X2)|v1_xboole_0(X1)|v1_xboole_0(X6)|v1_xboole_0(X7)|v1_xboole_0(X3)|~r2_hidden(X5,k1_relset_1(X1,k1_mesfunc2(X2,k2_mesfunc2(X3,X4))))|~r2_hidden(X5,k1_relset_1(X2,k2_mesfunc2(X3,X4)))|~v1_funct_1(k2_mesfunc2(X3,X4))|~m1_subset_1(k1_mesfunc2(X2,k2_mesfunc2(X3,X4)),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X5,X1)|~m1_subset_1(X5,X2)|~r2_hidden(X5,k1_relset_1(X6,k2_mesfunc2(X7,X4)))|~r2_hidden(X5,k1_relset_1(X3,X4))|~r2_hidden(X5,k1_relset_1(X7,X4))|~v1_funct_1(X4)|~m1_subset_1(k2_mesfunc2(X7,X4),k1_zfmisc_1(k2_zfmisc_1(X6,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X7,k7_numbers)))|~m1_subset_1(X5,X3)|~m1_subset_1(X5,X6)|~m1_subset_1(X5,X7))).#trainneg
cnf(i_0_4654, plain, (m1_subset_1(k12_supinf_2(k2_mesfunc2(X1,X2),X3),k7_numbers)|v1_xboole_0(X1)|~v1_relat_1(k2_mesfunc2(X1,X2))|~v1_funct_1(k2_mesfunc2(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_557, plain, (v1_valued_0(k2_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_membered(k7_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_242, plain, (~r2_hidden(X1,esk11_1(k1_zfmisc_1(X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4651, plain, (r1_xxreal_0(X1,k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X2))))))))))))))))|r1_xxreal_0(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X2))))))))))))))),X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k7_numbers))).#trainneg
cnf(i_0_3057, plain, (r1_xxreal_0(k10_mesfunc1(X1,X2),X3)|r1_xxreal_0(X3,k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v1_xxreal_0(X3)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_559, plain, (v5_relat_1(k2_mesfunc2(X1,X2),k7_numbers)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_4544, plain, (k12_supinf_2(k1_mesfunc2(X1,k1_mesfunc2(X2,k1_mesfunc2(X3,k2_mesfunc2(X4,X5)))),X6)=k1_extreal2(k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X7,X5),X6),np__0),np__0),np__0)|v1_xboole_0(X2)|v1_xboole_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X4)|v1_xboole_0(X7)|~r2_hidden(X6,k1_relset_1(X1,k1_mesfunc2(X2,k1_mesfunc2(X3,k2_mesfunc2(X4,X5)))))|~r2_hidden(X6,k1_relset_1(X2,k1_mesfunc2(X3,k2_mesfunc2(X4,X5))))|~v1_funct_1(k1_mesfunc2(X3,k2_mesfunc2(X4,X5)))|~m1_subset_1(k1_mesfunc2(X2,k1_mesfunc2(X3,k2_mesfunc2(X4,X5))),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X3,k2_mesfunc2(X4,X5)),k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X6,X1)|~m1_subset_1(X6,X2)|~r2_hidden(X6,k1_relset_1(X3,k2_mesfunc2(X4,X5)))|~r2_hidden(X6,k1_relset_1(X7,X5))|~r2_hidden(X6,k1_relset_1(X4,X5))|~v1_funct_1(X5)|~m1_subset_1(k2_mesfunc2(X4,X5),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,k7_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X6,X7)|~m1_subset_1(X6,X3)|~m1_subset_1(X6,X4))).#trainneg
cnf(i_0_3608, plain, (k1_extreal2(k1_extreal2(k1_extreal2(k3_extreal1(k12_supinf_2(X1,X2)),np__0),np__0),np__0)=k1_extreal2(k12_supinf_2(k1_mesfunc2(X3,k2_mesfunc2(X4,X1)),X2),np__0)|v1_xboole_0(X4)|v1_xboole_0(X3)|r1_xxreal_0(np__0,k12_supinf_2(X1,X2))|~r2_hidden(X2,k1_relset_1(X3,k2_mesfunc2(X4,X1)))|~r2_hidden(X2,k1_relset_1(X4,X1))|~v1_funct_1(X1)|~m1_subset_1(k2_mesfunc2(X4,X1),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,X3)|~m1_subset_1(X2,X4)|~m1_subset_1(k12_supinf_2(X1,X2),k7_numbers))).#trainneg
cnf(i_0_4866, plain, (v1_funct_1(k10_mesfunc1(X1,X2))|v1_xboole_0(X1)|~v5_relat_1(k1_mesfunc2(X1,X2),k7_numbers)|~v4_relat_1(k2_mesfunc2(X1,X2),X1)|~v4_relat_1(k1_mesfunc2(X1,X2),X1)|~v1_relat_1(k2_mesfunc2(X1,X2))|~v1_relat_1(k1_mesfunc2(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_4526, plain, (k1_extreal2(k12_supinf_2(k1_mesfunc2(X5,k2_mesfunc2(X6,X2)),X3),np__0)=k1_extreal2(k12_supinf_2(k1_mesfunc2(X4,k2_mesfunc2(X1,X2)),X3),np__0)|v1_xboole_0(X4)|v1_xboole_0(X5)|v1_xboole_0(X6)|v1_xboole_0(X1)|~r2_hidden(X3,k1_relset_1(X4,k2_mesfunc2(X1,X2)))|~v1_funct_1(k2_mesfunc2(X1,X2))|~m1_subset_1(k2_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X3,X4)|~r2_hidden(X3,k1_relset_1(X5,k2_mesfunc2(X6,X2)))|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X6,X2))|~v1_funct_1(X2)|~m1_subset_1(k2_mesfunc2(X6,X2),k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X6,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X5)|~m1_subset_1(X3,X6))).#trainneg
cnf(i_0_446, plain, (v1_finsub_1(esk14_1(X1)))).#trainneg
cnf(i_0_560, plain, (v1_relat_1(k2_mesfunc2(X1,X2))|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_3317, plain, (k1_relset_1(X1,np__0)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4864, plain, (r1_xxreal_0(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X1)))))))))))))))),X2)|r1_xxreal_0(X2,k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X1)))))))))))))))))|~v1_xxreal_0(X2)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_562, plain, (k1_relset_1(X1,esk10_2(X2,X1))=k9_xtuple_0(esk10_2(X2,X1))|v1_xboole_0(X1)|~v1_relat_1(esk10_2(X2,X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(X1,X2)|~v1_prob_1(X1,X2))).#trainneg
cnf(i_0_563, plain, (m1_subset_1(k1_relset_1(X1,esk10_2(X2,X1)),k1_zfmisc_1(X1))|v1_xboole_0(X1)|~v1_relat_1(esk10_2(X2,X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(X1,X2)|~v1_prob_1(X1,X2))).#trainneg
cnf(i_0_5077, plain, (m1_subset_1(k1_relset_1(X1,X2),k1_zfmisc_1(X1))|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_4576, plain, (k3_supinf_2(k6_mesfunc5(X1,X2,X3,X4),k6_mesfunc5(X1,X2,X3,X5))=k6_mesfunc5(X1,X2,X3,k3_mesfunc1(X1,X4,X5))|v1_xboole_0(X2)|v1_xboole_0(X1)|k1_relset_1(X1,X6)!=k1_relset_1(X1,X5)|k1_relset_1(X1,X6)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X5,k1_relset_1(X1,X6))|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,X6))|~r1_mesfunc5(X1,X2,X7,k1_mesfunc2(X1,X6))|~v4_measure1(X3,X1,X2)|~v4_measure1(X7,X1,X2)|~v6_supinf_2(X5)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X7)|~v10_valued_0(X3)|~v10_valued_0(X7)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X7,X2,k7_numbers)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3696, plain, (k1_extreal2(k1_extreal2(np__0,k12_supinf_2(k2_mesfunc2(X1,X2),X3)),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X5,X2),X3),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X5)|v1_xboole_0(X1)|~r2_hidden(X3,k1_relset_1(X4,k2_mesfunc2(X1,X2)))|~v1_funct_1(k2_mesfunc2(X1,X2))|~m1_subset_1(k2_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(k12_supinf_2(k2_mesfunc2(X1,X2),X3),k7_numbers)|~m1_subset_1(X3,X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X5,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X5))).#trainneg
cnf(i_0_4542, plain, (k1_extreal2(k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X6,X3),X4),np__0),np__0),np__0)=k1_extreal2(k12_supinf_2(k1_mesfunc2(X5,k1_mesfunc2(X1,k2_mesfunc2(X2,X3))),X4),np__0)|v1_xboole_0(X5)|v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X6)|~r2_hidden(X4,k1_relset_1(X5,k1_mesfunc2(X1,k2_mesfunc2(X2,X3))))|~v1_funct_1(k1_mesfunc2(X1,k2_mesfunc2(X2,X3)))|~m1_subset_1(k1_mesfunc2(X1,k2_mesfunc2(X2,X3)),k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X4,X5)|~r2_hidden(X4,k1_relset_1(X1,k2_mesfunc2(X2,X3)))|~r2_hidden(X4,k1_relset_1(X6,X3))|~r2_hidden(X4,k1_relset_1(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(k2_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X6,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,X6)|~m1_subset_1(X4,X1)|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_4357, plain, (k1_extreal2(k1_extreal2(k1_extreal2(np__0,k2_supinf_2(k12_supinf_2(X1,X2))),np__0),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X3,X1),X2),np__0),np__0)|v1_xboole_0(X3)|~r2_hidden(X2,k1_relset_1(X3,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X2,X3)|~m1_subset_1(k2_supinf_2(k12_supinf_2(X1,X2)),k7_numbers))).#trainneg
cnf(i_0_4079, plain, (k1_extreal2(k1_extreal2(k1_extreal2(np__0,k12_supinf_2(k1_mesfunc2(X1,X2),X3)),np__0),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X5,k1_mesfunc2(X4,X2)),X3),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X1)|v1_xboole_0(X5)|~r2_hidden(X3,k9_xtuple_0(k1_mesfunc2(X1,X2)))|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X4,X2))|~v1_funct_1(X2)|~m1_subset_1(k12_supinf_2(k1_mesfunc2(X1,X2),X3),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X4)|~r2_hidden(X3,k1_relset_1(X5,k1_mesfunc2(X4,X2)))|~v1_funct_1(k1_mesfunc2(X4,X2))|~m1_subset_1(k1_mesfunc2(X4,X2),k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X3,X5))).#trainneg
cnf(i_0_444, plain, (v2_finsub_1(esk14_1(X1)))).#trainneg
cnf(i_0_564, plain, (v1_xxreal_0(esk10_2(X1,X2))|v1_xboole_0(X2)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3254, plain, (~r2_hidden(X1,k1_relset_1(X2,np__0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5123, plain, (r1_xxreal_0(X1,k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X2))))))))))))))))))|r1_xxreal_0(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X2))))))))))))))))),X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k7_numbers))).#trainneg
cnf(i_0_2975, plain, (k10_mesfunc1(X1,X2)=k10_mesfunc1(X3,X4)|v1_xboole_0(X3)|v1_xboole_0(X1)|~v5_relat_1(X4,k7_numbers)|~v4_relat_1(X4,X3)|~v1_relat_1(X4)|~v1_funct_1(X4)|~v1_xboole_0(k2_zfmisc_1(X3,k7_numbers))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_xboole_0(k2_zfmisc_1(X1,k7_numbers)))).#trainneg
cnf(i_0_565, plain, (v1_xcmplx_0(esk10_2(X1,X2))|v1_xboole_0(X2)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_566, plain, (v1_xreal_0(esk10_2(X1,X2))|v1_xboole_0(X2)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3642, plain, (k1_extreal2(k1_extreal2(np__0,k12_supinf_2(k1_mesfunc2(X1,X2),X3)),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X4,X2),X3),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X1)|~r2_hidden(X3,k9_xtuple_0(k1_mesfunc2(X1,X2)))|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X4,X2))|~v1_funct_1(X2)|~m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k12_supinf_2(k1_mesfunc2(X1,X2),X3),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X4))).#trainneg
cnf(i_0_3637, plain, (k12_supinf_2(k1_mesfunc2(X1,X2),X3)=k1_extreal2(k12_supinf_2(X4,X3),np__0)|v1_xboole_0(X1)|k1_mesfunc2(X1,X2)!=k1_mesfunc2(X1,X4)|~r2_hidden(X3,k9_xtuple_0(k1_mesfunc2(X1,X2)))|~v1_funct_1(X4)|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_3641, plain, (k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X1,X2),X3),np__0),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X4,X2),X3),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X1)|~r2_hidden(X3,k9_xtuple_0(k1_mesfunc2(X1,X2)))|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X4,X2))|~v1_funct_1(X2)|~m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X4))).#trainneg
cnf(i_0_1269, plain, (v1_xcmplx_0(esk15_1(X1))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3467, plain, (k1_relset_1(X1,np__0)=k9_xtuple_0(np__0))).#trainneg
cnf(i_0_456, plain, (v3_finsub_1(esk14_1(X1)))).#trainneg
cnf(i_0_5434, negated_conjecture, (v1_xboole_0(X1)|k1_relset_1(esk1_0,k10_mesfunc1(X1,X2))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~v1_funct_1(k10_mesfunc1(X1,X2))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_xboole_0(k2_zfmisc_1(esk1_0,k7_numbers))|~v1_xboole_0(k2_zfmisc_1(X1,k7_numbers)))).#trainneg
cnf(i_0_567, plain, (v7_ordinal1(esk10_2(X1,X2))|v1_xboole_0(X2)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3967, plain, (m1_subset_1(k12_supinf_2(np__0,X1),k7_numbers)|~v1_funct_1(np__0))).#trainneg
cnf(i_0_5403, plain, (v6_supinf_2(np__0)|~v1_xboole_0(X1)|~v1_funct_1(np__0))).#trainneg
cnf(i_0_5424, plain, (r1_xxreal_0(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X1)))))))))))))))))),X2)|r1_xxreal_0(X2,k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X1)))))))))))))))))))|~v1_xxreal_0(X2)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_569, plain, (v1_xboole_0(esk10_2(X1,X2))|v1_xboole_0(X2)|~v1_xboole_0(k2_zfmisc_1(X2,k7_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_570, plain, (v3_membered(esk10_2(X1,X2))|v1_xboole_0(X2)|~v3_membered(k2_zfmisc_1(X2,k7_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_5632, plain, (k12_supinf_2(k1_mesfunc2(X1,X2),X3)=k1_extreal2(k12_supinf_2(X4,X3),np__0)|v1_xboole_0(X1)|k1_mesfunc2(X1,X2)!=k1_mesfunc2(X1,X4)|~r2_hidden(X3,k1_relset_1(X1,k1_mesfunc2(X1,X2)))|~v1_funct_1(X4)|~v1_funct_1(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_3594, plain, (k1_extreal2(k1_extreal2(k1_extreal2(k2_supinf_2(k12_supinf_2(X1,X2)),np__0),np__0),np__0)=k1_extreal2(k1_extreal2(np__0,k12_supinf_2(k2_mesfunc2(X4,X1),X2)),np__0)|v1_xboole_0(X4)|v1_xboole_0(X3)|~r2_hidden(X2,k1_relset_1(X3,k2_mesfunc2(X4,X1)))|~r2_hidden(X2,k1_relset_1(X4,X1))|~v1_funct_1(X1)|~m1_subset_1(k2_mesfunc2(X4,X1),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,X3)|~m1_subset_1(X2,X4)|~v1_funct_1(k2_mesfunc2(X4,X1))|~m1_subset_1(k12_supinf_2(k2_mesfunc2(X4,X1),X2),k7_numbers))).#trainneg
cnf(i_0_3681, plain, (r1_xxreal_0(k1_extreal2(k12_supinf_2(k2_mesfunc2(X5,X3),X4),np__0),np__0)|r1_xxreal_0(np__0,k1_extreal2(k12_supinf_2(k2_mesfunc2(X5,X3),X4),np__0))|v1_xboole_0(X5)|v1_xboole_0(X1)|v1_xboole_0(X2)|~v2_valued_0(k1_mesfunc2(X1,k2_mesfunc2(X2,X3)))|~v1_relat_1(k1_mesfunc2(X1,k2_mesfunc2(X2,X3)))|~v1_funct_1(k1_mesfunc2(X1,k2_mesfunc2(X2,X3)))|~r2_hidden(X4,k1_relset_1(X1,k2_mesfunc2(X2,X3)))|~r2_hidden(X4,k1_relset_1(X2,X3))|~r2_hidden(X4,k1_relset_1(X5,X3))|~v1_funct_1(X3)|~m1_subset_1(k2_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X4,X1)|~m1_subset_1(X4,X2)|~m1_subset_1(X4,X5))).#trainneg
cnf(i_0_4534, plain, (k1_extreal2(np__0,k12_supinf_2(k1_mesfunc2(X1,k2_mesfunc2(X2,X3)),X4))=k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X5,X3),X4),np__0),np__0)|v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X5)|~m1_subset_1(k12_supinf_2(k1_mesfunc2(X1,k2_mesfunc2(X2,X3)),X4),k7_numbers)|~r2_hidden(X4,k1_relset_1(X1,k2_mesfunc2(X2,X3)))|~r2_hidden(X4,k1_relset_1(X5,X3))|~r2_hidden(X4,k1_relset_1(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(k2_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,X5)|~m1_subset_1(X4,X1)|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_5738, plain, (k1_relset_1(X1,np__0)=k1_relset_1(X2,np__0))).#trainneg
cnf(i_0_5839, plain, (r1_xxreal_0(np__0,k1_extreal2(k12_supinf_2(k2_mesfunc2(X1,X2),X3),np__0))|r1_xxreal_0(k1_extreal2(k12_supinf_2(k2_mesfunc2(X1,X2),X3),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X1)|~v2_valued_0(k1_mesfunc2(X4,k2_mesfunc2(X4,X2)))|~r2_hidden(X3,k1_relset_1(X4,X2))|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_relat_1(k1_mesfunc2(X4,k2_mesfunc2(X4,X2)))|~v1_funct_1(k1_mesfunc2(X4,k2_mesfunc2(X4,X2)))|~v1_funct_1(X2)|~m1_subset_1(k2_mesfunc2(X4,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X4)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_5767, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k9_xtuple_0(np__0)))).#trainneg
cnf(i_0_5765, plain, (m1_subset_1(k9_xtuple_0(np__0),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5799, plain, (v6_supinf_2(np__0)|~v1_funct_1(np__0))).#trainneg
cnf(i_0_5766, plain, (k9_xtuple_0(np__0)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_571, plain, (v1_valued_0(esk10_2(X1,X2))|v1_xboole_0(X2)|~v1_valued_0(k2_zfmisc_1(X2,k7_numbers))|~v1_relat_1(k2_zfmisc_1(X2,k7_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_5764, plain, (v1_xboole_0(k9_xtuple_0(np__0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5768, plain, (~r2_hidden(X1,k9_xtuple_0(np__0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1042, negated_conjecture, (v1_xcmplx_0(k12_supinf_2(esk3_0,X1))|~v1_membered(k7_numbers))).#trainneg
cnf(i_0_572, plain, (v3_valued_0(esk10_2(X1,X2))|v1_xboole_0(X2)|~v3_valued_0(k2_zfmisc_1(X2,k7_numbers))|~v1_relat_1(k2_zfmisc_1(X2,k7_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_573, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,k7_numbers))|v1_xboole_0(X2)|~r2_hidden(X1,esk10_2(X3,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))|~v4_prob_1(X2,X3)|~v1_prob_1(X2,X3))).#trainneg
cnf(i_0_5196, plain, (m1_subset_1(k9_xtuple_0(k1_mesfunc2(X1,X2)),k1_zfmisc_1(X1))|v1_xboole_0(X1)|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_5924, plain, (k1_relset_1(X2,np__0)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5923, plain, (m1_subset_1(k1_relset_1(X2,np__0),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5926, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X3,np__0)))).#trainneg
cnf(i_0_5991, plain, (k9_xtuple_0(np__0)=np__0)).#trainneg
cnf(i_0_5922, plain, (v1_xboole_0(k1_relset_1(X2,np__0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5927, plain, (~r2_hidden(X1,k1_relset_1(X3,np__0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_576, plain, (v1_valued_0(esk10_2(X1,X2))|v1_xboole_0(X2)|~v1_membered(k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_3636, plain, (k12_supinf_2(k1_mesfunc2(X1,X2),X3)=k1_extreal2(k2_supinf_2(k12_supinf_2(X4,X3)),np__0)|v1_xboole_0(X1)|k1_mesfunc2(X1,X2)!=k2_mesfunc2(X1,X4)|~r2_hidden(X3,k9_xtuple_0(k1_mesfunc2(X1,X2)))|~v1_funct_1(X4)|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_3524, plain, (r1_xxreal_0(k1_extreal2(k1_extreal2(k2_supinf_2(k12_supinf_2(X3,X4)),np__0),np__0),np__0)|r1_xxreal_0(np__0,k1_extreal2(k1_extreal2(k2_supinf_2(k12_supinf_2(X3,X4)),np__0),np__0))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v2_valued_0(k1_mesfunc2(X1,k2_mesfunc2(X2,X3)))|~v1_relat_1(k1_mesfunc2(X1,k2_mesfunc2(X2,X3)))|~v1_funct_1(k1_mesfunc2(X1,k2_mesfunc2(X2,X3)))|~r2_hidden(X4,k1_relset_1(X1,k2_mesfunc2(X2,X3)))|~r2_hidden(X4,k1_relset_1(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(k2_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,X1)|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_5198, plain, (m1_subset_1(k9_xtuple_0(k10_mesfunc1(X1,X2)),k1_zfmisc_1(X1))|v1_xboole_0(X1)|~v1_funct_1(k10_mesfunc1(X1,X2))|~m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_4552, plain, (k1_extreal2(k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X6,X3),X4),np__0),np__0),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X5,k2_mesfunc2(X2,X3)),X4),np__0),np__0)|v1_xboole_0(X1)|v1_xboole_0(X5)|v1_xboole_0(X2)|v1_xboole_0(X6)|~r2_hidden(X4,k1_relset_1(X5,k2_mesfunc2(X2,X3)))|~r2_hidden(X4,k1_relset_1(X1,k2_mesfunc2(X2,X3)))|~v1_funct_1(k2_mesfunc2(X2,X3))|~m1_subset_1(k2_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k12_supinf_2(k2_mesfunc2(X2,X3),X4),k7_numbers)|~m1_subset_1(X4,X5)|~m1_subset_1(X4,X1)|~r2_hidden(X4,k1_relset_1(X6,X3))|~r2_hidden(X4,k1_relset_1(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X6,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,X6)|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_3539, plain, (k1_extreal2(k1_extreal2(k1_extreal2(np__0,k12_supinf_2(k2_mesfunc2(X1,X2),X3)),np__0),np__0)=k1_extreal2(k1_extreal2(k1_extreal2(k1_extreal2(k2_supinf_2(k12_supinf_2(X2,X3)),np__0),np__0),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X1)|~r2_hidden(X3,k1_relset_1(X4,k2_mesfunc2(X1,X2)))|~v1_funct_1(k2_mesfunc2(X1,X2))|~m1_subset_1(k2_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(k12_supinf_2(k2_mesfunc2(X1,X2),X3),k7_numbers)|~m1_subset_1(X3,X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_1368, plain, (v1_xcmplx_0(esk15_1(X1))|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5772, plain, (v1_xxreal_0(k12_supinf_2(np__0,X1))|~v1_funct_1(np__0))).#trainneg
cnf(i_0_1044, negated_conjecture, (v7_ordinal1(k12_supinf_2(esk3_0,X1))|~v6_membered(k7_numbers))).#trainneg
cnf(i_0_5426, negated_conjecture, (k9_xtuple_0(X1)=k1_relset_1(esk2_0,esk3_0)|v1_xboole_0(X2)|k10_mesfunc1(X2,X3)!=esk3_0|~v5_relat_1(X1,k7_numbers)|~v4_relat_1(X1,esk2_0)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_relat_1(X3,k7_numbers)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v1_xboole_0(k2_zfmisc_1(esk2_0,k7_numbers))|~v1_xboole_0(k2_zfmisc_1(X2,k7_numbers)))).#trainneg
cnf(i_0_1119, plain, (k2_xcmplx_0(np__0,X1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1429, negated_conjecture, (v1_xcmplx_0(k12_supinf_2(esk4_0,X1))|~v1_membered(k7_numbers))).#trainneg
cnf(i_0_4551, plain, (k1_extreal2(k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X7,X3),X4),np__0),np__0),np__0)=k1_extreal2(k12_supinf_2(k1_mesfunc2(X5,k1_mesfunc2(X6,k2_mesfunc2(X2,X3))),X4),np__0)|v1_xboole_0(X6)|v1_xboole_0(X1)|v1_xboole_0(X5)|v1_xboole_0(X2)|v1_xboole_0(X7)|~r2_hidden(X4,k1_relset_1(X5,k1_mesfunc2(X6,k2_mesfunc2(X2,X3))))|~r2_hidden(X4,k1_relset_1(X6,k2_mesfunc2(X2,X3)))|~r2_hidden(X4,k1_relset_1(X1,k2_mesfunc2(X2,X3)))|~v1_funct_1(k2_mesfunc2(X2,X3))|~m1_subset_1(k1_mesfunc2(X6,k2_mesfunc2(X2,X3)),k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X6,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,X5)|~m1_subset_1(X4,X6)|~m1_subset_1(X4,X1)|~r2_hidden(X4,k1_relset_1(X7,X3))|~r2_hidden(X4,k1_relset_1(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X7,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,X7)|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_1749, negated_conjecture, (r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)))|r1_xxreal_0(k6_mesfunc5(esk1_0,esk2_0,X1,k1_mesfunc2(esk1_0,X2)),np__0)|~v4_measure1(X1,esk1_0,esk2_0)|~v6_supinf_2(X1)|~v10_valued_0(X1)|~v1_funct_2(X1,esk2_0,k7_numbers)|~v1_funct_1(k1_mesfunc2(esk1_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k1_mesfunc2(esk1_0,X2),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk2_0,k7_numbers)))|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,X2))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,X2,k1_relset_1(esk1_0,esk4_0))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_1182, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|k9_xtuple_0(k2_mesfunc2(X1,X5))!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k9_xtuple_0(k2_mesfunc2(X1,X5)))|~r1_mesfunc5(X1,X2,X6,k2_mesfunc2(X1,X5))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(k2_mesfunc2(X1,X5))|~v1_funct_1(X6)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X1,X5),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v1_relat_1(k2_mesfunc2(X1,X5))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1841, plain, (m1_subset_1(k12_supinf_2(k1_mesfunc2(X3,X1),X2),k7_numbers)|v1_xboole_0(X3)|~m1_subset_1(k12_supinf_2(X1,X2),k7_numbers)|~r2_hidden(X2,k1_relset_1(X3,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_6074, plain, (v1_xboole_0(k1_relset_1(X1,np__0)))).#trainneg
cnf(i_0_6034, plain, (k1_relset_1(X1,np__0)=np__0)).#trainneg
cnf(i_0_1560, plain, (esk15_1(X1)=esk15_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6135, negated_conjecture, (v1_xcmplx_0(k12_supinf_2(esk3_0,X1))|~v6_membered(k7_numbers))).#trainneg
cnf(i_0_1543, plain, (k9_xtuple_0(X1)=np__0|v1_xboole_0(X2)|k10_mesfunc1(X2,X1)!=np__0|~v5_relat_1(X1,k7_numbers)|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(np__0)|~v1_funct_1(X1))).#trainneg
cnf(i_0_583, plain, (k1_relset_1(X1,esk15_1(k2_zfmisc_1(X1,k7_numbers)))=k1_relset_1(X1,X2)|v1_xboole_0(X1)|esk15_1(k2_zfmisc_1(X1,k7_numbers))!=k1_mesfunc2(X1,X2)|~v1_funct_1(X2)|~v1_funct_1(esk15_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_1431, negated_conjecture, (v7_ordinal1(k12_supinf_2(esk4_0,X1))|~v6_membered(k7_numbers))).#trainneg
cnf(i_0_585, plain, (k1_relset_1(X1,esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=k1_relset_1(X1,X2)|v1_xboole_0(X1)|esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))!=k1_mesfunc2(X1,X2)|~v1_funct_1(X2)|~v1_funct_1(esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_586, plain, (k1_relset_1(X1,esk10_2(X2,X1))=k1_relset_1(X1,X3)|v1_xboole_0(X1)|esk10_2(X2,X1)!=k1_mesfunc2(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(esk10_2(X2,X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(X1,X2)|~v1_prob_1(X1,X2))).#trainneg
cnf(i_0_6165, plain, (r1_xxreal_0(k12_supinf_2(k1_mesfunc2(X1,X2),X3),np__0)|r1_xxreal_0(np__0,k12_supinf_2(k1_mesfunc2(X1,X2),X3))|v1_xboole_0(X1)|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k12_supinf_2(X2,X3),k7_numbers)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_6084, plain, (k12_supinf_2(k1_mesfunc2(X1,X2),X3)=k1_extreal2(k2_supinf_2(k12_supinf_2(X4,X3)),np__0)|v1_xboole_0(X1)|k1_mesfunc2(X1,X2)!=k2_mesfunc2(X1,X4)|~r2_hidden(X3,k1_relset_1(X1,k1_mesfunc2(X1,X2)))|~v1_funct_1(X4)|~v1_funct_1(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_6139, plain, (k1_xxreal_3(np__0,X1)=X1|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2241, plain, (k12_supinf_2(k1_mesfunc2(X1,k1_mesfunc2(X2,k1_mesfunc2(X3,X4))),X5)=k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X6,X4),X5),np__0),np__0)|v1_xboole_0(X1)|v1_xboole_0(X3)|v1_xboole_0(X6)|v1_xboole_0(X2)|~r2_hidden(X5,k1_relset_1(X1,k1_mesfunc2(X2,k1_mesfunc2(X3,X4))))|~v1_funct_1(k1_mesfunc2(X2,k1_mesfunc2(X3,X4)))|~m1_subset_1(k1_mesfunc2(X2,k1_mesfunc2(X3,X4)),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X5,X1)|~r2_hidden(X5,k1_relset_1(X2,k1_mesfunc2(X3,X4)))|~r2_hidden(X5,k1_relset_1(X3,X4))|~r2_hidden(X5,k1_relset_1(X6,X4))|~v1_funct_1(X4)|~m1_subset_1(k1_mesfunc2(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X6,k7_numbers)))|~m1_subset_1(X5,X2)|~m1_subset_1(X5,X3)|~m1_subset_1(X5,X6))).#trainneg
cnf(i_0_2243, plain, (k12_supinf_2(k1_mesfunc2(X1,k1_mesfunc2(X2,k1_mesfunc2(X3,X4))),X5)=k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X7,X4),X5),np__0),np__0)|v1_xboole_0(X6)|v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X7)|~r2_hidden(X5,k1_relset_1(X1,k1_mesfunc2(X2,k1_mesfunc2(X3,X4))))|~r2_hidden(X5,k1_relset_1(X2,k1_mesfunc2(X3,X4)))|~r2_hidden(X5,k1_relset_1(X6,k1_mesfunc2(X3,X4)))|~v1_funct_1(k1_mesfunc2(X3,X4))|~m1_subset_1(k1_mesfunc2(X2,k1_mesfunc2(X3,X4)),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X6,k7_numbers)))|~m1_subset_1(X5,X1)|~m1_subset_1(X5,X2)|~m1_subset_1(X5,X6)|~r2_hidden(X5,k1_relset_1(X3,X4))|~r2_hidden(X5,k1_relset_1(X7,X4))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X7,k7_numbers)))|~m1_subset_1(X5,X3)|~m1_subset_1(X5,X7))).#trainneg
cnf(i_0_4519, plain, (k1_extreal2(np__0,k1_extreal2(k12_supinf_2(k2_mesfunc2(X1,X2),X3),np__0))=k1_extreal2(k12_supinf_2(k1_mesfunc2(X4,k2_mesfunc2(X5,X2)),X3),np__0)|v1_xboole_0(X4)|v1_xboole_0(X5)|v1_xboole_0(X1)|~m1_subset_1(k1_extreal2(k12_supinf_2(k2_mesfunc2(X1,X2),X3),np__0),k7_numbers)|~r2_hidden(X3,k1_relset_1(X4,k2_mesfunc2(X5,X2)))|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X5,X2))|~v1_funct_1(X2)|~m1_subset_1(k2_mesfunc2(X5,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X4)|~m1_subset_1(X3,X5))).#trainneg
cnf(i_0_3100, plain, (k7_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4))=k6_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_mesfunc1(X1,X2,X4,k9_xtuple_0(k10_mesfunc1(X1,X4)))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k9_xtuple_0(k10_mesfunc1(X1,X4)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~m1_subset_1(k9_xtuple_0(k10_mesfunc1(X1,X4)),k1_zfmisc_1(k1_relset_1(X1,X4))))).#trainneg
cnf(i_0_6289, negated_conjecture, (v1_xcmplx_0(k12_supinf_2(esk4_0,X1))|~v6_membered(k7_numbers))).#trainneg
cnf(i_0_592, plain, (k1_relset_1(X1,esk15_1(k2_zfmisc_1(X1,k7_numbers)))=k1_relset_1(X1,X2)|v1_xboole_0(X1)|esk15_1(k2_zfmisc_1(X1,k7_numbers))!=k2_mesfunc2(X1,X2)|~v1_funct_1(X2)|~v1_funct_1(esk15_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_6246, plain, (~v1_subset_1(esk15_1(X2),X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5477, plain, (v4_relat_1(k10_mesfunc1(X3,X4),X1)|v1_xboole_0(X1)|v1_xboole_0(X3)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v5_relat_1(X4,k7_numbers)|~v4_relat_1(X4,X3)|~v1_relat_1(X4)|~v1_funct_1(X4)|~v1_xboole_0(k2_zfmisc_1(X1,k7_numbers))|~v1_xboole_0(k2_zfmisc_1(X3,k7_numbers)))).#trainneg
cnf(i_0_594, plain, (k1_relset_1(X1,esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=k1_relset_1(X1,X2)|v1_xboole_0(X1)|esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))!=k2_mesfunc2(X1,X2)|~v1_funct_1(X2)|~v1_funct_1(esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_595, plain, (k1_relset_1(X1,esk10_2(X2,X1))=k1_relset_1(X1,X3)|v1_xboole_0(X1)|esk10_2(X2,X1)!=k2_mesfunc2(X1,X3)|~v1_funct_1(X3)|~v1_funct_1(esk10_2(X2,X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(X1,X2)|~v1_prob_1(X1,X2))).#trainneg
cnf(i_0_1141, plain, (r1_xxreal_0(k3_xxreal_3(X1,X2),k3_supinf_2(X2,k3_extreal1(X1)))|~r1_xxreal_0(np__0,X2)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)|~r1_xxreal_0(np__0,k3_xxreal_3(X1,X2))|~m1_subset_1(k3_xxreal_3(X1,X2),k7_numbers))).#trainneg
cnf(i_0_1922, plain, (k1_extreal2(k1_extreal2(k1_extreal2(k12_supinf_2(X2,X3),np__0),np__0),np__0)=k1_extreal2(k12_supinf_2(k1_mesfunc2(X4,k1_mesfunc2(X1,X2)),X3),np__0)|v1_xboole_0(X4)|v1_xboole_0(X1)|~r2_hidden(X3,k1_relset_1(X4,k1_mesfunc2(X1,X2)))|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X3,X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_6536, plain, (k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X1,X2),X3),np__0),np__0)=k1_extreal2(k1_extreal2(k1_extreal2(k12_supinf_2(X2,X3),np__0),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X1)|~r2_hidden(X3,k1_relset_1(X4,k1_mesfunc2(X1,X2)))|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X3,X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_6430, plain, (k12_supinf_2(k1_mesfunc2(X1,k1_mesfunc2(X2,k1_mesfunc2(X3,X4))),X5)=k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X6,X4),X5),np__0),np__0)|v1_xboole_0(X6)|v1_xboole_0(X3)|v1_xboole_0(X2)|v1_xboole_0(X1)|~r2_hidden(X5,k1_relset_1(X1,k1_mesfunc2(X2,k1_mesfunc2(X3,X4))))|~r2_hidden(X5,k1_relset_1(X2,k1_mesfunc2(X3,X4)))|~r2_hidden(X5,k1_relset_1(X3,X4))|~r2_hidden(X5,k1_relset_1(X6,X4))|~v1_funct_1(X4)|~m1_subset_1(k1_mesfunc2(X2,k1_mesfunc2(X3,X4)),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X6,k7_numbers)))|~m1_subset_1(X5,X1)|~m1_subset_1(X5,X2)|~m1_subset_1(X5,X3)|~m1_subset_1(X5,X6))).#trainneg
cnf(i_0_6517, plain, (r1_xxreal_0(k3_xxreal_3(X1,X2),k3_supinf_2(k3_extreal1(X1),X2))|~r1_xxreal_0(np__0,k3_xxreal_3(X1,X2))|~r1_xxreal_0(np__0,X2)|~m1_subset_1(X2,k7_numbers)|~m1_subset_1(X1,k7_numbers)|~m1_subset_1(k3_extreal1(X1),k7_numbers))).#trainneg
cnf(i_0_5773, plain, (v1_xcmplx_0(k12_supinf_2(np__0,X1))|~v1_membered(k7_numbers)|~v1_funct_1(np__0))).#trainneg
cnf(i_0_603, plain, (k9_xtuple_0(esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))=k9_xtuple_0(X2)|v1_xboole_0(X1)|esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))!=k10_mesfunc1(X1,X2)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_funct_1(esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))))).#trainneg
cnf(i_0_6507, negated_conjecture, (v4_relat_1(k10_mesfunc1(X1,X2),X3)|v1_xboole_0(X1)|v1_xboole_0(X3)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(esk4_0,X3)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_xboole_0(k2_zfmisc_1(X3,k7_numbers))|~v1_xboole_0(k2_zfmisc_1(X1,k7_numbers)))).#trainneg
cnf(i_0_6508, negated_conjecture, (v4_relat_1(k10_mesfunc1(X1,X2),X3)|v1_xboole_0(X1)|v1_xboole_0(X3)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(esk3_0,X3)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_xboole_0(k2_zfmisc_1(X3,k7_numbers))|~v1_xboole_0(k2_zfmisc_1(X1,k7_numbers)))).#trainneg
cnf(i_0_604, plain, (k9_xtuple_0(esk10_2(X1,X2))=k9_xtuple_0(X3)|v1_xboole_0(X2)|esk10_2(X1,X2)!=k10_mesfunc1(X2,X3)|~v5_relat_1(X3,k7_numbers)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v1_funct_1(esk10_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_607, plain, (k9_xtuple_0(k2_mesfunc2(X1,X2))=k9_xtuple_0(X3)|v1_xboole_0(X1)|k2_mesfunc2(X1,X2)!=k10_mesfunc1(X1,X3)|~v5_relat_1(X3,k7_numbers)|~v4_relat_1(X3,X1)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v1_funct_1(k2_mesfunc2(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_6638, plain, (k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X1,X2),X3),np__0),np__0)=k1_extreal2(k1_extreal2(k1_extreal2(k12_supinf_2(X2,X3),np__0),np__0),np__0)|v1_xboole_0(X1)|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_6493, plain, (k7_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4))=k6_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,k10_mesfunc1(X1,X4)))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(k1_relset_1(X1,k10_mesfunc1(X1,X4)),k1_zfmisc_1(k1_relset_1(X1,X4)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k1_relset_1(X1,k10_mesfunc1(X1,X4)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v5_relat_1(X4,k7_numbers)|~v4_relat_1(X4,X1)|~v1_relat_1(X4))).#trainneg
cnf(i_0_6245, plain, (m1_subset_1(esk15_1(X2),k1_zfmisc_1(X1))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_608, plain, (k9_xtuple_0(k1_mesfunc2(X1,X2))=k9_xtuple_0(X3)|v1_xboole_0(X1)|k1_mesfunc2(X1,X2)!=k10_mesfunc1(X1,X3)|~v5_relat_1(X3,k7_numbers)|~v4_relat_1(X3,X1)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v1_funct_1(k1_mesfunc2(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_5775, plain, (v7_ordinal1(k12_supinf_2(np__0,X1))|~v6_membered(k7_numbers)|~v1_funct_1(np__0))).#trainneg
cnf(i_0_5824, plain, (r1_xxreal_0(X1,k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X2))))))))))))))))))))|r1_xxreal_0(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X2))))))))))))))))))),X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k7_numbers))).#trainneg
cnf(i_0_610, plain, (k12_supinf_2(esk15_1(k2_zfmisc_1(X1,k7_numbers)),X2)=k3_extreal1(k12_supinf_2(X3,X2))|v1_xboole_0(X1)|esk15_1(k2_zfmisc_1(X1,k7_numbers))!=k10_mesfunc1(X1,X3)|~r2_hidden(X2,k9_xtuple_0(esk15_1(k2_zfmisc_1(X1,k7_numbers))))|~v5_relat_1(X3,k7_numbers)|~v4_relat_1(X3,X1)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v1_funct_1(esk15_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_938, plain, (r1_xxreal_0(k7_mesfunc5(X1,X2,X3,X4),k3_supinf_2(k3_extreal1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4))),k3_extreal1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)))))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k2_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,X4)),k7_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_xxreal_0(np__0,k7_mesfunc5(X1,X2,X3,X4))|~m1_subset_1(k7_mesfunc5(X1,X2,X3,X4),k7_numbers))).#trainneg
cnf(i_0_5327, plain, (k1_extreal2(k1_extreal2(k1_extreal2(np__0,k2_supinf_2(k12_supinf_2(k1_mesfunc2(X5,X2),X3))),np__0),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X4,k1_mesfunc2(X1,X2)),X3),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X1)|v1_xboole_0(X5)|~r2_hidden(X3,k1_relset_1(X4,k1_mesfunc2(X1,X2)))|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(k1_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(k2_supinf_2(k12_supinf_2(k1_mesfunc2(X5,X2),X3)),k7_numbers)|~m1_subset_1(X3,X4)|~r2_hidden(X3,k1_relset_1(X5,X2))|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X5)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_3297, plain, (k12_supinf_2(k2_mesfunc2(X1,k2_mesfunc2(X2,X3)),X4)=k1_extreal2(k2_supinf_2(k1_extreal2(k2_supinf_2(k12_supinf_2(X3,X4)),np__0)),np__0)|v1_xboole_0(X1)|v1_xboole_0(X2)|~r2_hidden(X4,k1_relset_1(X1,k2_mesfunc2(X2,X3)))|~v1_funct_1(k2_mesfunc2(X2,X3))|~m1_subset_1(k2_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,X1)|~r2_hidden(X4,k1_relset_1(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_4499, plain, (k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X1,k1_mesfunc2(X2,X3)),X4),np__0),np__0)=k1_extreal2(k1_extreal2(k1_extreal2(k2_supinf_2(k1_extreal2(np__0,k12_supinf_2(X3,X4))),np__0),np__0),np__0)|v1_xboole_0(X1)|v1_xboole_0(X5)|v1_xboole_0(X2)|~r2_hidden(X4,k1_relset_1(X5,k1_mesfunc2(X2,X3)))|~r2_hidden(X4,k1_relset_1(X1,k1_mesfunc2(X2,X3)))|~v1_funct_1(k1_mesfunc2(X2,X3))|~m1_subset_1(k1_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(k1_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,X5)|~m1_subset_1(X4,X1)|~r2_hidden(X4,k1_relset_1(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k12_supinf_2(X3,X4),k7_numbers)|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_5211, plain, (k1_extreal2(k1_extreal2(np__0,k12_supinf_2(k2_mesfunc2(X1,X2),X3)),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X4,X2),X3),np__0),np__0)|v1_xboole_0(X1)|v1_xboole_0(X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X4,X2))|~v1_funct_1(X2)|~m1_subset_1(k2_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k12_supinf_2(k2_mesfunc2(X1,X2),X3),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X4))).#trainneg
cnf(i_0_6945, plain, (v1_xcmplx_0(k12_supinf_2(np__0,X1))|~v6_membered(k7_numbers)|~v1_funct_1(np__0))).#trainneg
cnf(i_0_612, plain, (k12_supinf_2(esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))),X2)=k3_extreal1(k12_supinf_2(X3,X2))|v1_xboole_0(X1)|esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))!=k10_mesfunc1(X1,X3)|~r2_hidden(X2,k9_xtuple_0(esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))))|~v5_relat_1(X3,k7_numbers)|~v4_relat_1(X3,X1)|~v1_relat_1(X3)|~v1_funct_1(X3)|~v1_funct_1(esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_613, plain, (k12_supinf_2(esk10_2(X1,X2),X3)=k3_extreal1(k12_supinf_2(X4,X3))|v1_xboole_0(X2)|esk10_2(X1,X2)!=k10_mesfunc1(X2,X4)|~r2_hidden(X3,k9_xtuple_0(esk10_2(X1,X2)))|~v5_relat_1(X4,k7_numbers)|~v4_relat_1(X4,X2)|~v1_relat_1(X4)|~v1_funct_1(X4)|~v1_funct_1(esk10_2(X1,X2))|~m1_subset_1(X3,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_616, plain, (k12_supinf_2(k2_mesfunc2(X1,X2),X3)=k3_extreal1(k12_supinf_2(X4,X3))|v1_xboole_0(X1)|k2_mesfunc2(X1,X2)!=k10_mesfunc1(X1,X4)|~r2_hidden(X3,k9_xtuple_0(k2_mesfunc2(X1,X2)))|~v5_relat_1(X4,k7_numbers)|~v4_relat_1(X4,X1)|~v1_relat_1(X4)|~v1_funct_1(X4)|~v1_funct_1(k2_mesfunc2(X1,X2))|~m1_subset_1(X3,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_7048, plain, (k7_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4))=k6_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_tarski(k9_xtuple_0(X4),k1_relset_1(X1,X4))|~r1_mesfunc1(X1,X2,X4,k9_xtuple_0(X4))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k9_xtuple_0(X4),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v5_relat_1(X4,k7_numbers)|~v4_relat_1(X4,X1)|~v1_relat_1(X4))).#trainneg
cnf(i_0_7049, plain, (k7_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4))=k6_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|~r1_mesfunc1(X1,X2,X4,k9_xtuple_0(X4))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(k9_xtuple_0(X4),k1_zfmisc_1(k1_relset_1(X1,X4)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k9_xtuple_0(X4),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v5_relat_1(X4,k7_numbers)|~v4_relat_1(X4,X1)|~v1_relat_1(X4))).#trainneg
cnf(i_0_7099, plain, (k1_extreal2(k12_supinf_2(k2_mesfunc2(X1,k2_mesfunc2(X2,X3)),X4),np__0)=k1_extreal2(k1_extreal2(k2_supinf_2(k1_extreal2(k2_supinf_2(k12_supinf_2(X3,X4)),np__0)),np__0),np__0)|v1_xboole_0(X1)|v1_xboole_0(X5)|v1_xboole_0(X2)|~r2_hidden(X4,k1_relset_1(X5,k2_mesfunc2(X2,X3)))|~r2_hidden(X4,k1_relset_1(X1,k2_mesfunc2(X2,X3)))|~v1_funct_1(k2_mesfunc2(X2,X3))|~m1_subset_1(k2_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X4,X5)|~m1_subset_1(X4,X1)|~r2_hidden(X4,k1_relset_1(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_7136, plain, (k12_supinf_2(k2_mesfunc2(X1,k2_mesfunc2(X2,X3)),X4)=k1_extreal2(k2_supinf_2(k12_supinf_2(k2_mesfunc2(X5,X3),X4)),np__0)|v1_xboole_0(X2)|v1_xboole_0(X1)|v1_xboole_0(X5)|~r2_hidden(X4,k1_relset_1(X1,k2_mesfunc2(X2,X3)))|~r2_hidden(X4,k1_relset_1(X2,X3))|~v1_funct_1(X3)|~m1_subset_1(k2_mesfunc2(X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,X1)|~m1_subset_1(X4,X2)|~r2_hidden(X4,k1_relset_1(X5,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X4,X5))).#trainneg
cnf(i_0_7242, plain, (k12_supinf_2(k2_mesfunc2(X1,X2),X3)=k3_extreal1(k12_supinf_2(X4,X3))|v1_xboole_0(X1)|k2_mesfunc2(X1,X2)!=k10_mesfunc1(X1,X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~v5_relat_1(X4,k7_numbers)|~v4_relat_1(X4,X1)|~v1_relat_1(X4)|~v1_funct_1(X4)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1)|~v1_relat_1(k2_mesfunc2(X1,X2)))).#trainneg
cnf(i_0_7204, plain, (r1_xxreal_0(np__0,k6_mesfunc5(X1,X2,X3,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|k9_xtuple_0(X5)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k9_xtuple_0(X5))|~r1_mesfunc5(X1,X2,X6,k10_mesfunc1(X1,X5))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(k10_mesfunc1(X1,X5))|~v1_funct_1(X6)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k10_mesfunc1(X1,X5),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v5_relat_1(X5,k7_numbers)|~v4_relat_1(X5,X1)|~v1_relat_1(X5)|~v1_funct_1(X5))).#trainneg
cnf(i_0_7206, plain, (m1_subset_1(k9_xtuple_0(X2),k1_zfmisc_1(X1))|v1_xboole_0(X1)|~v1_funct_1(k10_mesfunc1(X1,X2))|~m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1562, plain, (esk11_1(k1_zfmisc_1(X1))=esk15_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7198, negated_conjecture, (k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k9_xtuple_0(X1)!=k1_relset_1(esk1_0,esk4_0)|~v5_relat_1(X1,k7_numbers)|~v4_relat_1(X1,esk1_0)|~v1_relat_1(X1)|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~v1_xboole_0(k2_zfmisc_1(esk1_0,k7_numbers)))).#trainneg
cnf(i_0_617, plain, (k12_supinf_2(k1_mesfunc2(X1,X2),X3)=k3_extreal1(k12_supinf_2(X4,X3))|v1_xboole_0(X1)|k1_mesfunc2(X1,X2)!=k10_mesfunc1(X1,X4)|~r2_hidden(X3,k9_xtuple_0(k1_mesfunc2(X1,X2)))|~v5_relat_1(X4,k7_numbers)|~v4_relat_1(X4,X1)|~v1_relat_1(X4)|~v1_funct_1(X4)|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(X3,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_6967, plain, (r1_xxreal_0(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X1)))))))))))))))))))),X2)|r1_xxreal_0(X2,k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X1)))))))))))))))))))))|~v1_xxreal_0(X2)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_5486, plain, (m1_subset_1(k10_mesfunc1(X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|v1_xboole_0(X1)|v1_xboole_0(X3)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v5_relat_1(X4,k7_numbers)|~v4_relat_1(X4,X3)|~v1_relat_1(X4)|~v1_funct_1(X4)|~v1_xboole_0(k2_zfmisc_1(X1,k7_numbers))|~v1_xboole_0(k2_zfmisc_1(X3,k7_numbers)))).#trainneg
cnf(i_0_618, plain, (k12_supinf_2(k10_mesfunc1(X1,X2),X3)=k3_extreal1(k12_supinf_2(X4,X3))|v1_xboole_0(X1)|k10_mesfunc1(X1,X2)!=k10_mesfunc1(X1,X4)|~r2_hidden(X3,k9_xtuple_0(k10_mesfunc1(X1,X2)))|~v5_relat_1(X4,k7_numbers)|~v4_relat_1(X4,X1)|~v1_relat_1(X4)|~v1_funct_1(X4)|~v1_funct_1(k10_mesfunc1(X1,X2))|~m1_subset_1(X3,X1)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_7462, plain, (k12_supinf_2(k1_mesfunc2(X1,X2),X3)=k3_extreal1(k12_supinf_2(X4,X3))|v1_xboole_0(X1)|k1_mesfunc2(X1,X2)!=k10_mesfunc1(X1,X4)|~r2_hidden(X3,k1_relset_1(X1,k1_mesfunc2(X1,X2)))|~v5_relat_1(X4,k7_numbers)|~v4_relat_1(X4,X1)|~v1_relat_1(X4)|~v1_funct_1(X4)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_7265, negated_conjecture, (k7_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,esk4_0))=k6_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,esk4_0))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_mesfunc1(X1,X2,esk4_0,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~m1_subset_1(k1_relset_1(esk1_0,esk4_0),k1_zfmisc_1(k1_relset_1(X1,esk4_0)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k1_relset_1(esk1_0,esk4_0),X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_7255, negated_conjecture, (k7_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,esk3_0))=k6_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,esk3_0))|v1_xboole_0(X2)|v1_xboole_0(X1)|~r1_tarski(k1_relset_1(esk2_0,esk3_0),k1_relset_1(X1,esk3_0))|~r1_mesfunc1(X1,X2,esk3_0,k1_relset_1(esk2_0,esk3_0))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k1_relset_1(esk2_0,esk3_0),X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_7254, negated_conjecture, (k7_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,esk4_0))=k6_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,esk4_0))|v1_xboole_0(X2)|v1_xboole_0(X1)|~r1_tarski(k1_relset_1(esk1_0,esk4_0),k1_relset_1(X1,esk4_0))|~r1_mesfunc1(X1,X2,esk4_0,k1_relset_1(esk1_0,esk4_0))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k1_relset_1(esk1_0,esk4_0),X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_7199, plain, (k7_mesfunc5(X1,X2,X3,X4)=k6_mesfunc5(X1,X2,X3,X4)|v1_xboole_0(X1)|v1_xboole_0(X2)|k9_xtuple_0(X5)!=k1_relset_1(X1,X4)|~r1_mesfunc1(X1,X2,X4,k9_xtuple_0(X5))|~r1_mesfunc5(X1,X2,X6,k10_mesfunc1(X1,X5))|~v4_measure1(X3,X1,X2)|~v4_measure1(X6,X1,X2)|~v6_supinf_2(X4)|~v6_supinf_2(X3)|~v6_supinf_2(X6)|~v10_valued_0(X3)|~v10_valued_0(X6)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_2(X6,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(k10_mesfunc1(X1,X5))|~v1_funct_1(X6)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k10_mesfunc1(X1,X5),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v5_relat_1(X5,k7_numbers)|~v4_relat_1(X5,X1)|~v1_relat_1(X5)|~v1_funct_1(X5))).#trainneg
cnf(i_0_7231, plain, (k12_supinf_2(k10_mesfunc1(X1,X2),X3)=k1_extreal2(k12_supinf_2(X4,X3),np__0)|v1_xboole_0(X1)|k10_mesfunc1(X1,X2)!=k1_mesfunc2(X1,X4)|~r2_hidden(X3,k9_xtuple_0(X2))|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X4)|~v1_funct_1(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_7266, negated_conjecture, (k7_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,esk3_0))=k6_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,esk3_0))|v1_xboole_0(X1)|v1_xboole_0(X2)|~r1_mesfunc1(X1,X2,esk3_0,k1_relset_1(esk2_0,esk3_0))|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~m1_subset_1(k1_relset_1(esk2_0,esk3_0),k1_zfmisc_1(k1_relset_1(X1,esk3_0)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k1_relset_1(esk2_0,esk3_0),X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_7516, negated_conjecture, (~r1_xxreal_0(k3_extreal1(k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0)),k6_mesfunc5(esk1_0,esk2_0,esk3_0,k10_mesfunc1(esk1_0,esk4_0))))).#trainneg
cnf(i_0_7515, negated_conjecture, (~r1_xxreal_0(k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0),k6_mesfunc5(esk1_0,esk2_0,esk3_0,k10_mesfunc1(esk1_0,esk4_0)))|~r1_xxreal_0(np__0,k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_subset_1(k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0),k7_numbers))).#trainneg
cnf(i_0_7514, negated_conjecture, (~r1_xxreal_0(k3_extreal1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0)),k6_mesfunc5(esk1_0,esk2_0,esk3_0,k10_mesfunc1(esk1_0,esk4_0)))|~v6_supinf_2(esk4_0))).#trainneg
cnf(i_0_7461, negated_conjecture, (k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k9_xtuple_0(X1)!=k1_relset_1(esk1_0,esk4_0)|~v5_relat_1(X1,k7_numbers)|~v4_relat_1(X1,esk1_0)|~v1_relat_1(X1)|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~v1_xboole_0(k2_zfmisc_1(esk1_0,k7_numbers)))).#trainneg
cnf(i_0_7493, negated_conjecture, (m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|v1_xboole_0(X1)|v1_xboole_0(X3)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(esk4_0,X3)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_xboole_0(k2_zfmisc_1(X3,k7_numbers))|~v1_xboole_0(k2_zfmisc_1(X1,k7_numbers)))).#trainneg
cnf(i_0_7494, negated_conjecture, (m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|v1_xboole_0(X1)|v1_xboole_0(X3)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(esk3_0,X3)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_xboole_0(k2_zfmisc_1(X3,k7_numbers))|~v1_xboole_0(k2_zfmisc_1(X1,k7_numbers)))).#trainneg
cnf(i_0_621, plain, (v1_xxreal_0(k3_mesfunc1(X1,X2,X3))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_622, plain, (v1_xcmplx_0(k3_mesfunc1(X1,X2,X3))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_7513, negated_conjecture, (~r1_xxreal_0(k6_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0),k6_mesfunc5(esk1_0,esk2_0,esk3_0,k10_mesfunc1(esk1_0,esk4_0)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0))|~v6_supinf_2(esk4_0)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0),k7_numbers))).#trainneg
cnf(i_0_7508, plain, (k12_supinf_2(k1_mesfunc2(X1,X2),X3)=k3_extreal1(k12_supinf_2(X4,X3))|v1_xboole_0(X1)|k1_mesfunc2(X1,X2)!=k10_mesfunc1(X1,X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~v5_relat_1(X4,k7_numbers)|~v4_relat_1(X4,X1)|~v1_relat_1(X4)|~v1_funct_1(X4)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_7241, plain, (k12_supinf_2(k2_mesfunc2(X1,X2),X3)=k3_extreal1(k12_supinf_2(X4,X3))|v1_xboole_0(X1)|k2_mesfunc2(X1,X2)!=k10_mesfunc1(X1,X4)|~r2_hidden(X3,k1_relset_1(X1,k2_mesfunc2(X1,X2)))|~v5_relat_1(X4,k7_numbers)|~v4_relat_1(X4,X1)|~v1_relat_1(X4)|~v1_funct_1(X4)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1)|~v1_relat_1(k2_mesfunc2(X1,X2)))).#trainneg
cnf(i_0_7228, plain, (k7_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4))=k6_mesfunc5(X1,X2,X3,k10_mesfunc1(X1,X4))|v1_xboole_0(X1)|v1_xboole_0(X2)|X5!=k9_xtuple_0(X4)|~r1_tarski(X5,k1_relset_1(X1,X4))|~r1_mesfunc1(X1,X2,X4,X5)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(k10_mesfunc1(X1,X4))|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(k10_mesfunc1(X1,X4),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X5,X2)|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~v5_relat_1(X4,k7_numbers)|~v4_relat_1(X4,X1)|~v1_relat_1(X4))).#trainneg
cnf(i_0_7222, plain, (k12_supinf_2(k10_mesfunc1(X1,X2),X3)=k1_extreal2(k2_supinf_2(k12_supinf_2(X4,X3)),np__0)|v1_xboole_0(X1)|k10_mesfunc1(X1,X2)!=k2_mesfunc2(X1,X4)|~r2_hidden(X3,k9_xtuple_0(X2))|~v1_funct_1(X4)|~v1_funct_1(k10_mesfunc1(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1442, negated_conjecture, (r1_xxreal_0(k1_supinf_1,k12_supinf_2(esk3_0,X1))|r1_xxreal_0(k12_supinf_2(esk3_0,X1),k1_supinf_1))).#trainneg
cnf(i_0_623, plain, (v1_xreal_0(k3_mesfunc1(X1,X2,X3))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_7422, plain, (~v1_subset_1(esk11_1(k1_zfmisc_1(X2)),X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_7487, plain, (r1_xxreal_0(X1,k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X2))))))))))))))))))))))|r1_xxreal_0(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(k3_extreal1(X2))))))))))))))))))))),X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k7_numbers))).#trainneg
cnf(i_0_624, plain, (v7_ordinal1(k3_mesfunc1(X1,X2,X3))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_626, plain, (v1_xboole_0(k3_mesfunc1(X1,X2,X3))|v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,k7_numbers))|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_7640, negated_conjecture, (~r1_xxreal_0(np__0,k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0))|~m1_subset_1(k7_mesfunc5(esk1_0,esk2_0,esk3_0,esk4_0),k7_numbers)|k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_7636, negated_conjecture, (k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers))))).#trainneg
cnf(i_0_1476, negated_conjecture, (r1_xxreal_0(k1_supinf_1,k12_supinf_2(esk4_0,X1))|r1_xxreal_0(k12_supinf_2(esk4_0,X1),k1_supinf_1))).#trainneg
cnf(i_0_7644, negated_conjecture, (k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k9_xtuple_0(X1)!=k1_relset_1(esk1_0,esk4_0)|~v5_relat_1(X1,k7_numbers)|~v4_relat_1(X1,esk1_0)|~v1_relat_1(X1)|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(X1)|~v1_xboole_0(k2_zfmisc_1(esk1_0,k7_numbers)))).#trainneg
cnf(i_0_627, plain, (v3_membered(k3_mesfunc1(X1,X2,X3))|v1_xboole_0(X1)|~v3_membered(k2_zfmisc_1(X1,k7_numbers))|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_628, plain, (v1_valued_0(k3_mesfunc1(X1,X2,X3))|v1_xboole_0(X1)|~v1_valued_0(k2_zfmisc_1(X1,k7_numbers))|~v1_relat_1(k2_zfmisc_1(X1,k7_numbers))|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_629, plain, (v3_valued_0(k3_mesfunc1(X1,X2,X3))|v1_xboole_0(X1)|~v3_valued_0(k2_zfmisc_1(X1,k7_numbers))|~v1_relat_1(k2_zfmisc_1(X1,k7_numbers))|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1535, plain, (k7_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,k2_mesfunc2(X1,X4)))=k6_mesfunc5(X1,X2,X3,k1_mesfunc2(X1,k2_mesfunc2(X1,X4)))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X3)|~v1_funct_1(k2_mesfunc2(X1,X4))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(k2_mesfunc2(X1,X4),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k1_relset_1(X1,k1_mesfunc2(X1,k2_mesfunc2(X1,X4))),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1)|~r1_tarski(k1_relset_1(X1,k1_mesfunc2(X1,k2_mesfunc2(X1,X4))),k1_relset_1(X1,X4))|~r1_mesfunc1(X1,X2,X4,k1_relset_1(X1,k1_mesfunc2(X1,k2_mesfunc2(X1,X4))))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_6125, plain, (k1_extreal2(k1_extreal2(k1_extreal2(np__0,k12_supinf_2(k2_mesfunc2(X1,X2),X3)),np__0),np__0)=k1_extreal2(k1_extreal2(k1_extreal2(k1_extreal2(k2_supinf_2(k12_supinf_2(X2,X3)),np__0),np__0),np__0),np__0)|v1_xboole_0(X1)|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k2_mesfunc2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(k12_supinf_2(k2_mesfunc2(X1,X2),X3),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_2693, plain, (r1_xxreal_0(k12_supinf_2(X1,X2),np__0)|r1_xxreal_0(np__0,k12_supinf_2(X1,X2))|~v5_relat_1(X1,k7_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_630, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,k7_numbers))|v1_xboole_0(X2)|~r2_hidden(X1,k3_mesfunc1(X2,X3,X4))|~v5_relat_1(X4,k7_numbers)|~v5_relat_1(X3,k7_numbers)|~v4_relat_1(X4,X2)|~v4_relat_1(X3,X2)|~v1_relat_1(X4)|~v1_relat_1(X3)|~v1_funct_1(X4)|~v1_funct_1(X3))).#trainneg
cnf(i_0_631, plain, (k1_relset_1(X1,k3_mesfunc1(X1,X2,X3))=k1_relset_1(X1,X4)|v1_xboole_0(X1)|k3_mesfunc1(X1,X2,X3)!=k1_mesfunc2(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(k3_mesfunc1(X1,X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_632, plain, (k1_relset_1(X1,k3_mesfunc1(X1,X2,X3))=k1_relset_1(X1,X4)|v1_xboole_0(X1)|k3_mesfunc1(X1,X2,X3)!=k2_mesfunc2(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(k3_mesfunc1(X1,X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1516, plain, (r1_xxreal_0(k7_mesfunc5(X1,X2,X3,X4),X5)|r1_xxreal_0(X5,k7_mesfunc5(X1,X2,X3,X4))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xxreal_0(X5)|~v4_measure1(X3,X1,X2)|~v6_supinf_2(X3)|~v10_valued_0(X3)|~v1_funct_2(X3,X2,k7_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v4_prob_1(X2,X1)|~v1_prob_1(X2,X1))).#trainneg
cnf(i_0_7892, plain, (k1_relset_1(X1,k10_mesfunc1(X1,X2))=k1_relset_1(X1,X3)|v1_xboole_0(X1)|k10_mesfunc1(X1,X2)!=k2_mesfunc2(X1,X3)|~v5_relat_1(k2_mesfunc2(X1,X2),k7_numbers)|~v5_relat_1(k1_mesfunc2(X1,X2),k7_numbers)|~v4_relat_1(k2_mesfunc2(X1,X2),X1)|~v4_relat_1(k1_mesfunc2(X1,X2),X1)|~v1_relat_1(k2_mesfunc2(X1,X2))|~v1_relat_1(k1_mesfunc2(X1,X2))|~v1_funct_1(X3)|~v1_funct_1(k2_mesfunc2(X1,X2))|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_7889, plain, (k1_relset_1(X1,k10_mesfunc1(X1,X2))=k1_relset_1(X1,X3)|v1_xboole_0(X1)|k10_mesfunc1(X1,X2)!=k1_mesfunc2(X1,X3)|~v5_relat_1(k2_mesfunc2(X1,X2),k7_numbers)|~v5_relat_1(k1_mesfunc2(X1,X2),k7_numbers)|~v4_relat_1(k2_mesfunc2(X1,X2),X1)|~v4_relat_1(k1_mesfunc2(X1,X2),X1)|~v1_relat_1(k2_mesfunc2(X1,X2))|~v1_relat_1(k1_mesfunc2(X1,X2))|~v1_funct_1(X3)|~v1_funct_1(k2_mesfunc2(X1,X2))|~v1_funct_1(k1_mesfunc2(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~v1_funct_1(X2)|~m1_subset_1(k10_mesfunc1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers))))).#trainneg
cnf(i_0_7851, plain, (k1_extreal2(k1_extreal2(k1_extreal2(np__0,k12_supinf_2(k2_mesfunc2(X1,X2),X3)),np__0),np__0)=k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X4,k2_mesfunc2(X5,X2)),X3),np__0),np__0)|v1_xboole_0(X1)|v1_xboole_0(X5)|v1_xboole_0(X4)|~r2_hidden(X3,k1_relset_1(X1,X2))|~v1_funct_1(X2)|~m1_subset_1(k12_supinf_2(k2_mesfunc2(X1,X2),X3),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X3,X1)|~r2_hidden(X3,k1_relset_1(X4,k2_mesfunc2(X5,X2)))|~r2_hidden(X3,k1_relset_1(X5,X2))|~m1_subset_1(k2_mesfunc2(X5,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X3,X4)|~m1_subset_1(X3,X5))).#trainneg
cnf(i_0_7784, negated_conjecture, (k1_relset_1(esk1_0,k2_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|~r1_mesfunc1(esk1_0,esk2_0,k2_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_mesfunc1(esk1_0,esk2_0,k1_mesfunc2(esk1_0,esk4_0),k1_relset_1(esk1_0,esk4_0))|~r1_xxreal_0(np__0,k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)))|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k2_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k6_mesfunc5(esk1_0,esk2_0,esk3_0,k1_mesfunc2(esk1_0,esk4_0)),k7_numbers)|~m1_subset_1(k2_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k1_mesfunc2(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~m1_subset_1(k10_mesfunc1(esk1_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(esk1_0,k7_numbers)))|~v6_supinf_2(k2_mesfunc2(esk1_0,esk4_0)))).#trainneg
cnf(i_0_7421, plain, (m1_subset_1(esk11_1(k1_zfmisc_1(X2)),k1_zfmisc_1(X1))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_7785, negated_conjecture, (k1_relset_1(esk1_0,k1_mesfunc2(esk1_0,esk4_0))!=k1_relset_1(esk1_0,esk4_0)|k9_xtuple_0(X1)!=k1_relset_1(esk1_0,esk4_0)|~v5_relat_1(X1,k7_numbers)|~v4_relat_1(X1,esk1_0)|~v1_relat_1(X1)|~v1_funct_1(k2_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(k1_mesfunc2(esk1_0,esk4_0))|~v1_funct_1(X1)|~v1_xboole_0(k2_zfmisc_1(esk1_0,k7_numbers)))).#trainneg
cnf(i_0_634, plain, (k9_xtuple_0(k3_mesfunc1(X1,X2,X3))=k9_xtuple_0(X4)|v1_xboole_0(X1)|k3_mesfunc1(X1,X2,X3)!=k10_mesfunc1(X1,X4)|~v5_relat_1(X4,k7_numbers)|~v4_relat_1(X4,X1)|~v1_relat_1(X4)|~v1_funct_1(X4)|~v1_funct_1(k3_mesfunc1(X1,X2,X3))|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1314, plain, (v1_xcmplx_0(esk14_1(X1))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_635, plain, (v2_valued_0(k3_mesfunc1(X1,X2,X3))|v1_xboole_0(X1)|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_636, plain, (v1_valued_0(k3_mesfunc1(X1,X2,X3))|v1_xboole_0(X1)|~v1_membered(k7_numbers)|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_3588, plain, (k1_extreal2(k12_supinf_2(k1_mesfunc2(X3,k2_mesfunc2(X4,X1)),X2),np__0)=k1_extreal2(np__0,k1_extreal2(k1_extreal2(k2_supinf_2(k12_supinf_2(X1,X2)),np__0),np__0))|v1_xboole_0(X4)|v1_xboole_0(X3)|~m1_subset_1(k1_extreal2(k1_extreal2(k2_supinf_2(k12_supinf_2(X1,X2)),np__0),np__0),k7_numbers)|~r2_hidden(X2,k1_relset_1(X3,k2_mesfunc2(X4,X1)))|~r2_hidden(X2,k1_relset_1(X4,X1))|~v1_funct_1(X1)|~m1_subset_1(k2_mesfunc2(X4,X1),k1_zfmisc_1(k2_zfmisc_1(X3,k7_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X2,X3)|~m1_subset_1(X2,X4))).#trainneg
cnf(i_0_5170, plain, (m1_subset_1(k9_xtuple_0(esk10_2(X2,X1)),k1_zfmisc_1(X1))|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v4_prob_1(X1,X2)|~v1_prob_1(X1,X2))).#trainneg
cnf(i_0_1573, negated_conjecture, (r1_xxreal_0(np__0,k12_supinf_2(esk4_0,X1))|r1_xxreal_0(k12_supinf_2(esk4_0,X1),np__0))).#trainneg
cnf(i_0_1373, plain, (v1_xcmplx_0(esk14_1(X1))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_638, plain, (v5_relat_1(k3_mesfunc1(X1,X2,X3),k7_numbers)|v1_xboole_0(X1)|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_639, plain, (v1_relat_1(k3_mesfunc1(X1,X2,X3))|v1_xboole_0(X1)|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_7979, plain, (k1_extreal2(k1_extreal2(k1_extreal2(np__0,k12_supinf_2(k2_mesfunc2(X1,X2),X3)),np__0),np__0)=k1_extreal2(k12_supinf_2(k1_mesfunc2(X6,k1_mesfunc2(X4,k2_mesfunc2(X5,X2))),X3),np__0)|v1_xboole_0(X4)|v1_xboole_0(X5)|v1_xboole_0(X1)|v1_xboole_0(X6)|~r2_hidden(X3,k1_relset_1(X4,k2_mesfunc2(X5,X2)))|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X5,X2))|~v1_funct_1(X2)|~m1_subset_1(k2_mesfunc2(X5,X2),k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(k12_supinf_2(k2_mesfunc2(X1,X2),X3),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X4)|~m1_subset_1(X3,X5)|~r2_hidden(X3,k1_relset_1(X6,k1_mesfunc2(X4,k2_mesfunc2(X5,X2))))|~v1_funct_1(k1_mesfunc2(X4,k2_mesfunc2(X5,X2)))|~m1_subset_1(k1_mesfunc2(X4,k2_mesfunc2(X5,X2)),k1_zfmisc_1(k2_zfmisc_1(X6,k7_numbers)))|~m1_subset_1(X3,X6))).#trainneg
cnf(i_0_1449, plain, (v1_xcmplx_0(k2_supinf_2(X1))|~v6_membered(k7_numbers)|~m1_subset_1(X1,k7_numbers))).#trainneg
cnf(i_0_640, plain, (v4_relat_1(k3_mesfunc1(X1,X2,X3),X1)|v1_xboole_0(X1)|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1106, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(esk1_0))|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_641, plain, (k12_supinf_2(k3_mesfunc1(X1,X2,X3),X4)=k3_extreal1(k12_supinf_2(X5,X4))|v1_xboole_0(X1)|k3_mesfunc1(X1,X2,X3)!=k10_mesfunc1(X1,X5)|~r2_hidden(X4,k9_xtuple_0(k3_mesfunc1(X1,X2,X3)))|~v5_relat_1(X5,k7_numbers)|~v4_relat_1(X5,X1)|~v1_relat_1(X5)|~v1_funct_1(X5)|~v1_funct_1(k3_mesfunc1(X1,X2,X3))|~m1_subset_1(X4,X1)|~v5_relat_1(X3,k7_numbers)|~v5_relat_1(X2,k7_numbers)|~v4_relat_1(X3,X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2))).#trainneg
cnf(i_0_7965, plain, (k1_extreal2(k1_extreal2(k12_supinf_2(k1_mesfunc2(X5,k2_mesfunc2(X6,X2)),X3),np__0),np__0)=k1_extreal2(k1_extreal2(k1_extreal2(k12_supinf_2(k2_mesfunc2(X4,X2),X3),np__0),np__0),np__0)|v1_xboole_0(X4)|v1_xboole_0(X1)|v1_xboole_0(X5)|v1_xboole_0(X6)|~r2_hidden(X3,k1_relset_1(X1,X2))|~r2_hidden(X3,k1_relset_1(X4,X2))|~v1_funct_1(X2)|~m1_subset_1(k12_supinf_2(k2_mesfunc2(X1,X2),X3),k7_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,k7_numbers)))|~m1_subset_1(X3,X1)|~m1_subset_1(X3,X4)|~r2_hidden(X3,k1_relset_1(X5,k2_mesfunc2(X6,X2)))|~r2_hidden(X3,k1_relset_1(X6,X2))|~m1_subset_1(k2_mesfunc2(X6,X2),k1_zfmisc_1(k2_zfmisc_1(X5,k7_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X6,k7_numbers)))|~m1_subset_1(X3,X5)|~m1_subset_1(X3,X6))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 156
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 230
# Removed in clause preprocessing      : 20
# Initial clauses in saturation        : 210
# Processed clauses                    : 1754
# ...of these trivial                  : 19
# ...subsumed                          : 608
# ...remaining for further processing  : 1127
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 78
# Backward-rewritten                   : 74
# Generated clauses                    : 8019
# ...of the previous two non-trivial   : 7874
# Contextual simplify-reflections      : 345
# Paramodulations                      : 8004
# Factorizations                       : 0
# Equation resolutions                 : 17
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 971
#    Positive orientable unit clauses  : 96
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 14
#    Non-unit-clauses                  : 861
# Current number of unprocessed clauses: 6057
# ...number of literals in the above   : 93018
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1279
# Proof object given clauses           : 132
# Proof search given clauses           : 1127
# Clause-clause subsumption calls (NU) : 1337934
# Rec. Clause-clause subsumption calls : 35059
# Non-unit clause-clause subsumptions  : 987
# Unit Clause-clause subsumption calls : 10379
# Rewrite failures with RHS unbound    : 43
# BW rewrite match attempts            : 69
# BW rewrite match successes           : 32
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 466077

# -------------------------------------------------
# User time                : 34.050 s
# System time              : 1.101 s
# Total time               : 35.151 s
# Maximum resident set size: 370192 pages
34.05user 1.17system 0:36.01elapsed 97%CPU (0avgtext+0avgdata 370192maxresident)k
0inputs+936outputs (0major+383394minor)pagefaults 0swaps
