2020-05-02 00:54:32.414303: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_8/premsel_enigma_01_2020_T10_loop01_epoch_43
2020-05-02 00:54:32.476302: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-05-02 00:54:32.612003: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-05-02 00:54:33.009179: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-05-02 00:54:34.668193: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 2253752 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_8/premsel_enigma_01_2020_T10_loop01_epoch_43' loaded (query=512; context=1024; weigths=real; len_mult=0.000000).
# Preprocessing time       : 2.776 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(cc14_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xreal_0(X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', cc14_membered)).
fof(dt_k10_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>m1_subset_1(k10_real_1(X1,X2),k1_numbers)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', dt_k10_real_1)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', t6_boole)).
fof(redefinition_k10_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', redefinition_k10_real_1)).
fof(involutiveness_k5_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>k5_xcmplx_0(k5_xcmplx_0(X1))=X1), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', involutiveness_k5_xcmplx_0)).
fof(spc3_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k7_xcmplx_0(np__1,X1)=k5_xcmplx_0(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', spc3_arithm)).
fof(dt_k5_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>v1_xcmplx_0(k5_xcmplx_0(X1))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', dt_k5_xcmplx_0)).
fof(fc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>(v1_xcmplx_0(k4_xcmplx_0(X1))&v1_xreal_0(k4_xcmplx_0(X1)))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', fc3_xreal_0)).
fof(rqRealNeg__k4_xcmplx_0__rnm1d2_rn1d2, axiom, k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k7_xcmplx_0(np__1,np__2), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', rqRealNeg__k4_xcmplx_0__rnm1d2_rn1d2)).
fof(rqRealDiv__k7_xcmplx_0__r1_rm2_rnm1d2, axiom, k7_xcmplx_0(np__1,k4_xcmplx_0(np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', rqRealDiv__k7_xcmplx_0__r1_rm2_rnm1d2)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', fc3_membered)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', redefinition_k6_numbers)).
fof(t1_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k2_xcmplx_0(X1,k6_numbers)=X1), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', t1_arithm)).
fof(spc0_boole, axiom, v1_xboole_0(np__0), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', spc0_boole)).
fof(reflexivity_r1_xxreal_0, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>r1_xxreal_0(X1,X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', reflexivity_r1_xxreal_0)).
fof(rc4_xxreal_0, axiom, ?[X1]:(v1_xboole_0(X1)&v1_xxreal_0(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', rc4_xxreal_0)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', spc1_numerals)).
fof(fc4_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>(v1_xcmplx_0(k5_xcmplx_0(X1))&v1_xreal_0(k5_xcmplx_0(X1)))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', fc4_xreal_0)).
fof(cc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xcmplx_0(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', cc3_xreal_0)).
fof(rqRealNeg__k4_xcmplx_0__rn1d2_rnm1d2, axiom, k4_xcmplx_0(k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', rqRealNeg__k4_xcmplx_0__rn1d2_rnm1d2)).
fof(spc2_numerals, axiom, (((v2_xxreal_0(np__2)&m2_subset_1(np__2,k1_numbers,k5_numbers))&m1_subset_1(np__2,k5_numbers))&m1_subset_1(np__2,k1_numbers)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', spc2_numerals)).
fof(cc17_membered, axiom, ![X1]:(v6_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v7_ordinal1(X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', cc17_membered)).
fof(fc6_membered, axiom, v6_membered(k4_ordinal1), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', fc6_membered)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', redefinition_k5_numbers)).
fof(t8_xreal_1, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>![X3]:(v1_xreal_0(X3)=>![X4]:(v1_xreal_0(X4)=>~(((~(r1_xxreal_0(X2,X1))&r1_xxreal_0(X3,X4))&r1_xxreal_0(k2_xcmplx_0(X2,X4),k2_xcmplx_0(X1,X3)))))))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', t8_xreal_1)).
fof(spc0_numerals, axiom, ((m2_subset_1(np__0,k1_numbers,k5_numbers)&m1_subset_1(np__0,k5_numbers))&m1_subset_1(np__0,k1_numbers)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', spc0_numerals)).
fof(cc4_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xxreal_0(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', cc4_xreal_0)).
fof(fc5_xreal_0, axiom, ![X1, X2]:((v1_xreal_0(X1)&v1_xreal_0(X2))=>v1_xreal_0(k2_xcmplx_0(X1,X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', fc5_xreal_0)).
fof(cc3_int_1, axiom, ![X1]:(v1_int_1(X1)=>v1_xreal_0(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', cc3_int_1)).
fof(fc3_int_1, axiom, ![X1]:(v1_int_1(X1)=>(v1_xcmplx_0(k4_xcmplx_0(X1))&v1_int_1(k4_xcmplx_0(X1)))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', fc3_int_1)).
fof(cc2_int_1, axiom, ![X1]:(v7_ordinal1(X1)=>v1_int_1(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', cc2_int_1)).
fof(rqRealAdd__k2_xcmplx_0__rnm1d2_r2_rn3d2, axiom, k2_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),np__2)=k7_xcmplx_0(np__3,np__2), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', rqRealAdd__k2_xcmplx_0__rnm1d2_r2_rn3d2)).
fof(rqRealAdd__k2_xcmplx_0__r1_rn1d2_rn3d2, axiom, k2_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(np__3,np__2), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', rqRealAdd__k2_xcmplx_0__r1_rn1d2_rn3d2)).
fof(rqLessOrEqual__r1_xxreal_0__r0_rn1d2, axiom, r1_xxreal_0(np__0,k7_xcmplx_0(np__1,np__2)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', rqLessOrEqual__r1_xxreal_0__r0_rn1d2)).
fof(dt_k1_fdiff_1, axiom, ![X1, X2]:(((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))&v1_xreal_0(X2))=>m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', dt_k1_fdiff_1)).
fof(t65_sin_cos, axiom, ![X1]:(v1_xreal_0(X1)=>(r1_fdiff_1(k24_sin_cos,X1)&k1_fdiff_1(k24_sin_cos,X1)=k1_seq_1(k24_sin_cos,X1))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', t65_sin_cos)).
fof(cc4_membered, axiom, ![X1]:(v3_membered(X1)=>v2_membered(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', cc4_membered)).
fof(t9_fdiff_1, axiom, ![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(r2_fdiff_1(X2,X1)<=>(r1_tarski(X1,k1_relset_1(k1_numbers,X2))&![X3]:(m1_subset_1(X3,k1_numbers)=>(r2_hidden(X3,X1)=>r1_fdiff_1(X2,X3))))))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', t9_fdiff_1)).
fof(t85_sin_cos9, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(r1_fdiff_1(X2,X1)=>(r1_xxreal_0(k1_seq_1(X2,X1),k1_real_1(np__1))|r1_xxreal_0(np__1,k1_seq_1(X2,X1))|(r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k1_sin_cos9),X1)&k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k1_sin_cos9),X1)=k10_real_1(k1_fdiff_1(X2,X1),k7_real_1(np__1,k5_square_1(k1_seq_1(X2,X1))))))))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', t85_sin_cos9)).
fof(connectedness_r1_xxreal_0, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>(r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', connectedness_r1_xxreal_0)).
fof(cc13_membered, axiom, ![X1]:(v2_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xxreal_0(X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', cc13_membered)).
fof(dt_k24_sin_cos, axiom, ((v1_funct_1(k24_sin_cos)&v1_funct_2(k24_sin_cos,k1_numbers,k1_numbers))&m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', dt_k24_sin_cos)).
fof(rqRealAdd__k2_xcmplx_0__r2_rm1_r1, axiom, k2_xcmplx_0(np__2,k4_xcmplx_0(np__1))=np__1, file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', rqRealAdd__k2_xcmplx_0__r2_rm1_r1)).
fof(rqLessOrEqual__r1_xxreal_0__r2_rn3d2, axiom, ~(r1_xxreal_0(np__2,k7_xcmplx_0(np__3,np__2))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', rqLessOrEqual__r1_xxreal_0__r2_rn3d2)).
fof(redefinition_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>k1_real_1(X1)=k4_xcmplx_0(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', redefinition_k1_real_1)).
fof(t54_sin_cos, axiom, ![X1]:(v1_xreal_0(X1)=>~(r1_xxreal_0(k1_seq_1(k24_sin_cos,X1),k6_numbers))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', t54_sin_cos)).
fof(t115_sin_cos9, conjecture, ![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>((r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9)))&![X2]:(m1_subset_1(X2,k1_numbers)=>~((r2_hidden(X2,X1)&r1_xxreal_0(np__1,k1_seq_1(k24_sin_cos,X2))))))=>(r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),X1)&![X2]:(m1_subset_1(X2,k1_numbers)=>(r2_hidden(X2,X1)=>k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),X1),X2)=k10_real_1(k1_seq_1(k24_sin_cos,X2),k7_real_1(np__1,k5_square_1(k1_seq_1(k24_sin_cos,X2))))))))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', t115_sin_cos9)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', t1_subset)).
fof(cc1_xreal_0, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>v1_xreal_0(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', cc1_xreal_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', t4_subset)).
fof(cc21_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v3_membered(X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', cc21_membered)).
fof(dt_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:((((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_funct_1(X6))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>(v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))&m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4))))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', dt_k1_partfun1)).
fof(redefinition_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:((((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_funct_1(X6))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>k1_partfun1(X1,X2,X3,X4,X5,X6)=k3_relat_1(X5,X6)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', redefinition_k1_partfun1)).
fof(d7_fdiff_1, axiom, ![X1]:((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>![X2]:(r2_fdiff_1(X1,X2)=>![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(X3=k2_fdiff_1(X1,X2)<=>(k1_relset_1(k1_numbers,X3)=X2&![X4]:(m1_subset_1(X4,k1_numbers)=>(r2_hidden(X4,X2)=>k1_seq_1(X3,X4)=k1_fdiff_1(X1,X4)))))))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', d7_fdiff_1)).
fof(dt_k2_fdiff_1, axiom, ![X1, X2]:((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(v1_funct_1(k2_fdiff_1(X1,X2))&m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', dt_k2_fdiff_1)).
fof(commutativity_k2_xcmplx_0, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', commutativity_k2_xcmplx_0)).
fof(dt_k1_sin_cos9, axiom, (v1_funct_1(k1_sin_cos9)&m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', dt_k1_sin_cos9)).
fof(rqRealNeg__k4_xcmplx_0__r0_r0, axiom, k4_xcmplx_0(np__0)=np__0, file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', rqRealNeg__k4_xcmplx_0__r0_r0)).
fof(rqLessOrEqual__r1_xxreal_0__r0_r0, axiom, r1_xxreal_0(np__0,np__0), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', rqLessOrEqual__r1_xxreal_0__r0_r0)).
fof(rqLessOrEqual__r1_xxreal_0__r0_rm1, axiom, ~(r1_xxreal_0(np__0,k4_xcmplx_0(np__1))), file('/home/mptp/big1/bushy_np/problems_small_np/t115_sin_cos9', rqLessOrEqual__r1_xxreal_0__r0_rm1)).
fof(c_0_60, plain, ![X30, X31]:(~v3_membered(X30)|(~m1_subset_1(X31,X30)|v1_xreal_0(X31))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc14_membered])])])).
fof(c_0_61, plain, ![X171, X172]:(~m1_subset_1(X171,k1_numbers)|~v1_xreal_0(X172)|m1_subset_1(k10_real_1(X171,X172),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k10_real_1])])).
fof(c_0_62, plain, ![X526]:(~v1_xboole_0(X526)|X526=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
fof(c_0_63, plain, ![X435, X436]:(~m1_subset_1(X435,k1_numbers)|~v1_xreal_0(X436)|k10_real_1(X435,X436)=k7_xcmplx_0(X435,X436)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k10_real_1])])).
fof(c_0_64, plain, ![X395]:(~v1_xcmplx_0(X395)|k5_xcmplx_0(k5_xcmplx_0(X395))=X395), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k5_xcmplx_0])])).
fof(c_0_65, plain, ![X479]:(~v1_xcmplx_0(X479)|k7_xcmplx_0(np__1,X479)=k5_xcmplx_0(X479)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc3_arithm])])).
fof(c_0_66, plain, ![X210]:(~v1_xcmplx_0(X210)|v1_xcmplx_0(k5_xcmplx_0(X210))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_xcmplx_0])])).
fof(c_0_67, plain, ![X294]:((v1_xcmplx_0(k4_xcmplx_0(X294))|~v1_xreal_0(X294))&(v1_xreal_0(k4_xcmplx_0(X294))|~v1_xreal_0(X294))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc3_xreal_0])])])).
cnf(c_0_68, plain, (k4_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2))=k7_xcmplx_0(np__1,np__2)), inference(split_conjunct,[status(thm)],[rqRealNeg__k4_xcmplx_0__rnm1d2_rn1d2])).
cnf(c_0_69, plain, (k7_xcmplx_0(np__1,k4_xcmplx_0(np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2)), inference(split_conjunct,[status(thm)],[rqRealDiv__k7_xcmplx_0__r1_rm2_rnm1d2])).
cnf(c_0_70, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_71, plain, (m1_subset_1(k10_real_1(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_72, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
cnf(c_0_73, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_74, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
cnf(c_0_75, plain, (k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_76, plain, (k5_xcmplx_0(k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_77, plain, (k7_xcmplx_0(np__1,X1)=k5_xcmplx_0(X1)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_78, plain, (v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_66])).
cnf(c_0_79, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_80, plain, (k4_xcmplx_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)))=k7_xcmplx_0(np__1,np__2)), inference(rw,[status(thm)],[c_0_68, c_0_69])).
cnf(c_0_81, plain, (v1_xreal_0(k10_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_72])])).
fof(c_0_82, plain, ![X496]:(~v1_xcmplx_0(X496)|k2_xcmplx_0(X496,k6_numbers)=X496), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_arithm])])).
cnf(c_0_83, plain, (X1=k6_numbers|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_73, c_0_74])).
cnf(c_0_84, plain, (v1_xboole_0(np__0)), inference(split_conjunct,[status(thm)],[spc0_boole])).
fof(c_0_85, plain, ![X468, X469]:(~v1_xxreal_0(X468)|~v1_xxreal_0(X469)|r1_xxreal_0(X468,X468)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[reflexivity_r1_xxreal_0])])).
fof(c_0_86, plain, (v1_xboole_0(esk31_0)&v1_xxreal_0(esk31_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc4_xxreal_0])])).
cnf(c_0_87, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(spm,[status(thm)],[c_0_71, c_0_75])).
cnf(c_0_88, plain, (k7_xcmplx_0(np__1,k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78])).
cnf(c_0_89, plain, (m1_subset_1(np__1,k1_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
fof(c_0_90, plain, ![X313]:((v1_xcmplx_0(k5_xcmplx_0(X313))|~v1_xreal_0(X313))&(v1_xreal_0(k5_xcmplx_0(X313))|~v1_xreal_0(X313))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc4_xreal_0])])])).
fof(c_0_91, plain, ![X115]:(~v1_xreal_0(X115)|v1_xcmplx_0(X115)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_xreal_0])])).
cnf(c_0_92, plain, (k4_xcmplx_0(k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(k4_xcmplx_0(np__1),np__2)), inference(split_conjunct,[status(thm)],[rqRealNeg__k4_xcmplx_0__rn1d2_rnm1d2])).
cnf(c_0_93, plain, (v1_xreal_0(k7_xcmplx_0(np__1,np__2))|~v1_xreal_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)))), inference(spm,[status(thm)],[c_0_79, c_0_80])).
cnf(c_0_94, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(spm,[status(thm)],[c_0_81, c_0_75])).
cnf(c_0_95, plain, (m1_subset_1(np__2,k1_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
fof(c_0_96, plain, ![X46, X47]:(~v6_membered(X46)|(~m1_subset_1(X47,X46)|v7_ordinal1(X47))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc17_membered])])])).
cnf(c_0_97, plain, (v6_membered(k4_ordinal1)), inference(split_conjunct,[status(thm)],[fc6_membered])).
cnf(c_0_98, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
fof(c_0_99, plain, ![X539, X540, X541, X542]:(~v1_xreal_0(X539)|(~v1_xreal_0(X540)|(~v1_xreal_0(X541)|(~v1_xreal_0(X542)|(r1_xxreal_0(X540,X539)|~r1_xxreal_0(X541,X542)|~r1_xxreal_0(k2_xcmplx_0(X540,X542),k2_xcmplx_0(X539,X541))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_xreal_1])])])])).
cnf(c_0_100, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_101, plain, (k6_numbers=np__0), inference(spm,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_102, plain, (m1_subset_1(np__0,k1_numbers)), inference(split_conjunct,[status(thm)],[spc0_numerals])).
cnf(c_0_103, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_104, plain, (v1_xxreal_0(esk31_0)), inference(split_conjunct,[status(thm)],[c_0_86])).
fof(c_0_105, plain, ![X125]:(~v1_xreal_0(X125)|v1_xxreal_0(X125)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_xreal_0])])).
cnf(c_0_106, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88]), c_0_89])])).
cnf(c_0_107, plain, (v1_xreal_0(k5_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_108, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
fof(c_0_109, plain, ![X326, X327]:(~v1_xreal_0(X326)|~v1_xreal_0(X327)|v1_xreal_0(k2_xcmplx_0(X326,X327))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc5_xreal_0])])).
cnf(c_0_110, plain, (k4_xcmplx_0(k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(np__1,k4_xcmplx_0(np__2))), inference(rw,[status(thm)],[c_0_92, c_0_69])).
cnf(c_0_111, plain, (v1_xreal_0(k7_xcmplx_0(np__1,np__2))|~v1_xreal_0(k4_xcmplx_0(np__2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94]), c_0_89])])).
cnf(c_0_112, plain, (v1_xreal_0(np__2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_95]), c_0_72])])).
fof(c_0_113, plain, ![X108]:(~v1_int_1(X108)|v1_xreal_0(X108)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_int_1])])).
fof(c_0_114, plain, ![X289]:((v1_xcmplx_0(k4_xcmplx_0(X289))|~v1_int_1(X289))&(v1_int_1(k4_xcmplx_0(X289))|~v1_int_1(X289))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc3_int_1])])])).
fof(c_0_115, plain, ![X93]:(~v7_ordinal1(X93)|v1_int_1(X93)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_int_1])])).
cnf(c_0_116, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_96])).
cnf(c_0_117, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_118, plain, (v6_membered(k5_numbers)), inference(rw,[status(thm)],[c_0_97, c_0_98])).
cnf(c_0_119, plain, (r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X4)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(k2_xcmplx_0(X2,X4),k2_xcmplx_0(X1,X3))), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_120, plain, (k2_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(X1)), inference(rw,[status(thm)],[c_0_100, c_0_101])).
cnf(c_0_121, plain, (v1_xreal_0(np__0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_102]), c_0_72])])).
cnf(c_0_122, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)), inference(spm,[status(thm)],[c_0_103, c_0_104])).
cnf(c_0_123, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_124, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_107]), c_0_108])).
cnf(c_0_125, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_109])).
cnf(c_0_126, plain, (k2_xcmplx_0(k7_xcmplx_0(k4_xcmplx_0(np__1),np__2),np__2)=k7_xcmplx_0(np__3,np__2)), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__rnm1d2_r2_rn3d2])).
cnf(c_0_127, plain, (v1_xreal_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)))|~v1_xreal_0(k7_xcmplx_0(np__1,np__2))), inference(spm,[status(thm)],[c_0_79, c_0_110])).
cnf(c_0_128, plain, (v1_xreal_0(k7_xcmplx_0(np__1,np__2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_79]), c_0_112])])).
cnf(c_0_129, plain, (v1_xreal_0(X1)|~v1_int_1(X1)), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_130, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1)), inference(split_conjunct,[status(thm)],[c_0_114])).
cnf(c_0_131, plain, (v1_int_1(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_115])).
cnf(c_0_132, plain, (v7_ordinal1(np__1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_118])])).
cnf(c_0_133, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X1,X3),X2)|~r1_xxreal_0(np__0,X3)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120]), c_0_121])]), c_0_108])).
cnf(c_0_134, plain, (k2_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(np__3,np__2)), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_rn1d2_rn3d2])).
cnf(c_0_135, plain, (v1_xreal_0(np__1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_89]), c_0_72])])).
cnf(c_0_136, plain, (r1_xxreal_0(np__0,k7_xcmplx_0(np__1,np__2))), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r0_rn1d2])).
cnf(c_0_137, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1)), inference(spm,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_138, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(spm,[status(thm)],[c_0_124, c_0_125])).
cnf(c_0_139, plain, (k2_xcmplx_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)),np__2)=k7_xcmplx_0(np__3,np__2)), inference(rw,[status(thm)],[c_0_126, c_0_69])).
cnf(c_0_140, plain, (v1_xreal_0(k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_127, c_0_128])])).
fof(c_0_141, plain, ![X173, X174]:(~v1_funct_1(X173)|~m1_subset_1(X173,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X174)|m1_subset_1(k1_fdiff_1(X173,X174),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_fdiff_1])])).
fof(c_0_142, plain, ![X524]:((r1_fdiff_1(k24_sin_cos,X524)|~v1_xreal_0(X524))&(k1_fdiff_1(k24_sin_cos,X524)=k1_seq_1(k24_sin_cos,X524)|~v1_xreal_0(X524))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t65_sin_cos])])])).
fof(c_0_143, plain, ![X119]:(~v3_membered(X119)|v2_membered(X119)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_membered])])).
cnf(c_0_144, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_int_1(X1)), inference(spm,[status(thm)],[c_0_129, c_0_130])).
cnf(c_0_145, plain, (v1_int_1(np__1)), inference(spm,[status(thm)],[c_0_131, c_0_132])).
cnf(c_0_146, plain, (r1_xxreal_0(np__1,X1)|~v1_xreal_0(X1)|~r1_xxreal_0(k7_xcmplx_0(np__3,np__2),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134]), c_0_128]), c_0_135]), c_0_136])])).
cnf(c_0_147, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(np__0,X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_137]), c_0_125])).
cnf(c_0_148, plain, (m1_subset_1(k7_xcmplx_0(np__3,np__2),k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_112]), c_0_140])])).
fof(c_0_149, plain, ![X543, X544, X545]:(((r1_tarski(X543,k1_relset_1(k1_numbers,X544))|~r2_fdiff_1(X544,X543)|(~v1_funct_1(X544)|~m1_subset_1(X544,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X543)|~m1_subset_1(X543,k1_zfmisc_1(k1_numbers))))&(~m1_subset_1(X545,k1_numbers)|(~r2_hidden(X545,X543)|r1_fdiff_1(X544,X545))|~r2_fdiff_1(X544,X543)|(~v1_funct_1(X544)|~m1_subset_1(X544,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X543)|~m1_subset_1(X543,k1_zfmisc_1(k1_numbers)))))&((m1_subset_1(esk40_2(X543,X544),k1_numbers)|~r1_tarski(X543,k1_relset_1(k1_numbers,X544))|r2_fdiff_1(X544,X543)|(~v1_funct_1(X544)|~m1_subset_1(X544,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X543)|~m1_subset_1(X543,k1_zfmisc_1(k1_numbers))))&((r2_hidden(esk40_2(X543,X544),X543)|~r1_tarski(X543,k1_relset_1(k1_numbers,X544))|r2_fdiff_1(X544,X543)|(~v1_funct_1(X544)|~m1_subset_1(X544,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X543)|~m1_subset_1(X543,k1_zfmisc_1(k1_numbers))))&(~r1_fdiff_1(X544,esk40_2(X543,X544))|~r1_tarski(X543,k1_relset_1(k1_numbers,X544))|r2_fdiff_1(X544,X543)|(~v1_funct_1(X544)|~m1_subset_1(X544,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X543)|~m1_subset_1(X543,k1_zfmisc_1(k1_numbers))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_fdiff_1])])])])])).
fof(c_0_150, plain, ![X533, X534]:((r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X534,k1_sin_cos9),X533)|(r1_xxreal_0(k1_seq_1(X534,X533),k1_real_1(np__1))|r1_xxreal_0(np__1,k1_seq_1(X534,X533)))|~r1_fdiff_1(X534,X533)|(~v1_funct_1(X534)|~m1_subset_1(X534,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~m1_subset_1(X533,k1_numbers))&(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X534,k1_sin_cos9),X533)=k10_real_1(k1_fdiff_1(X534,X533),k7_real_1(np__1,k5_square_1(k1_seq_1(X534,X533))))|(r1_xxreal_0(k1_seq_1(X534,X533),k1_real_1(np__1))|r1_xxreal_0(np__1,k1_seq_1(X534,X533)))|~r1_fdiff_1(X534,X533)|(~v1_funct_1(X534)|~m1_subset_1(X534,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~m1_subset_1(X533,k1_numbers))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t85_sin_cos9])])])])).
fof(c_0_151, plain, ![X159, X160]:(~v1_xxreal_0(X159)|~v1_xxreal_0(X160)|(r1_xxreal_0(X159,X160)|r1_xxreal_0(X160,X159))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[connectedness_r1_xxreal_0])])).
fof(c_0_152, plain, ![X25, X26]:(~v2_membered(X25)|(~m1_subset_1(X26,X25)|v1_xxreal_0(X26))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc13_membered])])])).
cnf(c_0_153, plain, (m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_141])).
cnf(c_0_154, plain, (k1_fdiff_1(k24_sin_cos,X1)=k1_seq_1(k24_sin_cos,X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_142])).
cnf(c_0_155, plain, (v1_funct_1(k24_sin_cos)), inference(split_conjunct,[status(thm)],[dt_k24_sin_cos])).
cnf(c_0_156, plain, (m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k24_sin_cos])).
cnf(c_0_157, plain, (v2_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_143])).
cnf(c_0_158, plain, (k2_xcmplx_0(np__2,k4_xcmplx_0(np__1))=np__1), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r2_rm1_r1])).
cnf(c_0_159, plain, (v1_xreal_0(k4_xcmplx_0(np__1))), inference(spm,[status(thm)],[c_0_144, c_0_145])).
cnf(c_0_160, plain, (r1_xxreal_0(np__1,k2_xcmplx_0(k7_xcmplx_0(np__3,np__2),X1))|~v1_xreal_0(k7_xcmplx_0(np__3,np__2))|~v1_xreal_0(X1)|~r1_xxreal_0(np__0,X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_147]), c_0_125])).
cnf(c_0_161, plain, (v1_xreal_0(k7_xcmplx_0(np__3,np__2))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_148]), c_0_72])])).
fof(c_0_162, plain, ~r1_xxreal_0(np__2,k7_xcmplx_0(np__3,np__2)), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r2_rn3d2])).
cnf(c_0_163, plain, (r2_fdiff_1(X1,X2)|~r1_fdiff_1(X1,esk40_2(X2,X1))|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_149])).
cnf(c_0_164, plain, (r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9),X2)|r1_xxreal_0(k1_seq_1(X1,X2),k1_real_1(np__1))|r1_xxreal_0(np__1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_150])).
cnf(c_0_165, plain, (m1_subset_1(esk40_2(X1,X2),k1_numbers)|r2_fdiff_1(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_149])).
fof(c_0_166, plain, ![X443]:(~m1_subset_1(X443,k1_numbers)|k1_real_1(X443)=k4_xcmplx_0(X443)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_real_1])])).
cnf(c_0_167, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_151])).
cnf(c_0_168, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_152])).
cnf(c_0_169, plain, (m1_subset_1(k1_seq_1(k24_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_154]), c_0_155]), c_0_156])])).
cnf(c_0_170, plain, (v2_membered(k1_numbers)), inference(spm,[status(thm)],[c_0_157, c_0_72])).
fof(c_0_171, plain, ![X517]:(~v1_xreal_0(X517)|~r1_xxreal_0(k1_seq_1(k24_sin_cos,X517),k6_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t54_sin_cos])])])).
cnf(c_0_172, plain, (r1_xxreal_0(np__2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(np__1,k2_xcmplx_0(X1,X2))|~r1_xxreal_0(X2,k4_xcmplx_0(np__1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_158]), c_0_159]), c_0_112])])).
cnf(c_0_173, plain, (r1_xxreal_0(np__1,k2_xcmplx_0(k7_xcmplx_0(np__3,np__2),X1))|~v1_xreal_0(X1)|~r1_xxreal_0(np__0,X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_160, c_0_161])])).
cnf(c_0_174, plain, (~r1_xxreal_0(np__2,k7_xcmplx_0(np__3,np__2))), inference(split_conjunct,[status(thm)],[c_0_162])).
cnf(c_0_175, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9),X2)|r1_xxreal_0(k1_seq_1(X1,esk40_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9))),k1_real_1(np__1))|r1_xxreal_0(np__1,k1_seq_1(X1,esk40_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9))))|~r1_fdiff_1(X1,esk40_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_164]), c_0_165])).
cnf(c_0_176, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_166])).
cnf(c_0_177, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xxreal_0(X2)), inference(spm,[status(thm)],[c_0_167, c_0_123])).
cnf(c_0_178, plain, (v1_xxreal_0(k1_seq_1(k24_sin_cos,X1))|~v1_xreal_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_169]), c_0_170])])).
cnf(c_0_179, plain, (~v1_xreal_0(X1)|~r1_xxreal_0(k1_seq_1(k24_sin_cos,X1),k6_numbers)), inference(split_conjunct,[status(thm)],[c_0_171])).
fof(c_0_180, negated_conjecture, ~(![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>((r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9)))&![X2]:(m1_subset_1(X2,k1_numbers)=>~((r2_hidden(X2,X1)&r1_xxreal_0(np__1,k1_seq_1(k24_sin_cos,X2))))))=>(r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),X1)&![X2]:(m1_subset_1(X2,k1_numbers)=>(r2_hidden(X2,X1)=>k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),X1),X2)=k10_real_1(k1_seq_1(k24_sin_cos,X2),k7_real_1(np__1,k5_square_1(k1_seq_1(k24_sin_cos,X2)))))))))), inference(assume_negation,[status(cth)],[t115_sin_cos9])).
fof(c_0_181, plain, ![X499, X500]:(~r2_hidden(X499,X500)|m1_subset_1(X499,X500)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
fof(c_0_182, plain, ![X70]:(~m1_subset_1(X70,k1_numbers)|v1_xreal_0(X70)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_xreal_0])])).
cnf(c_0_183, plain, (~v1_xreal_0(X1)|~r1_xxreal_0(X1,k4_xcmplx_0(np__1))|~r1_xxreal_0(np__0,X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_173]), c_0_161])]), c_0_174])).
cnf(c_0_184, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9),X2)|r1_xxreal_0(k1_seq_1(X1,esk40_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9))),k4_xcmplx_0(np__1))|r1_xxreal_0(np__1,k1_seq_1(X1,esk40_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9))))|~r1_fdiff_1(X1,esk40_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_176]), c_0_89])])).
cnf(c_0_185, plain, (r1_xxreal_0(k1_seq_1(k24_sin_cos,X1),X2)|r1_xxreal_0(X2,k1_seq_1(k24_sin_cos,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(spm,[status(thm)],[c_0_177, c_0_178])).
cnf(c_0_186, plain, (~v1_xreal_0(X1)|~r1_xxreal_0(k1_seq_1(k24_sin_cos,X1),np__0)), inference(rw,[status(thm)],[c_0_179, c_0_101])).
fof(c_0_187, plain, ![X514, X515, X516]:(~r2_hidden(X514,X515)|~m1_subset_1(X515,k1_zfmisc_1(X516))|m1_subset_1(X514,X516)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_188, negated_conjecture, ![X8]:((v3_rcomp_1(esk1_0)&m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers)))&((r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9)))&(~m1_subset_1(X8,k1_numbers)|(~r2_hidden(X8,esk1_0)|~r1_xxreal_0(np__1,k1_seq_1(k24_sin_cos,X8)))))&((m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),esk1_0))&((r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),esk1_0))&(k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),esk1_0),esk2_0)!=k10_real_1(k1_seq_1(k24_sin_cos,esk2_0),k7_real_1(np__1,k5_square_1(k1_seq_1(k24_sin_cos,esk2_0))))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),esk1_0)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_180])])])])])).
cnf(c_0_189, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_181])).
cnf(c_0_190, plain, (r2_hidden(esk40_2(X1,X2),X1)|r2_fdiff_1(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_149])).
fof(c_0_191, plain, ![X76, X77]:(~v3_membered(X76)|(~m1_subset_1(X77,k1_zfmisc_1(X76))|v3_membered(X77))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc21_membered])])])).
fof(c_0_192, plain, ![X175, X176, X177, X178, X179, X180]:((v1_funct_1(k1_partfun1(X175,X176,X177,X178,X179,X180))|(~v1_funct_1(X179)|~m1_subset_1(X179,k1_zfmisc_1(k2_zfmisc_1(X175,X176)))|~v1_funct_1(X180)|~m1_subset_1(X180,k1_zfmisc_1(k2_zfmisc_1(X177,X178)))))&(m1_subset_1(k1_partfun1(X175,X176,X177,X178,X179,X180),k1_zfmisc_1(k2_zfmisc_1(X175,X178)))|(~v1_funct_1(X179)|~m1_subset_1(X179,k1_zfmisc_1(k2_zfmisc_1(X175,X176)))|~v1_funct_1(X180)|~m1_subset_1(X180,k1_zfmisc_1(k2_zfmisc_1(X177,X178)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_partfun1])])])).
fof(c_0_193, plain, ![X437, X438, X439, X440, X441, X442]:(~v1_funct_1(X441)|~m1_subset_1(X441,k1_zfmisc_1(k2_zfmisc_1(X437,X438)))|~v1_funct_1(X442)|~m1_subset_1(X442,k1_zfmisc_1(k2_zfmisc_1(X439,X440)))|k1_partfun1(X437,X438,X439,X440,X441,X442)=k3_relat_1(X441,X442)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_partfun1])])).
cnf(c_0_194, plain, (k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9),X2)=k10_real_1(k1_fdiff_1(X1,X2),k7_real_1(np__1,k5_square_1(k1_seq_1(X1,X2))))|r1_xxreal_0(k1_seq_1(X1,X2),k1_real_1(np__1))|r1_xxreal_0(np__1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_150])).
cnf(c_0_195, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_182])).
fof(c_0_196, plain, ![X162, X163, X164, X165]:(((k1_relset_1(k1_numbers,X164)=X163|X164!=k2_fdiff_1(X162,X163)|(~v1_funct_1(X164)|~m1_subset_1(X164,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X162,X163)|(~v1_funct_1(X162)|~m1_subset_1(X162,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(~m1_subset_1(X165,k1_numbers)|(~r2_hidden(X165,X163)|k1_seq_1(X164,X165)=k1_fdiff_1(X162,X165))|X164!=k2_fdiff_1(X162,X163)|(~v1_funct_1(X164)|~m1_subset_1(X164,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X162,X163)|(~v1_funct_1(X162)|~m1_subset_1(X162,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))))&((m1_subset_1(esk3_3(X162,X163,X164),k1_numbers)|k1_relset_1(k1_numbers,X164)!=X163|X164=k2_fdiff_1(X162,X163)|(~v1_funct_1(X164)|~m1_subset_1(X164,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X162,X163)|(~v1_funct_1(X162)|~m1_subset_1(X162,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&((r2_hidden(esk3_3(X162,X163,X164),X163)|k1_relset_1(k1_numbers,X164)!=X163|X164=k2_fdiff_1(X162,X163)|(~v1_funct_1(X164)|~m1_subset_1(X164,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X162,X163)|(~v1_funct_1(X162)|~m1_subset_1(X162,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(k1_seq_1(X164,esk3_3(X162,X163,X164))!=k1_fdiff_1(X162,esk3_3(X162,X163,X164))|k1_relset_1(k1_numbers,X164)!=X163|X164=k2_fdiff_1(X162,X163)|(~v1_funct_1(X164)|~m1_subset_1(X164,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X162,X163)|(~v1_funct_1(X162)|~m1_subset_1(X162,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_fdiff_1])])])])])).
fof(c_0_197, plain, ![X188, X189]:((v1_funct_1(k2_fdiff_1(X188,X189))|(~v1_funct_1(X188)|~m1_subset_1(X188,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(m1_subset_1(k2_fdiff_1(X188,X189),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|(~v1_funct_1(X188)|~m1_subset_1(X188,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_fdiff_1])])])).
cnf(c_0_198, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9),X2)|r1_xxreal_0(np__1,k1_seq_1(X1,esk40_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9))))|~r1_fdiff_1(X1,esk40_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9))|~v1_funct_1(X1)|~v1_xreal_0(k1_seq_1(X1,esk40_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9))))|~r1_xxreal_0(np__0,k1_seq_1(X1,esk40_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9))))|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_sin_cos9),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)), inference(spm,[status(thm)],[c_0_183, c_0_184])).
cnf(c_0_199, plain, (r1_xxreal_0(np__0,k1_seq_1(k24_sin_cos,X1))|~v1_xreal_0(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_121]), c_0_186])).
cnf(c_0_200, plain, (v1_xreal_0(k1_seq_1(k24_sin_cos,X1))|~v1_xreal_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_169]), c_0_72])])).
cnf(c_0_201, plain, (r1_fdiff_1(k24_sin_cos,X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_142])).
cnf(c_0_202, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_187])).
cnf(c_0_203, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_188])).
cnf(c_0_204, plain, (r2_fdiff_1(X1,X2)|m1_subset_1(esk40_2(X2,X1),X2)|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)), inference(spm,[status(thm)],[c_0_189, c_0_190])).
cnf(c_0_205, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9)))), inference(split_conjunct,[status(thm)],[c_0_188])).
cnf(c_0_206, negated_conjecture, (v3_rcomp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_188])).
cnf(c_0_207, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_191])).
cnf(c_0_208, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_192])).
cnf(c_0_209, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_193])).
cnf(c_0_210, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_192])).
fof(c_0_211, plain, ![X153, X154]:(~v1_xcmplx_0(X153)|~v1_xcmplx_0(X154)|k2_xcmplx_0(X153,X154)=k2_xcmplx_0(X154,X153)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k2_xcmplx_0])])).
cnf(c_0_212, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),esk1_0),esk2_0)!=k10_real_1(k1_seq_1(k24_sin_cos,esk2_0),k7_real_1(np__1,k5_square_1(k1_seq_1(k24_sin_cos,esk2_0))))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_188])).
cnf(c_0_213, plain, (k10_real_1(k1_seq_1(k24_sin_cos,X1),k7_real_1(np__1,k5_square_1(k1_seq_1(k24_sin_cos,X1))))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),X1)|r1_xxreal_0(k1_seq_1(k24_sin_cos,X1),k1_real_1(np__1))|r1_xxreal_0(np__1,k1_seq_1(k24_sin_cos,X1))|~r1_fdiff_1(k24_sin_cos,X1)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_154]), c_0_155]), c_0_156])]), c_0_195])).
cnf(c_0_214, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_188])).
cnf(c_0_215, plain, (k1_seq_1(X3,X1)=k1_fdiff_1(X4,X1)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,X2)|X3!=k2_fdiff_1(X4,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X4,X2)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_196])).
cnf(c_0_216, plain, (m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_197])).
cnf(c_0_217, plain, (v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_197])).
cnf(c_0_218, negated_conjecture, (~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk1_0)|~r1_xxreal_0(np__1,k1_seq_1(k24_sin_cos,X1))), inference(split_conjunct,[status(thm)],[c_0_188])).
cnf(c_0_219, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),X1)|r1_xxreal_0(np__1,k1_seq_1(k24_sin_cos,esk40_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9))))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9))|~v1_xreal_0(esk40_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9)))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_199]), c_0_155]), c_0_156])]), c_0_200]), c_0_201])).
cnf(c_0_220, negated_conjecture, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk1_0)), inference(spm,[status(thm)],[c_0_202, c_0_203])).
cnf(c_0_221, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),esk1_0)|m1_subset_1(esk40_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9)),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_205]), c_0_203]), c_0_206])])).
cnf(c_0_222, negated_conjecture, (v3_membered(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_203]), c_0_72])])).
cnf(c_0_223, plain, (v1_funct_1(k3_relat_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(spm,[status(thm)],[c_0_208, c_0_209])).
cnf(c_0_224, plain, (m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k1_sin_cos9])).
cnf(c_0_225, plain, (v1_funct_1(k1_sin_cos9)), inference(split_conjunct,[status(thm)],[dt_k1_sin_cos9])).
cnf(c_0_226, plain, (m1_subset_1(k3_relat_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X6)))), inference(spm,[status(thm)],[c_0_210, c_0_209])).
cnf(c_0_227, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_228, plain, (k4_xcmplx_0(np__0)=np__0), inference(split_conjunct,[status(thm)],[rqRealNeg__k4_xcmplx_0__r0_r0])).
cnf(c_0_229, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_211])).
cnf(c_0_230, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~r1_xxreal_0(X3,np__0)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120]), c_0_121])]), c_0_108])).
cnf(c_0_231, negated_conjecture, (r1_xxreal_0(k1_seq_1(k24_sin_cos,esk2_0),k1_real_1(np__1))|r1_xxreal_0(np__1,k1_seq_1(k24_sin_cos,esk2_0))|k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),esk1_0),esk2_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),esk2_0)|~r1_fdiff_1(k24_sin_cos,esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),esk1_0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_213]), c_0_214])).
cnf(c_0_232, plain, (k1_seq_1(k2_fdiff_1(X1,X2),X3)=k1_fdiff_1(X1,X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r2_hidden(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_215]), c_0_216]), c_0_217])).
cnf(c_0_233, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_188])).
cnf(c_0_234, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9))|~v1_xreal_0(esk40_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9)))|~r2_hidden(esk40_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9)),esk1_0)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_219]), c_0_220])).
cnf(c_0_235, negated_conjecture, (v1_xreal_0(esk40_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_221]), c_0_222])])).
cnf(c_0_236, plain, (v1_funct_1(k3_relat_1(X1,k1_sin_cos9))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_224]), c_0_225])])).
cnf(c_0_237, plain, (m1_subset_1(k3_relat_1(X1,k1_sin_cos9),k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_224]), c_0_225])])).
cnf(c_0_238, plain, (v1_xcmplx_0(np__0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_228]), c_0_121])])).
cnf(c_0_239, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k2_xcmplx_0(X3,X1),X2)|~r1_xxreal_0(np__0,X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_229]), c_0_108]), c_0_108])).
cnf(c_0_240, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X2,np__0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_137]), c_0_125])).
cnf(c_0_241, negated_conjecture, (r1_xxreal_0(k1_seq_1(k24_sin_cos,esk2_0),k1_real_1(np__1))|r1_xxreal_0(np__1,k1_seq_1(k24_sin_cos,esk2_0))|~r1_fdiff_1(k24_sin_cos,esk2_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),esk1_0)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_232]), c_0_220]), c_0_233])).
cnf(c_0_242, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_190]), c_0_205]), c_0_203]), c_0_206])]), c_0_235])).
cnf(c_0_243, plain, (v1_funct_1(k3_relat_1(k24_sin_cos,k1_sin_cos9))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_156]), c_0_155])])).
cnf(c_0_244, plain, (m1_subset_1(k3_relat_1(k24_sin_cos,k1_sin_cos9),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_156]), c_0_155])])).
cnf(c_0_245, plain, (k2_xcmplx_0(np__0,X1)=X1|~v1_xcmplx_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_120]), c_0_238])])).
cnf(c_0_246, plain, (r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(np__0,X2)|~r1_xxreal_0(X1,np__0)), inference(spm,[status(thm)],[c_0_239, c_0_240])).
cnf(c_0_247, negated_conjecture, (r1_xxreal_0(k1_seq_1(k24_sin_cos,esk2_0),k4_xcmplx_0(np__1))|r1_xxreal_0(np__1,k1_seq_1(k24_sin_cos,esk2_0))|~r1_fdiff_1(k24_sin_cos,esk2_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),esk1_0)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_176]), c_0_89])])).
cnf(c_0_248, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k3_relat_1(k24_sin_cos,k1_sin_cos9),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_209]), c_0_225]), c_0_155]), c_0_224]), c_0_156])])).
cnf(c_0_249, negated_conjecture, (r2_fdiff_1(k3_relat_1(k24_sin_cos,k1_sin_cos9),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_209]), c_0_243]), c_0_244]), c_0_225]), c_0_155]), c_0_224]), c_0_156])])).
cnf(c_0_250, plain, (r1_xxreal_0(np__0,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(np__0,X2)|~r1_xxreal_0(X2,X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_245]), c_0_121])]), c_0_108])).
cnf(c_0_251, plain, (r1_xxreal_0(X1,k1_seq_1(k24_sin_cos,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,np__0)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_199]), c_0_200])).
cnf(c_0_252, plain, (r1_xxreal_0(np__0,np__0)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r0_r0])).
cnf(c_0_253, negated_conjecture, (r1_xxreal_0(k1_seq_1(k24_sin_cos,esk2_0),k4_xcmplx_0(np__1))|r1_xxreal_0(np__1,k1_seq_1(k24_sin_cos,esk2_0))|~r1_fdiff_1(k24_sin_cos,esk2_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(spm,[status(thm)],[c_0_247, c_0_242])).
fof(c_0_254, plain, ~r1_xxreal_0(np__0,k4_xcmplx_0(np__1)), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r0_rm1])).
cnf(c_0_255, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_248, c_0_249])])).
cnf(c_0_256, plain, (r1_xxreal_0(np__0,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k1_seq_1(k24_sin_cos,X2),X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_251]), c_0_121]), c_0_252])]), c_0_200])).
cnf(c_0_257, negated_conjecture, (r1_xxreal_0(k1_seq_1(k24_sin_cos,esk2_0),k4_xcmplx_0(np__1))|r1_xxreal_0(np__1,k1_seq_1(k24_sin_cos,esk2_0))|~r1_fdiff_1(k24_sin_cos,esk2_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_210]), c_0_225]), c_0_155]), c_0_224]), c_0_156])])).
cnf(c_0_258, plain, (~r1_xxreal_0(np__0,k4_xcmplx_0(np__1))), inference(split_conjunct,[status(thm)],[c_0_254])).
cnf(c_0_259, negated_conjecture, (v1_xreal_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_255]), c_0_72])])).
cnf(c_0_260, negated_conjecture, (r1_xxreal_0(np__1,k1_seq_1(k24_sin_cos,esk2_0))|~r1_fdiff_1(k24_sin_cos,esk2_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k24_sin_cos,k1_sin_cos9))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_257]), c_0_159])]), c_0_258]), c_0_259])])).
cnf(c_0_261, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k3_relat_1(k24_sin_cos,k1_sin_cos9),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_209]), c_0_225]), c_0_155]), c_0_224]), c_0_156])])).
cnf(c_0_262, negated_conjecture, (r1_xxreal_0(np__1,k1_seq_1(k24_sin_cos,esk2_0))|~r1_fdiff_1(k24_sin_cos,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_209]), c_0_243]), c_0_225]), c_0_155]), c_0_224]), c_0_156])])).
cnf(c_0_263, negated_conjecture, (r2_hidden(esk2_0,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_261, c_0_249])])).
cnf(c_0_264, negated_conjecture, (~r1_fdiff_1(k24_sin_cos,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_262]), c_0_263]), c_0_255])])).
cnf(c_0_265, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_201]), c_0_259])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 266
# Proof object clause steps            : 164
# Proof object formula steps           : 102
# Proof object conjectures             : 31
# Proof object clause conjectures      : 28
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 76
# Proof object initial formulas used   : 60
# Proof object generating inferences   : 77
# Proof object simplifying inferences  : 159
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 634
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 901
# Removed in clause preprocessing      : 82
# Initial clauses in saturation        : 819
# Processed clauses                    : 3115
# ...of these trivial                  : 141
# ...subsumed                          : 1042
# ...remaining for further processing  : 1932
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 46
# Backward-rewritten                   : 92
# Generated clauses                    : 23735
# ...of the previous two non-trivial   : 18126
# Contextual simplify-reflections      : 138
# Paramodulations                      : 23732
# Factorizations                       : 0
# Equation resolutions                 : 3
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1794
#    Positive orientable unit clauses  : 596
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 97
#    Non-unit-clauses                  : 1101
# Current number of unprocessed clauses: 15690
# ...number of literals in the above   : 70245
# Current number of archived formulas  : 0
# Current number of archived clauses   : 138
# Clause-clause subsumption calls (NU) : 480292
# Rec. Clause-clause subsumption calls : 123875
# Non-unit clause-clause subsumptions  : 670
# Unit Clause-clause subsumption calls : 39453
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 103
# BW rewrite match successes           : 27
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 560556

# -------------------------------------------------
# User time                : 20.156 s
# System time              : 0.960 s
# Total time               : 21.116 s
# Maximum resident set size: 327300 pages
20.15user 1.00system 0:22.60elapsed 93%CPU (0avgtext+0avgdata 327300maxresident)k
0inputs+136outputs (0major+345544minor)pagefaults 0swaps
