2020-06-03 01:00:06.550829: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_35
2020-06-03 01:00:06.647985: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-06-03 01:00:06.902006: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-06-03 01:00:07.444721: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-06-03 01:00:09.544744: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 2993946 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_35' loaded (query=512; context=768; weigths=real; len_mult=0.000000).
# Preprocessing time       : 3.900 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', t6_boole)).
fof(t117_seq_4, conjecture, ![X1]:(m2_subset_1(X1,k1_numbers,k5_numbers)=>![X2]:(m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))=>~((((~(r2_hidden(X2,X3))&X3!=k1_xboole_0)&v2_seq_4(X3,X1))&r1_xxreal_0(k24_seq_4(X1,X2,X3),k6_numbers)))))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', t117_seq_4)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', redefinition_k6_numbers)).
fof(t114_seq_4, axiom, ![X1]:(m2_subset_1(X1,k1_numbers,k5_numbers)=>![X2]:(m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))=>(X3!=k1_xboole_0=>r1_xxreal_0(k6_numbers,k24_seq_4(X1,X2,X3)))))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', t114_seq_4)).
fof(d15_seq_4, axiom, ![X1]:(m2_subset_1(X1,k1_numbers,k5_numbers)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))=>(v2_seq_4(X2,X1)<=>![X3]:(m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))=>(![X4]:(m1_subset_1(X4,k1_numbers)=>~((~(r1_xxreal_0(X4,k6_numbers))&![X5]:(m2_finseq_2(X5,k2_numbers,k14_seq_4(X1))=>~((~(r1_xxreal_0(X4,k22_seq_4(X5)))&r2_hidden(k15_seq_4(X1,X3,X5),X2)))))))=>r2_hidden(X3,X2)))))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', d15_seq_4)).
fof(t104_seq_4, axiom, ![X1]:(m2_subset_1(X1,k1_numbers,k5_numbers)=>![X2]:(m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))=>![X3]:(m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))=>k22_seq_4(k19_seq_4(X1,X2,X3))=k22_seq_4(k19_seq_4(X1,X3,X2))))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', t104_seq_4)).
fof(spc0_boole, axiom, v1_xboole_0(np__0), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', spc0_boole)).
fof(fraenkel_a_3_3_seq_4, axiom, ![X1, X2, X3, X4]:(((m2_subset_1(X2,k1_numbers,k5_numbers)&m2_finseq_2(X3,k2_numbers,k14_seq_4(X2)))&m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2))))=>(r2_hidden(X1,a_3_3_seq_4(X2,X3,X4))<=>?[X5]:((m2_finseq_2(X5,k2_numbers,k14_seq_4(X2))&X1=k22_seq_4(k19_seq_4(X2,X3,X5)))&r2_hidden(X5,X4)))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', fraenkel_a_3_3_seq_4)).
fof(t4_real, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>~(((r1_xxreal_0(X1,X2)&~(v2_xxreal_0(X2)))&v2_xxreal_0(X1))))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', t4_real)).
fof(cc14_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xreal_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', cc14_membered)).
fof(t5_real, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>(r1_xxreal_0(X1,X2)=>(v1_xboole_0(X2)|v3_xxreal_0(X1)|v2_xxreal_0(X2))))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', t5_real)).
fof(spc0_numerals, axiom, ((m2_subset_1(np__0,k1_numbers,k5_numbers)&m1_subset_1(np__0,k5_numbers))&m1_subset_1(np__0,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', spc0_numerals)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', fc3_membered)).
fof(fc16_xreal_0, axiom, ![X1]:((~(v3_xxreal_0(X1))&v1_xreal_0(X1))=>(v1_xcmplx_0(k4_xcmplx_0(X1))&~(v2_xxreal_0(k4_xcmplx_0(X1))))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', fc16_xreal_0)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', redefinition_m2_subset_1)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', cc1_subset_1)).
fof(fc6_ordinal1, axiom, (~(v1_xboole_0(k4_ordinal1))&v3_ordinal1(k4_ordinal1)), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', fc6_ordinal1)).
fof(cc20_membered, axiom, ![X1]:(v2_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v2_membered(X2))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', cc20_membered)).
fof(t80_seq_4, axiom, ![X1]:(m2_subset_1(X1,k1_numbers,k5_numbers)=>![X2]:(m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))=>![X3]:(m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))=>k15_seq_4(X1,X2,k19_seq_4(X1,X3,X2))=X3))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', t80_seq_4)).
fof(rqRealNeg__k4_xcmplx_0__r0_r0, axiom, k4_xcmplx_0(np__0)=np__0, file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', rqRealNeg__k4_xcmplx_0__r0_r0)).
fof(redefinition_m2_finseq_2, axiom, ![X1, X2]:(m1_finseq_2(X2,X1)=>![X3]:(m2_finseq_2(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', redefinition_m2_finseq_2)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', redefinition_k5_numbers)).
fof(fc15_xreal_0, axiom, ![X1]:((~(v2_xxreal_0(X1))&v1_xreal_0(X1))=>(v1_xcmplx_0(k4_xcmplx_0(X1))&~(v3_xxreal_0(k4_xcmplx_0(X1))))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', fc15_xreal_0)).
fof(cc13_membered, axiom, ![X1]:(v2_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xxreal_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', cc13_membered)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', dt_k5_numbers)).
fof(cc4_membered, axiom, ![X1]:(v3_membered(X1)=>v2_membered(X1)), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', cc4_membered)).
fof(fc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>(v1_xcmplx_0(k4_xcmplx_0(X1))&v1_xreal_0(k4_xcmplx_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', fc3_xreal_0)).
fof(dt_k24_seq_4, axiom, ![X1, X2, X3]:(((m1_subset_1(X1,k5_numbers)&m1_subset_1(X2,k14_seq_4(X1)))&m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1))))=>m1_subset_1(k24_seq_4(X1,X2,X3),k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', dt_k24_seq_4)).
fof(dt_k14_seq_4, axiom, ![X1]:(m1_subset_1(X1,k5_numbers)=>m1_finseq_2(k14_seq_4(X1),k2_numbers)), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', dt_k14_seq_4)).
fof(dt_k19_seq_4, axiom, ![X1, X2, X3]:(((m1_subset_1(X1,k5_numbers)&m1_subset_1(X2,k14_seq_4(X1)))&m1_subset_1(X3,k14_seq_4(X1)))=>m2_finseq_2(k19_seq_4(X1,X2,X3),k2_numbers,k14_seq_4(X1))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', dt_k19_seq_4)).
fof(d2_seq_4, axiom, ![X1]:(v3_membered(X1)=>(v3_xxreal_2(X1)=>(v1_xboole_0(X1)|![X2]:(v1_xreal_0(X2)=>(X2=k3_seq_4(X1)<=>(![X3]:(v1_xreal_0(X3)=>(r2_hidden(X3,X1)=>r1_xxreal_0(X2,X3)))&![X3]:(v1_xreal_0(X3)=>~((~(r1_xxreal_0(X3,k6_numbers))&![X4]:(v1_xreal_0(X4)=>~((r2_hidden(X4,X1)&~(r1_xxreal_0(k9_binop_2(X2,X3),X4)))))))))))))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', d2_seq_4)).
fof(t1_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k2_xcmplx_0(X1,k6_numbers)=X1), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', t1_arithm)).
fof(cc3_xxreal_0, axiom, ![X1]:((v1_xxreal_0(X1)&v2_xxreal_0(X1))=>((~(v1_xboole_0(X1))&v1_xxreal_0(X1))&~(v3_xxreal_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', cc3_xxreal_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', t4_subset)).
fof(redefinition_k9_binop_2, axiom, ![X1, X2]:((v1_xreal_0(X1)&v1_xreal_0(X2))=>k9_binop_2(X1,X2)=k2_xcmplx_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', redefinition_k9_binop_2)).
fof(commutativity_k2_xcmplx_0, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', commutativity_k2_xcmplx_0)).
fof(d17_seq_4, axiom, ![X1]:(m2_subset_1(X1,k1_numbers,k5_numbers)=>![X2]:(m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))=>![X4]:(m1_subset_1(X4,k1_numbers)=>(X4=k24_seq_4(X1,X2,X3)<=>![X5]:(m1_subset_1(X5,k1_zfmisc_1(k1_numbers))=>(X5=a_3_3_seq_4(X1,X2,X3)=>X4=k5_seq_4(X5)))))))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', d17_seq_4)).
fof(cc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xcmplx_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', cc3_xreal_0)).
fof(d2_xxreal_2, axiom, ![X1]:(v2_membered(X1)=>![X2]:(v1_xxreal_0(X2)=>(m2_xxreal_2(X2,X1)<=>![X3]:(v1_xxreal_0(X3)=>(r2_hidden(X3,X1)=>r1_xxreal_0(X2,X3)))))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', d2_xxreal_2)).
fof(fc12_xreal_0, axiom, ![X1, X2]:((((v2_xxreal_0(X1)&v1_xreal_0(X1))&~(v3_xxreal_0(X2)))&v1_xreal_0(X2))=>v2_xxreal_0(k2_xcmplx_0(X2,X1))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', fc12_xreal_0)).
fof(cc4_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xxreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', cc4_xreal_0)).
fof(cc2_xreal_0, axiom, ![X1]:(v7_ordinal1(X1)=>v1_xreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', cc2_xreal_0)).
fof(cc7_ordinal1, axiom, ![X1]:(v1_xboole_0(X1)=>v7_ordinal1(X1)), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', cc7_ordinal1)).
fof(redefinition_k5_seq_4, axiom, ![X1]:(m1_subset_1(X1,k1_zfmisc_1(k1_numbers))=>k5_seq_4(X1)=k3_seq_4(X1)), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', redefinition_k5_seq_4)).
fof(d9_xxreal_2, axiom, ![X1]:(v2_membered(X1)=>(v3_xxreal_2(X1)<=>?[X2]:(v1_xreal_0(X2)&m2_xxreal_2(X2,X1)))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', d9_xxreal_2)).
fof(connectedness_r1_xxreal_0, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>(r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', connectedness_r1_xxreal_0)).
fof(cc8_membered, axiom, ![X1]:(m1_subset_1(X1,k1_zfmisc_1(k1_numbers))=>v3_membered(X1)), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', cc8_membered)).
fof(cc2_xxreal_0, axiom, ![X1]:(v7_ordinal1(X1)=>v1_xxreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', cc2_xxreal_0)).
fof(cc8_ordinal1, axiom, ![X1]:(m1_subset_1(X1,k4_ordinal1)=>v7_ordinal1(X1)), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', cc8_ordinal1)).
fof(t8_real, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>~(((~(r1_xxreal_0(X1,X2))&~(v3_xxreal_0(X2)))&~(v2_xxreal_0(X1)))))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', t8_real)).
fof(redefinition_k19_seq_4, axiom, ![X1, X2, X3]:(((m1_subset_1(X1,k5_numbers)&m1_subset_1(X2,k14_seq_4(X1)))&m1_subset_1(X3,k14_seq_4(X1)))=>k19_seq_4(X1,X2,X3)=k45_valued_1(X2,X3)), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', redefinition_k19_seq_4)).
fof(dt_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)=>m1_subset_1(X3,X1))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', dt_m2_subset_1)).
fof(t7_boole, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', t7_boole)).
fof(t4_subset_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>~((X2!=k1_xboole_0&![X3]:(m1_subset_1(X3,X1)=>~(r2_hidden(X3,X2)))))), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', t4_subset_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', t3_subset)).
fof(reflexivity_r1_tarski, axiom, ![X1, X2]:r1_tarski(X1,X1), file('/scratch/mptp/allpred/preds__0.005/t117_seq_4', reflexivity_r1_tarski)).
fof(c_0_56, plain, ![X88]:(~v1_xboole_0(X88)|X88=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
fof(c_0_57, negated_conjecture, ~(![X1]:(m2_subset_1(X1,k1_numbers,k5_numbers)=>![X2]:(m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))=>~((((~(r2_hidden(X2,X3))&X3!=k1_xboole_0)&v2_seq_4(X3,X1))&r1_xxreal_0(k24_seq_4(X1,X2,X3),k6_numbers))))))), inference(assume_negation,[status(cth)],[t117_seq_4])).
cnf(c_0_58, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_59, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
fof(c_0_60, plain, ![X17, X18, X19]:(~m2_subset_1(X17,k1_numbers,k5_numbers)|(~m2_finseq_2(X18,k2_numbers,k14_seq_4(X17))|(~m1_subset_1(X19,k1_zfmisc_1(k14_seq_4(X17)))|(X19=k1_xboole_0|r1_xxreal_0(k6_numbers,k24_seq_4(X17,X18,X19)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t114_seq_4])])])).
fof(c_0_61, plain, ![X9, X10, X11, X13, X15]:(((m1_subset_1(esk4_3(X9,X10,X11),k1_numbers)|r2_hidden(X11,X10)|~m2_finseq_2(X11,k2_numbers,k14_seq_4(X9))|~v2_seq_4(X10,X9)|~m1_subset_1(X10,k1_zfmisc_1(k14_seq_4(X9)))|~m2_subset_1(X9,k1_numbers,k5_numbers))&((~r1_xxreal_0(esk4_3(X9,X10,X11),k6_numbers)|r2_hidden(X11,X10)|~m2_finseq_2(X11,k2_numbers,k14_seq_4(X9))|~v2_seq_4(X10,X9)|~m1_subset_1(X10,k1_zfmisc_1(k14_seq_4(X9)))|~m2_subset_1(X9,k1_numbers,k5_numbers))&(~m2_finseq_2(X13,k2_numbers,k14_seq_4(X9))|(r1_xxreal_0(esk4_3(X9,X10,X11),k22_seq_4(X13))|~r2_hidden(k15_seq_4(X9,X11,X13),X10))|r2_hidden(X11,X10)|~m2_finseq_2(X11,k2_numbers,k14_seq_4(X9))|~v2_seq_4(X10,X9)|~m1_subset_1(X10,k1_zfmisc_1(k14_seq_4(X9)))|~m2_subset_1(X9,k1_numbers,k5_numbers))))&((m2_finseq_2(esk5_2(X9,X10),k2_numbers,k14_seq_4(X9))|v2_seq_4(X10,X9)|~m1_subset_1(X10,k1_zfmisc_1(k14_seq_4(X9)))|~m2_subset_1(X9,k1_numbers,k5_numbers))&(((m2_finseq_2(esk6_3(X9,X10,X15),k2_numbers,k14_seq_4(X9))|r1_xxreal_0(X15,k6_numbers)|~m1_subset_1(X15,k1_numbers)|v2_seq_4(X10,X9)|~m1_subset_1(X10,k1_zfmisc_1(k14_seq_4(X9)))|~m2_subset_1(X9,k1_numbers,k5_numbers))&((~r1_xxreal_0(X15,k22_seq_4(esk6_3(X9,X10,X15)))|r1_xxreal_0(X15,k6_numbers)|~m1_subset_1(X15,k1_numbers)|v2_seq_4(X10,X9)|~m1_subset_1(X10,k1_zfmisc_1(k14_seq_4(X9)))|~m2_subset_1(X9,k1_numbers,k5_numbers))&(r2_hidden(k15_seq_4(X9,esk5_2(X9,X10),esk6_3(X9,X10,X15)),X10)|r1_xxreal_0(X15,k6_numbers)|~m1_subset_1(X15,k1_numbers)|v2_seq_4(X10,X9)|~m1_subset_1(X10,k1_zfmisc_1(k14_seq_4(X9)))|~m2_subset_1(X9,k1_numbers,k5_numbers))))&(~r2_hidden(esk5_2(X9,X10),X10)|v2_seq_4(X10,X9)|~m1_subset_1(X10,k1_zfmisc_1(k14_seq_4(X9)))|~m2_subset_1(X9,k1_numbers,k5_numbers))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d15_seq_4])])])])])])).
fof(c_0_62, plain, ![X50, X51, X52]:(~m2_subset_1(X50,k1_numbers,k5_numbers)|(~m2_finseq_2(X51,k2_numbers,k14_seq_4(X50))|(~m2_finseq_2(X52,k2_numbers,k14_seq_4(X50))|k22_seq_4(k19_seq_4(X50,X51,X52))=k22_seq_4(k19_seq_4(X50,X52,X51))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t104_seq_4])])])).
fof(c_0_63, negated_conjecture, (m2_subset_1(esk1_0,k1_numbers,k5_numbers)&(m2_finseq_2(esk2_0,k2_numbers,k14_seq_4(esk1_0))&(m1_subset_1(esk3_0,k1_zfmisc_1(k14_seq_4(esk1_0)))&(((~r2_hidden(esk2_0,esk3_0)&esk3_0!=k1_xboole_0)&v2_seq_4(esk3_0,esk1_0))&r1_xxreal_0(k24_seq_4(esk1_0,esk2_0,esk3_0),k6_numbers))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_57])])])])).
cnf(c_0_64, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_58])).
cnf(c_0_65, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_59])).
cnf(c_0_66, plain, (v1_xboole_0(np__0)), inference(split_conjunct,[status(thm)],[spc0_boole])).
cnf(c_0_67, plain, (X3=k1_xboole_0|r1_xxreal_0(k6_numbers,k24_seq_4(X1,X2,X3))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_68, plain, (r1_xxreal_0(esk4_3(X2,X3,X4),k22_seq_4(X1))|r2_hidden(X4,X3)|~m2_finseq_2(X1,k2_numbers,k14_seq_4(X2))|~r2_hidden(k15_seq_4(X2,X4,X1),X3)|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X2))|~v2_seq_4(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_69, plain, (k22_seq_4(k19_seq_4(X1,X2,X3))=k22_seq_4(k19_seq_4(X1,X3,X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))), inference(split_conjunct,[status(thm)],[c_0_62])).
fof(c_0_70, plain, ![X29, X30, X31, X32, X34]:((((m2_finseq_2(esk9_4(X29,X30,X31,X32),k2_numbers,k14_seq_4(X30))|~r2_hidden(X29,a_3_3_seq_4(X30,X31,X32))|(~m2_subset_1(X30,k1_numbers,k5_numbers)|~m2_finseq_2(X31,k2_numbers,k14_seq_4(X30))|~m1_subset_1(X32,k1_zfmisc_1(k14_seq_4(X30)))))&(X29=k22_seq_4(k19_seq_4(X30,X31,esk9_4(X29,X30,X31,X32)))|~r2_hidden(X29,a_3_3_seq_4(X30,X31,X32))|(~m2_subset_1(X30,k1_numbers,k5_numbers)|~m2_finseq_2(X31,k2_numbers,k14_seq_4(X30))|~m1_subset_1(X32,k1_zfmisc_1(k14_seq_4(X30))))))&(r2_hidden(esk9_4(X29,X30,X31,X32),X32)|~r2_hidden(X29,a_3_3_seq_4(X30,X31,X32))|(~m2_subset_1(X30,k1_numbers,k5_numbers)|~m2_finseq_2(X31,k2_numbers,k14_seq_4(X30))|~m1_subset_1(X32,k1_zfmisc_1(k14_seq_4(X30))))))&(~m2_finseq_2(X34,k2_numbers,k14_seq_4(X30))|X29!=k22_seq_4(k19_seq_4(X30,X31,X34))|~r2_hidden(X34,X32)|r2_hidden(X29,a_3_3_seq_4(X30,X31,X32))|(~m2_subset_1(X30,k1_numbers,k5_numbers)|~m2_finseq_2(X31,k2_numbers,k14_seq_4(X30))|~m1_subset_1(X32,k1_zfmisc_1(k14_seq_4(X30)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fraenkel_a_3_3_seq_4])])])])])).
fof(c_0_71, plain, ![X192, X193]:(~v1_xreal_0(X192)|(~v1_xreal_0(X193)|(~r1_xxreal_0(X192,X193)|v2_xxreal_0(X193)|~v2_xxreal_0(X192)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_real])])])])).
cnf(c_0_72, negated_conjecture, (r1_xxreal_0(k24_seq_4(esk1_0,esk2_0,esk3_0),k6_numbers)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_73, plain, (X1=k6_numbers|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_64, c_0_65])).
cnf(c_0_74, plain, (v1_xboole_0(np__0)), inference(evalgc,[status(thm)],[c_0_66])).
fof(c_0_75, plain, ![X125, X126]:(~v3_membered(X125)|(~m1_subset_1(X126,X125)|v1_xreal_0(X126))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc14_membered])])])).
fof(c_0_76, plain, ![X190, X191]:(~v1_xreal_0(X190)|(~v1_xreal_0(X191)|(~r1_xxreal_0(X190,X191)|(v1_xboole_0(X191)|v3_xxreal_0(X190)|v2_xxreal_0(X191))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_real])])])).
cnf(c_0_77, plain, (X3=k1_xboole_0|r1_xxreal_0(k6_numbers,k24_seq_4(X1,X2,X3))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))), inference(evalgc,[status(thm)],[c_0_67])).
cnf(c_0_78, plain, (r2_hidden(X4,X3)|r1_xxreal_0(esk4_3(X2,X3,X4),k22_seq_4(X1))|~v2_seq_4(X3,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X2)))|~r2_hidden(k15_seq_4(X2,X4,X1),X3)|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X2))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(X2))), inference(evalgc,[status(thm)],[c_0_68])).
cnf(c_0_79, plain, (k22_seq_4(k19_seq_4(X1,X2,X3))=k22_seq_4(k19_seq_4(X1,X3,X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))), inference(evalgc,[status(thm)],[c_0_69])).
cnf(c_0_80, plain, (X1=k22_seq_4(k19_seq_4(X2,X3,esk9_4(X1,X2,X3,X4)))|~r2_hidden(X1,a_3_3_seq_4(X2,X3,X4))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_81, plain, (v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~v2_xxreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_82, negated_conjecture, (r1_xxreal_0(k24_seq_4(esk1_0,esk2_0,esk3_0),k6_numbers)), inference(evalgc,[status(thm)],[c_0_72])).
cnf(c_0_83, plain, (k6_numbers=np__0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_73, c_0_74])])).
cnf(c_0_84, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_85, plain, (m1_subset_1(np__0,k1_numbers)), inference(split_conjunct,[status(thm)],[spc0_numerals])).
cnf(c_0_86, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
cnf(c_0_87, plain, (v1_xboole_0(X2)|v3_xxreal_0(X1)|v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_88, plain, (X1=k6_numbers|r1_xxreal_0(k6_numbers,k24_seq_4(X2,X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)), inference(rw,[status(thm)],[c_0_77, c_0_65])).
cnf(c_0_89, negated_conjecture, (esk3_0!=k1_xboole_0), inference(split_conjunct,[status(thm)],[c_0_63])).
fof(c_0_90, plain, ![X171]:((v1_xcmplx_0(k4_xcmplx_0(X171))|(v3_xxreal_0(X171)|~v1_xreal_0(X171)))&(~v2_xxreal_0(k4_xcmplx_0(X171))|(v3_xxreal_0(X171)|~v1_xreal_0(X171)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc16_xreal_0])])])])).
fof(c_0_91, plain, ![X77, X78, X79]:((~m2_subset_1(X79,X77,X78)|m1_subset_1(X79,X78)|(v1_xboole_0(X77)|v1_xboole_0(X78)|~m1_subset_1(X78,k1_zfmisc_1(X77))))&(~m1_subset_1(X79,X78)|m2_subset_1(X79,X77,X78)|(v1_xboole_0(X77)|v1_xboole_0(X78)|~m1_subset_1(X78,k1_zfmisc_1(X77))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])).
fof(c_0_92, plain, ![X53, X54]:(~v1_xboole_0(X53)|(~m1_subset_1(X54,k1_zfmisc_1(X53))|v1_xboole_0(X54))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
fof(c_0_93, plain, (~v1_xboole_0(k4_ordinal1)&v3_ordinal1(k4_ordinal1)), inference(fof_simplification,[status(thm)],[fc6_ordinal1])).
fof(c_0_94, plain, ![X160, X161]:(~v2_membered(X160)|(~m1_subset_1(X161,k1_zfmisc_1(X160))|v2_membered(X161))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc20_membered])])])).
cnf(c_0_95, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k19_seq_4(X4,X6,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k19_seq_4(X4,X5,X6)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(k19_seq_4(X4,X5,X6),k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(X6,k2_numbers,k14_seq_4(X4))|~m2_finseq_2(X5,k2_numbers,k14_seq_4(X4))|~m2_subset_1(X4,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79])])).
cnf(c_0_96, plain, (k22_seq_4(k19_seq_4(X2,X3,esk9_4(X1,X2,X3,X4)))=X1|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~r2_hidden(X1,a_3_3_seq_4(X2,X3,X4))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))), inference(evalgc,[status(thm)],[c_0_80])).
cnf(c_0_97, plain, (m2_finseq_2(esk9_4(X1,X2,X3,X4),k2_numbers,k14_seq_4(X2))|~r2_hidden(X1,a_3_3_seq_4(X2,X3,X4))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))), inference(split_conjunct,[status(thm)],[c_0_70])).
fof(c_0_98, plain, ![X41, X42, X43]:(~m2_subset_1(X41,k1_numbers,k5_numbers)|(~m2_finseq_2(X42,k2_numbers,k14_seq_4(X41))|(~m2_finseq_2(X43,k2_numbers,k14_seq_4(X41))|k15_seq_4(X41,X42,k19_seq_4(X41,X43,X42))=X43))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t80_seq_4])])])).
cnf(c_0_99, plain, (v2_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_81])).
cnf(c_0_100, negated_conjecture, (r1_xxreal_0(k24_seq_4(esk1_0,esk2_0,esk3_0),np__0)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_82, c_0_83])])).
cnf(c_0_101, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_84])).
cnf(c_0_102, plain, (m1_subset_1(np__0,k1_numbers)), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_103, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_104, plain, (v1_xboole_0(X2)|v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_87])).
cnf(c_0_105, plain, (X1=np__0|r1_xxreal_0(np__0,k24_seq_4(X2,X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_88, c_0_83]), c_0_83])])).
cnf(c_0_106, negated_conjecture, (esk3_0!=k1_xboole_0), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_107, plain, (v3_xxreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_108, plain, (k4_xcmplx_0(np__0)=np__0), inference(split_conjunct,[status(thm)],[rqRealNeg__k4_xcmplx_0__r0_r0])).
fof(c_0_109, plain, ![X55, X56, X57]:((~m2_finseq_2(X57,X55,X56)|m1_subset_1(X57,X56)|~m1_finseq_2(X56,X55))&(~m1_subset_1(X57,X56)|m2_finseq_2(X57,X55,X56)|~m1_finseq_2(X56,X55))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_m2_finseq_2])])])])).
cnf(c_0_110, plain, (m1_subset_1(X1,X3)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_111, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_112, plain, (~v1_xboole_0(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_113, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
fof(c_0_114, plain, ![X307]:((v1_xcmplx_0(k4_xcmplx_0(X307))|(v2_xxreal_0(X307)|~v1_xreal_0(X307)))&(~v3_xxreal_0(k4_xcmplx_0(X307))|(v2_xxreal_0(X307)|~v1_xreal_0(X307)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc15_xreal_0])])])])).
fof(c_0_115, plain, ![X201, X202]:(~v2_membered(X201)|(~m1_subset_1(X202,X201)|v1_xxreal_0(X202))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc13_membered])])])).
cnf(c_0_116, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_117, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
fof(c_0_118, plain, ![X236]:(~v3_membered(X236)|v2_membered(X236)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_membered])])).
cnf(c_0_119, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X6)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k19_seq_4(X4,esk9_4(X6,X4,X5,X7),X5)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(k19_seq_4(X4,esk9_4(X6,X4,X5,X7),X5),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X5,k2_numbers,k14_seq_4(X4))|~m2_finseq_2(esk9_4(X6,X4,X5,X7),k2_numbers,k14_seq_4(X4))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r2_hidden(X6,a_3_3_seq_4(X4,X5,X7))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X4)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96])])).
cnf(c_0_120, plain, (m2_finseq_2(esk9_4(X1,X2,X3,X4),k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~r2_hidden(X1,a_3_3_seq_4(X2,X3,X4))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))), inference(evalgc,[status(thm)],[c_0_97])).
cnf(c_0_121, plain, (k15_seq_4(X1,X2,k19_seq_4(X1,X3,X2))=X3|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))), inference(split_conjunct,[status(thm)],[c_0_98])).
fof(c_0_122, plain, ![X130]:((v1_xcmplx_0(k4_xcmplx_0(X130))|~v1_xreal_0(X130))&(v1_xreal_0(k4_xcmplx_0(X130))|~v1_xreal_0(X130))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc3_xreal_0])])])).
cnf(c_0_123, negated_conjecture, (v2_xxreal_0(np__0)|~v2_xxreal_0(k24_seq_4(esk1_0,esk2_0,esk3_0))|~v1_xreal_0(np__0)|~v1_xreal_0(k24_seq_4(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100])])).
cnf(c_0_124, plain, (v1_xreal_0(np__0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102]), c_0_103])])])).
cnf(c_0_125, plain, (v3_xxreal_0(np__0)|v2_xxreal_0(k24_seq_4(X1,X2,X3))|v1_xboole_0(k24_seq_4(X1,X2,X3))|X3=np__0|~v1_xreal_0(k24_seq_4(X1,X2,X3))|~v1_xreal_0(np__0)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105])])).
cnf(c_0_126, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k14_seq_4(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_127, negated_conjecture, (m2_finseq_2(esk2_0,k2_numbers,k14_seq_4(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_128, negated_conjecture, (m2_subset_1(esk1_0,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_129, negated_conjecture, (esk3_0!=k6_numbers), inference(rw,[status(thm)],[c_0_106, c_0_65])).
cnf(c_0_130, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X1))), inference(evalgc,[status(thm)],[c_0_107])).
cnf(c_0_131, plain, (k4_xcmplx_0(np__0)=np__0), inference(evalgc,[status(thm)],[c_0_108])).
fof(c_0_132, plain, ![X106, X107, X108]:(~m1_subset_1(X106,k5_numbers)|~m1_subset_1(X107,k14_seq_4(X106))|~m1_subset_1(X108,k1_zfmisc_1(k14_seq_4(X106)))|m1_subset_1(k24_seq_4(X106,X107,X108),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k24_seq_4])])).
cnf(c_0_133, plain, (m1_subset_1(X1,X3)|~m2_finseq_2(X1,X2,X3)|~m1_finseq_2(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_109])).
fof(c_0_134, plain, ![X87]:(~m1_subset_1(X87,k5_numbers)|m1_finseq_2(k14_seq_4(X87),k2_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k14_seq_4])])).
cnf(c_0_135, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_110])).
cnf(c_0_136, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_111])).
cnf(c_0_137, plain, (~v1_xboole_0(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_112])).
cnf(c_0_138, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_113])).
cnf(c_0_139, plain, (v2_xxreal_0(X1)|~v3_xxreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_114])).
cnf(c_0_140, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_115])).
cnf(c_0_141, plain, (m1_subset_1(np__0,k5_numbers)), inference(split_conjunct,[status(thm)],[spc0_numerals])).
cnf(c_0_142, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_116])).
cnf(c_0_143, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_117])).
cnf(c_0_144, plain, (v2_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_145, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k19_seq_4(X5,esk9_4(X4,X5,X6,X7),X6)),X2)|~r2_hidden(X4,a_3_3_seq_4(X5,X6,X7))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m2_finseq_2(k19_seq_4(X5,esk9_4(X4,X5,X6,X7),X6),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X6,k2_numbers,k14_seq_4(X5))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers)), inference(csr,[status(thm)],[c_0_119, c_0_120])).
cnf(c_0_146, plain, (k15_seq_4(X1,X2,k19_seq_4(X1,X3,X2))=X3|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))), inference(evalgc,[status(thm)],[c_0_121])).
fof(c_0_147, plain, ![X58, X59, X60]:(~m1_subset_1(X58,k5_numbers)|~m1_subset_1(X59,k14_seq_4(X58))|~m1_subset_1(X60,k14_seq_4(X58))|m2_finseq_2(k19_seq_4(X58,X59,X60),k2_numbers,k14_seq_4(X58))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k19_seq_4])])).
fof(c_0_148, plain, ![X69, X70, X71, X72, X76]:(((~v1_xreal_0(X71)|(~r2_hidden(X71,X69)|r1_xxreal_0(X70,X71))|X70!=k3_seq_4(X69)|~v1_xreal_0(X70)|v1_xboole_0(X69)|~v3_xxreal_2(X69)|~v3_membered(X69))&((v1_xreal_0(esk12_3(X69,X70,X72))|r1_xxreal_0(X72,k6_numbers)|~v1_xreal_0(X72)|X70!=k3_seq_4(X69)|~v1_xreal_0(X70)|v1_xboole_0(X69)|~v3_xxreal_2(X69)|~v3_membered(X69))&((r2_hidden(esk12_3(X69,X70,X72),X69)|r1_xxreal_0(X72,k6_numbers)|~v1_xreal_0(X72)|X70!=k3_seq_4(X69)|~v1_xreal_0(X70)|v1_xboole_0(X69)|~v3_xxreal_2(X69)|~v3_membered(X69))&(~r1_xxreal_0(k9_binop_2(X70,X72),esk12_3(X69,X70,X72))|r1_xxreal_0(X72,k6_numbers)|~v1_xreal_0(X72)|X70!=k3_seq_4(X69)|~v1_xreal_0(X70)|v1_xboole_0(X69)|~v3_xxreal_2(X69)|~v3_membered(X69)))))&(((v1_xreal_0(esk14_2(X69,X70))|v1_xreal_0(esk13_2(X69,X70))|X70=k3_seq_4(X69)|~v1_xreal_0(X70)|v1_xboole_0(X69)|~v3_xxreal_2(X69)|~v3_membered(X69))&((~r1_xxreal_0(esk14_2(X69,X70),k6_numbers)|v1_xreal_0(esk13_2(X69,X70))|X70=k3_seq_4(X69)|~v1_xreal_0(X70)|v1_xboole_0(X69)|~v3_xxreal_2(X69)|~v3_membered(X69))&(~v1_xreal_0(X76)|(~r2_hidden(X76,X69)|r1_xxreal_0(k9_binop_2(X70,esk14_2(X69,X70)),X76))|v1_xreal_0(esk13_2(X69,X70))|X70=k3_seq_4(X69)|~v1_xreal_0(X70)|v1_xboole_0(X69)|~v3_xxreal_2(X69)|~v3_membered(X69))))&(((v1_xreal_0(esk14_2(X69,X70))|r2_hidden(esk13_2(X69,X70),X69)|X70=k3_seq_4(X69)|~v1_xreal_0(X70)|v1_xboole_0(X69)|~v3_xxreal_2(X69)|~v3_membered(X69))&((~r1_xxreal_0(esk14_2(X69,X70),k6_numbers)|r2_hidden(esk13_2(X69,X70),X69)|X70=k3_seq_4(X69)|~v1_xreal_0(X70)|v1_xboole_0(X69)|~v3_xxreal_2(X69)|~v3_membered(X69))&(~v1_xreal_0(X76)|(~r2_hidden(X76,X69)|r1_xxreal_0(k9_binop_2(X70,esk14_2(X69,X70)),X76))|r2_hidden(esk13_2(X69,X70),X69)|X70=k3_seq_4(X69)|~v1_xreal_0(X70)|v1_xboole_0(X69)|~v3_xxreal_2(X69)|~v3_membered(X69))))&((v1_xreal_0(esk14_2(X69,X70))|~r1_xxreal_0(X70,esk13_2(X69,X70))|X70=k3_seq_4(X69)|~v1_xreal_0(X70)|v1_xboole_0(X69)|~v3_xxreal_2(X69)|~v3_membered(X69))&((~r1_xxreal_0(esk14_2(X69,X70),k6_numbers)|~r1_xxreal_0(X70,esk13_2(X69,X70))|X70=k3_seq_4(X69)|~v1_xreal_0(X70)|v1_xboole_0(X69)|~v3_xxreal_2(X69)|~v3_membered(X69))&(~v1_xreal_0(X76)|(~r2_hidden(X76,X69)|r1_xxreal_0(k9_binop_2(X70,esk14_2(X69,X70)),X76))|~r1_xxreal_0(X70,esk13_2(X69,X70))|X70=k3_seq_4(X69)|~v1_xreal_0(X70)|v1_xboole_0(X69)|~v3_xxreal_2(X69)|~v3_membered(X69))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_seq_4])])])])])])).
fof(c_0_149, plain, ![X172]:(~v1_xcmplx_0(X172)|k2_xcmplx_0(X172,k6_numbers)=X172), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_arithm])])).
cnf(c_0_150, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_122])).
cnf(c_0_151, negated_conjecture, (v2_xxreal_0(np__0)|~v2_xxreal_0(k24_seq_4(esk1_0,esk2_0,esk3_0))|~v1_xreal_0(k24_seq_4(esk1_0,esk2_0,esk3_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_123, c_0_124])])).
cnf(c_0_152, plain, (X1=np__0|v3_xxreal_0(np__0)|v2_xxreal_0(k24_seq_4(X2,X3,X1))|v1_xboole_0(k24_seq_4(X2,X3,X1))|~v1_xreal_0(k24_seq_4(X2,X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_125, c_0_124])])).
cnf(c_0_153, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k14_seq_4(esk1_0)))), inference(evalgc,[status(thm)],[c_0_126])).
cnf(c_0_154, negated_conjecture, (m2_finseq_2(esk2_0,k2_numbers,k14_seq_4(esk1_0))), inference(evalgc,[status(thm)],[c_0_127])).
cnf(c_0_155, negated_conjecture, (m2_subset_1(esk1_0,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_128])).
cnf(c_0_156, negated_conjecture, (esk3_0!=np__0), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_129, c_0_83])])).
cnf(c_0_157, plain, (v3_xxreal_0(np__0)|~v2_xxreal_0(np__0)|~v1_xreal_0(np__0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_131])])).
cnf(c_0_158, plain, (m1_subset_1(k24_seq_4(X1,X2,X3),k1_numbers)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_159, plain, (m1_subset_1(X1,X3)|~m1_finseq_2(X3,X2)|~m2_finseq_2(X1,X2,X3)), inference(evalgc,[status(thm)],[c_0_133])).
cnf(c_0_160, plain, (m1_finseq_2(k14_seq_4(X1),k2_numbers)|~m1_subset_1(X1,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_161, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m2_subset_1(X2,X3,X1)), inference(csr,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_162, plain, (~v1_xboole_0(k5_numbers)), inference(rw,[status(thm)],[c_0_137, c_0_138])).
fof(c_0_163, plain, ![X159]:(((~v1_xboole_0(X159)|(~v1_xxreal_0(X159)|~v2_xxreal_0(X159)))&(v1_xxreal_0(X159)|(~v1_xxreal_0(X159)|~v2_xxreal_0(X159))))&(~v3_xxreal_0(X159)|(~v1_xxreal_0(X159)|~v2_xxreal_0(X159)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc3_xxreal_0])])])])).
cnf(c_0_164, plain, (v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v3_xxreal_0(k4_xcmplx_0(X1))), inference(evalgc,[status(thm)],[c_0_139])).
cnf(c_0_165, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_140])).
cnf(c_0_166, plain, (m1_subset_1(np__0,k5_numbers)), inference(evalgc,[status(thm)],[c_0_141])).
cnf(c_0_167, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_143])])).
cnf(c_0_168, plain, (v2_membered(X1)|~v3_membered(X1)), inference(evalgc,[status(thm)],[c_0_144])).
cnf(c_0_169, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(esk9_4(X4,X1,X3,X5),X2)|~r2_hidden(X4,a_3_3_seq_4(X1,X3,X5))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(k19_seq_4(X1,esk9_4(X4,X1,X3,X5),X3),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(esk9_4(X4,X1,X3,X5),k2_numbers,k14_seq_4(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146])])).
cnf(c_0_170, plain, (m2_finseq_2(k19_seq_4(X1,X2,X3),k2_numbers,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X3,k14_seq_4(X1))), inference(split_conjunct,[status(thm)],[c_0_147])).
cnf(c_0_171, plain, (m2_finseq_2(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_finseq_2(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_109])).
fof(c_0_172, plain, ![X98, X99, X100]:(~r2_hidden(X98,X99)|~m1_subset_1(X99,k1_zfmisc_1(X100))|m1_subset_1(X98,X100)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_173, plain, (r1_xxreal_0(X2,k6_numbers)|v1_xboole_0(X3)|~r1_xxreal_0(k9_binop_2(X1,X2),esk12_3(X3,X1,X2))|~v1_xreal_0(X2)|X1!=k3_seq_4(X3)|~v1_xreal_0(X1)|~v3_xxreal_2(X3)|~v3_membered(X3)), inference(split_conjunct,[status(thm)],[c_0_148])).
fof(c_0_174, plain, ![X112, X113]:(~v1_xreal_0(X112)|~v1_xreal_0(X113)|k9_binop_2(X112,X113)=k2_xcmplx_0(X112,X113)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k9_binop_2])])).
fof(c_0_175, plain, ![X215, X216]:(~v1_xcmplx_0(X215)|~v1_xcmplx_0(X216)|k2_xcmplx_0(X215,X216)=k2_xcmplx_0(X216,X215)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k2_xcmplx_0])])).
cnf(c_0_176, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_149])).
cnf(c_0_177, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_150])).
cnf(c_0_178, negated_conjecture, (v2_xxreal_0(np__0)|v3_xxreal_0(np__0)|v1_xboole_0(k24_seq_4(esk1_0,esk2_0,esk3_0))|~v1_xreal_0(k24_seq_4(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152]), c_0_153]), c_0_154]), c_0_155])]), c_0_156])])).
cnf(c_0_179, plain, (v3_xxreal_0(np__0)|~v2_xxreal_0(np__0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_157, c_0_124])])])).
cnf(c_0_180, plain, (m1_subset_1(k24_seq_4(X1,X2,X3),k1_numbers)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))), inference(evalgc,[status(thm)],[c_0_158])).
cnf(c_0_181, negated_conjecture, (m1_subset_1(esk2_0,k14_seq_4(esk1_0))|~m1_finseq_2(k14_seq_4(esk1_0),k2_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_154])])).
cnf(c_0_182, plain, (m1_finseq_2(k14_seq_4(X1),k2_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_160])).
cnf(c_0_183, negated_conjecture, (m1_subset_1(esk1_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_155]), c_0_143])]), c_0_162])])).
cnf(c_0_184, plain, (~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v2_xxreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_163])).
cnf(c_0_185, plain, (v2_xxreal_0(np__0)|~v3_xxreal_0(np__0)|~v1_xreal_0(np__0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_131])])).
cnf(c_0_186, plain, (v1_xxreal_0(np__0)|~v2_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_166])])).
cnf(c_0_187, plain, (v2_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_168]), c_0_103])])])).
cnf(c_0_188, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(esk9_4(X4,X1,X3,X5),X2)|~r2_hidden(X4,a_3_3_seq_4(X1,X3,X5))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(k19_seq_4(X1,esk9_4(X4,X1,X3,X5),X3),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(csr,[status(thm)],[c_0_169, c_0_120])).
cnf(c_0_189, plain, (m2_finseq_2(k19_seq_4(X1,X2,X3),k2_numbers,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1))), inference(evalgc,[status(thm)],[c_0_170])).
cnf(c_0_190, plain, (m2_finseq_2(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_finseq_2(X2,X3)), inference(evalgc,[status(thm)],[c_0_171])).
cnf(c_0_191, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_172])).
cnf(c_0_192, plain, (v1_xboole_0(X3)|r1_xxreal_0(X2,k6_numbers)|X1!=k3_seq_4(X3)|~v3_membered(X3)|~v3_xxreal_2(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k9_binop_2(X1,X2),esk12_3(X3,X1,X2))), inference(evalgc,[status(thm)],[c_0_173])).
cnf(c_0_193, plain, (k9_binop_2(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_194, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_175])).
cnf(c_0_195, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_176])).
cnf(c_0_196, plain, (v1_xcmplx_0(np__0)|~v1_xreal_0(np__0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_131])])).
fof(c_0_197, plain, ![X20, X21, X22, X23, X24]:((X23!=k24_seq_4(X20,X21,X22)|(~m1_subset_1(X24,k1_zfmisc_1(k1_numbers))|(X24!=a_3_3_seq_4(X20,X21,X22)|X23=k5_seq_4(X24)))|~m1_subset_1(X23,k1_numbers)|~m1_subset_1(X22,k1_zfmisc_1(k14_seq_4(X20)))|~m2_finseq_2(X21,k2_numbers,k14_seq_4(X20))|~m2_subset_1(X20,k1_numbers,k5_numbers))&((m1_subset_1(esk7_4(X20,X21,X22,X23),k1_zfmisc_1(k1_numbers))|X23=k24_seq_4(X20,X21,X22)|~m1_subset_1(X23,k1_numbers)|~m1_subset_1(X22,k1_zfmisc_1(k14_seq_4(X20)))|~m2_finseq_2(X21,k2_numbers,k14_seq_4(X20))|~m2_subset_1(X20,k1_numbers,k5_numbers))&((esk7_4(X20,X21,X22,X23)=a_3_3_seq_4(X20,X21,X22)|X23=k24_seq_4(X20,X21,X22)|~m1_subset_1(X23,k1_numbers)|~m1_subset_1(X22,k1_zfmisc_1(k14_seq_4(X20)))|~m2_finseq_2(X21,k2_numbers,k14_seq_4(X20))|~m2_subset_1(X20,k1_numbers,k5_numbers))&(X23!=k5_seq_4(esk7_4(X20,X21,X22,X23))|X23=k24_seq_4(X20,X21,X22)|~m1_subset_1(X23,k1_numbers)|~m1_subset_1(X22,k1_zfmisc_1(k14_seq_4(X20)))|~m2_finseq_2(X21,k2_numbers,k14_seq_4(X20))|~m2_subset_1(X20,k1_numbers,k5_numbers))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d17_seq_4])])])])])).
cnf(c_0_198, negated_conjecture, (v3_xxreal_0(np__0)|v1_xboole_0(k24_seq_4(esk1_0,esk2_0,esk3_0))|~v1_xreal_0(k24_seq_4(esk1_0,esk2_0,esk3_0))), inference(csr,[status(thm)],[c_0_178, c_0_179])).
cnf(c_0_199, plain, (v1_xreal_0(k24_seq_4(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_180]), c_0_103])])])).
cnf(c_0_200, negated_conjecture, (m1_subset_1(esk2_0,k14_seq_4(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_182]), c_0_183])])])).
cnf(c_0_201, plain, (~v1_xxreal_0(X1)|~v1_xboole_0(X1)|~v2_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_184])).
cnf(c_0_202, plain, (v2_xxreal_0(np__0)|~v3_xxreal_0(np__0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_185, c_0_124])])).
cnf(c_0_203, plain, (v1_xxreal_0(np__0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_186, c_0_187])])).
cnf(c_0_204, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(esk9_4(X4,X1,X3,X5),X2)|~r2_hidden(X4,a_3_3_seq_4(X1,X3,X5))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(esk9_4(X4,X1,X3,X5),k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_189])])).
cnf(c_0_205, plain, (m2_finseq_2(X1,k2_numbers,k14_seq_4(X2))|~m1_subset_1(X1,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_182])])).
cnf(c_0_206, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_191])).
cnf(c_0_207, plain, (r2_hidden(esk9_4(X1,X2,X3,X4),X4)|~r2_hidden(X1,a_3_3_seq_4(X2,X3,X4))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_208, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,np__0)|X3!=k3_seq_4(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(k9_binop_2(X3,X2),esk12_3(X1,X3,X2))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_192, c_0_83])])).
cnf(c_0_209, plain, (k9_binop_2(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_193])).
cnf(c_0_210, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_194])).
cnf(c_0_211, plain, (k2_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_195, c_0_83])])).
cnf(c_0_212, plain, (v1_xcmplx_0(np__0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_196, c_0_124])])).
fof(c_0_213, plain, ![X183]:(~v1_xreal_0(X183)|v1_xcmplx_0(X183)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_xreal_0])])).
cnf(c_0_214, plain, (X1=k5_seq_4(X5)|X1!=k24_seq_4(X2,X3,X4)|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|X5!=a_3_3_seq_4(X2,X3,X4)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_197])).
cnf(c_0_215, negated_conjecture, (v3_xxreal_0(np__0)|v1_xboole_0(k24_seq_4(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_199]), c_0_153]), c_0_200]), c_0_183])])])).
cnf(c_0_216, plain, (~v3_xxreal_0(np__0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_202]), c_0_74]), c_0_203])])])).
fof(c_0_217, plain, ![X37, X38, X39]:((~m2_xxreal_2(X38,X37)|(~v1_xxreal_0(X39)|(~r2_hidden(X39,X37)|r1_xxreal_0(X38,X39)))|~v1_xxreal_0(X38)|~v2_membered(X37))&((v1_xxreal_0(esk10_2(X37,X38))|m2_xxreal_2(X38,X37)|~v1_xxreal_0(X38)|~v2_membered(X37))&((r2_hidden(esk10_2(X37,X38),X37)|m2_xxreal_2(X38,X37)|~v1_xxreal_0(X38)|~v2_membered(X37))&(~r1_xxreal_0(X38,esk10_2(X37,X38))|m2_xxreal_2(X38,X37)|~v1_xxreal_0(X38)|~v2_membered(X37))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_xxreal_2])])])])])).
cnf(c_0_218, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(esk9_4(X4,X1,X3,X5),X2)|~r2_hidden(X4,a_3_3_seq_4(X1,X3,X5))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_204, c_0_205]), c_0_206])).
cnf(c_0_219, plain, (r2_hidden(esk9_4(X1,X2,X3,X4),X4)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~r2_hidden(X1,a_3_3_seq_4(X2,X3,X4))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))), inference(evalgc,[status(thm)],[c_0_207])).
fof(c_0_220, plain, ![X364, X365]:(~v2_xxreal_0(X364)|~v1_xreal_0(X364)|v3_xxreal_0(X365)|~v1_xreal_0(X365)|v2_xxreal_0(k2_xcmplx_0(X365,X364))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc12_xreal_0])])])).
cnf(c_0_221, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,np__0)|X3!=k3_seq_4(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(k2_xcmplx_0(X3,X2),esk12_3(X1,X3,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_209])])).
cnf(c_0_222, plain, (k2_xcmplx_0(np__0,X1)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_211]), c_0_212])])])).
cnf(c_0_223, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_213])).
cnf(c_0_224, plain, (X1=k5_seq_4(X5)|X5!=a_3_3_seq_4(X2,X3,X4)|X1!=k24_seq_4(X2,X3,X4)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))), inference(evalgc,[status(thm)],[c_0_214])).
cnf(c_0_225, plain, (X1=np__0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_73, c_0_83])])).
cnf(c_0_226, negated_conjecture, (v1_xboole_0(k24_seq_4(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_215, c_0_216])])).
cnf(c_0_227, plain, (m2_xxreal_2(X1,X2)|~r1_xxreal_0(X1,esk10_2(X2,X1))|~v1_xxreal_0(X1)|~v2_membered(X2)), inference(split_conjunct,[status(thm)],[c_0_217])).
cnf(c_0_228, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(X4,a_3_3_seq_4(X1,X3,X2))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_219])])).
cnf(c_0_229, plain, (v3_xxreal_0(X2)|v2_xxreal_0(k2_xcmplx_0(X2,X1))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_220])).
fof(c_0_230, plain, ![X177]:(~v1_xreal_0(X177)|v1_xxreal_0(X177)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_xreal_0])])).
fof(c_0_231, plain, ![X129]:(~v7_ordinal1(X129)|v1_xreal_0(X129)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_xreal_0])])).
fof(c_0_232, plain, ![X148]:(~v1_xboole_0(X148)|v7_ordinal1(X148)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc7_ordinal1])])).
cnf(c_0_233, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,np__0)|k3_seq_4(X1)!=np__0|~v1_xreal_0(X2)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(X2,esk12_3(X1,np__0,X2))|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_222]), c_0_124])])])).
cnf(c_0_234, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_223])).
cnf(c_0_235, plain, (r2_hidden(esk12_3(X1,X2,X3),X1)|r1_xxreal_0(X3,k6_numbers)|v1_xboole_0(X1)|~v1_xreal_0(X3)|X2!=k3_seq_4(X1)|~v1_xreal_0(X2)|~v3_xxreal_2(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_148])).
fof(c_0_236, plain, ![X281]:(~m1_subset_1(X281,k1_zfmisc_1(k1_numbers))|k5_seq_4(X281)=k3_seq_4(X281)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k5_seq_4])])).
cnf(c_0_237, plain, (X1=k5_seq_4(a_3_3_seq_4(X2,X3,X4))|X1!=k24_seq_4(X2,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(a_3_3_seq_4(X2,X3,X4),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_224])])).
cnf(c_0_238, negated_conjecture, (k24_seq_4(esk1_0,esk2_0,esk3_0)=np__0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_226])])).
fof(c_0_239, plain, ![X92, X94]:(((v1_xreal_0(esk16_1(X92))|~v3_xxreal_2(X92)|~v2_membered(X92))&(m2_xxreal_2(esk16_1(X92),X92)|~v3_xxreal_2(X92)|~v2_membered(X92)))&(~v1_xreal_0(X94)|~m2_xxreal_2(X94,X92)|v3_xxreal_2(X92)|~v2_membered(X92))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_xxreal_2])])])])])).
cnf(c_0_240, plain, (m2_xxreal_2(X1,X2)|~v2_membered(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk10_2(X2,X1))), inference(evalgc,[status(thm)],[c_0_227])).
cnf(c_0_241, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(X4,a_3_3_seq_4(X1,X3,X2))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(csr,[status(thm)],[c_0_228, c_0_205])).
cnf(c_0_242, plain, (r2_hidden(esk10_2(X1,X2),X1)|m2_xxreal_2(X2,X1)|~v1_xxreal_0(X2)|~v2_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_217])).
cnf(c_0_243, plain, (v3_xxreal_0(X2)|v2_xxreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_229])).
cnf(c_0_244, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_230])).
cnf(c_0_245, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_231])).
cnf(c_0_246, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_232])).
cnf(c_0_247, plain, (m1_subset_1(esk7_4(X1,X2,X3,X4),k1_zfmisc_1(k1_numbers))|X4=k24_seq_4(X1,X2,X3)|~m1_subset_1(X4,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_197])).
cnf(c_0_248, plain, (esk7_4(X1,X2,X3,X4)=a_3_3_seq_4(X1,X2,X3)|X4=k24_seq_4(X1,X2,X3)|~m1_subset_1(X4,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_197])).
cnf(c_0_249, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,np__0)|k3_seq_4(X1)!=np__0|~v1_xreal_0(X2)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(X2,esk12_3(X1,np__0,X2))), inference(csr,[status(thm)],[c_0_233, c_0_234])).
cnf(c_0_250, plain, (v1_xboole_0(X1)|r1_xxreal_0(X3,k6_numbers)|r2_hidden(esk12_3(X1,X2,X3),X1)|X2!=k3_seq_4(X1)|~v3_membered(X1)|~v3_xxreal_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[c_0_235])).
cnf(c_0_251, plain, (k5_seq_4(X1)=k3_seq_4(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_236])).
cnf(c_0_252, negated_conjecture, (X1=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|X1!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_238]), c_0_153]), c_0_154]), c_0_155])])])).
cnf(c_0_253, plain, (v3_xxreal_2(X2)|~v1_xreal_0(X1)|~m2_xxreal_2(X1,X2)|~v2_membered(X2)), inference(split_conjunct,[status(thm)],[c_0_239])).
cnf(c_0_254, plain, (m2_xxreal_2(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v1_xxreal_0(esk4_3(X1,X2,X3))|~v2_membered(X4)|~v2_seq_4(X2,X1)|~r2_hidden(esk10_2(X4,esk4_3(X1,X2,X3)),a_3_3_seq_4(X1,X3,X2))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_241])])).
cnf(c_0_255, plain, (m2_xxreal_2(X2,X1)|r2_hidden(esk10_2(X1,X2),X1)|~v2_membered(X1)|~v1_xxreal_0(X2)), inference(evalgc,[status(thm)],[c_0_242])).
fof(c_0_256, plain, ![X127, X128]:(~v1_xxreal_0(X127)|~v1_xxreal_0(X128)|(r1_xxreal_0(X127,X128)|r1_xxreal_0(X128,X127))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[connectedness_r1_xxreal_0])])).
cnf(c_0_257, plain, (v3_xxreal_0(X1)|~v1_xboole_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_243])])).
cnf(c_0_258, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_244])).
cnf(c_0_259, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_245])).
cnf(c_0_260, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_246])).
cnf(c_0_261, plain, (X4=k24_seq_4(X1,X2,X3)|m1_subset_1(esk7_4(X1,X2,X3,X4),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X4,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))), inference(evalgc,[status(thm)],[c_0_247])).
cnf(c_0_262, plain, (X4=k24_seq_4(X1,X2,X3)|esk7_4(X1,X2,X3,X4)=a_3_3_seq_4(X1,X2,X3)|~m1_subset_1(X4,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))), inference(evalgc,[status(thm)],[c_0_248])).
cnf(c_0_263, plain, (m1_subset_1(esk4_3(X1,X2,X3),k1_numbers)|r2_hidden(X3,X2)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_264, plain, (v1_xboole_0(X1)|r1_xxreal_0(esk4_3(X2,X3,X4),np__0)|r2_hidden(X4,X3)|k3_seq_4(X1)!=np__0|~v1_xreal_0(esk4_3(X2,X3,X4))|~v3_xxreal_2(X1)|~v3_membered(X1)|~v2_seq_4(X3,X2)|~r2_hidden(esk12_3(X1,np__0,esk4_3(X2,X3,X4)),a_3_3_seq_4(X2,X4,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X4,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_241])])).
cnf(c_0_265, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,np__0)|r2_hidden(esk12_3(X1,X3,X2),X1)|X3!=k3_seq_4(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v3_xxreal_2(X1)|~v3_membered(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_250, c_0_83])])).
cnf(c_0_266, plain, (k5_seq_4(X1)=k3_seq_4(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_251])).
cnf(c_0_267, negated_conjecture, (k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_102])])).
cnf(c_0_268, negated_conjecture, (v2_seq_4(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_269, negated_conjecture, (~r2_hidden(esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_63])).
fof(c_0_270, plain, ![X286]:(~m1_subset_1(X286,k1_zfmisc_1(k1_numbers))|v3_membered(X286)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_membered])])).
cnf(c_0_271, plain, (v3_xxreal_2(X2)|~v2_membered(X2)|~v1_xreal_0(X1)|~m2_xxreal_2(X1,X2)), inference(evalgc,[status(thm)],[c_0_253])).
cnf(c_0_272, plain, (m2_xxreal_2(esk4_3(X1,X2,X3),a_3_3_seq_4(X1,X3,X2))|r2_hidden(X3,X2)|~v1_xxreal_0(esk4_3(X1,X2,X3))|~v2_membered(a_3_3_seq_4(X1,X3,X2))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_255])])).
cnf(c_0_273, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_256])).
fof(c_0_274, plain, ![X173]:(~v7_ordinal1(X173)|v1_xxreal_0(X173)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_xxreal_0])])).
fof(c_0_275, plain, ![X105]:(~m1_subset_1(X105,k4_ordinal1)|v7_ordinal1(X105)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_ordinal1])])).
cnf(c_0_276, plain, (v3_xxreal_0(X1)|~v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xboole_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X1,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_258])])).
cnf(c_0_277, plain, (v1_xreal_0(X1)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_260])])).
cnf(c_0_278, plain, (X1=k24_seq_4(X2,X3,X4)|m1_subset_1(a_3_3_seq_4(X2,X3,X4),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X1,k1_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_262])])).
cnf(c_0_279, plain, (r2_hidden(X3,X2)|m1_subset_1(esk4_3(X1,X2,X3),k1_numbers)|~v2_seq_4(X2,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))), inference(evalgc,[status(thm)],[c_0_263])).
cnf(c_0_280, plain, (v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|r1_xxreal_0(esk4_3(X1,X3,X2),np__0)|r2_hidden(X2,X3)|k3_seq_4(a_3_3_seq_4(X1,X2,X3))!=np__0|~v1_xreal_0(esk4_3(X1,X3,X2))|~v3_xxreal_2(a_3_3_seq_4(X1,X2,X3))|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~v2_seq_4(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_265]), c_0_124])])])).
cnf(c_0_281, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_267])])).
cnf(c_0_282, negated_conjecture, (v2_seq_4(esk3_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_268])).
cnf(c_0_283, negated_conjecture, (~r2_hidden(esk2_0,esk3_0)), inference(evalgc,[status(thm)],[c_0_269])).
cnf(c_0_284, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_270])).
cnf(c_0_285, plain, (v3_xxreal_2(a_3_3_seq_4(X1,X2,X3))|r2_hidden(X2,X3)|~v1_xreal_0(esk4_3(X1,X3,X2))|~v2_membered(a_3_3_seq_4(X1,X2,X3))|~v1_xxreal_0(esk4_3(X1,X3,X2))|~v2_seq_4(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_272])])).
cnf(c_0_286, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_273])).
cnf(c_0_287, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_274])).
cnf(c_0_288, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_275])).
fof(c_0_289, plain, ![X262, X263]:(~v1_xreal_0(X262)|(~v1_xreal_0(X263)|(r1_xxreal_0(X262,X263)|v3_xxreal_0(X263)|v2_xxreal_0(X262)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_real])])])])).
cnf(c_0_290, plain, (v3_xxreal_0(X1)|~v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xboole_0(k2_xcmplx_0(X1,X2))), inference(csr,[status(thm)],[c_0_276, c_0_277])).
cnf(c_0_291, plain, (esk4_3(X1,X2,X3)=k24_seq_4(X4,X5,X6)|m1_subset_1(a_3_3_seq_4(X4,X5,X6),k1_zfmisc_1(k1_numbers))|r2_hidden(X3,X2)|~m1_subset_1(X6,k1_zfmisc_1(k14_seq_4(X4)))|~m2_finseq_2(X5,k2_numbers,k14_seq_4(X4))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_279])])).
fof(c_0_292, plain, ![X66, X67, X68]:(~m1_subset_1(X66,k5_numbers)|~m1_subset_1(X67,k14_seq_4(X66))|~m1_subset_1(X68,k14_seq_4(X66))|k19_seq_4(X66,X67,X68)=k45_valued_1(X67,X68)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k19_seq_4])])).
cnf(c_0_293, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|r1_xxreal_0(esk4_3(esk1_0,esk3_0,esk2_0),np__0)|~v1_xreal_0(esk4_3(esk1_0,esk3_0,esk2_0))|~v3_xxreal_2(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|~v3_membered(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_281]), c_0_282]), c_0_153]), c_0_200]), c_0_183]), c_0_155])]), c_0_283])])).
cnf(c_0_294, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_284])).
cnf(c_0_295, plain, (v3_xxreal_2(a_3_3_seq_4(X1,X2,X3))|r2_hidden(X2,X3)|~v1_xreal_0(esk4_3(X1,X3,X2))|~v2_membered(a_3_3_seq_4(X1,X2,X3))|~v2_seq_4(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(csr,[status(thm)],[c_0_285, c_0_258])).
fof(c_0_296, plain, ![X143, X144, X145]:(v1_xboole_0(X143)|v1_xboole_0(X144)|~m1_subset_1(X144,k1_zfmisc_1(X143))|(~m2_subset_1(X145,X143,X144)|m1_subset_1(X145,X143))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_subset_1])])])])).
cnf(c_0_297, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_258])])).
cnf(c_0_298, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_287])).
cnf(c_0_299, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_288])).
cnf(c_0_300, plain, (r1_xxreal_0(X1,X2)|v3_xxreal_0(X2)|v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_289])).
cnf(c_0_301, plain, (v3_xxreal_0(np__0)|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v1_xboole_0(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_222]), c_0_124])])])).
cnf(c_0_302, plain, (esk4_3(X1,X2,X3)=k24_seq_4(X4,X5,X6)|r2_hidden(X3,X2)|m1_subset_1(a_3_3_seq_4(X4,X5,X6),k1_zfmisc_1(k1_numbers))|~v2_seq_4(X2,X1)|~m1_subset_1(X6,k1_zfmisc_1(k14_seq_4(X4)))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X5,k14_seq_4(X4))|~m1_subset_1(X4,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_205])])).
cnf(c_0_303, plain, (r2_hidden(X3,a_3_3_seq_4(X2,X4,X5))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(X2))|X3!=k22_seq_4(k19_seq_4(X2,X4,X1))|~r2_hidden(X1,X5)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X2))|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X2)))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_304, plain, (k19_seq_4(X1,X2,X3)=k45_valued_1(X2,X3)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X3,k14_seq_4(X1))), inference(split_conjunct,[status(thm)],[c_0_292])).
cnf(c_0_305, plain, (r2_hidden(X3,X2)|~r1_xxreal_0(esk4_3(X1,X2,X3),k6_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_306, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|r1_xxreal_0(esk4_3(esk1_0,esk3_0,esk2_0),np__0)|~v1_xreal_0(esk4_3(esk1_0,esk3_0,esk2_0))|~v3_xxreal_2(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))), inference(csr,[status(thm)],[c_0_293, c_0_294])).
cnf(c_0_307, plain, (v3_xxreal_2(a_3_3_seq_4(X1,X2,X3))|r2_hidden(X2,X3)|~v1_xreal_0(esk4_3(X1,X3,X2))|~v2_seq_4(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v3_membered(a_3_3_seq_4(X1,X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_295, c_0_168])])).
cnf(c_0_308, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m2_subset_1(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_296])).
cnf(c_0_309, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v7_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_297, c_0_298])])).
cnf(c_0_310, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_299, c_0_138])).
cnf(c_0_311, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_300])).
cnf(c_0_312, plain, (v3_xxreal_0(np__0)|~v2_xxreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xboole_0(X1)), inference(csr,[status(thm)],[c_0_301, c_0_277])).
cnf(c_0_313, plain, (esk4_3(X1,X2,X3)=k24_seq_4(X4,X5,X6)|r2_hidden(X3,X2)|m1_subset_1(a_3_3_seq_4(X4,X5,X6),k1_zfmisc_1(k1_numbers))|~v2_seq_4(X2,X1)|~m1_subset_1(X6,k1_zfmisc_1(k14_seq_4(X4)))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(X4))|~m1_subset_1(X4,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_205])])).
cnf(c_0_314, plain, (r2_hidden(X3,a_3_3_seq_4(X2,X4,X5))|X3!=k22_seq_4(k19_seq_4(X2,X4,X1))|~r2_hidden(X1,X5)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X2))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(X2))), inference(evalgc,[status(thm)],[c_0_303])).
cnf(c_0_315, plain, (k19_seq_4(X1,X2,X3)=k45_valued_1(X2,X3)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1))), inference(evalgc,[status(thm)],[c_0_304])).
cnf(c_0_316, plain, (r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~r1_xxreal_0(esk4_3(X1,X2,X3),k6_numbers)), inference(evalgc,[status(thm)],[c_0_305])).
cnf(c_0_317, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|r1_xxreal_0(esk4_3(esk1_0,esk3_0,esk2_0),np__0)|~v1_xreal_0(esk4_3(esk1_0,esk3_0,esk2_0))|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v3_membered(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_306, c_0_307]), c_0_282]), c_0_153]), c_0_200]), c_0_183]), c_0_155])]), c_0_283])])).
cnf(c_0_318, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_308])).
cnf(c_0_319, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_310])])).
cnf(c_0_320, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_277])])).
cnf(c_0_321, plain, (v3_xxreal_0(np__0)|~v2_xxreal_0(X1)|~v1_xboole_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_234])])).
cnf(c_0_322, negated_conjecture, (esk4_3(X1,X2,X3)=k24_seq_4(esk1_0,X4,esk3_0)|r2_hidden(X3,X2)|m1_subset_1(a_3_3_seq_4(esk1_0,X4,esk3_0),k1_zfmisc_1(k1_numbers))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X4,k14_seq_4(esk1_0))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_313, c_0_153]), c_0_183]), c_0_155])])])).
fof(c_0_323, plain, ![X117, X118]:(~r2_hidden(X117,X118)|~v1_xboole_0(X118)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
cnf(c_0_324, plain, (r2_hidden(X1,a_3_3_seq_4(X2,X3,X4))|X1!=k22_seq_4(k45_valued_1(X3,X5))|~r2_hidden(X5,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_finseq_2(X5,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X5,k14_seq_4(X2))|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_314, c_0_315])])).
cnf(c_0_325, plain, (r2_hidden(X1,X2)|~r1_xxreal_0(esk4_3(X3,X2,X1),np__0)|~v2_seq_4(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X3)))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(X3))|~m2_subset_1(X3,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_316, c_0_83])])).
cnf(c_0_326, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|r1_xxreal_0(esk4_3(esk1_0,esk3_0,esk2_0),np__0)|~v1_xreal_0(esk4_3(esk1_0,esk3_0,esk2_0))|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))), inference(csr,[status(thm)],[c_0_317, c_0_294])).
cnf(c_0_327, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m2_subset_1(X2,X3,X1)), inference(csr,[status(thm)],[c_0_318, c_0_136])).
fof(c_0_328, plain, ![X26, X27]:((m1_subset_1(esk8_2(X26,X27),X26)|X27=k1_xboole_0|~m1_subset_1(X27,k1_zfmisc_1(X26)))&(r2_hidden(esk8_2(X26,X27),X27)|X27=k1_xboole_0|~m1_subset_1(X27,k1_zfmisc_1(X26)))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_subset_1])])])])])).
cnf(c_0_329, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_319, c_0_166])])).
cnf(c_0_330, plain, (v1_xreal_0(esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_279]), c_0_103])])])).
cnf(c_0_331, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_320, c_0_277])])).
cnf(c_0_332, plain, (~v2_xxreal_0(X1)|~v1_xboole_0(X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[c_0_321, c_0_216]), c_0_277])).
cnf(c_0_333, negated_conjecture, (esk4_3(esk1_0,esk3_0,X1)=k24_seq_4(esk1_0,X2,esk3_0)|r2_hidden(X1,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_322, c_0_153]), c_0_282]), c_0_183]), c_0_155])])])).
cnf(c_0_334, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_323])).
cnf(c_0_335, plain, (r2_hidden(X1,a_3_3_seq_4(X2,X3,X4))|X1!=k22_seq_4(k45_valued_1(X3,X5))|~r2_hidden(X5,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_324, c_0_205]), c_0_205]), c_0_206])).
cnf(c_0_336, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|~v1_xreal_0(esk4_3(esk1_0,esk3_0,esk2_0))|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_326]), c_0_282]), c_0_153]), c_0_154]), c_0_155])]), c_0_283])])).
cnf(c_0_337, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_155]), c_0_143])]), c_0_162])])).
cnf(c_0_338, plain, (r2_hidden(esk8_2(X1,X2),X2)|X2=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_328])).
cnf(c_0_339, plain, (r1_xxreal_0(np__0,esk4_3(X1,X2,X3))|r1_xxreal_0(esk4_3(X1,X2,X3),np__0)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_330])])).
cnf(c_0_340, plain, (v3_xxreal_0(X1)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2)), inference(csr,[status(thm)],[c_0_331, c_0_332])).
cnf(c_0_341, negated_conjecture, (esk4_3(esk1_0,esk3_0,X1)=np__0|r2_hidden(X1,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_333, c_0_200]), c_0_238])])).
cnf(c_0_342, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_334])).
cnf(c_0_343, plain, (r2_hidden(k22_seq_4(k45_valued_1(X1,X2)),a_3_3_seq_4(X3,X1,X4))|~r2_hidden(X2,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X3)))|~m1_subset_1(X1,k14_seq_4(X3))|~m1_subset_1(X3,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_335])])).
cnf(c_0_344, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_330]), c_0_282]), c_0_153]), c_0_154]), c_0_155])]), c_0_283])])).
cnf(c_0_345, negated_conjecture, (k24_seq_4(X1,X2,X3)=esk1_0|m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_337])])).
cnf(c_0_346, plain, (X2=k1_xboole_0|r2_hidden(esk8_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_338])).
fof(c_0_347, plain, ![X114, X115]:((~m1_subset_1(X114,k1_zfmisc_1(X115))|r1_tarski(X114,X115))&(~r1_tarski(X114,X115)|m1_subset_1(X114,k1_zfmisc_1(X115)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_348, plain, ![X152]:r1_tarski(X152,X152), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r1_tarski])])).
cnf(c_0_349, plain, (r1_xxreal_0(np__0,esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(csr,[status(thm)],[c_0_339, c_0_325])).
cnf(c_0_350, plain, (r1_xxreal_0(X1,np__0)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_340, c_0_74]), c_0_216])])).
cnf(c_0_351, negated_conjecture, (esk4_3(esk1_0,esk3_0,esk2_0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_341, c_0_200]), c_0_283])])).
cnf(c_0_352, plain, (~v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_342, c_0_343])])).
cnf(c_0_353, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|esk1_0=np__0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_345]), c_0_238]), c_0_153]), c_0_154]), c_0_155])])])).
cnf(c_0_354, plain, (X1=k6_numbers|r2_hidden(esk8_2(X2,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(rw,[status(thm)],[c_0_346, c_0_65])).
cnf(c_0_355, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_347])).
cnf(c_0_356, plain, (r1_tarski(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_348])).
cnf(c_0_357, plain, (r2_hidden(k22_seq_4(k19_seq_4(X1,X2,X3)),a_3_3_seq_4(X1,X2,X4))|~r2_hidden(X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_314])])).
cnf(c_0_358, plain, (v2_xxreal_0(esk4_3(X1,X2,X3))|v1_xboole_0(esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|~v1_xreal_0(esk4_3(X1,X2,X3))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_349]), c_0_124])]), c_0_216])])).
cnf(c_0_359, plain, (r2_hidden(X1,X2)|~v2_seq_4(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X3)))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(X3))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_xboole_0(esk4_3(X3,X2,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_350])])).
cnf(c_0_360, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|esk4_3(esk1_0,esk3_0,esk2_0)=np__0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_351])])).
cnf(c_0_361, plain, (v3_xxreal_0(np__0)|~v2_xxreal_0(np__0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_212]), c_0_74])])])).
cnf(c_0_362, negated_conjecture, (esk1_0=np__0|~r2_hidden(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_352, c_0_353]), c_0_153]), c_0_200]), c_0_183]), c_0_155])])])).
cnf(c_0_363, plain, (X1=np__0|r2_hidden(esk8_2(X2,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_354, c_0_83])])).
cnf(c_0_364, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_355])).
cnf(c_0_365, plain, (r1_tarski(X1,X1)), inference(evalgc,[status(thm)],[c_0_356])).
cnf(c_0_366, plain, (~v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_342, c_0_357])])).
cnf(c_0_367, plain, (v2_xxreal_0(esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_358, c_0_330]), c_0_359])).
cnf(c_0_368, negated_conjecture, (a_3_3_seq_4(esk1_0,esk2_0,esk3_0)=np__0|esk4_3(esk1_0,esk3_0,esk2_0)=np__0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_360])])).
cnf(c_0_369, plain, (~v2_xxreal_0(np__0)), inference(sr,[status(thm)],[c_0_361, c_0_216])).
cnf(c_0_370, negated_conjecture, (esk1_0=np__0|~m1_subset_1(esk3_0,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_362, c_0_363]), c_0_156])])).
cnf(c_0_371, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_364, c_0_365])])).
cnf(c_0_372, plain, (~v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X4,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_205])])).
cnf(c_0_373, negated_conjecture, (a_3_3_seq_4(esk1_0,esk2_0,esk3_0)=np__0), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_368]), c_0_282]), c_0_153]), c_0_154]), c_0_155])]), c_0_369]), c_0_283])])).
cnf(c_0_374, negated_conjecture, (esk1_0=np__0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_370, c_0_371])])).
cnf(c_0_375, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[spc0_numerals])).
cnf(c_0_376, plain, (~v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X1,k5_numbers)|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(csr,[status(thm)],[c_0_372, c_0_206])).
cnf(c_0_377, negated_conjecture, (a_3_3_seq_4(np__0,esk2_0,esk3_0)=np__0), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_373, c_0_374])])).
cnf(c_0_378, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k14_seq_4(np__0)))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_153, c_0_374])])).
cnf(c_0_379, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_375])).
cnf(c_0_380, negated_conjecture, (~r2_hidden(X1,esk3_0)|~m2_finseq_2(esk2_0,k2_numbers,k14_seq_4(np__0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_376, c_0_377]), c_0_74]), c_0_378]), c_0_166]), c_0_379])])])).
cnf(c_0_381, negated_conjecture, (m2_finseq_2(esk2_0,k2_numbers,k14_seq_4(np__0))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_154, c_0_374])])).
cnf(c_0_382, negated_conjecture, (~r2_hidden(X1,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_380, c_0_381])])).
cnf(c_0_383, negated_conjecture, (~m1_subset_1(esk3_0,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_363]), c_0_156])])).
cnf(c_0_384, negated_conjecture, ($false), inference(sr,[status(thm)],[c_0_378, c_0_383]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 385
# Proof object clause steps            : 279
# Proof object formula steps           : 106
# Proof object conjectures             : 59
# Proof object clause conjectures      : 56
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 74
# Proof object initial formulas used   : 56
# Proof object generating inferences   : 84
# Proof object simplifying inferences  : 165
# Training examples: 172 positive, 1976 negative
# Training: Positive examples begin
cnf(c_0_65, plain, (k1_xboole_0=k6_numbers)).# trainpos
cnf(c_0_138, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_106, negated_conjecture, (esk3_0!=k1_xboole_0)).# trainpos
cnf(c_0_131, plain, (k4_xcmplx_0(np__0)=np__0)).# trainpos
cnf(c_0_74, plain, (v1_xboole_0(np__0))).# trainpos
cnf(c_0_283, negated_conjecture, (~r2_hidden(esk2_0,esk3_0))).# trainpos
cnf(c_0_282, negated_conjecture, (v2_seq_4(esk3_0,esk1_0))).# trainpos
cnf(c_0_103, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_155, negated_conjecture, (m2_subset_1(esk1_0,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_153, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k14_seq_4(esk1_0))))).# trainpos
cnf(c_0_342, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_379, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_137, plain, (~v1_xboole_0(k4_ordinal1))).# trainpos
cnf(c_0_168, plain, (v2_membered(X1)|~v3_membered(X1))).# trainpos
cnf(c_0_298, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_154, negated_conjecture, (m2_finseq_2(esk2_0,k2_numbers,k14_seq_4(esk1_0)))).# trainpos
cnf(c_0_82, negated_conjecture, (r1_xxreal_0(k24_seq_4(esk1_0,esk2_0,esk3_0),k6_numbers))).# trainpos
cnf(c_0_258, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_64, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_102, plain, (m1_subset_1(np__0,k1_numbers))).# trainpos
cnf(c_0_166, plain, (m1_subset_1(np__0,k5_numbers))).# trainpos
cnf(c_0_260, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_234, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_364, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_206, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_201, plain, (~v1_xxreal_0(X1)|~v1_xboole_0(X1)|~v2_xxreal_0(X1))).# trainpos
cnf(c_0_365, plain, (r1_tarski(X1,X1))).# trainpos
cnf(c_0_143, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_259, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_159, plain, (m1_subset_1(X1,X3)|~m1_finseq_2(X3,X2)|~m2_finseq_2(X1,X2,X3))).# trainpos
cnf(c_0_177, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_299, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).# trainpos
cnf(c_0_316, plain, (r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~r1_xxreal_0(esk4_3(X1,X2,X3),k6_numbers))).# trainpos
cnf(c_0_195, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_294, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_190, plain, (m2_finseq_2(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_finseq_2(X2,X3))).# trainpos
cnf(c_0_136, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_165, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_182, plain, (m1_finseq_2(k14_seq_4(X1),k2_numbers)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_266, plain, (k5_seq_4(X1)=k3_seq_4(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_164, plain, (v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v3_xxreal_0(k4_xcmplx_0(X1)))).# trainpos
cnf(c_0_130, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X1)))).# trainpos
cnf(c_0_101, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_142, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_286, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_311, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_209, plain, (k9_binop_2(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_210, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_99, plain, (v2_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_271, plain, (v3_xxreal_2(X2)|~v2_membered(X2)|~v1_xreal_0(X1)|~m2_xxreal_2(X1,X2))).# trainpos
cnf(c_0_104, plain, (v1_xboole_0(X2)|v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_346, plain, (X2=k1_xboole_0|r2_hidden(esk8_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_135, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_243, plain, (v3_xxreal_0(X2)|v2_xxreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1))).# trainpos
cnf(c_0_255, plain, (m2_xxreal_2(X2,X1)|r2_hidden(esk10_2(X1,X2),X1)|~v2_membered(X1)|~v1_xxreal_0(X2))).# trainpos
cnf(c_0_240, plain, (m2_xxreal_2(X1,X2)|~v2_membered(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,esk10_2(X2,X1)))).# trainpos
cnf(c_0_318, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_180, plain, (m1_subset_1(k24_seq_4(X1,X2,X3),k1_numbers)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1))))).# trainpos
cnf(c_0_77, plain, (X3=k1_xboole_0|r1_xxreal_0(k6_numbers,k24_seq_4(X1,X2,X3))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1)))).# trainpos
cnf(c_0_315, plain, (k19_seq_4(X1,X2,X3)=k45_valued_1(X2,X3)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1)))).# trainpos
cnf(c_0_192, plain, (v1_xboole_0(X3)|r1_xxreal_0(X2,k6_numbers)|X1!=k3_seq_4(X3)|~v3_membered(X3)|~v3_xxreal_2(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k9_binop_2(X1,X2),esk12_3(X3,X1,X2)))).# trainpos
cnf(c_0_189, plain, (m2_finseq_2(k19_seq_4(X1,X2,X3),k2_numbers,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1)))).# trainpos
cnf(c_0_146, plain, (k15_seq_4(X1,X2,k19_seq_4(X1,X3,X2))=X3|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1)))).# trainpos
cnf(c_0_79, plain, (k22_seq_4(k19_seq_4(X1,X2,X3))=k22_seq_4(k19_seq_4(X1,X3,X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1)))).# trainpos
cnf(c_0_314, plain, (r2_hidden(X3,a_3_3_seq_4(X2,X4,X5))|X3!=k22_seq_4(k19_seq_4(X2,X4,X1))|~r2_hidden(X1,X5)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X2))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(X2)))).# trainpos
cnf(c_0_224, plain, (X1=k5_seq_4(X5)|X5!=a_3_3_seq_4(X2,X3,X4)|X1!=k24_seq_4(X2,X3,X4)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2)))).# trainpos
cnf(c_0_250, plain, (v1_xboole_0(X1)|r1_xxreal_0(X3,k6_numbers)|r2_hidden(esk12_3(X1,X2,X3),X1)|X2!=k3_seq_4(X1)|~v3_membered(X1)|~v3_xxreal_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_279, plain, (r2_hidden(X3,X2)|m1_subset_1(esk4_3(X1,X2,X3),k1_numbers)|~v2_seq_4(X2,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1)))).# trainpos
cnf(c_0_261, plain, (X4=k24_seq_4(X1,X2,X3)|m1_subset_1(esk7_4(X1,X2,X3,X4),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X4,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1)))).# trainpos
cnf(c_0_219, plain, (r2_hidden(esk9_4(X1,X2,X3,X4),X4)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~r2_hidden(X1,a_3_3_seq_4(X2,X3,X4))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2)))).# trainpos
cnf(c_0_120, plain, (m2_finseq_2(esk9_4(X1,X2,X3,X4),k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~r2_hidden(X1,a_3_3_seq_4(X2,X3,X4))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2)))).# trainpos
cnf(c_0_96, plain, (k22_seq_4(k19_seq_4(X2,X3,esk9_4(X1,X2,X3,X4)))=X1|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~r2_hidden(X1,a_3_3_seq_4(X2,X3,X4))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2)))).# trainpos
cnf(c_0_262, plain, (X4=k24_seq_4(X1,X2,X3)|esk7_4(X1,X2,X3,X4)=a_3_3_seq_4(X1,X2,X3)|~m1_subset_1(X4,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1)))).# trainpos
cnf(c_0_78, plain, (r2_hidden(X4,X3)|r1_xxreal_0(esk4_3(X2,X3,X4),k22_seq_4(X1))|~v2_seq_4(X3,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X2)))|~r2_hidden(k15_seq_4(X2,X4,X1),X3)|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X2))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(X2)))).# trainpos
cnf(c_0_83, plain, (k6_numbers=np__0)).# trainpos
cnf(c_0_156, negated_conjecture, (esk3_0!=np__0)).# trainpos
cnf(c_0_371, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_325, plain, (r2_hidden(X1,X2)|~r1_xxreal_0(esk4_3(X3,X2,X1),np__0)|~v2_seq_4(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X3)))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(X3))|~m2_subset_1(X3,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_363, plain, (X1=np__0|r2_hidden(esk8_2(X2,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_95, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k19_seq_4(X4,X6,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k19_seq_4(X4,X5,X6)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(k19_seq_4(X4,X5,X6),k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(X6,k2_numbers,k14_seq_4(X4))|~m2_finseq_2(X5,k2_numbers,k14_seq_4(X4))|~m2_subset_1(X4,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_183, negated_conjecture, (m1_subset_1(esk1_0,k5_numbers))).# trainpos
cnf(c_0_225, plain, (X1=np__0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_100, negated_conjecture, (r1_xxreal_0(k24_seq_4(esk1_0,esk2_0,esk3_0),np__0))).# trainpos
cnf(c_0_357, plain, (r2_hidden(k22_seq_4(k19_seq_4(X1,X2,X3)),a_3_3_seq_4(X1,X2,X4))|~r2_hidden(X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_105, plain, (X1=np__0|r1_xxreal_0(np__0,k24_seq_4(X2,X3,X1))|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_330, plain, (v1_xreal_0(esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_199, plain, (v1_xreal_0(k24_seq_4(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_277, plain, (v1_xreal_0(X1)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_337, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers))).# trainpos
cnf(c_0_181, negated_conjecture, (m1_subset_1(esk2_0,k14_seq_4(esk1_0))|~m1_finseq_2(k14_seq_4(esk1_0),k2_numbers))).# trainpos
cnf(c_0_205, plain, (m2_finseq_2(X1,k2_numbers,k14_seq_4(X2))|~m1_subset_1(X1,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_237, plain, (X1=k5_seq_4(a_3_3_seq_4(X2,X3,X4))|X1!=k24_seq_4(X2,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(a_3_3_seq_4(X2,X3,X4),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_324, plain, (r2_hidden(X1,a_3_3_seq_4(X2,X3,X4))|X1!=k22_seq_4(k45_valued_1(X3,X5))|~r2_hidden(X5,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_finseq_2(X5,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X5,k14_seq_4(X2))|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_157, plain, (v3_xxreal_0(np__0)|~v2_xxreal_0(np__0)|~v1_xreal_0(np__0))).# trainpos
cnf(c_0_167, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers))).# trainpos
cnf(c_0_124, plain, (v1_xreal_0(np__0))).# trainpos
cnf(c_0_208, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,np__0)|X3!=k3_seq_4(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(k9_binop_2(X3,X2),esk12_3(X1,X3,X2)))).# trainpos
cnf(c_0_196, plain, (v1_xcmplx_0(np__0)|~v1_xreal_0(np__0))).# trainpos
cnf(c_0_265, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,np__0)|r2_hidden(esk12_3(X1,X3,X2),X1)|X3!=k3_seq_4(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v3_xxreal_2(X1)|~v3_membered(X1))).# trainpos
cnf(c_0_211, plain, (k2_xcmplx_0(X1,np__0)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_257, plain, (v3_xxreal_0(X1)|~v1_xboole_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_297, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_278, plain, (X1=k24_seq_4(X2,X3,X4)|m1_subset_1(a_3_3_seq_4(X2,X3,X4),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X1,k1_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_200, negated_conjecture, (m1_subset_1(esk2_0,k14_seq_4(esk1_0)))).# trainpos
cnf(c_0_179, plain, (v3_xxreal_0(np__0)|~v2_xxreal_0(np__0))).# trainpos
cnf(c_0_187, plain, (v2_membered(k5_numbers))).# trainpos
cnf(c_0_222, plain, (k2_xcmplx_0(np__0,X1)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_343, plain, (r2_hidden(k22_seq_4(k45_valued_1(X1,X2)),a_3_3_seq_4(X3,X1,X4))|~r2_hidden(X2,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X3)))|~m1_subset_1(X1,k14_seq_4(X3))|~m1_subset_1(X3,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_309, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v7_ordinal1(X2))).# trainpos
cnf(c_0_366, plain, (~v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_291, plain, (esk4_3(X1,X2,X3)=k24_seq_4(X4,X5,X6)|m1_subset_1(a_3_3_seq_4(X4,X5,X6),k1_zfmisc_1(k1_numbers))|r2_hidden(X3,X2)|~m1_subset_1(X6,k1_zfmisc_1(k14_seq_4(X4)))|~m2_finseq_2(X5,k2_numbers,k14_seq_4(X4))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_125, plain, (v3_xxreal_0(np__0)|v2_xxreal_0(k24_seq_4(X1,X2,X3))|v1_xboole_0(k24_seq_4(X1,X2,X3))|X3=np__0|~v1_xreal_0(k24_seq_4(X1,X2,X3))|~v1_xreal_0(np__0)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_123, negated_conjecture, (v2_xxreal_0(np__0)|~v2_xxreal_0(k24_seq_4(esk1_0,esk2_0,esk3_0))|~v1_xreal_0(np__0)|~v1_xreal_0(k24_seq_4(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_186, plain, (v1_xxreal_0(np__0)|~v2_membered(k5_numbers))).# trainpos
cnf(c_0_221, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,np__0)|X3!=k3_seq_4(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(k2_xcmplx_0(X3,X2),esk12_3(X1,X3,X2)))).# trainpos
cnf(c_0_276, plain, (v3_xxreal_0(X1)|~v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xboole_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X1,X2)))).# trainpos
cnf(c_0_119, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X6)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k19_seq_4(X4,esk9_4(X6,X4,X5,X7),X5)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(k19_seq_4(X4,esk9_4(X6,X4,X5,X7),X5),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X5,k2_numbers,k14_seq_4(X4))|~m2_finseq_2(esk9_4(X6,X4,X5,X7),k2_numbers,k14_seq_4(X4))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r2_hidden(X6,a_3_3_seq_4(X4,X5,X7))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X4))))).# trainpos
cnf(c_0_345, negated_conjecture, (k24_seq_4(X1,X2,X3)=esk1_0|m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_178, negated_conjecture, (v2_xxreal_0(np__0)|v3_xxreal_0(np__0)|v1_xboole_0(k24_seq_4(esk1_0,esk2_0,esk3_0))|~v1_xreal_0(k24_seq_4(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_301, plain, (v3_xxreal_0(np__0)|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v1_xboole_0(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_302, plain, (esk4_3(X1,X2,X3)=k24_seq_4(X4,X5,X6)|r2_hidden(X3,X2)|m1_subset_1(a_3_3_seq_4(X4,X5,X6),k1_zfmisc_1(k1_numbers))|~v2_seq_4(X2,X1)|~m1_subset_1(X6,k1_zfmisc_1(k14_seq_4(X4)))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X5,k14_seq_4(X4))|~m1_subset_1(X4,k5_numbers))).# trainpos
cnf(c_0_352, plain, (~v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_372, plain, (~v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X4,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_169, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(esk9_4(X4,X1,X3,X5),X2)|~r2_hidden(X4,a_3_3_seq_4(X1,X3,X5))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(k19_seq_4(X1,esk9_4(X4,X1,X3,X5),X3),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(esk9_4(X4,X1,X3,X5),k2_numbers,k14_seq_4(X1)))).# trainpos
cnf(c_0_185, plain, (v2_xxreal_0(np__0)|~v3_xxreal_0(np__0)|~v1_xreal_0(np__0))).# trainpos
cnf(c_0_215, negated_conjecture, (v3_xxreal_0(np__0)|v1_xboole_0(k24_seq_4(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_216, plain, (~v3_xxreal_0(np__0))).# trainpos
cnf(c_0_361, plain, (v3_xxreal_0(np__0)|~v2_xxreal_0(np__0))).# trainpos
cnf(c_0_321, plain, (v3_xxreal_0(np__0)|~v2_xxreal_0(X1)|~v1_xboole_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_204, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(esk9_4(X4,X1,X3,X5),X2)|~r2_hidden(X4,a_3_3_seq_4(X1,X3,X5))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(esk9_4(X4,X1,X3,X5),k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_313, plain, (esk4_3(X1,X2,X3)=k24_seq_4(X4,X5,X6)|r2_hidden(X3,X2)|m1_subset_1(a_3_3_seq_4(X4,X5,X6),k1_zfmisc_1(k1_numbers))|~v2_seq_4(X2,X1)|~m1_subset_1(X6,k1_zfmisc_1(k14_seq_4(X4)))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(X4))|~m1_subset_1(X4,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_233, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,np__0)|k3_seq_4(X1)!=np__0|~v1_xreal_0(X2)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(X2,esk12_3(X1,np__0,X2))|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_226, negated_conjecture, (v1_xboole_0(k24_seq_4(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_228, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(X4,a_3_3_seq_4(X1,X3,X2))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1)))).# trainpos
cnf(c_0_322, negated_conjecture, (esk4_3(X1,X2,X3)=k24_seq_4(esk1_0,X4,esk3_0)|r2_hidden(X3,X2)|m1_subset_1(a_3_3_seq_4(esk1_0,X4,esk3_0),k1_zfmisc_1(k1_numbers))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X4,k14_seq_4(esk1_0))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_320, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_319, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_238, negated_conjecture, (k24_seq_4(esk1_0,esk2_0,esk3_0)=np__0)).# trainpos
cnf(c_0_254, plain, (m2_xxreal_2(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v1_xxreal_0(esk4_3(X1,X2,X3))|~v2_membered(X4)|~v2_seq_4(X2,X1)|~r2_hidden(esk10_2(X4,esk4_3(X1,X2,X3)),a_3_3_seq_4(X1,X3,X2))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_264, plain, (v1_xboole_0(X1)|r1_xxreal_0(esk4_3(X2,X3,X4),np__0)|r2_hidden(X4,X3)|k3_seq_4(X1)!=np__0|~v1_xreal_0(esk4_3(X2,X3,X4))|~v3_xxreal_2(X1)|~v3_membered(X1)|~v2_seq_4(X3,X2)|~r2_hidden(esk12_3(X1,np__0,esk4_3(X2,X3,X4)),a_3_3_seq_4(X2,X4,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X4,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_333, negated_conjecture, (esk4_3(esk1_0,esk3_0,X1)=k24_seq_4(esk1_0,X2,esk3_0)|r2_hidden(X1,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).# trainpos
cnf(c_0_331, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).# trainpos
cnf(c_0_329, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_272, plain, (m2_xxreal_2(esk4_3(X1,X2,X3),a_3_3_seq_4(X1,X3,X2))|r2_hidden(X3,X2)|~v1_xxreal_0(esk4_3(X1,X2,X3))|~v2_membered(a_3_3_seq_4(X1,X3,X2))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_280, plain, (v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|r1_xxreal_0(esk4_3(X1,X3,X2),np__0)|r2_hidden(X2,X3)|k3_seq_4(a_3_3_seq_4(X1,X2,X3))!=np__0|~v1_xreal_0(esk4_3(X1,X3,X2))|~v3_xxreal_2(a_3_3_seq_4(X1,X2,X3))|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~v2_seq_4(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_252, negated_conjecture, (X1=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|X1!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_350, plain, (r1_xxreal_0(X1,np__0)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_341, negated_conjecture, (esk4_3(esk1_0,esk3_0,X1)=np__0|r2_hidden(X1,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).# trainpos
cnf(c_0_339, plain, (r1_xxreal_0(np__0,esk4_3(X1,X2,X3))|r1_xxreal_0(esk4_3(X1,X2,X3),np__0)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_285, plain, (v3_xxreal_2(a_3_3_seq_4(X1,X2,X3))|r2_hidden(X2,X3)|~v1_xreal_0(esk4_3(X1,X3,X2))|~v2_membered(a_3_3_seq_4(X1,X2,X3))|~v1_xxreal_0(esk4_3(X1,X3,X2))|~v2_seq_4(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_359, plain, (r2_hidden(X1,X2)|~v2_seq_4(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X3)))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(X3))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_xboole_0(esk4_3(X3,X2,X1)))).# trainpos
cnf(c_0_267, negated_conjecture, (k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_358, plain, (v2_xxreal_0(esk4_3(X1,X2,X3))|v1_xboole_0(esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|~v1_xreal_0(esk4_3(X1,X2,X3))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_351, negated_conjecture, (esk4_3(esk1_0,esk3_0,esk2_0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_281, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_307, plain, (v3_xxreal_2(a_3_3_seq_4(X1,X2,X3))|r2_hidden(X2,X3)|~v1_xreal_0(esk4_3(X1,X3,X2))|~v2_seq_4(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v3_membered(a_3_3_seq_4(X1,X2,X3)))).# trainpos
cnf(c_0_293, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|r1_xxreal_0(esk4_3(esk1_0,esk3_0,esk2_0),np__0)|~v1_xreal_0(esk4_3(esk1_0,esk3_0,esk2_0))|~v3_xxreal_2(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|~v3_membered(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_317, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|r1_xxreal_0(esk4_3(esk1_0,esk3_0,esk2_0),np__0)|~v1_xreal_0(esk4_3(esk1_0,esk3_0,esk2_0))|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v3_membered(a_3_3_seq_4(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_336, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|~v1_xreal_0(esk4_3(esk1_0,esk3_0,esk2_0))|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_344, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_360, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|esk4_3(esk1_0,esk3_0,esk2_0)=np__0)).# trainpos
cnf(c_0_353, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|esk1_0=np__0)).# trainpos
cnf(c_0_368, negated_conjecture, (a_3_3_seq_4(esk1_0,esk2_0,esk3_0)=np__0|esk4_3(esk1_0,esk3_0,esk2_0)=np__0)).# trainpos
cnf(c_0_362, negated_conjecture, (esk1_0=np__0|~r2_hidden(X1,esk3_0))).# trainpos
cnf(c_0_373, negated_conjecture, (a_3_3_seq_4(esk1_0,esk2_0,esk3_0)=np__0)).# trainpos
cnf(c_0_370, negated_conjecture, (esk1_0=np__0|~m1_subset_1(esk3_0,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_374, negated_conjecture, (esk1_0=np__0)).# trainpos
cnf(c_0_378, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k14_seq_4(np__0))))).# trainpos
cnf(c_0_377, negated_conjecture, (a_3_3_seq_4(np__0,esk2_0,esk3_0)=np__0)).# trainpos
cnf(c_0_381, negated_conjecture, (m2_finseq_2(esk2_0,k2_numbers,k14_seq_4(np__0)))).# trainpos
cnf(c_0_380, negated_conjecture, (~r2_hidden(X1,esk3_0)|~m2_finseq_2(esk2_0,k2_numbers,k14_seq_4(np__0)))).# trainpos
cnf(c_0_383, negated_conjecture, (~m1_subset_1(esk3_0,k1_zfmisc_1(X1)))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_147, plain, (o_0_0_xboole_0=k1_xboole_0)).#trainneg
cnf(i_0_74, plain, (v1_xxreal_0(esk15_0))).#trainneg
cnf(i_0_96, plain, (v1_xxreal_0(esk17_0))).#trainneg
cnf(i_0_221, plain, (v1_xxreal_0(esk22_0))).#trainneg
cnf(i_0_240, plain, (v1_xxreal_0(esk24_0))).#trainneg
cnf(i_0_413, plain, (v1_xxreal_0(esk39_0))).#trainneg
cnf(i_0_89, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_181, plain, (~r1_xxreal_0(np__1,np__0))).#trainneg
cnf(i_0_76, plain, (v1_xboole_0(esk15_0))).#trainneg
cnf(i_0_419, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_81, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_182, plain, (~v1_xboole_0(k2_numbers))).#trainneg
cnf(i_0_120, plain, (v1_membered(k2_numbers))).#trainneg
cnf(i_0_97, plain, (v1_xboole_0(esk17_0))).#trainneg
cnf(i_0_243, plain, (v1_xboole_0(esk25_0))).#trainneg
cnf(i_0_41, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_224, plain, (v3_ordinal1(esk22_0))).#trainneg
cnf(i_0_508, plain, (v3_ordinal1(esk46_0))).#trainneg
cnf(i_0_44, plain, (v6_membered(k4_ordinal1))).#trainneg
cnf(i_0_223, plain, (v7_ordinal1(esk22_0))).#trainneg
cnf(i_0_371, plain, (v7_ordinal1(esk35_0))).#trainneg
cnf(i_0_73, plain, (v1_xreal_0(esk15_0))).#trainneg
cnf(i_0_220, plain, (v1_xreal_0(esk22_0))).#trainneg
cnf(i_0_411, plain, (v1_xreal_0(esk39_0))).#trainneg
cnf(i_0_75, plain, (v1_xcmplx_0(esk15_0))).#trainneg
cnf(i_0_222, plain, (v1_xcmplx_0(esk22_0))).#trainneg
cnf(i_0_414, plain, (v1_xcmplx_0(esk39_0))).#trainneg
cnf(i_0_114, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_218, plain, (v1_card_1(esk22_0))).#trainneg
cnf(i_0_384, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_291, plain, (v1_relat_1(esk27_0))).#trainneg
cnf(i_0_330, plain, (v1_relat_1(esk30_0))).#trainneg
cnf(i_0_344, plain, (v1_relat_1(esk31_0))).#trainneg
cnf(i_0_458, plain, (v1_relat_1(esk42_0))).#trainneg
cnf(i_0_541, plain, (v1_relat_1(esk48_0))).#trainneg
cnf(i_0_239, plain, (v3_xxreal_0(esk24_0))).#trainneg
cnf(i_0_412, plain, (v3_xxreal_0(esk39_0))).#trainneg
cnf(i_0_196, plain, (v1_funct_1(k8_seq_4))).#trainneg
cnf(i_0_288, plain, (v1_funct_1(esk27_0))).#trainneg
cnf(i_0_327, plain, (v1_funct_1(esk30_0))).#trainneg
cnf(i_0_341, plain, (v1_funct_1(esk31_0))).#trainneg
cnf(i_0_456, plain, (v1_funct_1(esk42_0))).#trainneg
cnf(i_0_538, plain, (v1_funct_1(esk48_0))).#trainneg
cnf(i_0_285, plain, (v1_finseq_1(esk27_0))).#trainneg
cnf(i_0_219, plain, (v1_finset_1(esk22_0))).#trainneg
cnf(i_0_435, plain, (~r2_hidden(X3,X2)|~r2_hidden(X1,X2)|~r2_hidden(X1,esk40_2(X3,X2)))).#trainneg
cnf(i_0_286, plain, (v1_finset_1(esk27_0))).#trainneg
cnf(i_0_113, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_227, plain, (~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_287, plain, (~v1_xboole_0(esk27_0))).#trainneg
cnf(i_0_480, plain, (v3_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_452, plain, (v1_finset_1(esk41_0))).#trainneg
cnf(i_0_454, plain, (v1_finset_1(esk42_0))).#trainneg
cnf(i_0_324, plain, (v1_valued_0(esk30_0))).#trainneg
cnf(i_0_337, plain, (v1_valued_0(esk31_0))).#trainneg
cnf(i_0_534, plain, (v1_valued_0(esk48_0))).#trainneg
cnf(i_0_226, plain, (v1_ordinal1(esk22_0))).#trainneg
cnf(i_0_225, plain, (v2_ordinal1(esk22_0))).#trainneg
cnf(i_0_322, plain, (v3_valued_0(esk30_0))).#trainneg
cnf(i_0_335, plain, (v3_valued_0(esk31_0))).#trainneg
cnf(i_0_532, plain, (v3_valued_0(esk48_0))).#trainneg
cnf(i_0_290, plain, (v2_relat_1(esk27_0))).#trainneg
cnf(i_0_284, plain, (v2_finseq_1(esk27_0))).#trainneg
cnf(i_0_323, plain, (v2_valued_0(esk30_0))).#trainneg
cnf(i_0_336, plain, (v2_valued_0(esk31_0))).#trainneg
cnf(i_0_85, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_331, plain, (~v1_xboole_0(esk30_0))).#trainneg
cnf(i_0_533, plain, (v2_valued_0(esk48_0))).#trainneg
cnf(i_0_17, plain, (X1=k24_seq_4(X2,X3,X4)|k5_seq_4(esk7_4(X2,X3,X4,X1))!=X1|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2)))).#trainneg
cnf(i_0_321, plain, (v4_valued_0(esk30_0))).#trainneg
cnf(i_0_195, plain, (v1_funct_2(k8_seq_4,k2_numbers,k1_numbers))).#trainneg
cnf(i_0_340, plain, (~v1_xboole_0(esk31_0))).#trainneg
cnf(i_0_372, plain, (~v1_xboole_0(esk35_0))).#trainneg
cnf(i_0_520, plain, (v3_ordinal1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_334, plain, (v4_valued_0(esk31_0))).#trainneg
cnf(i_0_531, plain, (v4_valued_0(esk48_0))).#trainneg
cnf(i_0_320, plain, (v5_valued_0(esk30_0))).#trainneg
cnf(i_0_333, plain, (v5_valued_0(esk31_0))).#trainneg
cnf(i_0_530, plain, (v5_valued_0(esk48_0))).#trainneg
cnf(i_0_451, plain, (v5_finset_1(esk41_0))).#trainneg
cnf(i_0_453, plain, (~v1_xboole_0(esk41_0))).#trainneg
cnf(i_0_455, plain, (~v1_xboole_0(esk42_0))).#trainneg
cnf(i_0_509, plain, (~v1_xboole_0(esk46_0))).#trainneg
cnf(i_0_537, plain, (~v1_xboole_0(esk48_0))).#trainneg
cnf(i_0_346, plain, (~v1_finset_1(k4_ordinal1))).#trainneg
cnf(i_0_264, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_137, plain, (v1_xboole_0(esk19_1(X1)))).#trainneg
cnf(i_0_313, plain, (v1_xboole_0(esk29_1(X1)))).#trainneg
cnf(i_0_189, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_361, plain, (~v1_xboole_0(k3_finseq_2(X1)))).#trainneg
cnf(i_0_145, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_156, plain, (~v3_xxreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_325, plain, (v1_funct_2(esk30_0,k5_numbers,k5_numbers))).#trainneg
cnf(i_0_130, plain, (v1_card_1(k1_card_1(X1)))).#trainneg
cnf(i_0_262, plain, (v1_relat_1(esk26_1(X1)))).#trainneg
cnf(i_0_317, plain, (v1_relat_1(esk29_1(X1)))).#trainneg
cnf(i_0_360, plain, (v1_relat_1(esk33_1(X1)))).#trainneg
cnf(i_0_367, plain, (v1_relat_1(esk34_1(X1)))).#trainneg
cnf(i_0_378, plain, (v1_relat_1(esk36_1(X1)))).#trainneg
cnf(i_0_400, plain, (v1_relat_1(esk37_1(X1)))).#trainneg
cnf(i_0_563, plain, (v1_relat_1(esk49_1(X1)))).#trainneg
cnf(i_0_259, plain, (v1_funct_1(esk26_1(X1)))).#trainneg
cnf(i_0_314, plain, (v1_funct_1(esk29_1(X1)))).#trainneg
cnf(i_0_358, plain, (v1_funct_1(esk33_1(X1)))).#trainneg
cnf(i_0_364, plain, (v1_funct_1(esk34_1(X1)))).#trainneg
cnf(i_0_376, plain, (v1_funct_1(esk36_1(X1)))).#trainneg
cnf(i_0_398, plain, (v1_funct_1(esk37_1(X1)))).#trainneg
cnf(i_0_468, plain, (~v1_xboole_0(esk43_1(X1)))).#trainneg
cnf(i_0_111, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_121, plain, (~v7_ordinal1(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_561, plain, (v1_funct_1(esk49_1(X1)))).#trainneg
cnf(i_0_338, plain, (v1_funct_2(esk31_0,k5_numbers,k5_numbers))).#trainneg
cnf(i_0_112, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_228, plain, (m1_subset_1(esk22_0,k5_numbers))).#trainneg
cnf(i_0_401, plain, (v1_card_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_515, plain, (v1_card_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_311, plain, (v1_finseq_1(esk29_1(X1)))).#trainneg
cnf(i_0_312, plain, (v1_finset_1(esk29_1(X1)))).#trainneg
cnf(i_0_437, plain, (v4_funct_1(k3_finseq_2(X1)))).#trainneg
cnf(i_0_261, plain, (v3_relat_1(esk26_1(X1)))).#trainneg
cnf(i_0_366, plain, (v2_relat_1(esk34_1(X1)))).#trainneg
cnf(i_0_374, plain, (v4_valued_0(esk36_1(X1)))).#trainneg
cnf(i_0_396, plain, (v1_funcop_1(esk37_1(X1)))).#trainneg
cnf(i_0_565, plain, (v3_finseq_1(k3_finseq_2(X1)))).#trainneg
cnf(i_0_560, plain, (v5_ordinal1(esk49_1(X1)))).#trainneg
cnf(i_0_125, plain, (k2_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_208, plain, (k2_xcmplx_0(np__0,np__1)=np__1)).#trainneg
cnf(i_0_281, plain, (k2_xcmplx_0(np__1,np__0)=np__1)).#trainneg
cnf(i_0_117, plain, (r1_xxreal_0(np__0,np__0))).#trainneg
cnf(i_0_144, plain, (r1_xxreal_0(np__0,np__1))).#trainneg
cnf(i_0_199, plain, (~v1_subset_1(esk20_1(X1),X1))).#trainneg
cnf(i_0_209, plain, (r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_328, plain, (v5_relat_1(esk30_0,k5_numbers))).#trainneg
cnf(i_0_535, plain, (v1_funct_2(esk48_0,k5_numbers,k1_numbers))).#trainneg
cnf(i_0_342, plain, (v5_relat_1(esk31_0,k5_numbers))).#trainneg
cnf(i_0_539, plain, (v5_relat_1(esk48_0,k1_numbers))).#trainneg
cnf(i_0_256, plain, (v1_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_192, plain, (v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_289, plain, (v4_relat_1(esk27_0,k5_numbers))).#trainneg
cnf(i_0_329, plain, (v4_relat_1(esk30_0,k5_numbers))).#trainneg
cnf(i_0_343, plain, (v4_relat_1(esk31_0,k5_numbers))).#trainneg
cnf(i_0_457, plain, (v4_relat_1(esk42_0,k5_numbers))).#trainneg
cnf(i_0_540, plain, (v4_relat_1(esk48_0,k5_numbers))).#trainneg
cnf(i_0_326, plain, (v1_partfun1(esk30_0,k5_numbers))).#trainneg
cnf(i_0_339, plain, (v1_partfun1(esk31_0,k5_numbers))).#trainneg
cnf(i_0_536, plain, (v1_partfun1(esk48_0,k5_numbers))).#trainneg
cnf(i_0_263, plain, (k1_card_1(k1_card_1(X1))=k1_card_1(X1))).#trainneg
cnf(i_0_188, plain, (k2_xcmplx_0(np__1,k4_xcmplx_0(np__1))=np__0)).#trainneg
cnf(i_0_516, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),np__1)=np__0)).#trainneg
cnf(i_0_510, plain, (m1_subset_1(esk46_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_171, plain, (~v1_xboole_0(k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_136, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_194, plain, (m1_subset_1(k8_seq_4,k1_zfmisc_1(k2_zfmisc_1(k2_numbers,k1_numbers))))).#trainneg
cnf(i_0_349, plain, (~v1_xxreal_0(X1)|~v1_xboole_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_238, plain, (v3_ordinal1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_551, plain, (~v1_xboole_0(X1)|~v3_card_1(X1,np__1))).#trainneg
cnf(i_0_332, plain, (m1_subset_1(esk30_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_345, plain, (m1_subset_1(esk31_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_305, plain, (v1_finset_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_569, plain, (v1_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_369, plain, (v3_membered(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_316, plain, (v4_relat_1(esk29_1(X1),k5_numbers))).#trainneg
cnf(i_0_118, plain, (m1_subset_1(esk18_1(X1),X1))).#trainneg
cnf(i_0_217, plain, (m1_finseq_2(k3_finseq_2(X1),X1))).#trainneg
cnf(i_0_469, plain, (m1_finseq_2(esk43_1(X1),X1))).#trainneg
cnf(i_0_318, plain, (m1_finseq_1(esk29_1(X1),X1))).#trainneg
cnf(i_0_315, plain, (v5_relat_1(esk29_1(X1),X1))).#trainneg
cnf(i_0_562, plain, (v5_relat_1(esk49_1(X1),X1))).#trainneg
cnf(i_0_260, plain, (v4_relat_1(esk26_1(X1),X1))).#trainneg
cnf(i_0_359, plain, (v4_relat_1(esk33_1(X1),X1))).#trainneg
cnf(i_0_365, plain, (v4_relat_1(esk34_1(X1),X1))).#trainneg
cnf(i_0_377, plain, (v4_relat_1(esk36_1(X1),X1))).#trainneg
cnf(i_0_399, plain, (v4_relat_1(esk37_1(X1),X1))).#trainneg
cnf(i_0_107, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_211, plain, (v1_relat_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_542, plain, (m1_subset_1(esk48_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_139, plain, (~v1_xxreal_0(X1)|~v2_xxreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_207, plain, (v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_301, plain, (v1_xboole_0(X1)|~v3_card_1(X1,k1_xboole_0))).#trainneg
cnf(i_0_184, plain, (v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_568, plain, (v2_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_249, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_438, plain, (v4_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_258, plain, (v1_partfun1(esk26_1(X1),X1))).#trainneg
cnf(i_0_357, plain, (v1_partfun1(esk33_1(X1),X1))).#trainneg
cnf(i_0_375, plain, (v1_partfun1(esk36_1(X1),X1))).#trainneg
cnf(i_0_397, plain, (v1_partfun1(esk37_1(X1),X1))).#trainneg
cnf(i_0_205, plain, (v1_xboole_0(esk21_2(X1,X2)))).#trainneg
cnf(i_0_119, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_201, plain, (v1_relat_1(k3_relat_1(X1,X2)))).#trainneg
cnf(i_0_204, plain, (v1_relat_1(esk21_2(X1,X2)))).#trainneg
cnf(i_0_233, plain, (v1_relat_1(esk23_2(X1,X2)))).#trainneg
cnf(i_0_230, plain, (v1_funct_1(esk23_2(X1,X2)))).#trainneg
cnf(i_0_138, plain, (m1_subset_1(esk19_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_200, plain, (m1_subset_1(esk20_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_210, plain, (v1_funct_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_523, plain, (v1_xboole_0(X1)|v1_finset_1(esk47_1(X1)))).#trainneg
cnf(i_0_8, plain, (v2_seq_4(X2,X1)|~r2_hidden(esk5_2(X1,X2),X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1))))).#trainneg
cnf(i_0_306, plain, (v1_xboole_0(X1)|v1_zfmisc_1(esk28_1(X1)))).#trainneg
cnf(i_0_280, plain, (v3_valued_0(X1)|~m1_finseq_1(X1,k1_numbers))).#trainneg
cnf(i_0_159, plain, (v1_xboole_0(X1)|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_307, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk28_1(X1)))).#trainneg
cnf(i_0_146, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_462, plain, (v4_valued_0(X1)|~m1_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_101, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_202, plain, (v5_relat_1(esk21_2(X1,X2),X2))).#trainneg
cnf(i_0_231, plain, (v5_relat_1(esk23_2(X1,X2),X2))).#trainneg
cnf(i_0_203, plain, (v4_relat_1(esk21_2(X1,X2),X1))).#trainneg
cnf(i_0_232, plain, (v4_relat_1(esk23_2(X1,X2),X1))).#trainneg
cnf(i_0_235, plain, (v1_xxreal_0(k2_xxreal_2(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_116, plain, (v1_xboole_0(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_257, plain, (v1_xreal_0(k3_seq_4(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_105, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_213, plain, (v1_xreal_0(k6_square_1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_177, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_241, plain, (v4_funct_1(X1)|~m1_finseq_2(X1,X2))).#trainneg
cnf(i_0_404, plain, (v1_xboole_0(X1)|v3_card_1(esk38_1(X1),np__1))).#trainneg
cnf(i_0_193, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_350, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk32_1(X1)))).#trainneg
cnf(i_0_524, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk47_1(X1)))).#trainneg
cnf(i_0_390, plain, (v1_finseq_1(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_489, plain, (v1_relat_1(esk44_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_503, plain, (v1_relat_1(esk45_1(X1))|~v1_card_1(X1))).#trainneg
cnf(i_0_487, plain, (v1_funct_1(esk44_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_502, plain, (v1_funct_1(esk45_1(X1))|~v1_card_1(X1))).#trainneg
cnf(i_0_484, plain, (v1_finseq_1(esk44_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_187, plain, (v1_finset_1(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_486, plain, (v1_finset_1(esk44_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_483, plain, (v2_finseq_1(esk44_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_183, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_304, plain, (v1_finset_1(X1)|~v1_finset_1(k1_card_1(X1)))).#trainneg
cnf(i_0_550, plain, (v1_zfmisc_1(X1)|~v3_card_1(X1,np__1))).#trainneg
cnf(i_0_129, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_175, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_482, plain, (v1_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_numbers)))).#trainneg
cnf(i_0_172, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_197, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k2_numbers))).#trainneg
cnf(i_0_488, plain, (v4_relat_1(esk44_1(X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_108, plain, (m1_subset_1(k22_seq_4(X1),k1_numbers)|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_504, plain, (v3_relat_1(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_517, plain, (v4_valued_0(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_491, plain, (v3_xxreal_2(X1)|~v2_membered(X1)|~v5_xxreal_2(X1))).#trainneg
cnf(i_0_382, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_128, plain, (v1_funct_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_174, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_127, plain, (v1_finseq_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_229, plain, (v1_funct_2(esk23_2(X1,X2),X1,X2))).#trainneg
cnf(i_0_485, plain, (v3_card_1(esk44_1(X1),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_501, plain, (v3_card_1(esk45_1(X1),X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_296, plain, (k13_seq_4(X1)=k54_valued_1(X1)|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_479, plain, (k6_square_1(X1)=k7_square_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_370, plain, (k16_rvsum_1(X1)=k18_rvsum_1(X1)|~m1_finseq_1(X1,k1_numbers))).#trainneg
cnf(i_0_83, plain, (m2_finseq_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_84, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_124, plain, (v5_relat_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_99, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_362, plain, (m1_subset_1(k7_square_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_162, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_490, plain, (v4_xxreal_2(X1)|~v2_membered(X1)|~v5_xxreal_2(X1))).#trainneg
cnf(i_0_277, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_363, plain, (m1_subset_1(k18_rvsum_1(X1),k1_numbers)|~m1_finseq_1(X1,k1_numbers))).#trainneg
cnf(i_0_496, plain, (v1_xboole_0(X1)|v2_xxreal_0(X1)|v3_xxreal_0(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_292, plain, (v3_valued_0(k2_zfmisc_1(X2,X1))|~v3_membered(X1))).#trainneg
cnf(i_0_308, plain, (v1_xboole_0(X1)|m1_subset_1(esk28_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_351, plain, (v1_xboole_0(X1)|m1_subset_1(esk32_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_405, plain, (v1_xboole_0(X1)|m1_subset_1(esk38_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_525, plain, (v1_xboole_0(X1)|m1_subset_1(esk47_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_237, plain, (k13_seq_4(k13_seq_4(X1))=k13_seq_4(X1)|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_169, plain, (a_2_0_finseq_2(X1,X2)=k4_finseq_2(X1,X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_319, plain, (m2_finseq_1(k13_seq_4(X1),k1_numbers)|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_95, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_72, plain, (m1_finseq_2(k4_finseq_2(X1,X2),X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_206, plain, (m1_subset_1(esk21_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_168, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_215, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_178, plain, (m1_subset_1(k5_seq_4(X1),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_242, plain, (v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_355, plain, (v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_459, plain, (v1_finset_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_475, plain, (v1_seq_2(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_comseq_2(X1))).#trainneg
cnf(i_0_79, plain, (v1_xreal_0(esk16_1(X1))|~v2_membered(X1)|~v3_xxreal_2(X1))).#trainneg
cnf(i_0_297, plain, (k2_xxreal_2(X1)=k3_seq_4(X1)|v1_xboole_0(X1)|~v3_membered(X1)|~v3_xxreal_2(X1))).#trainneg
cnf(i_0_234, plain, (m1_subset_1(esk23_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_143, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m2_xxreal_2(X2,X1))).#trainneg
cnf(i_0_49, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_283, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_434, plain, (v1_finseq_1(X2)|~v3_finseq_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_436, plain, (r2_hidden(esk40_2(X1,X2),X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_68, plain, (k4_finseq_2(X1,k2_numbers)=k14_seq_4(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_78, plain, (m2_xxreal_2(esk16_1(X1),X1)|~v2_membered(X1)|~v3_xxreal_2(X1))).#trainneg
cnf(i_0_110, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_463, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_179, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_559, plain, (k7_square_1(k18_rvsum_1(k12_rvsum_1(k13_seq_4(X1))))=k22_seq_4(X1)|~m2_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_148, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_433, plain, (k3_finseq_1(X1)=k1_card_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_526, plain, (k4_finseqop(k2_numbers,k1_numbers,X1,k8_seq_4)=k13_seq_4(X1)|~m2_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_214, plain, (v1_xboole_0(X2)|~v7_ordinal1(X1)|~v1_xboole_0(k4_finseq_2(X1,X2)))).#trainneg
cnf(i_0_248, plain, (v1_relat_1(k54_valued_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_474, plain, (v2_seq_2(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_comseq_2(X1))).#trainneg
cnf(i_0_252, plain, (v1_relat_1(k30_valued_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_471, plain, (v1_xboole_0(k3_relat_1(X1,X2))|~v1_xboole_0(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_473, plain, (v1_xboole_0(k3_relat_1(X1,X2))|~v1_xboole_0(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_163, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_410, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_270, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_173, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_180, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_100, plain, (k9_binop_2(X1,X2)=k9_binop_2(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_28, plain, (r1_xxreal_0(k6_numbers,k22_seq_4(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1)))).#trainneg
cnf(i_0_409, plain, (v1_relat_1(k39_valued_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_149, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_418, plain, (v1_relat_1(k30_valued_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_273, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).#trainneg
cnf(i_0_247, plain, (v1_funct_1(k54_valued_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_555, plain, (v1_comseq_2(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_seq_2(X1)|~v2_seq_2(X1))).#trainneg
cnf(i_0_251, plain, (v1_funct_1(k30_valued_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_155, plain, (m1_subset_1(k9_binop_2(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_417, plain, (v1_funct_1(k30_valued_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_408, plain, (v1_funct_1(k39_valued_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_250, plain, (v1_valued_0(k30_valued_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_416, plain, (v1_valued_0(k30_valued_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_552, plain, (v1_valued_0(k39_valued_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_246, plain, (v3_valued_0(k54_valued_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_415, plain, (v3_valued_0(k30_valued_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_282, plain, (k30_valued_1(k30_valued_1(X1))=X1|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_135, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_69, plain, (m2_finseq_1(X3,X2)|~m1_finseq_2(X1,X2)|~m2_finseq_2(X3,X2,X1))).#trainneg
cnf(i_0_176, plain, (m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_368, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_564, plain, (v1_xcmplx_0(k16_rvsum_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_302, plain, (v1_valued_0(X2)|~v1_relat_1(X1)|~v1_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_276, plain, (v3_valued_0(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_65, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_22, plain, (X2=k1_xboole_0|m1_subset_1(esk8_2(X1,X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_373, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_514, plain, (k18_valued_1(X1,X1)=k39_valued_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_394, plain, (v1_xboole_0(X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xboole_0(k2_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_395, plain, (v1_xboole_0(X2)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v1_xboole_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_31, plain, (m2_xxreal_2(X2,X1)|v1_xxreal_0(esk10_2(X1,X2))|~v2_membered(X1)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_447, plain, (v1_finseq_1(k54_valued_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_386, plain, (v1_finseq_1(k30_valued_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_190, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X2)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_481, plain, (v5_relat_1(X3,X2)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_10, plain, (v2_seq_4(X3,X2)|r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X2)))|~r1_xxreal_0(X1,k22_seq_4(esk6_3(X2,X3,X1))))).#trainneg
cnf(i_0_269, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v4_relat_1(X1,X2)|~v1_partfun1(X1,X2))).#trainneg
cnf(i_0_441, plain, (v1_comseq_2(X1)|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k2_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers))))).#trainneg
cnf(i_0_164, plain, (v1_xboole_0(X1)|v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_385, plain, (v4_relat_1(X3,X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_566, plain, (v4_relat_1(k3_relat_1(X1,X2),X3)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X3))).#trainneg
cnf(i_0_236, plain, (v1_xboole_0(X1)|v3_card_1(X3,X2)|~v7_ordinal1(X2)|~m1_subset_1(X3,k4_finseq_2(X2,X1)))).#trainneg
cnf(i_0_216, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_519, plain, (v1_partfun1(X1,X2)|~v1_funct_2(X1,X2,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2))))).#trainneg
cnf(i_0_407, plain, (v3_xxreal_0(X2)|v2_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_381, plain, (v1_relat_1(k24_valued_1(X1,X2))|~v1_xcmplx_0(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_380, plain, (v1_funct_1(k24_valued_1(X1,X2))|~v1_xcmplx_0(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_402, plain, (v1_xboole_0(X1)|v1_partfun1(X2,X3)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_32, plain, (r1_xxreal_0(X1,X3)|~v2_membered(X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X1)|~r2_hidden(X3,X2)|~m2_xxreal_2(X1,X2))).#trainneg
cnf(i_0_300, plain, (v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k2_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers))))).#trainneg
cnf(i_0_545, plain, (v1_funct_1(k30_valued_1(X1))|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k2_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers))))).#trainneg
cnf(i_0_379, plain, (k1_valued_1(X1,X2)=k1_valued_1(X2,X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1))).#trainneg
cnf(i_0_298, plain, (v2_comseq_2(k54_valued_1(X1))|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k2_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers))))).#trainneg
cnf(i_0_543, plain, (v2_comseq_2(k30_valued_1(X1))|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k2_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers))))).#trainneg
cnf(i_0_403, plain, (k18_valued_1(X1,X2)=k18_valued_1(X2,X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1))).#trainneg
cnf(i_0_494, plain, (v1_relat_1(k45_valued_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1))).#trainneg
cnf(i_0_80, plain, (k15_seq_4(X1,X2,X3)=k1_valued_1(X2,X3)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1)))).#trainneg
cnf(i_0_547, plain, (v1_partfun1(k30_valued_1(X1),X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1)|~v4_relat_1(X1,X2)|~v1_partfun1(X1,X2))).#trainneg
cnf(i_0_50, plain, (k15_seq_4(X1,X2,X3)=k15_seq_4(X1,X3,X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1)))).#trainneg
cnf(i_0_507, plain, (v1_relat_1(k1_valued_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1))).#trainneg
cnf(i_0_493, plain, (v1_funct_1(k45_valued_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1))).#trainneg
cnf(i_0_506, plain, (v1_funct_1(k1_valued_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1))).#trainneg
cnf(i_0_511, plain, (v1_valued_0(k45_valued_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1))).#trainneg
cnf(i_0_12, plain, (v2_seq_4(X2,X1)|m2_finseq_2(esk5_2(X1,X2),k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1))))).#trainneg
cnf(i_0_294, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|~v1_xboole_0(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_67, plain, (m2_finseq_2(k15_seq_4(X1,X2,X3),k2_numbers,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1)))).#trainneg
cnf(i_0_299, plain, (v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k2_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers))))).#trainneg
cnf(i_0_464, plain, (k1_valued_1(X1,k30_valued_1(X2))=k45_valued_1(X1,X2)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1))).#trainneg
cnf(i_0_544, plain, (v1_funct_2(k30_valued_1(X1),k5_numbers,k2_numbers)|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k2_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers))))).#trainneg
cnf(i_0_64, plain, (v1_xboole_0(X2)|r1_xxreal_0(X3,X1)|X3!=k3_seq_4(X2)|~v3_membered(X2)|~v3_xxreal_2(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_465, plain, (v1_finseq_1(k45_valued_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1))).#trainneg
cnf(i_0_60, plain, (X2=k3_seq_4(X1)|v1_xboole_0(X1)|v1_xreal_0(esk13_2(X1,X2))|v1_xreal_0(esk14_2(X1,X2))|~v3_membered(X1)|~v3_xxreal_2(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_440, plain, (v1_xboole_0(X2)|v1_funct_1(k30_valued_1(X1))|~v3_membered(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_88, plain, (m1_subset_1(a_3_9_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1)))).#trainneg
cnf(i_0_57, plain, (X2=k3_seq_4(X1)|v1_xboole_0(X1)|v1_xreal_0(esk14_2(X1,X2))|r2_hidden(esk13_2(X1,X2),X1)|~v3_membered(X1)|~v3_xxreal_2(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_59, plain, (X2=k3_seq_4(X1)|v1_xboole_0(X1)|v1_xreal_0(esk13_2(X1,X2))|~v3_membered(X1)|~v3_xxreal_2(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(esk14_2(X1,X2),k6_numbers))).#trainneg
cnf(i_0_54, plain, (X2=k3_seq_4(X1)|v1_xboole_0(X1)|v1_xreal_0(esk14_2(X1,X2))|~v3_membered(X1)|~v3_xxreal_2(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,esk13_2(X1,X2)))).#trainneg
cnf(i_0_439, plain, (v1_xboole_0(X3)|v1_partfun1(k30_valued_1(X1),X2)|~v3_membered(X3)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_56, plain, (X2=k3_seq_4(X1)|v1_xboole_0(X1)|r2_hidden(esk13_2(X1,X2),X1)|~v3_membered(X1)|~v3_xxreal_2(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(esk14_2(X1,X2),k6_numbers))).#trainneg
cnf(i_0_53, plain, (X2=k3_seq_4(X1)|v1_xboole_0(X1)|~v3_membered(X1)|~v3_xxreal_2(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X2,esk13_2(X1,X2))|~r1_xxreal_0(esk14_2(X1,X2),k6_numbers))).#trainneg
cnf(i_0_527, plain, (v1_finseq_1(k1_valued_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1))).#trainneg
cnf(i_0_63, plain, (v1_xboole_0(X1)|r1_xxreal_0(X3,k6_numbers)|v1_xreal_0(esk12_3(X1,X2,X3))|X2!=k3_seq_4(X1)|~v3_membered(X1)|~v3_xxreal_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_160, plain, (k4_finseqop(X2,X4,X1,X3)=k3_relat_1(X1,X3)|~v1_funct_1(X3)|~m1_finseq_1(X1,X2)|~v1_funct_2(X3,X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X4))))).#trainneg
cnf(i_0_352, plain, (m2_finseq_1(k4_finseqop(X2,X4,X1,X3),X4)|~v1_funct_1(X3)|~m1_finseq_1(X1,X2)|~v1_funct_2(X3,X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X4))))).#trainneg
cnf(i_0_11, plain, (v2_seq_4(X2,X1)|r1_xxreal_0(X3,k6_numbers)|m2_finseq_2(esk6_3(X1,X2,X3),k2_numbers,k14_seq_4(X1))|~m1_subset_1(X3,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1))))).#trainneg
cnf(i_0_58, plain, (X3=k3_seq_4(X2)|v1_xboole_0(X2)|v1_xreal_0(esk13_2(X2,X3))|r1_xxreal_0(k9_binop_2(X3,esk14_2(X2,X3)),X1)|~v3_membered(X2)|~v3_xxreal_2(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_55, plain, (X3=k3_seq_4(X2)|v1_xboole_0(X2)|r2_hidden(esk13_2(X2,X3),X2)|r1_xxreal_0(k9_binop_2(X3,esk14_2(X2,X3)),X1)|~v3_membered(X2)|~v3_xxreal_2(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_37, plain, (r2_hidden(esk11_4(X1,X2,X3,X4),X4)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~r2_hidden(X1,a_3_9_seq_4(X2,X3,X4))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2)))).#trainneg
cnf(i_0_39, plain, (m1_subset_1(esk11_4(X1,X2,X3,X4),k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~r2_hidden(X1,a_3_9_seq_4(X2,X3,X4))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2)))).#trainneg
cnf(i_0_38, plain, (k22_seq_4(k19_seq_4(X2,X3,esk11_4(X1,X2,X3,X4)))=X1|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~r2_hidden(X1,a_3_9_seq_4(X2,X3,X4))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2)))).#trainneg
cnf(i_0_9, plain, (v2_seq_4(X2,X1)|r1_xxreal_0(X3,k6_numbers)|r2_hidden(k15_seq_4(X1,esk5_2(X1,X2),esk6_3(X1,X2,X3)),X2)|~m1_subset_1(X3,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1))))).#trainneg
cnf(i_0_52, plain, (X3=k3_seq_4(X2)|v1_xboole_0(X2)|r1_xxreal_0(k9_binop_2(X3,esk14_2(X2,X3)),X1)|~v3_membered(X2)|~v3_xxreal_2(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~r2_hidden(X1,X2)|~r1_xxreal_0(X3,esk13_2(X2,X3)))).#trainneg
cnf(i_0_36, plain, (r2_hidden(X3,a_3_9_seq_4(X2,X4,X5))|X3!=k22_seq_4(k19_seq_4(X2,X4,X1))|~r2_hidden(X1,X5)|~m1_subset_1(X1,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X2)))).#trainneg
cnf(i_0_652, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk19_1(X2)))).#trainneg
cnf(i_0_655, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk20_1(X2)))).#trainneg
cnf(i_0_27, plain, (k1_xboole_0=np__0)).#trainneg
cnf(i_0_580, negated_conjecture, (m1_subset_1(X1,k14_seq_4(esk1_0))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_147, plain, (o_0_0_xboole_0=np__0)).#trainneg
cnf(i_0_578, plain, (~v7_ordinal1(esk39_0))).#trainneg
cnf(i_0_579, plain, (~v7_ordinal1(esk24_0))).#trainneg
cnf(i_0_1062, plain, (v3_membered(k5_numbers))).#trainneg
cnf(i_0_586, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_608, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk18_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_605, plain, (v1_xreal_0(esk35_0))).#trainneg
cnf(i_0_634, plain, (v1_relat_1(esk25_0))).#trainneg
cnf(i_0_28, plain, (r1_xxreal_0(np__0,k22_seq_4(X1))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_587, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk46_0))).#trainneg
cnf(i_0_1099, plain, (v5_relat_1(k8_seq_4,k1_numbers))).#trainneg
cnf(i_0_596, plain, (~v1_xboole_0(esk39_0))).#trainneg
cnf(i_0_588, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_595, plain, (m1_subset_1(X1,k2_zfmisc_1(k2_numbers,k1_numbers))|~r2_hidden(X1,k8_seq_4))).#trainneg
cnf(i_0_1029, plain, (m1_finseq_2(k14_seq_4(X1),k2_numbers)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1116, plain, (~v7_ordinal1(X1)|~v1_xboole_0(k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_636, plain, (v1_relat_1(esk17_0))).#trainneg
cnf(i_0_637, plain, (v1_relat_1(esk15_0))).#trainneg
cnf(i_0_10, plain, (r1_xxreal_0(X1,np__0)|v2_seq_4(X2,X3)|~r1_xxreal_0(X1,k22_seq_4(esk6_3(X3,X2,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X3)))|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_599, plain, (v3_ordinal1(k1_card_1(X1)))).#trainneg
cnf(i_0_1161, plain, (v4_relat_1(k8_seq_4,k2_numbers))).#trainneg
cnf(i_0_597, plain, (~v1_xboole_0(esk24_0))).#trainneg
cnf(i_0_604, plain, (~v7_ordinal1(k5_numbers))).#trainneg
cnf(i_0_647, plain, (v4_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_697, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_602, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k5_numbers))|~r2_hidden(X1,esk30_0))).#trainneg
cnf(i_0_603, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k5_numbers))|~r2_hidden(X1,esk31_0))).#trainneg
cnf(i_0_1624, plain, (r2_hidden(X1,a_3_3_seq_4(X2,X3,X4))|X1!=k22_seq_4(k19_seq_4(X2,X5,X3))|~r2_hidden(X5,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_finseq_2(X5,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_638, plain, (v1_relat_1(np__0))).#trainneg
cnf(i_0_641, plain, (v1_funct_1(esk25_0))).#trainneg
cnf(i_0_643, plain, (v1_funct_1(esk17_0))).#trainneg
cnf(i_0_53, plain, (X1=k3_seq_4(X2)|v1_xboole_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_2(X2)|~v3_membered(X2)|~r1_xxreal_0(esk14_2(X2,X1),np__0)|~r1_xxreal_0(X1,esk13_2(X2,X1)))).#trainneg
cnf(i_0_616, plain, (v1_relat_1(esk46_0)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_644, plain, (v1_funct_1(esk15_0))).#trainneg
cnf(i_0_1078, plain, (~v1_finset_1(k1_numbers))).#trainneg
cnf(i_0_784, plain, (v1_relat_1(k5_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_619, plain, (v1_relat_1(esk22_0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_620, plain, (v1_relat_1(esk18_1(X1))|~v4_funct_1(X1))).#trainneg
cnf(i_0_1592, plain, (k15_seq_4(X1,X2,k45_valued_1(X3,X2))=X3|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1747, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X6)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k19_seq_4(X4,X5,esk9_4(X6,X4,X5,X7))),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(k19_seq_4(X4,X5,esk9_4(X6,X4,X5,X7)),k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r2_hidden(X6,a_3_3_seq_4(X4,X5,X7))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X4)))|~m2_finseq_2(X5,k2_numbers,k14_seq_4(X4))|~m2_subset_1(X4,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_722, negated_conjecture, (v1_xboole_0(esk3_0)|~v1_xboole_0(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_1732, plain, (r2_hidden(k22_seq_4(k19_seq_4(X1,X2,X3)),a_3_9_seq_4(X1,X2,X4))|~r2_hidden(X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1740, plain, (v3_xxreal_0(esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|~v3_xxreal_0(k22_seq_4(X4))|~v1_xreal_0(k22_seq_4(X4))|~v1_xreal_0(esk4_3(X1,X2,X3))|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,X4),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_942, plain, (v7_ordinal1(np__0))).#trainneg
cnf(i_0_613, negated_conjecture, (~v1_xboole_0(k14_seq_4(esk1_0))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_940, plain, (v7_ordinal1(np__1))).#trainneg
cnf(i_0_645, plain, (v1_funct_1(np__0))).#trainneg
cnf(i_0_783, plain, (v1_relat_1(k8_seq_4))).#trainneg
cnf(i_0_301, plain, (v1_xboole_0(X1)|~v3_card_1(X1,np__0))).#trainneg
cnf(i_0_621, negated_conjecture, (v1_relat_1(esk3_0)|~v4_funct_1(k1_zfmisc_1(k14_seq_4(esk1_0))))).#trainneg
cnf(i_0_875, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_795, plain, (v4_funct_1(k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_623, plain, (v1_relat_1(np__1)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_1617, plain, (k22_seq_4(k19_seq_4(X1,X2,X3))=k22_seq_4(k45_valued_1(X3,X2))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_681, plain, (X1=esk15_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_830, plain, (X1=k24_seq_4(X2,X3,X4)|k3_seq_4(esk7_4(X2,X3,X4,X1))!=X1|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X1,k1_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk7_4(X2,X3,X4,X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1101, plain, (k7_square_1(k18_rvsum_1(k12_rvsum_1(k13_seq_4(X1))))=k22_seq_4(k13_seq_4(X1))|~m2_finseq_1(k13_seq_4(X1),k2_numbers)|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_1411, plain, (m2_xxreal_2(X2,X1)|~v1_xboole_0(X1)|~v1_xxreal_0(X2)|~v2_membered(X1))).#trainneg
cnf(i_0_1709, plain, (k22_seq_4(k45_valued_1(X2,esk9_4(X3,X1,X2,X4)))=X3|~r2_hidden(X3,a_3_3_seq_4(X1,X2,X4))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(X3,X1,X2,X4),k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_837, plain, (m2_finseq_2(X1,X2,k4_finseq_2(X3,X2))|~m1_subset_1(X1,k4_finseq_2(X3,X2))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_22, plain, (X1=np__0|m1_subset_1(esk8_2(X2,X1),X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_877, plain, (v3_ordinal1(np__0))).#trainneg
cnf(i_0_967, plain, (v1_xreal_0(np__1))).#trainneg
cnf(i_0_624, plain, (v1_relat_1(np__1)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_1049, plain, (v3_membered(esk46_0))).#trainneg
cnf(i_0_901, plain, (v1_funct_1(k5_numbers)|~v1_funct_1(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1690, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,a_3_3_seq_4(X3,X4,X1))|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(X3)))|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X3))|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1708, plain, (r2_hidden(X1,a_3_3_seq_4(X2,X3,X4))|~r2_hidden(esk9_4(X1,X2,X3,X5),X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_finseq_2(esk9_4(X1,X2,X3,X5),k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r2_hidden(X1,a_3_3_seq_4(X2,X3,X5))|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X2))))).#trainneg
cnf(i_0_1102, plain, (k7_square_1(k18_rvsum_1(k12_rvsum_1(k54_valued_1(X1))))=k22_seq_4(X1)|~m2_finseq_1(X1,k2_numbers)|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_1722, plain, (X1=k24_seq_4(X2,X3,X4)|k5_seq_4(a_3_3_seq_4(X2,X3,X4))!=X1|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X1,k1_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1734, plain, (r2_hidden(X1,a_3_9_seq_4(X2,X3,X4))|X1!=k22_seq_4(k19_seq_4(X2,X5,X3))|~r2_hidden(X5,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_2(X5,k2_numbers,k14_seq_4(X2)))).#trainneg
cnf(i_0_1148, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_1531, plain, (~v1_xboole_0(k8_seq_4))).#trainneg
cnf(i_0_680, plain, (X1=esk17_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1262, plain, (v3_valued_0(k8_seq_4))).#trainneg
cnf(i_0_635, plain, (v1_relat_1(esk19_1(X1)))).#trainneg
cnf(i_0_59, plain, (X1=k3_seq_4(X2)|v1_xreal_0(esk13_2(X2,X1))|v1_xboole_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_2(X2)|~v3_membered(X2)|~r1_xxreal_0(esk14_2(X2,X1),np__0))).#trainneg
cnf(i_0_628, plain, (v1_relat_1(k5_numbers)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_662, plain, (v1_funct_1(esk22_0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_570, plain, (~m1_subset_1(esk39_0,k5_numbers))).#trainneg
cnf(i_0_571, plain, (~m1_subset_1(esk24_0,k5_numbers))).#trainneg
cnf(i_0_1280, plain, (v1_valued_0(k5_numbers)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1296, plain, (v3_valued_0(k5_numbers)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_630, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|~r2_hidden(X1,esk48_0))).#trainneg
cnf(i_0_642, plain, (v1_funct_1(esk19_1(X1)))).#trainneg
cnf(i_0_56, plain, (X1=k3_seq_4(X2)|v1_xboole_0(X2)|r2_hidden(esk13_2(X2,X1),X2)|~v1_xreal_0(X1)|~v3_xxreal_2(X2)|~v3_membered(X2)|~r1_xxreal_0(esk14_2(X2,X1),np__0))).#trainneg
cnf(i_0_1590, plain, (k1_valued_1(X2,k19_seq_4(X1,X3,X2))=X3|~m1_subset_1(k19_seq_4(X1,X3,X2),k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1748, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X6)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k19_seq_4(X4,X5,esk11_4(X6,X4,X5,X7))),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(k19_seq_4(X4,X5,esk11_4(X6,X4,X5,X7)),k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r2_hidden(X6,a_3_9_seq_4(X4,X5,X7))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X4)))|~m2_finseq_2(X5,k2_numbers,k14_seq_4(X4))|~m2_subset_1(X4,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_676, plain, (X1=esk29_1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_679, plain, (X1=esk19_1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_751, negated_conjecture, (v1_xboole_0(esk3_0)|~m1_subset_1(esk2_0,esk3_0))).#trainneg
cnf(i_0_1633, plain, (v1_xxreal_0(esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|~v2_membered(k1_numbers)|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1661, plain, (v3_membered(esk7_4(X1,X2,X3,X4))|X4=k24_seq_4(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X4,k1_numbers)|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1362, plain, (v3_xxreal_0(np__0)|~v3_xxreal_0(np__1)|~v1_xreal_0(np__1)|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_677, plain, (X1=esk25_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_674, plain, (v3_valued_0(esk29_1(k1_numbers)))).#trainneg
cnf(i_0_776, plain, (v1_relat_1(esk46_0)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_11, plain, (r1_xxreal_0(X1,np__0)|v2_seq_4(X2,X3)|m2_finseq_2(esk6_3(X3,X2,X1),k2_numbers,k14_seq_4(X3))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X3)))|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_668, plain, (v1_funct_1(np__1)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_701, plain, (~m1_subset_1(k5_numbers,k5_numbers))).#trainneg
cnf(i_0_1085, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_650, plain, (v1_funct_1(esk21_2(X1,X2)))).#trainneg
cnf(i_0_1509, plain, (k15_seq_4(X1,X3,X2)=k1_valued_1(X2,X3)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1712, plain, (r2_hidden(X1,a_3_3_seq_4(X2,X3,X4))|~r2_hidden(esk11_4(X1,X2,X3,X5),X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_finseq_2(esk11_4(X1,X2,X3,X5),k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r2_hidden(X1,a_3_9_seq_4(X2,X3,X5))|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X2))))).#trainneg
cnf(i_0_700, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk38_1(X1)))).#trainneg
cnf(i_0_1216, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1551, plain, (m2_finseq_1(k19_seq_4(X1,X2,X3),k2_numbers)|~m1_finseq_2(k14_seq_4(X1),k2_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1733, plain, (r2_hidden(X1,a_3_9_seq_4(X2,X3,X4))|X1!=k22_seq_4(k45_valued_1(X3,X5))|~r2_hidden(X5,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X5,k14_seq_4(X2))|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1695, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,a_3_9_seq_4(X3,X4,X1))|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(X3)))|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X3))|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1713, plain, (k22_seq_4(k45_valued_1(X2,esk11_4(X3,X1,X2,X4)))=X3|~r2_hidden(X3,a_3_9_seq_4(X1,X2,X4))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(esk11_4(X3,X1,X2,X4),k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1490, plain, (v1_xxreal_0(k24_seq_4(X1,X2,X3))|~v2_membered(k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_669, plain, (v1_funct_1(np__1)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_698, plain, (v4_funct_1(esk43_1(X1)))).#trainneg
cnf(i_0_9, plain, (r1_xxreal_0(X1,np__0)|v2_seq_4(X2,X3)|r2_hidden(k15_seq_4(X3,esk5_2(X3,X2),esk6_3(X3,X2,X1)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X3)))|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_654, plain, (v1_relat_1(esk20_1(X1))|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_632, plain, (~v3_xxreal_0(np__1)|~v1_xxreal_0(np__1))).#trainneg
cnf(i_0_715, plain, (~v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_673, plain, (v1_funct_1(k5_numbers)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_659, plain, (v1_funct_1(esk46_0)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_663, plain, (v1_funct_1(esk20_1(X1))|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1550, plain, (r1_xxreal_0(X1,X2)|v1_xboole_0(X3)|X1!=k3_seq_4(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v3_xxreal_2(X3)|~v3_membered(X3)|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_2013, plain, (esk17_0=esk15_0)).#trainneg
cnf(i_0_1841, plain, (esk15_0=np__0)).#trainneg
cnf(i_0_1932, plain, (v3_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_1831, negated_conjecture, (v1_xreal_0(esk1_0))).#trainneg
cnf(i_0_1920, negated_conjecture, (m1_subset_1(esk2_0,k14_seq_4(esk1_0))|~v7_ordinal1(esk1_0))).#trainneg
cnf(i_0_1837, plain, (esk25_0=np__0)).#trainneg
cnf(i_0_1887, plain, (r1_xxreal_0(esk39_0,esk39_0))).#trainneg
cnf(i_0_1895, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_665, plain, (v1_funct_1(esk18_1(X1))|~v4_funct_1(X1))).#trainneg
cnf(i_0_666, negated_conjecture, (v1_funct_1(esk3_0)|~v4_funct_1(k1_zfmisc_1(k14_seq_4(esk1_0))))).#trainneg
cnf(i_0_1828, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1964, plain, (k5_seq_4(a_3_3_seq_4(X1,X2,X3))=k24_seq_4(X1,X2,X3)|~m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(k24_seq_4(X1,X2,X3),k1_numbers)|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1854, plain, (r2_hidden(k22_seq_4(k19_seq_4(X1,X3,X2)),a_3_3_seq_4(X1,X2,X4))|~r2_hidden(X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1894, plain, (r1_xxreal_0(X1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2103, plain, (v1_valued_0(k8_seq_4))).#trainneg
cnf(i_0_1829, negated_conjecture, (v3_ordinal1(esk1_0))).#trainneg
cnf(i_0_63, plain, (v1_xreal_0(esk12_3(X1,X2,X3))|v1_xboole_0(X1)|r1_xxreal_0(X3,np__0)|X2!=k3_seq_4(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v3_xxreal_2(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_1888, plain, (r1_xxreal_0(esk24_0,esk24_0))).#trainneg
cnf(i_0_1959, plain, (~v7_ordinal1(k1_numbers))).#trainneg
cnf(i_0_2194, plain, (~v3_xxreal_0(np__1))).#trainneg
cnf(i_0_2253, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1766, negated_conjecture, (m1_subset_1(X1,k14_seq_4(esk1_0))|v1_xboole_0(esk3_0)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_678, plain, (X1=esk21_2(X2,X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_686, plain, (v1_relat_1(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1830, negated_conjecture, (v7_ordinal1(esk1_0))).#trainneg
cnf(i_0_1976, plain, (r2_hidden(k22_seq_4(k19_seq_4(X1,X3,X2)),a_3_9_seq_4(X1,X2,X4))|~r2_hidden(X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1)))).#trainneg
cnf(i_0_1839, plain, (esk19_1(X1)=np__0)).#trainneg
cnf(i_0_2193, plain, (X1=k24_seq_4(X2,X3,X4)|v3_membered(a_3_3_seq_4(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X1,k1_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1912, plain, (m2_finseq_2(X1,k2_numbers,k14_seq_4(X2))|~m1_subset_1(X1,k14_seq_4(X2))|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1889, plain, (r1_xxreal_0(esk22_0,esk22_0))).#trainneg
cnf(i_0_687, plain, (v1_funct_1(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_725, plain, (v7_ordinal1(esk46_0)|~v1_finset_1(esk46_0))).#trainneg
cnf(i_0_2218, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_688, plain, (X1=k1_card_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2018, plain, (X1=k24_seq_4(X2,X3,X4)|k3_seq_4(a_3_3_seq_4(X2,X3,X4))!=X1|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X1,k1_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1847, negated_conjecture, (v3_xxreal_0(k24_seq_4(esk1_0,esk2_0,esk3_0))|v2_xxreal_0(np__0)|v1_xboole_0(k24_seq_4(esk1_0,esk2_0,esk3_0))|~v1_xreal_0(np__0)|~v1_xreal_0(k24_seq_4(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_2230, plain, (k1_valued_1(k19_seq_4(X1,X3,X2),X2)=X3|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(k19_seq_4(X1,X3,X2),k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2093, plain, (m1_subset_1(esk29_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_730, plain, (v1_valued_0(esk49_1(k2_numbers)))).#trainneg
cnf(i_0_1033, plain, (v2_membered(esk46_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1960, plain, (~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_693, plain, (v1_xcmplx_0(np__1)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_765, plain, (v1_xxreal_0(np__1)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_1988, plain, (~v1_xboole_0(k14_seq_4(np__0)))).#trainneg
cnf(i_0_2229, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_695, plain, (v1_xcmplx_0(np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_703, plain, (v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2000, plain, (k22_seq_4(k19_seq_4(X1,X3,X2))=k22_seq_4(k45_valued_1(X3,X2))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2261, plain, (r2_hidden(k22_seq_4(k45_valued_1(X1,X2)),a_3_9_seq_4(X3,X1,X4))|~r2_hidden(X2,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X3)))|~m1_subset_1(X1,k14_seq_4(X3))|~m1_subset_1(X3,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1836, plain, (esk29_1(X1)=np__0)).#trainneg
cnf(i_0_1858, plain, (v3_xxreal_0(np__0)|X3=np__0|~v3_xxreal_0(k24_seq_4(X1,X2,X3))|~v1_xreal_0(k24_seq_4(X1,X2,X3))|~v1_xreal_0(np__0)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_785, plain, (v3_valued_0(esk49_1(k1_numbers)))).#trainneg
cnf(i_0_2179, negated_conjecture, (~m1_subset_1(esk2_0,esk3_0))).#trainneg
cnf(i_0_708, plain, (v1_finset_1(X1)|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_1991, plain, (~v1_xboole_0(k14_seq_4(np__1)))).#trainneg
cnf(i_0_717, plain, (~v1_xboole_0(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_771, plain, (v1_xxreal_0(k5_numbers)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_782, negated_conjecture, (v1_relat_1(esk3_0)|~v1_relat_1(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_709, plain, (v1_finset_1(X1)|~v7_ordinal1(k1_card_1(X1)))).#trainneg
cnf(i_0_710, plain, (v1_finset_1(X1)|~m1_subset_1(k1_card_1(X1),k5_numbers))).#trainneg
cnf(i_0_2282, plain, (k24_seq_4(X1,X2,X3)=np__0|m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1809, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2082, plain, (k5_seq_4(a_3_3_seq_4(X1,X2,X3))=k24_seq_4(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(k5_seq_4(a_3_3_seq_4(X1,X2,X3)),k1_numbers)|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2186, plain, (v3_xxreal_0(X1)|~v2_xxreal_0(np__0)|~v1_xreal_0(X1)|~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2192, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|~v2_membered(k1_numbers)|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1962, plain, (k1_valued_1(X2,k45_valued_1(X3,X2))=X3|~m1_subset_1(k45_valued_1(X3,X2),k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2002, plain, (r2_hidden(k22_seq_4(k45_valued_1(X3,X2)),a_3_3_seq_4(X1,X2,X4))|~r2_hidden(X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1194, plain, (v1_relat_1(k30_valued_1(esk48_0)))).#trainneg
cnf(i_0_1915, plain, (~v1_xboole_0(k14_seq_4(esk22_0))|~m2_subset_1(esk22_0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_713, plain, (m2_finseq_2(X1,X2,esk43_1(X2))|~m1_subset_1(X1,esk43_1(X2)))).#trainneg
cnf(i_0_947, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_973, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_714, plain, (m2_finseq_2(X1,X2,k3_finseq_2(X2))|~m1_subset_1(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_719, plain, (v1_xboole_0(esk20_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2281, plain, (k9_binop_2(X1,X2)=k24_seq_4(X3,X4,X5)|m1_subset_1(a_3_3_seq_4(X3,X4,X5),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X3)))|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X3))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2272, plain, (k22_seq_4(k45_valued_1(X1,esk11_4(X2,X3,X1,X4)))=X2|~r2_hidden(X2,a_3_9_seq_4(X3,X1,X4))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X3)))|~m1_subset_1(X1,k14_seq_4(X3))|~m1_subset_1(X3,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m2_finseq_2(X1,k2_numbers,k14_seq_4(X3)))).#trainneg
cnf(i_0_1917, plain, (~v1_xboole_0(k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1861, plain, (v2_xxreal_0(k24_seq_4(X1,X2,X3))|X3=np__0|~v2_xxreal_0(np__0)|~v1_xreal_0(k24_seq_4(X1,X2,X3))|~v1_xreal_0(np__0)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1875, plain, (r1_xxreal_0(np__0,k22_seq_4(k19_seq_4(X1,X2,X3)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1195, plain, (v1_relat_1(k30_valued_1(esk31_0)))).#trainneg
cnf(i_0_1914, plain, (~v1_xboole_0(k14_seq_4(esk35_0))|~m2_subset_1(esk35_0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_721, plain, (v1_xboole_0(esk18_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_741, plain, (~v1_xboole_0(esk44_1(np__1))|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_743, plain, (~v1_xboole_0(esk45_1(np__1))|~v1_card_1(np__1))).#trainneg
cnf(i_0_999, plain, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1025, plain, (v1_finseq_1(k5_numbers)|~v3_finseq_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_723, plain, (v1_xboole_0(k8_seq_4)|~v1_xboole_0(k2_zfmisc_1(k2_numbers,k1_numbers)))).#trainneg
cnf(i_0_732, plain, (v1_valued_0(esk23_2(X1,k2_numbers)))).#trainneg
cnf(i_0_2023, plain, (v3_xxreal_2(X1)|~v1_xreal_0(X2)|~v2_membered(X1)|~v1_xboole_0(X1)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_2074, plain, (m2_xxreal_2(X4,a_3_3_seq_4(X2,X3,X1))|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_xxreal_0(X4)|~v2_membered(a_3_3_seq_4(X2,X3,X1)))).#trainneg
cnf(i_0_2239, plain, (r2_hidden(X1,a_3_3_seq_4(X2,X3,X4))|~r2_hidden(esk11_4(X1,X2,X3,X5),X4)|~r2_hidden(X1,a_3_9_seq_4(X2,X3,X5))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk11_4(X1,X2,X3,X5),k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1951, plain, (r2_hidden(X1,a_3_3_seq_4(X2,X3,X4))|X1!=k22_seq_4(k45_valued_1(X5,X3))|~r2_hidden(X5,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_finseq_2(X5,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X5,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2237, plain, (k1_valued_1(k45_valued_1(X3,X2),X2)=X3|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k45_valued_1(X3,X2),k14_seq_4(X1)))).#trainneg
cnf(i_0_1997, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X6,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k19_seq_4(X4,X5,X6)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(k19_seq_4(X4,X5,X6),k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X5,k14_seq_4(X4))|~m1_subset_1(X6,k14_seq_4(X4))|~m1_subset_1(X4,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2235, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k1_valued_1(k19_seq_4(X5,X6,esk9_4(X4,X5,X6,X7)),X3),X2)|~r2_hidden(X4,a_3_3_seq_4(X5,X6,X7))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m2_finseq_2(k19_seq_4(X5,X6,esk9_4(X4,X5,X6,X7)),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X6,k2_numbers,k14_seq_4(X5))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(k19_seq_4(X5,X6,esk9_4(X4,X5,X6,X7)),k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2217, plain, (r1_xxreal_0(np__0,X1)|r1_xxreal_0(X1,np__0)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_2260, plain, (m2_finseq_1(k45_valued_1(X2,X3),k2_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1874, negated_conjecture, (r1_xxreal_0(np__0,k22_seq_4(esk2_0)))).#trainneg
cnf(i_0_992, plain, (v1_xcmplx_0(np__1)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_1196, plain, (v1_relat_1(k30_valued_1(esk30_0)))).#trainneg
cnf(i_0_733, plain, (v1_valued_0(esk21_2(X1,k2_numbers)))).#trainneg
cnf(i_0_993, plain, (v1_xcmplx_0(np__1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1342, plain, (~v7_ordinal1(np__1)|~v7_ordinal1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1044, negated_conjecture, (v2_membered(esk3_0)|~v2_membered(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_1060, negated_conjecture, (v3_membered(esk3_0)|~v3_membered(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_734, plain, (v1_relat_1(k22_seq_4(X1))|~v4_funct_1(k1_numbers)|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_735, plain, (v1_funct_1(k22_seq_4(X1))|~v4_funct_1(k1_numbers)|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_2236, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k1_valued_1(k19_seq_4(X5,X6,esk11_4(X4,X5,X6,X7)),X3),X2)|~r2_hidden(X4,a_3_9_seq_4(X5,X6,X7))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m2_finseq_2(k19_seq_4(X5,X6,esk11_4(X4,X5,X6,X7)),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X6,k2_numbers,k14_seq_4(X5))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(k19_seq_4(X5,X6,esk11_4(X4,X5,X6,X7)),k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2007, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X6,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k19_seq_4(X4,X6,X5)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(k19_seq_4(X4,X6,X5),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X5,k2_numbers,k14_seq_4(X4))|~m2_finseq_2(X6,k2_numbers,k14_seq_4(X4))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X5,k14_seq_4(X4))|~m1_subset_1(X6,k14_seq_4(X4))|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_2259, plain, (m1_subset_1(k19_seq_4(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2450, negated_conjecture, (~v1_xboole_0(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_2370, plain, (k3_seq_4(a_3_3_seq_4(X1,X2,X3))=k24_seq_4(X1,X2,X3)|~m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k24_seq_4(X1,X2,X3),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_652, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,np__0))).#trainneg
cnf(i_0_2013, plain, (esk17_0=np__0)).#trainneg
cnf(i_0_745, plain, (v1_xreal_0(k7_square_1(X1))|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2384, plain, (r1_xxreal_0(esk35_0,esk35_0))).#trainneg
cnf(i_0_2721, plain, (~v1_xboole_0(k14_seq_4(esk22_0)))).#trainneg
cnf(i_0_2389, plain, (r1_xxreal_0(X1,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_749, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_750, plain, (v2_seq_4(X1,X2)|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(esk5_2(X2,X1),X1))).#trainneg
cnf(i_0_2685, plain, (k5_seq_4(a_3_3_seq_4(X1,X2,X3))=k24_seq_4(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2519, plain, (k3_seq_4(a_3_3_seq_4(X1,X2,X3))=k24_seq_4(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(k3_seq_4(a_3_3_seq_4(X1,X2,X3)),k1_numbers)|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2859, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X3)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(X4,X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X4,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_finseq_2(k19_seq_4(X1,X4,X3),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X1)))).#trainneg
cnf(i_0_2390, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2580, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_312, plain, (v1_finset_1(np__0))).#trainneg
cnf(i_0_755, plain, (v1_xxreal_0(esk48_0)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2849, plain, (~v1_xboole_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2776, plain, (v1_xcmplx_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2364, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X1)|v1_xboole_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_756, plain, (v1_xxreal_0(esk46_0)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_757, plain, (v1_xxreal_0(esk31_0)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_2806, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X3,X4)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(X4,X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X4,k14_seq_4(X1))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_finseq_2(k19_seq_4(X1,X4,X3),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X1)))).#trainneg
cnf(i_0_2411, plain, (m1_subset_1(k3_seq_4(X1),k1_numbers)|~v3_membered(X1))).#trainneg
cnf(i_0_2692, plain, (v3_xxreal_0(X1)|~v2_xxreal_0(np__0)|~v1_xboole_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2518, plain, (k1_card_1(X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_311, plain, (v1_finseq_1(np__0))).#trainneg
cnf(i_0_674, plain, (v3_valued_0(np__0))).#trainneg
cnf(i_0_138, plain, (m1_subset_1(np__0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_758, plain, (v1_xxreal_0(esk30_0)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_2399, plain, (m1_subset_1(esk39_0,k1_numbers))).#trainneg
cnf(i_0_2850, plain, (~m1_subset_1(k4_xcmplx_0(np__1),k5_numbers))).#trainneg
cnf(i_0_2416, plain, (m1_subset_1(X1,k1_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_760, plain, (v1_xxreal_0(esk20_1(X1))|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_762, plain, (v1_xxreal_0(esk18_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_2409, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2684, plain, (k5_seq_4(a_3_3_seq_4(X1,X2,X3))=k24_seq_4(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k5_seq_4(a_3_3_seq_4(X1,X2,X3)),k5_numbers))).#trainneg
cnf(i_0_2778, plain, (v3_xxreal_2(X1)|~v1_xboole_0(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_2756, plain, (X3=np__0|~v3_xxreal_0(k24_seq_4(X1,X2,X3))|~v1_xxreal_0(k24_seq_4(X1,X2,X3))|~v2_xxreal_0(np__0)|~v1_xreal_0(k24_seq_4(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2400, plain, (m1_subset_1(esk35_0,k1_numbers))).#trainneg
cnf(i_0_2401, plain, (m1_subset_1(esk22_0,k1_numbers))).#trainneg
cnf(i_0_2486, plain, (v3_xxreal_0(esk22_0)|v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_316, plain, (v4_relat_1(np__0,k5_numbers))).#trainneg
cnf(i_0_763, negated_conjecture, (v1_xxreal_0(esk3_0)|~v2_membered(k1_zfmisc_1(k14_seq_4(esk1_0))))).#trainneg
cnf(i_0_2824, plain, (r1_xxreal_0(esk39_0,np__0)|r1_xxreal_0(np__0,esk39_0))).#trainneg
cnf(i_0_910, plain, (~v1_finseq_1(k5_numbers)|~v1_funct_1(k5_numbers)|~v1_relat_1(k5_numbers))).#trainneg
cnf(i_0_2777, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k2_numbers)))).#trainneg
cnf(i_0_2852, negated_conjecture, (v3_membered(esk3_0)|~m1_subset_1(k14_seq_4(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_764, plain, (v1_xxreal_0(k8_seq_4)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k2_numbers,k1_numbers))))).#trainneg
cnf(i_0_769, plain, (v1_xxreal_0(k22_seq_4(X1))|~v2_membered(k1_numbers)|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_2538, negated_conjecture, (v2_xxreal_0(np__0)|v1_xboole_0(k24_seq_4(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k24_seq_4(esk1_0,esk2_0,esk3_0),k5_numbers)|~v1_xreal_0(k24_seq_4(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_2755, plain, (X3=np__0|~v1_xboole_0(k24_seq_4(X1,X2,X3))|~v1_xxreal_0(k24_seq_4(X1,X2,X3))|~v2_xxreal_0(np__0)|~v1_xreal_0(k24_seq_4(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2536, negated_conjecture, (v2_xxreal_0(np__0)|v1_xboole_0(k24_seq_4(esk1_0,esk2_0,esk3_0))|~v7_ordinal1(k24_seq_4(esk1_0,esk2_0,esk3_0))|~v1_xreal_0(k24_seq_4(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_2796, plain, (v3_xxreal_2(a_3_3_seq_4(X1,X2,X3))|~v1_xreal_0(X4)|~v2_membered(a_3_3_seq_4(X1,X2,X3))|~v1_xboole_0(X3)|~v1_xxreal_0(X4)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1218, plain, (v1_funct_1(k30_valued_1(esk48_0)))).#trainneg
cnf(i_0_315, plain, (v5_relat_1(np__0,X1))).#trainneg
cnf(i_0_772, plain, (v1_xxreal_0(k7_square_1(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2825, plain, (r1_xxreal_0(esk24_0,np__0)|r1_xxreal_0(np__0,esk24_0))).#trainneg
cnf(i_0_1961, plain, (~v1_finseq_1(k1_numbers)|~v1_funct_1(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2712, plain, (~v2_xxreal_0(k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2502, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2503, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_773, plain, (v1_relat_1(k7_square_1(X1))|~v4_funct_1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_774, plain, (v1_funct_1(k7_square_1(X1))|~v4_funct_1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2761, plain, (v2_xxreal_0(k22_seq_4(k19_seq_4(X1,X2,X3)))|~v2_xxreal_0(np__0)|~v1_xreal_0(k22_seq_4(k19_seq_4(X1,X2,X3)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2701, plain, (v1_xboole_0(X1)|k3_seq_4(X1)!=np__0|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(np__1,esk12_3(X1,np__0,np__1)))).#trainneg
cnf(i_0_2907, plain, (~v1_funct_2(k19_seq_4(X1,X2,X3),k5_numbers,k2_numbers)|~v1_xboole_0(k19_seq_4(X1,X2,X3))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2407, plain, (m1_subset_1(k6_square_1(X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2808, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,k19_seq_4(X6,X5,X4),X3),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X4,k14_seq_4(X6))|~m1_subset_1(X6,k5_numbers)|~m2_finseq_2(k19_seq_4(X6,X5,X4),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~m1_subset_1(k19_seq_4(X6,X5,X4),k14_seq_4(X1))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2828, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),np__0)|r1_xxreal_0(np__0,esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|~v2_membered(k1_numbers)|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2809, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k1_valued_1(k19_seq_4(X6,X5,X4),X3),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X4,k14_seq_4(X6))|~m1_subset_1(X6,k5_numbers)|~m2_finseq_2(k19_seq_4(X6,X5,X4),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(k19_seq_4(X6,X5,X4),k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2919, plain, (m1_subset_1(k45_valued_1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1009, plain, (v1_finseq_1(esk22_0)|~v3_finseq_1(k5_numbers))).#trainneg
cnf(i_0_1219, plain, (v1_funct_1(k30_valued_1(esk31_0)))).#trainneg
cnf(i_0_318, plain, (m1_finseq_1(np__0,X1))).#trainneg
cnf(i_0_779, plain, (v1_relat_1(esk20_1(X1))|~v1_relat_1(X1))).#trainneg
cnf(i_0_2826, plain, (r1_xxreal_0(esk22_0,np__0)|r1_xxreal_0(np__0,esk22_0))).#trainneg
cnf(i_0_1018, plain, (v1_finseq_1(np__1)|~v3_finseq_1(k5_numbers))).#trainneg
cnf(i_0_781, plain, (v1_relat_1(esk18_1(k1_zfmisc_1(X1)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_788, plain, (v3_valued_0(esk23_2(X1,k1_numbers)))).#trainneg
cnf(i_0_2714, plain, (v3_xxreal_0(k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2863, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,k19_seq_4(X6,X4,X5),X3),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X4,k14_seq_4(X6))|~m1_subset_1(X6,k5_numbers)|~m2_finseq_2(k19_seq_4(X6,X4,X5),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(k19_seq_4(X6,X4,X5),k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2899, plain, (v1_relat_1(k19_seq_4(X1,X2,X3))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1019, plain, (v1_finseq_1(np__1)|~v3_finseq_1(k1_numbers))).#trainneg
cnf(i_0_1220, plain, (v1_funct_1(k30_valued_1(esk30_0)))).#trainneg
cnf(i_0_2455, plain, (~v1_xboole_0(a_3_9_seq_4(X1,X2,X3))|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_789, plain, (v3_valued_0(esk21_2(X1,k1_numbers)))).#trainneg
cnf(i_0_1076, negated_conjecture, (v1_finset_1(esk3_0)|~v1_finset_1(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_1771, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_791, plain, (v1_xxreal_0(k18_rvsum_1(X1))|~v2_membered(k1_numbers)|~m1_finseq_1(X1,k1_numbers))).#trainneg
cnf(i_0_2484, plain, (r1_xxreal_0(np__0,k22_seq_4(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k14_seq_4(X2)))).#trainneg
cnf(i_0_2811, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|v1_xboole_0(X2)|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X4,k14_seq_4(X6))|~m1_subset_1(X6,k5_numbers)|~m2_finseq_2(k19_seq_4(X6,X5,X4),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~m1_subset_1(k15_seq_4(X1,X3,k19_seq_4(X6,X5,X4)),X2))).#trainneg
cnf(i_0_2864, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|v1_xboole_0(X2)|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X4,k14_seq_4(X6))|~m1_subset_1(X6,k5_numbers)|~m2_finseq_2(k19_seq_4(X6,X4,X5),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~m1_subset_1(k15_seq_4(X1,X3,k19_seq_4(X6,X4,X5)),X2))).#trainneg
cnf(i_0_2862, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k1_valued_1(k19_seq_4(X6,X4,X5),X3),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X4,k14_seq_4(X6))|~m1_subset_1(X6,k5_numbers)|~m2_finseq_2(k19_seq_4(X6,X4,X5),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(k19_seq_4(X6,X4,X5),k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1827, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X6)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k19_seq_4(X4,esk11_4(X6,X4,X5,X7),X5)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(k19_seq_4(X4,esk11_4(X6,X4,X5,X7),X5),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X5,k2_numbers,k14_seq_4(X4))|~m2_finseq_2(esk11_4(X6,X4,X5,X7),k2_numbers,k14_seq_4(X4))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~r2_hidden(X6,a_3_9_seq_4(X4,X5,X7))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X4))))).#trainneg
cnf(i_0_2619, plain, (k22_seq_4(k45_valued_1(X2,X3))=k22_seq_4(k45_valued_1(X3,X2))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2226, plain, (r1_xxreal_0(X1,esk4_3(X2,X3,X4))|r1_xxreal_0(esk4_3(X2,X3,X4),X1)|r2_hidden(X4,X3)|~v1_xreal_0(X1)|~v2_membered(k1_numbers)|~v2_seq_4(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2805, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(X5,X4)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X4,k14_seq_4(X6))|~m1_subset_1(X6,k5_numbers)|~m2_finseq_2(k45_valued_1(X5,X4),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1993, plain, (v1_xboole_0(esk44_1(np__0)))).#trainneg
cnf(i_0_656, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk20_1(X1)))).#trainneg
cnf(i_0_792, plain, (v1_relat_1(k18_rvsum_1(X1))|~v4_funct_1(k1_numbers)|~m1_finseq_1(X1,k1_numbers))).#trainneg
cnf(i_0_889, plain, (v1_funct_1(esk46_0)|~v1_funct_1(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1832, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_1833, negated_conjecture, (v1_relat_1(esk1_0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_1772, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1773, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_793, plain, (v1_funct_1(k18_rvsum_1(X1))|~v4_funct_1(k1_numbers)|~m1_finseq_1(X1,k1_numbers))).#trainneg
cnf(i_0_798, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_1816, plain, (m1_subset_1(esk8_2(X1,X2),X2)|X2=np__0|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2346, negated_conjecture, (v3_xxreal_0(esk1_0)|v2_xxreal_0(X1)|r1_xxreal_0(X1,esk1_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2280, plain, (k2_xcmplx_0(X1,X2)=k24_seq_4(X3,X4,X5)|m1_subset_1(a_3_3_seq_4(X3,X4,X5),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X3)))|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X3))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2078, plain, (v1_xboole_0(a_3_3_seq_4(X3,X4,X1))|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(X3)))|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X3))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,a_3_3_seq_4(X3,X4,X1)))).#trainneg
cnf(i_0_2915, plain, (v1_valued_0(k19_seq_4(X1,X2,X3))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1974, plain, (~v1_xboole_0(a_3_9_seq_4(X1,X2,X3))|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1965, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(X6,esk9_4(X4,X5,X6,X7))),X2)|~r2_hidden(X4,a_3_3_seq_4(X5,X6,X7))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m2_finseq_2(k45_valued_1(X6,esk9_4(X4,X5,X6,X7)),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X6,k2_numbers,k14_seq_4(X5))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(X4,X5,X6,X7),k14_seq_4(X5))|~m1_subset_1(X6,k14_seq_4(X5))|~m1_subset_1(X5,k5_numbers))).#trainneg
cnf(i_0_1781, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2104, plain, (v1_relat_1(k30_valued_1(k8_seq_4)))).#trainneg
cnf(i_0_1482, plain, (v1_partfun1(k8_seq_4,k2_numbers))).#trainneg
cnf(i_0_612, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk18_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_799, plain, (v1_xxreal_0(esk28_1(X1))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1834, negated_conjecture, (v1_funct_1(esk1_0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_1835, negated_conjecture, (v1_finseq_1(esk1_0)|~v3_finseq_1(k5_numbers))).#trainneg
cnf(i_0_1774, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1775, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_800, plain, (v1_relat_1(esk28_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_801, plain, (v1_funct_1(esk28_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1923, plain, (r1_xxreal_0(np__0,k22_seq_4(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2076, plain, (a_3_3_seq_4(X3,X4,X1)=np__0|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(X3)))|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X3))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(a_3_3_seq_4(X3,X4,X1),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_2283, plain, (k5_seq_4(X1)=k24_seq_4(X2,X3,X4)|m1_subset_1(a_3_3_seq_4(X2,X3,X4),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2815, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,k19_seq_4(X5,esk9_4(X4,X5,X6,X7),X6),X3),X2)|~r2_hidden(X4,a_3_3_seq_4(X5,X6,X7))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m2_finseq_2(k19_seq_4(X5,esk9_4(X4,X5,X6,X7),X6),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X6,k2_numbers,k14_seq_4(X5))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(k19_seq_4(X5,esk9_4(X4,X5,X6,X7),X6),k14_seq_4(X1))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2991, plain, (k3_seq_4(a_3_3_seq_4(X1,X2,X3))=k24_seq_4(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v3_membered(a_3_3_seq_4(X1,X2,X3)))).#trainneg
cnf(i_0_2952, plain, (k3_seq_4(a_3_3_seq_4(X1,X2,X3))=k24_seq_4(X1,X2,X3)|~m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3096, plain, (v3_xxreal_2(X1)|~v1_xboole_0(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_3012, plain, (v1_valued_0(np__0))).#trainneg
cnf(i_0_3046, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,np__0))).#trainneg
cnf(i_0_802, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk28_1(X2)))).#trainneg
cnf(i_0_3052, plain, (~v6_membered(k1_numbers))).#trainneg
cnf(i_0_804, plain, (v1_relat_1(esk28_1(X1))|v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_806, plain, (v1_xxreal_0(esk32_1(X1))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2953, plain, (m1_subset_1(k24_seq_4(X1,X2,X3),k1_numbers)|~m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3496, plain, (a_3_3_seq_4(X1,X2,X3)=np__0|~v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3377, plain, (X1=esk44_1(np__0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3376, plain, (esk44_1(np__0)=np__0)).#trainneg
cnf(i_0_2955, plain, (k3_seq_4(a_3_3_seq_4(X1,X2,X3))=k24_seq_4(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k3_seq_4(a_3_3_seq_4(X1,X2,X3)),k5_numbers))).#trainneg
cnf(i_0_3147, negated_conjecture, (k24_seq_4(esk1_0,esk2_0,esk3_0)=np__0|v2_xxreal_0(np__0)|~m1_subset_1(k24_seq_4(esk1_0,esk2_0,esk3_0),k5_numbers))).#trainneg
cnf(i_0_3366, negated_conjecture, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(X5,X4)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(esk1_0))|~m1_subset_1(X4,k14_seq_4(esk1_0))|~m2_finseq_2(k45_valued_1(X5,X4),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3091, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2105, plain, (v1_funct_1(k30_valued_1(k8_seq_4)))).#trainneg
cnf(i_0_3388, negated_conjecture, (v1_xcmplx_0(esk1_0))).#trainneg
cnf(i_0_807, plain, (v1_relat_1(esk32_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3055, plain, (v1_relat_1(esk39_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_808, plain, (v1_funct_1(esk32_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_809, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk32_1(X2)))).#trainneg
cnf(i_0_3447, plain, (v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|~v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3462, plain, (k2_zfmisc_1(X1,X2)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3336, plain, (v1_xboole_0(X1)|r1_xxreal_0(esk4_3(X2,X1,X3),k22_seq_4(k45_valued_1(X4,X3)))|r2_hidden(X3,X1)|~v2_seq_4(X1,X2)|~m1_subset_1(X4,X1)|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X4,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers)|~m2_finseq_2(k19_seq_4(X2,X4,X3),k2_numbers,k14_seq_4(X2))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X2)))).#trainneg
cnf(i_0_3292, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(esk9_4(X4,X1,X3,X5),X2)|~r2_hidden(X4,a_3_3_seq_4(X1,X3,X5))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(k45_valued_1(esk9_4(X4,X1,X3,X5),X3),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(esk9_4(X4,X1,X3,X5),k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2997, plain, (k1_card_1(np__0)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3152, plain, (v3_xxreal_2(a_3_3_seq_4(X1,X2,X3))|~v1_xreal_0(X4)|~v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v3_membered(a_3_3_seq_4(X1,X2,X3)))).#trainneg
cnf(i_0_2996, plain, (v1_card_1(np__0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2637, plain, (v1_valued_0(esk49_1(k1_numbers)))).#trainneg
cnf(i_0_2939, plain, (v3_xxreal_0(esk35_0)|v2_xxreal_0(esk35_0))).#trainneg
cnf(i_0_811, plain, (v1_relat_1(esk32_1(X1))|v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3056, plain, (v1_funct_1(esk39_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_813, plain, (v1_xxreal_0(esk38_1(X1))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_814, plain, (v1_relat_1(esk38_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2934, plain, (m1_subset_1(k7_square_1(X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3293, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(esk9_4(X4,X1,X3,X5),X2)|~r2_hidden(X4,a_3_3_seq_4(X1,X3,X5))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k19_seq_4(X1,esk9_4(X4,X1,X3,X5),X3),k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3360, negated_conjecture, (k22_seq_4(k45_valued_1(X1,X2))=k22_seq_4(k45_valued_1(X2,X1))|~m1_subset_1(X1,k14_seq_4(esk1_0))|~m1_subset_1(X2,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_3463, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3239, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k19_seq_4(X6,X5,X4)),X2)|~m1_subset_1(k19_seq_4(X6,X5,X4),k14_seq_4(X1))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X4,k14_seq_4(X6))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X6,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3495, plain, (r1_xxreal_0(np__0,k22_seq_4(X1))|~m1_subset_1(X1,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3256, negated_conjecture, (m1_subset_1(k45_valued_1(X1,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_3323, plain, (v1_xboole_0(X1)|r1_xxreal_0(esk4_3(X2,X1,X3),k22_seq_4(k45_valued_1(X3,X4)))|r2_hidden(X3,X1)|~v2_seq_4(X1,X2)|~m1_subset_1(X4,X1)|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X4,k14_seq_4(X2))|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers)|~m2_finseq_2(k19_seq_4(X2,X4,X3),k2_numbers,k14_seq_4(X2))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X2)))).#trainneg
cnf(i_0_3013, plain, (v1_relat_1(k30_valued_1(np__0)))).#trainneg
cnf(i_0_2324, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|X1!=k3_seq_4(k5_numbers)|~v1_xreal_0(X1)|~v3_xxreal_2(k5_numbers))).#trainneg
cnf(i_0_815, plain, (v1_funct_1(esk38_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3057, plain, (v1_finseq_1(esk39_0)|~v3_finseq_1(k1_numbers))).#trainneg
cnf(i_0_3391, plain, (v3_ordinal1(X1)|~v7_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3489, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_816, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk38_1(X2)))).#trainneg
cnf(i_0_818, plain, (v1_relat_1(esk38_1(X1))|v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2968, plain, (v3_xxreal_0(X1)|v3_xxreal_0(k4_xcmplx_0(X1))|v1_xboole_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_3510, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k19_seq_4(X5,esk9_4(X4,X5,X6,X7),X6)),X2)|~r2_hidden(X4,a_3_3_seq_4(X5,X6,X7))|~m1_subset_1(k19_seq_4(X5,esk9_4(X4,X5,X6,X7),X6),k14_seq_4(X1))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_finseq_2(X6,k2_numbers,k14_seq_4(X5))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3294, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(esk9_4(X4,X1,X3,X5),X2)|~r2_hidden(X4,a_3_3_seq_4(X1,X3,X5))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(k19_seq_4(X1,esk9_4(X4,X1,X3,X5),X3),k14_seq_4(X1)))).#trainneg
cnf(i_0_3286, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k19_seq_4(X6,X4,X5)),X2)|~m1_subset_1(k19_seq_4(X6,X4,X5),k14_seq_4(X1))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X4,k14_seq_4(X6))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X6,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3148, negated_conjecture, (X1=k24_seq_4(esk1_0,esk2_0,esk3_0)|v2_xxreal_0(np__0)|~v1_xboole_0(X1)|~m1_subset_1(k24_seq_4(esk1_0,esk2_0,esk3_0),k5_numbers))).#trainneg
cnf(i_0_3426, plain, (v1_xreal_0(esk8_2(X1,X2))|X2=np__0|~v3_membered(X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3505, plain, (k5_seq_4(X1)=k24_seq_4(X2,X3,X4)|m1_subset_1(a_3_3_seq_4(X2,X3,X4),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3038, plain, (v2_membered(np__0)|~v2_membered(X1))).#trainneg
cnf(i_0_3014, plain, (v1_funct_1(k30_valued_1(np__0)))).#trainneg
cnf(i_0_3051, plain, (v3_ordinal1(esk39_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_2325, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|X1!=k3_seq_4(k1_numbers)|~v1_xreal_0(X1)|~v3_xxreal_2(k1_numbers))).#trainneg
cnf(i_0_820, plain, (v1_xxreal_0(esk47_1(X1))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3098, plain, (v1_xxreal_0(esk35_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1909, negated_conjecture, (v1_relat_1(esk1_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3309, negated_conjecture, (v1_finset_1(esk3_0)|~m1_subset_1(k14_seq_4(esk1_0),k5_numbers))).#trainneg
cnf(i_0_3490, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k2_numbers)))).#trainneg
cnf(i_0_821, plain, (v1_relat_1(esk47_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_822, plain, (v1_funct_1(esk47_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3452, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,k45_valued_1(X5,esk9_4(X4,X6,X5,X7)),X3),X2)|~r2_hidden(X4,a_3_3_seq_4(X6,X5,X7))|~m1_subset_1(esk9_4(X4,X6,X5,X7),k14_seq_4(X6))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X6)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X6,k5_numbers)|~m2_finseq_2(k45_valued_1(X5,esk9_4(X4,X6,X5,X7)),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~m1_subset_1(k45_valued_1(X5,esk9_4(X4,X6,X5,X7)),k14_seq_4(X1))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3456, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|v1_xboole_0(X2)|~v2_seq_4(X2,X1)|~r2_hidden(X4,a_3_3_seq_4(X6,X5,X7))|~m1_subset_1(esk9_4(X4,X6,X5,X7),k14_seq_4(X6))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X6)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X6,k5_numbers)|~m2_finseq_2(k45_valued_1(X5,esk9_4(X4,X6,X5,X7)),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~m1_subset_1(k15_seq_4(X1,X3,k45_valued_1(X5,esk9_4(X4,X6,X5,X7))),X2))).#trainneg
cnf(i_0_3514, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|v1_xboole_0(X2)|~v2_seq_4(X2,X1)|~r2_hidden(X4,a_3_3_seq_4(X5,X6,X7))|~m1_subset_1(k19_seq_4(X5,esk9_4(X4,X5,X6,X7),X6),k14_seq_4(X1))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_finseq_2(X6,k2_numbers,k14_seq_4(X5))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(k15_seq_4(X1,k19_seq_4(X5,esk9_4(X4,X5,X6,X7),X6),X3),X2))).#trainneg
cnf(i_0_2989, plain, (k3_seq_4(X1)=k24_seq_4(X2,X3,X4)|v3_membered(a_3_3_seq_4(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v3_membered(X1))).#trainneg
cnf(i_0_3445, plain, (k2_xcmplx_0(X1,X2)=k24_seq_4(X3,X4,X5)|m1_subset_1(a_3_3_seq_4(X3,X4,X5),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X3)))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X4,k14_seq_4(X3))|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_3453, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k1_valued_1(k45_valued_1(X5,esk9_4(X4,X6,X5,X7)),X3),X2)|~r2_hidden(X4,a_3_3_seq_4(X6,X5,X7))|~m1_subset_1(esk9_4(X4,X6,X5,X7),k14_seq_4(X6))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X6)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X6,k5_numbers)|~m2_finseq_2(k45_valued_1(X5,esk9_4(X4,X6,X5,X7)),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(k45_valued_1(X5,esk9_4(X4,X6,X5,X7)),k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2956, plain, (v3_xxreal_0(esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|~v3_xxreal_0(k22_seq_4(k45_valued_1(X4,X3)))|~v1_xreal_0(k22_seq_4(k45_valued_1(X4,X3)))|~v1_xreal_0(esk4_3(X1,X2,X3))|~v2_seq_4(X2,X1)|~r2_hidden(X4,X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2990, plain, (k3_seq_4(X1)=k24_seq_4(X2,X3,X4)|m1_subset_1(a_3_3_seq_4(X2,X3,X4),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v3_membered(X1))).#trainneg
cnf(i_0_3040, plain, (v3_membered(np__0)|~v3_membered(X1))).#trainneg
cnf(i_0_3027, plain, (v4_relat_1(np__0,X1))).#trainneg
cnf(i_0_1916, plain, (~v1_xboole_0(k14_seq_4(k2_xcmplx_0(X1,X2)))|~m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_823, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk47_1(X2)))).#trainneg
cnf(i_0_3102, plain, (v1_xcmplx_0(esk35_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1910, negated_conjecture, (v1_funct_1(esk1_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_1911, negated_conjecture, (v1_finseq_1(esk1_0)|~v3_finseq_1(k1_numbers))).#trainneg
cnf(i_0_3177, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_825, plain, (v1_relat_1(esk47_1(X1))|v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_827, plain, (k54_valued_1(k13_seq_4(X1))=k13_seq_4(X1)|~m1_finseq_1(k13_seq_4(X1),k2_numbers)|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_2240, plain, (v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2719, plain, (v3_xxreal_0(X1)|~v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xboole_0(k2_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2058, plain, (v1_xreal_0(esk8_2(X1,X2))|X2=np__0|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3554, plain, (m1_subset_1(k24_seq_4(X1,X2,X3),k1_numbers)|k24_seq_4(X1,X2,X3)=np__0|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3592, plain, (v2_finseq_1(np__0))).#trainneg
cnf(i_0_3718, plain, (v1_card_1(np__0))).#trainneg
cnf(i_0_3594, plain, (v3_card_1(np__0,np__0))).#trainneg
cnf(i_0_3846, plain, (v3_ordinal1(X1)|~v7_ordinal1(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_829, plain, (k13_seq_4(k54_valued_1(X1))=k54_valued_1(X1)|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_831, plain, (v1_funct_1(k13_seq_4(X1))|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_3909, plain, (v2_membered(np__0))).#trainneg
cnf(i_0_3918, plain, (v1_xxreal_0(esk35_0))).#trainneg
cnf(i_0_3575, plain, (X1=k5_seq_4(X2)|X2!=np__0|X1!=k24_seq_4(X3,X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X3)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X3))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_xboole_0(X5))).#trainneg
cnf(i_0_832, plain, (v1_finseq_1(k13_seq_4(X1))|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_3991, plain, (v3_membered(np__0))).#trainneg
cnf(i_0_3847, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~v7_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_833, plain, (m1_finseq_1(k13_seq_4(X1),k1_numbers)|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_3518, plain, (v1_xreal_0(k24_seq_4(X1,X2,X3))|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3921, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(X5,esk9_4(X4,X6,X5,X7))),X2)|~r2_hidden(X4,a_3_3_seq_4(X6,X5,X7))|~m1_subset_1(k45_valued_1(X5,esk9_4(X4,X6,X5,X7)),k14_seq_4(X1))|~m1_subset_1(esk9_4(X4,X6,X5,X7),k14_seq_4(X6))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X6)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X6,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4018, plain, (v1_xcmplx_0(esk35_0))).#trainneg
cnf(i_0_3688, plain, (k1_card_1(np__0)=np__0)).#trainneg
cnf(i_0_835, plain, (m2_finseq_1(k54_valued_1(X1),k1_numbers)|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_3103, plain, (v1_relat_1(esk35_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_1776, plain, (v1_relat_1(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_838, plain, (v4_funct_1(k4_finseq_2(X1,X2))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_839, plain, (v1_xxreal_0(esk21_2(X1,X2))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3865, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(esk9_4(X4,X5,X6,X7),X6)),X2)|~r2_hidden(X4,a_3_3_seq_4(X5,X6,X7))|~m1_subset_1(k45_valued_1(esk9_4(X4,X5,X6,X7),X6),k14_seq_4(X1))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_finseq_2(X6,k2_numbers,k14_seq_4(X5))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(X6,k14_seq_4(X5))|~m1_subset_1(esk9_4(X4,X5,X6,X7),k14_seq_4(X5))|~m1_subset_1(X5,k5_numbers))).#trainneg
cnf(i_0_3515, plain, (v1_xreal_0(k24_seq_4(X1,X2,X3))|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3516, plain, (m1_subset_1(k24_seq_4(X1,X2,X3),k1_numbers)|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3099, plain, (v3_ordinal1(esk35_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_3822, negated_conjecture, (~v1_xboole_0(k2_zfmisc_1(k5_numbers,k2_numbers))|~r2_hidden(X1,k45_valued_1(X2,esk2_0))|~m1_subset_1(X2,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_3104, plain, (v1_funct_1(esk35_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3911, plain, (~v1_finset_1(esk39_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_1777, plain, (v1_funct_1(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1778, plain, (v1_finseq_1(X1)|~v3_finseq_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_842, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,esk21_2(X2,X3)))).#trainneg
cnf(i_0_847, plain, (v1_relat_1(esk47_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3861, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X1)|v1_xboole_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3938, plain, (v1_xboole_0(X1)|r1_xxreal_0(esk4_3(X2,X1,X3),X4)|r2_hidden(X3,X1)|~v2_seq_4(X1,X2)|~r2_hidden(X4,a_3_3_seq_4(X5,X6,X7))|~m1_subset_1(k15_seq_4(X2,k45_valued_1(esk9_4(X4,X5,X6,X7),X6),X3),X1)|~m1_subset_1(k45_valued_1(esk9_4(X4,X5,X6,X7),X6),k14_seq_4(X2))|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers)|~m2_finseq_2(X6,k2_numbers,k14_seq_4(X5))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(X6,k14_seq_4(X5))|~m1_subset_1(esk9_4(X4,X5,X6,X7),k14_seq_4(X5))|~m1_subset_1(X5,k5_numbers))).#trainneg
cnf(i_0_3606, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~v1_xboole_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2638, plain, (v1_relat_1(k30_valued_1(esk49_1(k1_numbers))))).#trainneg
cnf(i_0_848, plain, (v1_relat_1(esk38_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3105, plain, (v1_finseq_1(esk35_0)|~v3_finseq_1(k1_numbers))).#trainneg
cnf(i_0_2646, plain, (v1_xxreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2851, negated_conjecture, (v2_membered(esk3_0)|~v3_membered(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_849, plain, (v1_relat_1(esk32_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_852, plain, (v1_relat_1(esk28_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_3869, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k1_valued_1(k19_seq_4(X5,esk9_4(X4,X5,X6,X7),X6),X3),X2)|~r2_hidden(X4,a_3_3_seq_4(X5,X6,X7))|~m1_subset_1(k19_seq_4(X5,esk9_4(X4,X5,X6,X7),X6),k14_seq_4(X1))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_finseq_2(X6,k2_numbers,k14_seq_4(X5))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3925, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|v1_xboole_0(X2)|~v2_seq_4(X2,X1)|~r2_hidden(X4,a_3_3_seq_4(X6,X5,X7))|~m1_subset_1(k45_valued_1(X5,esk9_4(X4,X6,X5,X7)),k14_seq_4(X1))|~m1_subset_1(esk9_4(X4,X6,X5,X7),k14_seq_4(X6))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X6)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X6,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~m1_subset_1(k15_seq_4(X1,k45_valued_1(X5,esk9_4(X4,X6,X5,X7)),X3),X2))).#trainneg
cnf(i_0_3605, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_3604, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers))).#trainneg
cnf(i_0_2639, plain, (v1_funct_1(k30_valued_1(esk49_1(k1_numbers))))).#trainneg
cnf(i_0_854, plain, (v1_relat_1(esk20_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3114, plain, (v1_relat_1(esk22_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3115, plain, (v1_funct_1(esk22_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_3308, negated_conjecture, (v1_finset_1(esk3_0)|~v7_ordinal1(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_3311, plain, (v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_856, plain, (v1_relat_1(esk18_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_858, plain, (v3_ordinal1(esk48_0)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_3941, plain, (v1_xboole_0(X1)|r1_xxreal_0(esk4_3(X2,X1,X3),X4)|r2_hidden(X3,X1)|~v2_seq_4(X1,X2)|~r2_hidden(X4,a_3_3_seq_4(X5,X6,X7))|~m1_subset_1(k15_seq_4(X2,X3,k19_seq_4(X5,esk9_4(X4,X5,X6,X7),X6)),X1)|~m1_subset_1(k19_seq_4(X5,esk9_4(X4,X5,X6,X7),X6),k14_seq_4(X2))|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers)|~m2_finseq_2(X6,k2_numbers,k14_seq_4(X5))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3791, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(X5,X4)),X2)|~m1_subset_1(k45_valued_1(X5,X4),k14_seq_4(X1))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X4,k14_seq_4(X6))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X6,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3752, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|v1_xboole_0(X2)|~v2_seq_4(X2,X1)|~r2_hidden(X4,a_3_3_seq_4(X1,X3,X5))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(X4,X1,X3,X5),X2))).#trainneg
cnf(i_0_3858, plain, (v3_xxreal_0(X1)|v1_xboole_0(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3116, plain, (v1_finseq_1(esk22_0)|~v3_finseq_1(k1_numbers))).#trainneg
cnf(i_0_3268, plain, (v1_valued_0(esk23_2(X1,k1_numbers)))).#trainneg
cnf(i_0_3621, negated_conjecture, (v3_xxreal_0(esk1_0)|r1_xxreal_0(np__0,esk1_0))).#trainneg
cnf(i_0_2339, plain, (r1_xxreal_0(X1,k5_numbers)|X1!=k3_seq_4(k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1)|~v3_xxreal_2(k1_zfmisc_1(k1_numbers))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_859, plain, (v3_ordinal1(esk47_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3127, plain, (v2_xxreal_0(np__0)|r1_xxreal_0(np__0,esk39_0)|~v2_xxreal_0(esk39_0))).#trainneg
cnf(i_0_950, plain, (v1_xreal_0(esk46_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_976, plain, (v1_xcmplx_0(esk46_0)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1913, plain, (v4_funct_1(k14_seq_4(X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_861, plain, (v3_ordinal1(esk38_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_862, plain, (v3_ordinal1(esk32_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3867, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k1_valued_1(X3,k19_seq_4(X5,esk9_4(X4,X5,X6,X7),X6)),X2)|~r2_hidden(X4,a_3_3_seq_4(X5,X6,X7))|~m1_subset_1(k19_seq_4(X5,esk9_4(X4,X5,X6,X7),X6),k14_seq_4(X1))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_finseq_2(X6,k2_numbers,k14_seq_4(X5))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3874, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(X4,X5)),X2)|~m1_subset_1(k45_valued_1(X4,X5),k14_seq_4(X1))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X4,k14_seq_4(X6))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X6,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3643, negated_conjecture, (v3_xxreal_0(esk1_0)|r1_xxreal_0(X1,esk1_0)|~v1_xboole_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3876, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k1_valued_1(X3,k19_seq_4(X6,X4,X5)),X2)|~m1_subset_1(k19_seq_4(X6,X4,X5),k14_seq_4(X1))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X4,k14_seq_4(X6))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X6,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3563, plain, (m1_subset_1(k24_seq_4(X1,X2,X3),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_3699, plain, (v3_xxreal_2(a_3_3_seq_4(X1,X2,X3))|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3922, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k1_valued_1(X3,k45_valued_1(X5,esk9_4(X4,X6,X5,X7))),X2)|~r2_hidden(X4,a_3_3_seq_4(X6,X5,X7))|~m1_subset_1(k45_valued_1(X5,esk9_4(X4,X6,X5,X7)),k14_seq_4(X1))|~m1_subset_1(esk9_4(X4,X6,X5,X7),k14_seq_4(X6))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X6)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X6,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3793, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k1_valued_1(X3,k19_seq_4(X6,X5,X4)),X2)|~m1_subset_1(k19_seq_4(X6,X5,X4),k14_seq_4(X1))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X4,k14_seq_4(X6))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X6,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3797, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|v1_xboole_0(X2)|~v2_seq_4(X2,X1)|~m1_subset_1(k19_seq_4(X6,X5,X4),k14_seq_4(X1))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X4,k14_seq_4(X6))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X6,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~m1_subset_1(k15_seq_4(X1,X3,k19_seq_4(X6,X5,X4)),X2))).#trainneg
cnf(i_0_3833, negated_conjecture, (v1_valued_0(k45_valued_1(X1,esk2_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_1002, plain, (v1_finseq_1(esk46_0)|~v3_finseq_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3304, plain, (v1_valued_0(esk21_2(X1,k1_numbers)))).#trainneg
cnf(i_0_2323, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|X1!=k3_seq_4(k1_zfmisc_1(k14_seq_4(esk1_0)))|~v1_xreal_0(X1)|~v3_xxreal_2(k1_zfmisc_1(k14_seq_4(esk1_0)))|~v3_membered(k1_zfmisc_1(k14_seq_4(esk1_0))))).#trainneg
cnf(i_0_863, plain, (v3_ordinal1(esk31_0)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_3267, plain, (v2_xxreal_0(np__0)|r1_xxreal_0(np__0,esk22_0)|~v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_2219, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_864, plain, (v3_ordinal1(esk30_0)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_865, plain, (v3_ordinal1(esk28_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3928, plain, (v1_xboole_0(X1)|r1_xxreal_0(esk4_3(X2,X1,X3),X4)|r2_hidden(X3,X1)|~v2_seq_4(X1,X2)|~r2_hidden(X4,a_3_3_seq_4(X5,X6,X7))|~m1_subset_1(k1_valued_1(k45_valued_1(X6,esk9_4(X4,X5,X6,X7)),X3),X1)|~m1_subset_1(esk9_4(X4,X5,X6,X7),k14_seq_4(X5))|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m1_subset_1(X6,k14_seq_4(X5))|~m1_subset_1(X5,k5_numbers)|~m2_finseq_2(k45_valued_1(X6,esk9_4(X4,X5,X6,X7)),k2_numbers,k14_seq_4(X2))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(k45_valued_1(X6,esk9_4(X4,X5,X6,X7)),k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3596, negated_conjecture, (v3_xxreal_0(esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|~v3_xxreal_0(k22_seq_4(k45_valued_1(X4,X5)))|~v1_xreal_0(k22_seq_4(k45_valued_1(X4,X5)))|~v1_xreal_0(esk4_3(X1,X2,X3))|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(X5,X4)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(esk1_0))|~m1_subset_1(X4,k14_seq_4(esk1_0))|~m2_finseq_2(k45_valued_1(X5,X4),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3860, plain, (v3_xxreal_0(X1)|v1_xboole_0(k4_xcmplx_0(X1))|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3863, plain, (v3_xxreal_0(X1)|v3_xxreal_0(k4_xcmplx_0(X1))|v1_xboole_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3553, negated_conjecture, (m1_subset_1(k24_seq_4(X1,X2,X3),k1_numbers)|k24_seq_4(X1,X2,X3)=esk1_0|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3683, plain, (v3_xxreal_0(esk4_3(X1,X2,X3))|v2_xxreal_0(k22_seq_4(k45_valued_1(X4,X3)))|v1_xboole_0(esk4_3(X1,X2,X3))|v1_xboole_0(X2)|r2_hidden(X3,X2)|~v1_xreal_0(k22_seq_4(k45_valued_1(X4,X3)))|~v1_xreal_0(esk4_3(X1,X2,X3))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X4,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X4,X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3534, plain, (m2_xxreal_2(X2,np__0)|~v1_xboole_0(X1)|~v1_xxreal_0(X2)|~v2_membered(np__0))).#trainneg
cnf(i_0_3598, negated_conjecture, (v3_xxreal_0(esk4_3(X1,X2,X3))|v2_xxreal_0(k22_seq_4(k45_valued_1(X4,X5)))|v1_xboole_0(esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|~v1_xreal_0(k22_seq_4(k45_valued_1(X4,X5)))|~v1_xreal_0(esk4_3(X1,X2,X3))|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(X5,X4)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(esk1_0))|~m1_subset_1(X4,k14_seq_4(esk1_0))|~m2_finseq_2(k45_valued_1(X5,X4),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3746, plain, (k7_square_1(X1)=k24_seq_4(X2,X3,X4)|m1_subset_1(a_3_3_seq_4(X2,X3,X4),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1383, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk37_1(X1)))).#trainneg
cnf(i_0_3061, plain, (m1_subset_1(esk21_2(X1,X2),k1_numbers))).#trainneg
cnf(i_0_2314, plain, (r1_xxreal_0(X1,esk22_0)|X1!=k3_seq_4(k5_numbers)|~v1_xreal_0(X1)|~v3_xxreal_2(k5_numbers))).#trainneg
cnf(i_0_867, plain, (v3_ordinal1(esk21_2(X1,X2))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2349, negated_conjecture, (v1_xxreal_0(esk2_0)|~v2_membered(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_2350, negated_conjecture, (v3_ordinal1(esk2_0)|~v3_ordinal1(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_868, plain, (v3_ordinal1(esk20_1(X1))|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_870, plain, (v3_ordinal1(esk18_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_3840, plain, (v3_xxreal_0(esk4_3(X1,X2,X3))|v2_xxreal_0(k22_seq_4(k45_valued_1(X3,X4)))|v1_xboole_0(esk4_3(X1,X2,X3))|v1_xboole_0(X2)|r2_hidden(X3,X2)|~v1_xreal_0(k22_seq_4(k45_valued_1(X3,X4)))|~v1_xreal_0(esk4_3(X1,X2,X3))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X4,k14_seq_4(X1))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X4,X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3569, plain, (v3_xxreal_2(np__0)|~v1_xreal_0(X4)|~v1_xboole_0(X3)|~v2_membered(np__0)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3971, negated_conjecture, (k2_xcmplx_0(X1,X2)=k24_seq_4(esk1_0,X3,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X3,esk3_0),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X3,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_3536, plain, (r1_xxreal_0(X3,np__0)|v2_seq_4(np__0,X2)|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3630, plain, (X1=a_3_3_seq_4(X2,X3,X4)|~v1_xboole_0(X1)|~v1_xboole_0(X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3987, plain, (k3_seq_4(X1)=k24_seq_4(X2,X3,X4)|m1_subset_1(a_3_3_seq_4(X2,X3,X4),k1_zfmisc_1(k1_numbers))|~v3_membered(X1)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4150, plain, (r2_hidden(X1,X2)|~v2_seq_4(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X3)))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(X3))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r2_hidden(np__0,a_3_3_seq_4(X3,X1,X2))|~m1_subset_1(X1,k14_seq_4(X3))|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_4106, plain, (v3_xxreal_2(np__0))).#trainneg
cnf(i_0_4515, plain, (m2_xxreal_2(X1,np__0)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_4545, plain, (v1_xreal_0(esk21_2(X1,X2)))).#trainneg
cnf(i_0_871, negated_conjecture, (v3_ordinal1(esk3_0)|~v3_ordinal1(k1_zfmisc_1(k14_seq_4(esk1_0))))).#trainneg
cnf(i_0_4109, plain, (r1_xxreal_0(esk35_0,np__0)|r1_xxreal_0(np__0,esk35_0))).#trainneg
cnf(i_0_872, plain, (v3_ordinal1(k18_rvsum_1(X1))|~v3_ordinal1(k1_numbers)|~m1_finseq_1(X1,k1_numbers))).#trainneg
cnf(i_0_873, plain, (v3_ordinal1(k7_square_1(X1))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4184, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(esk9_4(X4,X5,X3,X6),X2)|~r2_hidden(X4,a_3_3_seq_4(X5,X3,X6))|~m1_subset_1(k45_valued_1(esk9_4(X4,X5,X3,X6),X3),k14_seq_4(X1))|~m1_subset_1(esk9_4(X4,X5,X3,X6),k14_seq_4(X5))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X6,k1_zfmisc_1(k14_seq_4(X5)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X3,k14_seq_4(X5))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X5,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(X4,X5,X3,X6),k14_seq_4(X1)))).#trainneg
cnf(i_0_4328, negated_conjecture, (r1_xxreal_0(np__0,esk1_0))).#trainneg
cnf(i_0_4182, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k1_valued_1(k45_valued_1(esk9_4(X4,X5,X6,X7),X6),X3),X2)|~r2_hidden(X4,a_3_3_seq_4(X5,X6,X7))|~m1_subset_1(k45_valued_1(esk9_4(X4,X5,X6,X7),X6),k14_seq_4(X1))|~m1_subset_1(esk9_4(X4,X5,X6,X7),k14_seq_4(X5))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X6,k14_seq_4(X5))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X5,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4181, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,k45_valued_1(esk9_4(X4,X5,X6,X7),X6),X3),X2)|~r2_hidden(X4,a_3_3_seq_4(X5,X6,X7))|~m1_subset_1(k45_valued_1(esk9_4(X4,X5,X6,X7),X6),k14_seq_4(X1))|~m1_subset_1(esk9_4(X4,X5,X6,X7),k14_seq_4(X5))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X6,k14_seq_4(X5))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X5,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4180, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k1_valued_1(X3,k45_valued_1(esk9_4(X4,X5,X6,X7),X6)),X2)|~r2_hidden(X4,a_3_3_seq_4(X5,X6,X7))|~m1_subset_1(k45_valued_1(esk9_4(X4,X5,X6,X7),X6),k14_seq_4(X1))|~m1_subset_1(esk9_4(X4,X5,X6,X7),k14_seq_4(X5))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X6,k14_seq_4(X5))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X5,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1384, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk36_1(X1)))).#trainneg
cnf(i_0_1027, plain, (m1_subset_1(esk40_2(X1,X2),X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_4426, plain, (r1_xxreal_0(np__0,esk22_0)|v3_xxreal_0(esk22_0))).#trainneg
cnf(i_0_874, plain, (v3_ordinal1(k8_seq_4)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k2_numbers,k1_numbers))))).#trainneg
cnf(i_0_1267, plain, (v1_valued_0(esk46_0)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_4204, plain, (~v1_finset_1(esk39_0)|~v7_ordinal1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4203, plain, (~v1_finset_1(esk39_0)|~v7_ordinal1(k1_zfmisc_1(k1_zfmisc_1(k1_numbers))))).#trainneg
cnf(i_0_4247, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4248, negated_conjecture, (v2_membered(esk3_0)|~m1_subset_1(k14_seq_4(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_879, plain, (v3_ordinal1(k22_seq_4(X1))|~v3_ordinal1(k1_numbers)|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_882, plain, (v1_xxreal_0(k5_seq_4(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4148, plain, (v3_xxreal_0(esk4_3(X1,X2,X3))|v2_xxreal_0(X4)|v1_xboole_0(esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|~v1_xreal_0(X4)|~v1_xreal_0(esk4_3(X1,X2,X3))|~v2_seq_4(X2,X1)|~r2_hidden(X4,a_3_3_seq_4(X1,X3,X2))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4492, plain, (v3_xxreal_0(X1)|v1_xboole_0(X1)|~v1_xboole_0(k4_xcmplx_0(X1))|~v1_xxreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_4222, plain, (k4_xcmplx_0(X1)=np__0|v3_xxreal_0(X1)|v2_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4555, negated_conjecture, (v1_xxreal_0(esk2_0)|~v3_membered(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_4311, negated_conjecture, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(X5,X4)),X2)|~m1_subset_1(k45_valued_1(X5,X4),k14_seq_4(X1))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(esk1_0))|~m1_subset_1(X4,k14_seq_4(esk1_0))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1385, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk33_1(X1)))).#trainneg
cnf(i_0_964, plain, (v1_xreal_0(k7_square_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2329, plain, (r1_xxreal_0(X1,np__1)|X1!=k3_seq_4(k5_numbers)|~v1_xreal_0(X1)|~v3_xxreal_2(k5_numbers))).#trainneg
cnf(i_0_883, plain, (v3_ordinal1(k5_seq_4(X1))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1283, plain, (v3_valued_0(esk46_0)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2351, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_2352, negated_conjecture, (v1_xreal_0(esk2_0)|~v3_membered(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_4291, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_884, plain, (v1_relat_1(k5_seq_4(X1))|~v4_funct_1(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_885, plain, (v1_funct_1(k5_seq_4(X1))|~v4_funct_1(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4084, plain, (v1_xreal_0(k24_seq_4(X1,X2,X3))|k24_seq_4(X1,X2,X3)=np__0|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4491, plain, (v3_xxreal_0(X1)|v1_xboole_0(X1)|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_4147, plain, (v3_xxreal_0(esk4_3(X1,X2,X3))|v2_xxreal_0(X4)|v1_xboole_0(X4)|r2_hidden(X3,X2)|~v1_xreal_0(X4)|~v1_xreal_0(esk4_3(X1,X2,X3))|~v2_seq_4(X2,X1)|~r2_hidden(X4,a_3_3_seq_4(X1,X3,X2))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4501, negated_conjecture, (v1_xreal_0(k24_seq_4(X1,X2,X3))|k24_seq_4(X1,X2,X3)=esk1_0|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4146, plain, (v3_xxreal_0(esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|~v3_xxreal_0(X4)|~v1_xreal_0(X4)|~v1_xreal_0(esk4_3(X1,X2,X3))|~v2_seq_4(X2,X1)|~r2_hidden(X4,a_3_3_seq_4(X1,X3,X2))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4324, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|v1_xboole_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4489, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|v1_xboole_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k5_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1388, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk26_1(X1)))).#trainneg
cnf(i_0_970, plain, (v1_xreal_0(k5_seq_4(X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2330, plain, (r1_xxreal_0(X1,np__1)|X1!=k3_seq_4(k1_numbers)|~v1_xreal_0(X1)|~v3_xxreal_2(k1_numbers))).#trainneg
cnf(i_0_886, plain, (m1_subset_1(k3_seq_4(X1),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3167, plain, (v2_xxreal_0(np__0)|r1_xxreal_0(np__0,esk24_0)|~v2_xxreal_0(esk24_0)|~v1_xreal_0(esk24_0))).#trainneg
cnf(i_0_2353, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v1_membered(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_2354, negated_conjecture, (v1_relat_1(esk2_0)|~v4_funct_1(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_888, plain, (v1_funct_1(esk47_1(X1))|v1_xboole_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_890, plain, (v1_funct_1(esk38_1(X1))|v1_xboole_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4155, plain, (r1_xxreal_0(esk4_3(X1,X2,esk9_4(X3,X4,X5,X6)),X3)|r2_hidden(esk9_4(X3,X4,X5,X6),X2)|~v2_seq_4(X2,X1)|~r2_hidden(X5,X2)|~r2_hidden(X3,a_3_3_seq_4(X4,X5,X6))|~m1_subset_1(k45_valued_1(X5,esk9_4(X3,X4,X5,X6)),k14_seq_4(X1))|~m1_subset_1(esk9_4(X3,X4,X5,X6),k14_seq_4(X4))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X6,k1_zfmisc_1(k14_seq_4(X4)))|~m1_subset_1(X5,k14_seq_4(X4))|~m1_subset_1(esk9_4(X3,X4,X5,X6),k14_seq_4(X1))|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X5,k14_seq_4(X1)))).#trainneg
cnf(i_0_4626, negated_conjecture, (k3_seq_4(X1)=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~v3_membered(X1)|~m1_subset_1(X2,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_4493, plain, (v3_xxreal_0(X1)|v1_xboole_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_4483, plain, (k4_xcmplx_0(X1)=np__0|v3_xxreal_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers))).#trainneg
cnf(i_0_4592, plain, (r1_xxreal_0(X1,np__0)|v2_seq_4(np__0,X2)|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4540, plain, (k7_square_1(X1)=k24_seq_4(X2,X3,X4)|m1_subset_1(a_3_3_seq_4(X2,X3,X4),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1453, plain, (v1_partfun1(esk23_2(X1,X1),X1))).#trainneg
cnf(i_0_2369, negated_conjecture, (r1_xxreal_0(esk1_0,esk1_0))).#trainneg
cnf(i_0_2334, plain, (r1_xxreal_0(X1,np__0)|X1!=k3_seq_4(k5_numbers)|~v1_xreal_0(X1)|~v3_xxreal_2(k5_numbers))).#trainneg
cnf(i_0_891, plain, (v1_funct_1(esk32_1(X1))|v1_xboole_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2493, plain, (v7_ordinal1(esk46_0)|~v1_finseq_1(esk46_0)|~v1_funct_1(esk46_0)|~v1_relat_1(esk46_0))).#trainneg
cnf(i_0_2355, negated_conjecture, (v1_funct_1(esk2_0)|~v4_funct_1(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_2356, negated_conjecture, (v1_finseq_1(esk2_0)|~v3_finseq_1(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_4452, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_894, plain, (v1_funct_1(esk28_1(X1))|v1_xboole_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3880, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|v1_xboole_0(X2)|~v2_seq_4(X2,X1)|~m1_subset_1(k19_seq_4(X6,X4,X5),k14_seq_4(X1))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X4,k14_seq_4(X6))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X6,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X6,k1_numbers,k5_numbers)|~m1_subset_1(k15_seq_4(X1,X3,k19_seq_4(X6,X4,X5)),X2))).#trainneg
cnf(i_0_4149, plain, (v2_xxreal_0(X1)|r2_hidden(X4,X3)|~v2_xxreal_0(esk4_3(X2,X3,X4))|~v1_xreal_0(X1)|~v1_xreal_0(esk4_3(X2,X3,X4))|~v2_seq_4(X3,X2)|~r2_hidden(X1,a_3_3_seq_4(X2,X4,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X4,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4231, plain, (v1_xboole_0(X1)|r1_xxreal_0(esk4_3(X2,X1,X3),X4)|r2_hidden(X3,X1)|~v2_seq_4(X1,X2)|~r2_hidden(X4,a_3_3_seq_4(X5,X6,X7))|~m1_subset_1(k15_seq_4(X2,X3,k45_valued_1(esk9_4(X4,X5,X6,X7),X6)),X1)|~m1_subset_1(k45_valued_1(esk9_4(X4,X5,X6,X7),X6),k14_seq_4(X2))|~m1_subset_1(esk9_4(X4,X5,X6,X7),k14_seq_4(X5))|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X6,k14_seq_4(X5))|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X5,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3817, negated_conjecture, (v1_relat_1(k45_valued_1(X1,esk2_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_2973, plain, (v1_xxreal_0(esk46_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1782, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_2335, plain, (r1_xxreal_0(X1,np__0)|X1!=k3_seq_4(k1_numbers)|~v1_xreal_0(X1)|~v3_xxreal_2(k1_numbers))).#trainneg
cnf(i_0_896, plain, (v1_funct_1(esk20_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4337, plain, (v1_xcmplx_0(esk46_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4118, plain, (v7_ordinal1(esk42_0)|~v7_ordinal1(k1_zfmisc_1(esk42_0)))).#trainneg
cnf(i_0_4119, plain, (v7_ordinal1(esk41_0)|~v7_ordinal1(k1_zfmisc_1(esk41_0)))).#trainneg
cnf(i_0_2601, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3178, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_898, plain, (v1_funct_1(esk18_1(k1_zfmisc_1(X1)))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_899, negated_conjecture, (v1_funct_1(esk3_0)|~v1_funct_1(k14_seq_4(esk1_0))|~v1_relat_1(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_3940, plain, (v1_xboole_0(X1)|r1_xxreal_0(esk4_3(X2,X1,X3),X4)|r2_hidden(X3,X1)|~v2_seq_4(X1,X2)|~r2_hidden(X4,a_3_3_seq_4(X5,X6,X7))|~m1_subset_1(k1_valued_1(k19_seq_4(X5,esk9_4(X4,X5,X6,X7),X6),X3),X1)|~m1_subset_1(k19_seq_4(X5,esk9_4(X4,X5,X6,X7),X6),k14_seq_4(X2))|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers)|~m2_finseq_2(X6,k2_numbers,k14_seq_4(X5))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3905, negated_conjecture, (k5_seq_4(X1)=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_3370, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(X5,X4)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(X6))|~m1_subset_1(X4,k14_seq_4(X6))|~m1_subset_1(X6,k5_numbers)|~m2_finseq_2(k45_valued_1(X5,X4),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3930, plain, (v1_xboole_0(X1)|r1_xxreal_0(esk4_3(X2,X1,esk9_4(X3,X4,X5,X6)),X3)|r2_hidden(esk9_4(X3,X4,X5,X6),X1)|~v2_seq_4(X1,X2)|~r2_hidden(X3,a_3_3_seq_4(X4,X5,X6))|~m1_subset_1(X5,X1)|~m1_subset_1(esk9_4(X3,X4,X5,X6),k14_seq_4(X4))|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X6,k1_zfmisc_1(k14_seq_4(X4)))|~m1_subset_1(X5,k14_seq_4(X4))|~m1_subset_1(X4,k5_numbers)|~m2_finseq_2(k45_valued_1(X5,esk9_4(X3,X4,X5,X6)),k2_numbers,k14_seq_4(X2))|~m2_finseq_2(esk9_4(X3,X4,X5,X6),k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(X3,X4,X5,X6),k14_seq_4(X2))|~m1_subset_1(X5,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4120, plain, (v7_ordinal1(esk27_0)|~v7_ordinal1(k1_zfmisc_1(esk27_0)))).#trainneg
cnf(i_0_1789, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_3117, plain, (r1_xxreal_0(X1,esk22_0)|X1!=k3_seq_4(k1_numbers)|~v1_xreal_0(X1)|~v3_xxreal_2(k1_numbers))).#trainneg
cnf(i_0_1381, plain, (v5_relat_1(k5_numbers,X1)|~v5_relat_1(k1_numbers,X1)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_4557, negated_conjecture, (v3_ordinal1(esk2_0)|~v7_ordinal1(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_936, negated_conjecture, (v7_ordinal1(esk3_0)|~v6_membered(k1_zfmisc_1(k14_seq_4(esk1_0))))).#trainneg
cnf(i_0_3196, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_903, plain, (v2_xxreal_0(k4_xcmplx_0(X1))|~v3_xxreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5172, negated_conjecture, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,esk2_0)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(esk2_0,X4)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X4,k14_seq_4(esk1_0))|~m2_finseq_2(k45_valued_1(esk2_0,X4),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4784, plain, (v3_xxreal_0(esk4_3(X1,X2,X3))|v2_xxreal_0(esk10_2(a_3_3_seq_4(X1,X3,X2),X4))|v1_xboole_0(esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|m2_xxreal_2(X4,a_3_3_seq_4(X1,X3,X2))|~v1_xreal_0(esk4_3(X1,X2,X3))|~v1_xreal_0(esk10_2(a_3_3_seq_4(X1,X3,X2),X4))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xxreal_0(X4)|~v2_membered(a_3_3_seq_4(X1,X3,X2)))).#trainneg
cnf(i_0_4671, negated_conjecture, (v2_xxreal_0(esk1_0)|v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_4708, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk40_2(X3,k1_zfmisc_1(X2)))|~r2_hidden(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_5178, plain, (v1_valued_0(k2_zfmisc_1(X1,k2_numbers)))).#trainneg
cnf(i_0_4960, negated_conjecture, (v1_relat_1(esk2_0))).#trainneg
cnf(i_0_5041, plain, (~v1_xboole_0(k14_seq_4(esk46_0))|~m2_subset_1(esk46_0,k1_numbers,k5_numbers)|~v1_finseq_1(esk46_0)|~v1_funct_1(esk46_0)|~v1_relat_1(esk46_0))).#trainneg
cnf(i_0_906, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|~v2_xxreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4670, negated_conjecture, (~v3_xxreal_0(esk1_0))).#trainneg
cnf(i_0_4758, plain, (r1_xxreal_0(k5_numbers,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_909, plain, (v1_finset_1(X1)|~v1_finseq_1(k1_card_1(X1))|~v1_funct_1(k1_card_1(X1))|~v1_relat_1(k1_card_1(X1)))).#trainneg
cnf(i_0_5044, negated_conjecture, (v1_funct_1(esk2_0))).#trainneg
cnf(i_0_4980, negated_conjecture, (k3_seq_4(X1)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v3_membered(X1))).#trainneg
cnf(i_0_4951, plain, (k3_seq_4(X1)=k24_seq_4(X2,X3,X4)|m1_subset_1(a_3_3_seq_4(X2,X3,X4),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4858, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5179, plain, (v3_valued_0(k2_zfmisc_1(X1,k1_numbers)))).#trainneg
cnf(i_0_4659, plain, (v3_xxreal_0(esk35_0)|r1_xxreal_0(np__0,esk35_0))).#trainneg
cnf(i_0_2307, plain, (r1_xxreal_0(X1,esk46_0)|X1!=k3_seq_4(k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1)|~v3_xxreal_2(k1_zfmisc_1(k1_numbers))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_911, plain, (v1_xxreal_0(k3_seq_4(X1))|v1_xboole_0(X1)|~v2_membered(X1)|~v3_xxreal_2(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_5183, plain, (v5_relat_1(k5_numbers,X1)|~v1_relat_1(k1_numbers)|~m1_finseq_1(k1_numbers,X1))).#trainneg
cnf(i_0_912, plain, (v1_xxreal_0(esk23_2(X1,X2))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_913, plain, (v3_ordinal1(esk23_2(X1,X2))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4782, plain, (v3_xxreal_0(esk4_3(X1,X2,X3))|v2_xxreal_0(esk12_3(a_3_3_seq_4(X1,X3,X2),X4,X5))|v1_xboole_0(esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|v1_xboole_0(a_3_3_seq_4(X1,X3,X2))|r1_xxreal_0(X5,np__0)|~v1_xreal_0(esk4_3(X1,X2,X3))|~v1_xreal_0(esk12_3(a_3_3_seq_4(X1,X3,X2),X4,X5))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|X4!=k3_seq_4(a_3_3_seq_4(X1,X3,X2))|~v1_xreal_0(X5)|~v1_xreal_0(X4)|~v3_xxreal_2(a_3_3_seq_4(X1,X3,X2))|~v3_membered(a_3_3_seq_4(X1,X3,X2)))).#trainneg
cnf(i_0_5074, plain, (v2_xxreal_0(esk10_2(a_3_3_seq_4(X1,X2,X3),X4))|r2_hidden(X2,X3)|m2_xxreal_2(X4,a_3_3_seq_4(X1,X2,X3))|~v2_xxreal_0(esk4_3(X1,X3,X2))|~v1_xreal_0(esk4_3(X1,X3,X2))|~v1_xreal_0(esk10_2(a_3_3_seq_4(X1,X2,X3),X4))|~v2_seq_4(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xxreal_0(X4)|~v2_membered(a_3_3_seq_4(X1,X2,X3)))).#trainneg
cnf(i_0_5166, negated_conjecture, (k5_seq_4(k1_numbers)=k24_seq_4(esk1_0,X1,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_4781, plain, (v3_xxreal_0(esk4_3(X1,X2,X3))|v2_xxreal_0(esk13_2(a_3_3_seq_4(X1,X3,X2),X4))|v1_xboole_0(esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|X4=k3_seq_4(a_3_3_seq_4(X1,X3,X2))|v1_xboole_0(a_3_3_seq_4(X1,X3,X2))|~v1_xreal_0(esk4_3(X1,X2,X3))|~v1_xreal_0(esk13_2(a_3_3_seq_4(X1,X3,X2),X4))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xreal_0(X4)|~v3_xxreal_2(a_3_3_seq_4(X1,X3,X2))|~v3_membered(a_3_3_seq_4(X1,X3,X2))|~r1_xxreal_0(esk14_2(a_3_3_seq_4(X1,X3,X2),X4),np__0))).#trainneg
cnf(i_0_4887, plain, (v3_xxreal_0(esk4_3(X1,X2,X3))|v2_xxreal_0(esk10_2(a_3_3_seq_4(X1,X3,X2),X4))|v1_xboole_0(esk10_2(a_3_3_seq_4(X1,X3,X2),X4))|r2_hidden(X3,X2)|m2_xxreal_2(X4,a_3_3_seq_4(X1,X3,X2))|~v1_xreal_0(esk4_3(X1,X2,X3))|~v1_xreal_0(esk10_2(a_3_3_seq_4(X1,X3,X2),X4))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xxreal_0(X4)|~v2_membered(a_3_3_seq_4(X1,X3,X2)))).#trainneg
cnf(i_0_4739, plain, (r1_xxreal_0(np__0,esk22_0))).#trainneg
cnf(i_0_3058, plain, (r1_xxreal_0(X1,esk39_0)|X1!=k3_seq_4(k1_numbers)|~v1_xreal_0(X1)|~v3_xxreal_2(k1_numbers))).#trainneg
cnf(i_0_916, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,esk23_2(X2,X3)))).#trainneg
cnf(i_0_917, plain, (v1_xboole_0(esk23_2(X1,X2))|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_922, plain, (v7_ordinal1(esk48_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4939, plain, (v1_xreal_0(k3_seq_4(X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5010, plain, (r1_xxreal_0(np__0,k3_seq_4(X1))|r1_xxreal_0(k3_seq_4(X1),np__0)|~v3_membered(X1))).#trainneg
cnf(i_0_4794, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|v1_xboole_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xboole_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4885, plain, (v3_xxreal_0(esk4_3(X1,X2,X3))|v2_xxreal_0(esk12_3(a_3_3_seq_4(X1,X3,X2),X4,X5))|v1_xboole_0(esk12_3(a_3_3_seq_4(X1,X3,X2),X4,X5))|r2_hidden(X3,X2)|v1_xboole_0(a_3_3_seq_4(X1,X3,X2))|r1_xxreal_0(X5,np__0)|~v1_xreal_0(esk4_3(X1,X2,X3))|~v1_xreal_0(esk12_3(a_3_3_seq_4(X1,X3,X2),X4,X5))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|X4!=k3_seq_4(a_3_3_seq_4(X1,X3,X2))|~v1_xreal_0(X5)|~v1_xreal_0(X4)|~v3_xxreal_2(a_3_3_seq_4(X1,X3,X2))|~v3_membered(a_3_3_seq_4(X1,X3,X2)))).#trainneg
cnf(i_0_5162, negated_conjecture, (k5_seq_4(np__0)=k24_seq_4(esk1_0,X1,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_4950, plain, (k3_seq_4(X1)=k24_seq_4(X2,X3,X4)|v3_membered(a_3_3_seq_4(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4884, plain, (v3_xxreal_0(esk4_3(X1,X2,X3))|v2_xxreal_0(esk13_2(a_3_3_seq_4(X1,X3,X2),X4))|v1_xboole_0(esk13_2(a_3_3_seq_4(X1,X3,X2),X4))|r2_hidden(X3,X2)|X4=k3_seq_4(a_3_3_seq_4(X1,X3,X2))|v1_xboole_0(a_3_3_seq_4(X1,X3,X2))|~v1_xreal_0(esk4_3(X1,X2,X3))|~v1_xreal_0(esk13_2(a_3_3_seq_4(X1,X3,X2),X4))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xreal_0(X4)|~v3_xxreal_2(a_3_3_seq_4(X1,X3,X2))|~v3_membered(a_3_3_seq_4(X1,X3,X2))|~r1_xxreal_0(esk14_2(a_3_3_seq_4(X1,X3,X2),X4),np__0))).#trainneg
cnf(i_0_5087, plain, (v1_xboole_0(X1)|r1_xxreal_0(esk4_3(X2,X1,X3),X4)|r2_hidden(X3,X1)|~v2_seq_4(X1,X2)|~r2_hidden(X4,a_3_3_seq_4(X5,X3,X6))|~m1_subset_1(esk9_4(X4,X5,X3,X6),X1)|~m1_subset_1(k45_valued_1(esk9_4(X4,X5,X3,X6),X3),k14_seq_4(X2))|~m1_subset_1(esk9_4(X4,X5,X3,X6),k14_seq_4(X5))|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X6,k1_zfmisc_1(k14_seq_4(X5)))|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X3,k14_seq_4(X5))|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X5,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(esk9_4(X4,X5,X3,X6),k14_seq_4(X2)))).#trainneg
cnf(i_0_3269, plain, (v1_relat_1(k30_valued_1(esk23_2(X1,k1_numbers))))).#trainneg
cnf(i_0_3106, plain, (r1_xxreal_0(X1,esk35_0)|X1!=k3_seq_4(k1_numbers)|~v1_xreal_0(X1)|~v3_xxreal_2(k1_numbers))).#trainneg
cnf(i_0_923, plain, (v7_ordinal1(esk47_1(X1))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4752, plain, (~v1_finset_1(esk39_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)),k5_numbers))).#trainneg
cnf(i_0_925, plain, (v7_ordinal1(esk38_1(X1))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_926, plain, (v7_ordinal1(esk32_1(X1))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4786, plain, (v3_xxreal_0(esk4_3(X1,X2,X3))|v2_xxreal_0(esk8_2(X4,a_3_3_seq_4(X1,X3,X2)))|v1_xboole_0(esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|a_3_3_seq_4(X1,X3,X2)=np__0|~v1_xreal_0(esk4_3(X1,X2,X3))|~v1_xreal_0(esk8_2(X4,a_3_3_seq_4(X1,X3,X2)))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(a_3_3_seq_4(X1,X3,X2),k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_5164, negated_conjecture, (k5_seq_4(a_3_3_seq_4(X1,X2,X3))=k24_seq_4(esk1_0,X4,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X4,esk3_0),k1_zfmisc_1(k1_numbers))|k24_seq_4(X1,X2,X3)=np__0|~m1_subset_1(X4,k14_seq_4(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4908, plain, (v3_xxreal_0(esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|m2_xxreal_2(X4,a_3_3_seq_4(X1,X3,X2))|~v3_xxreal_0(esk10_2(a_3_3_seq_4(X1,X3,X2),X4))|~v1_xreal_0(esk4_3(X1,X2,X3))|~v1_xreal_0(esk10_2(a_3_3_seq_4(X1,X3,X2),X4))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xxreal_0(X4)|~v2_membered(a_3_3_seq_4(X1,X3,X2)))).#trainneg
cnf(i_0_5165, negated_conjecture, (k5_seq_4(k5_numbers)=k24_seq_4(esk1_0,X1,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_5071, plain, (v2_xxreal_0(esk13_2(a_3_3_seq_4(X1,X2,X3),X4))|r2_hidden(X2,X3)|X4=k3_seq_4(a_3_3_seq_4(X1,X2,X3))|v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|~v2_xxreal_0(esk4_3(X1,X3,X2))|~v1_xreal_0(esk4_3(X1,X3,X2))|~v1_xreal_0(esk13_2(a_3_3_seq_4(X1,X2,X3),X4))|~v2_seq_4(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xreal_0(X4)|~v3_xxreal_2(a_3_3_seq_4(X1,X2,X3))|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~r1_xxreal_0(esk14_2(a_3_3_seq_4(X1,X2,X3),X4),np__0))).#trainneg
cnf(i_0_4667, plain, (m2_xxreal_2(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|v1_xboole_0(a_3_3_seq_4(X1,X3,X2))|~v1_xxreal_0(esk4_3(X1,X2,X3))|~v2_membered(X4)|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(esk10_2(X4,esk4_3(X1,X2,X3)),a_3_3_seq_4(X1,X3,X2)))).#trainneg
cnf(i_0_5029, negated_conjecture, (k7_square_1(X1)=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1)|~m1_subset_1(X2,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_4776, plain, (v1_xxreal_0(k24_seq_4(X1,X2,X3))|~v2_membered(k1_numbers)|~m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3270, plain, (v1_funct_1(k30_valued_1(esk23_2(X1,k1_numbers))))).#trainneg
cnf(i_0_2305, plain, (r1_xxreal_0(X1,esk48_0)|X1!=k3_seq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(X1)|~v3_xxreal_2(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_927, plain, (v7_ordinal1(esk31_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_5042, plain, (r1_xxreal_0(esk46_0,esk46_0)|~v1_finseq_1(esk46_0)|~v1_funct_1(esk46_0)|~v1_relat_1(esk46_0))).#trainneg
cnf(i_0_4750, plain, (~v7_ordinal1(k1_zfmisc_1(k1_numbers))|~v1_finseq_1(esk39_0)|~v1_funct_1(esk39_0)|~v1_relat_1(esk39_0))).#trainneg
cnf(i_0_1124, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_928, plain, (v7_ordinal1(esk30_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_929, plain, (v7_ordinal1(esk28_1(X1))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5018, plain, (r1_xxreal_0(np__0,X1)|r1_xxreal_0(X1,np__0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4987, plain, (X1=np__0|v3_xxreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k5_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4926, plain, (v1_xboole_0(X1)|~v1_xboole_0(k4_xcmplx_0(X1))|~v1_xxreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5017, plain, (r1_xxreal_0(np__0,X1)|r1_xxreal_0(X1,np__0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3305, plain, (v1_relat_1(k30_valued_1(esk21_2(X1,k1_numbers))))).#trainneg
cnf(i_0_2310, plain, (r1_xxreal_0(X1,esk31_0)|X1!=k3_seq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)))|~v1_xreal_0(X1)|~v3_xxreal_2(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_930, plain, (v7_ordinal1(esk23_2(X1,X2))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4849, plain, (v1_valued_0(esk46_0)|~v1_relat_1(esk46_0)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_4958, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v3_membered(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_962, negated_conjecture, (v1_xreal_0(esk3_0)|~v3_membered(k1_zfmisc_1(k14_seq_4(esk1_0))))).#trainneg
cnf(i_0_1933, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_4358, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_932, plain, (v7_ordinal1(esk21_2(X1,X2))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_933, plain, (v7_ordinal1(esk20_1(X1))|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4875, plain, (v3_xxreal_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_5016, plain, (r1_xxreal_0(np__0,k24_seq_4(X1,X2,X3))|r1_xxreal_0(k24_seq_4(X1,X2,X3),np__0)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2001, plain, (r2_hidden(k22_seq_4(k45_valued_1(X3,X2)),a_3_9_seq_4(X1,X2,X4))|~r2_hidden(X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3681, plain, (v3_xxreal_0(esk4_3(X1,X2,X3))|v1_xboole_0(X2)|r2_hidden(X3,X2)|~v3_xxreal_0(k22_seq_4(k45_valued_1(X4,X3)))|~v1_xreal_0(k22_seq_4(k45_valued_1(X4,X3)))|~v1_xreal_0(esk4_3(X1,X2,X3))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X4,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X4,X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2477, plain, (k24_seq_4(X1,X2,X3)=np__0|v3_membered(a_3_3_seq_4(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2139, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(X6,esk11_4(X4,X5,X6,X7))),X2)|~r2_hidden(X4,a_3_9_seq_4(X5,X6,X7))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m2_finseq_2(k45_valued_1(X6,esk11_4(X4,X5,X6,X7)),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X6,k2_numbers,k14_seq_4(X5))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(esk11_4(X4,X5,X6,X7),k14_seq_4(X5))|~m1_subset_1(X6,k14_seq_4(X5))|~m1_subset_1(X5,k5_numbers))).#trainneg
cnf(i_0_2759, plain, (v3_xxreal_0(np__0)|v2_xxreal_0(k22_seq_4(k19_seq_4(X1,X2,X3)))|v1_xboole_0(k22_seq_4(k19_seq_4(X1,X2,X3)))|~v1_xreal_0(k22_seq_4(k19_seq_4(X1,X2,X3)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4864, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k24_seq_4(X2,X3,X4))|r1_xxreal_0(k24_seq_4(X2,X3,X4),X1)|k24_seq_4(X2,X3,X4)=np__0|~v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5100, plain, (v7_ordinal1(esk42_0)|~m1_subset_1(k1_zfmisc_1(esk42_0),k5_numbers))).#trainneg
cnf(i_0_3306, plain, (v1_funct_1(k30_valued_1(esk21_2(X1,k1_numbers))))).#trainneg
cnf(i_0_5102, plain, (v7_ordinal1(esk41_0)|~m1_subset_1(k1_zfmisc_1(esk41_0),k5_numbers))).#trainneg
cnf(i_0_2311, plain, (r1_xxreal_0(X1,esk30_0)|X1!=k3_seq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)))|~v1_xreal_0(X1)|~v3_xxreal_2(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_935, plain, (v7_ordinal1(esk18_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_5039, plain, (v1_xreal_0(esk46_0)|~v1_finseq_1(esk46_0)|~v1_funct_1(esk46_0)|~v1_relat_1(esk46_0))).#trainneg
cnf(i_0_988, negated_conjecture, (v1_xcmplx_0(esk3_0)|~v1_membered(k1_zfmisc_1(k14_seq_4(esk1_0))))).#trainneg
cnf(i_0_1014, negated_conjecture, (v1_finseq_1(esk3_0)|~v3_finseq_1(k1_zfmisc_1(k14_seq_4(esk1_0))))).#trainneg
cnf(i_0_4359, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_5548, plain, (r2_hidden(X3,X2)|~v3_xxreal_0(esk4_3(X1,X2,X3))|~v1_xreal_0(esk4_3(X1,X2,X3))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5519, negated_conjecture, (k5_seq_4(a_3_3_seq_4(X1,X2,X3))=k24_seq_4(esk1_0,X4,esk3_0)|k24_seq_4(X1,X2,X3)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X4,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X4,k14_seq_4(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k14_seq_4(X1)))).#trainneg
cnf(i_0_5324, plain, (k3_seq_4(X1)=k24_seq_4(X2,X3,X4)|m1_subset_1(a_3_3_seq_4(X2,X3,X4),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X3,k14_seq_4(X2)))).#trainneg
cnf(i_0_5323, plain, (k3_seq_4(X1)=k24_seq_4(X2,X3,X4)|m1_subset_1(a_3_3_seq_4(X2,X3,X4),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_5232, plain, (v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|r1_xxreal_0(esk4_3(X1,X3,X2),np__0)|r2_hidden(X2,X3)|k24_seq_4(X1,X2,X3)!=np__0|~v1_xreal_0(esk4_3(X1,X3,X2))|~v3_xxreal_2(a_3_3_seq_4(X1,X2,X3))|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~v2_seq_4(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1)))).#trainneg
cnf(i_0_5177, plain, (v7_ordinal1(esk27_0)|~m1_subset_1(k1_zfmisc_1(esk27_0),k5_numbers))).#trainneg
cnf(i_0_5329, plain, (v1_valued_0(k2_zfmisc_1(X1,k1_numbers)))).#trainneg
cnf(i_0_5392, plain, (v2_xxreal_0(esk22_0))).#trainneg
cnf(i_0_5278, negated_conjecture, (k7_square_1(X1)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k7_square_1(X1)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_939, plain, (v7_ordinal1(k8_seq_4)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k2_numbers,k1_numbers))))).#trainneg
cnf(i_0_5765, negated_conjecture, (v1_xcmplx_0(esk3_0)|~m1_subset_1(k1_zfmisc_1(k14_seq_4(esk1_0)),k1_zfmisc_1(k2_numbers)))).#trainneg
cnf(i_0_5275, negated_conjecture, (esk4_3(X1,X2,X3)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|r2_hidden(X3,X2)|esk4_3(X1,X2,X3)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5518, negated_conjecture, (k5_seq_4(a_3_3_seq_4(X1,X2,X3))=k24_seq_4(esk1_0,X4,esk3_0)|k24_seq_4(X1,X2,X3)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X4,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X4,k14_seq_4(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5233, plain, (v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|r1_xxreal_0(esk4_3(X1,X3,X2),np__0)|r2_hidden(X2,X3)|k24_seq_4(X1,X2,X3)!=np__0|~v1_xreal_0(esk4_3(X1,X3,X2))|~v3_xxreal_2(a_3_3_seq_4(X1,X2,X3))|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~v2_seq_4(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1)))).#trainneg
cnf(i_0_5515, negated_conjecture, (k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,X1))=k24_seq_4(esk1_0,X2,esk3_0)|k24_seq_4(esk1_0,esk2_0,X1)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(esk1_0)))|~m1_subset_1(X2,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_5691, plain, (v3_xxreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_5359, plain, (m2_xxreal_2(X4,a_3_3_seq_4(X1,X2,X3))|r2_hidden(X2,X3)|~v1_xboole_0(esk10_2(a_3_3_seq_4(X1,X2,X3),X4))|~v2_xxreal_0(esk4_3(X1,X3,X2))|~v1_xreal_0(esk10_2(a_3_3_seq_4(X1,X2,X3),X4))|~v1_xreal_0(esk4_3(X1,X3,X2))|~v1_xxreal_0(X4)|~v2_membered(a_3_3_seq_4(X1,X2,X3))|~v2_seq_4(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5742, plain, (v1_xreal_0(esk42_0)|~m1_subset_1(k1_zfmisc_1(esk42_0),k5_numbers))).#trainneg
cnf(i_0_1761, negated_conjecture, (m1_subset_1(esk40_2(X1,esk3_0),k14_seq_4(esk1_0))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_5285, negated_conjecture, (k3_seq_4(X1)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k3_seq_4(X1)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_948, plain, (v1_xreal_0(esk48_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_5391, plain, (~v3_xxreal_0(esk22_0))).#trainneg
cnf(i_0_5345, plain, (v1_valued_0(k5_numbers)|~v1_relat_1(k5_numbers)|~m1_finseq_1(k1_numbers,k2_numbers))).#trainneg
cnf(i_0_949, plain, (v1_xreal_0(esk47_1(X1))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_951, plain, (v1_xreal_0(esk38_1(X1))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5358, plain, (m2_xxreal_2(X4,a_3_3_seq_4(X1,X2,X3))|r2_hidden(X2,X3)|~v3_xxreal_0(esk10_2(a_3_3_seq_4(X1,X2,X3),X4))|~v1_xxreal_0(esk10_2(a_3_3_seq_4(X1,X2,X3),X4))|~v2_xxreal_0(esk4_3(X1,X3,X2))|~v1_xreal_0(esk10_2(a_3_3_seq_4(X1,X2,X3),X4))|~v1_xreal_0(esk4_3(X1,X3,X2))|~v1_xxreal_0(X4)|~v2_membered(a_3_3_seq_4(X1,X2,X3))|~v2_seq_4(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5709, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,X3))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|k24_seq_4(X1,X2,X3)=np__0|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5235, plain, (v3_xxreal_0(esk4_3(X1,X3,X2))|v1_xboole_0(esk4_3(X1,X3,X2))|m2_xxreal_2(X4,a_3_3_seq_4(X1,X2,X3))|r2_hidden(X2,X3)|~v1_xboole_0(esk10_2(a_3_3_seq_4(X1,X2,X3),X4))|~v1_xreal_0(esk10_2(a_3_3_seq_4(X1,X2,X3),X4))|~v1_xreal_0(esk4_3(X1,X3,X2))|~v1_xxreal_0(X4)|~v2_membered(a_3_3_seq_4(X1,X2,X3))|~v2_seq_4(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5294, negated_conjecture, (X1=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|X1!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5749, plain, (v1_xreal_0(esk41_0)|~m1_subset_1(k1_zfmisc_1(esk41_0),k5_numbers))).#trainneg
cnf(i_0_1804, plain, (m1_subset_1(esk40_2(X1,k5_numbers),k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_5333, plain, (r1_xxreal_0(np__0,esk35_0))).#trainneg
cnf(i_0_5288, negated_conjecture, (k5_seq_4(X1)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k5_seq_4(X1)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_952, plain, (v1_xreal_0(esk32_1(X1))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5346, plain, (v3_valued_0(k5_numbers)|~v1_relat_1(k5_numbers)|~m1_finseq_1(k1_numbers,k1_numbers))).#trainneg
cnf(i_0_5255, plain, (~v2_xxreal_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1428, plain, (v4_relat_1(k5_numbers,X1)|~v4_relat_1(k1_numbers,X1)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_953, plain, (v1_xreal_0(esk31_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_954, plain, (v1_xreal_0(esk30_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_5382, negated_conjecture, (k5_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|k5_seq_4(k1_numbers)=k24_seq_4(esk1_0,X1,esk3_0)|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_5424, plain, (r1_xxreal_0(np__0,k3_seq_4(X1))|r1_xxreal_0(k3_seq_4(X1),np__0)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5542, negated_conjecture, (k5_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|k5_seq_4(k5_numbers)=k24_seq_4(esk1_0,X1,esk3_0)|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_5630, plain, (r1_xxreal_0(k5_seq_4(X1),np__0)|r1_xxreal_0(np__0,k5_seq_4(X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5316, negated_conjecture, (k3_seq_4(X1)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5745, plain, (r1_xxreal_0(esk42_0,esk42_0)|~m1_subset_1(k1_zfmisc_1(esk42_0),k5_numbers))).#trainneg
cnf(i_0_1877, plain, (m1_subset_1(esk40_2(X1,esk46_0),k1_numbers)|~r2_hidden(X1,esk46_0))).#trainneg
cnf(i_0_4334, negated_conjecture, (r1_xxreal_0(np__0,esk39_0)|v3_xxreal_0(esk1_0)|r1_xxreal_0(esk39_0,esk1_0))).#trainneg
cnf(i_0_5279, negated_conjecture, (k7_square_1(X1)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k7_square_1(X1)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_955, plain, (v1_xreal_0(esk28_1(X1))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5054, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4753, plain, (r1_xxreal_0(X1,k5_numbers)|r1_xxreal_0(k5_numbers,X1)|~v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_956, plain, (v1_xreal_0(esk23_2(X1,X2))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5373, negated_conjecture, (v3_membered(a_3_3_seq_4(esk1_0,X1,esk3_0))|k5_seq_4(k1_numbers)=k24_seq_4(esk1_0,X1,esk3_0)|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_5286, negated_conjecture, (k3_seq_4(X1)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k3_seq_4(X1)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v3_membered(X1))).#trainneg
cnf(i_0_5533, negated_conjecture, (v3_membered(a_3_3_seq_4(esk1_0,X1,esk3_0))|k5_seq_4(k5_numbers)=k24_seq_4(esk1_0,X1,esk3_0)|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_5477, negated_conjecture, (k3_seq_4(X1)=k24_seq_4(esk1_0,X2,esk3_0)|v3_membered(a_3_3_seq_4(esk1_0,X2,esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_5314, negated_conjecture, (k3_seq_4(k5_numbers)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5456, negated_conjecture, (v3_membered(a_3_3_seq_4(esk1_0,X1,esk3_0))|k5_seq_4(np__0)=k24_seq_4(esk1_0,X1,esk3_0)|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_5433, plain, (v1_xboole_0(k4_xcmplx_0(X1))|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5313, negated_conjecture, (k3_seq_4(np__0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5582, negated_conjecture, (v2_xxreal_0(X1)|r1_xxreal_0(X1,esk1_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5752, plain, (r1_xxreal_0(esk41_0,esk41_0)|~m1_subset_1(k1_zfmisc_1(esk41_0),k5_numbers))).#trainneg
cnf(i_0_2922, plain, (m1_subset_1(esk40_2(X1,np__0),X2)|~r2_hidden(X1,np__0))).#trainneg
cnf(i_0_5280, negated_conjecture, (k6_square_1(X1)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k6_square_1(X1)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_959, plain, (v1_xreal_0(esk20_1(X1))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4853, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_961, plain, (v1_xreal_0(esk18_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_963, plain, (v1_xreal_0(k18_rvsum_1(X1))|~m1_finseq_1(X1,k1_numbers))).#trainneg
cnf(i_0_5941, negated_conjecture, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,esk40_2(X5,esk3_0))))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(esk40_2(X5,esk3_0),X4)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X4,k14_seq_4(esk1_0))|~m2_finseq_2(k45_valued_1(esk40_2(X5,esk3_0),X4),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r2_hidden(X5,esk3_0))).#trainneg
cnf(i_0_5935, negated_conjecture, (r1_xxreal_0(np__0,k22_seq_4(esk40_2(X1,esk3_0)))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_6285, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk40_2(X3,np__0))|~r2_hidden(X3,np__0))).#trainneg
cnf(i_0_5996, plain, (v1_xreal_0(esk40_2(X1,k5_numbers))|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_6011, plain, (v2_xxreal_0(esk35_0))).#trainneg
cnf(i_0_5277, negated_conjecture, (k18_rvsum_1(X1)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k18_rvsum_1(X1)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_finseq_1(X1,k1_numbers))).#trainneg
cnf(i_0_965, plain, (v1_xreal_0(k8_seq_4)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k2_numbers,k1_numbers))))).#trainneg
cnf(i_0_6154, plain, (r1_xxreal_0(k5_numbers,X1)|r1_xxreal_0(X1,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6369, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_971, plain, (v1_xreal_0(k22_seq_4(X1))|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_974, plain, (v1_xcmplx_0(esk48_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_5819, negated_conjecture, (k5_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=k24_seq_4(esk1_0,X2,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_5979, negated_conjecture, (esk40_2(X1,k5_numbers)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|esk40_2(X1,k5_numbers)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_5883, negated_conjecture, (k22_seq_4(X1)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k22_seq_4(X1)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k18_rvsum_1(k12_rvsum_1(k13_seq_4(X1))),k1_numbers)|~m2_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_6168, negated_conjecture, (k24_seq_4(X1,X2,X3)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k24_seq_4(X1,X2,X3)!=np__0|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6110, negated_conjecture, (k22_seq_4(X1)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k22_seq_4(X1)!=np__0|~v1_xreal_0(k18_rvsum_1(k12_rvsum_1(k13_seq_4(X1))))|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m2_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_5947, negated_conjecture, (k24_seq_4(X1,X2,X3)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k24_seq_4(X1,X2,X3)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5855, negated_conjecture, (k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))=k24_seq_4(esk1_0,X1,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|esk4_3(esk1_0,esk3_0,esk2_0)=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_5875, plain, (v1_xreal_0(esk27_0)|~m1_subset_1(k1_zfmisc_1(esk27_0),k5_numbers))).#trainneg
cnf(i_0_6092, plain, (v1_xreal_0(esk40_2(X1,esk46_0))|~r2_hidden(X1,esk46_0))).#trainneg
cnf(i_0_5289, negated_conjecture, (k22_seq_4(X1)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k22_seq_4(X1)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_975, plain, (v1_xcmplx_0(esk47_1(X1))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5878, plain, (r1_xxreal_0(esk27_0,esk27_0)|~m1_subset_1(k1_zfmisc_1(esk27_0),k5_numbers))).#trainneg
cnf(i_0_6010, plain, (~v3_xxreal_0(esk35_0))).#trainneg
cnf(i_0_1278, negated_conjecture, (v1_valued_0(esk3_0)|~v1_valued_0(k14_seq_4(esk1_0))|~v1_relat_1(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_1294, negated_conjecture, (v3_valued_0(esk3_0)|~v3_valued_0(k14_seq_4(esk1_0))|~v1_relat_1(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_977, plain, (v1_xcmplx_0(esk38_1(X1))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_978, plain, (v1_xcmplx_0(esk32_1(X1))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5975, negated_conjecture, (k24_seq_4(X1,X2,X3)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6084, negated_conjecture, (k3_seq_4(k1_numbers)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6250, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|~v1_xboole_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5821, negated_conjecture, (k5_seq_4(a_3_3_seq_4(X1,X2,k14_seq_4(X1)))=k24_seq_4(esk1_0,X3,esk3_0)|k24_seq_4(X1,X2,k14_seq_4(X1))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X3,esk3_0),k1_zfmisc_1(k1_numbers))|~v7_ordinal1(X1)|~m1_subset_1(X3,k14_seq_4(esk1_0))|~m1_subset_1(X2,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5820, negated_conjecture, (k5_seq_4(a_3_3_seq_4(X1,X2,np__0))=k24_seq_4(esk1_0,X3,esk3_0)|k24_seq_4(X1,X2,np__0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X3,esk3_0),k1_zfmisc_1(k1_numbers))|~v7_ordinal1(X1)|~m1_subset_1(X3,k14_seq_4(esk1_0))|~m1_subset_1(X2,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5465, negated_conjecture, (k5_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|k5_seq_4(np__0)=k24_seq_4(esk1_0,X1,esk3_0)|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_5637, plain, (X1=np__0|~v7_ordinal1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5984, negated_conjecture, (k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))=esk1_0|esk1_0!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5833, plain, (k3_seq_4(X1)=k24_seq_4(X2,X3,k14_seq_4(X2))|m1_subset_1(a_3_3_seq_4(X2,X3,k14_seq_4(X2)),k1_zfmisc_1(k1_numbers))|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5924, plain, (m1_subset_1(esk42_0,k1_numbers)|~m1_subset_1(k1_zfmisc_1(esk42_0),k5_numbers))).#trainneg
cnf(i_0_6257, plain, (v1_relat_1(esk40_2(X1,np__0))|~r2_hidden(X1,np__0))).#trainneg
cnf(i_0_5807, negated_conjecture, (k5_seq_4(np__0)=np__0|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_5267, negated_conjecture, (esk40_2(X1,k1_numbers)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|esk40_2(X1,k1_numbers)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,k1_numbers))).#trainneg
cnf(i_0_979, plain, (v1_xcmplx_0(esk31_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_5925, plain, (r1_xxreal_0(np__0,esk42_0)|r1_xxreal_0(esk42_0,np__0)|~m1_subset_1(k1_zfmisc_1(esk42_0),k5_numbers))).#trainneg
cnf(i_0_6027, plain, (~v2_xxreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2420, negated_conjecture, (v1_xxreal_0(X1)|v1_xboole_0(esk3_0)|~v2_membered(k14_seq_4(esk1_0))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_2421, negated_conjecture, (v3_ordinal1(X1)|v1_xboole_0(esk3_0)|~v3_ordinal1(k14_seq_4(esk1_0))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_980, plain, (v1_xcmplx_0(esk30_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_981, plain, (v1_xcmplx_0(esk28_1(X1))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_5581, plain, (v3_xxreal_0(esk4_3(X1,X2,X3))|v2_xxreal_0(X4)|r1_xxreal_0(X4,esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|~v1_xreal_0(X4)|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5831, negated_conjecture, (k3_seq_4(X1)=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_5726, plain, (v2_xxreal_0(X3)|v1_xboole_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(esk9_4(X3,X1,X2,X4),k14_seq_4(X1))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r2_hidden(X3,a_3_3_seq_4(X1,X2,X4))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1)))).#trainneg
cnf(i_0_5290, negated_conjecture, (k24_seq_4(X1,X2,X3)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k24_seq_4(X1,X2,X3)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5511, plain, (a_3_3_seq_4(X2,X3,X4)=np__0|v3_xxreal_0(esk4_3(X2,X4,X3))|v1_xboole_0(esk4_3(X2,X4,X3))|r2_hidden(X3,X4)|~v1_xboole_0(esk8_2(X1,a_3_3_seq_4(X2,X3,X4)))|~v1_xreal_0(esk8_2(X1,a_3_3_seq_4(X2,X3,X4)))|~v1_xreal_0(esk4_3(X2,X4,X3))|~v2_seq_4(X4,X2)|~m1_subset_1(a_3_3_seq_4(X2,X3,X4),k1_zfmisc_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6081, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,X3))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|k24_seq_4(X1,X2,X3)=esk1_0|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5991, plain, (m1_subset_1(esk41_0,k1_numbers)|~m1_subset_1(k1_zfmisc_1(esk41_0),k5_numbers))).#trainneg
cnf(i_0_987, plain, (v1_xcmplx_0(esk18_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_5273, negated_conjecture, (esk8_2(X1,k1_numbers)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|np__0=k1_numbers|esk8_2(X1,k1_numbers)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k1_numbers,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_982, plain, (v1_xcmplx_0(esk23_2(X1,X2))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5992, plain, (r1_xxreal_0(np__0,esk41_0)|r1_xxreal_0(esk41_0,np__0)|~m1_subset_1(k1_zfmisc_1(esk41_0),k5_numbers))).#trainneg
cnf(i_0_2422, negated_conjecture, (v7_ordinal1(X1)|v1_xboole_0(esk3_0)|~v6_membered(k14_seq_4(esk1_0))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_2423, negated_conjecture, (v1_xreal_0(X1)|v1_xboole_0(esk3_0)|~v3_membered(k14_seq_4(esk1_0))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_984, plain, (v1_xcmplx_0(esk21_2(X1,X2))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_985, plain, (v1_xcmplx_0(esk20_1(X1))|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6008, plain, (esk40_2(X1,k5_numbers)=k24_seq_4(X2,X3,X4)|m1_subset_1(a_3_3_seq_4(X2,X3,X4),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_5695, plain, (v3_xxreal_0(k24_seq_4(X1,X2,X3))|v1_xboole_0(k24_seq_4(X1,X2,X3))|r1_xxreal_0(np__0,k24_seq_4(X1,X2,X3))|~v1_xreal_0(k24_seq_4(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5252, plain, (~v1_xboole_0(k4_xcmplx_0(X1))|~v1_xxreal_0(k4_xcmplx_0(X1))|~v2_xxreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_5558, negated_conjecture, (k7_square_1(X1)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5593, plain, (v2_xxreal_0(X1)|r1_xxreal_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5284, negated_conjecture, (k9_binop_2(X1,X2)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k9_binop_2(X1,X2)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6249, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|~v3_xxreal_0(X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3674, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(np__0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1013, plain, (v1_finseq_1(esk18_1(X1))|~v3_finseq_1(X1))).#trainneg
cnf(i_0_1992, plain, (v1_xboole_0(esk45_1(np__0))|~v1_card_1(np__0))).#trainneg
cnf(i_0_5982, negated_conjecture, (esk8_2(X1,k5_numbers)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|np__0=k5_numbers|esk8_2(X1,k5_numbers)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k5_numbers,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_989, plain, (v1_xcmplx_0(k18_rvsum_1(X1))|~v1_membered(k1_numbers)|~m1_finseq_1(X1,k1_numbers))).#trainneg
cnf(i_0_1248, negated_conjecture, (m2_finseq_1(esk2_0,k2_numbers)|~m1_finseq_2(k14_seq_4(esk1_0),k2_numbers))).#trainneg
cnf(i_0_2424, negated_conjecture, (v1_xcmplx_0(X1)|v1_xboole_0(esk3_0)|~v1_membered(k14_seq_4(esk1_0))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_2425, negated_conjecture, (v1_relat_1(X1)|v1_xboole_0(esk3_0)|~v4_funct_1(k14_seq_4(esk1_0))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_990, plain, (v1_xcmplx_0(k7_square_1(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_991, plain, (v1_xcmplx_0(k8_seq_4)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k2_numbers,k1_numbers))))).#trainneg
cnf(i_0_5938, negated_conjecture, (v3_xxreal_0(esk4_3(esk1_0,X1,X2))|v2_xxreal_0(k22_seq_4(k45_valued_1(X2,esk40_2(X3,esk3_0))))|v1_xboole_0(esk4_3(esk1_0,X1,X2))|v1_xboole_0(X1)|r2_hidden(X2,X1)|~v1_xreal_0(k22_seq_4(k45_valued_1(X2,esk40_2(X3,esk3_0))))|~v1_xreal_0(esk4_3(esk1_0,X1,X2))|~v2_seq_4(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(esk1_0)))|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(esk40_2(X3,esk3_0),X1)|~r2_hidden(X3,esk3_0))).#trainneg
cnf(i_0_5320, negated_conjecture, (k3_seq_4(X1)=k24_seq_4(esk1_0,esk2_0,X2)|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,X2),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5832, plain, (k3_seq_4(X1)=k24_seq_4(X2,X3,np__0)|m1_subset_1(a_3_3_seq_4(X2,X3,np__0),k1_zfmisc_1(k1_numbers))|~v7_ordinal1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5343, plain, (v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|v1_xxreal_0(k24_seq_4(X1,X2,X3))|~v3_xxreal_2(a_3_3_seq_4(X1,X2,X3))|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5283, negated_conjecture, (k2_xcmplx_0(X1,X2)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k2_xcmplx_0(X1,X2)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5429, plain, (v3_xxreal_0(k3_seq_4(X1))|v1_xboole_0(k3_seq_4(X1))|r1_xxreal_0(np__0,k3_seq_4(X1))|~v1_xreal_0(k3_seq_4(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_5371, negated_conjecture, (v2_membered(a_3_3_seq_4(esk1_0,X1,esk3_0))|k5_seq_4(k1_numbers)=k24_seq_4(esk1_0,X1,esk3_0)|~v2_membered(k1_numbers)|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_3675, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(np__0)))).#trainneg
cnf(i_0_1041, plain, (v2_membered(esk20_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_6098, negated_conjecture, (esk40_2(X1,esk46_0)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|esk40_2(X1,esk46_0)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,esk46_0))).#trainneg
cnf(i_0_996, plain, (v1_xcmplx_0(k5_seq_4(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2426, negated_conjecture, (v1_funct_1(X1)|v1_xboole_0(esk3_0)|~v4_funct_1(k14_seq_4(esk1_0))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_2427, negated_conjecture, (v1_finseq_1(X1)|v1_xboole_0(esk3_0)|~v3_finseq_1(k14_seq_4(esk1_0))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_997, plain, (v1_xcmplx_0(k22_seq_4(X1))|~v1_membered(k1_numbers)|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_1000, plain, (v1_finseq_1(esk48_0)|~v3_finseq_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_5274, negated_conjecture, (esk8_2(k1_numbers,X1)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|X1=np__0|esk8_2(k1_numbers,X1)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5912, negated_conjecture, (k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,np__0))=k24_seq_4(esk1_0,X1,esk3_0)|k24_seq_4(esk1_0,esk2_0,np__0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_5282, negated_conjecture, (k4_xcmplx_0(X1)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k4_xcmplx_0(X1)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6065, negated_conjecture, (k3_seq_4(esk46_0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5913, negated_conjecture, (k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,k14_seq_4(esk1_0)))=k24_seq_4(esk1_0,X1,esk3_0)|k24_seq_4(esk1_0,esk2_0,k14_seq_4(esk1_0))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_5222, negated_conjecture, (v3_xxreal_0(esk4_3(X1,X2,X3))|v2_xxreal_0(k22_seq_4(k45_valued_1(X4,esk2_0)))|v1_xboole_0(k22_seq_4(k45_valued_1(X4,esk2_0)))|r2_hidden(X3,X2)|~v1_xreal_0(k22_seq_4(k45_valued_1(X4,esk2_0)))|~v1_xreal_0(esk4_3(X1,X2,X3))|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(esk2_0,X4)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X4,k14_seq_4(esk1_0))|~m2_finseq_2(k45_valued_1(esk2_0,X4),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6821, negated_conjecture, (m2_finseq_1(esk2_0,k2_numbers))).#trainneg
cnf(i_0_6528, plain, (m1_subset_1(esk27_0,k1_numbers)|~m1_subset_1(k1_zfmisc_1(esk27_0),k5_numbers))).#trainneg
cnf(i_0_6810, plain, (v1_relat_1(esk45_1(np__0)))).#trainneg
cnf(i_0_6818, plain, (esk45_1(np__0)=np__0)).#trainneg
cnf(i_0_6312, negated_conjecture, (esk40_2(X1,np__0)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|esk40_2(X1,np__0)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,np__0))).#trainneg
cnf(i_0_1001, plain, (v1_finseq_1(esk47_1(X1))|v1_xboole_0(X1)|~v3_finseq_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6529, plain, (r1_xxreal_0(np__0,esk27_0)|r1_xxreal_0(esk27_0,np__0)|~m1_subset_1(k1_zfmisc_1(esk27_0),k5_numbers))).#trainneg
cnf(i_0_6653, negated_conjecture, (v3_ordinal1(X1)|v1_xboole_0(esk3_0)|~m1_subset_1(X1,esk3_0)|~v7_ordinal1(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_1003, plain, (v1_finseq_1(esk38_1(X1))|v1_xboole_0(X1)|~v3_finseq_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1004, plain, (v1_finseq_1(esk32_1(X1))|v1_xboole_0(X1)|~v3_finseq_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6401, plain, (v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|r1_xxreal_0(esk4_3(X1,X3,X2),np__0)|r2_hidden(X2,X3)|k24_seq_4(X1,X2,X3)!=np__0|~v1_xreal_0(esk4_3(X1,X3,X2))|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~v2_seq_4(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1031, plain, (v2_membered(esk48_0)|~v2_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_6786, negated_conjecture, (r1_xxreal_0(esk39_0,esk1_0))).#trainneg
cnf(i_0_1005, plain, (v1_finseq_1(esk31_0)|~v3_finseq_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_6824, negated_conjecture, (v1_xcmplx_0(X1)|v1_xboole_0(esk3_0)|~m1_subset_1(X1,esk3_0)|~m1_subset_1(k14_seq_4(esk1_0),k1_zfmisc_1(k2_numbers)))).#trainneg
cnf(i_0_6651, negated_conjecture, (v1_xboole_0(esk3_0)|v1_xxreal_0(X1)|~m1_subset_1(X1,esk3_0)|~v3_membered(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_1006, plain, (v1_finseq_1(esk30_0)|~v3_finseq_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_1007, plain, (v1_finseq_1(esk28_1(X1))|v1_xboole_0(X1)|~v3_finseq_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6752, negated_conjecture, (k7_square_1(esk4_3(X1,X2,X3))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6869, negated_conjecture, (X1=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|X1!=np__0|~v1_xreal_0(X1)|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6925, negated_conjecture, (X1=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|X1!=np__0|~v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6556, negated_conjecture, (k24_seq_4(X1,X2,X3)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k14_seq_4(X1)))).#trainneg
cnf(i_0_6677, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))=k24_seq_4(esk1_0,X1,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|esk4_3(esk1_0,esk3_0,esk2_0)=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_1036, plain, (v2_membered(esk31_0)|~v2_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_6399, negated_conjecture, (k3_seq_4(np__0)=np__0|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_5983, negated_conjecture, (esk8_2(k5_numbers,X1)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|X1=np__0|esk8_2(k5_numbers,X1)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1008, plain, (v1_finseq_1(esk23_2(X1,X2))|~v3_finseq_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5572, plain, (~v1_finseq_1(esk39_0)|~v1_funct_1(esk39_0)|~v1_relat_1(esk39_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k5_numbers))).#trainneg
cnf(i_0_6823, negated_conjecture, (v1_xcmplx_0(X1)|v1_xboole_0(esk3_0)|~m1_subset_1(X1,esk3_0)|~v3_membered(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_1205, plain, (v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_1010, plain, (v1_finseq_1(esk21_2(X1,X2))|~v3_finseq_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1011, plain, (v1_finseq_1(esk20_1(X1))|~v3_finseq_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6685, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,X3))=k24_seq_4(esk1_0,X4,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X4,esk3_0),k1_zfmisc_1(k1_numbers))|k24_seq_4(X1,X2,X3)=np__0|~m1_subset_1(X4,k14_seq_4(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6409, negated_conjecture, (~v3_xxreal_0(k22_seq_4(esk40_2(X1,esk3_0)))|~v1_xreal_0(k22_seq_4(esk40_2(X1,esk3_0)))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_6676, negated_conjecture, (k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_6410, negated_conjecture, (v2_xxreal_0(k22_seq_4(esk40_2(X1,esk3_0)))|v1_xboole_0(k22_seq_4(esk40_2(X1,esk3_0)))|~v1_xreal_0(k22_seq_4(esk40_2(X1,esk3_0)))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_6684, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,X3))=k24_seq_4(esk1_0,X4,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X4,esk3_0),k1_zfmisc_1(k1_numbers))|k24_seq_4(X1,X2,X3)=esk1_0|~m1_subset_1(X4,k14_seq_4(esk1_0))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6687, negated_conjecture, (k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(k1_numbers)|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_6769, negated_conjecture, (k7_square_1(k24_seq_4(X1,X2,X3))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6753, negated_conjecture, (k7_square_1(esk1_0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1037, plain, (v2_membered(esk30_0)|~v2_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1057, plain, (v3_membered(esk20_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_5271, negated_conjecture, (esk21_2(X1,X2)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|esk21_2(X1,X2)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1015, plain, (v1_finseq_1(k18_rvsum_1(X1))|~v3_finseq_1(k1_numbers)|~m1_finseq_1(X1,k1_numbers))).#trainneg
cnf(i_0_4755, plain, (r1_xxreal_0(k5_numbers,np__0)|r1_xxreal_0(np__0,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5495, plain, (~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_numbers)),k5_numbers)|~v1_finseq_1(esk39_0)|~v1_funct_1(esk39_0)|~v1_relat_1(esk39_0))).#trainneg
cnf(i_0_5744, plain, (~v1_xboole_0(k14_seq_4(esk42_0))|~m2_subset_1(esk42_0,k1_numbers,k5_numbers)|~m1_subset_1(k1_zfmisc_1(esk42_0),k5_numbers))).#trainneg
cnf(i_0_4122, negated_conjecture, (v7_ordinal1(esk3_0)|~v7_ordinal1(k1_zfmisc_1(esk3_0))|~m1_subset_1(k14_seq_4(esk1_0),k5_numbers))).#trainneg
cnf(i_0_3123, negated_conjecture, (v1_xxreal_0(esk3_0)|~v3_membered(k1_zfmisc_1(k14_seq_4(esk1_0))))).#trainneg
cnf(i_0_1016, plain, (v1_finseq_1(k7_square_1(X1))|~v3_finseq_1(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1017, plain, (v1_finseq_1(k8_seq_4)|~v3_finseq_1(k1_zfmisc_1(k2_zfmisc_1(k2_numbers,k1_numbers))))).#trainneg
cnf(i_0_6719, plain, (v1_xboole_0(k24_seq_4(X1,X2,X3))|r1_xxreal_0(np__0,k24_seq_4(X1,X2,X3))|~v7_ordinal1(k24_seq_4(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6686, negated_conjecture, (k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(k5_numbers)|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_6552, negated_conjecture, (k24_seq_4(esk1_0,esk2_0,X1)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(esk1_0))))).#trainneg
cnf(i_0_6683, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|k5_seq_4(k1_numbers)=k24_seq_4(esk1_0,X1,esk3_0)|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_6841, negated_conjecture, (v3_xxreal_0(esk4_3(esk1_0,esk3_0,X1))|v2_xxreal_0(k22_seq_4(k45_valued_1(X1,esk40_2(X2,esk3_0))))|v1_xboole_0(esk4_3(esk1_0,esk3_0,X1))|v1_xboole_0(esk3_0)|r2_hidden(X1,esk3_0)|~v1_xreal_0(k22_seq_4(k45_valued_1(X1,esk40_2(X2,esk3_0))))|~v1_xreal_0(esk4_3(esk1_0,esk3_0,X1))|~r2_hidden(X2,esk3_0)|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_6678, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))=k24_seq_4(esk1_0,X1,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|k3_seq_4(k5_numbers)=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_6682, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|k5_seq_4(k5_numbers)=k24_seq_4(esk1_0,X1,esk3_0)|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_6405, negated_conjecture, (v3_xxreal_0(esk4_3(X1,X2,X3))|v2_xxreal_0(k22_seq_4(k45_valued_1(X4,esk40_2(X5,esk3_0))))|v1_xboole_0(esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|~v1_xreal_0(k22_seq_4(k45_valued_1(X4,esk40_2(X5,esk3_0))))|~v1_xreal_0(esk4_3(X1,X2,X3))|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(esk40_2(X5,esk3_0),X4)),X2)|~r2_hidden(X5,esk3_0)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X4,k14_seq_4(esk1_0))|~m2_finseq_2(k45_valued_1(esk40_2(X5,esk3_0),X4),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1045, plain, (v2_membered(k8_seq_4)|~v2_membered(k2_zfmisc_1(k2_numbers,k1_numbers)))).#trainneg
cnf(i_0_1073, plain, (v1_finset_1(esk20_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_6787, negated_conjecture, (r1_xxreal_0(esk24_0,esk1_0)|~v1_xreal_0(esk24_0))).#trainneg
cnf(i_0_1022, plain, (v1_finseq_1(k5_seq_4(X1))|~v3_finseq_1(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6128, plain, (r1_xxreal_0(k5_numbers,esk39_0)|r1_xxreal_0(esk39_0,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5751, plain, (~v1_xboole_0(k14_seq_4(esk41_0))|~m2_subset_1(esk41_0,k1_numbers,k5_numbers)|~m1_subset_1(k1_zfmisc_1(esk41_0),k5_numbers))).#trainneg
cnf(i_0_5877, plain, (~v1_xboole_0(k14_seq_4(esk27_0))|~m2_subset_1(esk27_0,k1_numbers,k5_numbers)|~m1_subset_1(k1_zfmisc_1(esk27_0),k5_numbers))).#trainneg
cnf(i_0_4654, negated_conjecture, (v3_ordinal1(esk3_0)|~v7_ordinal1(k1_zfmisc_1(k14_seq_4(esk1_0))))).#trainneg
cnf(i_0_5764, negated_conjecture, (v1_xcmplx_0(esk3_0)|~v3_membered(k1_zfmisc_1(k14_seq_4(esk1_0))))).#trainneg
cnf(i_0_1023, plain, (v1_finseq_1(k22_seq_4(X1))|~v3_finseq_1(k1_numbers)|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_1030, plain, (v1_xxreal_0(esk16_1(X1))|~v2_membered(X1)|~v3_xxreal_2(X1))).#trainneg
cnf(i_0_6555, negated_conjecture, (k24_seq_4(X1,X2,X3)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6966, negated_conjecture, (v2_xxreal_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|v1_xboole_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|r2_hidden(X2,X3)|~v1_xreal_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|~v2_seq_4(X3,X4)|~r2_hidden(k15_seq_4(X4,k45_valued_1(esk2_0,X1),X2),X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X4)))|~m1_subset_1(X1,k14_seq_4(esk1_0))|~m2_finseq_2(k45_valued_1(esk2_0,X1),k2_numbers,k14_seq_4(X4))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X4))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X2,k14_seq_4(X4))|~m1_subset_1(k45_valued_1(esk2_0,X1),k14_seq_4(X4))|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_6873, plain, (v1_xboole_0(k3_seq_4(X1))|r1_xxreal_0(np__0,k3_seq_4(X1))|~v7_ordinal1(k3_seq_4(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_6781, plain, (r1_xxreal_0(X1,np__0)|~v3_xxreal_0(X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6681, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|k5_seq_4(np__0)=k24_seq_4(esk1_0,X1,esk3_0)|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_6666, negated_conjecture, (k3_seq_4(esk40_2(X1,k1_zfmisc_1(k1_numbers)))=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k14_seq_4(esk1_0))|~r2_hidden(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6763, negated_conjecture, (k7_square_1(np__0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6882, negated_conjecture, (v3_xxreal_2(a_3_3_seq_4(esk1_0,X1,esk3_0))|r2_hidden(X1,esk3_0)|k5_seq_4(k1_numbers)=k24_seq_4(esk1_0,X1,esk3_0)|~v1_xreal_0(esk4_3(esk1_0,esk3_0,X1))|~m1_subset_1(X1,k14_seq_4(esk1_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1047, plain, (v3_membered(esk48_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_1198, plain, (v1_xboole_0(esk49_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1032, plain, (v2_membered(esk47_1(X1))|v1_xboole_0(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_6130, plain, (r1_xxreal_0(k5_numbers,esk35_0)|r1_xxreal_0(esk35_0,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1052, plain, (v3_membered(esk31_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1053, plain, (v3_membered(esk30_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_4128, plain, (v7_ordinal1(X1)|~v7_ordinal1(k1_zfmisc_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4095, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~v7_ordinal1(k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1034, plain, (v2_membered(esk38_1(X1))|v1_xboole_0(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_1035, plain, (v2_membered(esk32_1(X1))|v1_xboole_0(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_6965, negated_conjecture, (v2_xxreal_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|v1_xboole_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|r2_hidden(X2,X3)|~v1_xreal_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|~v2_seq_4(X3,X4)|~r2_hidden(k1_valued_1(k45_valued_1(esk2_0,X1),X2),X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X4)))|~m1_subset_1(X1,k14_seq_4(esk1_0))|~m2_finseq_2(k45_valued_1(esk2_0,X1),k2_numbers,k14_seq_4(X4))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X4))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X2,k14_seq_4(X4))|~m1_subset_1(k45_valued_1(esk2_0,X1),k14_seq_4(X4))|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_6862, plain, (v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|v1_xxreal_0(k24_seq_4(X1,X2,X3))|r2_hidden(X2,X3)|~m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xreal_0(esk4_3(X1,X3,X2))|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~v2_seq_4(X3,X1)|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6729, plain, (v3_xxreal_0(X1)|~v2_xxreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers))).#trainneg
cnf(i_0_6767, negated_conjecture, (k7_square_1(k24_seq_4(X1,X2,X3))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6499, negated_conjecture, (k22_seq_4(X1)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k22_seq_4(X1)!=np__0|~m2_finseq_1(X1,k2_numbers)|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k18_rvsum_1(k12_rvsum_1(k13_seq_4(X1))),k5_numbers))).#trainneg
cnf(i_0_6472, plain, (v3_xxreal_0(k22_seq_4(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k22_seq_4(X1))|~v1_xreal_0(X2)|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_6494, negated_conjecture, (k5_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=np__0|k24_seq_4(esk1_0,X1,esk3_0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_6381, plain, (v3_xxreal_0(k18_rvsum_1(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k18_rvsum_1(X1))|~v1_xreal_0(X2)|~m1_finseq_1(X1,k1_numbers))).#trainneg
cnf(i_0_6761, negated_conjecture, (k7_square_1(k3_seq_4(X1))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v3_membered(X1))).#trainneg
cnf(i_0_6680, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))=k24_seq_4(esk1_0,X1,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|k3_seq_4(k1_numbers)=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_1061, plain, (v3_membered(k8_seq_4)|~v3_membered(k2_zfmisc_1(k2_numbers,k1_numbers)))).#trainneg
cnf(i_0_1236, plain, (v1_xboole_0(esk37_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6624, negated_conjecture, (k5_seq_4(np__0)=esk1_0|esk1_0!=np__0|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_1038, plain, (v2_membered(esk28_1(X1))|v1_xboole_0(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_6133, plain, (r1_xxreal_0(k5_numbers,esk24_0)|r1_xxreal_0(esk24_0,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1063, plain, (v1_finset_1(esk48_0)|~v1_finset_1(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_1068, plain, (v1_finset_1(esk31_0)|~v1_finset_1(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1488, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1039, plain, (v2_membered(esk23_2(X1,X2))|~v2_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1040, plain, (v2_membered(esk21_2(X1,X2))|~v2_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_7117, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,X3))=k24_seq_4(esk1_0,X4,esk3_0)|k24_seq_4(X1,X2,X3)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X4,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X4,k14_seq_4(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7118, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,X3))=k24_seq_4(esk1_0,X4,esk3_0)|k24_seq_4(X1,X2,X3)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X4,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X4,k14_seq_4(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k14_seq_4(X1)))).#trainneg
cnf(i_0_7114, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,esk2_0,X1))=k24_seq_4(esk1_0,X2,esk3_0)|k24_seq_4(esk1_0,esk2_0,X1)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(esk1_0)))|~m1_subset_1(X2,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_6972, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers))))).#trainneg
cnf(i_0_7150, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,esk2_0,X1))=k24_seq_4(esk1_0,X2,esk3_0)|k24_seq_4(esk1_0,esk2_0,X1)=esk1_0|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k14_seq_4(esk1_0)))|~m1_subset_1(X2,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7177, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(k1_numbers)|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7141, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_6970, negated_conjecture, (v1_finseq_1(esk2_0))).#trainneg
cnf(i_0_1043, plain, (v2_membered(esk18_1(k1_zfmisc_1(X1)))|~v2_membered(X1))).#trainneg
cnf(i_0_7001, negated_conjecture, (v2_xxreal_0(esk1_0)|~v2_xxreal_0(esk39_0))).#trainneg
cnf(i_0_7335, negated_conjecture, (v3_ordinal1(esk3_0)|~m1_subset_1(k1_zfmisc_1(k14_seq_4(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_7336, negated_conjecture, (v1_xcmplx_0(esk3_0)|~m1_subset_1(k1_zfmisc_1(k14_seq_4(esk1_0)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1048, plain, (v3_membered(esk47_1(X1))|v1_xboole_0(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_1050, plain, (v3_membered(esk38_1(X1))|v1_xboole_0(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_6991, plain, (v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|r1_xxreal_0(esk4_3(X1,X3,X2),np__0)|r2_hidden(X2,X3)|k24_seq_4(X1,X2,X3)!=np__0|~v1_xreal_0(esk4_3(X1,X3,X2))|~v2_seq_4(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7153, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,X3))=k24_seq_4(esk1_0,X4,esk3_0)|k24_seq_4(X1,X2,X3)=esk1_0|m1_subset_1(a_3_3_seq_4(esk1_0,X4,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X4,k14_seq_4(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7178, negated_conjecture, (k5_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(k1_numbers)|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7142, negated_conjecture, (k5_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7262, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(k5_numbers)|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7167, negated_conjecture, (v3_membered(a_3_3_seq_4(esk1_0,X1,esk3_0))|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(k1_numbers)|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7154, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,X3))=k24_seq_4(esk1_0,X4,esk3_0)|k24_seq_4(X1,X2,X3)=esk1_0|m1_subset_1(a_3_3_seq_4(esk1_0,X4,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X4,k14_seq_4(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k14_seq_4(X1)))).#trainneg
cnf(i_0_7263, negated_conjecture, (k5_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(k5_numbers)|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7131, negated_conjecture, (v3_membered(a_3_3_seq_4(esk1_0,X1,esk3_0))|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_1051, plain, (v3_membered(esk32_1(X1))|v1_xboole_0(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_1054, plain, (v3_membered(esk28_1(X1))|v1_xboole_0(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_1055, plain, (v3_membered(esk23_2(X1,X2))|~v3_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_7412, negated_conjecture, (k3_seq_4(esk40_2(X1,k1_zfmisc_1(k1_numbers)))=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7252, negated_conjecture, (v3_membered(a_3_3_seq_4(esk1_0,X1,esk3_0))|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(k5_numbers)|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7068, negated_conjecture, (X1=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|X1!=np__0|~v1_xcmplx_0(X1)|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_7216, negated_conjecture, (k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))=k24_seq_4(esk1_0,X1,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|k7_square_1(esk1_0)=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7215, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))=k24_seq_4(esk1_0,X1,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|k7_square_1(esk1_0)=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7036, negated_conjecture, (esk18_1(X1)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|esk18_1(X1)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v3_membered(X1))).#trainneg
cnf(i_0_1056, plain, (v3_membered(esk21_2(X1,X2))|~v3_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_7584, plain, (v1_finset_1(esk48_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_7229, negated_conjecture, (v7_ordinal1(esk3_0)|~m1_subset_1(k14_seq_4(esk1_0),k5_numbers)|~m1_subset_1(k1_zfmisc_1(esk3_0),k5_numbers))).#trainneg
cnf(i_0_6985, negated_conjecture, (v3_ordinal1(X1)|v1_xboole_0(esk3_0)|~v7_ordinal1(k14_seq_4(esk1_0))|~m1_subset_1(k1_zfmisc_1(X1),esk3_0))).#trainneg
cnf(i_0_1059, plain, (v3_membered(esk18_1(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_7381, plain, (r1_xxreal_0(esk39_0,np__0))).#trainneg
cnf(i_0_7278, negated_conjecture, (k24_seq_4(esk1_0,esk2_0,k14_seq_4(esk1_0))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7239, plain, (v1_xboole_0(k24_seq_4(X1,X2,X3))|r1_xxreal_0(np__0,k24_seq_4(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(k24_seq_4(X1,X2,X3),k5_numbers))).#trainneg
cnf(i_0_7176, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(k1_numbers)|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7377, plain, (v1_xboole_0(k24_seq_4(X1,X2,X3))|r1_xxreal_0(np__0,k24_seq_4(X1,X2,X3))|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~v7_ordinal1(k24_seq_4(X1,X2,X3))|~m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7378, plain, (v1_xboole_0(k24_seq_4(X1,X2,X3))|r1_xxreal_0(np__0,k24_seq_4(X1,X2,X3))|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~v7_ordinal1(k24_seq_4(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7277, negated_conjecture, (k24_seq_4(esk1_0,esk2_0,np__0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7403, negated_conjecture, (k3_seq_4(esk40_2(esk40_2(X1,k1_zfmisc_1(k1_numbers)),k1_zfmisc_1(k1_numbers)))=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k14_seq_4(esk1_0))|~r2_hidden(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7482, plain, (r1_xxreal_0(k24_seq_4(X1,X2,X3),k24_seq_4(X1,X2,X3))|v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|r2_hidden(X2,X3)|~v1_xreal_0(esk4_3(X1,X3,X2))|~v2_seq_4(X3,X1)|~m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7140, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7385, plain, (r1_xxreal_0(k4_xcmplx_0(X1),np__0)|~v1_xreal_0(k4_xcmplx_0(X1))|~v2_xxreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1069, plain, (v1_finset_1(esk30_0)|~v1_finset_1(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1237, plain, (v1_xboole_0(esk36_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6971, negated_conjecture, (m1_finseq_1(esk2_0,k2_numbers))).#trainneg
cnf(i_0_1066, plain, (v1_finset_1(esk38_1(X1))|v1_xboole_0(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_7589, plain, (v1_finset_1(esk31_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k5_numbers))).#trainneg
cnf(i_0_1077, plain, (v1_finset_1(k8_seq_4)|~v1_finset_1(k2_zfmisc_1(k2_numbers,k1_numbers)))).#trainneg
cnf(i_0_1067, plain, (v1_finset_1(esk32_1(X1))|v1_xboole_0(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_1070, plain, (v1_finset_1(esk28_1(X1))|v1_xboole_0(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_7205, negated_conjecture, (v3_membered(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k7_square_1(esk1_0)=np__0)).#trainneg
cnf(i_0_7369, negated_conjecture, (v2_xxreal_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|v1_xboole_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|r2_hidden(X2,X3)|~v1_xreal_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|~v2_seq_4(X3,X4)|~r2_hidden(k15_seq_4(X4,X2,k45_valued_1(esk2_0,X1)),X3)|~m1_subset_1(k45_valued_1(esk2_0,X1),k14_seq_4(X4))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X4)))|~m1_subset_1(X1,k14_seq_4(esk1_0))|~m1_subset_1(X2,k14_seq_4(X4))|~m1_subset_1(X4,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7067, negated_conjecture, (X1=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|X1!=np__0|~v1_xcmplx_0(X1)|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers))).#trainneg
cnf(i_0_7518, negated_conjecture, (k24_seq_4(esk1_0,X1,esk3_0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7164, negated_conjecture, (m1_subset_1(X1,k1_numbers)|k24_seq_4(esk1_0,X2,esk3_0)=k3_seq_4(k1_numbers)|~r2_hidden(X1,a_3_3_seq_4(esk1_0,X2,esk3_0))|~m1_subset_1(X2,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_1238, plain, (v1_xboole_0(esk34_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7035, negated_conjecture, (esk20_1(X1)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|esk20_1(X1)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1071, plain, (v1_finset_1(esk23_2(X1,X2))|~v1_finset_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_4559, negated_conjecture, (v3_ordinal1(esk2_0)|~v7_ordinal1(k1_zfmisc_1(k14_seq_4(esk1_0))))).#trainneg
cnf(i_0_7004, negated_conjecture, (v1_xboole_0(esk3_0)|v1_xxreal_0(X1)|~m1_subset_1(X1,esk3_0)|~m1_subset_1(k14_seq_4(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7107, negated_conjecture, (v1_xcmplx_0(X1)|v1_xboole_0(esk3_0)|~m1_subset_1(X1,esk3_0)|~m1_subset_1(k14_seq_4(esk1_0),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1072, plain, (v1_finset_1(esk21_2(X1,X2))|~v1_finset_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1075, plain, (v1_finset_1(esk18_1(k1_zfmisc_1(X1)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_7540, negated_conjecture, (k7_square_1(k24_seq_4(X1,X2,X3))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7081, negated_conjecture, (k24_seq_4(esk1_0,X1,esk3_0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7380, plain, (v1_xboole_0(k3_seq_4(X1))|r1_xxreal_0(np__0,k3_seq_4(X1))|~v3_membered(X1)|~m1_subset_1(k3_seq_4(X1),k5_numbers))).#trainneg
cnf(i_0_6679, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))=k24_seq_4(esk1_0,X1,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|k3_seq_4(np__0)=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_5531, negated_conjecture, (v2_membered(a_3_3_seq_4(esk1_0,X1,esk3_0))|k5_seq_4(k5_numbers)=k24_seq_4(esk1_0,X1,esk3_0)|~v2_membered(k1_numbers)|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_6426, plain, (v3_xxreal_0(esk40_2(X1,k5_numbers))|v2_xxreal_0(X2)|r1_xxreal_0(X2,esk40_2(X1,k5_numbers))|~v1_xreal_0(X2)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_5032, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk23_2(X1,X1)))).#trainneg
cnf(i_0_1239, plain, (v1_xboole_0(esk33_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7382, plain, (r1_xxreal_0(esk24_0,np__0)|~v1_xreal_0(esk24_0))).#trainneg
cnf(i_0_1079, plain, (r1_xxreal_0(X1,esk39_0)|r1_xxreal_0(esk39_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_6997, plain, (v2_membered(esk48_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_7468, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k5_numbers))).#trainneg
cnf(i_0_6984, negated_conjecture, (v3_ordinal1(esk3_0)|v1_xboole_0(esk3_0)|~v7_ordinal1(k14_seq_4(esk1_0))|~m1_subset_1(k1_zfmisc_1(k14_seq_4(esk1_0)),esk3_0))).#trainneg
cnf(i_0_1080, plain, (r1_xxreal_0(X1,esk24_0)|r1_xxreal_0(esk24_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1081, plain, (r1_xxreal_0(X1,esk22_0)|r1_xxreal_0(esk22_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_7755, plain, (r2_hidden(X1,X2)|v1_xboole_0(a_3_3_seq_4(X3,X1,X2))|~v2_seq_4(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X3)))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(X3))|~m2_subset_1(X3,k1_numbers,k5_numbers)|k24_seq_4(X3,X1,X2)!=np__0|~v1_xreal_0(esk4_3(X3,X2,X1))|~m1_subset_1(a_3_3_seq_4(X3,X1,X2),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(X3))|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_7658, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,k14_seq_4(X1)))=k24_seq_4(esk1_0,X3,esk3_0)|k24_seq_4(X1,X2,k14_seq_4(X1))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X3,esk3_0),k1_zfmisc_1(k1_numbers))|~v7_ordinal1(X1)|~m1_subset_1(X3,k14_seq_4(esk1_0))|~m1_subset_1(X2,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7643, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=k24_seq_4(esk1_0,X2,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7766, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=k24_seq_4(esk1_0,X2,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)=esk1_0|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_8063, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|r1_xxreal_0(esk4_3(esk1_0,esk3_0,esk2_0),np__0)|k7_square_1(esk1_0)=np__0|~v1_xreal_0(esk4_3(esk1_0,esk3_0,esk2_0)))).#trainneg
cnf(i_0_7717, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))=k24_seq_4(esk1_0,X1,esk3_0)|esk1_0=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7688, negated_conjecture, (v1_valued_0(esk2_0))).#trainneg
cnf(i_0_1084, plain, (r1_xxreal_0(X1,k2_xxreal_2(X2))|r1_xxreal_0(k2_xxreal_2(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(X2))).#trainneg
cnf(i_0_8212, plain, (r1_xxreal_0(esk39_0,esk35_0)|r1_xxreal_0(esk35_0,esk39_0))).#trainneg
cnf(i_0_7894, plain, (~v2_xxreal_0(esk39_0))).#trainneg
cnf(i_0_8208, plain, (~v2_xxreal_0(esk24_0)|~v1_xreal_0(esk24_0))).#trainneg
cnf(i_0_1088, plain, (v5_relat_1(esk47_1(k2_zfmisc_1(X1,X2)),X2)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1089, plain, (v5_relat_1(esk38_1(k2_zfmisc_1(X1,X2)),X2)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_8179, negated_conjecture, (v3_xxreal_2(a_3_3_seq_4(esk1_0,X1,esk3_0))|r2_hidden(X1,esk3_0)|k5_seq_4(k5_numbers)=k24_seq_4(esk1_0,X1,esk3_0)|~v1_xreal_0(esk4_3(esk1_0,esk3_0,X1))|~m1_subset_1(X1,k14_seq_4(esk1_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_8024, plain, (~v2_xxreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X1))|~v2_xxreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7824, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,X1,esk3_0))|r1_xxreal_0(esk4_3(esk1_0,esk3_0,X1),np__0)|r2_hidden(X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~v1_xreal_0(esk4_3(esk1_0,esk3_0,X1))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7645, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,k14_seq_4(X1)))=k24_seq_4(esk1_0,X3,esk3_0)|k24_seq_4(X1,X2,k14_seq_4(X1))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X3,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k14_seq_4(esk1_0))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1244, plain, (v1_xboole_0(esk26_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1090, plain, (v5_relat_1(esk32_1(k2_zfmisc_1(X1,X2)),X2)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_8215, plain, (r1_xxreal_0(esk39_0,esk24_0)|r1_xxreal_0(esk24_0,esk39_0))).#trainneg
cnf(i_0_7694, negated_conjecture, (~v1_funct_2(esk2_0,k5_numbers,k2_numbers)|~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_7748, negated_conjecture, (v7_ordinal1(esk3_0)|~v1_finset_1(esk3_0)|~m1_subset_1(k1_zfmisc_1(k14_seq_4(esk1_0)),k5_numbers))).#trainneg
cnf(i_0_7878, negated_conjecture, (r1_xxreal_0(esk3_0,esk3_0)|~m1_subset_1(k14_seq_4(esk1_0),k5_numbers)|~m1_subset_1(k1_zfmisc_1(esk3_0),k5_numbers))).#trainneg
cnf(i_0_1093, plain, (v5_relat_1(esk28_1(k2_zfmisc_1(X1,X2)),X2)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1096, plain, (v5_relat_1(esk20_1(k2_zfmisc_1(X1,X2)),X2))).#trainneg
cnf(i_0_7956, plain, (v1_xboole_0(k24_seq_4(X1,X2,X3))|r1_xxreal_0(np__0,k24_seq_4(X1,X2,X3))|k24_seq_4(X1,X2,X3)=np__0|~v7_ordinal1(k24_seq_4(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7920, plain, (v2_xxreal_0(k24_seq_4(X1,X2,X3))|v1_xboole_0(k24_seq_4(X1,X2,X3))|~v1_xreal_0(k24_seq_4(X1,X2,X3))|~m1_subset_1(k24_seq_4(X1,X2,X3),k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7657, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,np__0))=k24_seq_4(esk1_0,X3,esk3_0)|k24_seq_4(X1,X2,np__0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X3,esk3_0),k1_zfmisc_1(k1_numbers))|~v7_ordinal1(X1)|~m1_subset_1(X3,k14_seq_4(esk1_0))|~m1_subset_1(X2,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7671, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,esk2_0,k14_seq_4(esk1_0)))=k24_seq_4(esk1_0,X1,esk3_0)|k24_seq_4(esk1_0,esk2_0,k14_seq_4(esk1_0))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7952, negated_conjecture, (v1_xboole_0(k24_seq_4(esk1_0,X1,esk3_0))|r1_xxreal_0(np__0,k24_seq_4(esk1_0,X1,esk3_0))|k5_seq_4(k5_numbers)=k24_seq_4(esk1_0,X1,esk3_0)|~v7_ordinal1(k24_seq_4(esk1_0,X1,esk3_0))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7813, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,k14_seq_4(X1)))=k24_seq_4(esk1_0,X3,esk3_0)|k24_seq_4(X1,X2,k14_seq_4(X1))=esk1_0|m1_subset_1(a_3_3_seq_4(esk1_0,X3,esk3_0),k1_zfmisc_1(k1_numbers))|~v7_ordinal1(X1)|~m1_subset_1(X3,k14_seq_4(esk1_0))|~m1_subset_1(X2,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7094, plain, (v2_membered(esk31_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_2727, plain, (v1_relat_1(esk20_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1098, plain, (v5_relat_1(esk18_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_8217, plain, (r1_xxreal_0(esk39_0,esk22_0)|r1_xxreal_0(esk22_0,esk39_0))).#trainneg
cnf(i_0_7703, negated_conjecture, (v3_valued_0(esk2_0)|~v3_membered(k2_numbers))).#trainneg
cnf(i_0_7220, plain, (v2_membered(esk30_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_8259, negated_conjecture, (v7_ordinal1(esk3_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk3_0)),k5_numbers)|~m1_subset_1(k14_seq_4(esk1_0),k5_numbers))).#trainneg
cnf(i_0_7885, negated_conjecture, (v3_ordinal1(X1)|v1_xboole_0(esk3_0)|~v7_ordinal1(k14_seq_4(esk1_0))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),esk3_0))).#trainneg
cnf(i_0_1100, plain, (m1_subset_1(k22_seq_4(X1),k1_numbers)|~m1_subset_1(k18_rvsum_1(k12_rvsum_1(k13_seq_4(X1))),k1_numbers)|~m2_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_1104, plain, (v1_xboole_0(esk47_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_8026, plain, (r1_xxreal_0(X1,np__0)|~v2_xxreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7825, negated_conjecture, (v1_xreal_0(k24_seq_4(esk1_0,X1,esk3_0))|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|~m2_finseq_2(X1,k2_numbers,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7950, negated_conjecture, (v1_xboole_0(k24_seq_4(esk1_0,X1,esk3_0))|r1_xxreal_0(np__0,k24_seq_4(esk1_0,X1,esk3_0))|k5_seq_4(k1_numbers)=k24_seq_4(esk1_0,X1,esk3_0)|~v7_ordinal1(k24_seq_4(esk1_0,X1,esk3_0))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7951, negated_conjecture, (v1_xboole_0(k24_seq_4(esk1_0,X1,esk3_0))|r1_xxreal_0(np__0,k24_seq_4(esk1_0,X1,esk3_0))|k5_seq_4(np__0)=k24_seq_4(esk1_0,X1,esk3_0)|~v7_ordinal1(k24_seq_4(esk1_0,X1,esk3_0))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7993, negated_conjecture, (k3_seq_4(esk40_2(esk40_2(X1,k1_zfmisc_1(k1_numbers)),k1_zfmisc_1(k1_numbers)))=k24_seq_4(esk1_0,X2,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7644, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,np__0))=k24_seq_4(esk1_0,X3,esk3_0)|k24_seq_4(X1,X2,np__0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X3,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k14_seq_4(esk1_0))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8159, plain, (v2_xxreal_0(k3_seq_4(X1))|v1_xboole_0(k3_seq_4(X1))|~v1_xreal_0(k3_seq_4(X1))|~v3_membered(X1)|~m1_subset_1(k3_seq_4(X1),k5_numbers))).#trainneg
cnf(i_0_7812, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,np__0))=k24_seq_4(esk1_0,X3,esk3_0)|k24_seq_4(X1,X2,np__0)=esk1_0|m1_subset_1(a_3_3_seq_4(esk1_0,X3,esk3_0),k1_zfmisc_1(k1_numbers))|~v7_ordinal1(X1)|~m1_subset_1(X3,k14_seq_4(esk1_0))|~m1_subset_1(X2,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7981, negated_conjecture, (k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))=k24_seq_4(esk1_0,X1,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|k24_seq_4(esk1_0,esk2_0,np__0)=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7317, plain, (v2_membered(k8_seq_4)|~v3_membered(k2_zfmisc_1(k2_numbers,k1_numbers)))).#trainneg
cnf(i_0_2728, plain, (v1_funct_1(esk20_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7037, negated_conjecture, (esk16_1(X1)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|esk16_1(X1)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v3_xxreal_2(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_8228, plain, (r1_xxreal_0(esk39_0,np__1)|r1_xxreal_0(np__1,esk39_0))).#trainneg
cnf(i_0_7583, plain, (v1_finset_1(esk48_0)|~v7_ordinal1(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_7466, plain, (v7_ordinal1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_8073, negated_conjecture, (v2_xxreal_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|v1_xboole_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|r2_hidden(X2,X3)|v1_xboole_0(X3)|~v1_xreal_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|~v2_seq_4(X3,X4)|~m1_subset_1(k45_valued_1(esk2_0,X1),k14_seq_4(X4))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X4)))|~m1_subset_1(X1,k14_seq_4(esk1_0))|~m1_subset_1(X2,k14_seq_4(X4))|~m1_subset_1(X4,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(k15_seq_4(X4,X2,k45_valued_1(esk2_0,X1)),X3))).#trainneg
cnf(i_0_8167, plain, (v1_xboole_0(k24_seq_4(X1,X2,X3))|r1_xxreal_0(np__0,k24_seq_4(X1,X2,X3))|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~m1_subset_1(k24_seq_4(X1,X2,X3),k5_numbers)|~m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8068, negated_conjecture, (v2_xxreal_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|v1_xboole_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|r2_hidden(X2,X3)|~v1_xreal_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|~v2_seq_4(X3,X4)|~r2_hidden(k1_valued_1(X2,k45_valued_1(esk2_0,X1)),X3)|~m1_subset_1(k45_valued_1(esk2_0,X1),k14_seq_4(X4))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X4)))|~m1_subset_1(X1,k14_seq_4(esk1_0))|~m1_subset_1(X2,k14_seq_4(X4))|~m1_subset_1(X4,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7980, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))=k24_seq_4(esk1_0,X1,esk3_0)|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|k24_seq_4(esk1_0,esk2_0,np__0)=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7718, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,esk2_0,np__0))=k24_seq_4(esk1_0,X1,esk3_0)|k24_seq_4(esk1_0,esk2_0,np__0)=esk1_0|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7719, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,esk2_0,k14_seq_4(esk1_0)))=k24_seq_4(esk1_0,X1,esk3_0)|k24_seq_4(esk1_0,esk2_0,k14_seq_4(esk1_0))=esk1_0|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_8151, negated_conjecture, (k24_seq_4(esk1_0,esk40_2(X1,esk3_0),esk3_0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_7768, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,k14_seq_4(X1)))=k24_seq_4(esk1_0,X3,esk3_0)|k24_seq_4(X1,X2,k14_seq_4(X1))=esk1_0|m1_subset_1(a_3_3_seq_4(esk1_0,X3,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k14_seq_4(esk1_0))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7588, plain, (v1_finset_1(esk31_0)|~v7_ordinal1(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_7697, negated_conjecture, (v4_relat_1(esk2_0,k5_numbers))).#trainneg
cnf(i_0_7877, negated_conjecture, (~v1_xboole_0(k14_seq_4(esk3_0))|~m2_subset_1(esk3_0,k1_numbers,k5_numbers)|~m1_subset_1(k14_seq_4(esk1_0),k5_numbers)|~m1_subset_1(k1_zfmisc_1(esk3_0),k5_numbers))).#trainneg
cnf(i_0_1110, plain, (v1_xboole_0(esk23_2(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_7682, negated_conjecture, (v2_membered(esk2_0)|~v2_membered(k2_zfmisc_1(k5_numbers,k2_numbers)))).#trainneg
cnf(i_0_7683, negated_conjecture, (v1_xboole_0(esk2_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k2_numbers)))).#trainneg
cnf(i_0_7884, negated_conjecture, (v3_ordinal1(esk3_0)|v1_xboole_0(esk3_0)|~v7_ordinal1(k14_seq_4(esk1_0))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k14_seq_4(esk1_0))),esk3_0))).#trainneg
cnf(i_0_8265, plain, (v7_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1112, plain, (v1_xboole_0(esk20_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1114, plain, (v1_xboole_0(esk18_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_7670, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,esk2_0,np__0))=k24_seq_4(esk1_0,X1,esk3_0)|k24_seq_4(esk1_0,esk2_0,np__0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_6726, plain, (~v2_xxreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xboole_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6768, negated_conjecture, (k7_square_1(k24_seq_4(X1,X2,X3))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|k24_seq_4(X1,X2,X3)=esk1_0|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6216, plain, (v1_xboole_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xboole_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers))).#trainneg
cnf(i_0_7953, negated_conjecture, (v1_xboole_0(k24_seq_4(esk1_0,X1,esk3_0))|r1_xxreal_0(np__0,k24_seq_4(esk1_0,X1,esk3_0))|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|~v7_ordinal1(k24_seq_4(esk1_0,X1,esk3_0))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_6718, plain, (esk40_2(X1,k5_numbers)=k24_seq_4(X2,X3,X4)|m1_subset_1(a_3_3_seq_4(X2,X3,X4),k1_zfmisc_1(k1_numbers))|~r2_hidden(X1,k5_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X3,k14_seq_4(X2)))).#trainneg
cnf(i_0_8586, negated_conjecture, (k24_seq_4(esk1_0,X1,esk3_0)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_1118, plain, (v3_xxreal_0(esk22_0)|v2_xxreal_0(X1)|r1_xxreal_0(X1,esk22_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8270, plain, (r1_xxreal_0(esk24_0,esk35_0)|r1_xxreal_0(esk35_0,esk24_0))).#trainneg
cnf(i_0_8539, negated_conjecture, (r1_xxreal_0(esk3_0,esk3_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk3_0)),k5_numbers)|~m1_subset_1(k14_seq_4(esk1_0),k5_numbers))).#trainneg
cnf(i_0_8546, negated_conjecture, (v3_ordinal1(X1)|v1_xboole_0(esk3_0)|~v7_ordinal1(k14_seq_4(esk1_0))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1))),esk3_0))).#trainneg
cnf(i_0_1119, plain, (v3_xxreal_0(esk16_1(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,esk16_1(X1))|~v1_xreal_0(X2)|~v3_xxreal_2(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_1121, plain, (v3_xxreal_0(k6_square_1(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k6_square_1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8340, plain, (v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|r2_hidden(X2,X3)|k24_seq_4(X1,X2,X3)!=np__0|~v2_seq_4(X3,X1)|~m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1)))).#trainneg
cnf(i_0_8368, negated_conjecture, (k7_square_1(esk1_0)=np__0|v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|~v1_xreal_0(esk4_3(esk1_0,esk3_0,esk2_0)))).#trainneg
cnf(i_0_8693, plain, (v2_xxreal_0(k24_seq_4(X1,X2,X3))|v1_xboole_0(k24_seq_4(X1,X2,X3))|~v1_xreal_0(k24_seq_4(X1,X2,X3))|~m1_subset_1(a_3_3_seq_4(X1,X2,X3),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k24_seq_4(X1,X2,X3),k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8446, negated_conjecture, (r2_hidden(X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|v1_xboole_0(a_3_3_seq_4(esk1_0,X1,esk3_0))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(esk1_0))|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~v1_xreal_0(esk4_3(esk1_0,esk3_0,X1))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_8538, negated_conjecture, (~v1_xboole_0(k14_seq_4(esk3_0))|~m2_subset_1(esk3_0,k1_numbers,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk3_0)),k5_numbers)|~m1_subset_1(k14_seq_4(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1122, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8275, plain, (r1_xxreal_0(esk24_0,esk22_0)|r1_xxreal_0(esk22_0,esk24_0))).#trainneg
cnf(i_0_8673, plain, (v1_xreal_0(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_1123, plain, (v3_xxreal_0(k3_seq_4(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k3_seq_4(X1))|~v1_xreal_0(X2)|~v3_membered(X1))).#trainneg
cnf(i_0_1127, plain, (v1_funct_1(k3_relat_1(X1,X2))|~v1_relat_1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_8445, negated_conjecture, (k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|v1_xboole_0(a_3_3_seq_4(esk1_0,X1,esk3_0))|r2_hidden(X1,esk3_0)|~v2_xxreal_0(esk4_3(esk1_0,esk3_0,X1))|~v1_xreal_0(esk4_3(esk1_0,esk3_0,X1))|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_8512, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,X1,k14_seq_4(esk1_0)))=k24_seq_4(esk1_0,X2,esk3_0)|k24_seq_4(esk1_0,X1,k14_seq_4(esk1_0))=esk1_0|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_8516, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,k14_seq_4(X1)))=k24_seq_4(esk1_0,X3,esk3_0)|k24_seq_4(X1,X2,k14_seq_4(X1))=esk1_0|m1_subset_1(a_3_3_seq_4(esk1_0,X3,esk3_0),k1_zfmisc_1(k1_numbers))|~v7_ordinal1(X1)|~m1_subset_1(X3,k14_seq_4(esk1_0))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_8345, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,k14_seq_4(X1)))=k24_seq_4(esk1_0,X3,esk3_0)|k24_seq_4(X1,X2,k14_seq_4(X1))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X3,esk3_0),k1_zfmisc_1(k1_numbers))|~v7_ordinal1(X1)|~m1_subset_1(X3,k14_seq_4(esk1_0))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_8618, plain, (v2_xxreal_0(k24_seq_4(X1,X2,X3))|v1_xboole_0(k24_seq_4(X1,X2,X3))|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~m1_subset_1(k24_seq_4(X1,X2,X3),k5_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8341, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,X1,k14_seq_4(esk1_0)))=k24_seq_4(esk1_0,X2,esk3_0)|k24_seq_4(esk1_0,X1,k14_seq_4(esk1_0))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_8554, plain, (m1_subset_1(k22_seq_4(X1),k1_numbers)|~m2_finseq_1(X1,k2_numbers)|~m1_subset_1(k18_rvsum_1(k12_rvsum_1(k13_seq_4(X1))),k5_numbers))).#trainneg
cnf(i_0_8609, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,np__0))=k24_seq_4(esk1_0,X3,esk3_0)|k24_seq_4(X1,X2,np__0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X3,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k14_seq_4(esk1_0))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_8485, plain, (v1_valued_0(esk20_1(k2_zfmisc_1(X1,k2_numbers))))).#trainneg
cnf(i_0_8412, plain, (v2_xxreal_0(esk39_0)|r1_xxreal_0(esk39_0,esk35_0))).#trainneg
cnf(i_0_1128, plain, (X1=k3_relat_1(X2,X3)|~v1_xboole_0(X1)|~v1_relat_1(X2)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_8286, plain, (r1_xxreal_0(esk24_0,np__1)|r1_xxreal_0(np__1,esk24_0))).#trainneg
cnf(i_0_8678, plain, (r1_xxreal_0(X1,X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_1131, plain, (v1_funct_1(k3_relat_1(X1,X2))|~v1_relat_1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1132, plain, (X1=k3_relat_1(X2,X3)|~v1_xboole_0(X1)|~v1_relat_1(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_8619, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,X1,np__0))=k24_seq_4(esk1_0,X2,esk3_0)|k24_seq_4(esk1_0,X1,np__0)=esk1_0|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_8564, plain, (k2_zfmisc_1(X1,X2)=np__0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_8623, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,np__0))=k24_seq_4(esk1_0,X3,esk3_0)|k24_seq_4(X1,X2,np__0)=esk1_0|m1_subset_1(a_3_3_seq_4(esk1_0,X3,esk3_0),k1_zfmisc_1(k1_numbers))|~v7_ordinal1(X1)|~m1_subset_1(X3,k14_seq_4(esk1_0))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_8362, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=np__0|k24_seq_4(esk1_0,X1,esk3_0)=esk1_0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_8681, negated_conjecture, (v2_xxreal_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|v1_xboole_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|v1_xboole_0(X2)|r2_hidden(X3,X2)|~v1_xreal_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|~v2_seq_4(X2,X4)|~m1_subset_1(k15_seq_4(X4,k45_valued_1(esk2_0,X1),X3),X2)|~m1_subset_1(k45_valued_1(esk2_0,X1),k14_seq_4(X4))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X4)))|~m1_subset_1(X1,k14_seq_4(esk1_0))|~m1_subset_1(X3,k14_seq_4(X4))|~m1_subset_1(X4,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8031, plain, (v1_finset_1(esk30_0)|~v7_ordinal1(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_8486, plain, (v3_valued_0(esk20_1(k2_zfmisc_1(X1,k1_numbers))))).#trainneg
cnf(i_0_7704, negated_conjecture, (v5_relat_1(esk2_0,k2_numbers))).#trainneg
cnf(i_0_1133, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_9061, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k7_square_1(esk1_0)=np__0)).#trainneg
cnf(i_0_9079, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(X1,X2,X3))|r2_hidden(X2,X3)|k24_seq_4(X1,X2,X3)=esk1_0|k24_seq_4(X1,X2,X3)!=np__0|~v2_seq_4(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1)))).#trainneg
cnf(i_0_9055, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k3_seq_4(k1_numbers)=np__0)).#trainneg
cnf(i_0_9054, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k3_seq_4(k5_numbers)=np__0)).#trainneg
cnf(i_0_9052, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k5_seq_4(k1_numbers)=np__0)).#trainneg
cnf(i_0_9521, negated_conjecture, (k2_xcmplx_0(X1,X2)=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k2_xcmplx_0(X1,X2)!=np__0|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1140, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(X3)|r1_xxreal_0(X3,k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8642, plain, (~v1_finseq_1(esk39_0)|~v1_funct_1(esk39_0)|~v1_relat_1(esk39_0)|~m1_subset_1(k1_zfmisc_1(esk39_0),k5_numbers))).#trainneg
cnf(i_0_9063, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k24_seq_4(esk1_0,esk2_0,np__0)=np__0)).#trainneg
cnf(i_0_9051, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k5_seq_4(k5_numbers)=np__0)).#trainneg
cnf(i_0_9206, negated_conjecture, (k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|v1_xboole_0(a_3_3_seq_4(esk1_0,X1,esk3_0))|r2_hidden(X1,esk3_0)|r1_xxreal_0(esk4_3(esk1_0,esk3_0,X1),esk1_0)|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~v1_xreal_0(esk4_3(esk1_0,esk3_0,X1))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_9091, negated_conjecture, (k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|v1_xboole_0(a_3_3_seq_4(esk1_0,X1,esk3_0))|r2_hidden(X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(esk1_0))|~m1_subset_1(esk4_3(esk1_0,esk3_0,X1),k1_numbers))).#trainneg
cnf(i_0_9093, negated_conjecture, (k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|v1_xboole_0(a_3_3_seq_4(esk1_0,X1,esk3_0))|r2_hidden(X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_9075, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,X1,esk3_0))|r2_hidden(X1,esk3_0)|k5_seq_4(k1_numbers)=k24_seq_4(esk1_0,X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_9078, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,X1,esk3_0))|r2_hidden(X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(k1_numbers)|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_9076, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,X1,esk3_0))|r2_hidden(X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_1147, plain, (m1_subset_1(k13_seq_4(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_finseq_1(X1,k2_numbers))).#trainneg
cnf(i_0_8303, plain, (r1_xxreal_0(esk22_0,esk35_0)|r1_xxreal_0(esk35_0,esk22_0))).#trainneg
cnf(i_0_8643, plain, (~v1_finseq_1(esk24_0)|~v1_funct_1(esk24_0)|~v1_relat_1(esk24_0)|~m1_subset_1(k1_zfmisc_1(esk24_0),k5_numbers))).#trainneg
cnf(i_0_8545, negated_conjecture, (v3_ordinal1(esk3_0)|v1_xboole_0(esk3_0)|~v7_ordinal1(k14_seq_4(esk1_0))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k14_seq_4(esk1_0)))),esk3_0))).#trainneg
cnf(i_0_1150, plain, (v4_relat_1(esk47_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1151, plain, (v4_relat_1(esk38_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_9077, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,X1,esk3_0))|r2_hidden(X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(k5_numbers)|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_9074, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,X1,esk3_0))|r2_hidden(X1,esk3_0)|k5_seq_4(k5_numbers)=k24_seq_4(esk1_0,X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_9053, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k3_seq_4(np__0)=np__0)).#trainneg
cnf(i_0_9050, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k5_seq_4(np__0)=np__0)).#trainneg
cnf(i_0_9073, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,X1,esk3_0))|r2_hidden(X1,esk3_0)|k5_seq_4(np__0)=k24_seq_4(esk1_0,X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_9062, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k7_square_1(np__0)=np__0)).#trainneg
cnf(i_0_9218, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,X1,k14_seq_4(esk1_0)))=np__0|k24_seq_4(esk1_0,X1,k14_seq_4(esk1_0))=esk1_0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_1152, plain, (v4_relat_1(esk32_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_8956, negated_conjecture, (v3_ordinal1(X1)|v1_xboole_0(esk3_0)|~v7_ordinal1(k14_seq_4(esk1_0))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1)))),esk3_0))).#trainneg
cnf(i_0_8823, plain, (v1_xreal_0(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k5_numbers))).#trainneg
cnf(i_0_1155, plain, (v4_relat_1(esk28_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1158, plain, (v4_relat_1(esk20_1(k2_zfmisc_1(X1,X2)),X1))).#trainneg
cnf(i_0_9059, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k3_seq_4(esk46_0)=np__0)).#trainneg
cnf(i_0_9479, negated_conjecture, (v1_xboole_0(a_3_3_seq_4(esk1_0,X1,esk3_0))|r1_xxreal_0(esk4_3(esk1_0,esk3_0,X1),np__0)|r2_hidden(X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)=esk1_0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(esk4_3(esk1_0,esk3_0,X1))|~v3_xxreal_2(a_3_3_seq_4(esk1_0,X1,esk3_0))|~v3_membered(a_3_3_seq_4(esk1_0,X1,esk3_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_9236, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,k14_seq_4(X1)))=np__0|k24_seq_4(X1,X2,k14_seq_4(X1))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_9250, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,X1,k14_seq_4(esk1_0)))=np__0|k24_seq_4(esk1_0,X1,k14_seq_4(esk1_0))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_9262, negated_conjecture, (k3_seq_4(a_3_3_seq_4(X1,X2,np__0))=np__0|k24_seq_4(X1,X2,np__0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_8495, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,X1,np__0))=k24_seq_4(esk1_0,X2,esk3_0)|k24_seq_4(esk1_0,X1,np__0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k14_seq_4(esk1_0))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_6537, negated_conjecture, (k22_seq_4(k45_valued_1(X2,X1))=k5_seq_4(a_3_3_seq_4(esk1_0,esk2_0,esk3_0))|k22_seq_4(k45_valued_1(X2,X1))!=np__0|~m1_finseq_1(k45_valued_1(X1,X2),k2_numbers)|~m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0))|~m1_subset_1(X2,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_1160, plain, (v4_relat_1(esk18_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_8319, plain, (r1_xxreal_0(esk22_0,np__1)|r1_xxreal_0(np__1,esk22_0))).#trainneg
cnf(i_0_1162, plain, (k9_binop_2(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7261, negated_conjecture, (k3_seq_4(a_3_3_seq_4(esk1_0,X1,esk3_0))=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(k5_numbers)|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7128, negated_conjecture, (m1_subset_1(X1,k1_numbers)|k24_seq_4(esk1_0,X2,esk3_0)=k3_seq_4(np__0)|~r2_hidden(X1,a_3_3_seq_4(esk1_0,X2,esk3_0))|~m1_subset_1(X2,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_9693, negated_conjecture, (a_3_3_seq_4(esk1_0,esk2_0,esk3_0)=np__0|esk1_0=np__0)).#trainneg
cnf(i_0_9575, negated_conjecture, (a_3_3_seq_4(X1,X2,X3)=np__0|k24_seq_4(X1,X2,X3)=esk1_0|r2_hidden(X2,X3)|k24_seq_4(X1,X2,X3)!=np__0|~v2_seq_4(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9591, negated_conjecture, (a_3_3_seq_4(esk1_0,esk2_0,esk3_0)=np__0|k3_seq_4(k1_numbers)=np__0)).#trainneg
cnf(i_0_9607, negated_conjecture, (a_3_3_seq_4(esk1_0,esk2_0,esk3_0)=np__0|k3_seq_4(k5_numbers)=np__0)).#trainneg
cnf(i_0_9894, negated_conjecture, (a_3_3_seq_4(esk1_0,esk2_0,esk3_0)=np__0|k3_seq_4(np__0)=np__0)).#trainneg
cnf(i_0_9623, negated_conjecture, (a_3_3_seq_4(esk1_0,esk2_0,esk3_0)=np__0|k5_seq_4(k1_numbers)=np__0)).#trainneg
cnf(i_0_8828, plain, (r1_xxreal_0(X1,X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k5_numbers))).#trainneg
cnf(i_0_6989, negated_conjecture, (v7_ordinal1(X1)|v1_xboole_0(esk3_0)|~v1_finset_1(X1)|~v7_ordinal1(k14_seq_4(esk1_0))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_9725, negated_conjecture, (a_3_3_seq_4(esk1_0,esk2_0,esk3_0)=np__0|k5_seq_4(k5_numbers)=np__0)).#trainneg
cnf(i_0_9559, negated_conjecture, (a_3_3_seq_4(esk1_0,esk2_0,esk3_0)=np__0|k7_square_1(esk1_0)=np__0)).#trainneg
cnf(i_0_9897, negated_conjecture, (k3_seq_4(np__0)=np__0|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_9610, negated_conjecture, (k3_seq_4(k5_numbers)=np__0|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_9594, negated_conjecture, (k3_seq_4(k1_numbers)=np__0|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_9779, negated_conjecture, (a_3_3_seq_4(esk1_0,X1,esk3_0)=np__0|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(k1_numbers)|r2_hidden(X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7876, negated_conjecture, (r1_xxreal_0(esk3_0,X1)|r1_xxreal_0(X1,esk3_0)|~v1_xxreal_0(X1)|~m1_subset_1(k14_seq_4(esk1_0),k5_numbers)|~m1_subset_1(k1_zfmisc_1(esk3_0),k5_numbers))).#trainneg
cnf(i_0_9709, negated_conjecture, (a_3_3_seq_4(esk1_0,esk2_0,esk3_0)=np__0|k24_seq_4(esk1_0,esk2_0,np__0)=np__0)).#trainneg
cnf(i_0_9862, negated_conjecture, (a_3_3_seq_4(esk1_0,X1,esk3_0)=np__0|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(k5_numbers)|r2_hidden(X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_9795, negated_conjecture, (a_3_3_seq_4(esk1_0,X1,esk3_0)=np__0|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|r2_hidden(X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_10330, negated_conjecture, (X1=k24_seq_4(X2,X3,X4)|k24_seq_4(X2,X3,X4)=esk1_0|r2_hidden(X3,X4)|k5_seq_4(np__0)!=X1|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X1,k1_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|k24_seq_4(X2,X3,X4)!=np__0|~v2_seq_4(X4,X2)|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_10348, negated_conjecture, (k24_seq_4(X1,X2,X3)=k3_seq_4(np__0)|k24_seq_4(X1,X2,X3)=esk1_0|r2_hidden(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|k24_seq_4(X1,X2,X3)!=np__0|~v2_seq_4(X3,X1)|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_10328, negated_conjecture, (X1=k24_seq_4(X2,X3,X4)|k24_seq_4(X2,X3,X4)=esk1_0|r2_hidden(X3,X4)|k3_seq_4(np__0)!=X1|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X1,k1_numbers)|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|k24_seq_4(X2,X3,X4)!=np__0|~v2_seq_4(X4,X2)|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_10817, negated_conjecture, (k3_seq_4(np__0)=np__0|~m1_subset_1(esk3_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_9505, plain, (v1_valued_0(esk20_1(k2_zfmisc_1(X1,k1_numbers))))).#trainneg
cnf(i_0_8928, negated_conjecture, (v2_xxreal_0(esk1_0)|r1_xxreal_0(esk39_0,esk22_0))).#trainneg
cnf(i_0_8791, plain, (~v1_finseq_1(esk39_0)|~v1_funct_1(esk39_0)|~v1_relat_1(esk39_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk39_0)),k5_numbers))).#trainneg
cnf(i_0_8792, plain, (~v1_finseq_1(esk24_0)|~v1_funct_1(esk24_0)|~v1_relat_1(esk24_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk24_0)),k5_numbers))).#trainneg
cnf(i_0_10747, negated_conjecture, (v7_ordinal1(X1)|v1_xboole_0(esk3_0)|~v1_finset_1(X1)|~m1_subset_1(X1,esk3_0)|~m1_subset_1(k14_seq_4(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1188, plain, (v1_xboole_0(esk23_2(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_10793, negated_conjecture, (esk1_0=np__0|m2_xxreal_2(X1,esk3_0)|~v1_xxreal_0(X1)|~v2_membered(esk3_0))).#trainneg
cnf(i_0_10837, negated_conjecture, (k3_seq_4(k1_numbers)=np__0|~m1_subset_1(esk3_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_10815, negated_conjecture, (k3_seq_4(np__0)=np__0|m2_xxreal_2(X1,esk3_0)|~v1_xxreal_0(X1)|~v2_membered(esk3_0))).#trainneg
cnf(i_0_10827, negated_conjecture, (k3_seq_4(k5_numbers)=np__0|~m1_subset_1(esk3_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_10835, negated_conjecture, (k3_seq_4(k1_numbers)=np__0|m2_xxreal_2(X1,esk3_0)|~v1_xxreal_0(X1)|~v2_membered(esk3_0))).#trainneg
cnf(i_0_10379, negated_conjecture, (r2_hidden(X1,X2)|k24_seq_4(X3,X1,X2)=esk1_0|~v2_seq_4(X2,X3)|~r2_hidden(np__0,np__0)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X3)))|~m1_subset_1(X1,k14_seq_4(X3))|~m1_subset_1(X3,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|k24_seq_4(X3,X1,X2)!=np__0)).#trainneg
cnf(i_0_5758, plain, (v1_xreal_0(esk18_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_8929, plain, (r1_xxreal_0(esk39_0,esk22_0))).#trainneg
cnf(i_0_8677, plain, (~v1_xboole_0(k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_1190, plain, (v1_xboole_0(esk20_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8032, plain, (v1_finset_1(esk30_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k5_numbers))).#trainneg
cnf(i_0_10980, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|r1_xxreal_0(esk3_0,X1)|~m1_subset_1(k14_seq_4(esk1_0),k5_numbers)|~m1_subset_1(k1_zfmisc_1(esk3_0),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1192, plain, (v1_xboole_0(esk18_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1197, plain, (v1_relat_1(k30_valued_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_10825, negated_conjecture, (k3_seq_4(k5_numbers)=np__0|m2_xxreal_2(X1,esk3_0)|~v1_xxreal_0(X1)|~v2_membered(esk3_0))).#trainneg
cnf(i_0_10375, negated_conjecture, (m2_xxreal_2(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|k24_seq_4(X1,X3,X2)=esk1_0|~v1_xxreal_0(esk4_3(X1,X2,X3))|~v2_membered(X4)|~v2_seq_4(X2,X1)|~r2_hidden(esk10_2(X4,esk4_3(X1,X2,X3)),np__0)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|k24_seq_4(X1,X3,X2)!=np__0)).#trainneg
cnf(i_0_11231, negated_conjecture, (k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|r2_hidden(X1,esk3_0)|~m2_finseq_2(X1,k2_numbers,k14_seq_4(esk1_0))|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_10353, negated_conjecture, (k5_seq_4(np__0)=k24_seq_4(X1,X2,X3)|k24_seq_4(X1,X2,X3)=esk1_0|r2_hidden(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|k24_seq_4(X1,X2,X3)!=np__0|~v2_seq_4(X3,X1)|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_11260, negated_conjecture, (r2_hidden(X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|~r2_hidden(np__0,np__0)|~m1_subset_1(X1,k14_seq_4(esk1_0))|k24_seq_4(esk1_0,X1,esk3_0)!=np__0)).#trainneg
cnf(i_0_8827, plain, (~v1_xboole_0(k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k5_numbers))).#trainneg
cnf(i_0_8058, plain, (v1_finset_1(k8_seq_4)|~m1_subset_1(k2_zfmisc_1(k2_numbers,k1_numbers),k5_numbers))).#trainneg
cnf(i_0_7684, negated_conjecture, (v3_membered(esk2_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k2_numbers)))).#trainneg
cnf(i_0_7686, negated_conjecture, (v1_finset_1(esk2_0)|~v1_finset_1(k2_zfmisc_1(k5_numbers,k2_numbers)))).#trainneg
cnf(i_0_9977, negated_conjecture, (v3_ordinal1(X1)|v1_xboole_0(esk3_0)|~v7_ordinal1(k14_seq_4(esk1_0))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1))))),esk3_0))).#trainneg
cnf(i_0_7873, negated_conjecture, (v1_xreal_0(esk3_0)|~m1_subset_1(k14_seq_4(esk1_0),k5_numbers)|~m1_subset_1(k1_zfmisc_1(esk3_0),k5_numbers))).#trainneg
cnf(i_0_1207, plain, (v1_xxreal_0(k9_binop_2(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1208, plain, (v3_ordinal1(k9_binop_2(X1,X2))|~v3_ordinal1(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_10894, negated_conjecture, (r2_hidden(X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(k1_numbers)|~r2_hidden(np__0,np__0)|~m1_subset_1(X1,k14_seq_4(esk1_0))|k24_seq_4(esk1_0,X1,esk3_0)!=np__0)).#trainneg
cnf(i_0_11102, negated_conjecture, (X1=k24_seq_4(esk1_0,X2,esk3_0)|k24_seq_4(esk1_0,X2,esk3_0)=k3_seq_4(k5_numbers)|r2_hidden(X2,esk3_0)|k5_seq_4(np__0)!=X1|~m1_subset_1(X1,k1_numbers)|~m2_finseq_2(X2,k2_numbers,k14_seq_4(esk1_0))|k24_seq_4(esk1_0,X2,esk3_0)!=np__0|~m1_subset_1(X2,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_10847, negated_conjecture, (X1=k24_seq_4(esk1_0,X2,esk3_0)|k24_seq_4(esk1_0,X2,esk3_0)=k3_seq_4(k1_numbers)|r2_hidden(X2,esk3_0)|k5_seq_4(np__0)!=X1|~m1_subset_1(X1,k1_numbers)|~m2_finseq_2(X2,k2_numbers,k14_seq_4(esk1_0))|k24_seq_4(esk1_0,X2,esk3_0)!=np__0|~m1_subset_1(X2,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_11149, negated_conjecture, (r2_hidden(X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(k5_numbers)|~r2_hidden(np__0,np__0)|~m1_subset_1(X1,k14_seq_4(esk1_0))|k24_seq_4(esk1_0,X1,esk3_0)!=np__0)).#trainneg
cnf(i_0_10349, negated_conjecture, (k24_seq_4(X1,X2,X3)=esk1_0|r2_hidden(X2,X3)|~r2_hidden(X4,X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|k24_seq_4(X1,X2,X3)!=np__0|~v2_seq_4(X3,X1))).#trainneg
cnf(i_0_10870, negated_conjecture, (k5_seq_4(np__0)=k24_seq_4(esk1_0,X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(k1_numbers)|r2_hidden(X1,esk3_0)|~m2_finseq_2(X1,k2_numbers,k14_seq_4(esk1_0))|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_8057, plain, (v1_finset_1(k8_seq_4)|~v7_ordinal1(k2_zfmisc_1(k2_numbers,k1_numbers)))).#trainneg
cnf(i_0_5760, plain, (~v1_xboole_0(k14_seq_4(esk18_1(X1)))|~m2_subset_1(esk18_1(X1),k1_numbers,k5_numbers)|~v6_membered(X1))).#trainneg
cnf(i_0_7689, negated_conjecture, (v3_valued_0(esk2_0)|~v3_valued_0(k2_zfmisc_1(k5_numbers,k2_numbers)))).#trainneg
cnf(i_0_2724, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3488, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1210, plain, (v1_xreal_0(k9_binop_2(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1211, plain, (v1_xcmplx_0(k9_binop_2(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_11125, negated_conjecture, (k5_seq_4(np__0)=k24_seq_4(esk1_0,X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(k5_numbers)|r2_hidden(X1,esk3_0)|~m2_finseq_2(X1,k2_numbers,k14_seq_4(esk1_0))|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_10819, negated_conjecture, (k3_seq_4(np__0)=np__0|v1_xboole_0(esk3_0)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_11100, negated_conjecture, (X1=k24_seq_4(esk1_0,X2,esk3_0)|k24_seq_4(esk1_0,X2,esk3_0)=k3_seq_4(k5_numbers)|r2_hidden(X2,esk3_0)|k3_seq_4(np__0)!=X1|~m1_subset_1(X1,k1_numbers)|~m2_finseq_2(X2,k2_numbers,k14_seq_4(esk1_0))|k24_seq_4(esk1_0,X2,esk3_0)!=np__0|~m1_subset_1(X2,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_10115, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8192, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_3062, plain, (m1_subset_1(esk20_1(X1),k1_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_9150, negated_conjecture, (v1_xreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(esk2_0),k5_numbers))).#trainneg
cnf(i_0_5684, plain, (~v1_xboole_0(k14_seq_4(esk20_1(X1)))|~m2_subset_1(esk20_1(X1),k1_numbers,k5_numbers)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1212, plain, (v1_relat_1(k9_binop_2(X1,X2))|~v4_funct_1(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_8786, negated_conjecture, (v2_membered(esk2_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k2_numbers)))).#trainneg
cnf(i_0_3310, negated_conjecture, (v1_finset_1(esk3_0)|~v1_finseq_1(k14_seq_4(esk1_0))|~v1_funct_1(k14_seq_4(esk1_0))|~v1_relat_1(k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_8124, negated_conjecture, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|v1_xboole_0(esk3_0)|~v1_xxreal_0(X1)|~m1_subset_1(k14_seq_4(esk1_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,esk3_0))).#trainneg
cnf(i_0_1213, plain, (v1_funct_1(k9_binop_2(X1,X2))|~v4_funct_1(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1214, plain, (v1_finseq_1(k9_binop_2(X1,X2))|~v3_finseq_1(k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_10845, negated_conjecture, (X1=k24_seq_4(esk1_0,X2,esk3_0)|k24_seq_4(esk1_0,X2,esk3_0)=k3_seq_4(k1_numbers)|r2_hidden(X2,esk3_0)|k3_seq_4(np__0)!=X1|~m1_subset_1(X1,k1_numbers)|~m2_finseq_2(X2,k2_numbers,k14_seq_4(esk1_0))|k24_seq_4(esk1_0,X2,esk3_0)!=np__0|~m1_subset_1(X2,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_10367, negated_conjecture, (X1=k5_seq_4(X2)|k24_seq_4(X3,X4,X5)=esk1_0|r2_hidden(X4,X5)|X2!=np__0|X1!=k24_seq_4(X3,X4,X5)|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X3)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X3))|~m2_subset_1(X3,k1_numbers,k5_numbers)|k24_seq_4(X3,X4,X5)!=np__0|~v2_seq_4(X5,X3)|~m1_subset_1(X4,k14_seq_4(X3))|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_9878, negated_conjecture, (a_3_3_seq_4(esk1_0,X1,esk3_0)=np__0|k5_seq_4(k5_numbers)=k24_seq_4(esk1_0,X1,esk3_0)|r2_hidden(X1,esk3_0)|k24_seq_4(esk1_0,X1,esk3_0)!=np__0|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_10839, negated_conjecture, (k3_seq_4(k1_numbers)=np__0|v1_xboole_0(esk3_0)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_11232, negated_conjecture, (k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|r2_hidden(X1,esk3_0)|~r2_hidden(X2,esk3_0)|~m1_subset_1(X1,k14_seq_4(esk1_0))|k24_seq_4(esk1_0,X1,esk3_0)!=np__0)).#trainneg
cnf(i_0_10866, negated_conjecture, (k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(k1_numbers)|r2_hidden(X1,esk3_0)|~r2_hidden(X2,esk3_0)|~m1_subset_1(X1,k14_seq_4(esk1_0))|k24_seq_4(esk1_0,X1,esk3_0)!=np__0)).#trainneg
cnf(i_0_4336, plain, (v1_xreal_0(esk46_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3067, plain, (m1_subset_1(k1_card_1(X1),k1_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8944, plain, (v2_xxreal_0(esk24_0)|r1_xxreal_0(esk24_0,esk35_0)|~v1_xreal_0(esk24_0))).#trainneg
cnf(i_0_2966, plain, (v1_xxreal_0(esk48_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2974, plain, (v1_xxreal_0(esk31_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_11485, negated_conjecture, (v7_ordinal1(esk3_0)|v1_xboole_0(esk3_0)|~m1_subset_1(k14_seq_4(esk1_0),k5_numbers)|~m1_subset_1(esk3_0,esk3_0))).#trainneg
cnf(i_0_1221, plain, (v1_funct_1(k30_valued_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_5796, negated_conjecture, (k3_seq_4(np__0)=np__0)).#trainneg
cnf(i_0_5287, negated_conjecture, (k5_seq_4(np__0)=np__0)).#trainneg
cnf(i_0_11509, negated_conjecture, (k3_seq_4(k1_numbers)=np__0)).#trainneg
cnf(i_0_11514, negated_conjecture, (k3_seq_4(k5_numbers)=np__0)).#trainneg
cnf(i_0_1222, plain, (v1_valued_0(k54_valued_1(X1))|~v1_relat_1(k54_valued_1(X1))|~v1_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8671, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(k1_zfmisc_1(X2),k5_numbers))).#trainneg
cnf(i_0_8676, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_1223, plain, (v1_relat_1(k30_valued_1(k54_valued_1(X1)))|~v1_funct_1(k54_valued_1(X1))|~v1_relat_1(k54_valued_1(X1))|~v1_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1224, plain, (v1_funct_1(k30_valued_1(k54_valued_1(X1)))|~v1_funct_1(k54_valued_1(X1))|~v1_relat_1(k54_valued_1(X1))|~v1_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11713, plain, (v1_xboole_0(X1)|r1_xxreal_0(X2,np__0)|X3!=k3_seq_4(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v3_xxreal_2(X1)|~v3_membered(X1)|~r1_xxreal_0(k2_xcmplx_0(X2,X3),esk12_3(X1,X3,X2)))).#trainneg
cnf(i_0_10731, plain, (r1_xxreal_0(esk27_0,esk27_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk27_0)),k5_numbers))).#trainneg
cnf(i_0_8679, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_4094, plain, (v3_ordinal1(X1)|~v7_ordinal1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_1226, plain, (v1_relat_1(k30_valued_1(k30_valued_1(X1)))|~v1_funct_1(k30_valued_1(X1))|~v1_relat_1(k30_valued_1(X1))|~v3_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1227, plain, (v1_funct_1(k30_valued_1(k30_valued_1(X1)))|~v1_funct_1(k30_valued_1(X1))|~v1_relat_1(k30_valued_1(X1))|~v3_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4338, plain, (v1_xcmplx_0(esk46_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k2_numbers)))).#trainneg
cnf(i_0_11739, plain, (v1_xreal_0(esk20_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1232, plain, (v3_valued_0(X1)|~v3_valued_0(k30_valued_1(X1))|~v1_funct_1(k30_valued_1(X1))|~v1_relat_1(k30_valued_1(X1))|~v1_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6135, plain, (r1_xxreal_0(k5_numbers,esk22_0)|r1_xxreal_0(esk22_0,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8672, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(k1_zfmisc_1(X2),k5_numbers))).#trainneg
cnf(i_0_8821, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X2)),k5_numbers))).#trainneg
cnf(i_0_6677, negated_conjecture, (k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|esk4_3(esk1_0,esk3_0,esk2_0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_5095, plain, (v1_xxreal_0(esk46_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_11834, plain, (v1_xreal_0(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1249, plain, (m2_subset_1(k1_card_1(X1),k1_numbers,k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6145, plain, (r1_xxreal_0(k5_numbers,np__1)|r1_xxreal_0(np__1,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8826, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k5_numbers))).#trainneg
cnf(i_0_8829, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k5_numbers))).#trainneg
cnf(i_0_1251, plain, (v3_valued_0(esk47_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_1252, plain, (v3_valued_0(esk38_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_7980, negated_conjecture, (k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|k24_seq_4(esk1_0,esk2_0,np__0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_8247, plain, (v7_ordinal1(esk42_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk42_0)),k5_numbers))).#trainneg
cnf(i_0_1253, plain, (v3_valued_0(esk32_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_12101, plain, (r1_xxreal_0(esk46_0,esk46_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_11999, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1))),k5_numbers))).#trainneg
cnf(i_0_1256, plain, (v3_valued_0(esk28_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_1257, plain, (v3_valued_0(esk23_2(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_4077, negated_conjecture, (k24_seq_4(np__0,esk2_0,esk3_0)=np__0)).#trainneg
cnf(i_0_2, negated_conjecture, (v2_seq_4(esk3_0,np__0))).#trainneg
cnf(i_0_3360, negated_conjecture, (k22_seq_4(k45_valued_1(X1,X2))=k22_seq_4(k45_valued_1(X2,X1))|~m1_subset_1(X1,k14_seq_4(np__0))|~m1_subset_1(X2,k14_seq_4(np__0)))).#trainneg
cnf(i_0_1919, negated_conjecture, (m1_subset_1(esk2_0,k14_seq_4(np__0)))).#trainneg
cnf(i_0_5172, negated_conjecture, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,esk2_0)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(esk2_0,X4)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X4,k14_seq_4(np__0))|~m2_finseq_2(k45_valued_1(esk2_0,X4),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3256, negated_conjecture, (m1_subset_1(k45_valued_1(X1,esk2_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~m1_subset_1(X1,k14_seq_4(np__0)))).#trainneg
cnf(i_0_3366, negated_conjecture, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(X5,X4)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(np__0))|~m1_subset_1(X4,k14_seq_4(np__0))|~m2_finseq_2(k45_valued_1(X5,X4),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8248, plain, (v7_ordinal1(esk41_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk41_0)),k5_numbers))).#trainneg
cnf(i_0_4557, negated_conjecture, (v3_ordinal1(esk2_0)|~v7_ordinal1(k14_seq_4(np__0)))).#trainneg
cnf(i_0_1258, plain, (v3_valued_0(esk21_2(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_2349, negated_conjecture, (v1_xxreal_0(esk2_0)|~v2_membered(k14_seq_4(np__0)))).#trainneg
cnf(i_0_1259, plain, (v3_valued_0(esk20_1(k2_zfmisc_1(X1,X2)))|~v3_membered(X2))).#trainneg
cnf(i_0_1261, plain, (v3_valued_0(esk18_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v3_membered(X2))).#trainneg
cnf(i_0_4311, negated_conjecture, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,X5)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(X5,X4)),X2)|~m1_subset_1(k45_valued_1(X5,X4),k14_seq_4(X1))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(np__0))|~m1_subset_1(X4,k14_seq_4(np__0))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2334, plain, (r1_xxreal_0(X1,np__0)|X1!=np__0|~v1_xreal_0(X1)|~v3_xxreal_2(k5_numbers))).#trainneg
cnf(i_0_8179, negated_conjecture, (k5_seq_4(k5_numbers)=k24_seq_4(np__0,X1,esk3_0)|v3_xxreal_2(a_3_3_seq_4(np__0,X1,esk3_0))|r2_hidden(X1,esk3_0)|~v1_xreal_0(esk4_3(np__0,esk3_0,X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k14_seq_4(np__0)))).#trainneg
cnf(i_0_8253, plain, (v7_ordinal1(esk27_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk27_0)),k5_numbers))).#trainneg
cnf(i_0_1263, plain, (v1_xcmplx_0(k18_rvsum_1(X1))|~v1_valued_0(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_finseq_1(X1,k1_numbers))).#trainneg
cnf(i_0_2851, negated_conjecture, (v2_membered(esk3_0)|~v3_membered(k14_seq_4(np__0)))).#trainneg
cnf(i_0_4555, negated_conjecture, (v1_xxreal_0(esk2_0)|~v3_membered(k14_seq_4(np__0)))).#trainneg
cnf(i_0_8822, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X2)),k5_numbers))).#trainneg
cnf(i_0_4754, plain, (r1_xxreal_0(X1,k5_numbers)|r1_xxreal_0(k5_numbers,X1)|~v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1266, plain, (v1_valued_0(esk47_1(X1))|v1_xboole_0(X1)|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1268, plain, (v1_valued_0(esk38_1(X1))|v1_xboole_0(X1)|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6882, negated_conjecture, (k5_seq_4(k1_numbers)=k24_seq_4(np__0,X1,esk3_0)|v3_xxreal_2(a_3_3_seq_4(np__0,X1,esk3_0))|r2_hidden(X1,esk3_0)|~v1_xreal_0(esk4_3(np__0,esk3_0,X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k14_seq_4(np__0)))).#trainneg
cnf(i_0_1044, negated_conjecture, (v2_membered(esk3_0)|~v2_membered(k14_seq_4(np__0)))).#trainneg
cnf(i_0_3889, negated_conjecture, (esk4_3(X1,X2,X3)=k24_seq_4(np__0,X4,esk3_0)|r2_hidden(X3,X2)|m1_subset_1(a_3_3_seq_4(np__0,X4,esk3_0),k1_zfmisc_1(k1_numbers))|~v2_seq_4(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X4,k14_seq_4(np__0))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5941, negated_conjecture, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,esk40_2(X5,esk3_0))))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(esk40_2(X5,esk3_0),X4)),X2)|~r2_hidden(X5,esk3_0)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X4,k14_seq_4(np__0))|~m2_finseq_2(k45_valued_1(esk40_2(X5,esk3_0),X4),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3817, negated_conjecture, (v1_relat_1(k45_valued_1(X1,esk2_0))|~m1_subset_1(X1,k14_seq_4(np__0)))).#trainneg
cnf(i_0_580, negated_conjecture, (m1_subset_1(X1,k14_seq_4(np__0))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_7953, negated_conjecture, (k24_seq_4(np__0,X1,esk3_0)=k3_seq_4(np__0)|v1_xboole_0(k24_seq_4(np__0,X1,esk3_0))|r1_xxreal_0(np__0,k24_seq_4(np__0,X1,esk3_0))|~v7_ordinal1(k24_seq_4(np__0,X1,esk3_0))|~m1_subset_1(X1,k14_seq_4(np__0))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(np__0)))).#trainneg
cnf(i_0_9878, negated_conjecture, (k5_seq_4(k5_numbers)=k24_seq_4(np__0,X1,esk3_0)|a_3_3_seq_4(np__0,X1,esk3_0)=np__0|r2_hidden(X1,esk3_0)|k24_seq_4(np__0,X1,esk3_0)!=np__0|~m1_subset_1(X1,k14_seq_4(np__0)))).#trainneg
cnf(i_0_9987, plain, (v1_xreal_0(esk27_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk27_0)),k5_numbers))).#trainneg
cnf(i_0_11618, negated_conjecture, (v1_finset_1(esk2_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k2_numbers)))).#trainneg
cnf(i_0_1269, plain, (v1_valued_0(esk32_1(X1))|v1_xboole_0(X1)|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3308, negated_conjecture, (v1_finset_1(esk3_0)|~v7_ordinal1(k14_seq_4(np__0)))).#trainneg
cnf(i_0_2353, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v1_membered(k14_seq_4(np__0)))).#trainneg
cnf(i_0_4756, plain, (r1_xxreal_0(k5_numbers,X1)|r1_xxreal_0(X1,k5_numbers)|~v1_xboole_0(X1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4854, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1272, plain, (v1_valued_0(esk28_1(X1))|v1_xboole_0(X1)|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1273, plain, (v1_valued_0(esk23_2(X1,X2))|~v1_valued_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2335, plain, (r1_xxreal_0(X1,np__0)|X1!=np__0|~v1_xreal_0(X1)|~v3_xxreal_2(k1_numbers))).#trainneg
cnf(i_0_1761, negated_conjecture, (m1_subset_1(esk40_2(X1,esk3_0),k14_seq_4(np__0))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_3049, plain, (v1_xxreal_0(esk30_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_6376, plain, (m1_subset_1(esk18_1(X1),k1_numbers)|~v3_membered(X1))).#trainneg
cnf(i_0_3309, negated_conjecture, (v1_finset_1(esk3_0)|~m1_subset_1(k14_seq_4(np__0),k5_numbers))).#trainneg
cnf(i_0_1274, plain, (v1_valued_0(esk21_2(X1,X2))|~v1_valued_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2351, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(k14_seq_4(np__0)))).#trainneg
cnf(i_0_2352, negated_conjecture, (v1_xreal_0(esk2_0)|~v3_membered(k14_seq_4(np__0)))).#trainneg
cnf(i_0_4856, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X2)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1275, plain, (v1_valued_0(esk20_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1277, plain, (v1_valued_0(esk18_1(k1_zfmisc_1(X1)))|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7950, negated_conjecture, (k5_seq_4(k1_numbers)=k24_seq_4(np__0,X1,esk3_0)|v1_xboole_0(k24_seq_4(np__0,X1,esk3_0))|r1_xxreal_0(np__0,k24_seq_4(np__0,X1,esk3_0))|~v7_ordinal1(k24_seq_4(np__0,X1,esk3_0))|~m1_subset_1(X1,k14_seq_4(np__0))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(np__0)))).#trainneg
cnf(i_0_5371, negated_conjecture, (k5_seq_4(k1_numbers)=k24_seq_4(np__0,X1,esk3_0)|v2_membered(a_3_3_seq_4(np__0,X1,esk3_0))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k14_seq_4(np__0)))).#trainneg
cnf(i_0_7215, negated_conjecture, (k24_seq_4(esk1_0,X1,esk3_0)=k3_seq_4(np__0)|k7_square_1(esk1_0)=np__0|m1_subset_1(a_3_3_seq_4(esk1_0,X1,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k14_seq_4(esk1_0)))).#trainneg
cnf(i_0_7952, negated_conjecture, (k5_seq_4(k5_numbers)=k24_seq_4(np__0,X1,esk3_0)|v1_xboole_0(k24_seq_4(np__0,X1,esk3_0))|r1_xxreal_0(np__0,k24_seq_4(np__0,X1,esk3_0))|~v7_ordinal1(k24_seq_4(np__0,X1,esk3_0))|~m1_subset_1(X1,k14_seq_4(np__0))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(np__0)))).#trainneg
cnf(i_0_5531, negated_conjecture, (k5_seq_4(k5_numbers)=k24_seq_4(np__0,X1,esk3_0)|v2_membered(a_3_3_seq_4(np__0,X1,esk3_0))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k14_seq_4(np__0)))).#trainneg
cnf(i_0_5251, plain, (~v7_ordinal1(k4_xcmplx_0(X1))|~v2_xxreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_5678, plain, (v3_xxreal_0(X1)|r1_xxreal_0(np__0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5430, plain, (r1_xxreal_0(np__0,k3_seq_4(X1))|~v2_xxreal_0(k3_seq_4(X1))|~v1_xreal_0(k3_seq_4(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_6880, plain, (v3_xxreal_0(k24_seq_4(X1,X2,X3))|v1_xboole_0(k24_seq_4(X1,X2,X3))|r1_xxreal_0(np__0,k24_seq_4(X1,X2,X3))|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3129, plain, (v1_xxreal_0(k8_seq_4)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k2_numbers,k1_numbers))))).#trainneg
cnf(i_0_7449, plain, (m1_subset_1(esk49_1(X1),k1_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1282, plain, (v3_valued_0(esk47_1(X1))|v1_xboole_0(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2350, negated_conjecture, (v3_ordinal1(esk2_0)|~v3_ordinal1(k14_seq_4(np__0)))).#trainneg
cnf(i_0_1076, negated_conjecture, (v1_finset_1(esk3_0)|~v1_finset_1(k14_seq_4(np__0)))).#trainneg
cnf(i_0_1284, plain, (v3_valued_0(esk38_1(X1))|v1_xboole_0(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1285, plain, (v3_valued_0(esk32_1(X1))|v1_xboole_0(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6310, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X4,esk40_2(X5,np__0))))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(esk40_2(X5,np__0),X4)),X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X4,k14_seq_4(X6))|~m1_subset_1(X6,k5_numbers)|~m2_finseq_2(k45_valued_1(esk40_2(X5,np__0),X4),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r2_hidden(X5,np__0))).#trainneg
cnf(i_0_2085, plain, (r2_hidden(X1,a_3_9_seq_4(X2,X3,X4))|X1!=k22_seq_4(k45_valued_1(X5,X3))|~r2_hidden(X5,X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X2)))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X2))|~m2_finseq_2(X5,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,k14_seq_4(X2))|~m1_subset_1(X5,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2606, plain, (v1_xreal_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1978, plain, (r2_hidden(X3,a_3_9_seq_4(X1,X2,X5))|~r2_hidden(esk9_4(X3,X1,X2,X4),X5)|~m1_subset_1(X5,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r2_hidden(X3,a_3_3_seq_4(X1,X2,X4))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X1))))).#trainneg
cnf(i_0_3801, plain, (r1_xxreal_0(np__0,k22_seq_4(esk8_2(k14_seq_4(X1),X2)))|X2=np__0|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1))))).#trainneg
cnf(i_0_7556, plain, (m1_subset_1(esk37_1(X1),k1_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7877, negated_conjecture, (~v1_xboole_0(k14_seq_4(esk3_0))|~m1_subset_1(k14_seq_4(np__0),k5_numbers)|~m1_subset_1(k1_zfmisc_1(esk3_0),k5_numbers)|~m2_subset_1(esk3_0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1288, plain, (v3_valued_0(esk28_1(X1))|v1_xboole_0(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1060, negated_conjecture, (v3_membered(esk3_0)|~v3_membered(k14_seq_4(np__0)))).#trainneg
cnf(i_0_782, negated_conjecture, (v1_relat_1(esk3_0)|~v1_relat_1(k14_seq_4(np__0)))).#trainneg
cnf(i_0_4757, plain, (r1_xxreal_0(k5_numbers,X1)|r1_xxreal_0(X1,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1289, plain, (v3_valued_0(esk23_2(X1,X2))|~v3_valued_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1290, plain, (v3_valued_0(esk21_2(X1,X2))|~v3_valued_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_12777, plain, (v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_8538, negated_conjecture, (~v1_xboole_0(k14_seq_4(esk3_0))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk3_0)),k5_numbers)|~m1_subset_1(k14_seq_4(np__0),k5_numbers)|~m2_subset_1(esk3_0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1291, plain, (v3_valued_0(esk20_1(X1))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12218, plain, (r1_xxreal_0(esk42_0,esk42_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk42_0)),k5_numbers))).#trainneg
cnf(i_0_12234, plain, (~v1_finset_1(esk39_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_numbers))),k5_numbers))).#trainneg
cnf(i_0_12233, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1)))),k5_numbers))).#trainneg
cnf(i_0_12235, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1))),k5_numbers))).#trainneg
cnf(i_0_1293, plain, (v3_valued_0(esk18_1(k1_zfmisc_1(X1)))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1329, plain, (v1_valued_0(esk47_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_12424, negated_conjecture, (m2_finseq_1(k45_valued_1(X1,esk2_0),k2_numbers)|~m1_subset_1(X1,k14_seq_4(np__0)))).#trainneg
cnf(i_0_12571, negated_conjecture, (m1_subset_1(esk10_2(esk3_0,X1),k14_seq_4(np__0))|m2_xxreal_2(X1,esk3_0)|~v1_xxreal_0(X1)|~v2_membered(esk3_0))).#trainneg
cnf(i_0_12213, plain, (v1_xreal_0(esk42_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk42_0)),k5_numbers))).#trainneg
cnf(i_0_12729, plain, (v1_xreal_0(esk49_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1330, plain, (v1_valued_0(esk38_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_12483, plain, (r1_xxreal_0(esk41_0,esk41_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk41_0)),k5_numbers))).#trainneg
cnf(i_0_12328, negated_conjecture, (~v2_seq_4(esk3_0,np__0)|~r2_hidden(np__0,np__0)|~m1_subset_1(esk2_0,k14_seq_4(np__0)))).#trainneg
cnf(i_0_4857, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1331, plain, (v1_valued_0(esk32_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_1334, plain, (v1_valued_0(esk28_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_12478, plain, (v1_xreal_0(esk41_0)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(esk41_0)),k5_numbers))).#trainneg
cnf(i_0_12788, plain, (v1_xreal_0(esk37_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3822, negated_conjecture, (~v1_xboole_0(k2_zfmisc_1(k5_numbers,k2_numbers))|~r2_hidden(X1,k45_valued_1(X2,esk2_0))|~m1_subset_1(X2,k14_seq_4(np__0)))).#trainneg
cnf(i_0_1335, plain, (v1_valued_0(esk23_2(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_4958, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v3_membered(k14_seq_4(np__0)))).#trainneg
cnf(i_0_8674, plain, (r1_xxreal_0(X1,k5_numbers)|r1_xxreal_0(k5_numbers,X1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_1336, plain, (v1_valued_0(esk21_2(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_1337, plain, (v1_valued_0(esk20_1(k2_zfmisc_1(X1,X2)))|~v1_membered(X2))).#trainneg
cnf(i_0_8038, plain, (m1_subset_1(esk36_1(X1),k1_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1339, plain, (v1_valued_0(esk18_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_membered(X2))).#trainneg
cnf(i_0_8675, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k5_numbers))).#trainneg
cnf(i_0_1921, plain, (m2_finseq_1(X1,k2_numbers)|~m1_finseq_2(k14_seq_4(X2),k2_numbers)|~m1_subset_1(X1,k14_seq_4(X2))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_12316, negated_conjecture, (X1=k5_seq_4(X2)|X2!=np__0|X1!=k24_seq_4(np__0,esk2_0,esk3_0)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_numbers)|~m2_finseq_2(esk2_0,k2_numbers,k14_seq_4(np__0)))).#trainneg
cnf(i_0_4295, plain, (v3_ordinal1(esk48_0)|~v7_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_8114, plain, (m1_subset_1(esk34_1(X1),k1_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_12716, plain, (r1_xxreal_0(np__0,k24_seq_4(X1,X2,X3))|~v7_ordinal1(k24_seq_4(X1,X2,X3))|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12324, negated_conjecture, (m2_xxreal_2(esk4_3(np__0,esk3_0,esk2_0),X1)|~v1_xxreal_0(esk4_3(np__0,esk3_0,esk2_0))|~v2_membered(X1)|~v2_seq_4(esk3_0,np__0)|~r2_hidden(esk10_2(X1,esk4_3(np__0,esk3_0,esk2_0)),np__0)|~m1_subset_1(esk2_0,k14_seq_4(np__0)))).#trainneg
cnf(i_0_12577, negated_conjecture, (k24_seq_4(np__0,X1,esk3_0)=np__0|r1_xxreal_0(np__0,k24_seq_4(np__0,X1,esk3_0))|~m1_subset_1(X1,k14_seq_4(np__0))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(np__0))|~m1_subset_1(k24_seq_4(np__0,X1,esk3_0),k5_numbers))).#trainneg
cnf(i_0_12718, plain, (r1_xxreal_0(np__0,k24_seq_4(X1,X2,X3))|~m1_subset_1(k24_seq_4(X1,X2,X3),k5_numbers)|~v3_membered(a_3_3_seq_4(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12505, negated_conjecture, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(X5,X4)))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(X5,X4)),X2)|~m1_subset_1(k45_valued_1(X5,X4),k14_seq_4(X1))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(np__0))|~m1_subset_1(X4,k14_seq_4(np__0))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12641, negated_conjecture, (v3_xxreal_2(a_3_3_seq_4(np__0,X1,esk3_0))|r2_hidden(X1,esk3_0)|k5_seq_4(k1_numbers)=k24_seq_4(np__0,X1,esk3_0)|~v1_xreal_0(esk4_3(np__0,esk3_0,X1))|~m1_subset_1(X1,k14_seq_4(np__0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_12767, negated_conjecture, (r1_xxreal_0(esk4_3(X1,X2,X3),k22_seq_4(k45_valued_1(esk2_0,esk40_2(X4,np__0))))|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(esk40_2(X4,np__0),esk2_0)),X2)|~r2_hidden(X4,np__0)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(k45_valued_1(esk40_2(X4,np__0),esk2_0),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8197, plain, (m1_subset_1(esk33_1(X1),k1_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2323, negated_conjecture, (r1_xxreal_0(X1,esk3_0)|X1!=k3_seq_4(k1_zfmisc_1(k14_seq_4(np__0)))|~v1_xreal_0(X1)|~v3_xxreal_2(k1_zfmisc_1(k14_seq_4(np__0)))|~v3_membered(k1_zfmisc_1(k14_seq_4(np__0))))).#trainneg
cnf(i_0_12094, plain, (r1_xxreal_0(k5_numbers,esk46_0)|r1_xxreal_0(esk46_0,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8824, plain, (r1_xxreal_0(X1,k5_numbers)|r1_xxreal_0(k5_numbers,X1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k5_numbers))).#trainneg
cnf(i_0_1359, plain, (r1_xxreal_0(X1,esk10_2(X2,X3))|r1_xxreal_0(esk10_2(X2,X3),X1)|m2_xxreal_2(X3,X2)|~v1_xxreal_0(X1)|~v1_xxreal_0(X3)|~v2_membered(X2))).#trainneg
cnf(i_0_1360, plain, (v1_finseq_1(X1)|~v1_valued_0(k30_valued_1(X1))|~v1_finseq_1(k30_valued_1(X1))|~v1_funct_1(k30_valued_1(X1))|~v1_relat_1(k30_valued_1(X1))|~v1_valued_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12271, negated_conjecture, (k3_seq_4(X1)=k24_seq_4(np__0,X2,esk3_0)|v3_membered(a_3_3_seq_4(np__0,X2,esk3_0))|~v3_membered(X1)|~m2_finseq_2(X2,k2_numbers,k14_seq_4(np__0)))).#trainneg
cnf(i_0_12703, plain, (r1_xxreal_0(np__0,X1)|~v1_xboole_0(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_12557, negated_conjecture, (esk4_3(np__0,esk3_0,X1)=k24_seq_4(np__0,X2,esk3_0)|r2_hidden(X1,esk3_0)|m1_subset_1(a_3_3_seq_4(np__0,X2,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k14_seq_4(np__0))|~m1_subset_1(X1,k14_seq_4(np__0)))).#trainneg
cnf(i_0_12677, negated_conjecture, (v3_xxreal_2(a_3_3_seq_4(np__0,X1,esk3_0))|r2_hidden(X1,esk3_0)|k5_seq_4(k5_numbers)=k24_seq_4(np__0,X1,esk3_0)|~v1_xreal_0(esk4_3(np__0,esk3_0,X1))|~m1_subset_1(X1,k14_seq_4(np__0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_12430, negated_conjecture, (v2_xxreal_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|r2_hidden(X4,X3)|~v2_xxreal_0(esk4_3(X2,X3,X4))|~v1_xreal_0(k22_seq_4(k45_valued_1(X1,esk2_0)))|~v1_xreal_0(esk4_3(X2,X3,X4))|~v2_seq_4(X3,X2)|~r2_hidden(k15_seq_4(X2,X4,k45_valued_1(esk2_0,X1)),X3)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X2)))|~m1_subset_1(X1,k14_seq_4(np__0))|~m2_finseq_2(k45_valued_1(esk2_0,X1),k2_numbers,k14_seq_4(X2))|~m2_finseq_2(X4,k2_numbers,k14_seq_4(X2))|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12682, plain, (~v2_xxreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12472, negated_conjecture, (v2_xxreal_0(k22_seq_4(k45_valued_1(X1,X2)))|r2_hidden(X5,X4)|~v2_xxreal_0(esk4_3(X3,X4,X5))|~v1_xreal_0(k22_seq_4(k45_valued_1(X1,X2)))|~v1_xreal_0(esk4_3(X3,X4,X5))|~v2_seq_4(X4,X3)|~r2_hidden(k15_seq_4(X3,X5,k45_valued_1(X2,X1)),X4)|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X3)))|~m1_subset_1(X2,k14_seq_4(np__0))|~m1_subset_1(X1,k14_seq_4(np__0))|~m2_finseq_2(k45_valued_1(X2,X1),k2_numbers,k14_seq_4(X3))|~m2_finseq_2(X5,k2_numbers,k14_seq_4(X3))|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12503, negated_conjecture, (v2_xxreal_0(k22_seq_4(k45_valued_1(X1,X2)))|r2_hidden(X5,X4)|~v2_xxreal_0(esk4_3(X3,X4,X5))|~v1_xreal_0(k22_seq_4(k45_valued_1(X1,X2)))|~v1_xreal_0(esk4_3(X3,X4,X5))|~v2_seq_4(X4,X3)|~r2_hidden(k15_seq_4(X3,X5,k45_valued_1(X2,X1)),X4)|~m1_subset_1(k45_valued_1(X2,X1),k14_seq_4(X3))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X3)))|~m1_subset_1(X2,k14_seq_4(np__0))|~m1_subset_1(X1,k14_seq_4(np__0))|~m1_subset_1(X5,k14_seq_4(X3))|~m1_subset_1(X3,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4423, plain, (v3_ordinal1(esk31_0)|~v7_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_8456, plain, (m1_subset_1(esk26_1(X1),k1_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1365, plain, (v5_relat_1(esk48_0,X1)|~v5_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers),X1))).#trainneg
cnf(i_0_12095, plain, (r1_xxreal_0(esk39_0,esk46_0)|r1_xxreal_0(esk46_0,esk39_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8825, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(X1)),k5_numbers))).#trainneg
cnf(i_0_1491, plain, (v3_ordinal1(k24_seq_4(X1,X2,X3))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1366, plain, (v5_relat_1(esk47_1(X1),X2)|v1_xboole_0(X1)|~v5_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1367, plain, (v5_relat_1(esk46_0,X1)|~v5_relat_1(k1_numbers,X1)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_12502, negated_conjecture, (v3_xxreal_0(esk4_3(X1,X2,X3))|v2_xxreal_0(k22_seq_4(k45_valued_1(X4,X5)))|v1_xboole_0(esk4_3(X1,X2,X3))|r2_hidden(X3,X2)|~v1_xreal_0(k22_seq_4(k45_valued_1(X4,X5)))|~v1_xreal_0(esk4_3(X1,X2,X3))|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,X3,k45_valued_1(X5,X4)),X2)|~m1_subset_1(k45_valued_1(X5,X4),k14_seq_4(X1))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X5,k14_seq_4(np__0))|~m1_subset_1(X4,k14_seq_4(np__0))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12370, negated_conjecture, (~v1_xboole_0(a_3_3_seq_4(np__0,X1,X2))|~r2_hidden(esk2_0,X2)|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(np__0)))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(np__0)))).#trainneg
cnf(i_0_12266, negated_conjecture, (k24_seq_4(np__0,X1,esk3_0)=np__0|v3_xxreal_0(X2)|v2_xxreal_0(k24_seq_4(np__0,X1,esk3_0))|r1_xxreal_0(k24_seq_4(np__0,X1,esk3_0),X2)|~v1_xboole_0(X2)|~m2_finseq_2(X1,k2_numbers,k14_seq_4(np__0)))).#trainneg
cnf(i_0_2004, plain, (k22_seq_4(k45_valued_1(esk9_4(X3,X1,X2,X4),X2))=X3|~r2_hidden(X3,a_3_3_seq_4(X1,X2,X4))|~m1_subset_1(X4,k1_zfmisc_1(k14_seq_4(X1)))|~m2_finseq_2(X2,k2_numbers,k14_seq_4(X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(esk9_4(X3,X1,X2,X4),k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1968, plain, (r1_xxreal_0(esk4_3(X1,X2,X3),X4)|r2_hidden(X3,X2)|~v2_seq_4(X2,X1)|~r2_hidden(k15_seq_4(X1,k19_seq_4(X5,X6,esk9_4(X4,X5,X6,X7)),X3),X2)|~r2_hidden(X4,a_3_3_seq_4(X5,X6,X7))|~m1_subset_1(X2,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X7,k1_zfmisc_1(k14_seq_4(X5)))|~m2_finseq_2(k19_seq_4(X5,X6,esk9_4(X4,X5,X6,X7)),k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X3,k2_numbers,k14_seq_4(X1))|~m2_finseq_2(X6,k2_numbers,k14_seq_4(X5))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(X5,k1_numbers,k5_numbers)|~m1_subset_1(k19_seq_4(X5,X6,esk9_4(X4,X5,X6,X7)),k14_seq_4(X1))|~m1_subset_1(X3,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6212, plain, (v1_xboole_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v7_ordinal1(X1)|~v1_xboole_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12637, negated_conjecture, (k5_seq_4(k1_numbers)=k24_seq_4(np__0,X1,esk3_0)|v1_xboole_0(k24_seq_4(np__0,X1,esk3_0))|r1_xxreal_0(np__0,k24_seq_4(np__0,X1,esk3_0))|~m1_subset_1(X1,k14_seq_4(np__0))|~m2_finseq_2(X1,k2_numbers,k14_seq_4(np__0))|~m1_subset_1(k24_seq_4(np__0,X1,esk3_0),k5_numbers))).#trainneg
cnf(i_0_1368, plain, (v5_relat_1(esk38_1(X1),X2)|v1_xboole_0(X1)|~v5_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12096, plain, (r1_xxreal_0(esk24_0,esk46_0)|r1_xxreal_0(esk46_0,esk24_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1494, plain, (v1_xcmplx_0(k24_seq_4(X1,X2,X3))|~v1_membered(k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k14_seq_4(X1)))|~m1_subset_1(X2,k14_seq_4(X1))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1369, plain, (v5_relat_1(esk32_1(X1),X2)|v1_xboole_0(X1)|~v5_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1370, plain, (v5_relat_1(esk31_0,X1)|~v5_relat_1(k2_zfmisc_1(k5_numbers,k5_numbers),X1))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 295
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 569
# Removed in clause preprocessing      : 54
# Initial clauses in saturation        : 515
# Processed clauses                    : 3362
# ...of these trivial                  : 136
# ...subsumed                          : 1078
# ...remaining for further processing  : 2148
# Other redundant clauses eliminated   : 12
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 62
# Backward-rewritten                   : 480
# Generated clauses                    : 12813
# ...of the previous two non-trivial   : 10954
# Contextual simplify-reflections      : 245
# Paramodulations                      : 12768
# Factorizations                       : 0
# Equation resolutions                 : 43
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1604
#    Positive orientable unit clauses  : 275
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 56
#    Non-unit-clauses                  : 1273
# Current number of unprocessed clauses: 7315
# ...number of literals in the above   : 42428
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2693
# Proof object given clauses           : 172
# Proof search given clauses           : 2148
# Clause-clause subsumption calls (NU) : 1244539
# Rec. Clause-clause subsumption calls : 148961
# Non-unit clause-clause subsumptions  : 1024
# Unit Clause-clause subsumption calls : 39447
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 94
# BW rewrite match successes           : 48
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 315613

# -------------------------------------------------
# User time                : 22.065 s
# System time              : 0.640 s
# Total time               : 22.706 s
# Maximum resident set size: 384060 pages
22.06user 0.72system 0:23.27elapsed 97%CPU (0avgtext+0avgdata 384060maxresident)k
0inputs+856outputs (0major+246860minor)pagefaults 0swaps
