2020-05-31 09:53:18.736975: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_55
2020-05-31 09:53:18.832230: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-05-31 09:53:19.023368: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-05-31 09:53:19.707203: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-05-31 09:53:22.582075: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 3845083 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_55' loaded (query=512; context=768; weigths=real; len_mult=0.000000).
# Preprocessing time       : 5.105 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:((v1_pre_topc(X3)&m1_pre_topc(X3,X1))=>(X3=k1_pre_topc(X1,X2)<=>k2_struct_0(X3)=X2)))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', d5_pre_topc)).
fof(dt_k1_pre_topc, axiom, ![X1, X2]:((l1_pre_topc(X1)&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>(v1_pre_topc(k1_pre_topc(X1,X2))&m1_pre_topc(k1_pre_topc(X1,X2),X1))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', dt_k1_pre_topc)).
fof(d1_t_0topsp, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(l1_pre_topc(X2)=>(r1_t_0topsp(X1,X2)<=>?[X3]:(((v1_funct_1(X3)&v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))&v3_tops_2(X3,X1,X2))))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', d1_t_0topsp)).
fof(t57_tops_2, axiom, ![X1]:((~(v2_struct_0(X1))&l1_pre_topc(X1))=>![X2]:((~(v2_struct_0(X2))&l1_pre_topc(X2))=>![X3]:((~(v2_struct_0(X3))&l1_pre_topc(X3))=>![X4]:(((v1_funct_1(X4)&v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2)))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))=>![X5]:(((v1_funct_1(X5)&v1_funct_2(X5,u1_struct_0(X2),u1_struct_0(X3)))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)))))=>((v3_tops_2(X4,X1,X2)&v3_tops_2(X5,X2,X3))=>v3_tops_2(k1_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X3),X4,X5),X1,X3))))))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', t57_tops_2)).
fof(d3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k2_struct_0(X1)=u1_struct_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', d3_struct_0)).
fof(t6_tmap_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>(v1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))&m1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)),X1)))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', t6_tmap_1)).
fof(dt_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:((((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_funct_1(X6))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>(v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))&m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4))))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', dt_k1_partfun1)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', dt_l1_pre_topc)).
fof(dt_m1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_pre_topc(X2,X1)=>l1_pre_topc(X2))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', dt_m1_pre_topc)).
fof(dt_l1_rltopsp1, axiom, ![X1]:(l1_rltopsp1(X1)=>(l1_rlvect_1(X1)&l1_pre_topc(X1))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', dt_l1_rltopsp1)).
fof(dt_k15_euclid, axiom, ![X1]:(v7_ordinal1(X1)=>(v5_rltopsp1(k15_euclid(X1))&l1_rltopsp1(k15_euclid(X1)))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', dt_k15_euclid)).
fof(cc17_membered, axiom, ![X1]:(v6_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v7_ordinal1(X2))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', cc17_membered)).
fof(fc6_membered, axiom, v6_membered(k4_ordinal1), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', fc6_membered)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', redefinition_k5_numbers)).
fof(redefinition_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:((((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_funct_1(X6))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>k1_partfun1(X1,X2,X3,X4,X5,X6)=k3_relat_1(X5,X6)), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', redefinition_k1_partfun1)).
fof(cc1_toprealb, axiom, ![X1]:(m1_pre_topc(X1,k15_euclid(np__2))=>(v1_toprealb(X1)=>((~(v2_struct_0(X1))&v1_compts_1(X1))&v1_borsuk_2(X1)))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', cc1_toprealb)).
fof(spc2_numerals, axiom, (((v2_xxreal_0(np__2)&m2_subset_1(np__2,k1_numbers,k5_numbers))&m1_subset_1(np__2,k5_numbers))&m1_subset_1(np__2,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', spc2_numerals)).
fof(d5_toprealb, axiom, ![X1]:(m1_pre_topc(X1,k15_euclid(np__2))=>(v1_toprealb(X1)<=>(v1_topreal2(u1_struct_0(X1))&m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', d5_toprealb)).
fof(fc8_funct_2, axiom, ![X1, X2, X3, X4, X5]:(((((((~(v1_xboole_0(X2))&v1_funct_1(X4))&v1_funct_2(X4,X1,X2))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_funct_1(X5))&v1_funct_2(X5,X2,X3))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))=>(v1_funct_1(k3_relat_1(X4,X5))&v1_funct_2(k3_relat_1(X4,X5),X1,X3))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', fc8_funct_2)).
fof(d1_topreal2, axiom, ![X1]:(m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))=>(v1_topreal2(X1)<=>?[X2]:(((v1_funct_1(X2)&v1_funct_2(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))))))&v3_tops_2(X2,k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(k15_euclid(np__2),X1))))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', d1_topreal2)).
fof(free_g1_pre_topc, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>![X3, X4]:(g1_pre_topc(X1,X2)=g1_pre_topc(X3,X4)=>(X1=X3&X2=X4))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', free_g1_pre_topc)).
fof(abstractness_v1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>(v1_pre_topc(X1)=>X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', abstractness_v1_pre_topc)).
fof(t11_toprealb, conjecture, ![X1]:((v1_toprealb(X1)&m1_pre_topc(X1,k15_euclid(np__2)))=>![X2]:((v1_toprealb(X2)&m1_pre_topc(X2,k15_euclid(np__2)))=>r2_borsuk_3(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', t11_toprealb)).
fof(cc3_funct_2, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v1_funct_2(X3,X1,X2)=>v1_partfun1(X3,X1)))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', cc3_funct_2)).
fof(dt_k2_tops_2, axiom, ![X1, X2, X3]:(((v1_funct_1(X3)&v1_funct_2(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))=>((v1_funct_1(k2_tops_2(X1,X2,X3))&v1_funct_2(k2_tops_2(X1,X2,X3),X2,X1))&m1_subset_1(k2_tops_2(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1))))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', dt_k2_tops_2)).
fof(dt_u1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', dt_u1_pre_topc)).
fof(dt_g1_pre_topc, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>(v1_pre_topc(g1_pre_topc(X1,X2))&l1_pre_topc(g1_pre_topc(X1,X2)))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', dt_g1_pre_topc)).
fof(cc1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v1_partfun1(X3,X1)=>v1_funct_2(X3,X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', cc1_funct_2)).
fof(t56_tops_2, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:((~(v2_struct_0(X2))&l1_pre_topc(X2))=>![X3]:(((v1_funct_1(X3)&v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))=>(v3_tops_2(X3,X1,X2)=>v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1))))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', t56_tops_2)).
fof(rc1_toprealb, axiom, ?[X1]:(((((m1_pre_topc(X1,k15_euclid(np__2))&v1_pre_topc(X1))&v2_pre_topc(X1))&v1_monoid_0(X1))&v2_monoid_0(X1))&v1_toprealb(X1)), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', rc1_toprealb)).
fof(cc1_pre_topc, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(m1_pre_topc(X2,X1)=>v2_pre_topc(X2))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', cc1_pre_topc)).
fof(redefinition_r2_borsuk_3, axiom, ![X1, X2]:((((~(v2_struct_0(X1))&l1_pre_topc(X1))&~(v2_struct_0(X2)))&l1_pre_topc(X2))=>(r2_borsuk_3(X1,X2)<=>r1_t_0topsp(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', redefinition_r2_borsuk_3)).
fof(cc1_topreal2, axiom, ![X1]:(m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))=>(v1_topreal2(X1)=>(~(v1_xboole_0(X1))&v2_compts_1(X1,k15_euclid(np__2))))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', cc1_topreal2)).
fof(fc1_struct_0, axiom, ![X1]:((v2_struct_0(X1)&l1_struct_0(X1))=>v1_xboole_0(u1_struct_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', fc1_struct_0)).
fof(cc8_funct_2, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((v1_funct_1(X3)&v1_funct_2(X3,X1,X2))=>((v1_funct_1(X3)&~(v1_xboole_0(X3)))&v1_funct_2(X3,X1,X2))))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', cc8_funct_2)).
fof(cc1_funct_1, axiom, ![X1]:(v1_xboole_0(X1)=>v1_funct_1(X1)), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', cc1_funct_1)).
fof(dt_k1_topreal1, axiom, m1_subset_1(k1_topreal1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', dt_k1_topreal1)).
fof(fc6_euclid, axiom, ![X1]:(v7_ordinal1(X1)=>(((((((((v2_pre_topc(k15_euclid(X1))&v13_algstr_0(k15_euclid(X1)))&v2_rlvect_1(k15_euclid(X1)))&v3_rlvect_1(k15_euclid(X1)))&v4_rlvect_1(k15_euclid(X1)))&v5_rlvect_1(k15_euclid(X1)))&v6_rlvect_1(k15_euclid(X1)))&v7_rlvect_1(k15_euclid(X1)))&v8_rlvect_1(k15_euclid(X1)))&v5_rltopsp1(k15_euclid(X1)))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', fc6_euclid)).
fof(fc3_topreal1, axiom, ~(v1_xboole_0(k1_topreal1)), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', fc3_topreal1)).
fof(t15_topreala, axiom, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:(((~(v2_struct_0(X2))&v2_pre_topc(X2))&l1_pre_topc(X2))=>(r1_t_0topsp(X1,X2)<=>r1_t_0topsp(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', t15_topreala)).
fof(cc3_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_xboole_0(X3))), file('/scratch/mptp/allpred/preds__0.005/t11_toprealb', cc3_relset_1)).
fof(c_0_41, plain, ![X16, X17, X18]:((X18!=k1_pre_topc(X16,X17)|k2_struct_0(X18)=X17|(~v1_pre_topc(X18)|~m1_pre_topc(X18,X16))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~l1_pre_topc(X16))&(k2_struct_0(X18)!=X17|X18=k1_pre_topc(X16,X17)|(~v1_pre_topc(X18)|~m1_pre_topc(X18,X16))|~m1_subset_1(X17,k1_zfmisc_1(u1_struct_0(X16)))|~l1_pre_topc(X16))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_pre_topc])])])])).
cnf(c_0_42, plain, (k2_struct_0(X1)=X3|X1!=k1_pre_topc(X2,X3)|~v1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_41])).
fof(c_0_43, plain, ![X61, X62]:((v1_pre_topc(k1_pre_topc(X61,X62))|(~l1_pre_topc(X61)|~m1_subset_1(X62,k1_zfmisc_1(u1_struct_0(X61)))))&(m1_pre_topc(k1_pre_topc(X61,X62),X61)|(~l1_pre_topc(X61)|~m1_subset_1(X62,k1_zfmisc_1(u1_struct_0(X61)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_pre_topc])])])).
fof(c_0_44, plain, ![X9, X10, X12]:(((((v1_funct_1(esk3_2(X9,X10))|~r1_t_0topsp(X9,X10)|~l1_pre_topc(X10)|~l1_pre_topc(X9))&(v1_funct_2(esk3_2(X9,X10),u1_struct_0(X9),u1_struct_0(X10))|~r1_t_0topsp(X9,X10)|~l1_pre_topc(X10)|~l1_pre_topc(X9)))&(m1_subset_1(esk3_2(X9,X10),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X10))))|~r1_t_0topsp(X9,X10)|~l1_pre_topc(X10)|~l1_pre_topc(X9)))&(v3_tops_2(esk3_2(X9,X10),X9,X10)|~r1_t_0topsp(X9,X10)|~l1_pre_topc(X10)|~l1_pre_topc(X9)))&(~v1_funct_1(X12)|~v1_funct_2(X12,u1_struct_0(X9),u1_struct_0(X10))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X10))))|~v3_tops_2(X12,X9,X10)|r1_t_0topsp(X9,X10)|~l1_pre_topc(X10)|~l1_pre_topc(X9))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_t_0topsp])])])])])).
fof(c_0_45, plain, ![X20, X21, X22, X23, X24]:(v2_struct_0(X20)|~l1_pre_topc(X20)|(v2_struct_0(X21)|~l1_pre_topc(X21)|(v2_struct_0(X22)|~l1_pre_topc(X22)|(~v1_funct_1(X23)|~v1_funct_2(X23,u1_struct_0(X20),u1_struct_0(X21))|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X20),u1_struct_0(X21))))|(~v1_funct_1(X24)|~v1_funct_2(X24,u1_struct_0(X21),u1_struct_0(X22))|~m1_subset_1(X24,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X21),u1_struct_0(X22))))|(~v3_tops_2(X23,X20,X21)|~v3_tops_2(X24,X21,X22)|v3_tops_2(k1_partfun1(u1_struct_0(X20),u1_struct_0(X21),u1_struct_0(X21),u1_struct_0(X22),X23,X24),X20,X22))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t57_tops_2])])])])).
fof(c_0_46, plain, ![X31]:(~l1_struct_0(X31)|k2_struct_0(X31)=u1_struct_0(X31)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_struct_0])])).
cnf(c_0_47, plain, (k2_struct_0(X1)=X3|X1!=k1_pre_topc(X2,X3)|~l1_pre_topc(X2)|~v1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(evalgc,[status(thm)],[c_0_42])).
cnf(c_0_48, plain, (m1_pre_topc(k1_pre_topc(X1,X2),X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49, plain, (v1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_43])).
fof(c_0_50, plain, ![X59, X60]:((v1_pre_topc(g1_pre_topc(u1_struct_0(X60),u1_pre_topc(X60)))|~m1_pre_topc(X60,X59)|~l1_pre_topc(X59))&(m1_pre_topc(g1_pre_topc(u1_struct_0(X60),u1_pre_topc(X60)),X59)|~m1_pre_topc(X60,X59)|~l1_pre_topc(X59))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_tmap_1])])])])).
cnf(c_0_51, plain, (r1_t_0topsp(X2,X3)|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v3_tops_2(X1,X2,X3)|~l1_pre_topc(X3)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_52, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v3_tops_2(k1_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X3),X4,X5),X1,X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X3)|~v1_funct_1(X4)|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_1(X5)|~v1_funct_2(X5,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v3_tops_2(X4,X1,X2)|~v3_tops_2(X5,X2,X3)), inference(split_conjunct,[status(thm)],[c_0_45])).
fof(c_0_53, plain, ![X72, X73, X74, X75, X76, X77]:((v1_funct_1(k1_partfun1(X72,X73,X74,X75,X76,X77))|(~v1_funct_1(X76)|~m1_subset_1(X76,k1_zfmisc_1(k2_zfmisc_1(X72,X73)))|~v1_funct_1(X77)|~m1_subset_1(X77,k1_zfmisc_1(k2_zfmisc_1(X74,X75)))))&(m1_subset_1(k1_partfun1(X72,X73,X74,X75,X76,X77),k1_zfmisc_1(k2_zfmisc_1(X72,X75)))|(~v1_funct_1(X76)|~m1_subset_1(X76,k1_zfmisc_1(k2_zfmisc_1(X72,X73)))|~v1_funct_1(X77)|~m1_subset_1(X77,k1_zfmisc_1(k2_zfmisc_1(X74,X75)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_partfun1])])])).
cnf(c_0_54, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_46])).
cnf(c_0_55, plain, (k2_struct_0(k1_pre_topc(X1,X2))=X2|~v1_pre_topc(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(k1_pre_topc(X1,X2),X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_47])])).
cnf(c_0_56, plain, (m1_pre_topc(k1_pre_topc(X1,X2),X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[c_0_48])).
cnf(c_0_57, plain, (v1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[c_0_49])).
fof(c_0_58, plain, ![X44]:(~l1_pre_topc(X44)|l1_struct_0(X44)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
fof(c_0_59, plain, ![X33, X34]:(~l1_pre_topc(X33)|(~m1_pre_topc(X34,X33)|l1_pre_topc(X34))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_pre_topc])])])).
fof(c_0_60, plain, ![X43]:((l1_rlvect_1(X43)|~l1_rltopsp1(X43))&(l1_pre_topc(X43)|~l1_rltopsp1(X43))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_rltopsp1])])])).
fof(c_0_61, plain, ![X46]:((v5_rltopsp1(k15_euclid(X46))|~v7_ordinal1(X46))&(l1_rltopsp1(k15_euclid(X46))|~v7_ordinal1(X46))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k15_euclid])])])).
fof(c_0_62, plain, ![X39, X40]:(~v6_membered(X39)|(~m1_subset_1(X40,X39)|v7_ordinal1(X40))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc17_membered])])])).
cnf(c_0_63, plain, (v6_membered(k4_ordinal1)), inference(split_conjunct,[status(thm)],[fc6_membered])).
cnf(c_0_64, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
cnf(c_0_65, plain, (X1=k1_pre_topc(X3,X2)|k2_struct_0(X1)!=X2|~v1_pre_topc(X1)|~m1_pre_topc(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_66, plain, (m1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),X2)|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_67, plain, (r1_t_0topsp(X2,X3)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v1_funct_1(X1)|~v3_tops_2(X1,X2,X3)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))), inference(evalgc,[status(thm)],[c_0_51])).
cnf(c_0_68, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v3_tops_2(k1_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X3),X4,X5),X1,X3)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v3_tops_2(X5,X2,X3)|~v3_tops_2(X4,X1,X2)|~v1_funct_2(X5,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_52])).
cnf(c_0_69, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_70, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_53])).
fof(c_0_71, plain, ![X48, X49, X50, X51, X52, X53]:(~v1_funct_1(X52)|~m1_subset_1(X52,k1_zfmisc_1(k2_zfmisc_1(X48,X49)))|~v1_funct_1(X53)|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(X50,X51)))|k1_partfun1(X48,X49,X50,X51,X52,X53)=k3_relat_1(X52,X53)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_partfun1])])).
cnf(c_0_72, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_54])).
cnf(c_0_73, plain, (k2_struct_0(k1_pre_topc(X1,X2))=X2|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_55, c_0_56]), c_0_57])).
cnf(c_0_74, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_75, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
fof(c_0_76, plain, ![X32]:(((~v2_struct_0(X32)|~v1_toprealb(X32)|~m1_pre_topc(X32,k15_euclid(np__2)))&(v1_compts_1(X32)|~v1_toprealb(X32)|~m1_pre_topc(X32,k15_euclid(np__2))))&(v1_borsuk_2(X32)|~v1_toprealb(X32)|~m1_pre_topc(X32,k15_euclid(np__2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_toprealb])])])])).
cnf(c_0_77, plain, (l1_pre_topc(X1)|~l1_rltopsp1(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_78, plain, (l1_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_79, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_80, plain, (m1_subset_1(np__2,k5_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
cnf(c_0_81, plain, (v6_membered(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_63])).
cnf(c_0_82, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_64])).
fof(c_0_83, plain, ![X19]:(((v1_topreal2(u1_struct_0(X19))|~v1_toprealb(X19)|~m1_pre_topc(X19,k15_euclid(np__2)))&(m1_subset_1(u1_struct_0(X19),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_toprealb(X19)|~m1_pre_topc(X19,k15_euclid(np__2))))&(~v1_topreal2(u1_struct_0(X19))|~m1_subset_1(u1_struct_0(X19),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|v1_toprealb(X19)|~m1_pre_topc(X19,k15_euclid(np__2)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_toprealb])])])).
cnf(c_0_84, plain, (X1=k1_pre_topc(X3,X2)|k2_struct_0(X1)!=X2|~l1_pre_topc(X3)|~v1_pre_topc(X1)|~m1_pre_topc(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_85, plain, (m1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),X2)|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2)), inference(evalgc,[status(thm)],[c_0_66])).
cnf(c_0_86, plain, (v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_pre_topc(X1,X2)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_87, plain, (r1_t_0topsp(X1,X2)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|~m1_subset_1(k1_partfun1(u1_struct_0(X1),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X2),X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k1_partfun1(u1_struct_0(X1),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X2),X4,X5),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k1_partfun1(u1_struct_0(X1),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X2),X4,X5))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v3_tops_2(X5,X3,X2)|~v3_tops_2(X4,X1,X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l1_pre_topc(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68])])).
cnf(c_0_88, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_69])).
cnf(c_0_89, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_70])).
cnf(c_0_90, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_91, plain, (u1_struct_0(k1_pre_topc(X1,X2))=X2|~l1_struct_0(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_73])])).
cnf(c_0_92, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_74])).
cnf(c_0_93, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(evalgc,[status(thm)],[c_0_75])).
fof(c_0_94, plain, ![X281, X282, X283, X284, X285]:((v1_funct_1(k3_relat_1(X284,X285))|(v1_xboole_0(X282)|~v1_funct_1(X284)|~v1_funct_2(X284,X281,X282)|~m1_subset_1(X284,k1_zfmisc_1(k2_zfmisc_1(X281,X282)))|~v1_funct_1(X285)|~v1_funct_2(X285,X282,X283)|~m1_subset_1(X285,k1_zfmisc_1(k2_zfmisc_1(X282,X283)))))&(v1_funct_2(k3_relat_1(X284,X285),X281,X283)|(v1_xboole_0(X282)|~v1_funct_1(X284)|~v1_funct_2(X284,X281,X282)|~m1_subset_1(X284,k1_zfmisc_1(k2_zfmisc_1(X281,X282)))|~v1_funct_1(X285)|~v1_funct_2(X285,X282,X283)|~m1_subset_1(X285,k1_zfmisc_1(k2_zfmisc_1(X282,X283)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc8_funct_2])])])])).
fof(c_0_95, plain, ![X13, X15]:(((((v1_funct_1(esk4_1(X13))|~v1_topreal2(X13)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))&(v1_funct_2(esk4_1(X13),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X13)))|~v1_topreal2(X13)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))))&(m1_subset_1(esk4_1(X13),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X13)))))|~v1_topreal2(X13)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))))&(v3_tops_2(esk4_1(X13),k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(k15_euclid(np__2),X13))|~v1_topreal2(X13)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))))&(~v1_funct_1(X15)|~v1_funct_2(X15,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X13)))|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X13)))))|~v3_tops_2(X15,k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(k15_euclid(np__2),X13))|v1_topreal2(X13)|~m1_subset_1(X13,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_topreal2])])])])])).
cnf(c_0_96, plain, (~v2_struct_0(X1)|~v1_toprealb(X1)|~m1_pre_topc(X1,k15_euclid(np__2))), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_97, plain, (l1_pre_topc(X1)|~l1_rltopsp1(X1)), inference(evalgc,[status(thm)],[c_0_77])).
cnf(c_0_98, plain, (l1_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_78])).
cnf(c_0_99, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_79])).
cnf(c_0_100, plain, (m1_subset_1(np__2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_80])).
cnf(c_0_101, plain, (v6_membered(k5_numbers)), inference(rw,[status(thm)],[c_0_81, c_0_82])).
cnf(c_0_102, plain, (v1_toprealb(X1)|~v1_topreal2(u1_struct_0(X1))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_pre_topc(X1,k15_euclid(np__2))), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_103, plain, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=k1_pre_topc(X2,X3)|k2_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))!=X3|~v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85])])).
cnf(c_0_104, plain, (v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2)), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_105, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|r1_t_0topsp(X3,X1)|~v3_tops_2(X4,X2,X1)|~v3_tops_2(X5,X3,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(k1_partfun1(u1_struct_0(X3),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X1),X5,X4),u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X5)|~l1_pre_topc(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_87, c_0_88]), c_0_89])).
cnf(c_0_106, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_107, plain, (u1_struct_0(k1_pre_topc(X1,X2))=X2|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~l1_pre_topc(k1_pre_topc(X1,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92])])).
cnf(c_0_108, plain, (l1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_56])])).
cnf(c_0_109, plain, (v1_funct_2(k3_relat_1(X1,X2),X3,X4)|v1_xboole_0(X5)|~v1_funct_1(X1)|~v1_funct_2(X1,X3,X5)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X5)))|~v1_funct_1(X2)|~v1_funct_2(X2,X5,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,X4)))), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_110, plain, (v1_funct_2(esk4_1(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_111, plain, (~v1_toprealb(X1)|~v2_struct_0(X1)|~m1_pre_topc(X1,k15_euclid(np__2))), inference(evalgc,[status(thm)],[c_0_96])).
cnf(c_0_112, plain, (l1_pre_topc(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98])])).
cnf(c_0_113, plain, (v7_ordinal1(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_101])])])).
cnf(c_0_114, plain, (v1_toprealb(X1)|~v1_topreal2(u1_struct_0(X1))|~m1_pre_topc(X1,k15_euclid(np__2))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[c_0_102])).
fof(c_0_115, plain, ![X165, X166, X167, X168]:((X165=X167|g1_pre_topc(X165,X166)!=g1_pre_topc(X167,X168)|~m1_subset_1(X166,k1_zfmisc_1(k1_zfmisc_1(X165))))&(X166=X168|g1_pre_topc(X165,X166)!=g1_pre_topc(X167,X168)|~m1_subset_1(X166,k1_zfmisc_1(k1_zfmisc_1(X165))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_pre_topc])])])])).
fof(c_0_116, plain, ![X132]:(~l1_pre_topc(X132)|(~v1_pre_topc(X132)|X132=g1_pre_topc(u1_struct_0(X132),u1_pre_topc(X132)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_pre_topc])])).
cnf(c_0_117, plain, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=k1_pre_topc(X2,X3)|k2_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))!=X3|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2)), inference(csr,[status(thm)],[c_0_103, c_0_104])).
fof(c_0_118, negated_conjecture, ~(![X1]:((v1_toprealb(X1)&m1_pre_topc(X1,k15_euclid(np__2)))=>![X2]:((v1_toprealb(X2)&m1_pre_topc(X2,k15_euclid(np__2)))=>r2_borsuk_3(X1,X2)))), inference(assume_negation,[status(cth)],[t11_toprealb])).
cnf(c_0_119, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|r1_t_0topsp(X3,X1)|~v3_tops_2(X4,X2,X1)|~v3_tops_2(X5,X3,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(k3_relat_1(X5,X4),u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X5)|~l1_pre_topc(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_106])])).
cnf(c_0_120, plain, (u1_struct_0(k1_pre_topc(X1,X2))=X2|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(csr,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_121, plain, (v1_xboole_0(X5)|v1_funct_2(k3_relat_1(X1,X2),X3,X4)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,X5,X4)|~v1_funct_2(X1,X3,X5)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X5)))), inference(evalgc,[status(thm)],[c_0_109])).
cnf(c_0_122, plain, (v1_funct_2(esk4_1(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[c_0_110])).
cnf(c_0_123, plain, (v1_funct_1(esk4_1(X1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_124, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_125, plain, (~v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|~v1_toprealb(k1_pre_topc(k15_euclid(np__2),X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_56])])).
cnf(c_0_126, plain, (l1_pre_topc(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113])])).
cnf(c_0_127, plain, (v1_toprealb(k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~m1_subset_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_56])])).
fof(c_0_128, plain, ![X190, X191, X192]:(v1_xboole_0(X191)|(~m1_subset_1(X192,k1_zfmisc_1(k2_zfmisc_1(X190,X191)))|(~v1_funct_2(X192,X190,X191)|v1_partfun1(X192,X190)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc3_funct_2])])])])).
fof(c_0_129, plain, ![X63, X64, X65]:(((v1_funct_1(k2_tops_2(X63,X64,X65))|(~v1_funct_1(X65)|~v1_funct_2(X65,X63,X64)|~m1_subset_1(X65,k1_zfmisc_1(k2_zfmisc_1(X63,X64)))))&(v1_funct_2(k2_tops_2(X63,X64,X65),X64,X63)|(~v1_funct_1(X65)|~v1_funct_2(X65,X63,X64)|~m1_subset_1(X65,k1_zfmisc_1(k2_zfmisc_1(X63,X64))))))&(m1_subset_1(k2_tops_2(X63,X64,X65),k1_zfmisc_1(k2_zfmisc_1(X64,X63)))|(~v1_funct_1(X65)|~v1_funct_2(X65,X63,X64)|~m1_subset_1(X65,k1_zfmisc_1(k2_zfmisc_1(X63,X64)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_tops_2])])])).
cnf(c_0_130, plain, (X1=X2|g1_pre_topc(X1,X3)!=g1_pre_topc(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_115])).
cnf(c_0_131, plain, (X1=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~l1_pre_topc(X1)|~v1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_116])).
cnf(c_0_132, plain, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=k1_pre_topc(X2,X3)|u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))!=X3|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2)|~l1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_72])])).
fof(c_0_133, plain, ![X146]:(~l1_pre_topc(X146)|m1_subset_1(u1_pre_topc(X146),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X146))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_pre_topc])])).
fof(c_0_134, negated_conjecture, ((v1_toprealb(esk1_0)&m1_pre_topc(esk1_0,k15_euclid(np__2)))&((v1_toprealb(esk2_0)&m1_pre_topc(esk2_0,k15_euclid(np__2)))&~r2_borsuk_3(esk1_0,esk2_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_118])])])).
cnf(c_0_135, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(X4)|r1_t_0topsp(k1_pre_topc(X1,X2),X4)|~v3_tops_2(X5,X3,X4)|~v3_tops_2(X6,k1_pre_topc(X1,X2),X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~v1_funct_2(k3_relat_1(X6,X5),X2,u1_struct_0(X4))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_2(X6,X2,u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X6)|~l1_pre_topc(X4)|~l1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_120])])).
cnf(c_0_136, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v1_funct_2(k3_relat_1(X1,esk4_1(X2)),X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X2)))|~m1_subset_1(esk4_1(X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X2)))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~v1_funct_2(X1,X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_funct_1(esk4_1(X2))|~v1_funct_1(X1)|~v1_topreal2(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_122])])).
cnf(c_0_137, plain, (v1_funct_1(esk4_1(X1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[c_0_123])).
cnf(c_0_138, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[c_0_124])).
cnf(c_0_139, plain, (~v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_toprealb(k1_pre_topc(k15_euclid(np__2),X1))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_125, c_0_126])])).
cnf(c_0_140, plain, (v1_toprealb(k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~m1_subset_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_127, c_0_126])])).
cnf(c_0_141, plain, (v1_xboole_0(X1)|v1_partfun1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v1_funct_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_128])).
cnf(c_0_142, plain, (v1_funct_2(k2_tops_2(X1,X2,X3),X2,X1)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_143, plain, (X1=X2|g1_pre_topc(X1,X3)!=g1_pre_topc(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(evalgc,[status(thm)],[c_0_130])).
cnf(c_0_144, plain, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=X1|~l1_pre_topc(X1)|~v1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_131])).
fof(c_0_145, plain, ![X268, X269]:((v1_pre_topc(g1_pre_topc(X268,X269))|~m1_subset_1(X269,k1_zfmisc_1(k1_zfmisc_1(X268))))&(l1_pre_topc(g1_pre_topc(X268,X269))|~m1_subset_1(X269,k1_zfmisc_1(k1_zfmisc_1(X268))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_g1_pre_topc])])])).
cnf(c_0_146, plain, (k1_pre_topc(X2,u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~l1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_subset_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_132])])).
cnf(c_0_147, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_133])).
cnf(c_0_148, negated_conjecture, (m1_pre_topc(esk2_0,k15_euclid(np__2))), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_149, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(X4)|r1_t_0topsp(k1_pre_topc(X1,X2),X3)|~v3_tops_2(X5,k1_pre_topc(X1,X2),X4)|~v3_tops_2(X6,X4,X3)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X3))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X4))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k3_relat_1(X5,X6),X2,u1_struct_0(X3))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X3))|~v1_funct_2(X5,X2,u1_struct_0(X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~l1_pre_topc(X3)|~l1_pre_topc(X4)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[c_0_135, c_0_108])).
cnf(c_0_150, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v1_funct_2(k3_relat_1(X1,esk4_1(X2)),X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X2)))|~v1_topreal2(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(X1,X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_136, c_0_137]), c_0_138])).
cnf(c_0_151, plain, (~v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_topreal2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~m1_subset_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140])])).
fof(c_0_152, plain, ![X120, X121, X122]:(~m1_subset_1(X122,k1_zfmisc_1(k2_zfmisc_1(X120,X121)))|(~v1_partfun1(X122,X120)|v1_funct_2(X122,X120,X121))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_2])])).
cnf(c_0_153, plain, (v1_xboole_0(X1)|v1_partfun1(X2,X3)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_141])).
cnf(c_0_154, plain, (v1_funct_2(k2_tops_2(X1,X2,X3),X2,X1)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_142])).
cnf(c_0_155, plain, (m1_subset_1(k2_tops_2(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_129])).
fof(c_0_156, plain, ![X35, X36, X37]:(~l1_pre_topc(X35)|(v2_struct_0(X36)|~l1_pre_topc(X36)|(~v1_funct_1(X37)|~v1_funct_2(X37,u1_struct_0(X35),u1_struct_0(X36))|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X35),u1_struct_0(X36))))|(~v3_tops_2(X37,X35,X36)|v3_tops_2(k2_tops_2(u1_struct_0(X35),u1_struct_0(X36),X37),X36,X35))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t56_tops_2])])])])).
cnf(c_0_157, plain, (X1=u1_struct_0(X2)|g1_pre_topc(X1,X3)!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v1_pre_topc(X2)|~l1_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144])])).
cnf(c_0_158, plain, (l1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_145])).
cnf(c_0_159, plain, (v1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(split_conjunct,[status(thm)],[c_0_145])).
cnf(c_0_160, plain, (k1_pre_topc(X1,u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))|~m1_subset_1(u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_92])])).
cnf(c_0_161, plain, (l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_85])])).
fof(c_0_162, plain, (((((m1_pre_topc(esk5_0,k15_euclid(np__2))&v1_pre_topc(esk5_0))&v2_pre_topc(esk5_0))&v1_monoid_0(esk5_0))&v2_monoid_0(esk5_0))&v1_toprealb(esk5_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc1_toprealb])])).
cnf(c_0_163, plain, (u1_struct_0(X1)=X2|X1!=g1_pre_topc(X2,X3)|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_pre_topc(X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_144])])).
cnf(c_0_164, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_147])).
cnf(c_0_165, negated_conjecture, (m1_pre_topc(esk2_0,k15_euclid(np__2))), inference(evalgc,[status(thm)],[c_0_148])).
fof(c_0_166, plain, ![X41, X42]:(~v2_pre_topc(X41)|~l1_pre_topc(X41)|(~m1_pre_topc(X42,X41)|v2_pre_topc(X42))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_pre_topc])])])).
cnf(c_0_167, negated_conjecture, (m1_pre_topc(esk1_0,k15_euclid(np__2))), inference(split_conjunct,[status(thm)],[c_0_134])).
fof(c_0_168, plain, ![X27, X28]:((~r2_borsuk_3(X27,X28)|r1_t_0topsp(X27,X28)|(v2_struct_0(X27)|~l1_pre_topc(X27)|v2_struct_0(X28)|~l1_pre_topc(X28)))&(~r1_t_0topsp(X27,X28)|r2_borsuk_3(X27,X28)|(v2_struct_0(X27)|~l1_pre_topc(X27)|v2_struct_0(X28)|~l1_pre_topc(X28)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_r2_borsuk_3])])])])).
cnf(c_0_169, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(k1_pre_topc(k15_euclid(np__2),X3))|v2_struct_0(X4)|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(k15_euclid(np__2),X3))|v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v3_tops_2(X5,k1_pre_topc(X1,X2),X4)|~v3_tops_2(esk4_1(X3),X4,k1_pre_topc(k15_euclid(np__2),X3))|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X3)))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X4))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk4_1(X3),u1_struct_0(X4),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X3)))|~v1_funct_2(X5,X2,u1_struct_0(X4))|~v1_funct_1(esk4_1(X3))|~v1_funct_1(X5)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X3))|~l1_pre_topc(X4)|~l1_pre_topc(X1)|~v1_topreal2(X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(X5,X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_150])])).
cnf(c_0_170, plain, (~v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_120]), c_0_126])])])).
cnf(c_0_171, plain, (v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_partfun1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_152])).
cnf(c_0_172, plain, (v1_partfun1(k2_tops_2(X1,X2,X3),X2)|v1_xboole_0(X1)|~m1_subset_1(k2_tops_2(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_154])])).
cnf(c_0_173, plain, (m1_subset_1(k2_tops_2(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_155])).
cnf(c_0_174, plain, (v2_struct_0(X2)|v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v3_tops_2(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_156])).
cnf(c_0_175, plain, (u1_struct_0(g1_pre_topc(X1,X2))=X1|~v1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~l1_pre_topc(g1_pre_topc(X1,X2))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_157])])).
cnf(c_0_176, plain, (l1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(evalgc,[status(thm)],[c_0_158])).
cnf(c_0_177, plain, (v1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(evalgc,[status(thm)],[c_0_159])).
cnf(c_0_178, plain, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_toprealb(X1)|~m1_pre_topc(X1,k15_euclid(np__2))), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_179, negated_conjecture, (v1_toprealb(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_180, plain, (k1_pre_topc(X1,u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))|~m1_subset_1(u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(csr,[status(thm)],[c_0_160, c_0_161])).
cnf(c_0_181, plain, (m1_pre_topc(esk5_0,k15_euclid(np__2))), inference(split_conjunct,[status(thm)],[c_0_162])).
cnf(c_0_182, plain, (v1_toprealb(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_162])).
fof(c_0_183, plain, ![X55]:((~v1_xboole_0(X55)|~v1_topreal2(X55)|~m1_subset_1(X55,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))&(v2_compts_1(X55,k15_euclid(np__2))|~v1_topreal2(X55)|~m1_subset_1(X55,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc1_topreal2])])])])).
cnf(c_0_184, plain, (v1_topreal2(u1_struct_0(X1))|~v1_toprealb(X1)|~m1_pre_topc(X1,k15_euclid(np__2))), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_185, plain, (u1_struct_0(X1)=X2|X1!=g1_pre_topc(X2,X3)|~v1_pre_topc(X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[c_0_163, c_0_164])).
cnf(c_0_186, negated_conjecture, (v1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~l1_pre_topc(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_165])])).
cnf(c_0_187, plain, (v2_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_166])).
cnf(c_0_188, negated_conjecture, (m1_pre_topc(esk1_0,k15_euclid(np__2))), inference(evalgc,[status(thm)],[c_0_167])).
cnf(c_0_189, negated_conjecture, (~r2_borsuk_3(esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_190, plain, (r2_borsuk_3(X1,X2)|v2_struct_0(X1)|v2_struct_0(X2)|~r1_t_0topsp(X1,X2)|~l1_pre_topc(X1)|~l1_pre_topc(X2)), inference(split_conjunct,[status(thm)],[c_0_168])).
cnf(c_0_191, negated_conjecture, (v1_toprealb(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_192, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(k15_euclid(np__2),X4))|~v1_topreal2(X4)|~v3_tops_2(esk4_1(X4),X3,k1_pre_topc(k15_euclid(np__2),X4))|~v3_tops_2(X5,k1_pre_topc(X1,X2),X3)|~m1_subset_1(esk4_1(X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk4_1(X4),u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))|~v1_funct_2(X5,X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_funct_2(X5,X2,u1_struct_0(X3))|~v1_funct_1(X5)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X4))|~l1_pre_topc(X3)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_169, c_0_137]), c_0_170])).
cnf(c_0_193, plain, (v1_funct_1(k2_tops_2(X1,X2,X3))|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_194, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_171])).
cnf(c_0_195, plain, (v1_partfun1(k2_tops_2(X1,X2,X3),X2)|v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)), inference(csr,[status(thm)],[c_0_172, c_0_173])).
cnf(c_0_196, plain, (v2_struct_0(X2)|v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v1_funct_1(X3)|~v3_tops_2(X3,X1,X2)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_174])).
fof(c_0_197, plain, ![X97]:(~v2_struct_0(X97)|~l1_struct_0(X97)|v1_xboole_0(u1_struct_0(X97))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_struct_0])])).
cnf(c_0_198, plain, (u1_struct_0(g1_pre_topc(X1,X2))=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_175, c_0_176]), c_0_177])).
cnf(c_0_199, plain, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_toprealb(X1)|~m1_pre_topc(X1,k15_euclid(np__2))), inference(evalgc,[status(thm)],[c_0_178])).
cnf(c_0_200, negated_conjecture, (v1_toprealb(esk2_0)), inference(evalgc,[status(thm)],[c_0_179])).
fof(c_0_201, plain, ![X235, X236, X237]:(((v1_funct_1(X237)|(~v1_funct_1(X237)|~v1_funct_2(X237,X235,X236))|~m1_subset_1(X237,k1_zfmisc_1(k2_zfmisc_1(X235,X236)))|(v1_xboole_0(X235)|v1_xboole_0(X236)))&(~v1_xboole_0(X237)|(~v1_funct_1(X237)|~v1_funct_2(X237,X235,X236))|~m1_subset_1(X237,k1_zfmisc_1(k2_zfmisc_1(X235,X236)))|(v1_xboole_0(X235)|v1_xboole_0(X236))))&(v1_funct_2(X237,X235,X236)|(~v1_funct_1(X237)|~v1_funct_2(X237,X235,X236))|~m1_subset_1(X237,k1_zfmisc_1(k2_zfmisc_1(X235,X236)))|(v1_xboole_0(X235)|v1_xboole_0(X236)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc8_funct_2])])])])])).
fof(c_0_202, plain, ![X265]:(~v1_xboole_0(X265)|v1_funct_1(X265)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_funct_1])])).
cnf(c_0_203, plain, (m1_subset_1(k1_topreal1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(split_conjunct,[status(thm)],[dt_k1_topreal1])).
cnf(c_0_204, plain, (k1_pre_topc(X1,u1_struct_0(X2))=X2|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~v1_pre_topc(X2)|~l1_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_144])])).
cnf(c_0_205, plain, (v1_pre_topc(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_162])).
cnf(c_0_206, plain, (m1_pre_topc(esk5_0,k15_euclid(np__2))), inference(evalgc,[status(thm)],[c_0_181])).
cnf(c_0_207, plain, (v1_toprealb(esk5_0)), inference(evalgc,[status(thm)],[c_0_182])).
cnf(c_0_208, plain, (~v1_xboole_0(X1)|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(split_conjunct,[status(thm)],[c_0_183])).
cnf(c_0_209, plain, (v1_topreal2(u1_struct_0(X1))|~v1_toprealb(X1)|~m1_pre_topc(X1,k15_euclid(np__2))), inference(evalgc,[status(thm)],[c_0_184])).
cnf(c_0_210, plain, (u1_struct_0(g1_pre_topc(X1,X2))=X1|~v1_pre_topc(g1_pre_topc(X1,X2))|~l1_pre_topc(g1_pre_topc(X1,X2))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_185])])).
cnf(c_0_211, negated_conjecture, (v1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_186, c_0_126])])).
cnf(c_0_212, negated_conjecture, (l1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_165]), c_0_126])])])).
cnf(c_0_213, plain, (v2_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(evalgc,[status(thm)],[c_0_187])).
fof(c_0_214, plain, ![X68]:((((((((((v2_pre_topc(k15_euclid(X68))|~v7_ordinal1(X68))&(v13_algstr_0(k15_euclid(X68))|~v7_ordinal1(X68)))&(v2_rlvect_1(k15_euclid(X68))|~v7_ordinal1(X68)))&(v3_rlvect_1(k15_euclid(X68))|~v7_ordinal1(X68)))&(v4_rlvect_1(k15_euclid(X68))|~v7_ordinal1(X68)))&(v5_rlvect_1(k15_euclid(X68))|~v7_ordinal1(X68)))&(v6_rlvect_1(k15_euclid(X68))|~v7_ordinal1(X68)))&(v7_rlvect_1(k15_euclid(X68))|~v7_ordinal1(X68)))&(v8_rlvect_1(k15_euclid(X68))|~v7_ordinal1(X68)))&(v5_rltopsp1(k15_euclid(X68))|~v7_ordinal1(X68))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc6_euclid])])])).
cnf(c_0_215, negated_conjecture, (v1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|~l1_pre_topc(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_188])])).
cnf(c_0_216, negated_conjecture, (~r2_borsuk_3(esk1_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_189])).
cnf(c_0_217, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r2_borsuk_3(X1,X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~r1_t_0topsp(X1,X2)), inference(evalgc,[status(thm)],[c_0_190])).
cnf(c_0_218, negated_conjecture, (v1_toprealb(esk1_0)), inference(evalgc,[status(thm)],[c_0_191])).
cnf(c_0_219, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(k15_euclid(np__2),X4))|~v1_topreal2(X4)|~v3_tops_2(esk4_1(X4),X3,k1_pre_topc(k15_euclid(np__2),X4))|~v3_tops_2(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2,X5),k1_pre_topc(X1,X2),X3)|~m1_subset_1(esk4_1(X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))))|~m1_subset_1(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2,X5),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2,X5),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk4_1(X4),u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))|~v1_funct_2(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2,X5),X2,u1_struct_0(X3))|~v1_funct_1(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2,X5))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X4))|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2)))|~v1_funct_2(X5,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2)|~v1_funct_1(X5)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_154])])).
cnf(c_0_220, plain, (v1_funct_1(k2_tops_2(X1,X2,X3))|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_193])).
cnf(c_0_221, plain, (v1_funct_2(k2_tops_2(X1,X2,X3),X2,X4)|v1_xboole_0(X1)|~m1_subset_1(k2_tops_2(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_195])])).
cnf(c_0_222, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v3_tops_2(k2_tops_2(u1_struct_0(X3),X2,X4),k1_pre_topc(X1,X2),X3)|~v3_tops_2(X4,X3,k1_pre_topc(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X2)))|~v1_funct_2(X4,u1_struct_0(X3),X2)|~v1_funct_1(X4)|~l1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_120])])).
cnf(c_0_223, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_197])).
cnf(c_0_224, plain, (k1_pre_topc(X1,u1_struct_0(X2))=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(u1_pre_topc(X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_198])])).
cnf(c_0_225, negated_conjecture, (m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_165]), c_0_200])])])).
cnf(c_0_226, plain, (v1_xboole_0(X2)|v1_xboole_0(X3)|~v1_xboole_0(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_201])).
cnf(c_0_227, plain, (v1_funct_1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_202])).
cnf(c_0_228, plain, (m1_subset_1(k1_topreal1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[c_0_203])).
cnf(c_0_229, plain, (k1_pre_topc(X1,u1_struct_0(X2))=X2|~v1_pre_topc(X2)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)), inference(csr,[status(thm)],[c_0_204, c_0_93])).
cnf(c_0_230, plain, (v1_pre_topc(esk5_0)), inference(evalgc,[status(thm)],[c_0_205])).
cnf(c_0_231, plain, (m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_206]), c_0_207])])])).
fof(c_0_232, plain, ~v1_xboole_0(k1_topreal1), inference(fof_simplification,[status(thm)],[fc3_topreal1])).
cnf(c_0_233, plain, (~v1_topreal2(X1)|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[c_0_208])).
cnf(c_0_234, plain, (v1_topreal2(u1_struct_0(esk5_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_206]), c_0_207])])])).
fof(c_0_235, plain, ![X25, X26]:((~r1_t_0topsp(X25,X26)|r1_t_0topsp(g1_pre_topc(u1_struct_0(X25),u1_pre_topc(X25)),g1_pre_topc(u1_struct_0(X26),u1_pre_topc(X26)))|(v2_struct_0(X26)|~v2_pre_topc(X26)|~l1_pre_topc(X26))|(~v2_pre_topc(X25)|~l1_pre_topc(X25)))&(~r1_t_0topsp(g1_pre_topc(u1_struct_0(X25),u1_pre_topc(X25)),g1_pre_topc(u1_struct_0(X26),u1_pre_topc(X26)))|r1_t_0topsp(X25,X26)|(v2_struct_0(X26)|~v2_pre_topc(X26)|~l1_pre_topc(X26))|(~v2_pre_topc(X25)|~l1_pre_topc(X25)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t15_topreala])])])])])).
cnf(c_0_236, negated_conjecture, (u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))=u1_struct_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_211]), c_0_212])])])).
cnf(c_0_237, negated_conjecture, (v2_pre_topc(esk2_0)|~v2_pre_topc(k15_euclid(np__2))|~l1_pre_topc(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_165])])).
cnf(c_0_238, plain, (v2_pre_topc(k15_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_214])).
cnf(c_0_239, negated_conjecture, (v1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_215, c_0_126])])).
cnf(c_0_240, negated_conjecture, (l1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_188]), c_0_126])])])).
cnf(c_0_241, negated_conjecture, (v2_struct_0(esk1_0)|v2_struct_0(esk2_0)|~r1_t_0topsp(esk1_0,esk2_0)|~l1_pre_topc(esk2_0)|~l1_pre_topc(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_217])])).
cnf(c_0_242, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_188]), c_0_218])])])).
cnf(c_0_243, negated_conjecture, (~v2_struct_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_165]), c_0_200])])])).
cnf(c_0_244, negated_conjecture, (l1_pre_topc(esk1_0)|~l1_pre_topc(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_188])])).
cnf(c_0_245, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(k15_euclid(np__2),X4))|~v1_topreal2(X4)|~v3_tops_2(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2,X5),k1_pre_topc(X1,X2),X3)|~v3_tops_2(esk4_1(X4),X3,k1_pre_topc(k15_euclid(np__2),X4))|~m1_subset_1(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2,X5),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(esk4_1(X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk4_1(X4),u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))|~v1_funct_2(X5,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2)|~v1_funct_1(X5)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X4))|~l1_pre_topc(X3)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_219, c_0_220]), c_0_221]), c_0_173])).
cnf(c_0_246, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v3_tops_2(k2_tops_2(u1_struct_0(X3),X2,X4),k1_pre_topc(X1,X2),X3)|~v3_tops_2(X4,X3,k1_pre_topc(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X4,u1_struct_0(X3),X2)|~v1_funct_1(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[c_0_222, c_0_108])).
cnf(c_0_247, plain, (v3_tops_2(esk4_1(X1),k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_248, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_223])).
cnf(c_0_249, negated_conjecture, (g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))=k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))|~m1_subset_1(u1_pre_topc(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_165]), c_0_225]), c_0_126])])])).
cnf(c_0_250, negated_conjecture, (l1_pre_topc(esk2_0)|~l1_pre_topc(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_165])])).
cnf(c_0_251, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|~v1_funct_1(X1)|~v1_xboole_0(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_226])).
cnf(c_0_252, plain, (v1_funct_1(X1)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_227])).
cnf(c_0_253, plain, (v1_funct_2(esk4_1(X1),k1_topreal1,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_120]), c_0_228]), c_0_126])])])).
cnf(c_0_254, plain, (k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk5_0))=esk5_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_206]), c_0_230]), c_0_231]), c_0_126])])])).
cnf(c_0_255, plain, (~v1_xboole_0(k1_topreal1)), inference(split_conjunct,[status(thm)],[c_0_232])).
cnf(c_0_256, plain, (~v1_xboole_0(u1_struct_0(esk5_0))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_234])])).
fof(c_0_257, plain, ![X208, X209, X210]:(~v1_xboole_0(X208)|(~m1_subset_1(X210,k1_zfmisc_1(k2_zfmisc_1(X208,X209)))|v1_xboole_0(X210))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_relset_1])])])).
cnf(c_0_258, plain, (r1_t_0topsp(X1,X2)|v2_struct_0(X2)|~r1_t_0topsp(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_235])).
cnf(c_0_259, negated_conjecture, (k1_pre_topc(X1,u1_struct_0(esk2_0))=g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(esk2_0,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_236])])).
cnf(c_0_260, negated_conjecture, (v2_pre_topc(esk2_0)|~v2_pre_topc(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_237, c_0_126])])])).
cnf(c_0_261, plain, (v2_pre_topc(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_238])).
cnf(c_0_262, negated_conjecture, (u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))=u1_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_239]), c_0_240])])])).
cnf(c_0_263, negated_conjecture, (v2_pre_topc(esk1_0)|~v2_pre_topc(k15_euclid(np__2))|~l1_pre_topc(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_188])])).
cnf(c_0_264, negated_conjecture, (~r1_t_0topsp(esk1_0,esk2_0)|~l1_pre_topc(esk2_0)|~l1_pre_topc(esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[c_0_241, c_0_242]), c_0_243])).
cnf(c_0_265, negated_conjecture, (l1_pre_topc(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_244, c_0_126])])])).
cnf(c_0_266, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(k15_euclid(np__2),X3))|~v1_topreal2(X3)|~v3_tops_2(esk4_1(X3),k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(k15_euclid(np__2),X3))|~m1_subset_1(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2,X4),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X3)))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk4_1(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X3)))|~v1_funct_2(X4,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2)|~v1_funct_1(X4)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X3))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(X1)|~v3_tops_2(X4,k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(X1,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_246])])).
cnf(c_0_267, plain, (v3_tops_2(esk4_1(X1),k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[c_0_247])).
cnf(c_0_268, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_248, c_0_92])])).
cnf(c_0_269, negated_conjecture, (u1_struct_0(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))=u1_struct_0(esk2_0)|~m1_subset_1(u1_pre_topc(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_249])])).
cnf(c_0_270, negated_conjecture, (l1_pre_topc(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_250, c_0_126])])])).
cnf(c_0_271, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)), inference(csr,[status(thm)],[c_0_251, c_0_252])).
cnf(c_0_272, plain, (v1_funct_2(esk4_1(u1_struct_0(esk5_0)),k1_topreal1,u1_struct_0(esk5_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_254]), c_0_234]), c_0_231])])])).
cnf(c_0_273, plain, (~v1_xboole_0(k1_topreal1)), inference(evalgc,[status(thm)],[c_0_255])).
cnf(c_0_274, plain, (~v1_xboole_0(u1_struct_0(esk5_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_256, c_0_231])])).
cnf(c_0_275, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(k1_topreal1,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_120]), c_0_228]), c_0_126])])])).
cnf(c_0_276, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(split_conjunct,[status(thm)],[c_0_257])).
cnf(c_0_277, plain, (v2_struct_0(X2)|r1_t_0topsp(X1,X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~r1_t_0topsp(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))), inference(evalgc,[status(thm)],[c_0_258])).
cnf(c_0_278, negated_conjecture, (g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))=k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_165]), c_0_225]), c_0_126])])])).
cnf(c_0_279, negated_conjecture, (v2_pre_topc(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_261]), c_0_113])])])).
cnf(c_0_280, negated_conjecture, (k1_pre_topc(X1,u1_struct_0(esk1_0))=g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(esk1_0,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_262])])).
cnf(c_0_281, negated_conjecture, (m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_188]), c_0_218])])])).
cnf(c_0_282, negated_conjecture, (v2_pre_topc(esk1_0)|~v2_pre_topc(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_263, c_0_126])])])).
cnf(c_0_283, negated_conjecture, (~r1_t_0topsp(esk1_0,esk2_0)|~l1_pre_topc(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_264, c_0_265])])])).
cnf(c_0_284, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(k15_euclid(np__2),X3))|~v1_topreal2(X3)|~v3_tops_2(X4,k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X4,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2)|~v1_funct_1(X4)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X3))|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_266, c_0_122]), c_0_138]), c_0_173]), c_0_267]), c_0_268])).
cnf(c_0_285, plain, (l1_pre_topc(g1_pre_topc(u1_struct_0(k1_pre_topc(X1,X2)),u1_pre_topc(k1_pre_topc(X1,X2))))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_56])])).
cnf(c_0_286, negated_conjecture, (u1_struct_0(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))=u1_struct_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_269, c_0_164]), c_0_270])])])).
cnf(c_0_287, plain, (g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2)))=k1_pre_topc(X1,X2)|~v1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_120])])).
cnf(c_0_288, plain, (~v1_xboole_0(esk4_1(u1_struct_0(esk5_0)))|~m1_subset_1(esk4_1(u1_struct_0(esk5_0)),k1_zfmisc_1(k2_zfmisc_1(k1_topreal1,u1_struct_0(esk5_0))))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_272]), c_0_273]), c_0_274])])).
cnf(c_0_289, plain, (m1_subset_1(esk4_1(u1_struct_0(esk5_0)),k1_zfmisc_1(k2_zfmisc_1(k1_topreal1,u1_struct_0(esk5_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_254]), c_0_234]), c_0_231])])])).
cnf(c_0_290, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(evalgc,[status(thm)],[c_0_276])).
cnf(c_0_291, negated_conjecture, (r1_t_0topsp(X1,esk2_0)|~v2_pre_topc(X1)|~r1_t_0topsp(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_277, c_0_278]), c_0_279]), c_0_270])]), c_0_243])])).
cnf(c_0_292, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))=k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_188]), c_0_281]), c_0_126])])])).
cnf(c_0_293, negated_conjecture, (v2_pre_topc(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_261]), c_0_113])])])).
cnf(c_0_294, negated_conjecture, (~r1_t_0topsp(esk1_0,esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_283, c_0_270])])).
cnf(c_0_295, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X1),k1_pre_topc(k15_euclid(np__2),X2))|~v1_topreal2(X2)|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(esk4_1(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X1)|~v1_funct_1(esk4_1(X1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X2))|~v1_topreal2(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_284, c_0_267]), c_0_126])])])).
cnf(c_0_296, plain, (v1_funct_2(esk4_1(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X1)|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_120]), c_0_126])])])).
cnf(c_0_297, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X1)))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_120]), c_0_126])])])).
cnf(c_0_298, negated_conjecture, (l1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_286]), c_0_225]), c_0_126])])])).
cnf(c_0_299, plain, (g1_pre_topc(X1,u1_pre_topc(k1_pre_topc(X2,X1)))=k1_pre_topc(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_287, c_0_108]), c_0_57])).
cnf(c_0_300, plain, (~v1_xboole_0(esk4_1(u1_struct_0(esk5_0)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_288, c_0_289])])).
cnf(c_0_301, plain, (v1_xboole_0(esk4_1(X1))|~v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_138])])).
cnf(c_0_302, negated_conjecture, (~r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)),k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_292]), c_0_293]), c_0_265])]), c_0_294])])).
cnf(c_0_303, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X1),k1_pre_topc(k15_euclid(np__2),X2))|~v1_topreal2(X2)|~v1_topreal2(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_295, c_0_137]), c_0_296]), c_0_297]), c_0_170])).
cnf(c_0_304, negated_conjecture, (v1_topreal2(u1_struct_0(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_165]), c_0_200])])])).
cnf(c_0_305, negated_conjecture, (v1_topreal2(u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_188]), c_0_218])])])).
cnf(c_0_306, negated_conjecture, (l1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_299]), c_0_225]), c_0_126])])])).
cnf(c_0_307, plain, (~v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_301]), c_0_234]), c_0_231])])])).
cnf(c_0_308, negated_conjecture, (~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_303]), c_0_304]), c_0_305]), c_0_225]), c_0_281]), c_0_306])]), c_0_307])])).
cnf(c_0_309, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_308, c_0_108]), c_0_228]), c_0_126])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 310
# Proof object clause steps            : 231
# Proof object formula steps           : 79
# Proof object conjectures             : 54
# Proof object clause conjectures      : 51
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 59
# Proof object initial formulas used   : 41
# Proof object generating inferences   : 81
# Proof object simplifying inferences  : 154
# Training examples: 144 positive, 1519 negative
# Training: Positive examples begin
cnf(c_0_82, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_218, negated_conjecture, (v1_toprealb(esk1_0))).# trainpos
cnf(c_0_200, negated_conjecture, (v1_toprealb(esk2_0))).# trainpos
cnf(c_0_207, plain, (v1_toprealb(esk5_0))).# trainpos
cnf(c_0_216, negated_conjecture, (~r2_borsuk_3(esk1_0,esk2_0))).# trainpos
cnf(c_0_188, negated_conjecture, (m1_pre_topc(esk1_0,k15_euclid(np__2)))).# trainpos
cnf(c_0_111, plain, (~v1_toprealb(X1)|~v2_struct_0(X1)|~m1_pre_topc(X1,k15_euclid(np__2)))).# trainpos
cnf(c_0_92, plain, (l1_struct_0(X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_230, plain, (v1_pre_topc(esk5_0))).# trainpos
cnf(c_0_165, negated_conjecture, (m1_pre_topc(esk2_0,k15_euclid(np__2)))).# trainpos
cnf(c_0_81, plain, (v6_membered(k4_ordinal1))).# trainpos
cnf(c_0_273, plain, (~v1_xboole_0(k1_topreal1))).# trainpos
cnf(c_0_143, plain, (X1=X2|g1_pre_topc(X1,X3)!=g1_pre_topc(X2,X4)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1))))).# trainpos
cnf(c_0_114, plain, (v1_toprealb(X1)|~v1_topreal2(u1_struct_0(X1))|~m1_pre_topc(X1,k15_euclid(np__2))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_233, plain, (~v1_topreal2(X1)|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_100, plain, (m1_subset_1(np__2,k5_numbers))).# trainpos
cnf(c_0_97, plain, (l1_pre_topc(X1)|~l1_rltopsp1(X1))).# trainpos
cnf(c_0_206, plain, (m1_pre_topc(esk5_0,k15_euclid(np__2)))).# trainpos
cnf(c_0_252, plain, (v1_funct_1(X1)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_228, plain, (m1_subset_1(k1_topreal1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_209, plain, (v1_topreal2(u1_struct_0(X1))|~v1_toprealb(X1)|~m1_pre_topc(X1,k15_euclid(np__2)))).# trainpos
cnf(c_0_72, plain, (k2_struct_0(X1)=u1_struct_0(X1)|~l1_struct_0(X1))).# trainpos
cnf(c_0_199, plain, (m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_toprealb(X1)|~m1_pre_topc(X1,k15_euclid(np__2)))).# trainpos
cnf(c_0_261, plain, (v2_pre_topc(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_98, plain, (l1_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_93, plain, (l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1))).# trainpos
cnf(c_0_99, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_213, plain, (v2_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_pre_topc(X2,X1))).# trainpos
cnf(c_0_248, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).# trainpos
cnf(c_0_137, plain, (v1_funct_1(esk4_1(X1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_217, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r2_borsuk_3(X1,X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~r1_t_0topsp(X1,X2))).# trainpos
cnf(c_0_164, plain, (m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).# trainpos
cnf(c_0_176, plain, (l1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).# trainpos
cnf(c_0_177, plain, (v1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).# trainpos
cnf(c_0_144, plain, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=X1|~l1_pre_topc(X1)|~v1_pre_topc(X1))).# trainpos
cnf(c_0_290, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).# trainpos
cnf(c_0_104, plain, (v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2))).# trainpos
cnf(c_0_267, plain, (v3_tops_2(esk4_1(X1),k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_122, plain, (v1_funct_2(esk4_1(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_138, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_57, plain, (v1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_85, plain, (m1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),X2)|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2))).# trainpos
cnf(c_0_56, plain, (m1_pre_topc(k1_pre_topc(X1,X2),X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_194, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_153, plain, (v1_xboole_0(X1)|v1_partfun1(X2,X3)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_47, plain, (k2_struct_0(X1)=X3|X1!=k1_pre_topc(X2,X3)|~l1_pre_topc(X2)|~v1_pre_topc(X1)|~m1_pre_topc(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))).# trainpos
cnf(c_0_251, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|~v1_funct_1(X1)|~v1_xboole_0(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_84, plain, (X1=k1_pre_topc(X3,X2)|k2_struct_0(X1)!=X2|~l1_pre_topc(X3)|~v1_pre_topc(X1)|~m1_pre_topc(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3))))).# trainpos
cnf(c_0_220, plain, (v1_funct_1(k2_tops_2(X1,X2,X3))|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_154, plain, (v1_funct_2(k2_tops_2(X1,X2,X3),X2,X1)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_173, plain, (m1_subset_1(k2_tops_2(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~v1_funct_1(X3)|~v1_funct_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_88, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_106, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_277, plain, (v2_struct_0(X2)|r1_t_0topsp(X1,X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~r1_t_0topsp(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))).# trainpos
cnf(c_0_89, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_67, plain, (r1_t_0topsp(X2,X3)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v1_funct_1(X1)|~v3_tops_2(X1,X2,X3)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)))))).# trainpos
cnf(c_0_121, plain, (v1_xboole_0(X5)|v1_funct_2(k3_relat_1(X1,X2),X3,X4)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,X5,X4)|~v1_funct_2(X1,X3,X5)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X5))))).# trainpos
cnf(c_0_196, plain, (v2_struct_0(X2)|v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),X2,X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v1_funct_1(X3)|~v3_tops_2(X3,X1,X2)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).# trainpos
cnf(c_0_68, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X1)|v3_tops_2(k1_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X3),X4,X5),X1,X3)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v3_tops_2(X5,X2,X3)|~v3_tops_2(X4,X1,X2)|~v1_funct_2(X5,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))).# trainpos
cnf(c_0_243, negated_conjecture, (~v2_struct_0(esk2_0))).# trainpos
cnf(c_0_112, plain, (l1_pre_topc(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_242, negated_conjecture, (~v2_struct_0(esk1_0))).# trainpos
cnf(c_0_113, plain, (v7_ordinal1(np__2))).# trainpos
cnf(c_0_157, plain, (X1=u1_struct_0(X2)|g1_pre_topc(X1,X3)!=X2|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v1_pre_topc(X2)|~l1_pre_topc(X2))).# trainpos
cnf(c_0_250, negated_conjecture, (l1_pre_topc(esk2_0)|~l1_pre_topc(k15_euclid(np__2)))).# trainpos
cnf(c_0_241, negated_conjecture, (v2_struct_0(esk1_0)|v2_struct_0(esk2_0)|~r1_t_0topsp(esk1_0,esk2_0)|~l1_pre_topc(esk2_0)|~l1_pre_topc(esk1_0))).# trainpos
cnf(c_0_244, negated_conjecture, (l1_pre_topc(esk1_0)|~l1_pre_topc(k15_euclid(np__2)))).# trainpos
cnf(c_0_237, negated_conjecture, (v2_pre_topc(esk2_0)|~v2_pre_topc(k15_euclid(np__2))|~l1_pre_topc(k15_euclid(np__2)))).# trainpos
cnf(c_0_263, negated_conjecture, (v2_pre_topc(esk1_0)|~v2_pre_topc(k15_euclid(np__2))|~l1_pre_topc(k15_euclid(np__2)))).# trainpos
cnf(c_0_234, plain, (v1_topreal2(u1_struct_0(esk5_0)))).# trainpos
cnf(c_0_304, negated_conjecture, (v1_topreal2(u1_struct_0(esk2_0)))).# trainpos
cnf(c_0_268, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_281, negated_conjecture, (m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_225, negated_conjecture, (m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_305, negated_conjecture, (v1_topreal2(u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_231, plain, (m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_126, plain, (l1_pre_topc(k15_euclid(np__2)))).# trainpos
cnf(c_0_108, plain, (l1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_55, plain, (k2_struct_0(k1_pre_topc(X1,X2))=X2|~v1_pre_topc(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(k1_pre_topc(X1,X2),X1))).# trainpos
cnf(c_0_87, plain, (r1_t_0topsp(X1,X2)|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|~m1_subset_1(k1_partfun1(u1_struct_0(X1),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X2),X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k1_partfun1(u1_struct_0(X1),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X2),X4,X5),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k1_partfun1(u1_struct_0(X1),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X2),X4,X5))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v3_tops_2(X5,X3,X2)|~v3_tops_2(X4,X1,X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l1_pre_topc(X3))).# trainpos
cnf(c_0_103, plain, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=k1_pre_topc(X2,X3)|k2_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))!=X3|~v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2))).# trainpos
cnf(c_0_175, plain, (u1_struct_0(g1_pre_topc(X1,X2))=X1|~v1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~l1_pre_topc(g1_pre_topc(X1,X2)))).# trainpos
cnf(c_0_161, plain, (l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2))).# trainpos
cnf(c_0_125, plain, (~v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|~v1_toprealb(k1_pre_topc(k15_euclid(np__2),X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(k15_euclid(np__2)))).# trainpos
cnf(c_0_186, negated_conjecture, (v1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~l1_pre_topc(k15_euclid(np__2)))).# trainpos
cnf(c_0_301, plain, (v1_xboole_0(esk4_1(X1))|~v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_127, plain, (v1_toprealb(k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~m1_subset_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(k15_euclid(np__2)))).# trainpos
cnf(c_0_215, negated_conjecture, (v1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|~l1_pre_topc(k15_euclid(np__2)))).# trainpos
cnf(c_0_256, plain, (~v1_xboole_0(u1_struct_0(esk5_0))|~m1_subset_1(u1_struct_0(esk5_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_136, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v1_funct_2(k3_relat_1(X1,esk4_1(X2)),X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X2)))|~m1_subset_1(esk4_1(X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X2)))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~v1_funct_2(X1,X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_funct_1(esk4_1(X2))|~v1_funct_1(X1)|~v1_topreal2(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_172, plain, (v1_partfun1(k2_tops_2(X1,X2,X3),X2)|v1_xboole_0(X1)|~m1_subset_1(k2_tops_2(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).# trainpos
cnf(c_0_265, negated_conjecture, (l1_pre_topc(esk1_0))).# trainpos
cnf(c_0_270, negated_conjecture, (l1_pre_topc(esk2_0))).# trainpos
cnf(c_0_91, plain, (u1_struct_0(k1_pre_topc(X1,X2))=X2|~l1_struct_0(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).# trainpos
cnf(c_0_282, negated_conjecture, (v2_pre_topc(esk1_0)|~v2_pre_topc(k15_euclid(np__2)))).# trainpos
cnf(c_0_119, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|r1_t_0topsp(X3,X1)|~v3_tops_2(X4,X2,X1)|~v3_tops_2(X5,X3,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(k3_relat_1(X5,X4),u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X5)|~l1_pre_topc(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2))).# trainpos
cnf(c_0_221, plain, (v1_funct_2(k2_tops_2(X1,X2,X3),X2,X4)|v1_xboole_0(X1)|~m1_subset_1(k2_tops_2(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).# trainpos
cnf(c_0_260, negated_conjecture, (v2_pre_topc(esk2_0)|~v2_pre_topc(k15_euclid(np__2)))).# trainpos
cnf(c_0_212, negated_conjecture, (l1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))).# trainpos
cnf(c_0_240, negated_conjecture, (l1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))))).# trainpos
cnf(c_0_132, plain, (g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))=k1_pre_topc(X2,X3)|u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))!=X3|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2)|~l1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))).# trainpos
cnf(c_0_151, plain, (~v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_topreal2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~m1_subset_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_285, plain, (l1_pre_topc(g1_pre_topc(u1_struct_0(k1_pre_topc(X1,X2)),u1_pre_topc(k1_pre_topc(X1,X2))))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_279, negated_conjecture, (v2_pre_topc(esk2_0))).# trainpos
cnf(c_0_293, negated_conjecture, (v2_pre_topc(esk1_0))).# trainpos
cnf(c_0_107, plain, (u1_struct_0(k1_pre_topc(X1,X2))=X2|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~l1_pre_topc(k1_pre_topc(X1,X2)))).# trainpos
cnf(c_0_283, negated_conjecture, (~r1_t_0topsp(esk1_0,esk2_0)|~l1_pre_topc(esk2_0))).# trainpos
cnf(c_0_146, plain, (k1_pre_topc(X2,u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~l1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_subset_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2))).# trainpos
cnf(c_0_253, plain, (v1_funct_2(esk4_1(X1),k1_topreal1,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_275, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(k1_topreal1,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_170, plain, (~v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_297, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X1)))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_296, plain, (v1_funct_2(esk4_1(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X1)|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_160, plain, (k1_pre_topc(X1,u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))|~m1_subset_1(u1_struct_0(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))).# trainpos
cnf(c_0_287, plain, (g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2)))=k1_pre_topc(X1,X2)|~v1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).# trainpos
cnf(c_0_135, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(X4)|r1_t_0topsp(k1_pre_topc(X1,X2),X4)|~v3_tops_2(X5,X3,X4)|~v3_tops_2(X6,k1_pre_topc(X1,X2),X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~v1_funct_2(k3_relat_1(X6,X5),X2,u1_struct_0(X4))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_2(X6,X2,u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X6)|~l1_pre_topc(X4)|~l1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).# trainpos
cnf(c_0_222, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v3_tops_2(k2_tops_2(u1_struct_0(X3),X2,X4),k1_pre_topc(X1,X2),X3)|~v3_tops_2(X4,X3,k1_pre_topc(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X2)))|~v1_funct_2(X4,u1_struct_0(X3),X2)|~v1_funct_1(X4)|~l1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).# trainpos
cnf(c_0_204, plain, (k1_pre_topc(X1,u1_struct_0(X2))=X2|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~v1_pre_topc(X2)|~l1_pre_topc(X2))).# trainpos
cnf(c_0_169, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(k1_pre_topc(k15_euclid(np__2),X3))|v2_struct_0(X4)|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(k15_euclid(np__2),X3))|v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v3_tops_2(X5,k1_pre_topc(X1,X2),X4)|~v3_tops_2(esk4_1(X3),X4,k1_pre_topc(k15_euclid(np__2),X3))|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X3)))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X4))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk4_1(X3),u1_struct_0(X4),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X3)))|~v1_funct_2(X5,X2,u1_struct_0(X4))|~v1_funct_1(esk4_1(X3))|~v1_funct_1(X5)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X3))|~l1_pre_topc(X4)|~l1_pre_topc(X1)|~v1_topreal2(X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(X5,X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))))).# trainpos
cnf(c_0_224, plain, (k1_pre_topc(X1,u1_struct_0(X2))=g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~m1_subset_1(u1_pre_topc(X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2)))))).# trainpos
cnf(c_0_254, plain, (k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk5_0))=esk5_0)).# trainpos
cnf(c_0_272, plain, (v1_funct_2(esk4_1(u1_struct_0(esk5_0)),k1_topreal1,u1_struct_0(esk5_0)))).# trainpos
cnf(c_0_289, plain, (m1_subset_1(esk4_1(u1_struct_0(esk5_0)),k1_zfmisc_1(k2_zfmisc_1(k1_topreal1,u1_struct_0(esk5_0)))))).# trainpos
cnf(c_0_219, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(k15_euclid(np__2),X4))|~v1_topreal2(X4)|~v3_tops_2(esk4_1(X4),X3,k1_pre_topc(k15_euclid(np__2),X4))|~v3_tops_2(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2,X5),k1_pre_topc(X1,X2),X3)|~m1_subset_1(esk4_1(X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))))|~m1_subset_1(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2,X5),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2,X5),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk4_1(X4),u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))|~v1_funct_2(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2,X5),X2,u1_struct_0(X3))|~v1_funct_1(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2,X5))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X4))|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2)))|~v1_funct_2(X5,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2)|~v1_funct_1(X5))).# trainpos
cnf(c_0_266, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(k15_euclid(np__2),X3))|~v1_topreal2(X3)|~v3_tops_2(esk4_1(X3),k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(k15_euclid(np__2),X3))|~m1_subset_1(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2,X4),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X3)))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk4_1(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X3)))|~v1_funct_2(X4,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2)|~v1_funct_1(X4)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X3))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(X1)|~v3_tops_2(X4,k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(X1,X2)))).# trainpos
cnf(c_0_249, negated_conjecture, (g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))=k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))|~m1_subset_1(u1_pre_topc(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0)))))).# trainpos
cnf(c_0_269, negated_conjecture, (u1_struct_0(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))=u1_struct_0(esk2_0)|~m1_subset_1(u1_pre_topc(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0)))))).# trainpos
cnf(c_0_295, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X1),k1_pre_topc(k15_euclid(np__2),X2))|~v1_topreal2(X2)|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(esk4_1(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X1)|~v1_funct_1(esk4_1(X1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X2))|~v1_topreal2(X1))).# trainpos
cnf(c_0_163, plain, (u1_struct_0(X1)=X2|X1!=g1_pre_topc(X2,X3)|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_pre_topc(X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_286, negated_conjecture, (u1_struct_0(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))=u1_struct_0(esk2_0))).# trainpos
cnf(c_0_288, plain, (~v1_xboole_0(esk4_1(u1_struct_0(esk5_0)))|~m1_subset_1(esk4_1(u1_struct_0(esk5_0)),k1_zfmisc_1(k2_zfmisc_1(k1_topreal1,u1_struct_0(esk5_0)))))).# trainpos
cnf(c_0_210, plain, (u1_struct_0(g1_pre_topc(X1,X2))=X1|~v1_pre_topc(g1_pre_topc(X1,X2))|~l1_pre_topc(g1_pre_topc(X1,X2)))).# trainpos
cnf(c_0_298, negated_conjecture, (l1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))))))).# trainpos
cnf(c_0_307, plain, (~v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))))).# trainpos
cnf(c_0_236, negated_conjecture, (u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))=u1_struct_0(esk2_0))).# trainpos
cnf(c_0_262, negated_conjecture, (u1_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))=u1_struct_0(esk1_0))).# trainpos
cnf(c_0_259, negated_conjecture, (k1_pre_topc(X1,u1_struct_0(esk2_0))=g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(esk2_0,X1))).# trainpos
cnf(c_0_280, negated_conjecture, (k1_pre_topc(X1,u1_struct_0(esk1_0))=g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(esk1_0,X1))).# trainpos
cnf(c_0_306, negated_conjecture, (l1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_278, negated_conjecture, (g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))=k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))).# trainpos
cnf(c_0_291, negated_conjecture, (r1_t_0topsp(X1,esk2_0)|~v2_pre_topc(X1)|~r1_t_0topsp(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~l1_pre_topc(X1))).# trainpos
cnf(c_0_292, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))=k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_302, negated_conjecture, (~r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)),k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))))).# trainpos
cnf(c_0_308, negated_conjecture, (~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_152, plain, (o_0_0_xboole_0=k1_xboole_0)).#trainneg
cnf(i_0_230, plain, (l1_pre_topc(esk17_0))).#trainneg
cnf(i_0_276, plain, (l1_pre_topc(esk22_0))).#trainneg
cnf(i_0_280, plain, (l1_pre_topc(esk23_0))).#trainneg
cnf(i_0_333, plain, (l1_pre_topc(esk33_0))).#trainneg
cnf(i_0_343, plain, (l1_pre_topc(esk35_0))).#trainneg
cnf(i_0_436, plain, (l1_pre_topc(esk41_0))).#trainneg
cnf(i_0_455, plain, (l1_pre_topc(esk46_0))).#trainneg
cnf(i_0_110, plain, (v1_funct_1(esk6_0))).#trainneg
cnf(i_0_131, plain, (v1_funct_1(esk8_0))).#trainneg
cnf(i_0_129, plain, (v1_funct_2(esk8_0,k5_numbers,k5_numbers))).#trainneg
cnf(i_0_275, plain, (~v2_struct_0(esk22_0))).#trainneg
cnf(i_0_279, plain, (~v2_struct_0(esk23_0))).#trainneg
cnf(i_0_37, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_313, plain, (v1_funct_1(esk29_0))).#trainneg
cnf(i_0_460, plain, (v1_funct_1(esk47_0))).#trainneg
cnf(i_0_498, plain, (v1_funct_1(esk55_0))).#trainneg
cnf(i_0_42, plain, (v1_topreal2(k1_topreal1))).#trainneg
cnf(i_0_177, plain, (v1_topreal2(esk13_0))).#trainneg
cnf(i_0_228, plain, (v1_pre_topc(esk17_0))).#trainneg
cnf(i_0_331, plain, (v1_pre_topc(esk33_0))).#trainneg
cnf(i_0_341, plain, (v1_pre_topc(esk35_0))).#trainneg
cnf(i_0_434, plain, (v1_pre_topc(esk41_0))).#trainneg
cnf(i_0_229, plain, (v2_struct_0(esk17_0))).#trainneg
cnf(i_0_332, plain, (v2_struct_0(esk33_0))).#trainneg
cnf(i_0_101, plain, (v2_pre_topc(esk5_0))).#trainneg
cnf(i_0_227, plain, (v2_pre_topc(esk17_0))).#trainneg
cnf(i_0_282, plain, (~v2_struct_0(esk24_0))).#trainneg
cnf(i_0_274, plain, (v2_pre_topc(esk22_0))).#trainneg
cnf(i_0_202, plain, (X1=X2|g1_pre_topc(X3,X1)!=g1_pre_topc(X4,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X3))))).#trainneg
cnf(i_0_278, plain, (v2_pre_topc(esk23_0))).#trainneg
cnf(i_0_310, plain, (v1_funct_2(esk29_0,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_342, plain, (~v2_struct_0(esk35_0))).#trainneg
cnf(i_0_435, plain, (~v2_struct_0(esk41_0))).#trainneg
cnf(i_0_45, plain, (l1_rlvect_1(X1)|~l1_rltopsp1(X1))).#trainneg
cnf(i_0_97, plain, (v1_xcmplx_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_340, plain, (v2_pre_topc(esk35_0))).#trainneg
cnf(i_0_433, plain, (v2_pre_topc(esk41_0))).#trainneg
cnf(i_0_443, plain, (l1_struct_0(esk45_0))).#trainneg
cnf(i_0_469, plain, (l1_struct_0(esk48_0))).#trainneg
cnf(i_0_495, plain, (l1_struct_0(esk54_0))).#trainneg
cnf(i_0_277, plain, (v1_compts_1(esk23_0))).#trainneg
cnf(i_0_142, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_398, plain, (v2_xxreal_0(esk38_0))).#trainneg
cnf(i_0_122, plain, (v1_xcmplx_0(esk7_0))).#trainneg
cnf(i_0_184, plain, (v1_xcmplx_0(esk14_0))).#trainneg
cnf(i_0_199, plain, (v1_xcmplx_0(esk15_0))).#trainneg
cnf(i_0_400, plain, (v1_xcmplx_0(esk38_0))).#trainneg
cnf(i_0_492, plain, (v1_xcmplx_0(esk53_0))).#trainneg
cnf(i_0_294, plain, (~v1_xboole_0(np__2))).#trainneg
cnf(i_0_120, plain, (v1_xreal_0(esk7_0))).#trainneg
cnf(i_0_183, plain, (v1_xreal_0(esk14_0))).#trainneg
cnf(i_0_36, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_158, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_419, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_119, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_145, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_196, plain, (v1_xreal_0(esk15_0))).#trainneg
cnf(i_0_397, plain, (v1_xreal_0(esk38_0))).#trainneg
cnf(i_0_437, plain, (v1_xreal_0(esk42_0))).#trainneg
cnf(i_0_65, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_109, plain, (v1_xboole_0(esk6_0))).#trainneg
cnf(i_0_123, plain, (v1_xboole_0(esk7_0))).#trainneg
cnf(i_0_170, plain, (v1_xboole_0(esk11_0))).#trainneg
cnf(i_0_330, plain, (v1_xboole_0(esk32_0))).#trainneg
cnf(i_0_147, plain, (v3_membered(k1_numbers))).#trainneg
cnf(i_0_323, plain, (v3_membered(esk31_0))).#trainneg
cnf(i_0_475, plain, (v3_membered(esk50_0))).#trainneg
cnf(i_0_99, plain, (v2_monoid_0(esk5_0))).#trainneg
cnf(i_0_432, plain, (v2_monoid_0(esk41_0))).#trainneg
cnf(i_0_135, plain, (~v1_xboole_0(esk8_0))).#trainneg
cnf(i_0_178, plain, (~v1_xboole_0(esk13_0))).#trainneg
cnf(i_0_494, plain, (v2_monoid_0(esk54_0))).#trainneg
cnf(i_0_113, plain, (v1_relat_1(esk6_0))).#trainneg
cnf(i_0_141, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_300, plain, (~v1_xboole_0(esk27_0))).#trainneg
cnf(i_0_312, plain, (~v1_xboole_0(esk29_0))).#trainneg
cnf(i_0_172, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_151, plain, (v1_int_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_134, plain, (v1_relat_1(esk8_0))).#trainneg
cnf(i_0_316, plain, (v1_relat_1(esk29_0))).#trainneg
cnf(i_0_463, plain, (v1_relat_1(esk47_0))).#trainneg
cnf(i_0_499, plain, (v1_relat_1(esk55_0))).#trainneg
cnf(i_0_121, plain, (v1_xxreal_0(esk7_0))).#trainneg
cnf(i_0_169, plain, (v1_xxreal_0(esk11_0))).#trainneg
cnf(i_0_185, plain, (v1_xxreal_0(esk14_0))).#trainneg
cnf(i_0_198, plain, (v1_xxreal_0(esk15_0))).#trainneg
cnf(i_0_272, plain, (v1_xxreal_0(esk21_0))).#trainneg
cnf(i_0_399, plain, (v1_xxreal_0(esk38_0))).#trainneg
cnf(i_0_405, plain, (v1_xxreal_0(esk39_0))).#trainneg
cnf(i_0_197, plain, (v3_xxreal_0(esk15_0))).#trainneg
cnf(i_0_271, plain, (v3_xxreal_0(esk21_0))).#trainneg
cnf(i_0_100, plain, (v1_monoid_0(esk5_0))).#trainneg
cnf(i_0_326, plain, (~v1_xboole_0(esk31_0))).#trainneg
cnf(i_0_478, plain, (~v1_xboole_0(esk50_0))).#trainneg
cnf(i_0_468, plain, (v1_monoid_0(esk48_0))).#trainneg
cnf(i_0_108, plain, (v1_valued_0(esk6_0))).#trainneg
cnf(i_0_34, plain, (m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_496, plain, (~v1_setfam_1(k1_numbers))).#trainneg
cnf(i_0_390, plain, (~v1_setfam_1(k4_ordinal1))).#trainneg
cnf(i_0_195, plain, (v1_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_205, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_128, plain, (v1_valued_0(esk8_0))).#trainneg
cnf(i_0_309, plain, (v1_valued_0(esk29_0))).#trainneg
cnf(i_0_459, plain, (v1_valued_0(esk47_0))).#trainneg
cnf(i_0_107, plain, (v2_valued_0(esk6_0))).#trainneg
cnf(i_0_127, plain, (v2_valued_0(esk8_0))).#trainneg
cnf(i_0_308, plain, (v2_valued_0(esk29_0))).#trainneg
cnf(i_0_458, plain, (v2_valued_0(esk47_0))).#trainneg
cnf(i_0_106, plain, (v3_valued_0(esk6_0))).#trainneg
cnf(i_0_126, plain, (v3_valued_0(esk8_0))).#trainneg
cnf(i_0_307, plain, (v3_valued_0(esk29_0))).#trainneg
cnf(i_0_457, plain, (v3_valued_0(esk47_0))).#trainneg
cnf(i_0_125, plain, (v4_valued_0(esk8_0))).#trainneg
cnf(i_0_124, plain, (v5_valued_0(esk8_0))).#trainneg
cnf(i_0_182, plain, (v1_int_1(esk14_0))).#trainneg
cnf(i_0_322, plain, (~v1_setfam_1(esk31_0))).#trainneg
cnf(i_0_497, plain, (~v3_funct_1(esk55_0))).#trainneg
cnf(i_0_306, plain, (v1_fcont_1(esk29_0))).#trainneg
cnf(i_0_96, plain, (~v2_xxreal_0(X1)|~v1_xboole_0(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_456, plain, (v1_fcont_1(esk47_0))).#trainneg
cnf(i_0_283, plain, (l2_struct_0(esk24_0))).#trainneg
cnf(i_0_442, plain, (l2_struct_0(esk44_0))).#trainneg
cnf(i_0_214, plain, (v2_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_439, plain, (v3_ordinal1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_325, plain, (v1_membered(esk31_0))).#trainneg
cnf(i_0_477, plain, (v1_membered(esk50_0))).#trainneg
cnf(i_0_179, plain, (v4_funct_1(esk13_0))).#trainneg
cnf(i_0_324, plain, (v2_membered(esk31_0))).#trainneg
cnf(i_0_476, plain, (v2_membered(esk50_0))).#trainneg
cnf(i_0_273, plain, (v7_pre_topc(esk22_0))).#trainneg
cnf(i_0_281, plain, (v7_struct_0(esk24_0))).#trainneg
cnf(i_0_299, plain, (v3_ordinal1(esk27_0))).#trainneg
cnf(i_0_474, plain, (v1_setfam_1(esk50_0))).#trainneg
cnf(i_0_250, plain, (v1_funct_1(esk19_1(X1)))).#trainneg
cnf(i_0_263, plain, (v1_funct_1(esk20_1(X1)))).#trainneg
cnf(i_0_487, plain, (v1_funct_1(esk52_1(X1)))).#trainneg
cnf(i_0_162, plain, (v1_xboole_0(esk10_1(X1)))).#trainneg
cnf(i_0_262, plain, (v2_funct_1(esk20_1(X1)))).#trainneg
cnf(i_0_201, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_338, plain, (~v1_subset_1(esk34_1(X1),X1))).#trainneg
cnf(i_0_153, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_78, plain, (~v7_ordinal1(X1)|~v2_struct_0(k15_euclid(X1)))).#trainneg
cnf(i_0_138, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_139, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_140, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_186, plain, (m1_subset_1(esk14_0,k1_numbers))).#trainneg
cnf(i_0_55, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_253, plain, (v1_relat_1(esk19_1(X1)))).#trainneg
cnf(i_0_266, plain, (v1_relat_1(esk20_1(X1)))).#trainneg
cnf(i_0_489, plain, (v1_relat_1(esk52_1(X1)))).#trainneg
cnf(i_0_485, plain, (v4_valued_0(esk52_1(X1)))).#trainneg
cnf(i_0_112, plain, (v4_relat_1(esk6_0,k1_numbers))).#trainneg
cnf(i_0_133, plain, (v4_relat_1(esk8_0,k5_numbers))).#trainneg
cnf(i_0_315, plain, (v4_relat_1(esk29_0,k1_numbers))).#trainneg
cnf(i_0_462, plain, (v4_relat_1(esk47_0,k1_numbers))).#trainneg
cnf(i_0_111, plain, (v5_relat_1(esk6_0,k1_numbers))).#trainneg
cnf(i_0_132, plain, (v5_relat_1(esk8_0,k5_numbers))).#trainneg
cnf(i_0_314, plain, (v5_relat_1(esk29_0,k1_numbers))).#trainneg
cnf(i_0_461, plain, (v5_relat_1(esk47_0,k1_numbers))).#trainneg
cnf(i_0_130, plain, (v1_partfun1(esk8_0,k5_numbers))).#trainneg
cnf(i_0_311, plain, (v1_partfun1(esk29_0,k1_numbers))).#trainneg
cnf(i_0_389, plain, (~v3_xxreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_30, plain, (v1_compts_1(X1)|~v1_toprealb(X1)|~m1_pre_topc(X1,k15_euclid(np__2)))).#trainneg
cnf(i_0_363, plain, (~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_xboole_0(u1_pre_topc(X1)))).#trainneg
cnf(i_0_159, plain, (l1_struct_0(X1)|~l1_algstr_0(X1))).#trainneg
cnf(i_0_168, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_301, plain, (m1_subset_1(esk27_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_215, plain, (l1_struct_0(X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_200, plain, (v1_xreal_0(X1)|~v1_int_1(X1))).#trainneg
cnf(i_0_470, plain, (v3_membered(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_93, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_327, plain, (m1_subset_1(esk31_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_479, plain, (m1_subset_1(esk50_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_146, plain, (m1_subset_1(esk9_1(X1),X1))).#trainneg
cnf(i_0_259, plain, (v2_funct_2(esk20_1(X1),X1))).#trainneg
cnf(i_0_252, plain, (v4_relat_1(esk19_1(X1),X1))).#trainneg
cnf(i_0_265, plain, (v4_relat_1(esk20_1(X1),X1))).#trainneg
cnf(i_0_488, plain, (v4_relat_1(esk52_1(X1),X1))).#trainneg
cnf(i_0_251, plain, (v5_relat_1(esk19_1(X1),X1))).#trainneg
cnf(i_0_264, plain, (v5_relat_1(esk20_1(X1),X1))).#trainneg
cnf(i_0_249, plain, (v1_partfun1(esk19_1(X1),X1))).#trainneg
cnf(i_0_261, plain, (v1_partfun1(esk20_1(X1),X1))).#trainneg
cnf(i_0_486, plain, (v1_partfun1(esk52_1(X1),X1))).#trainneg
cnf(i_0_143, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_29, plain, (v1_borsuk_2(X1)|~v1_toprealb(X1)|~m1_pre_topc(X1,k15_euclid(np__2)))).#trainneg
cnf(i_0_160, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_334, plain, (l2_algstr_0(X1)|~l1_rlvect_1(X1))).#trainneg
cnf(i_0_180, plain, (m1_subset_1(esk13_0,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_114, plain, (m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_161, plain, (l2_struct_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_144, plain, (v1_monoid_0(X1)|~l1_struct_0(X1)|~v2_monoid_0(X1))).#trainneg
cnf(i_0_258, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_286, plain, (v1_funct_1(esk25_2(X1,X2)))).#trainneg
cnf(i_0_350, plain, (v1_xboole_0(esk36_2(X1,X2)))).#trainneg
cnf(i_0_92, plain, (v1_relat_1(k3_relat_1(X1,X2)))).#trainneg
cnf(i_0_289, plain, (v1_relat_1(esk25_2(X1,X2)))).#trainneg
cnf(i_0_349, plain, (v1_relat_1(esk36_2(X1,X2)))).#trainneg
cnf(i_0_163, plain, (m1_subset_1(esk10_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_339, plain, (m1_subset_1(esk34_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_288, plain, (v4_relat_1(esk25_2(X1,X2),X1))).#trainneg
cnf(i_0_348, plain, (v4_relat_1(esk36_2(X1,X2),X1))).#trainneg
cnf(i_0_287, plain, (v5_relat_1(esk25_2(X1,X2),X2))).#trainneg
cnf(i_0_347, plain, (v5_relat_1(esk36_2(X1,X2),X2))).#trainneg
cnf(i_0_248, plain, (v1_funct_2(esk19_1(X1),X1,X1))).#trainneg
cnf(i_0_94, plain, (~v2_xxreal_0(X1)|~v1_xxreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_284, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_136, plain, (m1_subset_1(esk8_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_446, plain, (v1_funct_1(u1_algstr_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_47, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_317, plain, (m1_subset_1(esk29_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_464, plain, (m1_subset_1(esk47_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_376, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_465, plain, (v7_struct_0(X1)|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_209, plain, (v1_funct_1(esk16_1(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_260, plain, (v1_funct_2(esk20_1(X1),X1,X1))).#trainneg
cnf(i_0_247, plain, (v3_funct_2(esk19_1(X1),X1,X1))).#trainneg
cnf(i_0_49, plain, (v5_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_74, plain, (v13_algstr_0(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_73, plain, (v2_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_72, plain, (v3_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_71, plain, (v4_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_70, plain, (v5_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_403, plain, (v1_funct_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_50, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_69, plain, (v6_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_421, plain, (v1_int_1(X1)|~m1_subset_1(X1,k4_numbers))).#trainneg
cnf(i_0_408, plain, (~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_502, plain, (~l1_struct_0(X1)|~v1_subset_1(k2_struct_0(X1),u1_struct_0(X1)))).#trainneg
cnf(i_0_68, plain, (v7_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_67, plain, (v8_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_270, plain, (v2_monoid_0(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_212, plain, (v1_relat_1(esk16_1(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_471, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk49_1(X1)))).#trainneg
cnf(i_0_285, plain, (v1_funct_2(esk25_2(X1,X2),X1,X2))).#trainneg
cnf(i_0_440, plain, (v2_tops_1(esk43_1(X1),X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_137, plain, (v1_tops_1(k2_struct_0(X1),X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_500, plain, (v1_tops_1(esk56_1(X1),X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_256, plain, (v4_funct_1(u1_struct_0(k15_euclid(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_420, plain, (v6_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_302, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_76, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_188, plain, (v1_xboole_0(k2_zfmisc_1(X2,X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_447, plain, (v7_pre_topc(X1)|~l1_pre_topc(X1)|~v8_pre_topc(X1))).#trainneg
cnf(i_0_373, plain, (v1_connsp_1(X1)|~l1_pre_topc(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_450, plain, (v2_xxreal_0(X1)|v1_xboole_0(X1)|v3_xxreal_0(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_360, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_88, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_268, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(k2_struct_0(X1)))).#trainneg
cnf(i_0_257, plain, (v1_xboole_0(k2_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_187, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_352, plain, (v3_valued_0(k2_zfmisc_1(X2,X1))|~v3_membered(X1))).#trainneg
cnf(i_0_254, plain, (m1_subset_1(esk19_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_267, plain, (m1_subset_1(esk20_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_472, plain, (v1_xboole_0(X1)|m1_subset_1(esk49_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_171, plain, (m1_subset_1(u2_struct_0(X1),u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_211, plain, (v4_relat_1(esk16_1(X1),u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_210, plain, (v5_relat_1(esk16_1(X1),u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_208, plain, (v1_partfun1(esk16_1(X1),u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_117, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_118, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_206, plain, (v3_tops_2(esk16_1(X1),X1,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_167, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_64, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_175, plain, (v1_xboole_0(esk12_1(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_239, plain, (v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_xboole_0(esk18_1(X1)))).#trainneg
cnf(i_0_328, plain, (v1_compts_1(X1)|~l1_pre_topc(X1)|~v2_struct_0(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_336, plain, (v1_tdlat_3(X1)|~l1_pre_topc(X1)|~v2_struct_0(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_53, plain, (v2_compts_1(X1,k15_euclid(np__2))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_335, plain, (v2_tdlat_3(X1)|~l1_pre_topc(X1)|~v2_struct_0(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_173, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_411, plain, (v2_struct_0(X1)|v1_borsuk_2(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_424, plain, (v2_struct_0(X1)|v2_pscomp_1(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_compts_1(X1))).#trainneg
cnf(i_0_290, plain, (m1_subset_1(esk25_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_351, plain, (m1_subset_1(esk36_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_404, plain, (v1_relat_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_255, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_39, plain, (m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_441, plain, (m1_subset_1(esk43_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_501, plain, (m1_subset_1(esk56_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_87, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_52, plain, (v3_pre_topc(k2_struct_0(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_104, plain, (v4_pre_topc(k2_struct_0(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_417, plain, (v3_tops_1(esk40_1(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_115, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_204, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_473, plain, (v1_membered(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_490, plain, (v2_fcont_1(X1)|~v1_funct_1(X1)|~v3_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_292, plain, (v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_xboole_0(esk26_1(X1)))).#trainneg
cnf(i_0_304, plain, (v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_xboole_0(esk28_1(X1)))).#trainneg
cnf(i_0_374, plain, (v1_fcont_1(X1)|~v1_funct_1(X1)|~v2_fcont_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_409, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_320, plain, (v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_xboole_0(esk30_1(X1)))).#trainneg
cnf(i_0_56, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r2_borsuk_3(X1,X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_319, plain, (v2_struct_0(X1)|v3_pre_topc(esk30_1(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_303, plain, (v2_struct_0(X1)|v4_pre_topc(esk28_1(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_354, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_207, plain, (v1_funct_2(esk16_1(X1),u1_struct_0(X1),u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_166, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_165, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_191, plain, (v3_valued_0(X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1))))).#trainneg
cnf(i_0_241, plain, (v1_finseq_1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1))))).#trainneg
cnf(i_0_176, plain, (m1_subset_1(esk12_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_418, plain, (m1_subset_1(esk40_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_233, plain, (v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))).#trainneg
cnf(i_0_362, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_427, plain, (v1_fcont_1(X1)|~v1_funct_1(X1)|~v3_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_402, plain, (v1_xboole_0(X1)|~v2_struct_0(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_318, plain, (v2_struct_0(X1)|v4_pre_topc(esk30_1(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_89, plain, (v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v2_tops_1(k2_struct_0(X1),X1))).#trainneg
cnf(i_0_493, plain, (v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_236, plain, (v1_funct_1(X1)|~l1_struct_0(X2)|~v1_monoid_0(X2)|~m1_subset_1(X1,u1_struct_0(X2)))).#trainneg
cnf(i_0_25, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r1_t_0topsp(X1,X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~r2_borsuk_3(X1,X2))).#trainneg
cnf(i_0_26, plain, (v2_struct_0(X2)|v2_struct_0(X1)|r2_borsuk_3(X2,X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~r2_borsuk_3(X1,X2))).#trainneg
cnf(i_0_238, plain, (v2_struct_0(X1)|v1_tops_1(esk18_1(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_86, plain, (v2_monoid_0(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v2_monoid_0(X1)|~m1_pre_topc(X2,X1))).#trainneg
cnf(i_0_232, plain, (v2_struct_0(X1)|v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_346, plain, (v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_345, plain, (v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_116, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_483, plain, (v3_card_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1))))).#trainneg
cnf(i_0_245, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_298, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_213, plain, (m1_subset_1(esk16_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_237, plain, (v1_relat_1(X1)|~l1_struct_0(X2)|~v1_monoid_0(X2)|~m1_subset_1(X1,u1_struct_0(X2)))).#trainneg
cnf(i_0_482, plain, (v1_valued_0(X2)|~v1_relat_1(X1)|~v1_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_448, plain, (v3_valued_0(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_344, plain, (v1_finseq_1(X2)|~l1_struct_0(X1)|~v2_monoid_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_11, plain, (v1_topreal2(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v3_tops_2(X1,k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(k15_euclid(np__2),X2))|~v1_funct_2(X1,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X2)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X2))))))).#trainneg
cnf(i_0_291, plain, (v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v2_tops_1(esk26_1(X1),X1))).#trainneg
cnf(i_0_396, plain, (v2_tops_1(X2,X1)|~l1_pre_topc(X1)|~v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_225, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_226, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_391, plain, (v3_valued_0(k3_relat_1(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v3_valued_0(X2))).#trainneg
cnf(i_0_445, plain, (v1_funct_2(u1_algstr_0(X1),k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_10, plain, (v1_funct_1(esk3_2(X1,X2))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~r1_t_0topsp(X1,X2))).#trainneg
cnf(i_0_240, plain, (v2_struct_0(X1)|m1_subset_1(esk18_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_293, plain, (v2_struct_0(X1)|m1_subset_1(esk26_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_305, plain, (v2_struct_0(X1)|m1_subset_1(esk28_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_321, plain, (v2_struct_0(X1)|m1_subset_1(esk30_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_480, plain, (v1_connsp_1(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_connsp_1(X1))).#trainneg
cnf(i_0_223, plain, (v2_funct_1(X1)|~v1_funct_1(X1)|~v3_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_388, plain, (v1_partfun1(X1,X2)|~v1_funct_2(X1,X2,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2))))).#trainneg
cnf(i_0_149, plain, (v1_funct_1(k3_relat_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_91, plain, (v2_struct_0(X1)|v1_xboole_0(X2)|~l1_pre_topc(X1)|~v2_struct_0(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_410, plain, (v3_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_157, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_366, plain, (v2_struct_0(X2)|v1_funct_1(esk37_2(X1,X2))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_369, plain, (v2_struct_0(X2)|v1_relat_1(esk37_2(X1,X2))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_422, plain, (v2_funct_1(k3_relat_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_funct_1(X2)|~v2_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_444, plain, (m1_subset_1(u1_algstr_0(X1),k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)),u1_struct_0(X1))))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_217, plain, (v2_struct_0(k1_pre_topc(X1,X2))|~l1_pre_topc(X1)|~v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_189, plain, (v2_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_7, plain, (v3_tops_2(esk3_2(X1,X2),X1,X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~r1_t_0topsp(X1,X2))).#trainneg
cnf(i_0_484, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m2_subset_1(esk51_2(X1,X2),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_9, plain, (v1_funct_2(esk3_2(X1,X2),u1_struct_0(X1),u1_struct_0(X2))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~r1_t_0topsp(X1,X2))).#trainneg
cnf(i_0_222, plain, (v2_funct_2(X1,X2)|~v1_funct_1(X1)|~v3_funct_2(X1,X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_235, plain, (v1_funct_1(k3_relat_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_funct_2(X2,X3)|~v2_funct_2(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X3))))).#trainneg
cnf(i_0_368, plain, (v2_struct_0(X2)|v4_relat_1(esk37_2(X1,X2),u1_struct_0(X1))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_220, plain, (v3_funct_2(X1,X2,X3)|~v1_funct_1(X1)|~v2_funct_1(X1)|~v2_funct_2(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_156, plain, (v1_funct_1(k3_relat_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_fcont_1(X2)|~v1_fcont_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_155, plain, (v1_fcont_1(k3_relat_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_fcont_1(X2)|~v1_fcont_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_364, plain, (v2_struct_0(X2)|v5_pre_topc(esk37_2(X1,X2),X1,X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_367, plain, (v2_struct_0(X2)|v5_relat_1(esk37_2(X1,X2),u1_struct_0(X2))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_365, plain, (v2_struct_0(X2)|v1_funct_2(esk37_2(X1,X2),u1_struct_0(X1),u1_struct_0(X2))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_8, plain, (m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~r1_t_0topsp(X1,X2))).#trainneg
cnf(i_0_370, plain, (v2_struct_0(X2)|m1_subset_1(esk37_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_23, plain, (v2_struct_0(X2)|r1_t_0topsp(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~r1_t_0topsp(X1,X2))).#trainneg
cnf(i_0_84, plain, (v2_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v1_funct_1(X1)|~v3_tops_2(X1,X2,X3)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)))))).#trainneg
cnf(i_0_234, plain, (v2_funct_2(k3_relat_1(X1,X2),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_funct_2(X2,X3)|~v2_funct_2(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X3))))).#trainneg
cnf(i_0_82, plain, (v2_funct_2(X1,u1_struct_0(X2))|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v1_funct_1(X1)|~v3_tops_2(X1,X3,X2)|~v1_funct_2(X1,u1_struct_0(X3),u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))).#trainneg
cnf(i_0_414, plain, (v5_pre_topc(X1,X2,X3)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v1_funct_1(X1)|~v2_pre_topc(X2)|~v1_xboole_0(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)))))).#trainneg
cnf(i_0_372, plain, (v1_xboole_0(X3)|v1_funct_1(k3_relat_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,X3,X5)|~v1_funct_2(X1,X4,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X5)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X3))))).#trainneg
cnf(i_0_429, plain, (v5_pre_topc(X1,X2,X3)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v1_funct_1(X1)|~v2_pre_topc(X3)|~v2_pre_topc(X2)|~v1_tdlat_3(X2)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)))))).#trainneg
cnf(i_0_81, plain, (v5_pre_topc(X1,X2,X3)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v1_funct_1(X1)|~v3_tops_2(X1,X2,X3)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)))))).#trainneg
cnf(i_0_359, plain, (v1_funct_1(k3_relat_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,X3,X3)|~v1_funct_2(X1,X3,X3)|~v3_funct_2(X2,X3,X3)|~v3_funct_2(X1,X3,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X3))))).#trainneg
cnf(i_0_358, plain, (v1_funct_2(k3_relat_1(X1,X2),X3,X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,X3,X3)|~v1_funct_2(X1,X3,X3)|~v3_funct_2(X2,X3,X3)|~v3_funct_2(X1,X3,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X3))))).#trainneg
cnf(i_0_357, plain, (v3_funct_2(k3_relat_1(X1,X2),X3,X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,X3,X3)|~v1_funct_2(X1,X3,X3)|~v3_funct_2(X2,X3,X3)|~v3_funct_2(X1,X3,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X3))))).#trainneg
cnf(i_0_393, plain, (v2_struct_0(X3)|v2_struct_0(X2)|v1_t_0topsp(X1,X2,X3)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v1_funct_1(X1)|~v3_tops_2(X1,X2,X3)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3)))))).#trainneg
cnf(i_0_242, plain, (v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X1),X2),X1,X1)|~l1_pre_topc(X1)|~v1_funct_1(X2)|~v3_tops_2(X2,X1,X1)|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1)))))).#trainneg
cnf(i_0_194, plain, (v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|v1_funct_1(k3_relat_1(X1,X2))|~l1_pre_topc(X5)|~l1_pre_topc(X4)|~l1_pre_topc(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_pre_topc(X5)|~v2_pre_topc(X4)|~v2_pre_topc(X3)|~v5_pre_topc(X2,X4,X5)|~v5_pre_topc(X1,X3,X4)|~v1_funct_2(X2,u1_struct_0(X4),u1_struct_0(X5))|~v1_funct_2(X1,u1_struct_0(X3),u1_struct_0(X4))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X5))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4)))))).#trainneg
cnf(i_0_192, plain, (v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|v5_pre_topc(k3_relat_1(X1,X2),X3,X4)|~l1_pre_topc(X5)|~l1_pre_topc(X4)|~l1_pre_topc(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_pre_topc(X5)|~v2_pre_topc(X4)|~v2_pre_topc(X3)|~v5_pre_topc(X2,X5,X4)|~v5_pre_topc(X1,X3,X5)|~v1_funct_2(X2,u1_struct_0(X5),u1_struct_0(X4))|~v1_funct_2(X1,u1_struct_0(X3),u1_struct_0(X5))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X4))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X5)))))).#trainneg
cnf(i_0_193, plain, (v2_struct_0(X5)|v2_struct_0(X4)|v2_struct_0(X3)|v1_funct_2(k3_relat_1(X1,X2),u1_struct_0(X3),u1_struct_0(X4))|~l1_pre_topc(X5)|~l1_pre_topc(X4)|~l1_pre_topc(X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v2_pre_topc(X5)|~v2_pre_topc(X4)|~v2_pre_topc(X3)|~v5_pre_topc(X2,X5,X4)|~v5_pre_topc(X1,X3,X5)|~v1_funct_2(X2,u1_struct_0(X5),u1_struct_0(X4))|~v1_funct_2(X1,u1_struct_0(X3),u1_struct_0(X5))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X4))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X5)))))).#trainneg
cnf(i_0_897, plain, (v2_struct_0(X1)|r2_borsuk_3(X1,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_523, plain, (v1_funct_1(esk32_0))).#trainneg
cnf(i_0_524, plain, (v1_funct_1(esk11_0))).#trainneg
cnf(i_0_514, plain, (esk32_0=k1_xboole_0)).#trainneg
cnf(i_0_515, plain, (esk11_0=k1_xboole_0)).#trainneg
cnf(i_0_517, plain, (esk7_0=k1_xboole_0)).#trainneg
cnf(i_0_511, plain, (~v1_xboole_0(esk38_0))).#trainneg
cnf(i_0_522, plain, (~v2_struct_0(esk5_0))).#trainneg
cnf(i_0_709, plain, (v1_xreal_0(np__2))).#trainneg
cnf(i_0_516, plain, (esk10_1(X1)=k1_xboole_0)).#trainneg
cnf(i_0_518, plain, (k1_xboole_0=esk6_0)).#trainneg
cnf(i_0_542, plain, (l2_algstr_0(X1)|~l1_rltopsp1(X1))).#trainneg
cnf(i_0_1016, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r2_borsuk_3(X2,X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~r1_t_0topsp(X1,X2))).#trainneg
cnf(i_0_1019, negated_conjecture, (v2_monoid_0(esk1_0)|~v2_monoid_0(k15_euclid(np__2))|~v2_pre_topc(k15_euclid(np__2))|~l1_pre_topc(k15_euclid(np__2)))).#trainneg
cnf(i_0_1018, negated_conjecture, (v2_monoid_0(esk2_0)|~v2_monoid_0(k15_euclid(np__2))|~v2_pre_topc(k15_euclid(np__2))|~l1_pre_topc(k15_euclid(np__2)))).#trainneg
cnf(i_0_850, plain, (v3_membered(k5_numbers))).#trainneg
cnf(i_0_662, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_529, plain, (l1_struct_0(esk44_0))).#trainneg
cnf(i_0_530, plain, (l1_struct_0(esk24_0))).#trainneg
cnf(i_0_533, plain, (v1_xxreal_0(esk42_0))).#trainneg
cnf(i_0_946, plain, (X1=u1_pre_topc(X2)|g1_pre_topc(X3,X1)!=X2|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X3)))|~v1_pre_topc(X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_614, plain, (v7_ordinal1(np__1))).#trainneg
cnf(i_0_552, plain, (~v2_xxreal_0(esk21_0))).#trainneg
cnf(i_0_553, plain, (~v2_xxreal_0(esk15_0))).#trainneg
cnf(i_0_547, plain, (X1=esk6_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_549, plain, (X1=esk36_2(X2,X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_591, plain, (l1_pre_topc(esk5_0)|~l1_pre_topc(k15_euclid(np__2)))).#trainneg
cnf(i_0_1057, plain, (~v1_xboole_0(k5_numbers))).#trainneg
cnf(i_0_637, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_704, plain, (v1_xreal_0(np__1))).#trainneg
cnf(i_0_841, plain, (v3_membered(esk27_0))).#trainneg
cnf(i_0_577, plain, (v1_funct_1(esk14_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_595, plain, (~v1_xboole_0(esk21_0))).#trainneg
cnf(i_0_596, plain, (~v1_xboole_0(esk15_0))).#trainneg
cnf(i_0_599, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_648, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_551, plain, (v1_funct_1(esk36_2(X1,X2)))).#trainneg
cnf(i_0_1519, plain, (r1_t_0topsp(X1,X1)|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X1),X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X1),X2),u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X1),X2))|~l1_pre_topc(X1)|~v3_tops_2(X2,X1,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1523, plain, (r1_t_0topsp(X1,X2)|v2_struct_0(X1)|~m1_subset_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(X2),u1_struct_0(X1),X3))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v3_tops_2(X3,X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1542, plain, (v2_struct_0(X1)|r1_t_0topsp(X1,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1550, negated_conjecture, (~r1_t_0topsp(esk2_0,esk1_0)|~l1_pre_topc(esk1_0)|~l1_pre_topc(esk2_0))).#trainneg
cnf(i_0_1547, plain, (l1_algstr_0(X1)|~l1_rltopsp1(X1))).#trainneg
cnf(i_0_1481, plain, (r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(k15_euclid(np__2),X1))|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~v1_funct_2(esk4_1(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_funct_1(esk4_1(X1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_517, plain, (esk7_0=esk6_0)).#trainneg
cnf(i_0_571, plain, (v1_funct_1(esk50_0)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_515, plain, (esk11_0=esk6_0)).#trainneg
cnf(i_0_1546, plain, (v1_xxreal_0(np__2))).#trainneg
cnf(i_0_1548, plain, (l2_struct_0(X1)|~l1_rltopsp1(X1))).#trainneg
cnf(i_0_573, plain, (v1_funct_1(esk34_1(X1))|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_574, plain, (v1_funct_1(esk31_0)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1480, plain, (r1_t_0topsp(X1,X1)|~m1_subset_1(esk16_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~v1_funct_2(esk16_1(X1),u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(esk16_1(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1432, plain, (v2_struct_0(X1)|r1_t_0topsp(X2,X1)|~v2_pre_topc(X1)|~v2_pre_topc(X2)|~r1_t_0topsp(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)),X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v1_pre_topc(X1))).#trainneg
cnf(i_0_1433, plain, (v2_struct_0(X1)|r1_t_0topsp(X2,X1)|~v2_pre_topc(X1)|~v2_pre_topc(X2)|~r1_t_0topsp(X2,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v1_pre_topc(X2))).#trainneg
cnf(i_0_1353, plain, (k1_pre_topc(X1,X2)=k1_pre_topc(X1,X3)|k2_struct_0(k1_pre_topc(X1,X2))!=X3|~v1_pre_topc(k1_pre_topc(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1350, negated_conjecture, (k1_pre_topc(k15_euclid(np__2),X1)=esk2_0|k2_struct_0(esk2_0)!=X1|~v1_pre_topc(esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(k15_euclid(np__2)))).#trainneg
cnf(i_0_1307, plain, (v1_topreal2(X1)|~m1_subset_1(esk3_2(k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(k15_euclid(np__2),X1)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(esk3_2(k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(k15_euclid(np__2),X1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_funct_1(esk3_2(k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(k15_euclid(np__2),X1)))|~r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(k15_euclid(np__2),X1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))).#trainneg
cnf(i_0_1529, plain, (v1_topreal2(X1)|v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|v2_struct_0(X2)|v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|~m1_subset_1(k1_partfun1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),X3,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(k1_partfun1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),X3,X4),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_funct_1(k1_partfun1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),X3,X4))|~v3_tops_2(X4,X2,k1_pre_topc(k15_euclid(np__2),X1))|~v3_tops_2(X3,k1_pre_topc(k15_euclid(np__2),k1_topreal1),X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(X2))))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_funct_2(X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|~l1_pre_topc(X2)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))).#trainneg
cnf(i_0_1645, plain, (~v7_ordinal1(X1)|~v1_xboole_0(u1_struct_0(k15_euclid(X1)))|~l1_pre_topc(k15_euclid(X1)))).#trainneg
cnf(i_0_1522, plain, (v1_topreal2(X1)|v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~m1_subset_1(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_funct_1(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2))|~v3_tops_2(X2,k1_pre_topc(k15_euclid(np__2),X1),k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~v1_funct_2(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_funct_1(X2)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1)))).#trainneg
cnf(i_0_1552, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r1_t_0topsp(X1,X2)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~r1_t_0topsp(X2,X1))).#trainneg
cnf(i_0_169, plain, (v1_xxreal_0(k1_xboole_0))).#trainneg
cnf(i_0_122, plain, (v1_xcmplx_0(k1_xboole_0))).#trainneg
cnf(i_0_514, plain, (esk32_0=esk6_0)).#trainneg
cnf(i_0_152, plain, (o_0_0_xboole_0=esk6_0)).#trainneg
cnf(i_0_576, plain, (v1_funct_1(esk27_0)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_153, plain, (m1_subset_1(esk6_0,k5_numbers))).#trainneg
cnf(i_0_519, plain, (~m1_subset_1(esk21_0,k5_numbers))).#trainneg
cnf(i_0_643, plain, (X1=k2_struct_0(X2)|~v1_xboole_0(X1)|~l1_struct_0(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_578, plain, (v1_funct_1(esk13_0)|~v4_funct_1(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_580, plain, (v1_funct_1(esk9_1(X1))|~v4_funct_1(X1))).#trainneg
cnf(i_0_1351, negated_conjecture, (k1_pre_topc(k15_euclid(np__2),X1)=esk1_0|k2_struct_0(esk1_0)!=X1|~v1_pre_topc(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(k15_euclid(np__2)))).#trainneg
cnf(i_0_1349, plain, (k1_pre_topc(k15_euclid(np__2),X1)=esk5_0|k2_struct_0(esk5_0)!=X1|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(k15_euclid(np__2)))).#trainneg
cnf(i_0_1431, plain, (v2_struct_0(X1)|r1_t_0topsp(X2,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X1)|~v2_pre_topc(X2)|~r1_t_0topsp(X2,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v1_pre_topc(X2))).#trainneg
cnf(i_0_1201, plain, (m1_subset_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_toprealb(k1_pre_topc(k15_euclid(np__2),X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(k15_euclid(np__2)))).#trainneg
cnf(i_0_1194, plain, (v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2))).#trainneg
cnf(i_0_1192, plain, (m1_subset_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_toprealb(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(k15_euclid(np__2))|~m1_pre_topc(X1,k15_euclid(np__2)))).#trainneg
cnf(i_0_1205, plain, (v1_pre_topc(g1_pre_topc(u1_struct_0(k1_pre_topc(X1,X2)),u1_pre_topc(k1_pre_topc(X1,X2))))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1430, plain, (v2_struct_0(X1)|r1_t_0topsp(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)),X1)|~v2_pre_topc(X1)|~v2_pre_topc(X2)|~r1_t_0topsp(X2,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v1_pre_topc(X1))).#trainneg
cnf(i_0_1642, plain, (v2_struct_0(X1)|~l1_pre_topc(X1)|~v1_xboole_0(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))).#trainneg
cnf(i_0_1191, plain, (v1_topreal2(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))|~v1_toprealb(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(k15_euclid(np__2))|~m1_pre_topc(X1,k15_euclid(np__2)))).#trainneg
cnf(i_0_120, plain, (v1_xreal_0(k1_xboole_0))).#trainneg
cnf(i_0_1647, plain, (v1_xxreal_0(np__1))).#trainneg
cnf(i_0_516, plain, (esk10_1(X1)=esk6_0)).#trainneg
cnf(i_0_583, plain, (v1_funct_1(np__1)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_584, plain, (v1_funct_1(np__1)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_520, plain, (~m1_subset_1(esk15_0,k5_numbers))).#trainneg
cnf(i_0_1583, plain, (~v1_xboole_0(u1_struct_0(esk24_0)))).#trainneg
cnf(i_0_710, plain, (X1=esk12_1(X2)|~v1_xboole_0(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1591, plain, (k2_zfmisc_1(X1,X2)=esk6_0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_586, plain, (v1_funct_1(k5_numbers)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_587, plain, (v1_funct_1(k1_topreal1)|~v4_funct_1(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_1641, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))|~l1_pre_topc(esk1_0))).#trainneg
cnf(i_0_1521, plain, (v2_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X1),X2),u1_struct_0(X1))|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X1),X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X1),X2),u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X1),X2))|~l1_pre_topc(X1)|~v3_tops_2(X2,X1,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1524, plain, (v2_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|v2_struct_0(X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v3_tops_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1525, plain, (v2_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X1))|v2_struct_0(X2)|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X2),X3))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v3_tops_2(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1200, plain, (v1_topreal2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_toprealb(k1_pre_topc(k15_euclid(np__2),X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(k15_euclid(np__2)))).#trainneg
cnf(i_0_1436, plain, (v2_funct_1(esk4_1(X1))|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~v1_funct_2(esk4_1(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_funct_1(esk4_1(X1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_1497, plain, (v5_pre_topc(X1,X2,X3)|~v2_pre_topc(X3)|~v2_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~v2_struct_0(X2))).#trainneg
cnf(i_0_1520, plain, (v2_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X1),X2))|~m1_subset_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X1),X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X1),X2),u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),u1_struct_0(X1),X2))|~l1_pre_topc(X1)|~v3_tops_2(X2,X1,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1533, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v3_tops_2(k3_relat_1(X4,X5),X1,X3)|~v3_tops_2(X5,X2,X3)|~v3_tops_2(X4,X1,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1196, plain, (v1_pre_topc(g1_pre_topc(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))))|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2))).#trainneg
cnf(i_0_1532, plain, (v2_funct_2(k1_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X3),X4,X5),u1_struct_0(X3))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(k1_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X3),X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~v1_funct_2(k1_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X3),X4,X5),u1_struct_0(X1),u1_struct_0(X3))|~v1_funct_1(k1_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X3),X4,X5))|~l1_pre_topc(X1)|~l1_pre_topc(X3)|~v3_tops_2(X5,X2,X3)|~v3_tops_2(X4,X1,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_588, plain, (v1_funct_1(np__2)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_631, plain, (v3_valued_0(esk20_1(k1_numbers)))).#trainneg
cnf(i_0_163, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_589, plain, (v1_funct_1(np__2)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_608, plain, (v7_ordinal1(esk14_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1634, plain, (~v1_xboole_0(u1_struct_0(esk41_0)))).#trainneg
cnf(i_0_1635, plain, (~v1_xboole_0(u1_struct_0(esk35_0)))).#trainneg
cnf(i_0_1592, plain, (k2_zfmisc_1(X1,X2)=esk6_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_590, plain, (v1_xxreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_594, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k4_numbers))).#trainneg
cnf(i_0_1531, plain, (v2_funct_1(k1_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X3),X4,X5))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(k1_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X3),X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~v1_funct_2(k1_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X3),X4,X5),u1_struct_0(X1),u1_struct_0(X3))|~v1_funct_1(k1_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X3),X4,X5))|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v3_tops_2(X5,X2,X3)|~v3_tops_2(X4,X1,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1195, plain, (v2_monoid_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_monoid_0(X2)|~v2_pre_topc(X2)|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2))).#trainneg
cnf(i_0_1204, plain, (v2_monoid_0(k1_pre_topc(X1,X2))|~v2_monoid_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1640, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk2_0))|~l1_pre_topc(esk2_0))).#trainneg
cnf(i_0_1163, plain, (v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_xboole_0(k2_struct_0(X1))|~m1_subset_1(k2_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1190, plain, (~v2_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v1_toprealb(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(k15_euclid(np__2))|~m1_pre_topc(X1,k15_euclid(np__2)))).#trainneg
cnf(i_0_1584, plain, (u1_pre_topc(g1_pre_topc(X2,X1))=X1|~v1_pre_topc(g1_pre_topc(X2,X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~l1_pre_topc(g1_pre_topc(X2,X1)))).#trainneg
cnf(i_0_632, plain, (v3_valued_0(esk19_1(k1_numbers)))).#trainneg
cnf(i_0_1648, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_598, plain, (v1_tops_1(u1_struct_0(X1),X1)|~l1_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_680, plain, (v1_xcmplx_0(np__1)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_1637, plain, (~v1_xboole_0(u1_struct_0(esk23_0)))).#trainneg
cnf(i_0_1638, plain, (~v1_xboole_0(u1_struct_0(esk22_0)))).#trainneg
cnf(i_0_1597, plain, (k2_zfmisc_1(X1,X2)=esk36_2(X3,X4)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_601, plain, (v1_funct_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_602, plain, (v7_ordinal1(esk50_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1639, plain, (~v1_xboole_0(u1_struct_0(esk5_0))|~l1_pre_topc(esk5_0))).#trainneg
cnf(i_0_1484, plain, (v2_funct_2(esk4_1(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~v1_funct_2(esk4_1(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_funct_1(esk4_1(X1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_1644, plain, (v1_xboole_0(X1)|v2_struct_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~v1_xboole_0(u1_struct_0(k1_pre_topc(X2,X1)))|~l1_pre_topc(k1_pre_topc(X2,X1)))).#trainneg
cnf(i_0_1600, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk2_0))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_1282, plain, (v1_funct_2(esk16_1(X1),u1_struct_0(X1),X2)|~m1_subset_1(esk16_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),X2)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1185, plain, (v1_xboole_0(esk4_1(X1))|~v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_1590, plain, (k2_struct_0(X1)=esk6_0|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1587, plain, (l1_pre_topc(k15_euclid(np__1)))).#trainneg
cnf(i_0_681, plain, (v1_xcmplx_0(np__1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1267, plain, (v2_funct_1(esk19_1(X1)))).#trainneg
cnf(i_0_862, plain, (v1_membered(esk27_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_824, plain, (~v1_xboole_0(u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_603, plain, (v7_ordinal1(esk47_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_685, plain, (v1_xcmplx_0(np__2)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_1598, plain, (k2_zfmisc_1(X1,X2)=esk36_2(X3,X4)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_604, plain, (v7_ordinal1(esk34_1(X1))|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_605, plain, (v7_ordinal1(esk31_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1317, plain, (v1_partfun1(esk4_1(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_1382, plain, (v1_partfun1(k2_tops_2(X1,X1,X2),X1)|~m1_subset_1(k2_tops_2(X1,X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_2(X2,X1,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1189, plain, (v1_toprealb(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v1_topreal2(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))|~m1_subset_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(k15_euclid(np__2))|~m1_pre_topc(X1,k15_euclid(np__2)))).#trainneg
cnf(i_0_1485, plain, (v2_funct_2(esk3_2(X1,X2),u1_struct_0(X2))|~m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk3_2(X1,X2),u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(esk3_2(X1,X2))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~r1_t_0topsp(X1,X2))).#trainneg
cnf(i_0_1274, plain, (v1_partfun1(esk4_1(k1_topreal1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~m1_subset_1(esk4_1(k1_topreal1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))))))).#trainneg
cnf(i_0_1603, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(X1,X2)))|~l1_pre_topc(k1_pre_topc(X1,X2))|~v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_686, plain, (v1_xcmplx_0(np__2)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1601, plain, (v1_xboole_0(u1_struct_0(esk33_0)))).#trainneg
cnf(i_0_871, plain, (v1_membered(k5_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_606, plain, (v7_ordinal1(esk29_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_736, plain, (v1_relat_1(esk14_0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_819, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_827, plain, (~v1_xboole_0(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_1653, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X4)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_607, plain, (v7_ordinal1(esk27_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_609, plain, (v7_ordinal1(esk13_0)|~v6_membered(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_1171, plain, (v1_xboole_0(esk4_1(X1))|~v1_xboole_0(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_1429, plain, (v1_funct_1(k3_relat_1(X5,X6))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~v1_funct_1(X5))).#trainneg
cnf(i_0_1157, plain, (v1_pre_topc(g1_pre_topc(u1_struct_0(esk5_0),u1_pre_topc(esk5_0)))|~l1_pre_topc(k15_euclid(np__2)))).#trainneg
cnf(i_0_988, plain, (v2_struct_0(X1)|~v2_tops_1(u1_struct_0(X1),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_1633, negated_conjecture, (m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_1539, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v1_funct_2(k3_relat_1(X4,X5),u1_struct_0(X1),u1_struct_0(X2))|~v5_pre_topc(X4,X1,X3)|~v2_pre_topc(X3)|~v2_pre_topc(X2)|~v2_pre_topc(X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v3_tops_2(X5,X3,X2))).#trainneg
cnf(i_0_743, plain, (v1_relat_1(np__1)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_1602, plain, (v1_xboole_0(u1_struct_0(esk17_0)))).#trainneg
cnf(i_0_884, plain, (v2_membered(esk27_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_611, plain, (v7_ordinal1(esk9_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_744, plain, (v1_relat_1(np__1)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_1654, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X4)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_612, plain, (v7_ordinal1(esk8_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_617, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_591, plain, (l1_pre_topc(esk5_0))).#trainneg
cnf(i_0_1684, plain, (v2_struct_0(k1_pre_topc(X1,X2))|~v2_tops_1(X2,k1_pre_topc(X1,X2))|~v2_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1692, plain, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X1),k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~v1_funct_2(esk4_1(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_funct_1(esk4_1(X1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_1693, plain, (v2_struct_0(X1)|r1_t_0topsp(X1,X2)|~m1_subset_1(esk3_2(X2,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(esk3_2(X2,X1),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(esk3_2(X2,X1))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~r1_t_0topsp(X2,X1))).#trainneg
cnf(i_0_1755, plain, (v7_ordinal1(esk6_0))).#trainneg
cnf(i_0_2043, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_618, plain, (v7_ordinal1(k1_topreal1)|~v6_membered(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_2108, plain, (v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_1699, plain, (l1_algstr_0(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1700, plain, (l2_struct_0(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_628, plain, (v3_valued_0(esk36_2(X1,k1_numbers)))).#trainneg
cnf(i_0_630, plain, (v3_valued_0(esk25_2(X1,k1_numbers)))).#trainneg
cnf(i_0_1710, plain, (v2_struct_0(X1)|r1_t_0topsp(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),X1)|~v2_pre_topc(X1)|~v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1)|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))).#trainneg
cnf(i_0_1682, plain, (m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2))))|~l1_struct_0(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1810, plain, (v1_funct_2(esk4_1(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(esk5_0))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|k2_struct_0(esk5_0)!=X1)).#trainneg
cnf(i_0_2123, plain, (v1_funct_2(k2_tops_2(X1,X1,X2),X1,X3)|~m1_subset_1(k2_tops_2(X1,X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~v1_funct_2(X2,X1,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_1812, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(esk5_0))))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|k2_struct_0(esk5_0)!=X1)).#trainneg
cnf(i_0_2119, plain, (v1_funct_2(esk4_1(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2)|v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2)))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_748, plain, (v1_relat_1(np__2)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_1972, plain, (v4_relat_1(esk6_0,X1))).#trainneg
cnf(i_0_2172, plain, (u1_struct_0(esk33_0)=esk6_0)).#trainneg
cnf(i_0_634, plain, (v2_struct_0(esk54_0)|~v1_xboole_0(u1_struct_0(esk54_0)))).#trainneg
cnf(i_0_749, plain, (v1_relat_1(np__2)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_1019, negated_conjecture, (v2_monoid_0(esk1_0)|~v2_monoid_0(k15_euclid(np__2))|~v2_pre_topc(k15_euclid(np__2)))).#trainneg
cnf(i_0_1018, negated_conjecture, (v2_monoid_0(esk2_0)|~v2_monoid_0(k15_euclid(np__2))|~v2_pre_topc(k15_euclid(np__2)))).#trainneg
cnf(i_0_635, plain, (v2_struct_0(esk48_0)|~v1_xboole_0(u1_struct_0(esk48_0)))).#trainneg
cnf(i_0_636, plain, (v2_struct_0(esk45_0)|~v1_xboole_0(u1_struct_0(esk45_0)))).#trainneg
cnf(i_0_1703, plain, (v2_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|r1_t_0topsp(X1,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X1)|~v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1696, plain, (v2_struct_0(X1)|r1_t_0topsp(X1,X2)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(k1_partfun1(u1_struct_0(X2),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X1),X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k1_partfun1(u1_struct_0(X2),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X1),X4,X5),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(k1_partfun1(u1_struct_0(X2),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X1),X4,X5))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v3_tops_2(X5,X3,X1)|~v3_tops_2(X4,X2,X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_1809, plain, (v3_tops_2(esk4_1(X1),k1_pre_topc(k15_euclid(np__2),k1_topreal1),esk5_0)|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|k2_struct_0(esk5_0)!=X1)).#trainneg
cnf(i_0_1802, plain, (v1_funct_2(esk4_1(X1),u1_struct_0(esk5_0),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|k2_struct_0(esk5_0)!=k1_topreal1)).#trainneg
cnf(i_0_1801, plain, (v3_tops_2(esk4_1(X1),esk5_0,k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|k2_struct_0(esk5_0)!=k1_topreal1)).#trainneg
cnf(i_0_1747, plain, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|v1_topreal2(X1)|~m1_subset_1(esk3_2(k1_pre_topc(k15_euclid(np__2),X1),k1_pre_topc(k15_euclid(np__2),k1_topreal1)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(esk3_2(k1_pre_topc(k15_euclid(np__2),X1),k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_funct_1(esk3_2(k1_pre_topc(k15_euclid(np__2),X1),k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|~r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X1),k1_pre_topc(k15_euclid(np__2),k1_topreal1)))).#trainneg
cnf(i_0_1977, plain, (v5_relat_1(esk6_0,X1))).#trainneg
cnf(i_0_2239, plain, (u1_struct_0(esk17_0)=esk6_0)).#trainneg
cnf(i_0_1712, plain, (k1_pre_topc(X1,X2)=k1_pre_topc(X1,X3)|u1_struct_0(k1_pre_topc(X1,X2))!=X3|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~l1_struct_0(k1_pre_topc(X1,X2)))).#trainneg
cnf(i_0_639, plain, (v1_xboole_0(u1_struct_0(esk54_0))|~v2_struct_0(esk54_0))).#trainneg
cnf(i_0_1983, plain, (v3_membered(esk6_0)|~v3_membered(X1))).#trainneg
cnf(i_0_1984, plain, (v1_membered(esk6_0)|~v1_membered(X1))).#trainneg
cnf(i_0_640, plain, (v1_xboole_0(u1_struct_0(esk48_0))|~v2_struct_0(esk48_0))).#trainneg
cnf(i_0_641, plain, (v1_xboole_0(u1_struct_0(esk45_0))|~v2_struct_0(esk45_0))).#trainneg
cnf(i_0_1808, plain, (r1_t_0topsp(esk5_0,k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(esk5_0)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|k2_struct_0(esk5_0)!=k1_topreal1)).#trainneg
cnf(i_0_1724, negated_conjecture, (v3_tops_2(esk4_1(X1),esk2_0,k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|k2_struct_0(esk2_0)!=k1_topreal1|~v1_pre_topc(esk2_0))).#trainneg
cnf(i_0_1677, plain, (v1_xboole_0(X2)|~l1_struct_0(k1_pre_topc(X1,X2))|~v2_struct_0(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1813, plain, (v1_topreal2(X1)|~v3_tops_2(X2,k1_pre_topc(k15_euclid(np__2),k1_topreal1),esk5_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(esk5_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(esk5_0))|~v1_funct_1(X2)|k2_struct_0(esk5_0)!=X1)).#trainneg
cnf(i_0_1900, plain, (v2_struct_0(X1)|r1_t_0topsp(X2,X1)|v2_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~r1_t_0topsp(X2,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))).#trainneg
cnf(i_0_1897, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(g1_pre_topc(X3,X4))|r1_t_0topsp(g1_pre_topc(X3,X4),X1)|~v3_tops_2(X5,X2,X1)|~v3_tops_2(X6,g1_pre_topc(X3,X4),X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(X2))))|~v1_funct_2(k1_partfun1(X3,u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X1),X6,X5),X3,u1_struct_0(X1))|~v1_funct_2(X5,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_2(X6,X3,u1_struct_0(X2))|~v1_funct_1(X5)|~v1_funct_1(X6)|~l1_pre_topc(X1)|~l1_pre_topc(g1_pre_topc(X3,X4))|~l1_pre_topc(X2)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3))))).#trainneg
cnf(i_0_2228, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v1_funct_2(k3_relat_1(X4,X5),u1_struct_0(X3),u1_struct_0(X2))|~v2_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X3)|~v3_tops_2(X5,X1,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~l1_pre_topc(X3)|~v3_tops_2(X4,X3,X1))).#trainneg
cnf(i_0_893, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1320, plain, (v2_funct_2(esk19_1(X1),X1))).#trainneg
cnf(i_0_2018, plain, (~v2_struct_0(g1_pre_topc(X1,u1_pre_topc(g1_pre_topc(X1,X2))))|~m1_pre_topc(g1_pre_topc(X1,X2),k15_euclid(np__2))|~v1_toprealb(g1_pre_topc(X1,u1_pre_topc(g1_pre_topc(X1,X2))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_645, plain, (v1_funct_1(k2_struct_0(X1))|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_664, plain, (v1_xcmplx_0(esk50_0)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1985, plain, (v2_membered(esk6_0)|~v2_membered(X1))).#trainneg
cnf(i_0_2095, plain, (u1_struct_0(X1)=esk6_0|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_650, plain, (v1_funct_1(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_652, plain, (v7_ordinal1(esk19_1(X1))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_1898, plain, (v2_struct_0(X1)|r1_t_0topsp(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)),X1)|v2_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X1)|~v2_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~v1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~l1_pre_topc(X1)|~l1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X2)|~v1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~r1_t_0topsp(X2,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1775, negated_conjecture, (v3_tops_2(esk4_1(X1),esk1_0,k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|k2_struct_0(esk1_0)!=k1_topreal1|~v1_pre_topc(esk1_0))).#trainneg
cnf(i_0_1709, plain, (v2_struct_0(X1)|r1_t_0topsp(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)),X1)|~v2_pre_topc(X1)|~v2_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~v1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~l1_pre_topc(X1)|~l1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~v2_pre_topc(X2)|~r1_t_0topsp(X2,X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1820, plain, (v2_struct_0(X1)|r1_t_0topsp(X2,X1)|~v2_pre_topc(X1)|~v2_pre_topc(X2)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~v1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~r1_t_0topsp(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)),X1)|~l1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2))))).#trainneg
cnf(i_0_1909, plain, (l1_pre_topc(g1_pre_topc(u1_struct_0(esk5_0),u1_pre_topc(esk5_0))))).#trainneg
cnf(i_0_1711, plain, (k1_pre_topc(X1,k2_struct_0(k1_pre_topc(X1,X2)))=k1_pre_topc(X1,X2)|~m1_subset_1(k2_struct_0(k1_pre_topc(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1804, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|k2_struct_0(esk5_0)!=k1_topreal1)).#trainneg
cnf(i_0_1849, plain, (v2_struct_0(g1_pre_topc(X1,X2))|v3_tops_2(k2_tops_2(u1_struct_0(X3),X1,X4),g1_pre_topc(X1,X2),X3)|~v3_tops_2(X4,X3,g1_pre_topc(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X1)))|~v1_funct_2(X4,u1_struct_0(X3),X1)|~v1_funct_1(X4)|~l1_pre_topc(g1_pre_topc(X1,X2))|~l1_pre_topc(X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_668, plain, (v1_xcmplx_0(esk31_0)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1564, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1581, plain, (v2_struct_0(esk44_0)|~v1_xboole_0(u1_struct_0(esk44_0)))).#trainneg
cnf(i_0_2024, plain, (~v2_struct_0(g1_pre_topc(u1_struct_0(g1_pre_topc(X1,X2)),X2))|~m1_pre_topc(g1_pre_topc(X1,X2),k15_euclid(np__2))|~v1_toprealb(g1_pre_topc(u1_struct_0(g1_pre_topc(X1,X2)),X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_655, plain, (v7_ordinal1(esk20_1(X1))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_990, plain, (v1_funct_1(esk50_0)|~v1_relat_1(k1_numbers)|~v1_funct_1(k1_numbers))).#trainneg
cnf(i_0_670, plain, (v1_xcmplx_0(esk27_0)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_683, plain, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1997, plain, (v1_xxreal_0(X1)|~m1_subset_1(X1,k4_numbers))).#trainneg
cnf(i_0_657, plain, (v1_funct_1(esk49_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_658, plain, (v7_ordinal1(esk49_1(X1))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1728, negated_conjecture, (r1_t_0topsp(esk2_0,k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(esk2_0)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|k2_struct_0(esk2_0)!=k1_topreal1|~v1_pre_topc(esk2_0))).#trainneg
cnf(i_0_1850, plain, (v2_struct_0(g1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(X4)|v3_tops_2(k1_partfun1(X1,u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X4),X5,X6),g1_pre_topc(X1,X2),X4)|~v3_tops_2(X6,X3,X4)|~v3_tops_2(X5,g1_pre_topc(X1,X2),X3)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X3))))|~v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_2(X5,X1,u1_struct_0(X3))|~v1_funct_1(X6)|~v1_funct_1(X5)|~l1_pre_topc(X4)|~l1_pre_topc(X3)|~l1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1896, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(g1_pre_topc(X3,X4))|v3_tops_2(k1_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X2),X3,X5,X6),X1,g1_pre_topc(X3,X4))|~v3_tops_2(X6,X2,g1_pre_topc(X3,X4))|~v3_tops_2(X5,X1,X2)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),X3)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X2),X3)|~v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X6)|~v1_funct_1(X5)|~l1_pre_topc(g1_pre_topc(X3,X4))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3))))).#trainneg
cnf(i_0_1879, plain, (v2_struct_0(g1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(X4)|r1_t_0topsp(X4,g1_pre_topc(X1,X2))|~v3_tops_2(X5,X3,g1_pre_topc(X1,X2))|~v3_tops_2(X6,X4,X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X1)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X3))))|~v1_funct_2(k1_partfun1(u1_struct_0(X4),u1_struct_0(X3),u1_struct_0(X3),X1,X6,X5),u1_struct_0(X4),X1)|~v1_funct_2(X5,u1_struct_0(X3),X1)|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X6)|~l1_pre_topc(g1_pre_topc(X1,X2))|~l1_pre_topc(X4)|~l1_pre_topc(X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1949, plain, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|v2_struct_0(X2)|v3_tops_2(k3_relat_1(X3,esk4_1(X1)),X2,k1_pre_topc(k15_euclid(np__2),X1))|~v3_tops_2(X3,X2,k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~v1_funct_2(esk4_1(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_funct_1(esk4_1(X1))|~v1_funct_1(X3)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(X2)|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_1831, plain, (v2_pre_topc(g1_pre_topc(u1_struct_0(k1_pre_topc(X1,X2)),u1_pre_topc(k1_pre_topc(X1,X2))))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1782, negated_conjecture, (r1_t_0topsp(esk1_0,k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(esk1_0)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|k2_struct_0(esk1_0)!=k1_topreal1|~v1_pre_topc(esk1_0))).#trainneg
cnf(i_0_2583, plain, (v3_membered(esk6_0))).#trainneg
cnf(i_0_2589, plain, (v1_membered(esk6_0))).#trainneg
cnf(i_0_1550, negated_conjecture, (~r1_t_0topsp(esk2_0,esk1_0)|~l1_pre_topc(esk2_0))).#trainneg
cnf(i_0_660, plain, (v1_funct_1(u2_struct_0(X1))|~v4_funct_1(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_2358, plain, (v1_tops_1(esk6_0,esk33_0))).#trainneg
cnf(i_0_2432, negated_conjecture, (v2_monoid_0(esk1_0)|~v2_pre_topc(k15_euclid(np__2)))).#trainneg
cnf(i_0_2433, negated_conjecture, (v2_monoid_0(esk2_0)|~v2_pre_topc(k15_euclid(np__2)))).#trainneg
cnf(i_0_661, plain, (v7_ordinal1(u2_struct_0(X1))|~v6_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_665, plain, (v1_xcmplx_0(esk49_1(X1))|v1_xboole_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2290, plain, (v2_struct_0(X1)|r1_t_0topsp(X1,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_2283, plain, (l1_struct_0(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2608, plain, (X1=k1_pre_topc(X2,X3)|u1_struct_0(X1)!=X3|~l1_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2)|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2438, plain, (v2_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|r1_t_0topsp(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),X1)|~l1_pre_topc(X1)|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X1))).#trainneg
cnf(i_0_2299, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k1_pre_topc(X2,X1))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~l1_pre_topc(k1_pre_topc(X2,X1)))).#trainneg
cnf(i_0_2707, plain, (v2_struct_0(X1)|r1_t_0topsp(X1,g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~l1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v2_pre_topc(X2)|~r1_t_0topsp(X2,X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_2672, plain, (v2_membered(esk6_0))).#trainneg
cnf(i_0_2280, plain, (l1_pre_topc(k15_euclid(esk6_0)))).#trainneg
cnf(i_0_2501, plain, (v1_tops_1(esk6_0,esk17_0))).#trainneg
cnf(i_0_2388, plain, (g1_pre_topc(esk6_0,u1_pre_topc(esk33_0))=k1_pre_topc(X1,X2)|k2_struct_0(g1_pre_topc(esk6_0,u1_pre_topc(esk33_0)))!=X2|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(esk33_0,X1))).#trainneg
cnf(i_0_666, plain, (v1_xcmplx_0(esk47_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_997, plain, (v1_funct_1(esk31_0)|~v1_relat_1(k1_numbers)|~v1_funct_1(k1_numbers))).#trainneg
cnf(i_0_2367, plain, (~v1_subset_1(k2_struct_0(esk33_0),esk6_0)|~l1_struct_0(esk33_0))).#trainneg
cnf(i_0_2510, plain, (~v1_subset_1(k2_struct_0(esk17_0),esk6_0)|~l1_struct_0(esk17_0))).#trainneg
cnf(i_0_667, plain, (v1_xcmplx_0(esk34_1(X1))|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_669, plain, (v1_xcmplx_0(esk29_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2723, plain, (v2_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|r1_t_0topsp(X2,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|v2_struct_0(X1)|~v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X2)|~l1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~r1_t_0topsp(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2876, plain, (r1_t_0topsp(X1,k1_pre_topc(k15_euclid(np__2),X2))|v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|v2_struct_0(k1_pre_topc(k15_euclid(np__2),X2))|v2_struct_0(X1)|~m1_subset_1(k3_relat_1(X3,esk4_1(X2)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X2)))))|~v1_funct_2(k3_relat_1(X3,esk4_1(X2)),u1_struct_0(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X2)))|~v1_funct_1(k3_relat_1(X3,esk4_1(X2)))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X2))|~l1_pre_topc(X1)|~v1_topreal2(X2)|~v3_tops_2(X3,X1,k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_funct_1(X3)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))).#trainneg
cnf(i_0_2456, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|r1_t_0topsp(X3,X2)|~v3_tops_2(X4,X1,X3)|~v3_tops_2(X5,X2,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X3))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k3_relat_1(X5,X4),u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X3))|~v1_funct_2(X5,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X4)|~v1_funct_1(X5)|~l1_pre_topc(X2)|~l1_pre_topc(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2278, plain, (v2_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|r1_t_0topsp(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),X2)|v2_struct_0(X1)|~l1_pre_topc(X2)|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X1)|~v2_pre_topc(X2)|~v1_pre_topc(X2)|~r1_t_0topsp(X2,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2682, plain, (u1_struct_0(X1)=esk6_0|~v2_struct_0(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_687, plain, (v1_xreal_0(esk50_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2360, plain, (v4_relat_1(esk16_1(esk33_0),esk6_0))).#trainneg
cnf(i_0_1106, plain, (v1_valued_0(esk50_0)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2532, plain, (g1_pre_topc(esk6_0,u1_pre_topc(esk17_0))=k1_pre_topc(X1,X2)|k2_struct_0(g1_pre_topc(esk6_0,u1_pre_topc(esk17_0)))!=X2|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(esk17_0,X1))).#trainneg
cnf(i_0_671, plain, (v1_xcmplx_0(esk20_1(X1))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_999, plain, (v1_funct_1(esk27_0)|~v1_relat_1(k1_numbers)|~v1_funct_1(k1_numbers))).#trainneg
cnf(i_0_691, plain, (v1_xreal_0(esk31_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_693, plain, (v1_xreal_0(esk27_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2669, plain, (v1_funct_1(u1_struct_0(X1))|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_1660, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_672, plain, (v1_xcmplx_0(esk19_1(X1))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_674, plain, (v1_xcmplx_0(esk13_0)|~v1_membered(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_2721, plain, (v2_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|r1_t_0topsp(X2,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|v2_struct_0(X1)|~v2_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X2)|~l1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~v1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~r1_t_0topsp(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)),X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2621, plain, (v1_xboole_0(X1)|~v2_struct_0(k1_pre_topc(X2,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~l1_pre_topc(k1_pre_topc(X2,X1)))).#trainneg
cnf(i_0_2881, plain, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X1),X2)|v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|v2_struct_0(X2)|~m1_subset_1(k3_relat_1(X3,esk4_1(X1)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~v1_funct_2(k3_relat_1(X3,esk4_1(X1)),u1_struct_0(X2),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_funct_1(k3_relat_1(X3,esk4_1(X1)))|~l1_pre_topc(X2)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(X1)|~v3_tops_2(X3,X2,k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_funct_1(X3)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))).#trainneg
cnf(i_0_706, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2361, plain, (v5_relat_1(esk16_1(esk33_0),esk6_0))).#trainneg
cnf(i_0_1113, plain, (v1_valued_0(esk31_0)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2368, plain, (~v2_struct_0(g1_pre_topc(esk6_0,u1_pre_topc(esk33_0)))|~m1_pre_topc(esk33_0,k15_euclid(np__2))|~v1_toprealb(g1_pre_topc(esk6_0,u1_pre_topc(esk33_0))))).#trainneg
cnf(i_0_676, plain, (v1_xcmplx_0(esk9_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_1009, plain, (v1_funct_1(k5_numbers)|~v1_relat_1(k1_numbers)|~v1_funct_1(k1_numbers))).#trainneg
cnf(i_0_725, plain, (v1_relat_1(esk50_0)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_730, plain, (v1_relat_1(esk31_0)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1650, plain, (esk12_1(X1)=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X3)|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1652, plain, (k2_struct_0(X1)=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X3)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_677, plain, (v1_xcmplx_0(esk8_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_679, plain, (v1_xcmplx_0(u2_struct_0(X1))|~v1_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_2274, plain, (v2_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|r1_t_0topsp(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|v2_struct_0(X1)|~l1_pre_topc(g1_pre_topc(u1_struct_0(X2),u1_pre_topc(X2)))|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X1)|~v2_pre_topc(X2)|~r1_t_0topsp(X2,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_732, plain, (v1_relat_1(esk27_0)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2362, plain, (v1_partfun1(esk16_1(esk33_0),esk6_0))).#trainneg
cnf(i_0_1115, plain, (v1_valued_0(esk27_0)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2511, plain, (~v2_struct_0(g1_pre_topc(esk6_0,u1_pre_topc(esk17_0)))|~m1_pre_topc(esk17_0,k15_euclid(np__2))|~v1_toprealb(g1_pre_topc(esk6_0,u1_pre_topc(esk17_0))))).#trainneg
cnf(i_0_684, plain, (v1_xcmplx_0(k1_topreal1)|~v1_membered(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_1126, plain, (v1_valued_0(k5_numbers)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_746, plain, (v1_relat_1(k5_numbers)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_750, plain, (v1_xxreal_0(esk50_0)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1656, plain, (esk12_1(X1)=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X2)|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1658, plain, (k2_struct_0(X1)=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X2)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_688, plain, (v1_xreal_0(esk49_1(X1))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_689, plain, (v1_xreal_0(esk47_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2873, plain, (v2_struct_0(g1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(X4)|r1_t_0topsp(X3,g1_pre_topc(X1,X2))|~v3_tops_2(X5,X4,g1_pre_topc(X1,X2))|~v3_tops_2(X6,X3,X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v1_funct_2(k3_relat_1(X6,X5),u1_struct_0(X3),X1)|~v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_2(X5,u1_struct_0(X4),X1)|~v1_funct_1(X5)|~v1_funct_1(X6)|~l1_pre_topc(X3)|~l1_pre_topc(X4))).#trainneg
cnf(i_0_2302, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v1_funct_2(k3_relat_1(X1,esk4_1(X2)),X3,u1_struct_0(esk5_0))|~m1_subset_1(esk4_1(X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(esk5_0))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~v1_funct_2(X1,X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_funct_1(esk4_1(X2))|~v1_funct_1(X1)|k2_struct_0(esk5_0)!=X2|~v1_topreal2(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_2742, plain, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X1),k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v1_topreal2(k2_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~m1_subset_1(k2_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_2287, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(k1_pre_topc(k15_euclid(np__2),X3))|r1_t_0topsp(X1,k1_pre_topc(k15_euclid(np__2),X3))|v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v3_tops_2(esk4_1(X3),X2,k1_pre_topc(k15_euclid(np__2),X3))|~v3_tops_2(X4,X1,X2)|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X3)))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk4_1(X3),u1_struct_0(X2),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X3)))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(esk4_1(X3))|~v1_funct_1(X4)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X3))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v1_topreal2(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))))).#trainneg
cnf(i_0_2577, plain, (k1_pre_topc(X1,u1_struct_0(k1_pre_topc(X1,X2)))=k1_pre_topc(X1,X2)|~l1_struct_0(k1_pre_topc(X1,X2))|~m1_subset_1(u1_struct_0(k1_pre_topc(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2579, plain, (k1_pre_topc(k15_euclid(np__2),X2)=esk5_0|u1_struct_0(esk5_0)!=X2|~l1_struct_0(esk5_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|k2_struct_0(esk5_0)!=X1)).#trainneg
cnf(i_0_755, plain, (v1_xxreal_0(esk31_0)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2503, plain, (v4_relat_1(esk16_1(esk17_0),esk6_0))).#trainneg
cnf(i_0_1131, plain, (v3_valued_0(esk50_0)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_690, plain, (v1_xreal_0(esk34_1(X1))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1138, plain, (v3_valued_0(esk31_0)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_757, plain, (v1_xxreal_0(esk27_0)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_771, plain, (v1_xxreal_0(k5_numbers)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_692, plain, (v1_xreal_0(esk29_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_694, plain, (v1_xreal_0(esk20_1(X1))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_2635, plain, (v2_struct_0(g1_pre_topc(X1,X2))|r1_t_0topsp(X3,g1_pre_topc(X1,X2))|~v2_pre_topc(g1_pre_topc(X1,X2))|~v2_pre_topc(X3)|~r1_t_0topsp(X3,g1_pre_topc(X1,u1_pre_topc(g1_pre_topc(X1,X2))))|~l1_pre_topc(g1_pre_topc(X1,u1_pre_topc(g1_pre_topc(X1,X2))))|~l1_pre_topc(g1_pre_topc(X1,X2))|~l1_pre_topc(X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_2729, plain, (v1_funct_2(esk4_1(k2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_topreal2(k2_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~m1_subset_1(k2_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_2659, plain, (v2_struct_0(X1)|v2_struct_0(k1_pre_topc(k15_euclid(np__2),X2))|v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|v1_funct_2(k3_relat_1(X3,esk4_1(X2)),u1_struct_0(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X2)))|~v2_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v2_pre_topc(k1_pre_topc(k15_euclid(np__2),X2))|~v2_pre_topc(X1)|~v3_tops_2(X3,X1,k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~m1_subset_1(esk4_1(X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X2)))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~v1_funct_2(esk4_1(X2),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X2)))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_funct_1(esk4_1(X2))|~v1_funct_1(X3)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X2))|~l1_pre_topc(X1)|~v1_topreal2(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_2809, plain, (v2_struct_0(X1)|r1_t_0topsp(X1,g1_pre_topc(X2,X3))|v2_struct_0(g1_pre_topc(X2,X3))|~m1_subset_1(k2_tops_2(u1_struct_0(X1),X2,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(X2,X3)),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),X2,X4),u1_struct_0(g1_pre_topc(X2,X3)),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),X2,X4))|~l1_pre_topc(g1_pre_topc(X2,X3))|~l1_pre_topc(X1)|~v3_tops_2(X4,X1,g1_pre_topc(X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),X2)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v1_funct_2(X4,u1_struct_0(X1),X2)|~v1_funct_1(X4))).#trainneg
cnf(i_0_1811, plain, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|v1_topreal2(X1)|~v3_tops_2(X2,esk5_0,k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(X2,u1_struct_0(esk5_0),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_funct_1(X2)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(esk5_0)|k2_struct_0(esk5_0)!=X1)).#trainneg
cnf(i_0_2223, plain, (v2_struct_0(g1_pre_topc(X1,X2))|~v2_tops_1(X1,g1_pre_topc(X1,X2))|~v2_pre_topc(g1_pre_topc(X1,X2))|~l1_pre_topc(g1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_2124, plain, (~v2_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_pre_topc(X1,k15_euclid(np__2))|~v1_topreal2(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))|~m1_subset_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_2632, plain, (v2_struct_0(g1_pre_topc(X1,X2))|r1_t_0topsp(X3,g1_pre_topc(X1,X2))|~v2_pre_topc(g1_pre_topc(X1,X2))|~v2_pre_topc(X3)|~r1_t_0topsp(X3,g1_pre_topc(u1_struct_0(g1_pre_topc(X1,X2)),X2))|~l1_pre_topc(g1_pre_topc(u1_struct_0(g1_pre_topc(X1,X2)),X2))|~l1_pre_topc(g1_pre_topc(X1,X2))|~l1_pre_topc(X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1338, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk20_1(X1)))).#trainneg
cnf(i_0_2504, plain, (v5_relat_1(esk16_1(esk17_0),esk6_0))).#trainneg
cnf(i_0_1140, plain, (v3_valued_0(esk27_0)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_3362, plain, (k1_pre_topc(k15_euclid(np__2),X1)=esk5_0|u1_struct_0(esk5_0)!=X1|k2_struct_0(k15_euclid(np__2))!=k2_struct_0(esk5_0)|~l1_struct_0(esk5_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_695, plain, (v1_xreal_0(esk19_1(X1))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_1151, plain, (v3_valued_0(k5_numbers)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_3111, plain, (~v1_subset_1(esk6_0,esk6_0)|~l1_struct_0(esk33_0))).#trainneg
cnf(i_0_697, plain, (v1_xreal_0(esk13_0)|~v3_membered(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_699, plain, (v1_xreal_0(esk9_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_3222, plain, (k1_pre_topc(X1,u1_struct_0(X2))=X2|~l1_struct_0(X2)|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(X2,X1)|~v1_pre_topc(X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1339, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk19_1(X1)))).#trainneg
cnf(i_0_2505, plain, (v1_partfun1(esk16_1(esk17_0),esk6_0))).#trainneg
cnf(i_0_1990, plain, (v7_ordinal1(esk14_0)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3349, plain, (k1_pre_topc(k15_euclid(np__2),X1)=esk5_0|u1_struct_0(esk5_0)!=X1|esk56_1(k15_euclid(np__2))!=k2_struct_0(esk5_0)|~l1_struct_0(esk5_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_700, plain, (v1_xreal_0(esk8_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_2677, plain, (u1_struct_0(esk54_0)=esk6_0|~v2_struct_0(esk54_0))).#trainneg
cnf(i_0_3113, plain, (~v1_subset_1(esk6_0,esk6_0)|~l1_struct_0(esk17_0))).#trainneg
cnf(i_0_702, plain, (v1_xreal_0(u2_struct_0(X1))|~v3_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_707, plain, (v1_xreal_0(k1_topreal1)|~v3_membered(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_3145, plain, (v2_struct_0(X1)|v2_struct_0(k1_pre_topc(X2,X3))|v2_struct_0(X4)|r1_t_0topsp(X1,k1_pre_topc(X2,X3))|~v3_tops_2(X5,X4,X1)|~v3_tops_2(X6,k1_pre_topc(X2,X3),X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(X4))))|~v1_funct_2(k3_relat_1(X6,X5),X3,u1_struct_0(X1))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_2(X6,X3,u1_struct_0(X4))|~v1_funct_1(X5)|~v1_funct_1(X6)|~l1_pre_topc(k1_pre_topc(X2,X3))|~l1_pre_topc(X1)|~l1_pre_topc(X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_2967, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(X4)|r1_t_0topsp(X4,k1_pre_topc(X1,X2))|~v3_tops_2(X5,X3,k1_pre_topc(X1,X2))|~v3_tops_2(X6,X4,X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X3))))|~v1_funct_2(k1_partfun1(u1_struct_0(X4),u1_struct_0(X3),u1_struct_0(X3),X2,X6,X5),u1_struct_0(X4),X2)|~v1_funct_2(X5,u1_struct_0(X3),X2)|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X6)|~l1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X4)|~l1_pre_topc(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3002, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(k1_pre_topc(X3,X4))|r1_t_0topsp(k1_pre_topc(X3,X4),X2)|~v3_tops_2(X5,X1,k1_pre_topc(X3,X4))|~v3_tops_2(X6,X2,X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k1_partfun1(u1_struct_0(X2),u1_struct_0(X1),u1_struct_0(X1),X4,X6,X5),u1_struct_0(X2),X4)|~v1_funct_2(X5,u1_struct_0(X1),X4)|~v1_funct_2(X6,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X5)|~v1_funct_1(X6)|~l1_pre_topc(X2)|~l1_pre_topc(k1_pre_topc(X3,X4))|~l1_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3))).#trainneg
cnf(i_0_2991, plain, (v2_struct_0(X1)|v2_struct_0(k1_pre_topc(X2,X3))|v2_struct_0(X4)|r1_t_0topsp(X4,X1)|~v3_tops_2(X5,k1_pre_topc(X2,X3),X1)|~v3_tops_2(X6,X4,k1_pre_topc(X2,X3))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(X1))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),X3)))|~v1_funct_2(k1_partfun1(u1_struct_0(X4),X3,X3,u1_struct_0(X1),X6,X5),u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_2(X5,X3,u1_struct_0(X1))|~v1_funct_2(X6,u1_struct_0(X4),X3)|~v1_funct_1(X5)|~v1_funct_1(X6)|~l1_pre_topc(X1)|~l1_pre_topc(X4)|~l1_pre_topc(k1_pre_topc(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_2918, plain, (v3_tops_2(k2_tops_2(X2,X2,X3),k1_pre_topc(X1,X2),k1_pre_topc(X1,X2))|~v3_tops_2(X3,k1_pre_topc(X1,X2),k1_pre_topc(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~v1_funct_2(X3,X2,X2)|~v1_funct_1(X3)|~l1_pre_topc(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2986, plain, (v2_struct_0(X1)|v2_struct_0(k1_pre_topc(X2,X3))|v2_struct_0(X4)|r1_t_0topsp(X4,k1_pre_topc(X2,X3))|~v3_tops_2(X5,X1,X4)|~v3_tops_2(X6,k1_pre_topc(X2,X3),X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X4))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(X1))))|~v1_funct_2(k1_partfun1(X3,u1_struct_0(X1),u1_struct_0(X1),u1_struct_0(X4),X6,X5),X3,u1_struct_0(X4))|~v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X4))|~v1_funct_2(X6,X3,u1_struct_0(X1))|~v1_funct_1(X5)|~v1_funct_1(X6)|~l1_pre_topc(k1_pre_topc(X2,X3))|~l1_pre_topc(X4)|~l1_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_3003, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(k1_pre_topc(X3,X4))|r1_t_0topsp(X1,k1_pre_topc(X3,X4))|~v3_tops_2(X5,X2,k1_pre_topc(X3,X4))|~v3_tops_2(X6,X1,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),X4)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(k3_relat_1(X6,X5),u1_struct_0(X1),X4)|~v1_funct_2(X5,u1_struct_0(X2),X4)|~v1_funct_2(X6,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X5)|~v1_funct_1(X6)|~l1_pre_topc(k1_pre_topc(X3,X4))|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3))).#trainneg
cnf(i_0_3348, plain, (k1_pre_topc(X1,u1_struct_0(k1_pre_topc(X1,X2)))=k1_pre_topc(X1,X2)|~m1_subset_1(u1_struct_0(k1_pre_topc(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~l1_pre_topc(k1_pre_topc(X1,X2)))).#trainneg
cnf(i_0_1556, plain, (v1_funct_1(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_800, plain, (v1_relat_1(esk34_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2678, plain, (u1_struct_0(esk48_0)=esk6_0|~v2_struct_0(esk48_0))).#trainneg
cnf(i_0_3351, plain, (k1_pre_topc(k15_euclid(np__2),X1)=esk5_0|u1_struct_0(esk5_0)!=X1|esk43_1(k15_euclid(np__2))!=k2_struct_0(esk5_0)|~l1_struct_0(esk5_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_712, plain, (v1_funct_1(esk12_1(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2679, plain, (u1_struct_0(esk45_0)=esk6_0|~v2_struct_0(esk45_0))).#trainneg
cnf(i_0_2671, plain, (v1_xcmplx_0(esk50_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2814, plain, (v1_xcmplx_0(esk31_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3182, plain, (v1_funct_1(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_716, plain, (v7_ordinal1(esk25_2(X1,X2))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_717, plain, (v1_xcmplx_0(esk25_2(X1,X2))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2916, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(X3,k1_pre_topc(X1,X2))|~v2_pre_topc(k1_pre_topc(X1,X2))|~v2_pre_topc(X3)|~r1_t_0topsp(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3)),g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~l1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2955, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(X4)|r1_t_0topsp(X4,X3)|~v3_tops_2(X5,k1_pre_topc(X1,X2),X4)|~v3_tops_2(X6,X3,k1_pre_topc(X1,X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X4))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X2)))|~v1_funct_2(k1_partfun1(u1_struct_0(X3),X2,X2,u1_struct_0(X4),X6,X5),u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_2(X5,X2,u1_struct_0(X4))|~v1_funct_2(X6,u1_struct_0(X3),X2)|~v1_funct_1(X5)|~v1_funct_1(X6)|~l1_pre_topc(X3)|~l1_pre_topc(X4)|~l1_pre_topc(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3038, plain, (k2_struct_0(k1_pre_topc(X1,X2))=u1_struct_0(k1_pre_topc(X1,X2))|~m1_subset_1(k2_struct_0(k1_pre_topc(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_2978, plain, (v2_struct_0(X1)|v3_tops_2(k2_tops_2(X3,u1_struct_0(X1),X4),X1,k1_pre_topc(X2,X3))|~v3_tops_2(X4,k1_pre_topc(X2,X3),X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(X1))))|~v1_funct_2(X4,X3,u1_struct_0(X1))|~v1_funct_1(X4)|~l1_pre_topc(X1)|~l1_pre_topc(k1_pre_topc(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_3014, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(k1_topreal1,u1_struct_0(esk5_0))))|k2_struct_0(esk5_0)!=X1|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_1557, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1269, plain, (v1_partfun1(esk25_2(X1,X1),X1))).#trainneg
cnf(i_0_2680, plain, (u1_struct_0(esk44_0)=esk6_0|~v2_struct_0(esk44_0))).#trainneg
cnf(i_0_3352, plain, (k1_pre_topc(k15_euclid(np__2),X1)=esk5_0|u1_struct_0(esk5_0)!=X1|esk40_1(k15_euclid(np__2))!=k2_struct_0(esk5_0)|~l1_struct_0(esk5_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v2_pre_topc(k15_euclid(np__2)))).#trainneg
cnf(i_0_718, plain, (v1_xreal_0(esk25_2(X1,X2))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2819, plain, (v1_xcmplx_0(esk27_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2820, plain, (v1_xcmplx_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_721, plain, (v7_ordinal1(esk36_2(X1,X2))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3013, plain, (v1_funct_2(esk4_1(X1),k1_topreal1,u1_struct_0(esk5_0))|k2_struct_0(esk5_0)!=X1|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_3006, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(k1_pre_topc(X3,X4))|r1_t_0topsp(k1_pre_topc(X3,X4),X1)|~v3_tops_2(X5,X2,X1)|~v3_tops_2(X6,k1_pre_topc(X3,X4),X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,u1_struct_0(X2))))|~v1_funct_2(k1_partfun1(X4,u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X1),X6,X5),X4,u1_struct_0(X1))|~v1_funct_2(X5,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_2(X6,X4,u1_struct_0(X2))|~v1_funct_1(X5)|~v1_funct_1(X6)|~l1_pre_topc(X1)|~l1_pre_topc(k1_pre_topc(X3,X4))|~l1_pre_topc(X2)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3))).#trainneg
cnf(i_0_2920, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(X4)|v3_tops_2(k1_partfun1(X2,u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X4),X5,X6),k1_pre_topc(X1,X2),X4)|~v3_tops_2(X6,X3,X4)|~v3_tops_2(X5,k1_pre_topc(X1,X2),X3)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_2(X5,X2,u1_struct_0(X3))|~v1_funct_1(X6)|~v1_funct_1(X5)|~l1_pre_topc(X4)|~l1_pre_topc(X3)|~l1_pre_topc(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2968, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(X3,k1_pre_topc(X1,X2))|~v2_pre_topc(k1_pre_topc(X1,X2))|~v2_pre_topc(X3)|~v1_pre_topc(X3)|~r1_t_0topsp(X3,g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~l1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2996, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(k1_pre_topc(X3,X4))|v3_tops_2(k1_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X2),X4,X5,X6),X1,k1_pre_topc(X3,X4))|~v3_tops_2(X6,X2,k1_pre_topc(X3,X4))|~v3_tops_2(X5,X1,X2)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X2),X4)|~v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X6)|~v1_funct_1(X5)|~l1_pre_topc(k1_pre_topc(X3,X4))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3))).#trainneg
cnf(i_0_2990, plain, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|v2_struct_0(k1_pre_topc(X2,X3))|v1_topreal2(X1)|~v3_tops_2(X4,k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(X2,X3))|~v3_tops_2(X5,k1_pre_topc(X2,X3),k1_pre_topc(k15_euclid(np__2),X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X3)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(k1_partfun1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X3,X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),X4,X5),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_funct_2(X4,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X3)|~v1_funct_2(X5,X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|~l1_pre_topc(k1_pre_topc(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_3469, plain, (v1_funct_2(esk4_1(X1),k1_topreal1,X1)|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_3497, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(k1_topreal1,X1)))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_1558, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1562, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_1580, plain, (v1_xboole_0(u1_struct_0(esk44_0))|~v2_struct_0(esk44_0))).#trainneg
cnf(i_0_722, plain, (v1_xcmplx_0(esk36_2(X1,X2))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1024, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk50_0))).#trainneg
cnf(i_0_3277, plain, (v1_xxreal_0(esk50_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3368, plain, (v1_xxreal_0(esk31_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_723, plain, (v1_xreal_0(esk36_2(X1,X2))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_726, plain, (v1_relat_1(esk49_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4046, plain, (v2_struct_0(X1)|r1_t_0topsp(X1,k1_pre_topc(X2,X3))|v2_struct_0(k1_pre_topc(X2,X3))|~m1_subset_1(k2_tops_2(u1_struct_0(X1),X3,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X2,X3)),u1_struct_0(X1))))|~v1_funct_2(k2_tops_2(u1_struct_0(X1),X3,X4),u1_struct_0(k1_pre_topc(X2,X3)),u1_struct_0(X1))|~v1_funct_1(k2_tops_2(u1_struct_0(X1),X3,X4))|~l1_pre_topc(k1_pre_topc(X2,X3))|~l1_pre_topc(X1)|~v3_tops_2(X4,X1,k1_pre_topc(X2,X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),X3)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X4,u1_struct_0(X1),X3)|~v1_funct_1(X4)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_3464, plain, (v1_funct_2(k3_relat_1(X1,esk4_1(X2)),X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X2)))|~m1_subset_1(esk4_1(X2),k1_zfmisc_1(k2_zfmisc_1(k1_topreal1,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X2)))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k1_topreal1)))|~v1_funct_2(X1,X3,k1_topreal1)|~v1_funct_1(esk4_1(X2))|~v1_funct_1(X1)|~v1_topreal2(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_4041, plain, (r1_t_0topsp(k1_pre_topc(X1,X2),X3)|v2_struct_0(k1_pre_topc(X1,X2))|~m1_subset_1(k2_tops_2(u1_struct_0(X3),X2,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X3))))|~v1_funct_2(k2_tops_2(u1_struct_0(X3),X2,X4),u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X3))|~v1_funct_1(k2_tops_2(u1_struct_0(X3),X2,X4))|~l1_pre_topc(X3)|~l1_pre_topc(k1_pre_topc(X1,X2))|~v3_tops_2(X4,X3,k1_pre_topc(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X4,u1_struct_0(X3),X2)|~v1_funct_1(X4)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3848, plain, (r1_t_0topsp(X1,k1_pre_topc(X2,X3))|v2_struct_0(X1)|~m1_subset_1(k2_tops_2(X3,u1_struct_0(X1),X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k1_pre_topc(X2,X3)))))|~v1_funct_2(k2_tops_2(X3,u1_struct_0(X1),X4),u1_struct_0(X1),u1_struct_0(k1_pre_topc(X2,X3)))|~v1_funct_1(k2_tops_2(X3,u1_struct_0(X1),X4))|~l1_pre_topc(k1_pre_topc(X2,X3))|~l1_pre_topc(X1)|~v3_tops_2(X4,k1_pre_topc(X2,X3),X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X4,X3,u1_struct_0(X1))|~v1_funct_1(X4)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1559, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1567, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_729, plain, (v1_relat_1(esk34_1(X1))|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1031, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk31_0))).#trainneg
cnf(i_0_3370, plain, (v1_xxreal_0(esk27_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3371, plain, (v1_xxreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_739, plain, (v1_relat_1(esk9_1(X1))|~v4_funct_1(X1))).#trainneg
cnf(i_0_737, plain, (v1_relat_1(esk13_0)|~v4_funct_1(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_742, plain, (v1_relat_1(u2_struct_0(X1))|~v4_funct_1(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_3720, plain, (v3_tops_2(esk4_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))),k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~m1_subset_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_3993, plain, (r1_t_0topsp(X1,k1_pre_topc(X2,X3))|v2_struct_0(k1_pre_topc(X2,X3))|v2_struct_0(X4)|v2_struct_0(X1)|~m1_subset_1(k1_partfun1(u1_struct_0(X1),u1_struct_0(X4),u1_struct_0(X4),X3,X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k1_pre_topc(X2,X3)))))|~v1_funct_2(k1_partfun1(u1_struct_0(X1),u1_struct_0(X4),u1_struct_0(X4),X3,X5,X6),u1_struct_0(X1),u1_struct_0(k1_pre_topc(X2,X3)))|~v1_funct_1(k1_partfun1(u1_struct_0(X1),u1_struct_0(X4),u1_struct_0(X4),X3,X5,X6))|~l1_pre_topc(k1_pre_topc(X2,X3))|~l1_pre_topc(X1)|~v3_tops_2(X6,X4,k1_pre_topc(X2,X3))|~v3_tops_2(X5,X1,X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X4))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),X3)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X4))|~v1_funct_2(X6,u1_struct_0(X4),X3)|~v1_funct_1(X6)|~v1_funct_1(X5)|~l1_pre_topc(X4)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_3853, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(k1_pre_topc(X1,X2),X3)|v2_struct_0(X3)|~m1_subset_1(k2_tops_2(X2,u1_struct_0(X3),X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k1_pre_topc(X1,X2)))))|~v1_funct_2(k2_tops_2(X2,u1_struct_0(X3),X4),u1_struct_0(X3),u1_struct_0(k1_pre_topc(X1,X2)))|~v1_funct_1(k2_tops_2(X2,u1_struct_0(X3),X4))|~l1_pre_topc(X3)|~l1_pre_topc(k1_pre_topc(X1,X2))|~v3_tops_2(X4,k1_pre_topc(X1,X2),X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X4,X2,u1_struct_0(X3))|~v1_funct_1(X4)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3663, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(k1_pre_topc(X4,X5))|r1_t_0topsp(X3,k1_pre_topc(X4,X5))|~v3_tops_2(X6,k1_pre_topc(X1,X2),k1_pre_topc(X4,X5))|~v3_tops_2(X7,X3,k1_pre_topc(X1,X2))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X5)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k1_partfun1(u1_struct_0(X3),X2,X2,X5,X7,X6),u1_struct_0(X3),X5)|~v1_funct_2(X6,X2,X5)|~v1_funct_2(X7,u1_struct_0(X3),X2)|~v1_funct_1(X6)|~v1_funct_1(X7)|~l1_pre_topc(k1_pre_topc(X4,X5))|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4)))|~l1_pre_topc(X4))).#trainneg
cnf(i_0_3954, plain, (v2_struct_0(X1)|r1_t_0topsp(X1,k1_pre_topc(X2,X3))|v2_struct_0(k1_pre_topc(X2,X3))|v2_struct_0(X4)|~m1_subset_1(k1_partfun1(X3,u1_struct_0(X4),u1_struct_0(X4),u1_struct_0(X1),X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X2,X3)),u1_struct_0(X1))))|~v1_funct_2(k1_partfun1(X3,u1_struct_0(X4),u1_struct_0(X4),u1_struct_0(X1),X5,X6),u1_struct_0(k1_pre_topc(X2,X3)),u1_struct_0(X1))|~v1_funct_1(k1_partfun1(X3,u1_struct_0(X4),u1_struct_0(X4),u1_struct_0(X1),X5,X6))|~l1_pre_topc(k1_pre_topc(X2,X3))|~l1_pre_topc(X1)|~v3_tops_2(X5,k1_pre_topc(X2,X3),X4)|~v3_tops_2(X6,X4,X1)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(X4))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_2(X5,X3,u1_struct_0(X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~l1_pre_topc(X4)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1560, plain, (v1_relat_1(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2365, plain, (m1_subset_1(esk43_1(esk33_0),k1_zfmisc_1(esk6_0)))).#trainneg
cnf(i_0_3176, plain, (v1_funct_1(u1_struct_0(esk54_0))|~v2_struct_0(esk54_0))).#trainneg
cnf(i_0_3359, plain, (k1_pre_topc(k15_euclid(np__2),X1)=esk5_0|u1_struct_0(esk5_0)!=X1|esk12_1(k15_euclid(np__2))!=k2_struct_0(esk5_0)|~l1_struct_0(esk5_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v2_pre_topc(k15_euclid(np__2)))).#trainneg
cnf(i_0_747, plain, (v1_relat_1(k1_topreal1)|~v4_funct_1(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_1033, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk27_0))).#trainneg
cnf(i_0_835, plain, (v3_membered(esk47_0)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_764, plain, (v1_xxreal_0(esk9_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_817, plain, (v1_xboole_0(esk34_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_751, plain, (v1_xxreal_0(esk49_1(X1))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_752, plain, (v1_xxreal_0(esk47_0)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3859, plain, (v2_struct_0(X1)|v3_tops_2(k2_tops_2(u1_struct_0(k1_pre_topc(X2,X3)),u1_struct_0(X1),X4),X1,k1_pre_topc(X2,X3))|~v3_tops_2(X4,k1_pre_topc(X2,X3),X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X2,X3)),u1_struct_0(X1))))|~m1_subset_1(u1_struct_0(k1_pre_topc(X2,X3)),k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X4,u1_struct_0(k1_pre_topc(X2,X3)),u1_struct_0(X1))|~v1_funct_1(X4)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_3998, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(k1_pre_topc(X1,X2),X3)|v2_struct_0(X4)|v2_struct_0(X3)|~m1_subset_1(k1_partfun1(u1_struct_0(X3),u1_struct_0(X4),u1_struct_0(X4),X2,X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k1_pre_topc(X1,X2)))))|~v1_funct_2(k1_partfun1(u1_struct_0(X3),u1_struct_0(X4),u1_struct_0(X4),X2,X5,X6),u1_struct_0(X3),u1_struct_0(k1_pre_topc(X1,X2)))|~v1_funct_1(k1_partfun1(u1_struct_0(X3),u1_struct_0(X4),u1_struct_0(X4),X2,X5,X6))|~l1_pre_topc(X3)|~l1_pre_topc(k1_pre_topc(X1,X2))|~v3_tops_2(X6,X4,k1_pre_topc(X1,X2))|~v3_tops_2(X5,X3,X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_2(X6,u1_struct_0(X4),X2)|~v1_funct_1(X6)|~v1_funct_1(X5)|~l1_pre_topc(X4)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3815, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(X3,k1_pre_topc(X1,X2))|v2_struct_0(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~v2_pre_topc(k1_pre_topc(X1,X2))|~v2_pre_topc(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v2_pre_topc(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~r1_t_0topsp(X3,g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~l1_pre_topc(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3)))|~l1_pre_topc(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2)))))).#trainneg
cnf(i_0_3863, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v3_tops_2(k2_tops_2(X3,X2,X4),k1_pre_topc(X1,X2),k1_pre_topc(X5,X3))|~v3_tops_2(X4,k1_pre_topc(X5,X3),k1_pre_topc(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X5)))|~v1_funct_2(X4,X3,X2)|~v1_funct_1(X4)|~l1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X5)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3775, plain, (v3_tops_2(k2_tops_2(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(k1_pre_topc(X1,X2)),X3),k1_pre_topc(X1,X2),k1_pre_topc(X1,X2))|~v3_tops_2(X3,k1_pre_topc(X1,X2),k1_pre_topc(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(k1_pre_topc(X1,X2)))))|~m1_subset_1(u1_struct_0(k1_pre_topc(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(k1_pre_topc(X1,X2)))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_3883, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(X3,k1_pre_topc(X1,X2))|~v2_pre_topc(k1_pre_topc(X1,X2))|~v2_pre_topc(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r1_t_0topsp(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3)),k1_pre_topc(X1,X2))|~l1_pre_topc(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3949, plain, (r1_t_0topsp(k1_pre_topc(X1,X2),X3)|v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(X4)|~m1_subset_1(k1_partfun1(X2,u1_struct_0(X4),u1_struct_0(X4),u1_struct_0(X3),X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X3))))|~v1_funct_2(k1_partfun1(X2,u1_struct_0(X4),u1_struct_0(X4),u1_struct_0(X3),X5,X6),u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X3))|~v1_funct_1(k1_partfun1(X2,u1_struct_0(X4),u1_struct_0(X4),u1_struct_0(X3),X5,X6))|~l1_pre_topc(X3)|~l1_pre_topc(k1_pre_topc(X1,X2))|~v3_tops_2(X5,k1_pre_topc(X1,X2),X4)|~v3_tops_2(X6,X4,X3)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X3))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X4))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X3))|~v1_funct_2(X5,X2,u1_struct_0(X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~l1_pre_topc(X4)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1561, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2366, plain, (m1_subset_1(esk56_1(esk33_0),k1_zfmisc_1(esk6_0)))).#trainneg
cnf(i_0_3177, plain, (v1_funct_1(u1_struct_0(esk48_0))|~v2_struct_0(esk48_0))).#trainneg
cnf(i_0_3350, plain, (k1_pre_topc(k15_euclid(np__2),X1)=esk5_0|u1_struct_0(esk5_0)!=X1|esk49_1(u1_struct_0(k15_euclid(np__2)))!=k2_struct_0(esk5_0)|~l1_struct_0(esk5_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_753, plain, (v1_xxreal_0(esk36_2(X1,X2))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1043, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_840, plain, (v3_membered(esk29_0)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_845, plain, (v3_membered(esk13_0)|~v3_membered(u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_838, plain, (v3_membered(esk34_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_859, plain, (v1_membered(esk34_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_754, plain, (v1_xxreal_0(esk34_1(X1))|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_756, plain, (v1_xxreal_0(esk29_0)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3676, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(X4)|r1_t_0topsp(X3,X4)|~v3_tops_2(X5,k1_pre_topc(X1,X2),X4)|~v3_tops_2(X6,X3,k1_pre_topc(X1,X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X4))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k3_relat_1(X6,X5),u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_2(X5,X2,u1_struct_0(X4))|~v1_funct_2(X6,u1_struct_0(X3),X2)|~v1_funct_1(X5)|~v1_funct_1(X6)|~l1_pre_topc(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3739, plain, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X1),k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v1_topreal2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~m1_subset_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_3529, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),X1)))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|k2_struct_0(esk5_0)!=k1_topreal1)).#trainneg
cnf(i_0_3819, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(k1_pre_topc(X3,X4))|v2_struct_0(X5)|r1_t_0topsp(k1_pre_topc(X3,X4),X5)|~v3_tops_2(X6,k1_pre_topc(X1,X2),k1_pre_topc(X3,X4))|~v3_tops_2(X7,X5,k1_pre_topc(X1,X2))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k1_partfun1(u1_struct_0(X5),X2,X2,X4,X7,X6),u1_struct_0(X5),X4)|~v1_funct_2(X6,X2,X4)|~v1_funct_2(X7,u1_struct_0(X5),X2)|~v1_funct_1(X6)|~v1_funct_1(X7)|~l1_pre_topc(X5)|~l1_pre_topc(k1_pre_topc(X3,X4))|~l1_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3))).#trainneg
cnf(i_0_3931, plain, (v1_funct_2(k3_relat_1(X1,esk4_1(X2)),X3,u1_struct_0(esk5_0))|~m1_subset_1(esk4_1(X2),k1_zfmisc_1(k2_zfmisc_1(k1_topreal1,u1_struct_0(esk5_0))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k1_topreal1)))|~v1_funct_2(X1,X3,k1_topreal1)|~v1_funct_1(esk4_1(X2))|~v1_funct_1(X1)|k2_struct_0(esk5_0)!=X2|~v1_topreal2(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_3874, plain, (v2_struct_0(k1_pre_topc(X2,X1))|r1_t_0topsp(X3,k1_pre_topc(X2,X1))|~v2_pre_topc(k1_pre_topc(X2,X1))|~v2_pre_topc(X3)|~m1_subset_1(u1_pre_topc(k1_pre_topc(X2,X1)),k1_zfmisc_1(k1_zfmisc_1(X1)))|~r1_t_0topsp(X3,g1_pre_topc(X1,u1_pre_topc(k1_pre_topc(X2,X1))))|~l1_pre_topc(g1_pre_topc(X1,u1_pre_topc(k1_pre_topc(X2,X1))))|~l1_pre_topc(X3)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_2381, plain, (v1_funct_2(esk16_1(esk33_0),esk6_0,esk6_0))).#trainneg
cnf(i_0_3178, plain, (v1_funct_1(u1_struct_0(esk45_0))|~v2_struct_0(esk45_0))).#trainneg
cnf(i_0_3353, plain, (k1_pre_topc(k15_euclid(np__2),X1)=esk5_0|u1_struct_0(esk5_0)!=X1|esk34_1(u1_struct_0(k15_euclid(np__2)))!=k2_struct_0(esk5_0)|~l1_struct_0(esk5_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_758, plain, (v1_xxreal_0(esk25_2(X1,X2))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2412, plain, (v1_funct_1(X1)|~v1_monoid_0(esk33_0)|~l1_struct_0(esk33_0)|~m1_subset_1(X1,esk6_0))).#trainneg
cnf(i_0_848, plain, (v3_membered(esk8_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_852, plain, (v3_membered(k1_topreal1)|~v3_membered(u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_881, plain, (v2_membered(esk34_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_981, plain, (v1_valued_0(esk20_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_759, plain, (v1_xxreal_0(esk20_1(X1))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_760, plain, (v1_xxreal_0(esk19_1(X1))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_3689, plain, (v3_tops_2(k2_tops_2(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),X2),g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v3_tops_2(X2,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))))|~m1_subset_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),k1_zfmisc_1(u1_struct_0(X3)))|~v1_funct_2(X2,u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))|~v1_funct_1(X2)|~l1_pre_topc(X3)|~m1_pre_topc(X1,X3))).#trainneg
cnf(i_0_3888, plain, (g1_pre_topc(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))=g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))|~m1_subset_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~m1_pre_topc(X1,X2))).#trainneg
cnf(i_0_3989, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))),k1_pre_topc(X1,X2))|~v2_pre_topc(k1_pre_topc(X1,X2))|~v2_pre_topc(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~v1_pre_topc(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4210, plain, (v1_topreal2(X1)|~m1_subset_1(esk4_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(esk4_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_funct_1(esk4_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))))|~v1_topreal2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~m1_subset_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_4408, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(X3,X4))|v2_struct_0(k1_pre_topc(X3,X4))|~m1_subset_1(k2_tops_2(X2,X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X3,X4)),u1_struct_0(k1_pre_topc(X1,X2)))))|~v1_funct_2(k2_tops_2(X2,X4,X5),u1_struct_0(k1_pre_topc(X3,X4)),u1_struct_0(k1_pre_topc(X1,X2)))|~v1_funct_1(k2_tops_2(X2,X4,X5))|~l1_pre_topc(k1_pre_topc(X3,X4))|~l1_pre_topc(k1_pre_topc(X1,X2))|~v3_tops_2(X5,k1_pre_topc(X1,X2),k1_pre_topc(X3,X4))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~v1_funct_2(X5,X2,X4)|~v1_funct_1(X5)|~l1_pre_topc(X1)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_4403, plain, (r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(X3,X4))|v2_struct_0(k1_pre_topc(X1,X2))|~m1_subset_1(k2_tops_2(X4,X2,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(k1_pre_topc(X3,X4)))))|~v1_funct_2(k2_tops_2(X4,X2,X5),u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(k1_pre_topc(X3,X4)))|~v1_funct_1(k2_tops_2(X4,X2,X5))|~l1_pre_topc(k1_pre_topc(X3,X4))|~l1_pre_topc(k1_pre_topc(X1,X2))|~v3_tops_2(X5,k1_pre_topc(X3,X4),k1_pre_topc(X1,X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X5,X4,X2)|~v1_funct_1(X5)|~l1_pre_topc(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4182, plain, (v2_struct_0(X1)|r1_t_0topsp(X1,k1_pre_topc(X2,X3))|~v3_tops_2(X4,k1_pre_topc(X2,X3),X1)|~m1_subset_1(k2_tops_2(X3,u1_struct_0(X1),X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(k2_tops_2(X3,u1_struct_0(X1),X4),u1_struct_0(X1),X3)|~v1_funct_2(X4,X3,u1_struct_0(X1))|~v1_funct_1(X4)|~l1_pre_topc(X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_4088, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4308, plain, (v1_xboole_0(esk43_1(esk33_0)))).#trainneg
cnf(i_0_4309, plain, (v3_membered(esk43_1(esk33_0)))).#trainneg
cnf(i_0_3360, plain, (k1_pre_topc(k15_euclid(np__2),X1)=esk5_0|u1_struct_0(esk5_0)!=X1|esk9_1(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))!=k2_struct_0(esk5_0)|~l1_struct_0(esk5_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_762, plain, (v1_xxreal_0(esk13_0)|~v2_membered(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_2413, plain, (v1_relat_1(X1)|~v1_monoid_0(esk33_0)|~l1_struct_0(esk33_0)|~m1_subset_1(X1,esk6_0))).#trainneg
cnf(i_0_856, plain, (v1_membered(esk47_0)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_861, plain, (v1_membered(esk29_0)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4320, plain, (X1=esk34_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4326, plain, (esk34_1(X1)=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_765, plain, (v1_xxreal_0(esk8_0)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_767, plain, (v1_xxreal_0(u2_struct_0(X1))|~v2_membered(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_4138, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(X3,X4))|~v3_tops_2(X5,k1_pre_topc(X3,X4),k1_pre_topc(X1,X2))|~m1_subset_1(k2_tops_2(X4,X2,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(X4,X2,X5),u1_struct_0(k1_pre_topc(X1,X2)),X4)|~v1_funct_2(X5,X4,X2)|~v1_funct_1(X5)|~l1_pre_topc(k1_pre_topc(X3,X4))|~l1_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3))).#trainneg
cnf(i_0_4197, plain, (g1_pre_topc(u1_struct_0(k1_pre_topc(X1,X2)),u1_pre_topc(k1_pre_topc(X1,X2)))=k1_pre_topc(X1,u1_struct_0(k1_pre_topc(X1,X2)))|~m1_subset_1(u1_pre_topc(k1_pre_topc(X1,X2)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)))))|~m1_subset_1(u1_struct_0(k1_pre_topc(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_4103, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(k1_pre_topc(X3,X4))|r1_t_0topsp(k1_pre_topc(X3,X4),k1_pre_topc(X1,X2))|~v3_tops_2(X5,k1_pre_topc(X3,X4),k1_pre_topc(X1,X2))|~m1_subset_1(k2_tops_2(X4,X2,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(X4,X2,X5),u1_struct_0(k1_pre_topc(X1,X2)),X4)|~v1_funct_2(X5,X4,X2)|~v1_funct_1(X5)|~l1_pre_topc(k1_pre_topc(X3,X4))|~l1_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3))).#trainneg
cnf(i_0_4498, plain, (v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4310, plain, (v1_membered(esk43_1(esk33_0)))).#trainneg
cnf(i_0_4311, plain, (v2_membered(esk43_1(esk33_0)))).#trainneg
cnf(i_0_3447, plain, (~v2_struct_0(X1)|~v1_topreal2(u1_struct_0(X1))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_pre_topc(X1,k15_euclid(np__2))|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_772, plain, (v1_xxreal_0(k1_topreal1)|~v2_membered(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_2557, plain, (v1_funct_1(X1)|~v1_monoid_0(esk17_0)|~l1_struct_0(esk17_0)|~m1_subset_1(X1,esk6_0))).#trainneg
cnf(i_0_866, plain, (v1_membered(esk13_0)|~v1_membered(u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_869, plain, (v1_membered(esk8_0)|~v1_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_4328, plain, (esk34_1(X1)=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_982, plain, (v1_valued_0(esk19_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_775, plain, (v1_funct_1(k2_struct_0(X1))|~v4_funct_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_776, plain, (v7_ordinal1(k2_struct_0(X1))|~v6_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_4540, plain, (r1_t_0topsp(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v1_topreal2(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))|~m1_subset_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_pre_topc(X1,k15_euclid(np__2)))).#trainneg
cnf(i_0_4261, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(k1_pre_topc(X3,X4))|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(X3,X4))|~v3_tops_2(X5,k1_pre_topc(X1,X2),k1_pre_topc(X3,X4))|~m1_subset_1(k2_tops_2(X2,X4,X5),k1_zfmisc_1(k2_zfmisc_1(X4,u1_struct_0(k1_pre_topc(X1,X2)))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(X2,X4,X5),X4,u1_struct_0(k1_pre_topc(X1,X2)))|~v1_funct_2(X5,X2,X4)|~v1_funct_1(X5)|~l1_pre_topc(k1_pre_topc(X3,X4))|~l1_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3))).#trainneg
cnf(i_0_4241, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(k1_pre_topc(X4,X5))|r1_t_0topsp(X3,k1_pre_topc(X1,X2))|~v3_tops_2(X6,k1_pre_topc(X4,X5),k1_pre_topc(X1,X2))|~v3_tops_2(X7,X3,k1_pre_topc(X4,X5))|~m1_subset_1(k1_partfun1(u1_struct_0(X3),X5,X5,X2,X7,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k1_pre_topc(X1,X2)))))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k1_partfun1(u1_struct_0(X3),X5,X5,X2,X7,X6),u1_struct_0(X3),u1_struct_0(k1_pre_topc(X1,X2)))|~v1_funct_2(X7,u1_struct_0(X3),X5)|~v1_funct_2(X6,X5,X2)|~v1_funct_1(X6)|~v1_funct_1(X7)|~l1_pre_topc(X3)|~l1_pre_topc(k1_pre_topc(X4,X5))|~l1_pre_topc(X1)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4)))|~l1_pre_topc(X4))).#trainneg
cnf(i_0_4300, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(X4)|r1_t_0topsp(X4,k1_pre_topc(X1,X2))|~v3_tops_2(X5,k1_pre_topc(X1,X2),X3)|~v3_tops_2(X6,X3,X4)|~m1_subset_1(k3_relat_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X4))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k3_relat_1(X5,X6),u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X4))|~v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_2(X5,X2,u1_struct_0(X3))|~v1_funct_1(X6)|~v1_funct_1(X5)|~l1_pre_topc(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4312, plain, (v1_funct_1(esk43_1(esk33_0)))).#trainneg
cnf(i_0_4313, plain, (v1_valued_0(esk43_1(esk33_0)))).#trainneg
cnf(i_0_4314, plain, (v3_valued_0(esk43_1(esk33_0)))).#trainneg
cnf(i_0_777, plain, (v1_xcmplx_0(k2_struct_0(X1))|~v1_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_2558, plain, (v1_relat_1(X1)|~v1_monoid_0(esk17_0)|~l1_struct_0(esk17_0)|~m1_subset_1(X1,esk6_0))).#trainneg
cnf(i_0_873, plain, (v1_membered(k1_topreal1)|~v1_membered(u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_878, plain, (v2_membered(esk47_0)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1053, plain, (v1_xboole_0(esk20_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1054, plain, (v1_xboole_0(esk19_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_778, plain, (v1_xreal_0(k2_struct_0(X1))|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_779, plain, (v1_relat_1(k2_struct_0(X1))|~v4_funct_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_4643, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(g1_pre_topc(k2_struct_0(k1_pre_topc(X1,X2)),u1_pre_topc(k1_pre_topc(X1,X2))),k1_pre_topc(X1,X2))|~v2_pre_topc(g1_pre_topc(k2_struct_0(k1_pre_topc(X1,X2)),u1_pre_topc(k1_pre_topc(X1,X2))))|~v2_pre_topc(k1_pre_topc(X1,X2))|~v1_pre_topc(g1_pre_topc(k2_struct_0(k1_pre_topc(X1,X2)),u1_pre_topc(k1_pre_topc(X1,X2))))|~m1_subset_1(k2_struct_0(k1_pre_topc(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(g1_pre_topc(k2_struct_0(k1_pre_topc(X1,X2)),u1_pre_topc(k1_pre_topc(X1,X2))))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_4110, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(X3,k1_pre_topc(X1,X2))|~v3_tops_2(X4,X3,k1_pre_topc(X1,X2))|~m1_subset_1(k2_tops_2(u1_struct_0(X3),X2,X4),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(u1_struct_0(X3),X2,X4),X2,u1_struct_0(X3))|~v1_funct_2(X4,u1_struct_0(X3),X2)|~v1_funct_1(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4506, plain, (v1_xboole_0(esk56_1(esk33_0)))).#trainneg
cnf(i_0_4507, plain, (v3_membered(esk56_1(esk33_0)))).#trainneg
cnf(i_0_4508, plain, (v1_membered(esk56_1(esk33_0)))).#trainneg
cnf(i_0_3879, plain, (X1=X2|g1_pre_topc(X1,X3)!=k1_pre_topc(X4,X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X4)))|~l1_pre_topc(X4))).#trainneg
cnf(i_0_780, plain, (v1_xxreal_0(k2_struct_0(X1))|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_4608, plain, (v1_funct_1(X1)|~l1_struct_0(esk33_0)|~m1_subset_1(X1,esk6_0)|~v2_monoid_0(esk33_0))).#trainneg
cnf(i_0_883, plain, (v2_membered(esk29_0)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_888, plain, (v2_membered(esk13_0)|~v2_membered(u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_783, plain, (v1_funct_1(esk43_1(X1))|~v4_funct_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_784, plain, (v7_ordinal1(esk43_1(X1))|~v6_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4625, plain, (v3_tops_2(k2_tops_2(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),X2),g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v3_tops_2(X2,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))))|~v1_funct_2(X2,u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))|~v1_funct_1(X2)|~l1_pre_topc(k1_pre_topc(X3,u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))))|~m1_pre_topc(X1,k1_pre_topc(X3,u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))))|~m1_subset_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3))).#trainneg
cnf(i_0_4530, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(k1_pre_topc(k15_euclid(np__2),X3))|v2_struct_0(X4)|r1_t_0topsp(X4,k1_pre_topc(k15_euclid(np__2),X3))|v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v3_tops_2(esk4_1(X3),k1_pre_topc(X1,X2),k1_pre_topc(k15_euclid(np__2),X3))|~v3_tops_2(X5,X4,k1_pre_topc(X1,X2))|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X3)))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk4_1(X3),X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X3)))|~v1_funct_2(X5,u1_struct_0(X4),X2)|~v1_funct_1(esk4_1(X3))|~v1_funct_1(X5)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X3))|~l1_pre_topc(X4)|~l1_pre_topc(X1)|~v1_topreal2(X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))))).#trainneg
cnf(i_0_4696, plain, (g1_pre_topc(u1_struct_0(esk5_0),u1_pre_topc(esk5_0))=esk5_0)).#trainneg
cnf(i_0_4689, plain, (v2_struct_0(X1)|v3_tops_2(k2_tops_2(u1_struct_0(esk5_0),u1_struct_0(X1),X2),X1,esk5_0)|~v3_tops_2(X2,esk5_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(esk5_0),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4509, plain, (v2_membered(esk56_1(esk33_0)))).#trainneg
cnf(i_0_4510, plain, (v1_funct_1(esk56_1(esk33_0)))).#trainneg
cnf(i_0_4815, plain, (esk43_1(esk33_0)=esk6_0)).#trainneg
cnf(i_0_4718, negated_conjecture, (esk1_0=esk5_0|k2_struct_0(esk1_0)!=u1_struct_0(esk5_0)|~v1_pre_topc(esk1_0))).#trainneg
cnf(i_0_785, plain, (v1_xcmplx_0(esk43_1(X1))|~v1_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4714, plain, (r1_t_0topsp(esk5_0,esk5_0)|k2_struct_0(esk5_0)!=k1_topreal1)).#trainneg
cnf(i_0_4674, plain, (~v2_tops_1(u1_struct_0(esk5_0),esk5_0))).#trainneg
cnf(i_0_5042, plain, (X1=esk20_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5065, plain, (X1=esk19_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_786, plain, (v1_xreal_0(esk43_1(X1))|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_787, plain, (v1_relat_1(esk43_1(X1))|~v4_funct_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4656, plain, (k2_struct_0(esk5_0)=u1_struct_0(esk5_0))).#trainneg
cnf(i_0_4678, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v3_tops_2(k1_partfun1(u1_struct_0(esk5_0),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(X1),X3,X4),esk5_0,X1)|~v3_tops_2(X3,esk5_0,X2)|~v3_tops_2(X4,X2,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(X2))))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(esk5_0),u1_struct_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_4673, plain, (k1_pre_topc(k15_euclid(np__2),X1)=esk5_0|k2_struct_0(esk5_0)!=X1|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_4788, plain, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X1),k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(k1_topreal1,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~v1_funct_2(esk4_1(X1),k1_topreal1,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_funct_1(esk4_1(X1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_4665, plain, (r1_t_0topsp(g1_pre_topc(u1_struct_0(esk5_0),u1_pre_topc(esk5_0)),esk5_0)|~v2_pre_topc(g1_pre_topc(u1_struct_0(esk5_0),u1_pre_topc(esk5_0))))).#trainneg
cnf(i_0_4676, plain, (r1_t_0topsp(X1,esk5_0)|~v2_pre_topc(X1)|~r1_t_0topsp(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),g1_pre_topc(u1_struct_0(esk5_0),u1_pre_topc(esk5_0)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4779, plain, (v2_struct_0(X1)|r1_t_0topsp(X1,esk5_0)|~v3_tops_2(X2,esk5_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(esk5_0),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5178, plain, (u1_struct_0(X1)=X2|X1!=k1_pre_topc(X3,X2)|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3)|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4511, plain, (v1_valued_0(esk56_1(esk33_0)))).#trainneg
cnf(i_0_4512, plain, (v3_valued_0(esk56_1(esk33_0)))).#trainneg
cnf(i_0_5160, plain, (esk56_1(esk33_0)=esk6_0)).#trainneg
cnf(i_0_4734, negated_conjecture, (esk2_0=esk5_0|k2_struct_0(esk2_0)!=u1_struct_0(esk5_0)|~v1_pre_topc(esk2_0))).#trainneg
cnf(i_0_788, plain, (v1_xxreal_0(esk43_1(X1))|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4707, negated_conjecture, (r1_t_0topsp(esk2_0,esk5_0)|k2_struct_0(esk2_0)!=k1_topreal1|~v1_pre_topc(esk2_0))).#trainneg
cnf(i_0_891, plain, (v2_membered(esk8_0)|~v2_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_895, plain, (v2_membered(k1_topreal1)|~v2_membered(u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_790, plain, (v1_funct_1(esk56_1(X1))|~v4_funct_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_791, plain, (v7_ordinal1(esk56_1(X1))|~v6_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4787, plain, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X1),k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~m1_subset_1(esk4_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))),k1_zfmisc_1(k2_zfmisc_1(k1_topreal1,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~v1_funct_2(esk4_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))),k1_topreal1,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_funct_1(esk4_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~m1_subset_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_4701, plain, (m1_subset_1(esk4_1(u1_struct_0(esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(esk5_0)))))).#trainneg
cnf(i_0_4658, plain, (v3_tops_2(k2_tops_2(u1_struct_0(esk5_0),u1_struct_0(esk5_0),X1),esk5_0,esk5_0)|~v3_tops_2(X1,esk5_0,esk5_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(esk5_0))))|~v1_funct_2(X1,u1_struct_0(esk5_0),u1_struct_0(esk5_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_4680, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v3_tops_2(k1_partfun1(u1_struct_0(X2),u1_struct_0(X1),u1_struct_0(X1),u1_struct_0(esk5_0),X3,X4),X2,esk5_0)|~v3_tops_2(X4,X1,esk5_0)|~v3_tops_2(X3,X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk5_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(esk5_0))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_5120, plain, (v2_struct_0(X1)|r1_t_0topsp(X1,esk5_0)|~v3_tops_2(X2,X1,esk5_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk5_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk5_0))|~v1_funct_1(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4698, plain, (v3_tops_2(esk4_1(u1_struct_0(esk5_0)),k1_pre_topc(k15_euclid(np__2),k1_topreal1),esk5_0))).#trainneg
cnf(i_0_4826, plain, (v1_membered(esk47_0)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4199, plain, (v3_valued_0(k2_zfmisc_1(X1,k1_numbers)))).#trainneg
cnf(i_0_3179, plain, (v1_funct_1(u1_struct_0(esk44_0))|~v2_struct_0(esk44_0))).#trainneg
cnf(i_0_5170, plain, (X1=u1_struct_0(esk5_0)|g1_pre_topc(X1,X2)!=esk5_0|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_792, plain, (v1_xcmplx_0(esk56_1(X1))|~v1_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4710, negated_conjecture, (r1_t_0topsp(esk1_0,esk5_0)|k2_struct_0(esk1_0)!=k1_topreal1|~v1_pre_topc(esk1_0))).#trainneg
cnf(i_0_4827, plain, (v1_membered(esk29_0)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_4955, plain, (v1_membered(esk13_0)|~v3_membered(u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5049, plain, (esk20_1(X1)=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5052, plain, (esk20_1(X1)=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_793, plain, (v1_xreal_0(esk56_1(X1))|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_794, plain, (v1_relat_1(esk56_1(X1))|~v4_funct_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4699, plain, (v1_funct_2(esk4_1(u1_struct_0(esk5_0)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(esk5_0)))).#trainneg
cnf(i_0_4677, plain, (v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(esk5_0),X2),esk5_0,X1)|~v3_tops_2(X2,X1,esk5_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk5_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk5_0))|~v1_funct_1(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4660, plain, (v3_tops_2(k2_tops_2(X1,u1_struct_0(esk5_0),X2),esk5_0,k1_pre_topc(X3,X1))|~v3_tops_2(X2,k1_pre_topc(X3,X1),esk5_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(esk5_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X3)))|~v1_funct_2(X2,X1,u1_struct_0(esk5_0))|~v1_funct_1(X2)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_4663, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r1_t_0topsp(X2,esk5_0)|~v3_tops_2(X3,esk5_0,X1)|~v3_tops_2(X4,X1,X2)|~m1_subset_1(k1_partfun1(u1_struct_0(esk5_0),u1_struct_0(X1),u1_struct_0(X1),u1_struct_0(X2),X3,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(X1))))|~v1_funct_2(k1_partfun1(u1_struct_0(esk5_0),u1_struct_0(X1),u1_struct_0(X1),u1_struct_0(X2),X3,X4),u1_struct_0(esk5_0),u1_struct_0(X2))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(esk5_0),u1_struct_0(X1))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4666, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r1_t_0topsp(X1,esk5_0)|~v3_tops_2(X3,X2,esk5_0)|~v3_tops_2(X4,X1,X2)|~m1_subset_1(k1_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(esk5_0),X4,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk5_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk5_0))))|~v1_funct_2(k1_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(esk5_0),X4,X3),u1_struct_0(X1),u1_struct_0(esk5_0))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk5_0))|~v1_funct_1(X3)|~v1_funct_1(X4)|~l1_pre_topc(X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_4664, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r1_t_0topsp(esk5_0,X2)|~v3_tops_2(X3,esk5_0,X1)|~v3_tops_2(X4,X1,X2)|~m1_subset_1(k1_partfun1(u1_struct_0(esk5_0),u1_struct_0(X1),u1_struct_0(X1),u1_struct_0(X2),X3,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(X1))))|~v1_funct_2(k1_partfun1(u1_struct_0(esk5_0),u1_struct_0(X1),u1_struct_0(X1),u1_struct_0(X2),X3,X4),u1_struct_0(esk5_0),u1_struct_0(X2))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_2(X3,u1_struct_0(esk5_0),u1_struct_0(X1))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l1_pre_topc(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5174, plain, (X1=k2_struct_0(k1_pre_topc(X2,X3))|g1_pre_topc(X1,X4)!=k1_pre_topc(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(k2_struct_0(k1_pre_topc(X2,X3)),k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_4956, plain, (v1_membered(esk8_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_2508, plain, (m1_subset_1(esk43_1(esk17_0),k1_zfmisc_1(esk6_0)))).#trainneg
cnf(i_0_795, plain, (v1_xxreal_0(esk56_1(X1))|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5288, plain, (v2_tops_1(esk6_0,esk33_0))).#trainneg
cnf(i_0_5072, plain, (esk19_1(X1)=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5075, plain, (esk19_1(X1)=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_797, plain, (v1_relat_1(esk49_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_806, plain, (v1_relat_1(esk9_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_3013, plain, (v1_funct_2(esk4_1(X1),k1_topreal1,u1_struct_0(esk5_0))|u1_struct_0(esk5_0)!=X1|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_5455, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(esk5_0))))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|u1_struct_0(esk5_0)!=X1)).#trainneg
cnf(i_0_5809, plain, (v1_xboole_0(esk43_1(esk17_0)))).#trainneg
cnf(i_0_5810, plain, (v3_membered(esk43_1(esk17_0)))).#trainneg
cnf(i_0_809, plain, (v3_pre_topc(u1_struct_0(X1),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_4714, plain, (r1_t_0topsp(esk5_0,esk5_0)|u1_struct_0(esk5_0)!=k1_topreal1)).#trainneg
cnf(i_0_810, plain, (v4_pre_topc(u1_struct_0(X1),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_811, plain, (v1_xboole_0(esk56_1(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5500, plain, (r1_t_0topsp(esk5_0,k1_pre_topc(k15_euclid(np__2),k1_topreal1)))).#trainneg
cnf(i_0_1810, plain, (v1_funct_2(esk4_1(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(esk5_0))|u1_struct_0(esk5_0)!=X1|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_5264, plain, (v2_struct_0(X1)|r1_t_0topsp(esk5_0,X1)|~v2_pre_topc(X1)|~r1_t_0topsp(esk5_0,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5811, plain, (v1_membered(esk43_1(esk17_0)))).#trainneg
cnf(i_0_5812, plain, (v2_membered(esk43_1(esk17_0)))).#trainneg
cnf(i_0_5813, plain, (v1_funct_1(esk43_1(esk17_0)))).#trainneg
cnf(i_0_815, plain, (v1_xboole_0(esk43_1(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5292, negated_conjecture, (esk1_0=esk5_0|u1_struct_0(esk1_0)!=u1_struct_0(esk5_0)|~v1_pre_topc(esk1_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_5332, plain, (~v1_subset_1(u1_struct_0(esk5_0),u1_struct_0(esk5_0))|~l1_struct_0(esk5_0))).#trainneg
cnf(i_0_821, plain, (v1_xboole_0(esk25_2(X1,X2))|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_822, plain, (v1_xboole_0(esk20_1(X1))|~v1_xboole_0(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_3014, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(k1_topreal1,u1_struct_0(esk5_0))))|u1_struct_0(esk5_0)!=X1|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_5275, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(esk5_0,k1_pre_topc(X1,X2))|~v2_pre_topc(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r1_t_0topsp(esk5_0,g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5626, plain, (r1_t_0topsp(esk5_0,X1)|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(k1_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(esk5_0),X3,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk5_0))))|~v1_funct_2(k1_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(esk5_0),X3,X4),u1_struct_0(X1),u1_struct_0(esk5_0))|~v1_funct_1(k1_partfun1(u1_struct_0(X1),u1_struct_0(X2),u1_struct_0(X2),u1_struct_0(esk5_0),X3,X4))|~l1_pre_topc(X1)|~v3_tops_2(X4,X2,esk5_0)|~v3_tops_2(X3,X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk5_0))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(esk5_0))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_5650, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(k1_pre_topc(X1,X2),esk5_0)|v2_struct_0(X3)|~m1_subset_1(k1_partfun1(X2,u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(esk5_0),X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk5_0))))|~v1_funct_2(k1_partfun1(X2,u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(esk5_0),X4,X5),u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk5_0))|~v1_funct_1(k1_partfun1(X2,u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(esk5_0),X4,X5))|~l1_pre_topc(k1_pre_topc(X1,X2))|~v3_tops_2(X4,k1_pre_topc(X1,X2),X3)|~v3_tops_2(X5,X3,esk5_0)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk5_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(esk5_0))|~v1_funct_2(X4,X2,u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5752, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r1_t_0topsp(X1,esk5_0)|~v3_tops_2(X3,esk5_0,X2)|~v3_tops_2(X4,X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k3_relat_1(X3,X4),u1_struct_0(esk5_0),u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(esk5_0),u1_struct_0(X2))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_5769, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r1_t_0topsp(X2,esk5_0)|~v3_tops_2(X3,X1,esk5_0)|~v3_tops_2(X4,X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk5_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k3_relat_1(X4,X3),u1_struct_0(X2),u1_struct_0(esk5_0))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(esk5_0))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~v1_funct_1(X4)|~l1_pre_topc(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5382, plain, (r1_t_0topsp(g1_pre_topc(X1,u1_pre_topc(esk5_0)),esk5_0)|~v2_pre_topc(g1_pre_topc(X1,u1_pre_topc(esk5_0)))|~v1_pre_topc(g1_pre_topc(X1,u1_pre_topc(esk5_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(g1_pre_topc(X1,u1_pre_topc(esk5_0)))|u1_struct_0(esk5_0)!=X1)).#trainneg
cnf(i_0_5644, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(k1_pre_topc(X1,X2),esk5_0)|~m1_subset_1(k2_tops_2(u1_struct_0(esk5_0),X2,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk5_0))))|~v1_funct_2(k2_tops_2(u1_struct_0(esk5_0),X2,X3),u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk5_0))|~v1_funct_1(k2_tops_2(u1_struct_0(esk5_0),X2,X3))|~l1_pre_topc(k1_pre_topc(X1,X2))|~v3_tops_2(X3,esk5_0,k1_pre_topc(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(esk5_0),X2)|~v1_funct_1(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5786, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r1_t_0topsp(esk5_0,X1)|~v3_tops_2(X3,esk5_0,X2)|~v3_tops_2(X4,X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k3_relat_1(X3,X4),u1_struct_0(esk5_0),u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(esk5_0),u1_struct_0(X2))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_5814, plain, (v1_valued_0(esk43_1(esk17_0)))).#trainneg
cnf(i_0_5815, plain, (v3_valued_0(esk43_1(esk17_0)))).#trainneg
cnf(i_0_823, plain, (v1_xboole_0(esk19_1(X1))|~v1_xboole_0(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_5540, negated_conjecture, (esk2_0=esk5_0|u1_struct_0(esk2_0)!=u1_struct_0(esk5_0)|~v1_pre_topc(esk2_0)|~l1_struct_0(esk2_0))).#trainneg
cnf(i_0_5040, plain, (v1_membered(k1_topreal1)|~v3_membered(u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1077, plain, (v3_valued_0(esk20_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_1078, plain, (v3_valued_0(esk19_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_826, plain, (v1_xboole_0(esk9_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_830, plain, (v1_xboole_0(k2_struct_0(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_5638, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|v3_tops_2(k1_partfun1(X2,u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(esk5_0),X4,X5),k1_pre_topc(X1,X2),esk5_0)|~v3_tops_2(X5,X3,esk5_0)|~v3_tops_2(X4,k1_pre_topc(X1,X2),X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk5_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(esk5_0))|~v1_funct_2(X4,X2,u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l1_pre_topc(X3)|~l1_pre_topc(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5498, plain, (r1_t_0topsp(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v1_topreal2(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))|~m1_subset_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_pre_topc(X1,k15_euclid(np__2)))).#trainneg
cnf(i_0_5760, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(X3,esk5_0)|~v3_tops_2(X4,k1_pre_topc(X1,X2),esk5_0)|~v3_tops_2(X5,X3,k1_pre_topc(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(esk5_0))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X2)))|~v1_funct_2(k1_partfun1(u1_struct_0(X3),X2,X2,u1_struct_0(esk5_0),X5,X4),u1_struct_0(X3),u1_struct_0(esk5_0))|~v1_funct_2(X4,X2,u1_struct_0(esk5_0))|~v1_funct_2(X5,u1_struct_0(X3),X2)|~v1_funct_1(X4)|~v1_funct_1(X5)|~l1_pre_topc(X3)|~l1_pre_topc(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5041, plain, (v2_membered(esk47_0)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_5327, plain, (v1_tops_1(u1_struct_0(esk5_0),esk5_0))).#trainneg
cnf(i_0_2050, plain, (v7_ordinal1(esk50_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_832, plain, (v3_membered(esk56_1(X1))|~v3_membered(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5183, plain, (v2_membered(esk29_0)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_5184, plain, (v2_membered(esk13_0)|~v3_membered(u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1301, plain, (v1_relat_1(u1_algstr_0(X1))|~l1_algstr_0(X1))).#trainneg
cnf(i_0_834, plain, (v3_membered(esk49_1(X1))|v1_xboole_0(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_836, plain, (v3_membered(esk43_1(X1))|~v3_membered(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5756, plain, (v2_struct_0(X1)|v2_struct_0(k1_pre_topc(X2,X3))|r1_t_0topsp(k1_pre_topc(X2,X3),esk5_0)|~v3_tops_2(X4,X1,esk5_0)|~v3_tops_2(X5,k1_pre_topc(X2,X3),X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk5_0))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(X1))))|~v1_funct_2(k1_partfun1(X3,u1_struct_0(X1),u1_struct_0(X1),u1_struct_0(esk5_0),X5,X4),X3,u1_struct_0(esk5_0))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(esk5_0))|~v1_funct_2(X5,X3,u1_struct_0(X1))|~v1_funct_1(X4)|~v1_funct_1(X5)|~l1_pre_topc(k1_pre_topc(X2,X3))|~l1_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_5374, plain, (r1_t_0topsp(esk5_0,k1_pre_topc(X2,X3))|~v3_tops_2(X4,k1_pre_topc(X2,X3),esk5_0)|~m1_subset_1(k2_tops_2(X3,X1,X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(k1_pre_topc(X2,X3)))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(k2_tops_2(X3,X1,X4),u1_struct_0(esk5_0),u1_struct_0(k1_pre_topc(X2,X3)))|~v1_funct_2(X4,X3,X1)|~v1_funct_1(X4)|~l1_pre_topc(X2)|u1_struct_0(esk5_0)!=X1)).#trainneg
cnf(i_0_4691, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v3_tops_2(k2_tops_2(u1_struct_0(esk5_0),X2,X3),k1_pre_topc(X1,X2),esk5_0)|~v3_tops_2(X3,esk5_0,k1_pre_topc(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(esk5_0),X2)|~v1_funct_1(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5578, plain, (v1_xboole_0(esk4_1(u1_struct_0(esk5_0)))|~v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))))).#trainneg
cnf(i_0_1809, plain, (v3_tops_2(esk4_1(X1),k1_pre_topc(k15_euclid(np__2),k1_topreal1),esk5_0)|u1_struct_0(esk5_0)!=X1|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_5544, plain, (v2_membered(esk8_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k5_numbers)))).#trainneg
cnf(i_0_5329, plain, (v3_pre_topc(u1_struct_0(esk5_0),esk5_0))).#trainneg
cnf(i_0_2118, plain, (v7_ordinal1(esk31_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_837, plain, (v3_membered(esk36_2(X1,X2))|~v3_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_5545, plain, (v2_membered(k1_topreal1)|~v3_membered(u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2193, plain, (v7_ordinal1(esk27_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3461, plain, (v1_xxreal_0(esk9_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_4325, plain, (v1_funct_1(esk34_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_842, plain, (v3_membered(esk25_2(X1,X2))|~v3_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_843, plain, (v3_membered(esk20_1(X1))|~v3_membered(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_5277, plain, (r1_t_0topsp(esk5_0,k1_pre_topc(X1,X2))|v2_struct_0(k1_pre_topc(X1,X2))|~m1_subset_1(k2_tops_2(u1_struct_0(esk5_0),u1_struct_0(k1_pre_topc(X1,X2)),X3),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(esk5_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(u1_struct_0(esk5_0),u1_struct_0(k1_pre_topc(X1,X2)),X3),X2,u1_struct_0(esk5_0))|~v1_funct_1(k2_tops_2(u1_struct_0(esk5_0),u1_struct_0(k1_pre_topc(X1,X2)),X3))|~l1_pre_topc(X1)|~v3_tops_2(X3,esk5_0,k1_pre_topc(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(k1_pre_topc(X1,X2)))))|~v1_funct_2(X3,u1_struct_0(esk5_0),u1_struct_0(k1_pre_topc(X1,X2)))|~v1_funct_1(X3)|~l1_pre_topc(k1_pre_topc(X1,X2)))).#trainneg
cnf(i_0_5137, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(k1_pre_topc(X3,X4))|r1_t_0topsp(k1_pre_topc(X3,X4),k1_pre_topc(X1,X2))|~m1_subset_1(k2_tops_2(u1_struct_0(k1_pre_topc(X1,X2)),X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X3,X4)),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(u1_struct_0(k1_pre_topc(X1,X2)),X4,X5),u1_struct_0(k1_pre_topc(X3,X4)),X2)|~v1_funct_1(k2_tops_2(u1_struct_0(k1_pre_topc(X1,X2)),X4,X5))|~l1_pre_topc(k1_pre_topc(X3,X4))|~l1_pre_topc(X1)|~v3_tops_2(X5,k1_pre_topc(X1,X2),k1_pre_topc(X3,X4))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),X4)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~v1_funct_2(X5,u1_struct_0(k1_pre_topc(X1,X2)),X4)|~v1_funct_1(X5)|~l1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X3))).#trainneg
cnf(i_0_4795, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(X3,X4))|v2_struct_0(k1_pre_topc(X3,X4))|~m1_subset_1(k2_tops_2(X2,u1_struct_0(k1_pre_topc(X3,X4)),X5),k1_zfmisc_1(k2_zfmisc_1(X4,u1_struct_0(k1_pre_topc(X1,X2)))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~v1_funct_2(k2_tops_2(X2,u1_struct_0(k1_pre_topc(X3,X4)),X5),X4,u1_struct_0(k1_pre_topc(X1,X2)))|~v1_funct_1(k2_tops_2(X2,u1_struct_0(k1_pre_topc(X3,X4)),X5))|~l1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X3)|~v3_tops_2(X5,k1_pre_topc(X1,X2),k1_pre_topc(X3,X4))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(X3,X4)))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X5,X2,u1_struct_0(k1_pre_topc(X3,X4)))|~v1_funct_1(X5)|~l1_pre_topc(k1_pre_topc(X3,X4))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5497, plain, (r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X1),k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v1_topreal2(k2_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~m1_subset_1(k2_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_2257, plain, (v7_ordinal1(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_5330, plain, (v4_pre_topc(u1_struct_0(esk5_0),esk5_0))).#trainneg
cnf(i_0_3169, plain, (v1_xreal_0(esk50_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_844, plain, (v3_membered(esk19_1(X1))|~v3_membered(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_4825, plain, (v1_relat_1(X1)|~l1_struct_0(esk33_0)|~m1_subset_1(X1,esk6_0)|~v2_monoid_0(esk33_0))).#trainneg
cnf(i_0_3110, plain, (v1_xcmplx_0(esk47_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3115, plain, (v1_xcmplx_0(esk29_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_855, plain, (v1_membered(esk49_1(X1))|v1_xboole_0(X1)|~v1_membered(X1))).#trainneg
cnf(i_0_877, plain, (v2_membered(esk49_1(X1))|v1_xboole_0(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_847, plain, (v3_membered(esk9_1(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_851, plain, (v3_membered(k2_struct_0(X1))|~v3_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_4747, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(esk5_0,k1_pre_topc(X1,X2))|~v3_tops_2(X3,esk5_0,k1_pre_topc(X1,X2))|~m1_subset_1(k2_tops_2(u1_struct_0(esk5_0),X2,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk5_0))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(u1_struct_0(esk5_0),X2,X3),u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk5_0))|~v1_funct_2(X3,u1_struct_0(esk5_0),X2)|~v1_funct_1(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5263, plain, (v2_struct_0(X1)|r1_t_0topsp(esk5_0,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X1)|~r1_t_0topsp(esk5_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4203, plain, (r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(k15_euclid(np__2),X1))|~m1_subset_1(esk4_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~v1_funct_2(esk4_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_funct_1(esk4_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v1_topreal2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~m1_subset_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_5426, plain, (g1_pre_topc(X1,u1_pre_topc(esk5_0))=esk5_0|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|u1_struct_0(esk5_0)!=X1)).#trainneg
cnf(i_0_5901, plain, (v3_pre_topc(esk6_0,esk17_0))).#trainneg
cnf(i_0_5888, plain, (esk43_1(esk17_0)=esk6_0)).#trainneg
cnf(i_0_853, plain, (v1_membered(esk56_1(X1))|~v1_membered(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5906, plain, (v4_pre_topc(esk6_0,esk17_0))).#trainneg
cnf(i_0_6124, plain, (X1=esk9_1(k1_zfmisc_1(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_857, plain, (v1_membered(esk43_1(X1))|~v1_membered(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_858, plain, (v1_membered(esk36_2(X1,X2))|~v1_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_6058, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(esk5_0,X3)|~v3_tops_2(X4,k1_pre_topc(X1,X2),esk5_0)|~v3_tops_2(X5,X3,k1_pre_topc(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(esk5_0))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X2)))|~v1_funct_2(k1_partfun1(u1_struct_0(X3),X2,X2,u1_struct_0(esk5_0),X5,X4),u1_struct_0(X3),u1_struct_0(esk5_0))|~v1_funct_2(X4,X2,u1_struct_0(esk5_0))|~v1_funct_2(X5,u1_struct_0(X3),X2)|~v1_funct_1(X4)|~v1_funct_1(X5)|~l1_pre_topc(X3)|~l1_pre_topc(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6340, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(k1_pre_topc(X3,X4))|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(X3,X4))|~v3_tops_2(X5,k1_pre_topc(X3,X4),k1_pre_topc(X1,X2))|~m1_subset_1(k2_tops_2(u1_struct_0(k1_pre_topc(X3,X4)),X2,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(k1_pre_topc(X3,X4)))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X3,X4)),X2)))|~m1_subset_1(u1_struct_0(k1_pre_topc(X3,X4)),k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(u1_struct_0(k1_pre_topc(X3,X4)),X2,X5),u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(k1_pre_topc(X3,X4)))|~v1_funct_2(X5,u1_struct_0(k1_pre_topc(X3,X4)),X2)|~v1_funct_1(X5)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3))))).#trainneg
cnf(i_0_6221, plain, (r1_t_0topsp(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))),k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v1_topreal2(u1_struct_0(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2)))))|~m1_subset_1(u1_struct_0(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2)))),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~m1_pre_topc(k1_pre_topc(X1,X2),k15_euclid(np__2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6213, plain, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|r1_t_0topsp(X1,k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v2_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_topreal2(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))|~m1_subset_1(u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_pre_topc(X1,k15_euclid(np__2)))).#trainneg
cnf(i_0_6067, plain, (v2_struct_0(X1)|v2_struct_0(X2)|r1_t_0topsp(esk5_0,X2)|~v3_tops_2(X3,X1,esk5_0)|~v3_tops_2(X4,X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk5_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(k3_relat_1(X4,X3),u1_struct_0(X2),u1_struct_0(esk5_0))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(esk5_0))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~v1_funct_1(X4)|~l1_pre_topc(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3174, plain, (v1_xreal_0(esk31_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5579, plain, (v1_relat_1(esk4_1(u1_struct_0(esk5_0))))).#trainneg
cnf(i_0_5900, plain, (v3_pre_topc(esk6_0,esk33_0)|~v2_pre_topc(esk33_0))).#trainneg
cnf(i_0_863, plain, (v1_membered(esk25_2(X1,X2))|~v1_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_5905, plain, (v4_pre_topc(esk6_0,esk33_0)|~v2_pre_topc(esk33_0))).#trainneg
cnf(i_0_3175, plain, (v1_xreal_0(esk27_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3193, plain, (v1_xcmplx_0(esk13_0)|~v3_membered(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_864, plain, (v1_membered(esk20_1(X1))|~v1_membered(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_865, plain, (v1_membered(esk19_1(X1))|~v1_membered(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_5176, plain, (X1=u1_struct_0(k1_pre_topc(X2,X3))|g1_pre_topc(X1,X4)!=k1_pre_topc(X2,X3)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(u1_struct_0(k1_pre_topc(X2,X3)),k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_5781, plain, (v2_struct_0(X1)|v2_struct_0(k1_pre_topc(X2,X3))|r1_t_0topsp(esk5_0,X1)|~v3_tops_2(X4,esk5_0,k1_pre_topc(X2,X3))|~v3_tops_2(X5,k1_pre_topc(X2,X3),X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),X3)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(X1))))|~v1_funct_2(k1_partfun1(u1_struct_0(esk5_0),X3,X3,u1_struct_0(X1),X4,X5),u1_struct_0(esk5_0),u1_struct_0(X1))|~v1_funct_2(X4,u1_struct_0(esk5_0),X3)|~v1_funct_2(X5,X3,u1_struct_0(X1))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l1_pre_topc(X1)|~l1_pre_topc(k1_pre_topc(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_5938, plain, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),k1_topreal1),esk5_0)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))).#trainneg
cnf(i_0_3235, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3788, plain, (v1_xcmplx_0(esk50_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1813, plain, (v1_topreal2(X1)|u1_struct_0(esk5_0)!=X1|~v3_tops_2(X2,k1_pre_topc(k15_euclid(np__2),k1_topreal1),esk5_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(esk5_0))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(esk5_0))|~v1_funct_1(X2))).#trainneg
cnf(i_0_868, plain, (v1_membered(esk9_1(k1_zfmisc_1(X1)))|~v1_membered(X1))).#trainneg
cnf(i_0_4954, plain, (v1_funct_1(X1)|~l1_struct_0(esk17_0)|~m1_subset_1(X1,esk6_0)|~v2_monoid_0(esk17_0))).#trainneg
cnf(i_0_3254, plain, (v1_xcmplx_0(esk8_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_3276, plain, (v1_xcmplx_0(k1_topreal1)|~v3_membered(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_872, plain, (v1_membered(k2_struct_0(X1))|~v1_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_4687, plain, (k2_struct_0(X1)=u1_struct_0(esk5_0)|X1!=esk5_0|~v1_pre_topc(X1)|~m1_pre_topc(X1,k15_euclid(np__2)))).#trainneg
cnf(i_0_4174, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(k15_euclid(np__2),X4))|~v1_topreal2(X4)|~v3_tops_2(esk4_1(X4),X3,k1_pre_topc(k15_euclid(np__2),X4))|~v3_tops_2(k2_tops_2(X2,X2,X5),k1_pre_topc(X1,X2),X3)|~m1_subset_1(esk4_1(X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))))|~m1_subset_1(k2_tops_2(X2,X2,X5),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(k2_tops_2(X2,X2,X5),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk4_1(X4),u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))|~v1_funct_2(k2_tops_2(X2,X2,X5),X2,u1_struct_0(X3))|~v1_funct_1(k2_tops_2(X2,X2,X5))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X4))|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~v1_funct_2(X5,X2,X2)|~v1_funct_1(X5))).#trainneg
cnf(i_0_5527, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(k1_pre_topc(X1,X2),esk5_0)|v2_struct_0(X3)|~m1_subset_1(k1_partfun1(u1_struct_0(esk5_0),u1_struct_0(X3),u1_struct_0(X3),X2,X4,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(k1_pre_topc(X1,X2)))))|~v1_funct_2(k1_partfun1(u1_struct_0(esk5_0),u1_struct_0(X3),u1_struct_0(X3),X2,X4,X5),u1_struct_0(esk5_0),u1_struct_0(k1_pre_topc(X1,X2)))|~v1_funct_1(k1_partfun1(u1_struct_0(esk5_0),u1_struct_0(X3),u1_struct_0(X3),X2,X4,X5))|~l1_pre_topc(k1_pre_topc(X1,X2))|~v3_tops_2(X5,X3,k1_pre_topc(X1,X2))|~v3_tops_2(X4,esk5_0,X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(X3))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X4,u1_struct_0(esk5_0),u1_struct_0(X3))|~v1_funct_2(X5,u1_struct_0(X3),X2)|~v1_funct_1(X5)|~v1_funct_1(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4175, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(k15_euclid(np__2),X4))|v1_xboole_0(X5)|~v1_topreal2(X4)|~v3_tops_2(esk4_1(X4),X3,k1_pre_topc(k15_euclid(np__2),X4))|~v3_tops_2(k2_tops_2(X5,X2,X6),k1_pre_topc(X1,X2),X3)|~m1_subset_1(esk4_1(X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))))|~m1_subset_1(k2_tops_2(X5,X2,X6),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(k2_tops_2(X5,X2,X6),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk4_1(X4),u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))|~v1_funct_2(k2_tops_2(X5,X2,X6),X2,u1_struct_0(X3))|~v1_funct_1(k2_tops_2(X5,X2,X6))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X4))|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,X2)))|~v1_funct_2(X6,X5,X2)|~v1_funct_1(X6))).#trainneg
cnf(i_0_3535, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v1_funct_2(k3_relat_1(X1,esk4_1(X2)),X3,X2)|~m1_subset_1(esk4_1(X2),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X2)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~v1_funct_2(X1,X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_funct_1(esk4_1(X2))|~v1_funct_1(X1)|~v1_topreal2(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_3624, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(k1_pre_topc(k15_euclid(np__2),X4))|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X4),k1_pre_topc(X1,X2))|v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v3_tops_2(X5,k1_pre_topc(X1,X2),X3)|~v3_tops_2(esk4_1(X4),X3,k1_pre_topc(k15_euclid(np__2),X4))|~m1_subset_1(esk4_1(X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk4_1(X4),u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))|~v1_funct_2(X5,X2,u1_struct_0(X3))|~v1_funct_1(esk4_1(X4))|~v1_funct_1(X5)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X4))|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v1_topreal2(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(X5,X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))))).#trainneg
cnf(i_0_3789, plain, (v1_xcmplx_0(esk31_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2509, plain, (m1_subset_1(esk56_1(esk17_0),k1_zfmisc_1(esk6_0)))).#trainneg
cnf(i_0_3926, plain, (v1_xcmplx_0(esk27_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1811, plain, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|v1_topreal2(X1)|u1_struct_0(esk5_0)!=X1|~v3_tops_2(X2,esk5_0,k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(X2,u1_struct_0(esk5_0),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_funct_1(X2)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))).#trainneg
cnf(i_0_875, plain, (v2_membered(esk56_1(X1))|~v2_membered(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5039, plain, (v1_relat_1(X1)|~l1_struct_0(esk17_0)|~m1_subset_1(X1,esk6_0)|~v2_monoid_0(esk17_0))).#trainneg
cnf(i_0_3927, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4092, plain, (v1_xxreal_0(esk50_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_996, plain, (v1_funct_1(esk34_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1112, plain, (v1_valued_0(esk34_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_879, plain, (v2_membered(esk43_1(X1))|~v2_membered(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_880, plain, (v2_membered(esk36_2(X1,X2))|~v2_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_4623, plain, (v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X1),X2),g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v3_tops_2(X2,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X3)))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X3)|~m1_pre_topc(X1,X3)|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_5363, plain, (v1_partfun1(esk4_1(u1_struct_0(esk5_0)),k1_topreal1)|~m1_subset_1(esk4_1(u1_struct_0(esk5_0)),k1_zfmisc_1(k2_zfmisc_1(k1_topreal1,u1_struct_0(esk5_0)))))).#trainneg
cnf(i_0_6193, plain, (v2_struct_0(X1)|v3_tops_2(k1_partfun1(u1_struct_0(esk5_0),u1_struct_0(X1),u1_struct_0(X1),u1_struct_0(esk5_0),X2,X3),esk5_0,esk5_0)|~v3_tops_2(X2,esk5_0,X1)|~v3_tops_2(X3,X1,esk5_0)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk5_0))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(esk5_0))|~v1_funct_2(X2,u1_struct_0(esk5_0),u1_struct_0(X1))|~v1_funct_1(X3)|~v1_funct_1(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5032, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(k1_pre_topc(k15_euclid(np__2),X3))|v2_struct_0(X4)|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X3),k1_pre_topc(X1,X2))|~v3_tops_2(X5,k1_pre_topc(X1,X2),X4)|~v3_tops_2(esk4_1(X3),X4,k1_pre_topc(k15_euclid(np__2),X3))|~m1_subset_1(k3_relat_1(X5,esk4_1(X3)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X3)))))|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X3)))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X4))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk4_1(X3),u1_struct_0(X4),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X3)))|~v1_funct_2(X5,X2,u1_struct_0(X4))|~v1_funct_1(esk4_1(X3))|~v1_funct_1(X5)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X3))|~l1_pre_topc(X4)|~l1_pre_topc(X1)|~v1_topreal2(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),k1_topreal1)))|~v1_funct_2(X5,u1_struct_0(k1_pre_topc(X1,X2)),k1_topreal1))).#trainneg
cnf(i_0_6407, plain, (r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),k1_topreal1),esk5_0)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))).#trainneg
cnf(i_0_5268, plain, (v2_struct_0(X1)|r1_t_0topsp(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),esk5_0)|~v2_pre_topc(X1)|~r1_t_0topsp(esk5_0,X1)|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3921, plain, (v7_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2512, plain, (m1_subset_1(esk12_1(esk17_0),k1_zfmisc_1(esk6_0)))).#trainneg
cnf(i_0_4093, plain, (v1_xxreal_0(esk31_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1808, plain, (r1_t_0topsp(esk5_0,k1_pre_topc(k15_euclid(np__2),X1))|u1_struct_0(esk5_0)!=k1_topreal1|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1)))).#trainneg
cnf(i_0_885, plain, (v2_membered(esk25_2(X1,X2))|~v2_membered(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_4201, plain, (v1_xxreal_0(esk27_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4202, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1137, plain, (v3_valued_0(esk34_1(X1))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_977, plain, (v1_valued_0(esk36_2(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_886, plain, (v2_membered(esk20_1(X1))|~v2_membered(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_887, plain, (v2_membered(esk19_1(X1))|~v2_membered(k2_zfmisc_1(X1,X1)))).#trainneg
cnf(i_0_6349, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(k1_pre_topc(X3,X4))|r1_t_0topsp(k1_pre_topc(X3,X4),k1_pre_topc(X1,X2))|~v3_tops_2(X5,k1_pre_topc(X3,X4),k1_pre_topc(X1,X2))|~m1_subset_1(k2_tops_2(X4,u1_struct_0(k1_pre_topc(X1,X2)),X5),k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,u1_struct_0(k1_pre_topc(X1,X2)))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~v1_funct_2(k2_tops_2(X4,u1_struct_0(k1_pre_topc(X1,X2)),X5),X2,X4)|~v1_funct_2(X5,X4,u1_struct_0(k1_pre_topc(X1,X2)))|~v1_funct_1(X5)|~l1_pre_topc(X1)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_5535, plain, (u1_struct_0(X1)=k2_struct_0(k1_pre_topc(X2,X3))|X1!=k1_pre_topc(X2,X3)|~v1_pre_topc(X1)|~m1_subset_1(k2_struct_0(k1_pre_topc(X2,X3)),k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_4523, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(k1_pre_topc(X3,X4))|v2_struct_0(X5)|r1_t_0topsp(X5,k1_pre_topc(X3,X4))|~v3_tops_2(X6,k1_pre_topc(X1,X2),k1_pre_topc(X3,X4))|~v3_tops_2(X7,X5,k1_pre_topc(X1,X2))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k3_relat_1(X7,X6),u1_struct_0(X5),X4)|~v1_funct_2(X6,X2,X4)|~v1_funct_2(X7,u1_struct_0(X5),X2)|~v1_funct_1(X6)|~v1_funct_1(X7)|~l1_pre_topc(k1_pre_topc(X3,X4))|~l1_pre_topc(X5)|~l1_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3))).#trainneg
cnf(i_0_5402, plain, (v2_struct_0(X2)|v2_struct_0(X3)|v3_tops_2(k1_partfun1(X1,u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(X2),X4,X5),esk5_0,X2)|~v3_tops_2(X4,esk5_0,X3)|~v3_tops_2(X5,X3,X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(X4,X1,u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l1_pre_topc(X2)|~l1_pre_topc(X3)|u1_struct_0(esk5_0)!=X1)).#trainneg
cnf(i_0_5515, plain, (r1_t_0topsp(X1,esk5_0)|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~r1_t_0topsp(esk5_0,X1)|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1))))).#trainneg
cnf(i_0_4639, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(k1_pre_topc(X1,X2),g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~l1_pre_topc(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~l1_pre_topc(k1_pre_topc(X1,X2))|~v2_pre_topc(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~v2_pre_topc(k1_pre_topc(X1,X2))|~v1_pre_topc(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6333, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(k1_pre_topc(X3,X4))|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(X3,X4))|~v3_tops_2(X5,k1_pre_topc(X3,X4),k1_pre_topc(X1,X2))|~m1_subset_1(k2_tops_2(u1_struct_0(k1_pre_topc(X3,X4)),X2,X5),k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X3,X4)),X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(u1_struct_0(k1_pre_topc(X3,X4)),X2,X5),X2,X4)|~v1_funct_2(X5,u1_struct_0(k1_pre_topc(X3,X4)),X2)|~v1_funct_1(X5)|~l1_pre_topc(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4198, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2513, plain, (m1_subset_1(esk40_1(esk17_0),k1_zfmisc_1(esk6_0)))).#trainneg
cnf(i_0_1801, plain, (v3_tops_2(esk4_1(X1),esk5_0,k1_pre_topc(k15_euclid(np__2),X1))|u1_struct_0(esk5_0)!=k1_topreal1|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_890, plain, (v2_membered(esk9_1(k1_zfmisc_1(X1)))|~v2_membered(X1))).#trainneg
cnf(i_0_4319, plain, (v3_membered(esk47_0)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4339, plain, (v1_xxreal_0(esk47_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4519, plain, (v1_xxreal_0(esk29_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_980, plain, (v1_valued_0(esk25_2(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_1052, plain, (v1_xboole_0(esk25_2(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_894, plain, (v2_membered(k2_struct_0(X1))|~v2_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_906, plain, (v1_funct_1(u1_pre_topc(X1))|~v4_funct_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_4353, plain, (v2_struct_0(X1)|v3_tops_2(k2_tops_2(u1_struct_0(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))),u1_struct_0(X1),X4),X1,g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3)))|~v3_tops_2(X4,g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3)),X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))),u1_struct_0(X1))))|~m1_subset_1(u1_struct_0(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))),k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X4,u1_struct_0(g1_pre_topc(u1_struct_0(X3),u1_pre_topc(X3))),u1_struct_0(X1))|~v1_funct_1(X4)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~m1_pre_topc(X3,X2))).#trainneg
cnf(i_0_5620, plain, (r1_t_0topsp(esk5_0,k1_pre_topc(X1,X2))|v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|~m1_subset_1(k1_partfun1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(esk5_0),X4,X5),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(esk5_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k1_partfun1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(esk5_0),X4,X5),X2,u1_struct_0(esk5_0))|~v1_funct_1(k1_partfun1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X3),u1_struct_0(X3),u1_struct_0(esk5_0),X4,X5))|~l1_pre_topc(X1)|~v3_tops_2(X5,X3,esk5_0)|~v3_tops_2(X4,k1_pre_topc(X1,X2),X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk5_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X3))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(esk5_0))|~v1_funct_2(X4,u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l1_pre_topc(X3)|~l1_pre_topc(k1_pre_topc(X1,X2)))).#trainneg
cnf(i_0_4804, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2519, plain, (v1_pre_topc(g1_pre_topc(esk6_0,u1_pre_topc(esk17_0))))).#trainneg
cnf(i_0_4517, plain, (v3_membered(esk29_0)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3529, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),X1)))|u1_struct_0(esk5_0)!=k1_topreal1|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_907, plain, (v7_ordinal1(u1_pre_topc(X1))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6484, plain, (v2_tops_1(esk6_0,esk17_0))).#trainneg
cnf(i_0_6709, plain, (v1_xboole_0(esk56_1(esk17_0)))).#trainneg
cnf(i_0_6710, plain, (v3_membered(esk56_1(esk17_0)))).#trainneg
cnf(i_0_6967, plain, (X1=esk25_2(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_6976, plain, (esk25_2(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X4)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_908, plain, (v1_xcmplx_0(u1_pre_topc(X1))|~v1_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_909, plain, (v1_xreal_0(u1_pre_topc(X1))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6735, negated_conjecture, (v3_tops_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1),g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|~v3_tops_2(X1,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~m1_subset_1(u1_pre_topc(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_6568, plain, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|r1_t_0topsp(X1,k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v2_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v2_pre_topc(X1)|~v1_topreal2(u1_struct_0(X1))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(X1)|~m1_pre_topc(X1,k15_euclid(np__2))|~v1_pre_topc(X1))).#trainneg
cnf(i_0_6734, negated_conjecture, (v3_tops_2(k2_tops_2(u1_struct_0(esk2_0),u1_struct_0(esk2_0),X1),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~v3_tops_2(X1,g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))|~m1_subset_1(u1_pre_topc(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk2_0),u1_struct_0(esk2_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_6559, plain, (r1_t_0topsp(k1_pre_topc(X2,X1),k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v1_topreal2(u1_struct_0(k1_pre_topc(X2,X1)))|~m1_subset_1(u1_struct_0(k1_pre_topc(X2,X1)),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(k1_pre_topc(X2,X1))|~l1_pre_topc(X2)|~m1_pre_topc(k1_pre_topc(X2,X1),k15_euclid(np__2)))).#trainneg
cnf(i_0_6671, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|r1_t_0topsp(X1,X3)|v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v3_tops_2(esk4_1(u1_struct_0(X3)),X2,X3)|~v3_tops_2(X4,X1,X2)|~m1_subset_1(esk4_1(u1_struct_0(X3)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(esk4_1(u1_struct_0(X3)),u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(esk4_1(u1_struct_0(X3)))|~v1_funct_1(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v1_topreal2(u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(u1_struct_0(X3),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))))).#trainneg
cnf(i_0_6564, plain, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v2_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v2_pre_topc(k1_pre_topc(X1,X2))|~v1_topreal2(u1_struct_0(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2)))))|~m1_subset_1(u1_struct_0(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2)))),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~l1_pre_topc(k1_pre_topc(X1,X2))|~m1_pre_topc(k1_pre_topc(X1,X2),k15_euclid(np__2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6711, plain, (v1_membered(esk56_1(esk17_0)))).#trainneg
cnf(i_0_6712, plain, (v2_membered(esk56_1(esk17_0)))).#trainneg
cnf(i_0_6713, plain, (v1_funct_1(esk56_1(esk17_0)))).#trainneg
cnf(i_0_1802, plain, (v1_funct_2(esk4_1(X1),u1_struct_0(esk5_0),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|u1_struct_0(esk5_0)!=k1_topreal1|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_910, plain, (v1_relat_1(u1_pre_topc(X1))|~v4_funct_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6714, plain, (v1_valued_0(esk56_1(esk17_0)))).#trainneg
cnf(i_0_6715, plain, (v3_valued_0(esk56_1(esk17_0)))).#trainneg
cnf(i_0_6981, plain, (esk25_2(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1064, plain, (v1_xboole_0(esk25_2(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_911, plain, (v1_xxreal_0(u1_pre_topc(X1))|~v2_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6556, plain, (r1_t_0topsp(g1_pre_topc(k2_struct_0(k1_pre_topc(X1,X2)),u1_pre_topc(k1_pre_topc(X1,X2))),k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v1_topreal2(u1_struct_0(g1_pre_topc(k2_struct_0(k1_pre_topc(X1,X2)),u1_pre_topc(k1_pre_topc(X1,X2)))))|~m1_subset_1(u1_struct_0(g1_pre_topc(k2_struct_0(k1_pre_topc(X1,X2)),u1_pre_topc(k1_pre_topc(X1,X2)))),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(k2_struct_0(k1_pre_topc(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(g1_pre_topc(k2_struct_0(k1_pre_topc(X1,X2)),u1_pre_topc(k1_pre_topc(X1,X2))))|~l1_pre_topc(X1)|~m1_pre_topc(k1_pre_topc(X1,X2),k15_euclid(np__2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_6700, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X4),k1_pre_topc(X1,X2))|~v1_topreal2(X4)|~v3_tops_2(esk4_1(X4),X3,k1_pre_topc(k15_euclid(np__2),X4))|~v3_tops_2(k2_tops_2(X2,X2,X5),k1_pre_topc(X1,X2),X3)|~m1_subset_1(esk4_1(X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))))|~m1_subset_1(k2_tops_2(X2,X2,X5),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(k2_tops_2(X2,X2,X5),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk4_1(X4),u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))|~v1_funct_2(k2_tops_2(X2,X2,X5),X2,u1_struct_0(X3))|~v1_funct_1(k2_tops_2(X2,X2,X5))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X4))|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))|~v1_funct_2(X5,X2,X2)|~v1_funct_1(X5))).#trainneg
cnf(i_0_6737, plain, (v3_tops_2(k2_tops_2(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(k1_pre_topc(X1,X2)),X3),g1_pre_topc(u1_struct_0(k1_pre_topc(X1,X2)),u1_pre_topc(k1_pre_topc(X1,X2))),g1_pre_topc(u1_struct_0(k1_pre_topc(X1,X2)),u1_pre_topc(k1_pre_topc(X1,X2))))|~v3_tops_2(X3,g1_pre_topc(u1_struct_0(k1_pre_topc(X1,X2)),u1_pre_topc(k1_pre_topc(X1,X2))),g1_pre_topc(u1_struct_0(k1_pre_topc(X1,X2)),u1_pre_topc(k1_pre_topc(X1,X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(k1_pre_topc(X1,X2)))))|~m1_subset_1(u1_pre_topc(k1_pre_topc(X1,X2)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)))))|~m1_subset_1(u1_struct_0(k1_pre_topc(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(k1_pre_topc(X1,X2)))|~v1_funct_1(X3)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_5193, plain, (v3_tops_2(k2_tops_2(u1_struct_0(X1),u1_struct_0(X1),X2),g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v3_tops_2(X2,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X3)))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X3)|~m1_pre_topc(X1,k1_pre_topc(X3,u1_struct_0(X1)))|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_3832, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(X4)|r1_t_0topsp(X3,X4)|~v3_tops_2(X5,k1_pre_topc(X1,X2),X3)|~v3_tops_2(X6,X4,k1_pre_topc(X1,X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k3_relat_1(X6,X5),u1_struct_0(X4),u1_struct_0(X3))|~v1_funct_2(X5,X2,u1_struct_0(X3))|~v1_funct_2(X6,u1_struct_0(X4),X2)|~v1_funct_1(X5)|~v1_funct_1(X6)|~l1_pre_topc(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_3777, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(X1,X2)))|~v2_struct_0(k1_pre_topc(X1,X2))|~m1_subset_1(u1_struct_0(k1_pre_topc(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_6798, plain, (v1_xboole_0(esk12_1(esk17_0)))).#trainneg
cnf(i_0_6799, plain, (v3_membered(esk12_1(esk17_0)))).#trainneg
cnf(i_0_6800, plain, (v1_membered(esk12_1(esk17_0)))).#trainneg
cnf(i_0_1804, plain, (m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk5_0),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|u1_struct_0(esk5_0)!=k1_topreal1|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_914, plain, (v3_membered(u1_pre_topc(X1))|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6801, plain, (v2_membered(esk12_1(esk17_0)))).#trainneg
cnf(i_0_6802, plain, (v1_funct_1(esk12_1(esk17_0)))).#trainneg
cnf(i_0_1073, plain, (v3_valued_0(esk36_2(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_1076, plain, (v3_valued_0(esk25_2(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_915, plain, (v1_membered(u1_pre_topc(X1))|~v1_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_916, plain, (v2_membered(u1_pre_topc(X1))|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6697, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X4),k1_pre_topc(X1,X2))|~v1_topreal2(X4)|~v3_tops_2(esk4_1(X4),X3,k1_pre_topc(k15_euclid(np__2),X4))|~v3_tops_2(k3_relat_1(X5,esk4_1(k1_topreal1)),k1_pre_topc(X1,X2),X3)|~m1_subset_1(esk4_1(X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))))|~m1_subset_1(k3_relat_1(X5,esk4_1(k1_topreal1)),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(k3_relat_1(X5,esk4_1(k1_topreal1)),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk4_1(X4),u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))|~v1_funct_2(k3_relat_1(X5,esk4_1(k1_topreal1)),X2,u1_struct_0(X3))|~v1_funct_1(k3_relat_1(X5,esk4_1(k1_topreal1)))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X4))|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~v1_funct_2(X5,X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_funct_1(X5))).#trainneg
cnf(i_0_4145, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(k1_pre_topc(X1,X2),X3)|~v3_tops_2(X4,X3,k1_pre_topc(X1,X2))|~m1_subset_1(k2_tops_2(u1_struct_0(X3),X2,X4),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(u1_struct_0(X3),X2,X4),X2,u1_struct_0(X3))|~v1_funct_2(X4,u1_struct_0(X3),X2)|~v1_funct_1(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_7525, plain, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X1),k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~m1_subset_1(esk4_1(X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X1)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(esk4_1(X1),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X1)|~v1_funct_1(esk4_1(X1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v1_topreal2(X1))).#trainneg
cnf(i_0_6803, plain, (v1_valued_0(esk12_1(esk17_0)))).#trainneg
cnf(i_0_6804, plain, (v3_valued_0(esk12_1(esk17_0)))).#trainneg
cnf(i_0_7097, plain, (esk56_1(esk17_0)=esk6_0)).#trainneg
cnf(i_0_3931, plain, (v1_funct_2(k3_relat_1(X1,esk4_1(X2)),X3,u1_struct_0(esk5_0))|u1_struct_0(esk5_0)!=X2|~v1_topreal2(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,k1_topreal1)))|~v1_funct_2(X1,X3,k1_topreal1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_918, plain, (v4_relat_1(esk49_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_4609, plain, (v3_membered(esk8_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6941, plain, (v1_xboole_0(esk40_1(esk17_0)))).#trainneg
cnf(i_0_6942, plain, (v3_membered(esk40_1(esk17_0)))).#trainneg
cnf(i_0_7266, plain, (X1=esk25_2(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_7443, plain, (X1=esk12_1(esk17_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_921, plain, (v4_relat_1(esk34_1(k2_zfmisc_1(X1,X2)),X1))).#trainneg
cnf(i_0_927, plain, (v4_relat_1(esk9_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_7617, negated_conjecture, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(esk2_0,k1_pre_topc(X1,X2))|~v2_pre_topc(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)),g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~l1_pre_topc(X1)|~m1_subset_1(u1_pre_topc(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_7538, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(k1_pre_topc(X1,X2),X3)|v2_struct_0(X3)|~m1_subset_1(k2_tops_2(X2,u1_struct_0(X3),X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(X2,u1_struct_0(X3),X4),u1_struct_0(X3),X2)|~v1_funct_1(k2_tops_2(X2,u1_struct_0(X3),X4))|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v3_tops_2(X4,k1_pre_topc(X1,X2),X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~v1_funct_2(X4,X2,u1_struct_0(X3))|~v1_funct_1(X4))).#trainneg
cnf(i_0_7331, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X1),k1_pre_topc(k15_euclid(np__2),X2))|~v1_topreal2(X2)|~m1_subset_1(esk4_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(esk4_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X1)|~v1_funct_1(esk4_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X2))|~v1_topreal2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~m1_subset_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_7524, plain, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X1),k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~m1_subset_1(esk4_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X1)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(esk4_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),X1)|~v1_funct_1(esk4_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v1_topreal2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~m1_subset_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_7543, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(k1_pre_topc(X1,X2),X3)|v2_struct_0(X3)|v2_struct_0(X4)|~m1_subset_1(k1_partfun1(u1_struct_0(X3),u1_struct_0(X4),u1_struct_0(X4),u1_struct_0(k1_pre_topc(X1,X2)),X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k1_partfun1(u1_struct_0(X3),u1_struct_0(X4),u1_struct_0(X4),u1_struct_0(k1_pre_topc(X1,X2)),X5,X6),u1_struct_0(X3),X2)|~v1_funct_1(k1_partfun1(u1_struct_0(X3),u1_struct_0(X4),u1_struct_0(X4),u1_struct_0(k1_pre_topc(X1,X2)),X5,X6))|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v3_tops_2(X6,X4,k1_pre_topc(X1,X2))|~v3_tops_2(X5,X3,X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(k1_pre_topc(X1,X2)))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X4))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(k1_pre_topc(X1,X2)))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~l1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X4))).#trainneg
cnf(i_0_7428, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(k1_pre_topc(k15_euclid(np__2),X4))|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X4),X3)|v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v3_tops_2(esk4_1(X4),k1_pre_topc(X1,X2),k1_pre_topc(k15_euclid(np__2),X4))|~v3_tops_2(X5,X3,k1_pre_topc(X1,X2))|~m1_subset_1(esk4_1(X4),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk4_1(X4),X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))|~v1_funct_2(X5,u1_struct_0(X3),X2)|~v1_funct_1(esk4_1(X4))|~v1_funct_1(X5)|~l1_pre_topc(X3)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X4))|~l1_pre_topc(X1)|~v1_topreal2(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))))).#trainneg
cnf(i_0_6943, plain, (v1_membered(esk40_1(esk17_0)))).#trainneg
cnf(i_0_6944, plain, (v2_membered(esk40_1(esk17_0)))).#trainneg
cnf(i_0_7461, plain, (esk12_1(esk17_0)=esk6_0)).#trainneg
cnf(i_0_2302, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v1_funct_2(k3_relat_1(X1,esk4_1(X2)),X3,u1_struct_0(esk5_0))|u1_struct_0(esk5_0)!=X2|~v1_topreal2(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(X1,X3,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_930, plain, (v5_relat_1(esk49_1(k2_zfmisc_1(X1,X2)),X2)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_5664, plain, (v1_membered(esk47_0)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6945, plain, (v1_funct_1(esk40_1(esk17_0)))).#trainneg
cnf(i_0_6946, plain, (v1_valued_0(esk40_1(esk17_0)))).#trainneg
cnf(i_0_933, plain, (v5_relat_1(esk34_1(k2_zfmisc_1(X1,X2)),X2))).#trainneg
cnf(i_0_939, plain, (v5_relat_1(esk9_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_7226, plain, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v2_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v2_pre_topc(k1_pre_topc(X1,X2))|~v1_topreal2(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~l1_pre_topc(X1)|~m1_pre_topc(k1_pre_topc(X1,X2),k15_euclid(np__2))|~m1_subset_1(u1_pre_topc(k1_pre_topc(X1,X2)),k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_7421, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(k1_pre_topc(X4,X5))|r1_t_0topsp(k1_pre_topc(X4,X5),X3)|~v3_tops_2(X6,k1_pre_topc(X1,X2),k1_pre_topc(X4,X5))|~v3_tops_2(X7,X3,k1_pre_topc(X1,X2))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X5)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k3_relat_1(X7,X6),u1_struct_0(X3),X5)|~v1_funct_2(X6,X2,X5)|~v1_funct_2(X7,u1_struct_0(X3),X2)|~v1_funct_1(X6)|~v1_funct_1(X7)|~l1_pre_topc(X3)|~l1_pre_topc(k1_pre_topc(X4,X5))|~l1_pre_topc(X1)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X4)))|~l1_pre_topc(X4))).#trainneg
cnf(i_0_7198, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(k1_pre_topc(X3,X4))|r1_t_0topsp(k1_pre_topc(X3,X4),X1)|~v1_topreal2(u1_struct_0(X1))|~v3_tops_2(esk4_1(u1_struct_0(X1)),X2,X1)|~v3_tops_2(X5,k1_pre_topc(X3,X4),X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(esk4_1(u1_struct_0(X1)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,u1_struct_0(X2))))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(X5,X4,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_funct_2(esk4_1(u1_struct_0(X1)),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_2(X5,X4,u1_struct_0(X2))|~v1_funct_1(X5)|~l1_pre_topc(X1)|~l1_pre_topc(k1_pre_topc(X3,X4))|~l1_pre_topc(X2)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3))).#trainneg
cnf(i_0_7519, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(k1_pre_topc(X1,X2),X3)|~v3_tops_2(X4,X3,k1_pre_topc(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k1_pre_topc(X1,X2)))))|~m1_subset_1(u1_struct_0(k1_pre_topc(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(k1_pre_topc(X1,X2)))|~v1_funct_1(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_7583, negated_conjecture, (r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)),esk2_0)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~m1_subset_1(u1_pre_topc(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_7536, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(k1_pre_topc(X1,X2),X3)|v2_struct_0(X3)|~m1_subset_1(k2_tops_2(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X3),X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X3),X4),u1_struct_0(X3),X2)|~v1_funct_1(k2_tops_2(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X3),X4))|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v3_tops_2(X4,k1_pre_topc(X1,X2),X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X3))))|~m1_subset_1(u1_struct_0(k1_pre_topc(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X4,u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X3))|~v1_funct_1(X4))).#trainneg
cnf(i_0_6947, plain, (v3_valued_0(esk40_1(esk17_0)))).#trainneg
cnf(i_0_2520, plain, (v2_pre_topc(g1_pre_topc(esk6_0,u1_pre_topc(esk17_0))))).#trainneg
cnf(i_0_944, plain, (u1_pre_topc(X1)=X2|X1!=g1_pre_topc(X3,X2)|~m1_subset_1(u1_pre_topc(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5672, plain, (v1_membered(esk29_0)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4824, plain, (v1_xxreal_0(esk13_0)|~v3_membered(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_4850, plain, (v1_xxreal_0(esk8_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers))))).#trainneg
cnf(i_0_7275, plain, (esk25_2(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X4)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7280, plain, (esk25_2(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_948, plain, (v3_valued_0(esk9_1(u1_struct_0(k15_euclid(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_7204, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|r1_t_0topsp(X3,X1)|~v1_topreal2(u1_struct_0(X1))|~v3_tops_2(esk4_1(u1_struct_0(X1)),X2,X1)|~v3_tops_2(esk16_1(X3),X3,X2)|~m1_subset_1(esk16_1(X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(esk4_1(u1_struct_0(X1)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(esk16_1(X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(esk4_1(u1_struct_0(X1)),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_2(esk16_1(X3),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(esk16_1(X3))|~l1_pre_topc(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_7194, plain, (v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|r1_t_0topsp(X3,X1)|~v1_topreal2(u1_struct_0(X1))|~v3_tops_2(esk4_1(u1_struct_0(X1)),X2,X1)|~v3_tops_2(X4,X3,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),k1_topreal1)))|~m1_subset_1(esk4_1(u1_struct_0(X1)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(X4,u1_struct_0(X3),k1_topreal1)|~v1_funct_2(esk4_1(u1_struct_0(X1)),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X4)|~l1_pre_topc(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_7559, negated_conjecture, (u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))=u1_pre_topc(esk2_0)|~m1_subset_1(u1_pre_topc(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_7596, negated_conjecture, (r1_t_0topsp(esk2_0,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~m1_subset_1(u1_pre_topc(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_7215, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|r1_t_0topsp(X3,X1)|~v1_topreal2(u1_struct_0(X1))|~v3_tops_2(esk4_1(u1_struct_0(X1)),X2,X1)|~v3_tops_2(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(X3),X4),X3,X2)|~m1_subset_1(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(X3),X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(esk4_1(u1_struct_0(X1)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(X3),X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(esk4_1(u1_struct_0(X1)),u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_2(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(X3),X4),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(k2_tops_2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(X3),X4))|~l1_pre_topc(X1)|~l1_pre_topc(X3)|~l1_pre_topc(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(X3))))|~v1_funct_2(X4,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(X3))|~v1_funct_1(X4))).#trainneg
cnf(i_0_7517, plain, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(k1_pre_topc(X1,X2),X3)|~v3_tops_2(X4,X3,k1_pre_topc(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),k2_struct_0(k1_pre_topc(X1,X2)))))|~m1_subset_1(k2_struct_0(k1_pre_topc(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X4,u1_struct_0(X3),k2_struct_0(k1_pre_topc(X1,X2)))|~v1_funct_1(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_6701, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X4),k1_pre_topc(X1,X2))|v1_xboole_0(X5)|~v1_topreal2(X4)|~v3_tops_2(esk4_1(X4),X3,k1_pre_topc(k15_euclid(np__2),X4))|~v3_tops_2(k2_tops_2(X5,X2,X6),k1_pre_topc(X1,X2),X3)|~m1_subset_1(esk4_1(X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))))|~m1_subset_1(k2_tops_2(X5,X2,X6),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(k2_tops_2(X5,X2,X6),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk4_1(X4),u1_struct_0(X3),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X4)))|~v1_funct_2(k2_tops_2(X5,X2,X6),X2,u1_struct_0(X3))|~v1_funct_1(k2_tops_2(X5,X2,X6))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X4))|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X5,X2)))|~v1_funct_2(X6,X5,X2)|~v1_funct_1(X6))).#trainneg
cnf(i_0_3917, plain, (v2_struct_0(k1_pre_topc(X1,u1_struct_0(X2)))|v2_struct_0(k1_pre_topc(k15_euclid(np__2),X3))|v2_struct_0(X4)|r1_t_0topsp(k1_pre_topc(X1,u1_struct_0(X2)),k1_pre_topc(k15_euclid(np__2),X3))|v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|v2_struct_0(X2)|~v3_tops_2(X5,k1_pre_topc(X1,u1_struct_0(X2)),X4)|~v3_tops_2(esk4_1(X3),X4,k1_pre_topc(k15_euclid(np__2),X3))|~m1_subset_1(esk4_1(X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X3)))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X4))))|~m1_subset_1(u1_struct_0(X2),k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk4_1(X3),u1_struct_0(X4),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X3)))|~v1_funct_2(X5,u1_struct_0(X2),u1_struct_0(X4))|~v1_funct_1(esk4_1(X3))|~v1_funct_1(X5)|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X3))|~l1_pre_topc(X4)|~l1_pre_topc(X1)|~v2_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v2_pre_topc(k1_pre_topc(k15_euclid(np__2),X3))|~v2_pre_topc(X2)|~v1_topreal2(X3)|~v3_tops_2(X5,X2,k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(X5,u1_struct_0(X2),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_7699, plain, (esk40_1(esk17_0)=esk6_0)).#trainneg
cnf(i_0_5229, plain, (X1=u1_pre_topc(esk5_0)|g1_pre_topc(X2,X1)!=esk5_0|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_949, plain, (v3_valued_0(u2_struct_0(k15_euclid(X1)))|~v7_ordinal1(X1)|~l2_struct_0(k15_euclid(X1)))).#trainneg
cnf(i_0_5801, plain, (v1_membered(esk8_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4953, plain, (v1_xxreal_0(k1_topreal1)|~v3_membered(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_950, plain, (v1_finseq_1(esk9_1(u1_struct_0(k15_euclid(X1))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_951, plain, (v1_finseq_1(u2_struct_0(k15_euclid(X1)))|~v7_ordinal1(X1)|~l2_struct_0(k15_euclid(X1)))).#trainneg
cnf(i_0_7987, plain, (u1_pre_topc(g1_pre_topc(X1,X2))=X2|~v1_pre_topc(g1_pre_topc(X1,X2))|~l1_pre_topc(g1_pre_topc(X1,X2)))).#trainneg
cnf(i_0_7774, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(k1_pre_topc(X1,X2),X3)|~v3_tops_2(X4,k1_pre_topc(X1,X2),X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X3))))|~m1_subset_1(u1_struct_0(k1_pre_topc(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X4,u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X3))|~v1_funct_1(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_7790, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(k1_pre_topc(X3,X4))|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(X3,X4))|~m1_subset_1(k2_tops_2(X4,X2,X5),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(X3,X4)))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(X4,X2,X5),X2,u1_struct_0(k1_pre_topc(X3,X4)))|~v1_funct_1(k2_tops_2(X4,X2,X5))|~l1_pre_topc(k1_pre_topc(X3,X4))|~l1_pre_topc(X1)|~v3_tops_2(X5,k1_pre_topc(X3,X4),k1_pre_topc(X1,X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~v1_funct_2(X5,X4,X2)|~v1_funct_1(X5)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_8041, negated_conjecture, (u1_struct_0(X1)=u1_struct_0(esk2_0)|X1!=k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))|~v1_pre_topc(X1)|~l1_pre_topc(X1)|~m1_subset_1(u1_pre_topc(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_4947, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2524, plain, (v1_funct_2(esk16_1(esk17_0),esk6_0,esk6_0))).#trainneg
cnf(i_0_6234, plain, (v2_membered(esk47_0)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7547, negated_conjecture, (~v2_struct_0(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~m1_subset_1(u1_pre_topc(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_1565, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_953, plain, (v7_ordinal1(esk12_1(X1))|~v6_membered(k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_954, plain, (v1_xcmplx_0(esk12_1(X1))|~v1_membered(k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_7772, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(k1_pre_topc(X1,X2),X3)|~v3_tops_2(X4,k1_pre_topc(X1,X2),X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k2_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X3))))|~m1_subset_1(k2_struct_0(k1_pre_topc(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X4,k2_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X3))|~v1_funct_1(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_7918, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(k1_pre_topc(X3,X4))|v2_struct_0(X5)|r1_t_0topsp(k1_pre_topc(X1,X2),X5)|~v3_tops_2(X7,X5,k1_pre_topc(X3,X4))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),X4)))|~m1_subset_1(k2_tops_2(X2,u1_struct_0(k1_pre_topc(X3,X4)),X6),k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k3_relat_1(X7,k2_tops_2(X2,u1_struct_0(k1_pre_topc(X3,X4)),X6)),u1_struct_0(X5),X2)|~v1_funct_2(X7,u1_struct_0(X5),X4)|~v1_funct_2(k2_tops_2(X2,u1_struct_0(k1_pre_topc(X3,X4)),X6),X4,X2)|~v1_funct_1(k2_tops_2(X2,u1_struct_0(k1_pre_topc(X3,X4)),X6))|~v1_funct_1(X7)|~l1_pre_topc(X5)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v3_tops_2(X6,k1_pre_topc(X1,X2),k1_pre_topc(X3,X4))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(X3,X4)))))|~v1_funct_2(X6,X2,u1_struct_0(k1_pre_topc(X3,X4)))|~v1_funct_1(X6)|~l1_pre_topc(k1_pre_topc(X3,X4)))).#trainneg
cnf(i_0_7850, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(X4)|r1_t_0topsp(k1_pre_topc(X1,X2),X4)|~v3_tops_2(X5,X3,k1_pre_topc(X1,X2))|~v3_tops_2(X6,X4,X3)|~m1_subset_1(k3_relat_1(X6,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(k1_pre_topc(X1,X2)))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X3))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k3_relat_1(X6,X5),u1_struct_0(X4),X2)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(k1_pre_topc(X1,X2)))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X6)|~l1_pre_topc(X4)|~l1_pre_topc(X1)|~l1_pre_topc(X3))).#trainneg
cnf(i_0_7804, plain, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),X1),k1_pre_topc(k15_euclid(np__2),X2))|~v1_topreal2(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~v1_topreal2(X2)|~m1_subset_1(esk4_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))))|~m1_subset_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_funct_2(esk4_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1))),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),X1)))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_7792, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(k1_pre_topc(X1,X2),X3)|v2_struct_0(X4)|~m1_subset_1(k1_partfun1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X4),u1_struct_0(X4),u1_struct_0(X3),X5,X6),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k1_partfun1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X4),u1_struct_0(X4),u1_struct_0(X3),X5,X6),X2,u1_struct_0(X3))|~v1_funct_1(k1_partfun1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X4),u1_struct_0(X4),u1_struct_0(X3),X5,X6))|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v3_tops_2(X6,X4,X3)|~v3_tops_2(X5,k1_pre_topc(X1,X2),X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X3))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X4))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X3))|~v1_funct_2(X5,u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~l1_pre_topc(X4)|~l1_pre_topc(k1_pre_topc(X1,X2)))).#trainneg
cnf(i_0_5604, plain, (v4_relat_1(esk4_1(u1_struct_0(esk5_0)),k1_topreal1))).#trainneg
cnf(i_0_6627, negated_conjecture, (k2_struct_0(esk2_0)=u1_struct_0(esk5_0)|esk2_0!=esk5_0|~v1_pre_topc(esk2_0))).#trainneg
cnf(i_0_955, plain, (v1_xreal_0(esk12_1(X1))|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6237, plain, (v2_membered(esk29_0)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4518, plain, (v3_membered(esk13_0)|~m1_subset_1(u1_struct_0(k15_euclid(np__2)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_956, plain, (v1_relat_1(esk12_1(X1))|~v4_funct_1(k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_957, plain, (v1_xxreal_0(esk12_1(X1))|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_7793, plain, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(k1_pre_topc(X3,X4))|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(X3,X4))|v2_struct_0(X5)|~m1_subset_1(k1_partfun1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X5),u1_struct_0(X5),X4,X6,X7),k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(X3,X4)))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k1_partfun1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X5),u1_struct_0(X5),X4,X6,X7),X2,u1_struct_0(k1_pre_topc(X3,X4)))|~v1_funct_1(k1_partfun1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X5),u1_struct_0(X5),X4,X6,X7))|~l1_pre_topc(k1_pre_topc(X3,X4))|~l1_pre_topc(X1)|~v3_tops_2(X7,X5,k1_pre_topc(X3,X4))|~v3_tops_2(X6,k1_pre_topc(X1,X2),X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X5))))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),X4)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X3)))|~v1_funct_2(X6,u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(X5))|~v1_funct_2(X7,u1_struct_0(X5),X4)|~v1_funct_1(X7)|~v1_funct_1(X6)|~l1_pre_topc(X5)|~l1_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(X3))).#trainneg
cnf(i_0_8428, negated_conjecture, (~v2_struct_0(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_8474, negated_conjecture, (v1_funct_2(esk4_1(u1_struct_0(esk2_0)),k1_topreal1,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_8461, negated_conjecture, (m1_subset_1(esk4_1(u1_struct_0(esk2_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_8460, negated_conjecture, (v1_funct_2(esk4_1(u1_struct_0(esk2_0)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_8484, plain, (v3_tops_1(esk6_0,esk17_0))).#trainneg
cnf(i_0_7887, plain, (v3_valued_0(esk34_1(k2_zfmisc_1(X1,k1_numbers))))).#trainneg
cnf(i_0_6628, negated_conjecture, (k2_struct_0(esk1_0)=u1_struct_0(esk5_0)|esk1_0!=esk5_0|~v1_pre_topc(esk1_0))).#trainneg
cnf(i_0_8483, negated_conjecture, (u1_struct_0(esk2_0)=u1_struct_0(esk1_0)|k2_struct_0(esk1_0)!=u1_struct_0(esk2_0)|~v1_pre_topc(esk1_0))).#trainneg
cnf(i_0_2244, plain, (l1_pre_topc(k15_euclid(esk9_1(X1)))|~v6_membered(X1))).#trainneg
cnf(i_0_960, plain, (v3_membered(esk12_1(X1))|~v3_membered(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_961, plain, (v1_membered(esk12_1(X1))|~v1_membered(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_8593, negated_conjecture, (u1_struct_0(X1)=u1_struct_0(esk2_0)|X1!=k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_8415, negated_conjecture, (v2_struct_0(X1)|v3_tops_2(k2_tops_2(u1_struct_0(esk2_0),u1_struct_0(X1),X2),X1,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~v3_tops_2(X2,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(esk2_0),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_8849, negated_conjecture, (g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))=g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))).#trainneg
cnf(i_0_9115, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))))=g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))).#trainneg
cnf(i_0_9072, negated_conjecture, (~v2_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))).#trainneg
cnf(i_0_9339, negated_conjecture, (~v2_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))))).#trainneg
cnf(i_0_4610, plain, (v3_membered(k1_topreal1)|~m1_subset_1(u1_struct_0(k15_euclid(np__2)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8775, negated_conjecture, (v1_relat_1(esk4_1(u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_962, plain, (v2_membered(esk12_1(X1))|~v2_membered(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_991, plain, (v1_funct_1(esk49_1(X1))|v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_963, plain, (v1_funct_1(esk40_1(X1))|~v4_funct_1(k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_964, plain, (v7_ordinal1(esk40_1(X1))|~v6_membered(k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_8784, negated_conjecture, (m1_subset_1(esk4_1(u1_struct_0(esk2_0)),k1_zfmisc_1(k2_zfmisc_1(k1_topreal1,u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_8946, negated_conjecture, (v2_struct_0(X1)|v3_tops_2(k2_tops_2(u1_struct_0(esk2_0),u1_struct_0(X1),X2),X1,g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~v3_tops_2(X2,g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(esk2_0),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_9213, negated_conjecture, (v2_struct_0(X1)|v3_tops_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(X1),X2),X1,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|~v3_tops_2(X2,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(esk1_0),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_8462, negated_conjecture, (v3_tops_2(esk4_1(u1_struct_0(esk2_0)),k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_9070, negated_conjecture, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|r1_t_0topsp(esk2_0,k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v2_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))).#trainneg
cnf(i_0_9337, negated_conjecture, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1))|r1_t_0topsp(esk1_0,k1_pre_topc(k15_euclid(np__2),k1_topreal1))|~v2_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))).#trainneg
cnf(i_0_5673, plain, (v1_membered(esk13_0)|~m1_subset_1(u1_struct_0(k15_euclid(np__2)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1593, plain, (esk36_2(X1,X2)=esk36_2(X3,X4))).#trainneg
cnf(i_0_6901, negated_conjecture, (r1_t_0topsp(esk2_0,esk5_0)|~r1_t_0topsp(esk5_0,esk2_0))).#trainneg
cnf(i_0_965, plain, (v1_xcmplx_0(esk40_1(X1))|~v1_membered(k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1107, plain, (v1_valued_0(esk49_1(X1))|v1_xboole_0(X1)|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1132, plain, (v3_valued_0(esk49_1(X1))|v1_xboole_0(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_966, plain, (v1_xreal_0(esk40_1(X1))|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_967, plain, (v1_relat_1(esk40_1(X1))|~v4_funct_1(k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_8456, negated_conjecture, (v1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))))))).#trainneg
cnf(i_0_8434, negated_conjecture, (g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))))=k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))|~m1_subset_1(u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_8229, negated_conjecture, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~v2_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_8443, negated_conjecture, (v2_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))))|~v2_pre_topc(k15_euclid(np__2)))).#trainneg
cnf(i_0_9020, negated_conjecture, (v2_struct_0(k1_pre_topc(X1,X2))|v3_tops_2(k2_tops_2(u1_struct_0(esk2_0),X2,X3),k1_pre_topc(X1,X2),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~v3_tops_2(X3,g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)),k1_pre_topc(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(esk2_0),X2)|~v1_funct_1(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_8922, negated_conjecture, (v2_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))))|r1_t_0topsp(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~v2_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~l1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0))))))).#trainneg
cnf(i_0_9189, negated_conjecture, (v2_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))))|r1_t_0topsp(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))),g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|~v2_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|~l1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))))))).#trainneg
cnf(i_0_6123, plain, (v1_membered(k1_topreal1)|~m1_subset_1(u1_struct_0(k15_euclid(np__2)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2046, plain, (k2_zfmisc_1(X1,esk6_0)=esk36_2(X2,X3))).#trainneg
cnf(i_0_6902, negated_conjecture, (r1_t_0topsp(esk1_0,esk5_0)|~r1_t_0topsp(esk5_0,esk1_0))).#trainneg
cnf(i_0_968, plain, (v1_xxreal_0(esk40_1(X1))|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_8703, negated_conjecture, (~v2_tops_1(u1_struct_0(esk5_0),esk2_0)|esk2_0!=esk5_0|~v1_pre_topc(esk2_0))).#trainneg
cnf(i_0_2185, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,esk6_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2247, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,esk6_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_970, plain, (v1_xboole_0(esk40_1(X1))|~v1_xboole_0(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_971, plain, (v3_membered(esk40_1(X1))|~v3_membered(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_8223, negated_conjecture, (m1_subset_1(u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0))))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_8222, negated_conjecture, (v2_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))))|~v2_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_8479, negated_conjecture, (v1_funct_2(k3_relat_1(X1,esk4_1(u1_struct_0(esk2_0))),X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_topreal1)))|~v1_funct_2(X1,X2,k1_topreal1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_9287, negated_conjecture, (v2_struct_0(k1_pre_topc(X1,X2))|v3_tops_2(k2_tops_2(u1_struct_0(esk1_0),X2,X3),k1_pre_topc(X1,X2),g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|~v3_tops_2(X3,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),k1_pre_topc(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(esk1_0),X2)|~v1_funct_1(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_9130, negated_conjecture, (m1_subset_1(u1_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_9054, negated_conjecture, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|r1_t_0topsp(X3,k1_pre_topc(X1,X2))|~v3_tops_2(X4,g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)),k1_pre_topc(X1,X2))|~v3_tops_2(X5,X3,g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(esk2_0)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk2_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k1_partfun1(u1_struct_0(X3),u1_struct_0(esk2_0),u1_struct_0(esk2_0),X2,X5,X4),u1_struct_0(X3),X2)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(esk2_0))|~v1_funct_2(X4,u1_struct_0(esk2_0),X2)|~v1_funct_1(X4)|~v1_funct_1(X5)|~l1_pre_topc(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_8279, negated_conjecture, (v2_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))))|r1_t_0topsp(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))),k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~v2_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~l1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_9321, negated_conjecture, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|v2_struct_0(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|r1_t_0topsp(X3,k1_pre_topc(X1,X2))|~v3_tops_2(X4,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),k1_pre_topc(X1,X2))|~v3_tops_2(X5,X3,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk1_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k1_partfun1(u1_struct_0(X3),u1_struct_0(esk1_0),u1_struct_0(esk1_0),X2,X5,X4),u1_struct_0(X3),X2)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(esk1_0))|~v1_funct_2(X4,u1_struct_0(esk1_0),X2)|~v1_funct_1(X4)|~v1_funct_1(X5)|~l1_pre_topc(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_8500, plain, (u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))=u1_struct_0(X1)|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6238, plain, (v2_membered(esk13_0)|~m1_subset_1(u1_struct_0(k15_euclid(np__2)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2111, plain, (k2_zfmisc_1(esk6_0,X1)=esk36_2(X2,X3))).#trainneg
cnf(i_0_8229, negated_conjecture, (~v2_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))))))).#trainneg
cnf(i_0_972, plain, (v1_membered(esk40_1(X1))|~v1_membered(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_6313, plain, (v2_membered(esk8_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k5_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_9505, plain, (v1_valued_0(esk36_2(X3,X4))|~v1_membered(X2))).#trainneg
cnf(i_0_9506, plain, (v3_valued_0(esk36_2(X3,X4))|~v3_membered(X2))).#trainneg
cnf(i_0_973, plain, (v2_membered(esk40_1(X1))|~v2_membered(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_975, plain, (v1_valued_0(esk49_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_9533, negated_conjecture, (u1_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))))=u1_struct_0(esk2_0))).#trainneg
cnf(i_0_9842, negated_conjecture, (m1_subset_1(u1_pre_topc(esk1_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_9892, plain, (u1_struct_0(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))=X2|~v2_pre_topc(k1_pre_topc(X1,X2))|~l1_pre_topc(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~l1_pre_topc(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_8434, negated_conjecture, (g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))))=k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_9882, negated_conjecture, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(X3,k1_pre_topc(X1,X2))|~v3_tops_2(X4,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),k1_pre_topc(X1,X2))|~v3_tops_2(X5,X3,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk1_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k3_relat_1(X5,X4),u1_struct_0(X3),X2)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(esk1_0))|~v1_funct_2(X4,u1_struct_0(esk1_0),X2)|~v1_funct_1(X4)|~v1_funct_1(X5)|~l1_pre_topc(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_9629, negated_conjecture, (r1_t_0topsp(esk1_0,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|~v2_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))))).#trainneg
cnf(i_0_6315, plain, (v2_membered(k1_topreal1)|~m1_subset_1(u1_struct_0(k15_euclid(np__2)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_9492, plain, (v3_valued_0(esk36_2(X2,X3)))).#trainneg
cnf(i_0_7596, negated_conjecture, (r1_t_0topsp(esk2_0,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~m1_subset_1(u1_pre_topc(esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_978, plain, (v1_valued_0(esk34_1(k2_zfmisc_1(X1,X2)))|~v1_membered(X2))).#trainneg
cnf(i_0_9486, negated_conjecture, (k2_struct_0(esk1_0)!=k1_topreal1|~v1_pre_topc(esk1_0))).#trainneg
cnf(i_0_9491, negated_conjecture, (k2_struct_0(esk2_0)!=k1_topreal1|~v1_pre_topc(esk2_0))).#trainneg
cnf(i_0_984, plain, (v1_valued_0(esk9_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_membered(X2))).#trainneg
cnf(i_0_9626, negated_conjecture, (r1_t_0topsp(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),esk1_0)|~v2_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))))).#trainneg
cnf(i_0_9457, negated_conjecture, (v3_tops_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(esk1_0),X1),g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|~v3_tops_2(X1,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk1_0))))|~v1_funct_2(X1,u1_struct_0(esk1_0),u1_struct_0(esk1_0))|~v1_funct_1(X1))).#trainneg
cnf(i_0_8252, negated_conjecture, (v2_struct_0(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))))|r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)),g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))))|~v2_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~l1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_8513, plain, (m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(g1_pre_topc(X1,X2)))))|~l1_pre_topc(g1_pre_topc(X1,X2))|~v1_pre_topc(g1_pre_topc(X1,X2)))).#trainneg
cnf(i_0_8480, negated_conjecture, (v1_xboole_0(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|v1_funct_2(k3_relat_1(X1,esk4_1(u1_struct_0(esk2_0))),X2,u1_struct_0(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))))|~v1_funct_2(X1,X2,u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_8499, plain, (u1_struct_0(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))=u1_struct_0(X1)|v2_struct_0(X1)|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_10314, negated_conjecture, (v1_toprealb(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_9679, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk36_2(X3,X4))))).#trainneg
cnf(i_0_9636, plain, (v1_xboole_0(k2_zfmisc_1(X3,esk6_0)))).#trainneg
cnf(i_0_9633, negated_conjecture, (r1_t_0topsp(esk1_0,esk1_0)|~v1_pre_topc(esk1_0))).#trainneg
cnf(i_0_9066, negated_conjecture, (k1_pre_topc(X1,u1_struct_0(esk2_0))=k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))|~m1_subset_1(u1_struct_0(esk2_0),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(esk2_0,X1))).#trainneg
cnf(i_0_989, plain, (v1_funct_1(esk56_1(X1))|~v1_relat_1(u1_struct_0(X1))|~v1_funct_1(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1989, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk6_0))).#trainneg
cnf(i_0_993, plain, (v1_funct_1(esk43_1(X1))|~v1_relat_1(u1_struct_0(X1))|~v1_funct_1(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_994, plain, (v1_funct_1(esk40_1(X1))|~v1_relat_1(u1_struct_0(X1))|~v1_funct_1(u1_struct_0(X1))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_10322, negated_conjecture, (r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)),esk2_0))).#trainneg
cnf(i_0_10321, negated_conjecture, (r1_t_0topsp(esk2_0,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_10345, negated_conjecture, (v2_struct_0(X1)|r1_t_0topsp(esk2_0,X1)|~v2_pre_topc(X1)|~r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)),g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_10329, negated_conjecture, (r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)),k1_pre_topc(k15_euclid(np__2),k1_topreal1)))).#trainneg
cnf(i_0_9682, plain, (v1_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk36_2(X3,X4))))).#trainneg
cnf(i_0_9651, plain, (v1_funct_1(k2_zfmisc_1(X3,esk6_0)))).#trainneg
cnf(i_0_9054, negated_conjecture, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(X3,k1_pre_topc(X1,X2))|~v3_tops_2(X4,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)),k1_pre_topc(X1,X2))|~v3_tops_2(X5,X3,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk2_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k1_partfun1(u1_struct_0(X3),u1_struct_0(esk2_0),u1_struct_0(esk2_0),X2,X5,X4),u1_struct_0(X3),X2)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(esk2_0))|~v1_funct_2(X4,u1_struct_0(esk2_0),X2)|~v1_funct_1(X4)|~v1_funct_1(X5)|~l1_pre_topc(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_2414, plain, (v1_finseq_1(X1)|~v2_monoid_0(esk33_0)|~l1_struct_0(esk33_0)|~m1_subset_1(X1,esk6_0))).#trainneg
cnf(i_0_10429, negated_conjecture, (u1_struct_0(esk1_0)!=k1_topreal1|~v1_pre_topc(esk1_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_10432, negated_conjecture, (u1_struct_0(esk2_0)!=k1_topreal1|~v1_pre_topc(esk2_0)|~l1_struct_0(esk2_0))).#trainneg
cnf(i_0_9984, plain, (v1_valued_0(esk36_2(X1,X2))|~v3_membered(X3))).#trainneg
cnf(i_0_1003, plain, (v1_funct_1(esk13_0)|~v1_relat_1(u1_struct_0(k15_euclid(np__2)))|~v1_funct_1(u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_10324, negated_conjecture, (r1_t_0topsp(X1,esk2_0)|~v2_pre_topc(X1)|~r1_t_0topsp(X1,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1158, negated_conjecture, (v1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_10409, negated_conjecture, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(X3,k1_pre_topc(X1,X2))|~v3_tops_2(X5,X3,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk1_0))))|~m1_subset_1(k2_tops_2(X2,u1_struct_0(esk1_0),X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k3_relat_1(X5,k2_tops_2(X2,u1_struct_0(esk1_0),X4)),u1_struct_0(X3),X2)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(esk1_0))|~v1_funct_2(k2_tops_2(X2,u1_struct_0(esk1_0),X4),u1_struct_0(esk1_0),X2)|~v1_funct_1(k2_tops_2(X2,u1_struct_0(esk1_0),X4))|~v1_funct_1(X5)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v3_tops_2(X4,k1_pre_topc(X1,X2),g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(esk1_0))))|~v1_funct_2(X4,X2,u1_struct_0(esk1_0))|~v1_funct_1(X4))).#trainneg
cnf(i_0_10353, negated_conjecture, (v2_struct_0(X1)|r1_t_0topsp(esk2_0,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X1)|~r1_t_0topsp(esk2_0,X1)|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_10312, negated_conjecture, (v2_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_10364, negated_conjecture, (v2_struct_0(X1)|v2_struct_0(X2)|v3_tops_2(k1_partfun1(u1_struct_0(X2),u1_struct_0(X1),u1_struct_0(X1),u1_struct_0(esk2_0),X3,X4),X2,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~v3_tops_2(X4,X1,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~v3_tops_2(X3,X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_2(X4,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l1_pre_topc(X1)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_10495, plain, (u1_struct_0(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))=X2|v2_struct_0(k1_pre_topc(X1,X2))|~l1_pre_topc(g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~l1_pre_topc(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_10351, negated_conjecture, (v2_struct_0(X1)|r1_t_0topsp(X1,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~v2_pre_topc(X1)|~r1_t_0topsp(esk2_0,X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_9683, plain, (v3_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk36_2(X3,X4))))).#trainneg
cnf(i_0_1006, plain, (v1_funct_1(esk9_1(k1_zfmisc_1(X1)))|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2559, plain, (v1_finseq_1(X1)|~v2_monoid_0(esk17_0)|~l1_struct_0(esk17_0)|~m1_subset_1(X1,esk6_0))).#trainneg
cnf(i_0_8810, negated_conjecture, (~v2_tops_1(u1_struct_0(esk5_0),esk1_0)|esk1_0!=esk5_0|~v1_pre_topc(esk1_0))).#trainneg
cnf(i_0_10427, plain, (v1_valued_0(esk34_1(esk36_2(X2,X3)))|~v1_membered(X1))).#trainneg
cnf(i_0_9674, plain, (v1_valued_0(esk36_2(X2,X3)))).#trainneg
cnf(i_0_1010, plain, (v1_funct_1(u1_pre_topc(X1))|~v1_relat_1(k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1011, plain, (v1_funct_1(k2_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~v1_funct_1(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_10304, negated_conjecture, (u1_pre_topc(X1)=u1_pre_topc(esk2_0)|X1!=k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_10347, negated_conjecture, (v2_struct_0(X1)|r1_t_0topsp(esk2_0,X1)|~v2_pre_topc(X1)|~r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)),X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_10408, negated_conjecture, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(X3,k1_pre_topc(X1,X2))|~v3_tops_2(X5,X3,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk1_0))))|~m1_subset_1(k2_tops_2(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk1_0),X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k3_relat_1(X5,k2_tops_2(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk1_0),X4)),u1_struct_0(X3),X2)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(esk1_0))|~v1_funct_2(k2_tops_2(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk1_0),X4),u1_struct_0(esk1_0),X2)|~v1_funct_1(k2_tops_2(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk1_0),X4))|~v1_funct_1(X5)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v3_tops_2(X4,k1_pre_topc(X1,X2),g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk1_0))))|~v1_funct_2(X4,u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk1_0))|~v1_funct_1(X4)|~l1_pre_topc(k1_pre_topc(X1,X2)))).#trainneg
cnf(i_0_10357, negated_conjecture, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(esk2_0,k1_pre_topc(X1,X2))|~v2_pre_topc(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)),g1_pre_topc(X2,u1_pre_topc(k1_pre_topc(X1,X2))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_10356, negated_conjecture, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(esk2_0,k1_pre_topc(X1,X2))|~v2_pre_topc(k1_pre_topc(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)),k1_pre_topc(X1,X2))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_9685, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk36_2(X3,X4))))).#trainneg
cnf(i_0_9900, plain, (v1_xboole_0(k2_zfmisc_1(esk6_0,X3)))).#trainneg
cnf(i_0_1012, plain, (v1_funct_1(k1_topreal1)|~v1_relat_1(u1_struct_0(k15_euclid(np__2)))|~v1_funct_1(u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_10431, negated_conjecture, (u1_struct_0(esk5_0)!=k1_topreal1|~v1_pre_topc(esk1_0)|esk1_0!=esk5_0)).#trainneg
cnf(i_0_9495, plain, (v4_relat_1(esk36_2(X3,X4),X1))).#trainneg
cnf(i_0_9496, plain, (v5_relat_1(esk36_2(X3,X4),X2))).#trainneg
cnf(i_0_1013, plain, (v1_funct_1(esk9_1(u1_struct_0(X1)))|~v1_monoid_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_1014, plain, (v1_funct_1(u2_struct_0(X1))|~v1_monoid_0(X1)|~l1_struct_0(X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_10323, negated_conjecture, (r1_t_0topsp(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),esk2_0)|~v2_pre_topc(X1)|~r1_t_0topsp(X1,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_10306, negated_conjecture, (r1_t_0topsp(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)),k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~v2_pre_topc(X1)|~r1_t_0topsp(X1,esk2_0)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_8253, negated_conjecture, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|r1_t_0topsp(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))),k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~v2_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~l1_pre_topc(g1_pre_topc(u1_struct_0(esk2_0),u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_9632, negated_conjecture, (r1_t_0topsp(esk1_0,esk1_0)|~v2_pre_topc(g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0))))).#trainneg
cnf(i_0_9809, negated_conjecture, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(k1_pre_topc(X1,X2),esk2_0)|~v3_tops_2(X4,k1_pre_topc(X1,X2),X3)|~v3_tops_2(esk4_1(u1_struct_0(esk2_0)),X3,esk2_0)|~m1_subset_1(esk4_1(u1_struct_0(esk2_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk2_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(esk4_1(u1_struct_0(esk2_0)),u1_struct_0(X3),u1_struct_0(esk2_0))|~v1_funct_2(X4,X2,u1_struct_0(X3))|~v1_funct_1(esk4_1(u1_struct_0(esk2_0)))|~v1_funct_1(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,k1_topreal1)))|~v1_funct_2(X4,X2,k1_topreal1))).#trainneg
cnf(i_0_9915, plain, (v1_funct_1(k2_zfmisc_1(esk6_0,X3)))).#trainneg
cnf(i_0_9485, negated_conjecture, (r1_t_0topsp(esk2_0,esk5_0)|u1_struct_0(esk5_0)!=k1_topreal1)).#trainneg
cnf(i_0_1023, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,esk56_1(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_8697, negated_conjecture, (v1_tops_1(u1_struct_0(esk5_0),esk2_0)|esk2_0!=esk5_0|~v1_pre_topc(esk2_0))).#trainneg
cnf(i_0_10434, negated_conjecture, (u1_struct_0(esk5_0)!=k1_topreal1|~v1_pre_topc(esk2_0)|esk2_0!=esk5_0)).#trainneg
cnf(i_0_1025, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk49_1(X2)))).#trainneg
cnf(i_0_1026, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,esk47_0))).#trainneg
cnf(i_0_9810, negated_conjecture, (r1_t_0topsp(k1_pre_topc(X1,X2),g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)))|v2_struct_0(k1_pre_topc(X1,X2))|~m1_subset_1(k2_tops_2(u1_struct_0(esk1_0),X2,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk1_0))))|~v1_funct_2(k2_tops_2(u1_struct_0(esk1_0),X2,X3),u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk1_0))|~v1_funct_1(k2_tops_2(u1_struct_0(esk1_0),X2,X3))|~l1_pre_topc(k1_pre_topc(X1,X2))|~v3_tops_2(X3,g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(esk1_0)),k1_pre_topc(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(X3,u1_struct_0(esk1_0),X2)|~v1_funct_1(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_8503, negated_conjecture, (u1_struct_0(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)))=u1_struct_0(esk1_0))).#trainneg
cnf(i_0_10651, negated_conjecture, (r1_t_0topsp(esk2_0,esk2_0))).#trainneg
cnf(i_0_9115, negated_conjecture, (g1_pre_topc(u1_struct_0(esk1_0),u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0))))=k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1911, negated_conjecture, (l1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_10554, plain, (v1_xboole_0(esk20_1(esk6_0)))).#trainneg
cnf(i_0_10555, plain, (v1_xboole_0(esk19_1(esk6_0)))).#trainneg
cnf(i_0_9333, negated_conjecture, (k1_pre_topc(X1,u1_struct_0(esk1_0))=k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~m1_pre_topc(esk1_0,X1))).#trainneg
cnf(i_0_1027, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,esk43_1(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_9339, negated_conjecture, (~v2_struct_0(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_10516, plain, (X1=k2_zfmisc_1(X2,esk6_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_10916, plain, (X1=k2_zfmisc_1(esk6_0,X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1028, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,esk40_1(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1029, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,esk36_2(X2,X3)))).#trainneg
cnf(i_0_11074, negated_conjecture, (v2_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1159, negated_conjecture, (v1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9629, negated_conjecture, (r1_t_0topsp(esk1_0,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)))|~v2_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9632, negated_conjecture, (r1_t_0topsp(esk1_0,esk1_0)|~v2_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_10420, plain, (v3_valued_0(k2_zfmisc_1(X3,esk6_0)))).#trainneg
cnf(i_0_9489, negated_conjecture, (r1_t_0topsp(esk1_0,esk5_0)|u1_struct_0(esk5_0)!=k1_topreal1)).#trainneg
cnf(i_0_9810, negated_conjecture, (v2_struct_0(k1_pre_topc(X1,X2))|r1_t_0topsp(k1_pre_topc(X1,X2),k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)))|~v3_tops_2(X3,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)),k1_pre_topc(X1,X2))|~m1_subset_1(k2_tops_2(u1_struct_0(esk1_0),X2,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk1_0))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k2_tops_2(u1_struct_0(esk1_0),X2,X3),u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk1_0))|~v1_funct_2(X3,u1_struct_0(esk1_0),X2)|~v1_funct_1(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1030, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk34_1(X2)))).#trainneg
cnf(i_0_10566, negated_conjecture, (k2_struct_0(esk1_0)!=u1_struct_0(esk2_0)|~v1_pre_topc(esk1_0))).#trainneg
cnf(i_0_1032, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,esk29_0))).#trainneg
cnf(i_0_1034, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,esk25_2(X2,X3)))).#trainneg
cnf(i_0_10994, negated_conjecture, (r1_t_0topsp(X1,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~r1_t_0topsp(X1,esk2_0))).#trainneg
cnf(i_0_10809, negated_conjecture, (u1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))=u1_pre_topc(esk2_0))).#trainneg
cnf(i_0_11430, negated_conjecture, (v1_funct_2(esk4_1(u1_struct_0(esk1_0)),k1_topreal1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11418, negated_conjecture, (m1_subset_1(esk4_1(u1_struct_0(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_11417, negated_conjecture, (v1_funct_2(esk4_1(u1_struct_0(esk1_0)),u1_struct_0(k1_pre_topc(k15_euclid(np__2),k1_topreal1)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11630, negated_conjecture, (r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)),esk1_0))).#trainneg
cnf(i_0_11483, plain, (X1=esk20_1(esk6_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_10421, plain, (v3_valued_0(k2_zfmisc_1(esk6_0,X3)))).#trainneg
cnf(i_0_11503, plain, (esk20_1(esk6_0)=esk6_0)).#trainneg
cnf(i_0_9882, negated_conjecture, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(X3,k1_pre_topc(X1,X2))|~v3_tops_2(X4,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)),k1_pre_topc(X1,X2))|~v3_tops_2(X5,X3,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk1_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k3_relat_1(X5,X4),u1_struct_0(X3),X2)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(esk1_0))|~v1_funct_2(X4,u1_struct_0(esk1_0),X2)|~v1_funct_1(X4)|~v1_funct_1(X5)|~l1_pre_topc(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1035, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X2))|~r2_hidden(X1,esk20_1(X2)))).#trainneg
cnf(i_0_11668, negated_conjecture, (~v1_xboole_0(u1_pre_topc(esk2_0)))).#trainneg
cnf(i_0_11521, plain, (X1=esk19_1(esk6_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_10795, plain, (v1_valued_0(esk34_1(esk36_2(X1,X2)))|~v3_membered(X3))).#trainneg
cnf(i_0_1036, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X2))|~r2_hidden(X1,esk19_1(X2)))).#trainneg
cnf(i_0_1037, plain, (m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,esk13_0))).#trainneg
cnf(i_0_11419, negated_conjecture, (v3_tops_2(esk4_1(u1_struct_0(esk1_0)),k1_pre_topc(k15_euclid(np__2),k1_topreal1),k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_11367, negated_conjecture, (v2_struct_0(X1)|v3_tops_2(k2_tops_2(u1_struct_0(esk1_0),u1_struct_0(X1),X2),X1,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)))|~v3_tops_2(X2,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(esk1_0),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_11461, negated_conjecture, (u1_struct_0(X1)=u1_struct_0(esk1_0)|X1!=k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0))|~v1_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_10760, negated_conjecture, (r1_t_0topsp(X1,esk2_0)|v2_struct_0(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~r1_t_0topsp(esk2_0,X1))).#trainneg
cnf(i_0_11541, plain, (esk19_1(esk6_0)=esk6_0)).#trainneg
cnf(i_0_10409, negated_conjecture, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(X3,k1_pre_topc(X1,X2))|~v3_tops_2(X4,k1_pre_topc(X1,X2),k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)))|~v3_tops_2(X5,X3,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk1_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k3_relat_1(X5,k2_tops_2(X2,u1_struct_0(esk1_0),X4)),u1_struct_0(X3),X2)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(esk1_0))|~v1_funct_2(X4,X2,u1_struct_0(esk1_0))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1038, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,esk12_1(X2))|~v2_pre_topc(X2)|~l1_pre_topc(X2))).#trainneg
cnf(i_0_11633, negated_conjecture, (k2_struct_0(esk2_0)!=u1_struct_0(esk1_0)|~v1_pre_topc(esk2_0))).#trainneg
cnf(i_0_1040, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk9_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_1041, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k5_numbers))|~r2_hidden(X1,esk8_0))).#trainneg
cnf(i_0_11431, negated_conjecture, (m1_subset_1(esk4_1(u1_struct_0(esk1_0)),k1_zfmisc_1(k2_zfmisc_1(k1_topreal1,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_11091, negated_conjecture, (~r1_t_0topsp(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)),k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_11106, negated_conjecture, (r1_t_0topsp(X1,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)))|~v2_pre_topc(X1)|~v1_pre_topc(X1)|~r1_t_0topsp(X1,esk1_0)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_10972, negated_conjecture, (r1_t_0topsp(esk2_0,g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~l1_pre_topc(g1_pre_topc(u1_struct_0(X1),u1_pre_topc(X1)))|~v2_pre_topc(X1)|~r1_t_0topsp(X1,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_11754, negated_conjecture, (u1_pre_topc(esk2_0)=u1_pre_topc(esk5_0)|u1_struct_0(esk2_0)!=u1_struct_0(esk5_0))).#trainneg
cnf(i_0_9321, negated_conjecture, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(X3,k1_pre_topc(X1,X2))|~v3_tops_2(X4,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)),k1_pre_topc(X1,X2))|~v3_tops_2(X5,X3,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk1_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X2)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k1_partfun1(u1_struct_0(X3),u1_struct_0(esk1_0),u1_struct_0(esk1_0),X2,X5,X4),u1_struct_0(X3),X2)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(esk1_0))|~v1_funct_2(X4,u1_struct_0(esk1_0),X2)|~v1_funct_1(X4)|~v1_funct_1(X5)|~l1_pre_topc(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1044, plain, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,u1_pre_topc(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_11648, negated_conjecture, (u1_struct_0(esk2_0)!=u1_struct_0(esk1_0)|~v1_pre_topc(esk1_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_1045, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,k2_struct_0(X2))|~l1_struct_0(X2))).#trainneg
cnf(i_0_1046, plain, (m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k1_topreal1))).#trainneg
cnf(i_0_10830, negated_conjecture, (v2_struct_0(k1_pre_topc(k15_euclid(np__2),X1))|r1_t_0topsp(esk2_0,k1_pre_topc(k15_euclid(np__2),X1))|~v2_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),X1))|~v1_topreal2(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),k1_topreal1)))).#trainneg
cnf(i_0_11435, negated_conjecture, (v1_funct_2(k3_relat_1(X1,esk4_1(u1_struct_0(esk1_0))),X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,k1_topreal1)))|~v1_funct_2(X1,X2,k1_topreal1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_10728, negated_conjecture, (v2_struct_0(X1)|v3_tops_2(k1_partfun1(u1_struct_0(X1),u1_struct_0(esk2_0),u1_struct_0(esk2_0),u1_struct_0(esk2_0),X2,X3),X1,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~v3_tops_2(X3,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)),k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~v3_tops_2(X2,X1,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk2_0)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk2_0),u1_struct_0(esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(esk2_0),u1_struct_0(esk2_0))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~v1_funct_1(X2)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_11086, negated_conjecture, (r1_t_0topsp(X1,esk1_0)|~v2_pre_topc(X1)|~r1_t_0topsp(X1,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)))|~l1_pre_topc(k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_10797, plain, (v1_valued_0(k2_zfmisc_1(X3,esk6_0)))).#trainneg
cnf(i_0_10650, negated_conjecture, (r1_t_0topsp(esk5_0,esk2_0)|u1_struct_0(esk5_0)!=k1_topreal1)).#trainneg
cnf(i_0_10408, negated_conjecture, (v2_struct_0(k1_pre_topc(X1,X2))|v2_struct_0(X3)|r1_t_0topsp(X3,k1_pre_topc(X1,X2))|~v3_tops_2(X4,k1_pre_topc(X1,X2),k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)))|~v3_tops_2(X5,X3,k1_pre_topc(k15_euclid(np__2),u1_struct_0(esk1_0)))|~m1_subset_1(k2_tops_2(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk1_0),X4),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),X2)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk1_0))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_funct_2(k3_relat_1(X5,k2_tops_2(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk1_0),X4)),u1_struct_0(X3),X2)|~v1_funct_2(k2_tops_2(u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk1_0),X4),u1_struct_0(esk1_0),X2)|~v1_funct_2(X4,u1_struct_0(k1_pre_topc(X1,X2)),u1_struct_0(esk1_0))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(esk1_0))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l1_pre_topc(X3)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_8699, negated_conjecture, (v3_pre_topc(u1_struct_0(esk5_0),esk2_0)|esk2_0!=esk5_0|~v1_pre_topc(esk2_0))).#trainneg
cnf(i_0_10435, plain, (v1_valued_0(esk9_1(k1_zfmisc_1(esk36_2(X2,X3))))|~v1_membered(X1))).#trainneg
cnf(i_0_3185, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(esk6_0,X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1050, plain, (v1_xboole_0(esk34_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1056, plain, (v1_xboole_0(esk9_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 238
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 502
# Removed in clause preprocessing      : 39
# Initial clauses in saturation        : 463
# Processed clauses                    : 2305
# ...of these trivial                  : 115
# ...subsumed                          : 527
# ...remaining for further processing  : 1663
# Other redundant clauses eliminated   : 6
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 46
# Backward-rewritten                   : 177
# Generated clauses                    : 11917
# ...of the previous two non-trivial   : 10578
# Contextual simplify-reflections      : 325
# Paramodulations                      : 11894
# Factorizations                       : 0
# Equation resolutions                 : 23
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1440
#    Positive orientable unit clauses  : 310
#    Positive unorientable unit clauses: 3
#    Negative unit clauses             : 55
#    Non-unit-clauses                  : 1072
# Current number of unprocessed clauses: 8357
# ...number of literals in the above   : 67878
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1886
# Proof object given clauses           : 144
# Proof search given clauses           : 1663
# Clause-clause subsumption calls (NU) : 1229039
# Rec. Clause-clause subsumption calls : 86148
# Non-unit clause-clause subsumptions  : 763
# Unit Clause-clause subsumption calls : 61392
# Rewrite failures with RHS unbound    : 60
# BW rewrite match attempts            : 231
# BW rewrite match successes           : 63
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 482335

# -------------------------------------------------
# User time                : 33.921 s
# System time              : 1.064 s
# Total time               : 34.985 s
# Maximum resident set size: 416020 pages
33.92user 1.17system 0:35.60elapsed 98%CPU (0avgtext+0avgdata 416020maxresident)k
0inputs+816outputs (0major+331037minor)pagefaults 0swaps
