# ENIGMA: Connected to the TF server '127.0.0.1:8888'.
# Preprocessing time       : 0.032 s
#TF#SERVER: Sending data to the TF server (conj=7, context=165, query=256, total=428).
#TF#SERVER: Sent 131049 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1688 bytes expected).
#TF#SERVER: Received 1688 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=165, query=5, total=177).
#TF#SERVER: Sent 47682 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (684 bytes expected).
#TF#SERVER: Received 684 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=255, query=256, total=518).
#TF#SERVER: Sent 279809 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2048 bytes expected).
#TF#SERVER: Received 2048 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=327, query=256, total=590).
#TF#SERVER: Sent 283576 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2336 bytes expected).
#TF#SERVER: Received 2336 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=391, query=256, total=654).
#TF#SERVER: Sent 378029 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2592 bytes expected).
#TF#SERVER: Received 2592 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=391, query=6, total=404).
#TF#SERVER: Sent 161378 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1592 bytes expected).
#TF#SERVER: Received 1592 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=444, query=256, total=707).
#TF#SERVER: Sent 362544 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2804 bytes expected).
#TF#SERVER: Received 2804 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=444, query=19, total=470).
#TF#SERVER: Sent 202954 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1856 bytes expected).
#TF#SERVER: Received 1856 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=488, query=256, total=751).
#TF#SERVER: Sent 386803 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2980 bytes expected).
#TF#SERVER: Received 2980 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=488, query=9, total=504).
#TF#SERVER: Sent 196148 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1992 bytes expected).
#TF#SERVER: Received 1992 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 512039 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=2, total=521).
#TF#SERVER: Sent 210952 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2060 bytes expected).
#TF#SERVER: Received 2060 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 485292 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=18, total=537).
#TF#SERVER: Sent 221810 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2124 bytes expected).
#TF#SERVER: Received 2124 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 487205 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=3, total=522).
#TF#SERVER: Sent 215612 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2064 bytes expected).
#TF#SERVER: Received 2064 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 628559 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=4, total=523).
#TF#SERVER: Sent 216451 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2068 bytes expected).
#TF#SERVER: Received 2068 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 574185 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=7, total=526).
#TF#SERVER: Sent 223698 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2080 bytes expected).
#TF#SERVER: Received 2080 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 679216 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=4, total=523).
#TF#SERVER: Sent 216592 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2068 bytes expected).
#TF#SERVER: Received 2068 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 627875 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=15, total=534).
#TF#SERVER: Sent 248725 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2112 bytes expected).
#TF#SERVER: Received 2112 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 654701 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=72, total=591).
#TF#SERVER: Sent 369916 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2340 bytes expected).
#TF#SERVER: Received 2340 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 576894 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=8, total=527).
#TF#SERVER: Sent 228401 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2084 bytes expected).
#TF#SERVER: Received 2084 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 507821 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=10, total=529).
#TF#SERVER: Sent 217819 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2092 bytes expected).
#TF#SERVER: Received 2092 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 617468 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=7, total=526).
#TF#SERVER: Sent 213131 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2080 bytes expected).
#TF#SERVER: Received 2080 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 578330 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=12, total=531).
#TF#SERVER: Sent 239496 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2100 bytes expected).
#TF#SERVER: Received 2100 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 651649 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=20, total=539).
#TF#SERVER: Sent 245121 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2132 bytes expected).
#TF#SERVER: Received 2132 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 474398 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=24, total=543).
#TF#SERVER: Sent 220366 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2148 bytes expected).
#TF#SERVER: Received 2148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 595434 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=13, total=532).
#TF#SERVER: Sent 226060 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2104 bytes expected).
#TF#SERVER: Received 2104 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 593165 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=2, total=521).
#TF#SERVER: Sent 210724 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2060 bytes expected).
#TF#SERVER: Received 2060 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 567454 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=177, total=696).
#TF#SERVER: Sent 551357 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2760 bytes expected).
#TF#SERVER: Received 2760 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 491439 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=4, total=523).
#TF#SERVER: Sent 212940 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2068 bytes expected).
#TF#SERVER: Received 2068 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 447417 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=5, total=524).
#TF#SERVER: Sent 212891 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2072 bytes expected).
#TF#SERVER: Received 2072 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 390071 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=5, total=524).
#TF#SERVER: Sent 212179 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2072 bytes expected).
#TF#SERVER: Received 2072 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 372736 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=14, total=533).
#TF#SERVER: Sent 220715 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2108 bytes expected).
#TF#SERVER: Received 2108 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 465958 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=3, total=522).
#TF#SERVER: Sent 213206 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2064 bytes expected).
#TF#SERVER: Received 2064 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 456910 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=4, total=523).
#TF#SERVER: Sent 213601 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2068 bytes expected).
#TF#SERVER: Received 2068 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 457573 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=14, total=533).
#TF#SERVER: Sent 220782 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2108 bytes expected).
#TF#SERVER: Received 2108 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 495151 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=64, total=583).
#TF#SERVER: Sent 254124 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2308 bytes expected).
#TF#SERVER: Received 2308 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 536820 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=11, total=530).
#TF#SERVER: Sent 221002 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2096 bytes expected).
#TF#SERVER: Received 2096 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 445480 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=80, total=599).
#TF#SERVER: Sent 280635 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2372 bytes expected).
#TF#SERVER: Received 2372 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 499527 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=6, total=525).
#TF#SERVER: Sent 216791 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2076 bytes expected).
#TF#SERVER: Received 2076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 640308 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=1, total=520).
#TF#SERVER: Sent 211497 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2056 bytes expected).
#TF#SERVER: Received 2056 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 601808 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=6, total=525).
#TF#SERVER: Sent 222102 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2076 bytes expected).
#TF#SERVER: Received 2076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 525354 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=2, total=521).
#TF#SERVER: Sent 212446 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2060 bytes expected).
#TF#SERVER: Received 2060 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 457287 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=6, total=525).
#TF#SERVER: Sent 216411 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2076 bytes expected).
#TF#SERVER: Received 2076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 632278 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=5, total=524).
#TF#SERVER: Sent 218173 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2072 bytes expected).
#TF#SERVER: Received 2072 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 605505 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=7, total=526).
#TF#SERVER: Sent 229361 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2080 bytes expected).
#TF#SERVER: Received 2080 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 650855 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=5, total=524).
#TF#SERVER: Sent 213648 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2072 bytes expected).
#TF#SERVER: Received 2072 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 607766 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=11, total=530).
#TF#SERVER: Sent 216683 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2096 bytes expected).
#TF#SERVER: Received 2096 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 549944 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=67, total=586).
#TF#SERVER: Sent 338101 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2320 bytes expected).
#TF#SERVER: Received 2320 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 550720 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=2, total=521).
#TF#SERVER: Sent 211017 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2060 bytes expected).
#TF#SERVER: Received 2060 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 636043 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=25, total=544).
#TF#SERVER: Sent 255597 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2152 bytes expected).
#TF#SERVER: Received 2152 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 502362 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=104, total=623).
#TF#SERVER: Sent 413463 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2468 bytes expected).
#TF#SERVER: Received 2468 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 713016 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=34, total=553).
#TF#SERVER: Sent 262655 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2188 bytes expected).
#TF#SERVER: Received 2188 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 715292 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=9, total=528).
#TF#SERVER: Sent 224433 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2088 bytes expected).
#TF#SERVER: Received 2088 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 582181 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=12, total=531).
#TF#SERVER: Sent 215128 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2100 bytes expected).
#TF#SERVER: Received 2100 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=256, total=775).
#TF#SERVER: Sent 709872 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3076 bytes expected).
#TF#SERVER: Received 3076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=7, context=512, query=2, total=521).
#TF#SERVER: Sent 212267 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2060 bytes expected).
#TF#SERVER: Received 2060 bytes from the TF server.

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d1_lp_space, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>![X2]:(m1_subset_1(X2,k1_numbers)=>![X3]:(((v1_funct_1(X3)&v1_funct_2(X3,k5_numbers,k1_numbers))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>(X3=k1_lp_space(X1,X2)<=>![X4]:(m2_subset_1(X4,k1_numbers,k5_numbers)=>k1_seq_1(X3,X4)=k4_power(k18_complex1(k1_seq_1(X1,X4)),X2)))))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', d1_lp_space)).
fof(t2_absvalue, axiom, ![X1]:(v1_xreal_0(X1)=>(X1=k6_numbers<=>k18_complex1(X1)=k6_numbers)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', t2_absvalue)).
fof(t11_lp_space, conjecture, ![X1]:(m1_subset_1(X1,k1_numbers)=>(r1_xxreal_0(np__1,X1)=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,k1_numbers))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>(![X3]:(m2_subset_1(X3,k1_numbers,k5_numbers)=>k1_seq_1(X2,X3)=k6_numbers)=>(v1_series_1(k1_lp_space(X2,X1))&k4_power(k4_series_1(k1_lp_space(X2,X1)),k10_real_1(np__1,X1))=k6_numbers))))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', t11_lp_space)).
fof(cc14_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xreal_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', cc14_membered)).
fof(t1_numerals, axiom, m1_subset_1(k1_xboole_0,k4_ordinal1), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', t1_numerals)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', redefinition_k6_numbers)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', redefinition_k5_numbers)).
fof(cc21_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v3_membered(X2))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', cc21_membered)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', dt_k5_numbers)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', fc3_membered)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', t6_boole)).
fof(dt_k1_lp_space, axiom, ![X1, X2]:((((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))&m1_subset_1(X2,k1_numbers))=>((v1_funct_1(k1_lp_space(X1,X2))&v1_funct_2(k1_lp_space(X1,X2),k5_numbers,k1_numbers))&m1_subset_1(k1_lp_space(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', dt_k1_lp_space)).
fof(rc4_xreal_0, axiom, ?[X1]:(((v1_xboole_0(X1)&v1_xcmplx_0(X1))&v1_xxreal_0(X1))&v1_xreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', rc4_xreal_0)).
fof(projectivity_k18_complex1, axiom, ![X1]:(v1_xcmplx_0(X1)=>k18_complex1(k18_complex1(X1))=k18_complex1(X1)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', projectivity_k18_complex1)).
fof(dt_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)=>m1_subset_1(X3,X1))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', dt_m2_subset_1)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', cc1_subset_1)).
fof(fc6_ordinal1, axiom, (~(v1_xboole_0(k4_ordinal1))&v3_ordinal1(k4_ordinal1)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', fc6_ordinal1)).
fof(d2_power, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>![X3]:(v1_xreal_0(X3)=>(((~(r1_xxreal_0(X1,k6_numbers))=>(X3=k3_power(X1,X2)<=>X3=k9_prepower(X1,X2)))&(X1=k6_numbers=>(r1_xxreal_0(X2,k6_numbers)|(X3=k3_power(X1,X2)<=>X3=k6_numbers))))&(v1_int_1(X2)=>(X3=k3_power(X1,X2)<=>?[X4]:((v1_int_1(X4)&X4=X2)&X3=k4_prepower(X1,X4)))))))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', d2_power)).
fof(t1_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k2_xcmplx_0(X1,k6_numbers)=X1), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', t1_arithm)).
fof(commutativity_k2_xcmplx_0, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', commutativity_k2_xcmplx_0)).
fof(cc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xcmplx_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', cc3_xreal_0)).
fof(dt_k7_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>m1_subset_1(k7_real_1(X1,X2),k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', dt_k7_real_1)).
fof(redefinition_k7_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k7_real_1(X1,X2)=k2_xcmplx_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', redefinition_k7_real_1)).
fof(d7_seq_2, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>(v2_comseq_2(X1)=>![X2]:(v1_xreal_0(X2)=>(X2=k1_seq_2(X1)<=>![X3]:(v1_xreal_0(X3)=>~((~(r1_xxreal_0(X3,k6_numbers))&![X4]:(m2_subset_1(X4,k1_numbers,k5_numbers)=>?[X5]:((m2_subset_1(X5,k1_numbers,k5_numbers)&r1_xxreal_0(X4,X5))&r1_xxreal_0(X3,k18_complex1(k9_real_1(k1_seq_1(X1,X5),X2)))))))))))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', d7_seq_2)).
fof(redefinition_k4_power, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&m1_subset_1(X2,k1_numbers))=>k4_power(X1,X2)=k3_power(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', redefinition_k4_power)).
fof(dt_k6_numbers, axiom, m2_subset_1(k6_numbers,k1_numbers,k5_numbers), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', dt_k6_numbers)).
fof(d1_series_1, axiom, ![X1]:(((((v1_relat_1(X1)&v4_relat_1(X1,k5_numbers))&v1_funct_1(X1))&v1_partfun1(X1,k5_numbers))&v1_valued_0(X1))=>![X2]:(((((v1_relat_1(X2)&v4_relat_1(X2,k5_numbers))&v1_funct_1(X2))&v1_partfun1(X2,k5_numbers))&v1_valued_0(X2))=>(X2=k2_series_1(X1)<=>(k1_funct_1(X2,k6_numbers)=k1_funct_1(X1,k6_numbers)&![X3]:(m2_subset_1(X3,k1_numbers,k5_numbers)=>k1_funct_1(X2,k2_nat_1(X3,np__1))=k2_xcmplx_0(k1_funct_1(X2,X3),k1_funct_1(X1,k2_nat_1(X3,np__1)))))))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', d1_series_1)).
fof(redefinition_k1_seq_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>k1_seq_1(X1,X2)=k1_funct_1(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', redefinition_k1_seq_1)).
fof(cc6_valued_0, axiom, ![X1]:((v1_relat_1(X1)&v3_valued_0(X1))=>(v1_relat_1(X1)&v1_valued_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', cc6_valued_0)).
fof(cc3_funct_2, axiom, ![X1, X2]:(~(v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v1_funct_2(X3,X1,X2)=>v1_partfun1(X3,X1)))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', cc3_funct_2)).
fof(fc1_numbers, axiom, ~(v1_xboole_0(k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', fc1_numbers)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', cc2_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', cc1_relset_1)).
fof(cc18_valued_0, axiom, ![X1, X2]:(v3_membered(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v3_valued_0(X3))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', cc18_valued_0)).
fof(redefinition_k9_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k9_real_1(X1,X2)=k6_xcmplx_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', redefinition_k9_real_1)).
fof(spc0_boole, axiom, v1_xboole_0(np__0), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', spc0_boole)).
fof(spc8_arithm, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k2_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', spc8_arithm)).
fof(involutiveness_k4_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>k4_xcmplx_0(k4_xcmplx_0(X1))=X1), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', involutiveness_k4_xcmplx_0)).
fof(d6_seq_2, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>(v2_comseq_2(X1)<=>?[X2]:(v1_xreal_0(X2)&![X3]:(v1_xreal_0(X3)=>~((~(r1_xxreal_0(X3,k6_numbers))&![X4]:(m2_subset_1(X4,k1_numbers,k5_numbers)=>?[X5]:((m2_subset_1(X5,k1_numbers,k5_numbers)&r1_xxreal_0(X4,X5))&r1_xxreal_0(X3,k18_complex1(k9_real_1(k1_seq_1(X1,X5),X2))))))))))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', d6_seq_2)).
fof(fc7_xreal_0, axiom, ![X1, X2]:((v1_xreal_0(X1)&v1_xreal_0(X2))=>v1_xreal_0(k6_xcmplx_0(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', fc7_xreal_0)).
fof(spc1_arithm, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', spc1_arithm)).
fof(rqRealNeg__k4_xcmplx_0__r0_r0, axiom, k4_xcmplx_0(np__0)=np__0, file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', rqRealNeg__k4_xcmplx_0__r0_r0)).
fof(dt_k4_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>v1_xcmplx_0(k4_xcmplx_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', dt_k4_xcmplx_0)).
fof(d3_series_1, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>k4_series_1(X1)=k2_seq_2(k3_series_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', d3_series_1)).
fof(redefinition_k2_seq_2, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>k2_seq_2(X1)=k1_seq_2(X1)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', redefinition_k2_seq_2)).
fof(fc61_valued_0, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_valued_0(X1))=>v1_xcmplx_0(k1_funct_1(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', fc61_valued_0)).
fof(existence_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>?[X3]:m2_subset_1(X3,X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', existence_m2_subset_1)).
fof(dt_k3_series_1, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>((v1_funct_1(k3_series_1(X1))&v1_funct_2(k3_series_1(X1),k5_numbers,k1_numbers))&m1_subset_1(k3_series_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', dt_k3_series_1)).
fof(t4_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k6_xcmplx_0(X1,k6_numbers)=X1), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', t4_arithm)).
fof(redefinition_k3_series_1, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>k3_series_1(X1)=k2_series_1(X1)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', redefinition_k3_series_1)).
fof(d2_series_1, axiom, ![X1]:(((v1_funct_1(X1)&v1_funct_2(X1,k5_numbers,k1_numbers))&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>(v1_series_1(X1)<=>v2_comseq_2(k3_series_1(X1)))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', d2_series_1)).
fof(rqLessOrEqual__r1_xxreal_0__r1_r0, axiom, ~(r1_xxreal_0(np__1,np__0)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', rqLessOrEqual__r1_xxreal_0__r1_r0)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', spc1_numerals)).
fof(rc4_xxreal_0, axiom, ?[X1]:(v1_xboole_0(X1)&v1_xxreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', rc4_xxreal_0)).
fof(redefinition_k10_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', redefinition_k10_real_1)).
fof(cc13_membered, axiom, ![X1]:(v2_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xxreal_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', cc13_membered)).
fof(t4_real, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>~(((r1_xxreal_0(X1,X2)&~(v2_xxreal_0(X2)))&v2_xxreal_0(X1))))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', t4_real)).
fof(cc2_xxreal_0, axiom, ![X1]:(v7_ordinal1(X1)=>v1_xxreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', cc2_xxreal_0)).
fof(cc7_ordinal1, axiom, ![X1]:(v1_xboole_0(X1)=>v7_ordinal1(X1)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', cc7_ordinal1)).
fof(t139_xreal_1, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>~(((~(r1_xxreal_0(X1,k6_numbers))&~(r1_xxreal_0(X2,k6_numbers)))&r1_xxreal_0(k7_xcmplx_0(X1,X2),k6_numbers))))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', t139_xreal_1)).
fof(cc4_membered, axiom, ![X1]:(v3_membered(X1)=>v2_membered(X1)), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', cc4_membered)).
fof(fc16_xreal_0, axiom, ![X1]:((~(v3_xxreal_0(X1))&v1_xreal_0(X1))=>(v1_xcmplx_0(k4_xcmplx_0(X1))&~(v2_xxreal_0(k4_xcmplx_0(X1))))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', fc16_xreal_0)).
fof(cc5_xxreal_0, axiom, ![X1]:((v1_xxreal_0(X1)&v3_xxreal_0(X1))=>((~(v1_xboole_0(X1))&v1_xxreal_0(X1))&~(v2_xxreal_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', cc5_xxreal_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', fc1_xboole_0)).
fof(cc3_xxreal_0, axiom, ![X1]:((v1_xxreal_0(X1)&v2_xxreal_0(X1))=>((~(v1_xboole_0(X1))&v1_xxreal_0(X1))&~(v3_xxreal_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', cc3_xxreal_0)).
fof(fc8_xreal_0, axiom, ![X1, X2]:((v1_xreal_0(X1)&v1_xreal_0(X2))=>v1_xreal_0(k7_xcmplx_0(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', fc8_xreal_0)).
fof(t6_real, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>(r1_xxreal_0(X1,X2)=>(v1_xboole_0(X1)|v2_xxreal_0(X2)|v3_xxreal_0(X1))))), file('/scratch/mptp/allpred/preds__0.005/t11_lp_space', t6_real)).
fof(c_0_67, plain, ![X18, X19, X20, X21]:((X20!=k1_lp_space(X18,X19)|(~m2_subset_1(X21,k1_numbers,k5_numbers)|k1_seq_1(X20,X21)=k4_power(k18_complex1(k1_seq_1(X18,X21)),X19))|(~v1_funct_1(X20)|~v1_funct_2(X20,k5_numbers,k1_numbers)|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|~m1_subset_1(X19,k1_numbers)|(~v1_funct_1(X18)|~v1_funct_2(X18,k5_numbers,k1_numbers)|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&((m2_subset_1(esk6_3(X18,X19,X20),k1_numbers,k5_numbers)|X20=k1_lp_space(X18,X19)|(~v1_funct_1(X20)|~v1_funct_2(X20,k5_numbers,k1_numbers)|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|~m1_subset_1(X19,k1_numbers)|(~v1_funct_1(X18)|~v1_funct_2(X18,k5_numbers,k1_numbers)|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(k1_seq_1(X20,esk6_3(X18,X19,X20))!=k4_power(k18_complex1(k1_seq_1(X18,esk6_3(X18,X19,X20))),X19)|X20=k1_lp_space(X18,X19)|(~v1_funct_1(X20)|~v1_funct_2(X20,k5_numbers,k1_numbers)|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|~m1_subset_1(X19,k1_numbers)|(~v1_funct_1(X18)|~v1_funct_2(X18,k5_numbers,k1_numbers)|~m1_subset_1(X18,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_lp_space])])])])])).
fof(c_0_68, plain, ![X32]:((X32!=k6_numbers|k18_complex1(X32)=k6_numbers|~v1_xreal_0(X32))&(k18_complex1(X32)!=k6_numbers|X32=k6_numbers|~v1_xreal_0(X32))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_absvalue])])])).
fof(c_0_69, negated_conjecture, ~(![X1]:(m1_subset_1(X1,k1_numbers)=>(r1_xxreal_0(np__1,X1)=>![X2]:(((v1_funct_1(X2)&v1_funct_2(X2,k5_numbers,k1_numbers))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))=>(![X3]:(m2_subset_1(X3,k1_numbers,k5_numbers)=>k1_seq_1(X2,X3)=k6_numbers)=>(v1_series_1(k1_lp_space(X2,X1))&k4_power(k4_series_1(k1_lp_space(X2,X1)),k10_real_1(np__1,X1))=k6_numbers)))))), inference(assume_negation,[status(cth)],[t11_lp_space])).
fof(c_0_70, plain, ![X80, X81]:(~v3_membered(X80)|(~m1_subset_1(X81,X80)|v1_xreal_0(X81))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc14_membered])])])).
cnf(c_0_71, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1)), inference(split_conjunct,[status(thm)],[t1_numerals])).
cnf(c_0_72, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
cnf(c_0_73, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
fof(c_0_74, plain, ![X120, X121]:(~v3_membered(X120)|(~m1_subset_1(X121,k1_zfmisc_1(X120))|v3_membered(X121))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc21_membered])])])).
cnf(c_0_75, plain, (X1=k1_lp_space(X2,X3)|k1_seq_1(X1,esk6_3(X2,X3,X1))!=k4_power(k18_complex1(k1_seq_1(X2,esk6_3(X2,X3,X1))),X3)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_76, plain, (k18_complex1(X1)=k6_numbers|X1!=k6_numbers|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
fof(c_0_77, negated_conjecture, ![X8]:(m1_subset_1(esk1_0,k1_numbers)&(r1_xxreal_0(np__1,esk1_0)&(((v1_funct_1(esk2_0)&v1_funct_2(esk2_0,k5_numbers,k1_numbers))&m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))&((~m2_subset_1(X8,k1_numbers,k5_numbers)|k1_seq_1(esk2_0,X8)=k6_numbers)&(~v1_series_1(k1_lp_space(esk2_0,esk1_0))|k4_power(k4_series_1(k1_lp_space(esk2_0,esk1_0)),k10_real_1(np__1,esk1_0))!=k6_numbers))))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_69])])])])).
cnf(c_0_78, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_79, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_71])).
cnf(c_0_80, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_72])).
cnf(c_0_81, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_73])).
cnf(c_0_82, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_83, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
cnf(c_0_84, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
fof(c_0_85, plain, ![X79]:(~v1_xboole_0(X79)|X79=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
fof(c_0_86, plain, ![X42, X43]:(((v1_funct_1(k1_lp_space(X42,X43))|(~v1_funct_1(X42)|~v1_funct_2(X42,k5_numbers,k1_numbers)|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X43,k1_numbers)))&(v1_funct_2(k1_lp_space(X42,X43),k5_numbers,k1_numbers)|(~v1_funct_1(X42)|~v1_funct_2(X42,k5_numbers,k1_numbers)|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X43,k1_numbers))))&(m1_subset_1(k1_lp_space(X42,X43),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|(~v1_funct_1(X42)|~v1_funct_2(X42,k5_numbers,k1_numbers)|~m1_subset_1(X42,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X43,k1_numbers)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_lp_space])])])).
cnf(c_0_87, plain, (X1=k1_lp_space(X2,X3)|k1_seq_1(X1,esk6_3(X2,X3,X1))!=k4_power(k18_complex1(k1_seq_1(X2,esk6_3(X2,X3,X1))),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_75])).
cnf(c_0_88, plain, (k18_complex1(X1)=k6_numbers|X1!=k6_numbers|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_89, negated_conjecture, (k1_seq_1(esk2_0,X1)=k6_numbers|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_90, negated_conjecture, (v1_funct_2(esk2_0,k5_numbers,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_91, negated_conjecture, (v1_funct_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_92, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_93, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_78])).
cnf(c_0_94, plain, (m1_subset_1(k6_numbers,k5_numbers)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_79, c_0_80]), c_0_81])).
cnf(c_0_95, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_82])).
cnf(c_0_96, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_83])).
cnf(c_0_97, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_84])).
fof(c_0_98, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)<=>(((~(r1_xxreal_0(X1,k6_numbers))=>(X3=k3_power(X1,X2)<=>X3=k9_prepower(X1,X2)))&(X1=k6_numbers=>(r1_xxreal_0(X2,k6_numbers)|(X3=k3_power(X1,X2)<=>X3=k6_numbers))))&(v1_int_1(X2)=>(X3=k3_power(X1,X2)<=>?[X4]:((v1_int_1(X4)&X4=X2)&X3=k4_prepower(X1,X4)))))), introduced(definition)).
cnf(c_0_99, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
fof(c_0_100, plain, (((v1_xboole_0(esk13_0)&v1_xcmplx_0(esk13_0))&v1_xxreal_0(esk13_0))&v1_xreal_0(esk13_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc4_xreal_0])])).
cnf(c_0_101, plain, (k1_seq_1(X1,X4)=k4_power(k18_complex1(k1_seq_1(X2,X4)),X3)|X1!=k1_lp_space(X2,X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_funct_1(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_102, plain, (m1_subset_1(k1_lp_space(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_86])).
fof(c_0_103, plain, ![X163]:(~v1_xcmplx_0(X163)|k18_complex1(k18_complex1(X163))=k18_complex1(X163)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[projectivity_k18_complex1])])).
cnf(c_0_104, plain, (X1=k1_lp_space(X2,X3)|k1_seq_1(X1,esk6_3(X2,X3,X1))!=k4_power(k6_numbers,X3)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|k1_seq_1(X2,esk6_3(X2,X3,X1))!=k6_numbers|~v1_xreal_0(k1_seq_1(X2,esk6_3(X2,X3,X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88])])).
cnf(c_0_105, negated_conjecture, (k1_seq_1(esk2_0,X1)=k6_numbers|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_106, negated_conjecture, (v1_funct_2(esk2_0,k5_numbers,k1_numbers)), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_107, negated_conjecture, (v1_funct_1(esk2_0)), inference(evalgc,[status(thm)],[c_0_91])).
cnf(c_0_108, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_92])).
cnf(c_0_109, plain, (v1_xreal_0(k6_numbers)|~v3_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94])])).
cnf(c_0_110, plain, (v3_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_97])])])).
fof(c_0_111, plain, ![X108, X109, X110]:(v1_xboole_0(X108)|v1_xboole_0(X109)|~m1_subset_1(X109,k1_zfmisc_1(X108))|(~m2_subset_1(X110,X108,X109)|m1_subset_1(X110,X108))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_subset_1])])])])).
fof(c_0_112, plain, ![X75, X76]:(~v1_xboole_0(X75)|(~m1_subset_1(X76,k1_zfmisc_1(X75))|v1_xboole_0(X76))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
fof(c_0_113, plain, (~v1_xboole_0(k4_ordinal1)&v3_ordinal1(k4_ordinal1)), inference(fof_simplification,[status(thm)],[fc6_ordinal1])).
fof(c_0_114, plain, ![X3, X2, X1]:(epred1_3(X1,X2,X3)=>(((~(r1_xxreal_0(X1,k6_numbers))=>(X3=k3_power(X1,X2)<=>X3=k9_prepower(X1,X2)))&(X1=k6_numbers=>(r1_xxreal_0(X2,k6_numbers)|(X3=k3_power(X1,X2)<=>X3=k6_numbers))))&(v1_int_1(X2)=>(X3=k3_power(X1,X2)<=>?[X4]:((v1_int_1(X4)&X4=X2)&X3=k4_prepower(X1,X4)))))), inference(split_equiv,[status(thm)],[c_0_98])).
fof(c_0_115, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>![X3]:(v1_xreal_0(X3)=>epred1_3(X1,X2,X3)))), inference(apply_def,[status(thm)],[d2_power, c_0_98])).
fof(c_0_116, plain, ![X102]:(~v1_xcmplx_0(X102)|k2_xcmplx_0(X102,k6_numbers)=X102), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_arithm])])).
fof(c_0_117, plain, ![X89, X90]:(~v1_xcmplx_0(X89)|~v1_xcmplx_0(X90)|k2_xcmplx_0(X89,X90)=k2_xcmplx_0(X90,X89)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k2_xcmplx_0])])).
cnf(c_0_118, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_99])).
cnf(c_0_119, plain, (v1_xboole_0(esk13_0)), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_120, plain, (k1_seq_1(X1,X4)=k4_power(k18_complex1(k1_seq_1(X2,X4)),X3)|X1!=k1_lp_space(X2,X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_101])).
cnf(c_0_121, plain, (m1_subset_1(k1_lp_space(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_102])).
cnf(c_0_122, plain, (v1_funct_1(k1_lp_space(X1,X2))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_123, plain, (v1_funct_2(k1_lp_space(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_124, plain, (k18_complex1(k18_complex1(X1))=k18_complex1(X1)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_103])).
fof(c_0_125, plain, ![X129]:(~v1_xreal_0(X129)|v1_xcmplx_0(X129)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_xreal_0])])).
cnf(c_0_126, negated_conjecture, (X1=k1_lp_space(esk2_0,X2)|k1_seq_1(X1,esk6_3(esk2_0,X2,X1))!=k4_power(k6_numbers,X2)|~v1_xreal_0(k6_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(esk6_3(esk2_0,X2,X1),k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_105]), c_0_106]), c_0_107]), c_0_108])])])).
cnf(c_0_127, plain, (v1_xreal_0(k6_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_109, c_0_110])])).
cnf(c_0_128, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m2_subset_1(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_129, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_130, plain, (~v1_xboole_0(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_113])).
fof(c_0_131, plain, ![X256, X257, X258, X260]:((((X256!=k3_power(X258,X257)|X256=k9_prepower(X258,X257)|r1_xxreal_0(X258,k6_numbers)|~epred1_3(X258,X257,X256))&(X256!=k9_prepower(X258,X257)|X256=k3_power(X258,X257)|r1_xxreal_0(X258,k6_numbers)|~epred1_3(X258,X257,X256)))&((X256!=k3_power(X258,X257)|X256=k6_numbers|r1_xxreal_0(X257,k6_numbers)|X258!=k6_numbers|~epred1_3(X258,X257,X256))&(X256!=k6_numbers|X256=k3_power(X258,X257)|r1_xxreal_0(X257,k6_numbers)|X258!=k6_numbers|~epred1_3(X258,X257,X256))))&((((v1_int_1(esk17_3(X256,X257,X258))|X256!=k3_power(X258,X257)|~v1_int_1(X257)|~epred1_3(X258,X257,X256))&(esk17_3(X256,X257,X258)=X257|X256!=k3_power(X258,X257)|~v1_int_1(X257)|~epred1_3(X258,X257,X256)))&(X256=k4_prepower(X258,esk17_3(X256,X257,X258))|X256!=k3_power(X258,X257)|~v1_int_1(X257)|~epred1_3(X258,X257,X256)))&(~v1_int_1(X260)|X260!=X257|X256!=k4_prepower(X258,X260)|X256=k3_power(X258,X257)|~v1_int_1(X257)|~epred1_3(X258,X257,X256)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_114])])])])])])).
fof(c_0_132, plain, ![X37, X38, X39]:(~v1_xreal_0(X37)|(~v1_xreal_0(X38)|(~v1_xreal_0(X39)|epred1_3(X37,X38,X39)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_115])])])).
fof(c_0_133, plain, ![X100, X101]:(~m1_subset_1(X100,k1_numbers)|~v1_xreal_0(X101)|m1_subset_1(k7_real_1(X100,X101),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_real_1])])).
fof(c_0_134, plain, ![X131, X132]:(~m1_subset_1(X131,k1_numbers)|~v1_xreal_0(X132)|k7_real_1(X131,X132)=k2_xcmplx_0(X131,X132)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_real_1])])).
cnf(c_0_135, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_116])).
cnf(c_0_136, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_117])).
cnf(c_0_137, plain, (v1_xcmplx_0(esk13_0)), inference(split_conjunct,[status(thm)],[c_0_100])).
cnf(c_0_138, plain, (X1=k6_numbers|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_118, c_0_80])).
cnf(c_0_139, plain, (v1_xboole_0(esk13_0)), inference(evalgc,[status(thm)],[c_0_119])).
fof(c_0_140, plain, ![X10, X11, X12, X14, X16]:(((m2_subset_1(esk3_3(X10,X11,X12),k1_numbers,k5_numbers)|r1_xxreal_0(X12,k6_numbers)|~v1_xreal_0(X12)|X11!=k1_seq_2(X10)|~v1_xreal_0(X11)|~v2_comseq_2(X10)|(~v1_funct_1(X10)|~v1_funct_2(X10,k5_numbers,k1_numbers)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(~m2_subset_1(X14,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_3(X10,X11,X12),X14)|~r1_xxreal_0(X12,k18_complex1(k9_real_1(k1_seq_1(X10,X14),X11)))|r1_xxreal_0(X12,k6_numbers)|~v1_xreal_0(X12)|X11!=k1_seq_2(X10)|~v1_xreal_0(X11)|~v2_comseq_2(X10)|(~v1_funct_1(X10)|~v1_funct_2(X10,k5_numbers,k1_numbers)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))&((v1_xreal_0(esk4_2(X10,X11))|X11=k1_seq_2(X10)|~v1_xreal_0(X11)|~v2_comseq_2(X10)|(~v1_funct_1(X10)|~v1_funct_2(X10,k5_numbers,k1_numbers)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&((~r1_xxreal_0(esk4_2(X10,X11),k6_numbers)|X11=k1_seq_2(X10)|~v1_xreal_0(X11)|~v2_comseq_2(X10)|(~v1_funct_1(X10)|~v1_funct_2(X10,k5_numbers,k1_numbers)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(((m2_subset_1(esk5_3(X10,X11,X16),k1_numbers,k5_numbers)|~m2_subset_1(X16,k1_numbers,k5_numbers)|X11=k1_seq_2(X10)|~v1_xreal_0(X11)|~v2_comseq_2(X10)|(~v1_funct_1(X10)|~v1_funct_2(X10,k5_numbers,k1_numbers)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(r1_xxreal_0(X16,esk5_3(X10,X11,X16))|~m2_subset_1(X16,k1_numbers,k5_numbers)|X11=k1_seq_2(X10)|~v1_xreal_0(X11)|~v2_comseq_2(X10)|(~v1_funct_1(X10)|~v1_funct_2(X10,k5_numbers,k1_numbers)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))&(r1_xxreal_0(esk4_2(X10,X11),k18_complex1(k9_real_1(k1_seq_1(X10,esk5_3(X10,X11,X16)),X11)))|~m2_subset_1(X16,k1_numbers,k5_numbers)|X11=k1_seq_2(X10)|~v1_xreal_0(X11)|~v2_comseq_2(X10)|(~v1_funct_1(X10)|~v1_funct_2(X10,k5_numbers,k1_numbers)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_seq_2])])])])])])).
cnf(c_0_141, plain, (k1_seq_1(k1_lp_space(X1,X2),X3)=k4_power(k18_complex1(k1_seq_1(X4,X3)),X5)|k1_lp_space(X1,X2)!=k1_lp_space(X4,X5)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X4,k5_numbers,k1_numbers)|~v1_funct_2(k1_lp_space(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(X4)|~v1_funct_1(k1_lp_space(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X5,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_121])])).
cnf(c_0_142, plain, (v1_funct_1(k1_lp_space(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_122])).
cnf(c_0_143, plain, (v1_funct_2(k1_lp_space(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_123])).
cnf(c_0_144, plain, (k18_complex1(k18_complex1(X1))=k18_complex1(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_124])).
cnf(c_0_145, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_125])).
cnf(c_0_146, negated_conjecture, (X1=k1_lp_space(esk2_0,X2)|k1_seq_1(X1,esk6_3(esk2_0,X2,X1))!=k4_power(k6_numbers,X2)|~m2_subset_1(esk6_3(esk2_0,X2,X1),k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_126, c_0_127])])).
cnf(c_0_147, plain, (m2_subset_1(esk6_3(X1,X2,X3),k1_numbers,k5_numbers)|X3=k1_lp_space(X1,X2)|~v1_funct_1(X3)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_67])).
fof(c_0_148, plain, ![X57, X58]:(~m1_subset_1(X57,k1_numbers)|~m1_subset_1(X58,k1_numbers)|k4_power(X57,X58)=k3_power(X57,X58)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_power])])).
cnf(c_0_149, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_128])).
cnf(c_0_150, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_129])).
cnf(c_0_151, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[dt_k6_numbers])).
cnf(c_0_152, plain, (~v1_xboole_0(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_130])).
cnf(c_0_153, plain, (X1=k3_power(X2,X3)|r1_xxreal_0(X3,k6_numbers)|X1!=k6_numbers|X2!=k6_numbers|~epred1_3(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_131])).
cnf(c_0_154, plain, (epred1_3(X1,X2,X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_155, plain, (m1_subset_1(k7_real_1(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_133])).
cnf(c_0_156, plain, (k7_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_134])).
cnf(c_0_157, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_135])).
cnf(c_0_158, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_136])).
cnf(c_0_159, plain, (v1_xcmplx_0(esk13_0)), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_160, plain, (esk13_0=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139])])).
fof(c_0_161, plain, ![X59, X60, X61]:(((k1_funct_1(X60,k6_numbers)=k1_funct_1(X59,k6_numbers)|X60!=k2_series_1(X59)|(~v1_relat_1(X60)|~v4_relat_1(X60,k5_numbers)|~v1_funct_1(X60)|~v1_partfun1(X60,k5_numbers)|~v1_valued_0(X60))|(~v1_relat_1(X59)|~v4_relat_1(X59,k5_numbers)|~v1_funct_1(X59)|~v1_partfun1(X59,k5_numbers)|~v1_valued_0(X59)))&(~m2_subset_1(X61,k1_numbers,k5_numbers)|k1_funct_1(X60,k2_nat_1(X61,np__1))=k2_xcmplx_0(k1_funct_1(X60,X61),k1_funct_1(X59,k2_nat_1(X61,np__1)))|X60!=k2_series_1(X59)|(~v1_relat_1(X60)|~v4_relat_1(X60,k5_numbers)|~v1_funct_1(X60)|~v1_partfun1(X60,k5_numbers)|~v1_valued_0(X60))|(~v1_relat_1(X59)|~v4_relat_1(X59,k5_numbers)|~v1_funct_1(X59)|~v1_partfun1(X59,k5_numbers)|~v1_valued_0(X59))))&((m2_subset_1(esk12_2(X59,X60),k1_numbers,k5_numbers)|k1_funct_1(X60,k6_numbers)!=k1_funct_1(X59,k6_numbers)|X60=k2_series_1(X59)|(~v1_relat_1(X60)|~v4_relat_1(X60,k5_numbers)|~v1_funct_1(X60)|~v1_partfun1(X60,k5_numbers)|~v1_valued_0(X60))|(~v1_relat_1(X59)|~v4_relat_1(X59,k5_numbers)|~v1_funct_1(X59)|~v1_partfun1(X59,k5_numbers)|~v1_valued_0(X59)))&(k1_funct_1(X60,k2_nat_1(esk12_2(X59,X60),np__1))!=k2_xcmplx_0(k1_funct_1(X60,esk12_2(X59,X60)),k1_funct_1(X59,k2_nat_1(esk12_2(X59,X60),np__1)))|k1_funct_1(X60,k6_numbers)!=k1_funct_1(X59,k6_numbers)|X60=k2_series_1(X59)|(~v1_relat_1(X60)|~v4_relat_1(X60,k5_numbers)|~v1_funct_1(X60)|~v1_partfun1(X60,k5_numbers)|~v1_valued_0(X60))|(~v1_relat_1(X59)|~v4_relat_1(X59,k5_numbers)|~v1_funct_1(X59)|~v1_partfun1(X59,k5_numbers)|~v1_valued_0(X59))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_series_1])])])])])).
fof(c_0_162, plain, ![X70, X71]:(~v1_relat_1(X70)|~v1_funct_1(X70)|~v3_valued_0(X70)|k1_seq_1(X70,X71)=k1_funct_1(X70,X71)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_seq_1])])).
cnf(c_0_163, plain, (r1_xxreal_0(esk4_2(X1,X2),k18_complex1(k9_real_1(k1_seq_1(X1,esk5_3(X1,X2,X3)),X2)))|X2=k1_seq_2(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_140])).
cnf(c_0_164, plain, (k1_seq_1(k1_lp_space(X1,X2),X3)=k4_power(k18_complex1(k1_seq_1(X4,X3)),X5)|k1_lp_space(X1,X2)!=k1_lp_space(X4,X5)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X4,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X4)|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X5,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_141, c_0_142]), c_0_143])).
cnf(c_0_165, plain, (k18_complex1(k6_numbers)=k6_numbers|~v1_xcmplx_0(X1)|X1!=k6_numbers|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_88])])).
cnf(c_0_166, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_145])).
cnf(c_0_167, negated_conjecture, (k1_lp_space(esk2_0,X1)=esk2_0|k4_power(k6_numbers,X1)!=k6_numbers|~m2_subset_1(esk6_3(esk2_0,X1,esk2_0),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_146, c_0_105]), c_0_106]), c_0_107]), c_0_108])])])).
cnf(c_0_168, plain, (X3=k1_lp_space(X1,X2)|m2_subset_1(esk6_3(X1,X2,X3),k1_numbers,k5_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_147])).
cnf(c_0_169, plain, (k4_power(X1,X2)=k3_power(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_148])).
cnf(c_0_170, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~m2_subset_1(X2,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(csr,[status(thm)],[c_0_149, c_0_150])).
cnf(c_0_171, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_151])).
cnf(c_0_172, plain, (~v1_xboole_0(k5_numbers)), inference(rw,[status(thm)],[c_0_152, c_0_81])).
cnf(c_0_173, plain, (X1=k3_power(X2,X3)|r1_xxreal_0(X3,k6_numbers)|X2!=k6_numbers|X1!=k6_numbers|~epred1_3(X2,X3,X1)), inference(evalgc,[status(thm)],[c_0_153])).
cnf(c_0_174, plain, (epred1_3(X1,X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_154])).
cnf(c_0_175, plain, (m1_subset_1(k7_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_155])).
cnf(c_0_176, plain, (k7_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_156])).
cnf(c_0_177, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_158])])).
cnf(c_0_178, plain, (v1_xcmplx_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_159, c_0_160])])).
cnf(c_0_179, plain, (X1=k2_series_1(X2)|k1_funct_1(X1,k2_nat_1(esk12_2(X2,X1),np__1))!=k2_xcmplx_0(k1_funct_1(X1,esk12_2(X2,X1)),k1_funct_1(X2,k2_nat_1(esk12_2(X2,X1),np__1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X1)|~v1_relat_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_funct_1(X2)|~v1_partfun1(X2,k5_numbers)|~v1_valued_0(X2)), inference(split_conjunct,[status(thm)],[c_0_161])).
cnf(c_0_180, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_162])).
fof(c_0_181, plain, ![X172]:((v1_relat_1(X172)|(~v1_relat_1(X172)|~v3_valued_0(X172)))&(v1_valued_0(X172)|(~v1_relat_1(X172)|~v3_valued_0(X172)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc6_valued_0])])])).
fof(c_0_182, plain, ![X117, X118, X119]:(v1_xboole_0(X118)|(~m1_subset_1(X119,k1_zfmisc_1(k2_zfmisc_1(X117,X118)))|(~v1_funct_2(X119,X117,X118)|v1_partfun1(X119,X117)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc3_funct_2])])])])).
fof(c_0_183, plain, ~v1_xboole_0(k1_numbers), inference(fof_simplification,[status(thm)],[fc1_numbers])).
fof(c_0_184, plain, ![X91, X92, X93]:((v4_relat_1(X93,X91)|~m1_subset_1(X93,k1_zfmisc_1(k2_zfmisc_1(X91,X92))))&(v5_relat_1(X93,X92)|~m1_subset_1(X93,k1_zfmisc_1(k2_zfmisc_1(X91,X92))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_185, plain, ![X65, X66, X67]:(~m1_subset_1(X67,k1_zfmisc_1(k2_zfmisc_1(X65,X66)))|v1_relat_1(X67)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_186, plain, ![X54, X55, X56]:(~v3_membered(X55)|(~m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(X54,X55)))|v3_valued_0(X56))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc18_valued_0])])])).
cnf(c_0_187, plain, (X2=k1_seq_2(X1)|r1_xxreal_0(esk4_2(X1,X2),k18_complex1(k9_real_1(k1_seq_1(X1,esk5_3(X1,X2,X3)),X2)))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_163])).
cnf(c_0_188, plain, (X2=k1_seq_2(X1)|~r1_xxreal_0(esk4_2(X1,X2),k6_numbers)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_140])).
fof(c_0_189, plain, ![X82, X83]:(~m1_subset_1(X82,k1_numbers)|~v1_xreal_0(X83)|k9_real_1(X82,X83)=k6_xcmplx_0(X82,X83)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k9_real_1])])).
cnf(c_0_190, plain, (v1_xboole_0(np__0)), inference(split_conjunct,[status(thm)],[spc0_boole])).
fof(c_0_191, plain, ![X225, X226]:(~v1_xcmplx_0(X225)|~v1_xcmplx_0(X226)|k2_xcmplx_0(k4_xcmplx_0(X225),k4_xcmplx_0(X226))=k4_xcmplx_0(k2_xcmplx_0(X225,X226))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc8_arithm])])).
fof(c_0_192, plain, ![X211]:(~v1_xcmplx_0(X211)|k4_xcmplx_0(k4_xcmplx_0(X211))=X211), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k4_xcmplx_0])])).
fof(c_0_193, plain, ![X23, X25, X27, X28, X30]:(((v1_xreal_0(esk7_1(X23))|~v2_comseq_2(X23)|(~v1_funct_1(X23)|~v1_funct_2(X23,k5_numbers,k1_numbers)|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&((m2_subset_1(esk8_2(X23,X25),k1_numbers,k5_numbers)|r1_xxreal_0(X25,k6_numbers)|~v1_xreal_0(X25)|~v2_comseq_2(X23)|(~v1_funct_1(X23)|~v1_funct_2(X23,k5_numbers,k1_numbers)|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(~m2_subset_1(X27,k1_numbers,k5_numbers)|~r1_xxreal_0(esk8_2(X23,X25),X27)|~r1_xxreal_0(X25,k18_complex1(k9_real_1(k1_seq_1(X23,X27),esk7_1(X23))))|r1_xxreal_0(X25,k6_numbers)|~v1_xreal_0(X25)|~v2_comseq_2(X23)|(~v1_funct_1(X23)|~v1_funct_2(X23,k5_numbers,k1_numbers)|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))))&((v1_xreal_0(esk9_2(X23,X28))|~v1_xreal_0(X28)|v2_comseq_2(X23)|(~v1_funct_1(X23)|~v1_funct_2(X23,k5_numbers,k1_numbers)|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&((~r1_xxreal_0(esk9_2(X23,X28),k6_numbers)|~v1_xreal_0(X28)|v2_comseq_2(X23)|(~v1_funct_1(X23)|~v1_funct_2(X23,k5_numbers,k1_numbers)|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(((m2_subset_1(esk10_3(X23,X28,X30),k1_numbers,k5_numbers)|~m2_subset_1(X30,k1_numbers,k5_numbers)|~v1_xreal_0(X28)|v2_comseq_2(X23)|(~v1_funct_1(X23)|~v1_funct_2(X23,k5_numbers,k1_numbers)|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(r1_xxreal_0(X30,esk10_3(X23,X28,X30))|~m2_subset_1(X30,k1_numbers,k5_numbers)|~v1_xreal_0(X28)|v2_comseq_2(X23)|(~v1_funct_1(X23)|~v1_funct_2(X23,k5_numbers,k1_numbers)|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))&(r1_xxreal_0(esk9_2(X23,X28),k18_complex1(k9_real_1(k1_seq_1(X23,esk10_3(X23,X28,X30)),X28)))|~m2_subset_1(X30,k1_numbers,k5_numbers)|~v1_xreal_0(X28)|v2_comseq_2(X23)|(~v1_funct_1(X23)|~v1_funct_2(X23,k5_numbers,k1_numbers)|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_seq_2])])])])])])).
cnf(c_0_194, plain, (k4_power(k18_complex1(k1_seq_1(X1,X3)),X2)=k1_seq_1(k1_lp_space(X1,X2),X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_164])])).
cnf(c_0_195, plain, (k18_complex1(k6_numbers)=k6_numbers|X1!=k6_numbers|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_165, c_0_166])).
cnf(c_0_196, negated_conjecture, (k1_lp_space(esk2_0,X1)=esk2_0|k4_power(k6_numbers,X1)!=k6_numbers|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_168]), c_0_106]), c_0_107]), c_0_108])])])).
cnf(c_0_197, plain, (k4_power(X1,X2)=k3_power(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_169])).
cnf(c_0_198, plain, (m1_subset_1(k6_numbers,k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_171]), c_0_96])]), c_0_172])])).
cnf(c_0_199, plain, (X1=k3_power(X2,X3)|r1_xxreal_0(X3,k6_numbers)|X2!=k6_numbers|X1!=k6_numbers|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_174])])).
cnf(c_0_200, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_176])])).
cnf(c_0_201, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_xcmplx_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_177, c_0_178])])).
cnf(c_0_202, plain, (X1=k2_series_1(X2)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|k2_xcmplx_0(k1_funct_1(X1,esk12_2(X2,X1)),k1_funct_1(X2,k2_nat_1(esk12_2(X2,X1),np__1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(X2,X1),np__1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_179])).
cnf(c_0_203, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[c_0_180])).
cnf(c_0_204, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_181])).
cnf(c_0_205, plain, (v1_xboole_0(X1)|v1_partfun1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))|~v1_funct_2(X2,X3,X1)), inference(split_conjunct,[status(thm)],[c_0_182])).
cnf(c_0_206, plain, (~v1_xboole_0(k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_183])).
cnf(c_0_207, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_184])).
cnf(c_0_208, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_185])).
cnf(c_0_209, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_186])).
cnf(c_0_210, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk4_2(X2,X1),k6_numbers)|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|k9_real_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1)!=k6_numbers|~v1_xreal_0(k9_real_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_88])])).
cnf(c_0_211, plain, (X2=k1_seq_2(X1)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~r1_xxreal_0(esk4_2(X1,X2),k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_188])).
cnf(c_0_212, plain, (k9_real_1(X1,X2)=k6_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_189])).
fof(c_0_213, plain, ![X164, X165]:(~v1_xreal_0(X164)|~v1_xreal_0(X165)|v1_xreal_0(k6_xcmplx_0(X164,X165))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc7_xreal_0])])).
fof(c_0_214, plain, ![X229, X230]:(~v1_xcmplx_0(X229)|~v1_xcmplx_0(X230)|k2_xcmplx_0(X229,k4_xcmplx_0(X230))=k6_xcmplx_0(X229,X230)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc1_arithm])])).
cnf(c_0_215, plain, (k4_xcmplx_0(np__0)=np__0), inference(split_conjunct,[status(thm)],[rqRealNeg__k4_xcmplx_0__r0_r0])).
cnf(c_0_216, plain, (v1_xboole_0(np__0)), inference(evalgc,[status(thm)],[c_0_190])).
cnf(c_0_217, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_191])).
cnf(c_0_218, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_192])).
fof(c_0_219, plain, ![X152]:(~v1_xcmplx_0(X152)|v1_xcmplx_0(k4_xcmplx_0(X152))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_xcmplx_0])])).
cnf(c_0_220, plain, (r1_xxreal_0(esk9_2(X1,X2),k18_complex1(k9_real_1(k1_seq_1(X1,esk10_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_193])).
cnf(c_0_221, negated_conjecture, (k4_power(k18_complex1(k6_numbers),X2)=k1_seq_1(k1_lp_space(esk2_0,X2),X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_105]), c_0_106]), c_0_107]), c_0_108])])])).
cnf(c_0_222, plain, (k18_complex1(k6_numbers)=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_127])])).
cnf(c_0_223, negated_conjecture, (k1_lp_space(esk2_0,X1)=esk2_0|k3_power(k6_numbers,X1)!=k6_numbers|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_197]), c_0_198])])])).
cnf(c_0_224, plain, (k3_power(X1,X2)=k6_numbers|r1_xxreal_0(X2,k6_numbers)|X1!=k6_numbers|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_127])])).
cnf(c_0_225, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_201]), c_0_198])])])).
cnf(c_0_226, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_seq_1(X1,esk12_2(X2,X1)),k1_funct_1(X2,k2_nat_1(esk12_2(X2,X1),np__1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(X2,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_203])])).
cnf(c_0_227, plain, (v1_valued_0(X1)|~v3_valued_0(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[c_0_204])).
cnf(c_0_228, plain, (v1_xboole_0(X1)|v1_partfun1(X2,X3)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_205])).
cnf(c_0_229, plain, (~v1_xboole_0(k1_numbers)), inference(evalgc,[status(thm)],[c_0_206])).
cnf(c_0_230, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_207])).
cnf(c_0_231, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_208])).
cnf(c_0_232, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_209])).
cnf(c_0_233, plain, (X1=k1_seq_2(X2)|k9_real_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1)!=k6_numbers|~v2_comseq_2(X2)|~v1_xreal_0(k9_real_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_210, c_0_211])).
cnf(c_0_234, plain, (k9_real_1(X1,X2)=k6_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_212])).
cnf(c_0_235, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_213])).
cnf(c_0_236, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_214])).
cnf(c_0_237, plain, (k4_xcmplx_0(np__0)=np__0), inference(evalgc,[status(thm)],[c_0_215])).
cnf(c_0_238, plain, (np__0=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_216])])).
cnf(c_0_239, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_217])).
cnf(c_0_240, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_218])).
cnf(c_0_241, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_219])).
cnf(c_0_242, plain, (v2_comseq_2(X1)|r1_xxreal_0(esk9_2(X1,X2),k18_complex1(k9_real_1(k1_seq_1(X1,esk10_3(X1,X2,X3)),X2)))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_220])).
cnf(c_0_243, negated_conjecture, (k4_power(k6_numbers,X1)=k1_seq_1(k1_lp_space(esk2_0,X1),X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers)), inference(rw,[status(thm)],[c_0_221, c_0_222])).
cnf(c_0_244, plain, (m2_subset_1(esk10_3(X1,X2,X3),k1_numbers,k5_numbers)|v2_comseq_2(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_193])).
cnf(c_0_245, plain, (k1_lp_space(esk2_0,X1)=esk2_0|r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_224]), c_0_127])])])).
cnf(c_0_246, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_225, c_0_166])).
fof(c_0_247, plain, ![X44]:(~v1_funct_1(X44)|~v1_funct_2(X44,k5_numbers,k1_numbers)|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|k4_series_1(X44)=k2_seq_2(k3_series_1(X44))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_series_1])])).
fof(c_0_248, plain, ![X53]:(~v1_funct_1(X53)|~v1_funct_2(X53,k5_numbers,k1_numbers)|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|k2_seq_2(X53)=k1_seq_2(X53)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_seq_2])])).
cnf(c_0_249, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_seq_1(X1,esk12_2(X2,X1)),k1_funct_1(X2,k2_nat_1(esk12_2(X2,X1),np__1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(X2,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)), inference(csr,[status(thm)],[c_0_226, c_0_227])).
cnf(c_0_250, negated_conjecture, (v1_partfun1(esk2_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_108]), c_0_106])]), c_0_229])])).
cnf(c_0_251, negated_conjecture, (v4_relat_1(esk2_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_108])])).
cnf(c_0_252, negated_conjecture, (v1_relat_1(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_108])])).
cnf(c_0_253, negated_conjecture, (v3_valued_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_108]), c_0_97])])])).
fof(c_0_254, plain, ![X192, X193]:(~v1_relat_1(X192)|~v1_funct_1(X192)|~v1_valued_0(X192)|v1_xcmplx_0(k1_funct_1(X192,X193))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc61_valued_0])])).
cnf(c_0_255, plain, (X1=k1_seq_2(X2)|k6_xcmplx_0(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1)!=k6_numbers|~v2_comseq_2(X2)|~v1_xreal_0(k6_xcmplx_0(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_234])])).
cnf(c_0_256, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_235])).
cnf(c_0_257, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_236])).
cnf(c_0_258, plain, (k4_xcmplx_0(k6_numbers)=k6_numbers), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_237, c_0_238]), c_0_238])])).
cnf(c_0_259, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)))=k2_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_239, c_0_240])])).
cnf(c_0_260, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_241])).
cnf(c_0_261, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),X2),k18_complex1(k9_real_1(k4_power(k6_numbers,X1),X2)))|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk10_3(k1_lp_space(esk2_0,X1),X2,X3),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_243])])).
cnf(c_0_262, plain, (v2_comseq_2(X1)|m2_subset_1(esk10_3(X1,X2,X3),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_244])).
cnf(c_0_263, plain, (k1_lp_space(esk2_0,X1)=esk2_0|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_245, c_0_246])).
fof(c_0_264, plain, ![X244, X245]:(v1_xboole_0(X244)|v1_xboole_0(X245)|~m1_subset_1(X245,k1_zfmisc_1(X244))|m2_subset_1(esk16_2(X244,X245),X244,X245)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[existence_m2_subset_1])])])])).
cnf(c_0_265, plain, (k4_series_1(X1)=k2_seq_2(k3_series_1(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_247])).
cnf(c_0_266, plain, (k2_seq_2(X1)=k1_seq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_248])).
fof(c_0_267, plain, ![X51]:(((v1_funct_1(k3_series_1(X51))|(~v1_funct_1(X51)|~v1_funct_2(X51,k5_numbers,k1_numbers)|~m1_subset_1(X51,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(v1_funct_2(k3_series_1(X51),k5_numbers,k1_numbers)|(~v1_funct_1(X51)|~v1_funct_2(X51,k5_numbers,k1_numbers)|~m1_subset_1(X51,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))&(m1_subset_1(k3_series_1(X51),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|(~v1_funct_1(X51)|~v1_funct_2(X51,k5_numbers,k1_numbers)|~m1_subset_1(X51,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_series_1])])])).
cnf(c_0_268, negated_conjecture, (k2_series_1(X1)=esk2_0|k2_xcmplx_0(k6_numbers,k1_funct_1(X1,k2_nat_1(esk12_2(X1,esk2_0),np__1)))!=k1_funct_1(esk2_0,k2_nat_1(esk12_2(X1,esk2_0),np__1))|k1_funct_1(esk2_0,k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m2_subset_1(esk12_2(X1,esk2_0),k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_105]), c_0_250]), c_0_251]), c_0_252]), c_0_253]), c_0_107])])])).
cnf(c_0_269, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_254])).
cnf(c_0_270, plain, (X1=k1_seq_2(X2)|k6_xcmplx_0(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1)!=k6_numbers|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(k1_seq_1(X2,esk5_3(X2,X1,X3)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_256])])).
cnf(c_0_271, plain, (k6_xcmplx_0(X1,k6_numbers)=k2_xcmplx_0(X1,k6_numbers)|~v1_xcmplx_0(k6_numbers)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_258])])).
cnf(c_0_272, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)))=k2_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(csr,[status(thm)],[c_0_259, c_0_260])).
cnf(c_0_273, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),X2),k18_complex1(k9_real_1(k4_power(k6_numbers,X1),X2)))|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_261, c_0_262])).
fof(c_0_274, plain, ![X191]:(~v1_xcmplx_0(X191)|k6_xcmplx_0(X191,k6_numbers)=X191), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_arithm])])).
cnf(c_0_275, negated_conjecture, (k4_power(k6_numbers,X1)=k1_seq_1(esk2_0,X2)|r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_263])])).
cnf(c_0_276, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m2_subset_1(esk16_2(X1,X2),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_264])).
cnf(c_0_277, plain, (k2_seq_2(k3_series_1(X1))=k4_series_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_265])).
cnf(c_0_278, plain, (k2_seq_2(X1)=k1_seq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_266])).
cnf(c_0_279, plain, (m1_subset_1(k3_series_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_267])).
cnf(c_0_280, plain, (v1_funct_1(k3_series_1(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_267])).
cnf(c_0_281, plain, (v1_funct_2(k3_series_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_267])).
fof(c_0_282, plain, ![X52]:(~v1_funct_1(X52)|~v1_funct_2(X52,k5_numbers,k1_numbers)|~m1_subset_1(X52,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|k3_series_1(X52)=k2_series_1(X52)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k3_series_1])])).
cnf(c_0_283, negated_conjecture, (k2_series_1(X1)=esk2_0|k1_funct_1(X1,k2_nat_1(esk12_2(X1,esk2_0),np__1))!=k1_funct_1(esk2_0,k2_nat_1(esk12_2(X1,esk2_0),np__1))|k1_funct_1(esk2_0,k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m2_subset_1(esk12_2(X1,esk2_0),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_xcmplx_0(k1_funct_1(X1,k2_nat_1(esk12_2(X1,esk2_0),np__1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_201])])).
cnf(c_0_284, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_269])).
cnf(c_0_285, plain, (X1=k1_seq_2(X2)|k6_xcmplx_0(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1)!=k6_numbers|~v2_comseq_2(X2)|~v1_xreal_0(k1_seq_1(X2,esk5_3(X2,X1,X3)))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_270, c_0_246])).
cnf(c_0_286, plain, (k6_xcmplx_0(X1,k6_numbers)=k2_xcmplx_0(X1,k6_numbers)|~v1_xcmplx_0(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_271, c_0_178])])).
cnf(c_0_287, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k2_xcmplx_0(k6_numbers,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_201]), c_0_258]), c_0_178])])])).
cnf(c_0_288, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),X2),k18_complex1(k9_real_1(k4_power(k6_numbers,X1),X2)))|~v1_xreal_0(X2)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_171])])).
cnf(c_0_289, plain, (k6_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_274])).
cnf(c_0_290, negated_conjecture, (k4_power(k6_numbers,X1)=k1_seq_1(esk2_0,X2)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)), inference(csr,[status(thm)],[c_0_275, c_0_246])).
cnf(c_0_291, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m2_subset_1(esk16_2(X1,X2),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_276])).
cnf(c_0_292, plain, (k1_seq_2(k3_series_1(X1))=k4_series_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k3_series_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(X1))|~m1_subset_1(k3_series_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_277, c_0_278])])).
cnf(c_0_293, plain, (m1_subset_1(k3_series_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_279])).
cnf(c_0_294, plain, (v1_funct_1(k3_series_1(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_280])).
cnf(c_0_295, plain, (v1_funct_2(k3_series_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_281])).
cnf(c_0_296, plain, (k3_series_1(X1)=k2_series_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_282])).
cnf(c_0_297, negated_conjecture, (k2_series_1(X1)=esk2_0|k1_funct_1(X1,k2_nat_1(esk12_2(X1,esk2_0),np__1))!=k1_funct_1(esk2_0,k2_nat_1(esk12_2(X1,esk2_0),np__1))|k1_funct_1(esk2_0,k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m2_subset_1(esk12_2(X1,esk2_0),k1_numbers,k5_numbers)|~v1_funct_1(X1)), inference(csr,[status(thm)],[c_0_283, c_0_284])).
cnf(c_0_298, negated_conjecture, (v1_valued_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_253]), c_0_252])])])).
cnf(c_0_299, plain, (m2_subset_1(esk12_2(X1,X2),k1_numbers,k5_numbers)|X2=k2_series_1(X1)|k1_funct_1(X2,k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v1_relat_1(X2)|~v4_relat_1(X2,k5_numbers)|~v1_funct_1(X2)|~v1_partfun1(X2,k5_numbers)|~v1_valued_0(X2)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_funct_1(X1)|~v1_partfun1(X1,k5_numbers)|~v1_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_161])).
cnf(c_0_300, plain, (k1_seq_2(X1)=k6_numbers|k2_xcmplx_0(k1_seq_1(X1,esk5_3(X1,k6_numbers,X2)),k6_numbers)!=k6_numbers|~v2_comseq_2(X1)|~v1_xreal_0(k1_seq_1(X1,esk5_3(X1,k6_numbers,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xcmplx_0(k1_seq_1(X1,esk5_3(X1,k6_numbers,X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_286]), c_0_127])])])).
cnf(c_0_301, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k2_xcmplx_0(k6_numbers,X1)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_287, c_0_260])).
cnf(c_0_302, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),X2),k18_complex1(k6_xcmplx_0(k4_power(k6_numbers,X1),X2)))|~v1_xreal_0(X2)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k4_power(k6_numbers,X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_234])])).
cnf(c_0_303, plain, (k6_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_289])).
cnf(c_0_304, negated_conjecture, (k4_power(k6_numbers,X2)=k6_numbers|r1_xxreal_0(X2,k6_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_290])])).
cnf(c_0_305, plain, (v1_xboole_0(X1)|m2_subset_1(esk16_2(X2,X1),X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(csr,[status(thm)],[c_0_291, c_0_150])).
cnf(c_0_306, plain, (k1_seq_2(k3_series_1(X1))=k4_series_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_292, c_0_293]), c_0_294]), c_0_295])).
cnf(c_0_307, plain, (k3_series_1(X1)=k2_series_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_296])).
cnf(c_0_308, negated_conjecture, (k2_series_1(esk2_0)=esk2_0|~m2_subset_1(esk12_2(esk2_0,esk2_0),k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_297]), c_0_298]), c_0_250]), c_0_251]), c_0_252]), c_0_107])])])).
cnf(c_0_309, plain, (X2=k2_series_1(X1)|m2_subset_1(esk12_2(X1,X2),k1_numbers,k5_numbers)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_299])).
cnf(c_0_310, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_311, plain, (k1_seq_2(X1)=k6_numbers|k2_xcmplx_0(k1_seq_1(X1,esk5_3(X1,k6_numbers,X2)),k6_numbers)!=k6_numbers|~v2_comseq_2(X1)|~v1_xreal_0(k1_seq_1(X1,esk5_3(X1,k6_numbers,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_300, c_0_166])).
cnf(c_0_312, plain, (k2_xcmplx_0(k6_numbers,k6_numbers)=k6_numbers), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_301, c_0_258]), c_0_258]), c_0_178])])])).
cnf(c_0_313, plain, (m2_subset_1(esk5_3(X1,X2,X3),k1_numbers,k5_numbers)|X2=k1_seq_2(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_140])).
fof(c_0_314, plain, ![X50]:((~v1_series_1(X50)|v2_comseq_2(k3_series_1(X50))|(~v1_funct_1(X50)|~v1_funct_2(X50,k5_numbers,k1_numbers)|~m1_subset_1(X50,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))&(~v2_comseq_2(k3_series_1(X50))|v1_series_1(X50)|(~v1_funct_1(X50)|~v1_funct_2(X50,k5_numbers,k1_numbers)|~m1_subset_1(X50,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_series_1])])])).
cnf(c_0_315, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),k6_numbers),k18_complex1(k4_power(k6_numbers,X1)))|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k4_power(k6_numbers,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(k4_power(k6_numbers,X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_303]), c_0_127])])])).
cnf(c_0_316, negated_conjecture, (k4_power(k6_numbers,X1)=k6_numbers|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_304, c_0_305]), c_0_96])]), c_0_172])])).
cnf(c_0_317, negated_conjecture, (~v1_series_1(k1_lp_space(esk2_0,esk1_0))|k4_power(k4_series_1(k1_lp_space(esk2_0,esk1_0)),k10_real_1(np__1,esk1_0))!=k6_numbers), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_318, plain, (k1_seq_2(k2_series_1(X1))=k4_series_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_306, c_0_307])])).
cnf(c_0_319, negated_conjecture, (k2_series_1(esk2_0)=esk2_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_308, c_0_309]), c_0_298]), c_0_250]), c_0_251]), c_0_252]), c_0_107])])])).
cnf(c_0_320, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers)), inference(evalgc,[status(thm)],[c_0_310])).
cnf(c_0_321, negated_conjecture, (k1_seq_2(esk2_0)=k6_numbers|~v2_comseq_2(esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(esk5_3(esk2_0,k6_numbers,X1),k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_105]), c_0_312]), c_0_127]), c_0_106]), c_0_107]), c_0_108])])])).
cnf(c_0_322, plain, (X2=k1_seq_2(X1)|m2_subset_1(esk5_3(X1,X2,X3),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_313])).
cnf(c_0_323, plain, (v1_series_1(X1)|~v2_comseq_2(k3_series_1(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_314])).
cnf(c_0_324, plain, (v2_comseq_2(X1)|~r1_xxreal_0(esk9_2(X1,X2),k6_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_193])).
cnf(c_0_325, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),k6_numbers),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_316]), c_0_222]), c_0_178]), c_0_198])])])).
fof(c_0_326, plain, ~r1_xxreal_0(np__1,np__0), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r1_r0])).
cnf(c_0_327, negated_conjecture, (k4_power(k4_series_1(k1_lp_space(esk2_0,esk1_0)),k10_real_1(np__1,esk1_0))!=k6_numbers|~v1_series_1(k1_lp_space(esk2_0,esk1_0))), inference(evalgc,[status(thm)],[c_0_317])).
cnf(c_0_328, negated_conjecture, (k4_series_1(esk2_0)=k1_seq_2(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_319]), c_0_106]), c_0_107]), c_0_108])])])).
cnf(c_0_329, negated_conjecture, (v1_xreal_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_320]), c_0_97])])])).
cnf(c_0_330, negated_conjecture, (k1_seq_2(esk2_0)=k6_numbers|~v2_comseq_2(esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_321, c_0_322]), c_0_127]), c_0_106]), c_0_107]), c_0_108])])])).
cnf(c_0_331, plain, (v1_series_1(X1)|~v1_funct_1(X1)|~v2_comseq_2(k3_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_323])).
cnf(c_0_332, plain, (v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~r1_xxreal_0(esk9_2(X1,X2),k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_324])).
cnf(c_0_333, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),k6_numbers),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_325, c_0_246])).
cnf(c_0_334, plain, (m1_subset_1(np__1,k1_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_335, plain, (~r1_xxreal_0(np__1,np__0)), inference(split_conjunct,[status(thm)],[c_0_326])).
cnf(c_0_336, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|k4_power(k1_seq_2(esk2_0),k10_real_1(np__1,esk1_0))!=k6_numbers|~v1_series_1(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_263]), c_0_328]), c_0_329])])])).
cnf(c_0_337, negated_conjecture, (k1_seq_2(esk2_0)=k6_numbers|~v2_comseq_2(esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_305]), c_0_96])]), c_0_172])])).
cnf(c_0_338, plain, (v1_series_1(X1)|~v2_comseq_2(k2_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_331, c_0_307])])).
cnf(c_0_339, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(X1,k6_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_332, c_0_333]), c_0_127])])])).
cnf(c_0_340, plain, (m1_subset_1(np__1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_334])).
cnf(c_0_341, plain, (~r1_xxreal_0(np__1,np__0)), inference(evalgc,[status(thm)],[c_0_335])).
cnf(c_0_342, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|k4_power(k6_numbers,k10_real_1(np__1,esk1_0))!=k6_numbers|~v1_series_1(esk2_0)|~v2_comseq_2(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_337])])).
cnf(c_0_343, negated_conjecture, (v1_series_1(esk2_0)|~v2_comseq_2(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_338, c_0_319]), c_0_106]), c_0_107]), c_0_108])])])).
cnf(c_0_344, plain, (v2_comseq_2(esk2_0)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_339, c_0_263]), c_0_106]), c_0_107]), c_0_108])])])).
cnf(c_0_345, plain, (v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_340]), c_0_97])])])).
cnf(c_0_346, plain, (~r1_xxreal_0(np__1,k6_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_341, c_0_238])])).
fof(c_0_347, plain, (v1_xboole_0(esk14_0)&v1_xxreal_0(esk14_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc4_xxreal_0])])).
cnf(c_0_348, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|k4_power(k6_numbers,k10_real_1(np__1,esk1_0))!=k6_numbers|~v2_comseq_2(esk2_0)), inference(csr,[status(thm)],[c_0_342, c_0_343])).
cnf(c_0_349, plain, (v2_comseq_2(esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_345]), c_0_346])])).
fof(c_0_350, plain, ![X72, X73]:(~m1_subset_1(X72,k1_numbers)|~v1_xreal_0(X73)|k10_real_1(X72,X73)=k7_xcmplx_0(X72,X73)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k10_real_1])])).
fof(c_0_351, plain, ![X122, X123]:(~v2_membered(X122)|(~m1_subset_1(X123,X122)|v1_xxreal_0(X123))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc13_membered])])])).
fof(c_0_352, plain, ![X182, X183]:(~v1_xreal_0(X182)|(~v1_xreal_0(X183)|(~r1_xxreal_0(X182,X183)|v2_xxreal_0(X183)|~v2_xxreal_0(X182)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_real])])])])).
cnf(c_0_353, plain, (v1_xboole_0(esk14_0)), inference(split_conjunct,[status(thm)],[c_0_347])).
fof(c_0_354, plain, ![X157]:(~v7_ordinal1(X157)|v1_xxreal_0(X157)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_xxreal_0])])).
fof(c_0_355, plain, ![X181]:(~v1_xboole_0(X181)|v7_ordinal1(X181)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc7_ordinal1])])).
fof(c_0_356, plain, ![X40, X41]:(~v1_xreal_0(X40)|(~v1_xreal_0(X41)|(r1_xxreal_0(X40,k6_numbers)|r1_xxreal_0(X41,k6_numbers)|~r1_xxreal_0(k7_xcmplx_0(X40,X41),k6_numbers)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t139_xreal_1])])])])).
cnf(c_0_357, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|k4_power(k6_numbers,k10_real_1(np__1,esk1_0))!=k6_numbers), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_348, c_0_349])])])).
cnf(c_0_358, plain, (k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_350])).
cnf(c_0_359, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_351])).
fof(c_0_360, plain, ![X128]:(~v3_membered(X128)|v2_membered(X128)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_membered])])).
cnf(c_0_361, plain, (v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~v2_xxreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_352])).
cnf(c_0_362, negated_conjecture, (r1_xxreal_0(np__1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_363, plain, (v2_xxreal_0(np__1)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
fof(c_0_364, plain, ![X130]:((v1_xcmplx_0(k4_xcmplx_0(X130))|(v3_xxreal_0(X130)|~v1_xreal_0(X130)))&(~v2_xxreal_0(k4_xcmplx_0(X130))|(v3_xxreal_0(X130)|~v1_xreal_0(X130)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc16_xreal_0])])])])).
fof(c_0_365, plain, ![X190]:(((~v1_xboole_0(X190)|(~v1_xxreal_0(X190)|~v3_xxreal_0(X190)))&(v1_xxreal_0(X190)|(~v1_xxreal_0(X190)|~v3_xxreal_0(X190))))&(~v2_xxreal_0(X190)|(~v1_xxreal_0(X190)|~v3_xxreal_0(X190)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_xxreal_0])])])])).
cnf(c_0_366, plain, (v1_xxreal_0(esk14_0)), inference(split_conjunct,[status(thm)],[c_0_347])).
cnf(c_0_367, plain, (v1_xboole_0(esk14_0)), inference(evalgc,[status(thm)],[c_0_353])).
cnf(c_0_368, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
fof(c_0_369, plain, ![X124]:(((~v1_xboole_0(X124)|(~v1_xxreal_0(X124)|~v2_xxreal_0(X124)))&(v1_xxreal_0(X124)|(~v1_xxreal_0(X124)|~v2_xxreal_0(X124))))&(~v3_xxreal_0(X124)|(~v1_xxreal_0(X124)|~v2_xxreal_0(X124)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc3_xxreal_0])])])])).
cnf(c_0_370, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_354])).
cnf(c_0_371, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_355])).
cnf(c_0_372, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(k7_xcmplx_0(X1,X2),k6_numbers)), inference(split_conjunct,[status(thm)],[c_0_356])).
cnf(c_0_373, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|r1_xxreal_0(k10_real_1(np__1,esk1_0),k6_numbers)|~v1_xreal_0(k10_real_1(np__1,esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_357, c_0_316])])).
cnf(c_0_374, plain, (k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_358])).
fof(c_0_375, plain, ![X159, X160]:(~v1_xreal_0(X159)|~v1_xreal_0(X160)|v1_xreal_0(k7_xcmplx_0(X159,X160))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc8_xreal_0])])).
cnf(c_0_376, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_359])).
cnf(c_0_377, plain, (v2_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_360])).
cnf(c_0_378, plain, (v2_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_361])).
cnf(c_0_379, negated_conjecture, (r1_xxreal_0(np__1,esk1_0)), inference(evalgc,[status(thm)],[c_0_362])).
cnf(c_0_380, plain, (v2_xxreal_0(np__1)), inference(evalgc,[status(thm)],[c_0_363])).
cnf(c_0_381, plain, (v3_xxreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_364])).
cnf(c_0_382, plain, (~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v3_xxreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_365])).
cnf(c_0_383, plain, (v1_xxreal_0(esk14_0)), inference(evalgc,[status(thm)],[c_0_366])).
cnf(c_0_384, plain, (esk14_0=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_367])])).
cnf(c_0_385, plain, (v1_xboole_0(k1_xboole_0)), inference(evalgc,[status(thm)],[c_0_368])).
cnf(c_0_386, plain, (~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v2_xxreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_369])).
cnf(c_0_387, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_370])).
cnf(c_0_388, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_371])).
fof(c_0_389, plain, ![X235, X236]:(~v1_xreal_0(X235)|(~v1_xreal_0(X236)|(~r1_xxreal_0(X235,X236)|(v1_xboole_0(X235)|v2_xxreal_0(X236)|v3_xxreal_0(X235))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_real])])])).
cnf(c_0_390, plain, (r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k7_xcmplx_0(X1,X2),k6_numbers)), inference(evalgc,[status(thm)],[c_0_372])).
cnf(c_0_391, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(np__1,esk1_0),k6_numbers)|r1_xxreal_0(esk1_0,k6_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_374]), c_0_329]), c_0_340])])])).
cnf(c_0_392, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_375])).
cnf(c_0_393, plain, (~v3_xxreal_0(X1)|~v1_xxreal_0(X1)|~v2_xxreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_369])).
cnf(c_0_394, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_376, c_0_320])])).
cnf(c_0_395, plain, (v2_membered(X1)|~v3_membered(X1)), inference(evalgc,[status(thm)],[c_0_377])).
cnf(c_0_396, negated_conjecture, (v2_xxreal_0(esk1_0)|~v1_xreal_0(esk1_0)|~v1_xreal_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_379]), c_0_380])])])).
cnf(c_0_397, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X1))), inference(evalgc,[status(thm)],[c_0_381])).
cnf(c_0_398, plain, (~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v3_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_382])).
cnf(c_0_399, plain, (v1_xxreal_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_383, c_0_384])])).
cnf(c_0_400, plain, (v1_xboole_0(k6_numbers)), inference(rw,[status(thm)],[c_0_385, c_0_80])).
cnf(c_0_401, plain, (~v1_xboole_0(X1)|~v2_xxreal_0(X1)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_386])).
cnf(c_0_402, plain, (v1_xxreal_0(X1)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_388])])).
cnf(c_0_403, plain, (v1_xboole_0(X1)|v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_389])).
cnf(c_0_404, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_390, c_0_391]), c_0_329]), c_0_345])]), c_0_346])])).
cnf(c_0_405, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_392])).
cnf(c_0_406, plain, (~v2_xxreal_0(X1)|~v1_xxreal_0(X1)|~v3_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_393])).
cnf(c_0_407, negated_conjecture, (v1_xxreal_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_394, c_0_395]), c_0_97])])])).
cnf(c_0_408, negated_conjecture, (v2_xxreal_0(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_396, c_0_329]), c_0_345])])).
cnf(c_0_409, plain, (v3_xxreal_0(k6_numbers)|~v2_xxreal_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_397, c_0_258]), c_0_127])])])).
cnf(c_0_410, plain, (~v3_xxreal_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_398, c_0_399]), c_0_400])])])).
cnf(c_0_411, plain, (~v2_xxreal_0(X1)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_402])])).
cnf(c_0_412, plain, (v1_xboole_0(X1)|v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_403])).
cnf(c_0_413, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_404, c_0_405]), c_0_329]), c_0_345])])])).
cnf(c_0_414, negated_conjecture, (~v3_xxreal_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_407]), c_0_408])])])).
cnf(c_0_415, plain, (~v2_xxreal_0(k6_numbers)), inference(sr,[status(thm)],[c_0_409, c_0_410])).
cnf(c_0_416, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_411, c_0_408])])).
cnf(c_0_417, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_412, c_0_413]), c_0_127]), c_0_329])]), c_0_414]), c_0_415]), c_0_416]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 418
# Proof object clause steps            : 289
# Proof object formula steps           : 129
# Proof object conjectures             : 67
# Proof object clause conjectures      : 64
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 89
# Proof object initial formulas used   : 67
# Proof object generating inferences   : 77
# Proof object simplifying inferences  : 177
# Training examples: 170 positive, 1176 negative
# Training: Positive examples begin
cnf(c_0_80, plain, (k1_xboole_0=k6_numbers)).# trainpos
cnf(c_0_81, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_237, plain, (k4_xcmplx_0(np__0)=np__0)).# trainpos
cnf(c_0_107, negated_conjecture, (v1_funct_1(esk2_0))).# trainpos
cnf(c_0_97, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_385, plain, (v1_xboole_0(k1_xboole_0))).# trainpos
cnf(c_0_216, plain, (v1_xboole_0(np__0))).# trainpos
cnf(c_0_139, plain, (v1_xboole_0(esk13_0))).# trainpos
cnf(c_0_367, plain, (v1_xboole_0(esk14_0))).# trainpos
cnf(c_0_380, plain, (v2_xxreal_0(np__1))).# trainpos
cnf(c_0_327, negated_conjecture, (k4_power(k4_series_1(k1_lp_space(esk2_0,esk1_0)),k10_real_1(np__1,esk1_0))!=k6_numbers|~v1_series_1(k1_lp_space(esk2_0,esk1_0)))).# trainpos
cnf(c_0_320, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers))).# trainpos
cnf(c_0_159, plain, (v1_xcmplx_0(esk13_0))).# trainpos
cnf(c_0_106, negated_conjecture, (v1_funct_2(esk2_0,k5_numbers,k1_numbers))).# trainpos
cnf(c_0_229, plain, (~v1_xboole_0(k1_numbers))).# trainpos
cnf(c_0_340, plain, (m1_subset_1(np__1,k1_numbers))).# trainpos
cnf(c_0_383, plain, (v1_xxreal_0(esk14_0))).# trainpos
cnf(c_0_152, plain, (~v1_xboole_0(k4_ordinal1))).# trainpos
cnf(c_0_79, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).# trainpos
cnf(c_0_379, negated_conjecture, (r1_xxreal_0(np__1,esk1_0))).# trainpos
cnf(c_0_341, plain, (~r1_xxreal_0(np__1,np__0))).# trainpos
cnf(c_0_108, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_118, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_96, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_171, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_388, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_166, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_387, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_395, plain, (v2_membered(X1)|~v3_membered(X1))).# trainpos
cnf(c_0_260, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_105, negated_conjecture, (k1_seq_1(esk2_0,X1)=k6_numbers|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_211, plain, (X2=k1_seq_2(X1)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~r1_xxreal_0(esk4_2(X1,X2),k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_88, plain, (k18_complex1(X1)=k6_numbers|X1!=k6_numbers|~v1_xreal_0(X1))).# trainpos
cnf(c_0_401, plain, (~v1_xboole_0(X1)|~v2_xxreal_0(X1)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_240, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_157, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_303, plain, (k6_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_144, plain, (k18_complex1(k18_complex1(X1))=k18_complex1(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_227, plain, (v1_valued_0(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_398, plain, (~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v3_xxreal_0(X1))).# trainpos
cnf(c_0_406, plain, (~v2_xxreal_0(X1)|~v1_xxreal_0(X1)|~v3_xxreal_0(X1))).# trainpos
cnf(c_0_202, plain, (X1=k2_series_1(X2)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|k2_xcmplx_0(k1_funct_1(X1,esk12_2(X2,X1)),k1_funct_1(X2,k2_nat_1(esk12_2(X2,X1),np__1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(X2,X1),np__1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers))).# trainpos
cnf(c_0_397, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X1)))).# trainpos
cnf(c_0_93, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_376, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_231, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_95, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_150, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_331, plain, (v1_series_1(X1)|~v1_funct_1(X1)|~v2_comseq_2(k3_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_294, plain, (v1_funct_1(k3_series_1(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_87, plain, (X1=k1_lp_space(X2,X3)|k1_seq_1(X1,esk6_3(X2,X3,X1))!=k4_power(k18_complex1(k1_seq_1(X2,esk6_3(X2,X3,X1))),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_174, plain, (epred1_3(X1,X2,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_158, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_405, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_256, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_230, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_257, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_374, plain, (k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_295, plain, (v1_funct_2(k3_series_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_378, plain, (v2_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_197, plain, (k4_power(X1,X2)=k3_power(X1,X2)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_293, plain, (m1_subset_1(k3_series_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_234, plain, (k9_real_1(X1,X2)=k6_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_176, plain, (k7_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_175, plain, (m1_subset_1(k7_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_232, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_277, plain, (k2_seq_2(k3_series_1(X1))=k4_series_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_278, plain, (k2_seq_2(X1)=k1_seq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_203, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_412, plain, (v1_xboole_0(X1)|v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_284, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).# trainpos
cnf(c_0_239, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_173, plain, (X1=k3_power(X2,X3)|r1_xxreal_0(X3,k6_numbers)|X2!=k6_numbers|X1!=k6_numbers|~epred1_3(X2,X3,X1))).# trainpos
cnf(c_0_142, plain, (v1_funct_1(k1_lp_space(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_143, plain, (v1_funct_2(k1_lp_space(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_307, plain, (k3_series_1(X1)=k2_series_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_390, plain, (r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(k7_xcmplx_0(X1,X2),k6_numbers))).# trainpos
cnf(c_0_149, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_291, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m2_subset_1(esk16_2(X1,X2),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_228, plain, (v1_xboole_0(X1)|v1_partfun1(X2,X3)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_332, plain, (v2_comseq_2(X1)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~r1_xxreal_0(esk9_2(X1,X2),k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_121, plain, (m1_subset_1(k1_lp_space(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_262, plain, (v2_comseq_2(X1)|m2_subset_1(esk10_3(X1,X2,X3),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_168, plain, (X3=k1_lp_space(X1,X2)|m2_subset_1(esk6_3(X1,X2,X3),k1_numbers,k5_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_120, plain, (k1_seq_1(X1,X4)=k4_power(k18_complex1(k1_seq_1(X2,X4)),X3)|X1!=k1_lp_space(X2,X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_309, plain, (X2=k2_series_1(X1)|m2_subset_1(esk12_2(X1,X2),k1_numbers,k5_numbers)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers))).# trainpos
cnf(c_0_322, plain, (X2=k1_seq_2(X1)|m2_subset_1(esk5_3(X1,X2,X3),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_242, plain, (v2_comseq_2(X1)|r1_xxreal_0(esk9_2(X1,X2),k18_complex1(k9_real_1(k1_seq_1(X1,esk10_3(X1,X2,X3)),X2)))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_187, plain, (X2=k1_seq_2(X1)|r1_xxreal_0(esk4_2(X1,X2),k18_complex1(k9_real_1(k1_seq_1(X1,esk5_3(X1,X2,X3)),X2)))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_329, negated_conjecture, (v1_xreal_0(esk1_0))).# trainpos
cnf(c_0_252, negated_conjecture, (v1_relat_1(esk2_0))).# trainpos
cnf(c_0_238, plain, (np__0=k6_numbers)).# trainpos
cnf(c_0_384, plain, (esk14_0=k6_numbers)).# trainpos
cnf(c_0_160, plain, (esk13_0=k6_numbers)).# trainpos
cnf(c_0_251, negated_conjecture, (v4_relat_1(esk2_0,k5_numbers))).# trainpos
cnf(c_0_345, plain, (v1_xreal_0(np__1))).# trainpos
cnf(c_0_402, plain, (v1_xxreal_0(X1)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_110, plain, (v3_membered(k5_numbers))).# trainpos
cnf(c_0_292, plain, (k1_seq_2(k3_series_1(X1))=k4_series_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k3_series_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k3_series_1(X1))|~m1_subset_1(k3_series_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_104, plain, (X1=k1_lp_space(X2,X3)|k1_seq_1(X1,esk6_3(X2,X3,X1))!=k4_power(k6_numbers,X3)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|k1_seq_1(X2,esk6_3(X2,X3,X1))!=k6_numbers|~v1_xreal_0(k1_seq_1(X2,esk6_3(X2,X3,X1))))).# trainpos
cnf(c_0_253, negated_conjecture, (v3_valued_0(esk2_0))).# trainpos
cnf(c_0_396, negated_conjecture, (v2_xxreal_0(esk1_0)|~v1_xreal_0(esk1_0)|~v1_xreal_0(np__1))).# trainpos
cnf(c_0_109, plain, (v1_xreal_0(k6_numbers)|~v3_membered(k5_numbers))).# trainpos
cnf(c_0_338, plain, (v1_series_1(X1)|~v2_comseq_2(k2_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_298, negated_conjecture, (v1_valued_0(esk2_0))).# trainpos
cnf(c_0_198, plain, (m1_subset_1(k6_numbers,k1_numbers))).# trainpos
cnf(c_0_258, plain, (k4_xcmplx_0(k6_numbers)=k6_numbers)).# trainpos
cnf(c_0_399, plain, (v1_xxreal_0(k6_numbers))).# trainpos
cnf(c_0_346, plain, (~r1_xxreal_0(np__1,k6_numbers))).# trainpos
cnf(c_0_178, plain, (v1_xcmplx_0(k6_numbers))).# trainpos
cnf(c_0_141, plain, (k1_seq_1(k1_lp_space(X1,X2),X3)=k4_power(k18_complex1(k1_seq_1(X4,X3)),X5)|k1_lp_space(X1,X2)!=k1_lp_space(X4,X5)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X4,k5_numbers,k1_numbers)|~v1_funct_2(k1_lp_space(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(X4)|~v1_funct_1(k1_lp_space(X1,X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X5,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_411, plain, (~v2_xxreal_0(X1)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_250, negated_conjecture, (v1_partfun1(esk2_0,k5_numbers))).# trainpos
cnf(c_0_318, plain, (k1_seq_2(k2_series_1(X1))=k4_series_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_165, plain, (k18_complex1(k6_numbers)=k6_numbers|~v1_xcmplx_0(X1)|X1!=k6_numbers|~v1_xreal_0(X1))).# trainpos
cnf(c_0_210, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk4_2(X2,X1),k6_numbers)|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|k9_real_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1)!=k6_numbers|~v1_xreal_0(k9_real_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1)))).# trainpos
cnf(c_0_177, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k6_numbers))).# trainpos
cnf(c_0_394, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k1_numbers))).# trainpos
cnf(c_0_200, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_199, plain, (X1=k3_power(X2,X3)|r1_xxreal_0(X3,k6_numbers)|X2!=k6_numbers|X1!=k6_numbers|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_226, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_seq_1(X1,esk12_2(X2,X1)),k1_funct_1(X2,k2_nat_1(esk12_2(X2,X1),np__1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(X2,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_407, negated_conjecture, (v1_xxreal_0(esk1_0))).# trainpos
cnf(c_0_410, plain, (~v3_xxreal_0(k6_numbers))).# trainpos
cnf(c_0_416, negated_conjecture, (~v1_xboole_0(esk1_0))).# trainpos
cnf(c_0_194, plain, (k4_power(k18_complex1(k1_seq_1(X1,X3)),X2)=k1_seq_1(k1_lp_space(X1,X2),X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_222, plain, (k18_complex1(k6_numbers)=k6_numbers)).# trainpos
cnf(c_0_224, plain, (k3_power(X1,X2)=k6_numbers|r1_xxreal_0(X2,k6_numbers)|X1!=k6_numbers|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_225, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_409, plain, (v3_xxreal_0(k6_numbers)|~v2_xxreal_0(k6_numbers))).# trainpos
cnf(c_0_268, negated_conjecture, (k2_series_1(X1)=esk2_0|k2_xcmplx_0(k6_numbers,k1_funct_1(X1,k2_nat_1(esk12_2(X1,esk2_0),np__1)))!=k1_funct_1(esk2_0,k2_nat_1(esk12_2(X1,esk2_0),np__1))|k1_funct_1(esk2_0,k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m2_subset_1(esk12_2(X1,esk2_0),k1_numbers,k5_numbers))).# trainpos
cnf(c_0_271, plain, (k6_xcmplx_0(X1,k6_numbers)=k2_xcmplx_0(X1,k6_numbers)|~v1_xcmplx_0(k6_numbers)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_126, negated_conjecture, (X1=k1_lp_space(esk2_0,X2)|k1_seq_1(X1,esk6_3(esk2_0,X2,X1))!=k4_power(k6_numbers,X2)|~v1_xreal_0(k6_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(esk6_3(esk2_0,X2,X1),k1_numbers,k5_numbers))).# trainpos
cnf(c_0_255, plain, (X1=k1_seq_2(X2)|k6_xcmplx_0(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1)!=k6_numbers|~v2_comseq_2(X2)|~v1_xreal_0(k6_xcmplx_0(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),k1_numbers))).# trainpos
cnf(c_0_414, negated_conjecture, (~v3_xxreal_0(esk1_0))).# trainpos
cnf(c_0_283, negated_conjecture, (k2_series_1(X1)=esk2_0|k1_funct_1(X1,k2_nat_1(esk12_2(X1,esk2_0),np__1))!=k1_funct_1(esk2_0,k2_nat_1(esk12_2(X1,esk2_0),np__1))|k1_funct_1(esk2_0,k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m2_subset_1(esk12_2(X1,esk2_0),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_xcmplx_0(k1_funct_1(X1,k2_nat_1(esk12_2(X1,esk2_0),np__1))))).# trainpos
cnf(c_0_270, plain, (X1=k1_seq_2(X2)|k6_xcmplx_0(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1)!=k6_numbers|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(k1_seq_1(X2,esk5_3(X2,X1,X3))))).# trainpos
cnf(c_0_167, negated_conjecture, (k1_lp_space(esk2_0,X1)=esk2_0|k4_power(k6_numbers,X1)!=k6_numbers|~m2_subset_1(esk6_3(esk2_0,X1,esk2_0),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_221, negated_conjecture, (k4_power(k18_complex1(k6_numbers),X2)=k1_seq_1(k1_lp_space(esk2_0,X2),X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_196, negated_conjecture, (k1_lp_space(esk2_0,X1)=esk2_0|k4_power(k6_numbers,X1)!=k6_numbers|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_308, negated_conjecture, (k2_series_1(esk2_0)=esk2_0|~m2_subset_1(esk12_2(esk2_0,esk2_0),k1_numbers,k5_numbers))).# trainpos
cnf(c_0_319, negated_conjecture, (k2_series_1(esk2_0)=esk2_0)).# trainpos
cnf(c_0_223, negated_conjecture, (k1_lp_space(esk2_0,X1)=esk2_0|k3_power(k6_numbers,X1)!=k6_numbers|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_343, negated_conjecture, (v1_series_1(esk2_0)|~v2_comseq_2(esk2_0))).# trainpos
cnf(c_0_328, negated_conjecture, (k4_series_1(esk2_0)=k1_seq_2(esk2_0))).# trainpos
cnf(c_0_245, plain, (k1_lp_space(esk2_0,X1)=esk2_0|r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_275, negated_conjecture, (k4_power(k6_numbers,X1)=k1_seq_1(esk2_0,X2)|r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_304, negated_conjecture, (k4_power(k6_numbers,X2)=k6_numbers|r1_xxreal_0(X2,k6_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_316, negated_conjecture, (k4_power(k6_numbers,X1)=k6_numbers|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_336, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|k4_power(k1_seq_2(esk2_0),k10_real_1(np__1,esk1_0))!=k6_numbers|~v1_series_1(esk2_0))).# trainpos
cnf(c_0_259, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)))=k2_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_287, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k2_xcmplx_0(k6_numbers,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).# trainpos
cnf(c_0_312, plain, (k2_xcmplx_0(k6_numbers,k6_numbers)=k6_numbers)).# trainpos
cnf(c_0_300, plain, (k1_seq_2(X1)=k6_numbers|k2_xcmplx_0(k1_seq_1(X1,esk5_3(X1,k6_numbers,X2)),k6_numbers)!=k6_numbers|~v2_comseq_2(X1)|~v1_xreal_0(k1_seq_1(X1,esk5_3(X1,k6_numbers,X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xcmplx_0(k1_seq_1(X1,esk5_3(X1,k6_numbers,X2))))).# trainpos
cnf(c_0_321, negated_conjecture, (k1_seq_2(esk2_0)=k6_numbers|~v2_comseq_2(esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m2_subset_1(esk5_3(esk2_0,k6_numbers,X1),k1_numbers,k5_numbers))).# trainpos
cnf(c_0_261, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),X2),k18_complex1(k9_real_1(k4_power(k6_numbers,X1),X2)))|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk10_3(k1_lp_space(esk2_0,X1),X2,X3),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_288, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),X2),k18_complex1(k9_real_1(k4_power(k6_numbers,X1),X2)))|~v1_xreal_0(X2)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_330, negated_conjecture, (k1_seq_2(esk2_0)=k6_numbers|~v2_comseq_2(esk2_0)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_337, negated_conjecture, (k1_seq_2(esk2_0)=k6_numbers|~v2_comseq_2(esk2_0))).# trainpos
cnf(c_0_302, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),X2),k18_complex1(k6_xcmplx_0(k4_power(k6_numbers,X1),X2)))|~v1_xreal_0(X2)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k4_power(k6_numbers,X1),k1_numbers))).# trainpos
cnf(c_0_315, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),k6_numbers),k18_complex1(k4_power(k6_numbers,X1)))|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k4_power(k6_numbers,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(k4_power(k6_numbers,X1)))).# trainpos
cnf(c_0_325, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),k6_numbers),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_339, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(X1,k6_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_342, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|k4_power(k6_numbers,k10_real_1(np__1,esk1_0))!=k6_numbers|~v1_series_1(esk2_0)|~v2_comseq_2(esk2_0))).# trainpos
cnf(c_0_344, plain, (v2_comseq_2(esk2_0)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_349, plain, (v2_comseq_2(esk2_0))).# trainpos
cnf(c_0_357, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|k4_power(k6_numbers,k10_real_1(np__1,esk1_0))!=k6_numbers)).# trainpos
cnf(c_0_373, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|r1_xxreal_0(k10_real_1(np__1,esk1_0),k6_numbers)|~v1_xreal_0(k10_real_1(np__1,esk1_0)))).# trainpos
cnf(c_0_391, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(np__1,esk1_0),k6_numbers)|r1_xxreal_0(esk1_0,k6_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,esk1_0)))).# trainpos
cnf(c_0_404, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,esk1_0)))).# trainpos
cnf(c_0_413, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_69, plain, (v1_xreal_0(esk13_0))).#trainneg
cnf(i_0_81, plain, (v6_membered(k4_ordinal1))).#trainneg
cnf(i_0_28, plain, (X1=k6_numbers|k18_complex1(X1)!=k6_numbers|~v1_xreal_0(X1))).#trainneg
cnf(i_0_70, plain, (v1_xxreal_0(esk13_0))).#trainneg
cnf(i_0_207, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_64, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_79, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_118, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_197, plain, (k2_xcmplx_0(np__1,np__0)=np__1)).#trainneg
cnf(i_0_170, plain, (k2_xcmplx_0(np__0,np__1)=np__1)).#trainneg
cnf(i_0_186, plain, (k6_xcmplx_0(np__1,np__1)=np__0)).#trainneg
cnf(i_0_112, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_113, plain, (m1_subset_1(np__0,k5_numbers))).#trainneg
cnf(i_0_182, plain, (r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_162, plain, (r1_xxreal_0(np__0,np__1))).#trainneg
cnf(i_0_179, plain, (r1_xxreal_0(np__0,np__0))).#trainneg
cnf(i_0_189, plain, (k6_xcmplx_0(np__0,np__1)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_199, plain, (~r1_xxreal_0(np__0,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_65, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_171, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_131, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_169, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_150, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_215, plain, (m1_subset_1(esk15_1(X1),X1))).#trainneg
cnf(i_0_114, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_85, plain, (v1_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_84, plain, (k16_complex1(X1)=k18_complex1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_132, plain, (v1_xreal_0(k16_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_98, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_99, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_172, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_105, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_82, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_163, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_213, plain, (v2_valued_0(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_183, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_188, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_222, plain, (k7_xcmplx_0(k6_numbers,X1)=k6_numbers|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_155, plain, (k7_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_104, plain, (m1_subset_1(k18_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_201, plain, (v3_valued_0(k2_zfmisc_1(X2,X1))|~v3_membered(X1))).#trainneg
cnf(i_0_216, plain, (v1_valued_0(k2_zfmisc_1(X2,X1))|~v1_membered(X1))).#trainneg
cnf(i_0_102, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_97, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_95, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_184, plain, (v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v3_xxreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_202, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_103, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_8, plain, (k16_complex1(X1)=k4_xcmplx_0(X1)|r1_xxreal_0(k6_numbers,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_152, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_9, plain, (k16_complex1(X1)=X1|~v1_xreal_0(X1)|~r1_xxreal_0(k6_numbers,X1))).#trainneg
cnf(i_0_60, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_101, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_198, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_136, plain, (v1_membered(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_178, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_156, plain, (v1_xreal_0(k3_power(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_116, plain, (v1_xreal_0(k1_seq_2(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_203, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_206, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_90, plain, (m1_subset_1(k4_series_1(X1),k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_91, plain, (m1_subset_1(k2_seq_2(X1),k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_78, plain, (m1_subset_1(k10_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_160, plain, (v1_xreal_0(k3_power(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_164, plain, (v1_xreal_0(k9_prepower(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_165, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_161, plain, (v7_ordinal1(k3_power(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_87, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_41, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_62, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_129, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_134, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_15, plain, (r1_xxreal_0(X4,k6_numbers)|X3!=k1_seq_2(X2)|~v1_funct_1(X2)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v2_comseq_2(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r1_xxreal_0(esk3_3(X2,X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r1_xxreal_0(X4,k18_complex1(k9_real_1(k1_seq_1(X2,X1),X3))))).#trainneg
cnf(i_0_143, plain, (v3_valued_0(X2)|~v3_valued_0(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_74, plain, (m1_subset_1(k4_power(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_44, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_218, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X2)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_128, plain, (k7_real_1(X1,X2)=k7_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_92, plain, (m1_subset_1(k9_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_180, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_181, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_157, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_148, plain, (v2_valued_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_190, plain, (v1_valued_0(X2)|~v1_relat_1(X1)|~v1_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_210, plain, (v2_valued_0(X2)|~v1_relat_1(X1)|~v2_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_48, plain, (v2_comseq_2(k3_series_1(X1))|~v1_funct_1(X1)|~v1_series_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_25, plain, (r1_xxreal_0(X3,k6_numbers)|~v1_funct_1(X2)|~v1_xreal_0(X3)|~v2_comseq_2(X2)|~r1_xxreal_0(esk8_2(X2,X3),X1)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r1_xxreal_0(X3,k18_complex1(k9_real_1(k1_seq_1(X2,X1),esk7_1(X2)))))).#trainneg
cnf(i_0_228, plain, (X1=k6_numbers|r1_xxreal_0(X3,k6_numbers)|X2!=k6_numbers|X1!=k3_power(X2,X3)|~epred1_3(X2,X3,X1))).#trainneg
cnf(i_0_27, plain, (v1_xreal_0(esk7_1(X1))|~v1_funct_1(X1)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_140, plain, (v1_funct_1(k2_series_1(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_89, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_212, plain, (v1_xboole_0(X2)|v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_135, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_200, plain, (v1_xxreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v2_valued_0(X1))).#trainneg
cnf(i_0_209, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),k4_xcmplx_0(X2))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_115, plain, (m2_subset_1(k2_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_220, plain, (v3_xxreal_0(X2)|v2_xxreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_221, plain, (v3_xxreal_0(X2)|v2_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_229, plain, (X1=k3_power(X2,X3)|r1_xxreal_0(X2,k6_numbers)|X1!=k9_prepower(X2,X3)|~epred1_3(X2,X3,X1))).#trainneg
cnf(i_0_230, plain, (X1=k9_prepower(X2,X3)|r1_xxreal_0(X2,k6_numbers)|X1!=k3_power(X2,X3)|~epred1_3(X2,X3,X1))).#trainneg
cnf(i_0_225, plain, (esk17_3(X1,X2,X3)=X2|X1!=k3_power(X3,X2)|~v1_int_1(X2)|~epred1_3(X3,X2,X1))).#trainneg
cnf(i_0_194, plain, (v1_funct_1(k2_series_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_141, plain, (v4_relat_1(k2_series_1(X1),k5_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_137, plain, (v3_valued_0(k2_series_1(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_142, plain, (v1_relat_1(k2_series_1(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_45, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_196, plain, (v1_relat_1(k2_series_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_138, plain, (v1_valued_0(k2_series_1(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_226, plain, (v1_int_1(esk17_3(X1,X2,X3))|X1!=k3_power(X3,X2)|~v1_int_1(X2)|~epred1_3(X3,X2,X1))).#trainneg
cnf(i_0_187, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_224, plain, (k4_prepower(X2,esk17_3(X1,X3,X2))=X1|X1!=k3_power(X2,X3)|~v1_int_1(X3)|~epred1_3(X2,X3,X1))).#trainneg
cnf(i_0_192, plain, (v1_valued_0(k2_series_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_223, plain, (X3=k3_power(X4,X2)|X1!=X2|X3!=k4_prepower(X4,X1)|~v1_int_1(X2)|~v1_int_1(X1)|~epred1_3(X4,X2,X3))).#trainneg
cnf(i_0_195, plain, (v4_relat_1(k2_series_1(X1),k5_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_139, plain, (v1_partfun1(k2_series_1(X1),k5_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_193, plain, (v1_partfun1(k2_series_1(X1),k5_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_125, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|~v1_funct_1(X1)|~v1_xboole_0(X1)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_24, plain, (v2_comseq_2(X1)|v1_xreal_0(esk9_2(X1,X2))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_26, plain, (r1_xxreal_0(X2,k6_numbers)|m2_subset_1(esk8_2(X1,X2),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_59, plain, (k1_funct_1(X1,k6_numbers)=k1_funct_1(X2,k6_numbers)|X1!=k2_series_1(X2)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_14, plain, (X2=k1_seq_2(X1)|v1_xreal_0(esk4_2(X1,X2))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_16, plain, (r1_xxreal_0(X3,k6_numbers)|m2_subset_1(esk3_3(X1,X2,X3),k1_numbers,k5_numbers)|X2!=k1_seq_2(X1)|~v1_funct_1(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_21, plain, (v2_comseq_2(X2)|r1_xxreal_0(X1,esk10_3(X2,X3,X1))|~v1_funct_1(X2)|~v1_xreal_0(X3)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_58, plain, (k2_xcmplx_0(k1_funct_1(X2,X1),k1_funct_1(X3,k2_nat_1(X1,np__1)))=k1_funct_1(X2,k2_nat_1(X1,np__1))|X2!=k2_series_1(X3)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_valued_0(X3)|~v1_valued_0(X2)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v1_partfun1(X3,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_11, plain, (X3=k1_seq_2(X2)|r1_xxreal_0(X1,esk5_3(X2,X3,X1))|~v1_funct_1(X2)|~v1_xreal_0(X3)|~v2_comseq_2(X2)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_32, plain, (k1_seq_1(k3_series_1(k1_lp_space(X2,X1)),X3)=k1_seq_1(k1_lp_space(X2,X1),X3)|m2_subset_1(esk11_2(X1,X2),k1_numbers,k5_numbers)|k1_seq_1(k3_series_1(k1_lp_space(X2,X1)),k6_numbers)!=k1_seq_1(k1_lp_space(X2,X1),k6_numbers)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_31, plain, (k1_seq_1(k3_series_1(k1_lp_space(X1,X2)),X3)=k1_seq_1(k1_lp_space(X1,X2),X3)|k1_seq_1(k3_series_1(k1_lp_space(X1,X2)),esk11_2(X2,X1))=k1_seq_1(k1_lp_space(X1,X2),esk11_2(X2,X1))|k1_seq_1(k3_series_1(k1_lp_space(X1,X2)),k6_numbers)!=k1_seq_1(k1_lp_space(X1,X2),k6_numbers)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_30, plain, (k1_seq_1(k3_series_1(k1_lp_space(X1,X2)),X3)=k1_seq_1(k1_lp_space(X1,X2),X3)|k1_seq_1(k3_series_1(k1_lp_space(X1,X2)),k6_numbers)!=k1_seq_1(k1_lp_space(X1,X2),k6_numbers)|k1_seq_1(k3_series_1(k1_lp_space(X1,X2)),k2_nat_1(esk11_2(X2,X1),np__1))!=k1_seq_1(k1_lp_space(X1,X2),k2_nat_1(esk11_2(X2,X1),np__1))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_413, negated_conjecture, (k4_power(k4_series_1(k1_lp_space(esk2_0,esk1_0)),k7_xcmplx_0(np__1,esk1_0))!=k6_numbers|~v1_series_1(k1_lp_space(esk2_0,esk1_0))|~v1_xreal_0(esk1_0))).#trainneg
cnf(i_0_288, plain, (k18_complex1(X1)=X1|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~r1_xxreal_0(k6_numbers,X1))).#trainneg
cnf(i_0_420, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_431, negated_conjecture, (k3_power(k4_series_1(k1_lp_space(esk2_0,esk1_0)),k10_real_1(np__1,esk1_0))!=k6_numbers|~v1_series_1(k1_lp_space(esk2_0,esk1_0))|~m1_subset_1(k10_real_1(np__1,esk1_0),k1_numbers)|~m1_subset_1(k4_series_1(k1_lp_space(esk2_0,esk1_0)),k1_numbers))).#trainneg
cnf(i_0_234, plain, (v1_xreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_307, plain, (v7_ordinal1(k6_numbers))).#trainneg
cnf(i_0_309, plain, (v7_ordinal1(np__1))).#trainneg
cnf(i_0_236, plain, (v1_xreal_0(k18_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_283, plain, (k18_complex1(X1)=k4_xcmplx_0(X1)|r1_xxreal_0(k6_numbers,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_287, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_373, plain, (v1_xreal_0(k10_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_271, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_418, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk15_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_337, plain, (v3_ordinal1(k6_numbers))).#trainneg
cnf(i_0_355, negated_conjecture, (X1=k1_lp_space(esk2_0,X2)|k1_seq_1(X1,esk6_3(esk2_0,X2,X1))!=k4_power(k18_complex1(k6_numbers),X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(esk6_3(esk2_0,X2,X1),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_404, negated_conjecture, (v5_relat_1(esk2_0,k1_numbers))).#trainneg
cnf(i_0_339, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_333, negated_conjecture, (v3_ordinal1(esk1_0)|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_240, plain, (v1_xcmplx_0(np__1)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_242, plain, (v1_xcmplx_0(np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_363, plain, (v1_xreal_0(k4_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_461, plain, (k7_real_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_430, plain, (X1=k1_lp_space(X2,X3)|k1_seq_1(X1,esk6_3(X2,X3,X1))!=k3_power(k18_complex1(k1_seq_1(X2,esk6_3(X2,X3,X1))),X3)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(k18_complex1(k1_seq_1(X2,esk6_3(X2,X3,X1))),k1_numbers))).#trainneg
cnf(i_0_244, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_570, plain, (v1_funct_2(k2_series_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_566, plain, (v1_funct_1(k2_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_245, plain, (v1_xxreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_179, plain, (r1_xxreal_0(k6_numbers,k6_numbers))).#trainneg
cnf(i_0_610, negated_conjecture, (~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_199, plain, (~r1_xxreal_0(k6_numbers,k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_253, plain, (v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_571, plain, (m1_subset_1(k2_series_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_673, plain, (m1_subset_1(esk5_3(X1,X2,X3),k5_numbers)|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_641, plain, (m1_subset_1(esk10_3(X1,X2,X3),k5_numbers)|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_634, plain, (v4_relat_1(k1_lp_space(X1,X2),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_567, plain, (k2_seq_2(k2_series_1(X1))=k4_series_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_712, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~r1_xxreal_0(k6_numbers,X1))).#trainneg
cnf(i_0_720, plain, (v1_xxreal_0(np__1))).#trainneg
cnf(i_0_223, plain, (X1=k3_power(X2,X3)|X1!=k4_prepower(X2,X3)|~epred1_3(X2,X3,X1)|~v1_int_1(X3))).#trainneg
cnf(i_0_162, plain, (r1_xxreal_0(k6_numbers,np__1))).#trainneg
cnf(i_0_787, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_303, negated_conjecture, (v7_ordinal1(esk1_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_691, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk4_2(X2,X1),k18_complex1(k6_xcmplx_0(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1)))|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),k1_numbers))).#trainneg
cnf(i_0_752, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_674, plain, (m1_subset_1(esk5_3(X1,X2,X3),k1_numbers)|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_703, negated_conjecture, (k3_power(k4_series_1(k1_lp_space(esk2_0,esk1_0)),k7_xcmplx_0(np__1,esk1_0))!=k6_numbers|~v1_series_1(k1_lp_space(esk2_0,esk1_0))|~m1_subset_1(k7_xcmplx_0(np__1,esk1_0),k1_numbers)|~m1_subset_1(k4_series_1(k1_lp_space(esk2_0,esk1_0)),k1_numbers))).#trainneg
cnf(i_0_568, plain, (v2_comseq_2(k2_series_1(X1))|~v1_series_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_639, plain, (k1_funct_1(k2_series_1(X1),k6_numbers)=k1_funct_1(X1,k6_numbers)|~v1_valued_0(X1)|~v1_valued_0(k2_series_1(X1))|~v1_partfun1(X1,k5_numbers)|~v1_partfun1(k2_series_1(X1),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(k2_series_1(X1),k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(k2_series_1(X1))|~v1_funct_1(X1)|~v1_funct_1(k2_series_1(X1)))).#trainneg
cnf(i_0_642, plain, (m1_subset_1(esk10_3(X1,X2,X3),k1_numbers)|v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_707, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk4_2(X2,X1),k9_real_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1))|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(k9_real_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1))|~r1_xxreal_0(k6_numbers,k9_real_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1)))).#trainneg
cnf(i_0_724, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_186, plain, (k6_xcmplx_0(np__1,np__1)=k6_numbers)).#trainneg
cnf(i_0_490, negated_conjecture, (v2_valued_0(esk2_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_170, plain, (k2_xcmplx_0(k6_numbers,np__1)=np__1)).#trainneg
cnf(i_0_263, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_313, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_272, plain, (v1_funct_2(esk15_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2)|~v1_partfun1(esk15_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_690, negated_conjecture, (X1=k1_seq_2(esk2_0)|r1_xxreal_0(esk4_2(esk2_0,X1),k18_complex1(k9_real_1(k6_numbers,X1)))|~v2_comseq_2(esk2_0)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_subset_1(esk5_3(esk2_0,X1,X2),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_636, plain, (v1_partfun1(k1_lp_space(X1,X2),k5_numbers)|~v1_funct_2(k1_lp_space(X1,X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_632, plain, (v1_relat_1(k1_lp_space(X1,X2))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_653, plain, (k1_seq_1(k3_series_1(X1),X2)=k4_power(k18_complex1(k1_seq_1(X3,X2)),X4)|k3_series_1(X1)!=k1_lp_space(X3,X4)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(k3_series_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(k3_series_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X4,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_626, plain, (~v1_xboole_0(k1_lp_space(X1,X2))|~v1_funct_2(k1_lp_space(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(X1,X2))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_444, plain, (v4_relat_1(k3_series_1(X1),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_462, plain, (k7_real_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_693, plain, (k1_seq_1(k2_series_1(k1_lp_space(X1,X2)),X3)=k1_seq_1(k1_lp_space(X1,X2),X3)|m2_subset_1(esk11_2(X2,X1),k1_numbers,k5_numbers)|k1_seq_1(k2_series_1(k1_lp_space(X1,X2)),k6_numbers)!=k1_seq_1(k1_lp_space(X1,X2),k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(k1_lp_space(X1,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(X1,X2))|~m1_subset_1(k1_lp_space(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_749, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_197, plain, (k2_xcmplx_0(np__1,k6_numbers)=np__1)).#trainneg
cnf(i_0_189, plain, (k6_xcmplx_0(k6_numbers,np__1)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_275, plain, (v2_xxreal_0(k4_xcmplx_0(np__1))|~v3_xxreal_0(np__1)|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_351, plain, (v1_membered(k5_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_354, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_276, plain, (v2_xxreal_0(k4_xcmplx_0(X1))|~v3_xxreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_278, plain, (v3_xxreal_0(k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_537, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_416, plain, (k2_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_442, plain, (v5_relat_1(k3_series_1(X1),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_443, plain, (v1_relat_1(k3_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_793, plain, (k7_real_1(X2,X1)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1053, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_973, negated_conjecture, (v1_xcmplx_0(esk1_0))).#trainneg
cnf(i_0_910, plain, (k18_complex1(k6_numbers)=k6_numbers|k4_series_1(X1)!=k6_numbers|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_279, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|~v2_xxreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1021, plain, (v1_membered(k5_numbers))).#trainneg
cnf(i_0_1022, plain, (v2_membered(k5_numbers))).#trainneg
cnf(i_0_291, plain, (v1_xreal_0(esk15_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_1050, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_943, plain, (k1_seq_1(k2_series_1(X1),k6_numbers)=k1_funct_1(X1,k6_numbers)|~v1_relat_1(k2_series_1(X1))|~v3_valued_0(k2_series_1(X1))|~v1_funct_1(k2_series_1(X1))|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1097, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_seq_1(X1,esk12_2(X2,X1)),k1_seq_1(X2,k2_nat_1(esk12_2(X2,X1),np__1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(X2,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v3_valued_0(X2))).#trainneg
cnf(i_0_823, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1018, plain, (v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_292, negated_conjecture, (v1_xreal_0(esk2_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_889, plain, (~v3_xxreal_0(np__1))).#trainneg
cnf(i_0_992, negated_conjecture, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_298, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_869, plain, (v4_relat_1(k2_series_1(X1),k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_863, plain, (v3_valued_0(k2_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_957, negated_conjecture, (X1=k1_seq_2(esk2_0)|r1_xxreal_0(esk4_2(esk2_0,X1),k9_real_1(k6_numbers,X1))|~v2_comseq_2(esk2_0)|~v1_xreal_0(k9_real_1(k6_numbers,X1))|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(k6_numbers,k9_real_1(k6_numbers,X1))|~m2_subset_1(esk5_3(esk2_0,X1,X2),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1090, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_976, plain, (k1_seq_1(k2_series_1(X1),X2)=k4_power(k18_complex1(k1_seq_1(X3,X2)),X4)|k2_series_1(X1)!=k1_lp_space(X3,X4)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X4,k1_numbers))).#trainneg
cnf(i_0_871, plain, (v1_partfun1(k2_series_1(X1),k5_numbers)|~v1_funct_2(k2_series_1(X1),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_924, negated_conjecture, (X1=k1_seq_2(esk2_0)|r1_xxreal_0(esk4_2(esk2_0,X1),k18_complex1(k6_xcmplx_0(k6_numbers,X1)))|~v2_comseq_2(esk2_0)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m2_subset_1(esk5_3(esk2_0,X1,X2),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_892, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_368, plain, (v1_xreal_0(k2_seq_2(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_301, plain, (v7_ordinal1(esk15_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_766, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_302, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_306, plain, (v7_ordinal1(k18_complex1(X1))|~v6_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_824, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1054, plain, (m2_subset_1(k2_xcmplx_0(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_874, plain, (v7_ordinal1(esk5_3(X1,X2,X3))|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_932, plain, (X1=k1_seq_2(X2)|k9_real_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1)!=k6_numbers|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k9_real_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1),k5_numbers))).#trainneg
cnf(i_0_931, plain, (X1=k1_seq_2(X2)|k9_real_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1)!=k6_numbers|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k9_real_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1),k1_numbers))).#trainneg
cnf(i_0_867, plain, (v1_relat_1(k2_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_485, plain, (v3_valued_0(k3_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_893, plain, (X1=k3_power(X2,X3)|X1!=k4_prepower(X2,X3)|~v1_int_1(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_308, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_449, plain, (v3_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v3_valued_0(k1_numbers))).#trainneg
cnf(i_0_318, plain, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_328, plain, (v1_xxreal_0(k5_numbers)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_970, plain, (v1_valued_0(k2_zfmisc_1(X1,k5_numbers)))).#trainneg
cnf(i_0_971, plain, (v1_valued_0(k2_zfmisc_1(X1,k1_numbers)))).#trainneg
cnf(i_0_311, plain, (v1_xcmplx_0(esk15_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_312, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_975, negated_conjecture, (X1=k1_seq_2(esk2_0)|r1_xxreal_0(esk4_2(esk2_0,X1),k18_complex1(k9_real_1(k6_numbers,X1)))|~v2_comseq_2(esk2_0)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1056, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X1,esk12_2(X2,X1)),k1_funct_1(X2,k2_xcmplx_0(np__1,esk12_2(X2,X1))))!=k1_funct_1(X1,k2_xcmplx_0(np__1,esk12_2(X2,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v7_ordinal1(esk12_2(X2,X1)))).#trainneg
cnf(i_0_1096, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_seq_1(X1,esk12_2(X2,X1)),k1_funct_1(X2,k2_nat_1(np__1,esk12_2(X2,X1))))!=k1_funct_1(X1,k2_nat_1(np__1,esk12_2(X2,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v7_ordinal1(esk12_2(X2,X1)))).#trainneg
cnf(i_0_1095, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_seq_1(X1,esk12_2(X2,X1)),k1_funct_1(X2,k2_xcmplx_0(np__1,esk12_2(X2,X1))))!=k1_funct_1(X1,k2_xcmplx_0(np__1,esk12_2(X2,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v7_ordinal1(esk12_2(X2,X1)))).#trainneg
cnf(i_0_958, plain, (X1=k1_seq_2(X2)|r1_xxreal_0(esk4_2(X2,X1),k6_xcmplx_0(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1))|~v2_comseq_2(X2)|~v1_xreal_0(k6_xcmplx_0(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~r1_xxreal_0(k6_numbers,k6_xcmplx_0(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),k1_numbers))).#trainneg
cnf(i_0_827, plain, (~v3_xxreal_0(X1)|~v2_xxreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1060, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1173, negated_conjecture, (r1_xxreal_0(esk1_0,esk1_0))).#trainneg
cnf(i_0_316, plain, (v1_xcmplx_0(k18_complex1(X1))|~v1_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1268, negated_conjecture, (r1_xxreal_0(k6_numbers,esk1_0)|r1_xxreal_0(esk1_0,k6_numbers))).#trainneg
cnf(i_0_1195, negated_conjecture, (v3_xxreal_0(X1)|v2_xxreal_0(X1)|v1_xboole_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_321, plain, (v1_xxreal_0(esk15_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_322, negated_conjecture, (v1_xxreal_0(esk2_0)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1125, plain, (k4_power(k1_seq_1(X1,X2),X3)=k1_seq_1(k1_lp_space(X1,X3),X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_xxreal_0(k6_numbers,k1_seq_1(X1,X2)))).#trainneg
cnf(i_0_1157, plain, (m1_subset_1(k1_funct_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1135, plain, (k1_seq_1(k2_series_1(X1),k6_numbers)=k1_funct_1(X1,k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_1201, plain, (k1_seq_1(k2_series_1(X1),k6_numbers)=k1_funct_1(X1,k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1138, negated_conjecture, (r1_xxreal_0(k10_real_1(np__1,esk1_0),k6_numbers)|~v1_series_1(k1_lp_space(esk2_0,esk1_0))|~m1_subset_1(k4_series_1(k1_lp_space(esk2_0,esk1_0)),k1_numbers)|~m1_subset_1(k10_real_1(np__1,esk1_0),k1_numbers)|k4_series_1(k1_lp_space(esk2_0,esk1_0))!=k6_numbers|~v1_xreal_0(k10_real_1(np__1,esk1_0))|~v1_xreal_0(k4_series_1(k1_lp_space(esk2_0,esk1_0))))).#trainneg
cnf(i_0_1129, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xcmplx_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1179, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_278, plain, (v3_xxreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_326, plain, (v1_xxreal_0(k18_complex1(X1))|~v2_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_495, plain, (v1_valued_0(k5_numbers)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1303, plain, (v1_xcmplx_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1304, plain, (v1_xxreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_331, plain, (v3_ordinal1(esk15_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_332, negated_conjecture, (v3_ordinal1(esk2_0)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1371, plain, (m2_subset_1(k2_xcmplx_0(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1219, plain, (m2_subset_1(k2_xcmplx_0(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1272, plain, (r1_xxreal_0(k6_numbers,X1)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1137, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(np__1,esk1_0),k6_numbers)|~v1_series_1(k1_lp_space(esk2_0,esk1_0))|~m1_subset_1(k4_series_1(k1_lp_space(esk2_0,esk1_0)),k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,esk1_0),k1_numbers)|k4_series_1(k1_lp_space(esk2_0,esk1_0))!=k6_numbers|~v1_xreal_0(k7_xcmplx_0(np__1,esk1_0))|~v1_xreal_0(k4_series_1(k1_lp_space(esk2_0,esk1_0))))).#trainneg
cnf(i_0_1159, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1178, plain, (r1_xxreal_0(X1,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1147, negated_conjecture, (k2_series_1(X1)=esk2_0|k2_xcmplx_0(k6_numbers,k1_seq_1(X1,k2_nat_1(esk12_2(X1,esk2_0),np__1)))!=k1_funct_1(esk2_0,k2_nat_1(esk12_2(X1,esk2_0),np__1))|k1_funct_1(esk2_0,k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1)|~m2_subset_1(esk12_2(X1,esk2_0),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_336, plain, (v3_ordinal1(k18_complex1(X1))|~v3_ordinal1(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_499, plain, (v2_valued_0(k5_numbers)|~v2_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_344, negated_conjecture, (v3_membered(esk2_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_347, negated_conjecture, (v1_xboole_0(esk2_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_735, negated_conjecture, (m1_subset_1(k1_seq_1(esk2_0,X1),k1_numbers))).#trainneg
cnf(i_0_341, plain, (v1_relat_1(esk15_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_343, plain, (v3_membered(esk15_1(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_1158, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1166, plain, (m1_subset_1(k16_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1119, plain, (k3_power(k18_complex1(k1_seq_1(X1,X2)),X3)=k1_seq_1(k1_lp_space(X1,X3),X2)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(k18_complex1(k1_seq_1(X1,X2)),k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1249, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1182, plain, (m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).#trainneg
cnf(i_0_346, plain, (v1_xboole_0(esk15_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_350, negated_conjecture, (v1_membered(esk2_0)|~v1_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_353, negated_conjecture, (v2_membered(esk2_0)|~v2_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_1113, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1247, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_349, plain, (v1_membered(esk15_1(k1_zfmisc_1(X1)))|~v1_membered(X1))).#trainneg
cnf(i_0_352, plain, (v2_membered(esk15_1(k1_zfmisc_1(X1)))|~v2_membered(X1))).#trainneg
cnf(i_0_1207, plain, (k2_xcmplx_0(k1_funct_1(X3,k2_nat_1(X2,np__1)),k1_funct_1(X1,X2))=k1_funct_1(X1,k2_nat_1(X2,np__1))|X1!=k2_series_1(X3)|~v1_valued_0(X3)|~v1_valued_0(X1)|~v1_partfun1(X3,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X3)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,X2),k1_numbers)|~m1_subset_1(k1_funct_1(X3,k2_nat_1(X2,np__1)),k1_numbers))).#trainneg
cnf(i_0_1274, plain, (r1_xxreal_0(k6_numbers,X1)|r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1419, plain, (k3_power(k1_seq_1(X1,X2),X3)=k1_seq_1(k1_lp_space(X1,X3),X2)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~r1_xxreal_0(k6_numbers,k1_seq_1(X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1551, negated_conjecture, (v1_xreal_0(k1_seq_1(esk2_0,X1)))).#trainneg
cnf(i_0_1406, negated_conjecture, (r1_xxreal_0(k6_numbers,esk1_0))).#trainneg
cnf(i_0_503, plain, (m1_subset_1(k1_seq_2(X1),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1549, plain, (v2_valued_0(k5_numbers)|~v1_relat_1(k1_numbers)|~v3_valued_0(k1_numbers))).#trainneg
cnf(i_0_1542, negated_conjecture, (k2_series_1(esk2_0)=esk2_0|k1_funct_1(esk2_0,k2_nat_1(esk12_2(esk2_0,esk2_0),np__1))!=k2_xcmplx_0(k6_numbers,k6_numbers)|~m2_subset_1(esk12_2(esk2_0,esk2_0),k1_numbers,k5_numbers)|~m2_subset_1(k2_nat_1(esk12_2(esk2_0,esk2_0),np__1),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_358, plain, (v3_xxreal_0(k3_power(X1,X2))|v2_xxreal_0(X3)|r1_xxreal_0(X3,k3_power(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1453, plain, (~v2_xxreal_0(k4_xcmplx_0(np__1))|~m1_subset_1(k4_xcmplx_0(np__1),k5_numbers))).#trainneg
cnf(i_0_1616, negated_conjecture, (v1_membered(esk2_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_1553, negated_conjecture, (v1_xcmplx_0(k1_seq_1(esk2_0,X1))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_359, plain, (v3_xxreal_0(k1_seq_2(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k1_seq_2(X1))|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_360, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1441, negated_conjecture, (r1_xxreal_0(k10_real_1(np__1,esk1_0),k6_numbers)|k4_series_1(k1_lp_space(esk2_0,esk1_0))!=k6_numbers|~v1_xreal_0(k10_real_1(np__1,esk1_0))|~v1_series_1(k1_lp_space(esk2_0,esk1_0))|~v1_funct_2(k1_lp_space(esk2_0,esk1_0),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,esk1_0))|~m1_subset_1(k1_lp_space(esk2_0,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1440, negated_conjecture, (r1_xxreal_0(k10_real_1(np__1,esk1_0),k6_numbers)|k4_series_1(k1_lp_space(esk2_0,esk1_0))!=k6_numbers|~v1_xreal_0(k10_real_1(np__1,esk1_0))|~v1_series_1(k1_lp_space(esk2_0,esk1_0))|~m1_subset_1(k4_series_1(k1_lp_space(esk2_0,esk1_0)),k5_numbers))).#trainneg
cnf(i_0_1439, negated_conjecture, (r1_xxreal_0(k10_real_1(np__1,esk1_0),k6_numbers)|k4_series_1(k1_lp_space(esk2_0,esk1_0))!=k6_numbers|~v1_xreal_0(k10_real_1(np__1,esk1_0))|~v1_series_1(k1_lp_space(esk2_0,esk1_0))|~m1_subset_1(k4_series_1(k1_lp_space(esk2_0,esk1_0)),k1_numbers))).#trainneg
cnf(i_0_1691, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|r1_xxreal_0(esk4_2(k1_lp_space(esk2_0,X2),X1),k18_complex1(k9_real_1(k4_power(k6_numbers,X2),X1)))|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1529, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(np__1,esk1_0),k6_numbers)|k4_series_1(k1_lp_space(esk2_0,esk1_0))!=k6_numbers|~v1_xreal_0(k7_xcmplx_0(np__1,esk1_0))|~v1_series_1(k1_lp_space(esk2_0,esk1_0))|~v1_funct_2(k1_lp_space(esk2_0,esk1_0),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,esk1_0))|~m1_subset_1(k1_lp_space(esk2_0,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1528, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(np__1,esk1_0),k6_numbers)|k4_series_1(k1_lp_space(esk2_0,esk1_0))!=k6_numbers|~v1_xreal_0(k7_xcmplx_0(np__1,esk1_0))|~v1_series_1(k1_lp_space(esk2_0,esk1_0))|~m1_subset_1(k4_series_1(k1_lp_space(esk2_0,esk1_0)),k5_numbers))).#trainneg
cnf(i_0_1527, negated_conjecture, (r1_xxreal_0(k7_xcmplx_0(np__1,esk1_0),k6_numbers)|k4_series_1(k1_lp_space(esk2_0,esk1_0))!=k6_numbers|~v1_xreal_0(k7_xcmplx_0(np__1,esk1_0))|~v1_series_1(k1_lp_space(esk2_0,esk1_0))|~m1_subset_1(k4_series_1(k1_lp_space(esk2_0,esk1_0)),k1_numbers))).#trainneg
cnf(i_0_1693, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|k9_real_1(k4_power(k6_numbers,X2),X1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(k9_real_1(k4_power(k6_numbers,X2),X1))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1650, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_866, plain, (v5_relat_1(k2_series_1(X1),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1608, plain, (v7_ordinal1(k4_xcmplx_0(np__1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1701, negated_conjecture, (k2_series_1(k1_lp_space(esk2_0,X1))=esk2_0|k1_funct_1(esk2_0,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X1),esk2_0),np__1))!=k2_xcmplx_0(k6_numbers,k4_power(k6_numbers,X1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X1),esk2_0),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X1),esk2_0),np__1),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_361, plain, (v3_xxreal_0(k16_complex1(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k16_complex1(X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1617, negated_conjecture, (v2_membered(esk2_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_1198, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1467, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1248, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_362, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_364, plain, (v7_ordinal1(k4_series_1(X1))|~v6_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1696, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|k9_real_1(k4_power(k6_numbers,X2),X1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k9_real_1(k4_power(k6_numbers,X2),X1),k1_numbers)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1525, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1596, negated_conjecture, (k1_seq_1(k1_lp_space(esk2_0,X2),X1)=k3_power(k6_numbers,X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1609, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1563, negated_conjecture, (k2_series_1(X1)=esk2_0|k1_funct_1(X1,k2_nat_1(esk12_2(X1,esk2_0),np__1))!=k1_seq_1(esk2_0,k2_nat_1(esk12_2(X1,esk2_0),np__1))|k1_funct_1(esk2_0,k6_numbers)!=k1_funct_1(X1,k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m2_subset_1(esk12_2(X1,esk2_0),k1_numbers,k5_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_365, plain, (v1_xcmplx_0(k4_series_1(X1))|~v1_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1550, negated_conjecture, (v3_membered(esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1305, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1418, negated_conjecture, (v1_xxreal_0(esk2_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1273, plain, (r1_xxreal_0(k6_numbers,X1)|r1_xxreal_0(X1,k6_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1624, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|r1_xxreal_0(X1,esk1_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_366, plain, (v1_xxreal_0(k4_series_1(X1))|~v2_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_367, plain, (v3_ordinal1(k4_series_1(X1))|~v3_ordinal1(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1958, negated_conjecture, (k4_power(k6_numbers,X1)=k3_power(k6_numbers,X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1911, plain, (r1_xxreal_0(k6_numbers,esk5_3(X1,X2,X3))|r1_xxreal_0(esk5_3(X1,X2,X3),k6_numbers)|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1771, negated_conjecture, (v1_xcmplx_0(k1_seq_1(esk2_0,X1)))).#trainneg
cnf(i_0_1950, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k3_power(k6_numbers,X2))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1692, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|r1_xxreal_0(esk4_2(k1_lp_space(esk2_0,X2),X1),k18_complex1(k6_xcmplx_0(k4_power(k6_numbers,X2),X1)))|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k4_power(k6_numbers,X2),k1_numbers)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1686, negated_conjecture, (X1=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|k1_seq_1(X1,esk6_3(k1_lp_space(esk2_0,X2),X3,X1))!=k4_power(k18_complex1(k4_power(k6_numbers,X2)),X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m2_subset_1(esk6_3(k1_lp_space(esk2_0,X2),X3,X1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1730, negated_conjecture, (k1_lp_space(esk2_0,X1)=esk2_0|k1_seq_1(k1_lp_space(esk2_0,X1),X2)!=k6_numbers|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_369, plain, (v7_ordinal1(k2_seq_2(X1))|~v6_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1770, negated_conjecture, (v1_membered(esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1625, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|r1_xxreal_0(X1,esk1_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_761, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_370, plain, (v1_xcmplx_0(k2_seq_2(X1))|~v1_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_371, plain, (v1_xxreal_0(k2_seq_2(X1))|~v2_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1936, negated_conjecture, (X1=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|k1_seq_1(X1,esk6_3(k1_lp_space(esk2_0,X2),X3,X1))!=k4_power(k18_complex1(k3_power(k6_numbers,X2)),X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m2_subset_1(esk6_3(k1_lp_space(esk2_0,X2),X3,X1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1520, plain, (r1_xxreal_0(k16_complex1(X1),k6_numbers)|r1_xxreal_0(k6_numbers,k16_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1954, negated_conjecture, (k4_power(k18_complex1(k3_power(k6_numbers,X1)),X3)=k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X1),X3),X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1700, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k4_power(k6_numbers,X2))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1699, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|k6_xcmplx_0(k4_power(k6_numbers,X2),X1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(k4_power(k6_numbers,X2))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1995, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_765, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_1610, plain, (v1_xxreal_0(k4_xcmplx_0(np__1))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1857, negated_conjecture, (k2_series_1(k1_lp_space(esk2_0,X1))=esk2_0|k1_seq_1(esk2_0,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X1),esk2_0),np__1))!=k2_xcmplx_0(k6_numbers,k4_power(k6_numbers,X1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X1),esk2_0),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X1),esk2_0),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_372, plain, (v3_ordinal1(k2_seq_2(X1))|~v3_ordinal1(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1867, negated_conjecture, (v2_membered(esk2_0)|~m1_subset_1(k2_zfmisc_1(k5_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1869, plain, (~v3_xxreal_0(k5_numbers)|~v2_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1466, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_764, plain, (v2_valued_0(k2_zfmisc_1(X1,X2))|~v2_membered(X2))).#trainneg
cnf(i_0_374, plain, (v7_ordinal1(k10_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_375, plain, (v1_xcmplx_0(k10_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1697, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|k9_real_1(k4_power(k6_numbers,X2),X1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k9_real_1(k4_power(k6_numbers,X2),X1),k5_numbers)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1518, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k6_numbers))).#trainneg
cnf(i_0_1715, plain, (k1_seq_1(k1_lp_space(k2_series_1(X1),X2),k6_numbers)=k3_power(k1_funct_1(X1,k6_numbers),X2)|~v1_xreal_0(k1_funct_1(X1,k6_numbers))|~v1_funct_2(k2_series_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k2_series_1(X1))|~r1_xxreal_0(k6_numbers,k1_funct_1(X1,k6_numbers))|~m1_subset_1(k2_series_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1507, plain, (r1_xxreal_0(k6_xcmplx_0(X1,X2),k6_numbers)|r1_xxreal_0(k6_numbers,k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1694, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|r1_xxreal_0(esk4_2(k1_lp_space(esk2_0,X2),X1),k9_real_1(k4_power(k6_numbers,X2),X1))|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(k9_real_1(k4_power(k6_numbers,X2),X1))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~r1_xxreal_0(k6_numbers,k9_real_1(k4_power(k6_numbers,X2),X1))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1695, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|k6_xcmplx_0(k4_power(k6_numbers,X2),X1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(k6_xcmplx_0(k4_power(k6_numbers,X2),X1))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k4_power(k6_numbers,X2),k1_numbers)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1511, plain, (r1_xxreal_0(k7_xcmplx_0(X1,X2),k6_numbers)|r1_xxreal_0(k6_numbers,k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2102, plain, (k2_zfmisc_1(X1,X2)=k6_numbers|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2193, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k6_numbers|k2_xcmplx_0(k4_power(k6_numbers,X1),k6_numbers)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~v1_xreal_0(k4_power(k6_numbers,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(k4_power(k6_numbers,X1)))).#trainneg
cnf(i_0_376, plain, (v1_xxreal_0(k10_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1626, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|r1_xxreal_0(X1,esk1_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2244, plain, (~v2_xxreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_2216, plain, (~v1_xboole_0(k4_xcmplx_0(np__1))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_2215, plain, (v3_valued_0(k2_zfmisc_1(X1,k1_numbers)))).#trainneg
cnf(i_0_767, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_377, plain, (v3_ordinal1(k10_real_1(X1,X2))|~v3_ordinal1(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_378, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X2,k2_nat_1(esk12_2(X2,X1),np__1)),k1_funct_1(X1,esk12_2(X2,X1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(X2,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_xcmplx_0(k1_funct_1(X1,esk12_2(X2,X1)))|~v1_xcmplx_0(k1_funct_1(X2,k2_nat_1(esk12_2(X2,X1),np__1))))).#trainneg
cnf(i_0_2015, negated_conjecture, (k1_funct_1(esk2_0,k6_numbers)=k1_seq_1(esk2_0,k6_numbers))).#trainneg
cnf(i_0_2289, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(k6_numbers,k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2031, negated_conjecture, (k1_funct_1(X1,k6_numbers)=k1_funct_1(esk2_0,k6_numbers)|X1!=esk2_0|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2017, negated_conjecture, (k4_series_1(esk2_0)=k2_seq_2(esk2_0))).#trainneg
cnf(i_0_2032, negated_conjecture, (k4_power(k6_numbers,X1)=k3_power(k6_numbers,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2055, plain, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))!=k2_xcmplx_0(k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k6_numbers)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2179, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k1_seq_1(k1_lp_space(esk2_0,X2),X3))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_768, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_2018, negated_conjecture, (v2_comseq_2(esk2_0)|~v1_series_1(esk2_0))).#trainneg
cnf(i_0_2180, negated_conjecture, (k2_series_1(k1_lp_space(esk2_0,X1))=esk2_0|k1_funct_1(esk2_0,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X1),esk2_0),np__1))!=k2_xcmplx_0(k6_numbers,k4_power(k6_numbers,X1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X1),esk2_0),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X1),esk2_0),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_384, plain, (v3_xxreal_0(k3_power(X1,X2))|v2_xxreal_0(X3)|r1_xxreal_0(X3,k3_power(X1,X2))|~v1_xreal_0(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1267, plain, (v7_ordinal1(X1)|~v1_xcmplx_0(X1)|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~r1_xxreal_0(k6_numbers,X1))).#trainneg
cnf(i_0_1552, negated_conjecture, (v7_ordinal1(k1_seq_1(esk2_0,X1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1554, negated_conjecture, (v1_xxreal_0(k1_seq_1(esk2_0,X1))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_385, plain, (v3_xxreal_0(k9_prepower(X1,X2))|v2_xxreal_0(X3)|r1_xxreal_0(X3,k9_prepower(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2058, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k3_power(k6_numbers,X1),k3_power(k6_numbers,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2173, negated_conjecture, (X1=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|k1_seq_1(X1,esk6_3(k1_lp_space(esk2_0,X2),X3,X1))!=k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),X4)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2181, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k4_power(k6_numbers,X1),k4_power(k6_numbers,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2057, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k4_power(k6_numbers,X1),k3_power(k6_numbers,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2182, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k3_power(k6_numbers,X1),k4_power(k6_numbers,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2131, negated_conjecture, (k1_lp_space(esk2_0,X1)=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|k3_power(k6_numbers,X1)!=k4_power(k18_complex1(k3_power(k6_numbers,X2)),X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(esk6_3(k1_lp_space(esk2_0,X2),X3,k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2251, negated_conjecture, (k3_power(k1_funct_1(esk2_0,k6_numbers),X1)=k1_seq_1(k1_lp_space(esk2_0,X1),k6_numbers)|~v1_xreal_0(k1_funct_1(esk2_0,k6_numbers))|~r1_xxreal_0(k6_numbers,k1_funct_1(esk2_0,k6_numbers))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2089, negated_conjecture, (k1_lp_space(esk2_0,X1)=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|k3_power(k6_numbers,X1)!=k4_power(k18_complex1(k4_power(k6_numbers,X2)),X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(esk6_3(k1_lp_space(esk2_0,X2),X3,k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2052, plain, (X2=k1_seq_2(X1)|r1_xxreal_0(k6_numbers,esk5_3(X1,X2,X3))|~v2_xxreal_0(esk5_3(X1,X2,X3))|~v1_xreal_0(esk5_3(X1,X2,X3))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_753, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1257, plain, (v1_xreal_0(esk15_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_1611, plain, (v3_ordinal1(k4_xcmplx_0(np__1))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_2056, negated_conjecture, (k2_series_1(k1_lp_space(esk2_0,X1))=esk2_0|k1_funct_1(esk2_0,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X1),esk2_0),np__1))!=k2_xcmplx_0(k6_numbers,k3_power(k6_numbers,X1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X1),esk2_0),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X1),esk2_0),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_386, plain, (v3_xxreal_0(k7_xcmplx_0(X1,X2))|v2_xxreal_0(X3)|r1_xxreal_0(X3,k7_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2463, negated_conjecture, (m2_subset_1(k1_seq_1(esk2_0,X1),k1_numbers,k5_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2293, plain, (k2_zfmisc_1(X1,k6_numbers)=k6_numbers)).#trainneg
cnf(i_0_2321, plain, (k2_zfmisc_1(X1,X2)=k6_numbers|~v1_xboole_0(X1))).#trainneg
cnf(i_0_389, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(X3)|r1_xxreal_0(X3,k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2545, negated_conjecture, (k1_seq_1(k1_lp_space(esk2_0,X1),k6_numbers)=k3_power(k6_numbers,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2403, negated_conjecture, (k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2)=k1_seq_1(k1_lp_space(esk2_0,X2),X3)|r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2426, plain, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k1_seq_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))!=k2_xcmplx_0(k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k6_numbers)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1))).#trainneg
cnf(i_0_2586, plain, (v3_xxreal_0(k7_xcmplx_0(X1,X2))|r1_xxreal_0(k6_numbers,k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_754, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1258, plain, (v1_xxreal_0(esk15_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_2358, negated_conjecture, (v1_xreal_0(k1_seq_2(esk2_0)))).#trainneg
cnf(i_0_1857, negated_conjecture, (k2_series_1(k1_lp_space(esk2_0,X1))=esk2_0|k1_seq_1(esk2_0,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X1),esk2_0),np__1))!=k2_xcmplx_0(k6_numbers,k4_power(k6_numbers,X1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_seq_1(esk2_0,k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X1),esk2_0),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X1),esk2_0),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2318, plain, (~v1_xboole_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1555, negated_conjecture, (v3_ordinal1(k1_seq_1(esk2_0,X1))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_2103, plain, (v1_xreal_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_395, plain, (v3_xxreal_0(k6_xcmplx_0(X1,X2))|v2_xxreal_0(X3)|r1_xxreal_0(X3,k6_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2125, plain, (X1=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|r1_xxreal_0(X2,k6_numbers)|k1_seq_1(X1,esk6_3(k1_lp_space(esk2_0,X2),X3,X1))!=k4_power(k6_numbers,X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2399, negated_conjecture, (X1=k1_lp_space(esk2_0,X3)|r1_xxreal_0(X2,k6_numbers)|k1_seq_1(X1,esk6_3(esk2_0,X3,X1))!=k4_power(k18_complex1(k3_power(k6_numbers,X2)),X3)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2130, negated_conjecture, (k1_lp_space(esk2_0,X1)=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|k4_power(k6_numbers,X1)!=k4_power(k18_complex1(k3_power(k6_numbers,X2)),X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(esk6_3(k1_lp_space(esk2_0,X2),X3,k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2397, negated_conjecture, (X1=k1_lp_space(esk2_0,X3)|r1_xxreal_0(X2,k6_numbers)|k1_seq_1(X1,esk6_3(esk2_0,X3,X1))!=k4_power(k18_complex1(k4_power(k6_numbers,X2)),X3)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2728, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X3,k6_numbers)|k2_xcmplx_0(k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k4_power(k18_complex1(k3_power(k6_numbers,X3)),X2))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~v1_xreal_0(X3)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2696, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k2_xcmplx_0(k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k1_seq_1(esk2_0,X3))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2719, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|r1_xxreal_0(esk4_2(k1_lp_space(esk2_0,X2),X1),k18_complex1(k9_real_1(k4_power(k18_complex1(k3_power(k6_numbers,X4)),X2),X1)))|r1_xxreal_0(X4,k6_numbers)|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(X4)|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2721, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X4,k6_numbers)|k9_real_1(k4_power(k18_complex1(k3_power(k6_numbers,X4)),X2),X1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(k9_real_1(k4_power(k18_complex1(k3_power(k6_numbers,X4)),X2),X1))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(X4)|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2601, plain, (v4_relat_1(k6_numbers,X1))).#trainneg
cnf(i_0_2613, plain, (k2_zfmisc_1(k6_numbers,X1)=k6_numbers)).#trainneg
cnf(i_0_2599, plain, (v1_relat_1(k6_numbers))).#trainneg
cnf(i_0_2598, plain, (v5_relat_1(k6_numbers,k6_numbers))).#trainneg
cnf(i_0_2597, plain, (v2_valued_0(k6_numbers)|~v2_membered(k6_numbers))).#trainneg
cnf(i_0_400, plain, (v1_xxreal_0(k3_power(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2591, plain, (v3_valued_0(k6_numbers)|~v3_membered(k6_numbers))).#trainneg
cnf(i_0_1851, plain, (v1_xxreal_0(k4_xcmplx_0(np__1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2359, negated_conjecture, (v7_ordinal1(k1_seq_2(esk2_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2602, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2607, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_401, plain, (v4_relat_1(esk15_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_403, plain, (v5_relat_1(esk15_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_2757, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k2_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X1),esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1))),k6_numbers)!=k4_power(k6_numbers,X1)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2703, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X3,k6_numbers)|k2_xcmplx_0(k4_power(k18_complex1(k3_power(k6_numbers,X3)),X1),k3_power(k6_numbers,X2))!=k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X3)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2702, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X3,k6_numbers)|k2_xcmplx_0(k4_power(k18_complex1(k3_power(k6_numbers,X3)),X1),k4_power(k6_numbers,X2))!=k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X3)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2758, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k2_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X1),esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1))),k6_numbers)!=k3_power(k6_numbers,X1)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2659, negated_conjecture, (v2_comseq_2(esk2_0)|r1_xxreal_0(esk9_2(esk2_0,X1),k18_complex1(k9_real_1(k4_power(k6_numbers,X3),X1)))|r1_xxreal_0(X3,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_xreal_0(X3)|~m2_subset_1(esk10_3(esk2_0,X1,X2),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2714, negated_conjecture, (X1=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|r1_xxreal_0(X4,k6_numbers)|k1_seq_1(X1,esk6_3(k1_lp_space(esk2_0,X2),X3,X1))!=k4_power(k18_complex1(k4_power(k18_complex1(k3_power(k6_numbers,X4)),X2)),X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_xreal_0(X4)|~m2_subset_1(esk6_3(k1_lp_space(esk2_0,X2),X3,X1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2651, plain, (k3_power(k6_numbers,X1)=k1_seq_1(esk2_0,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2645, negated_conjecture, (k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2)=k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k6_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2727, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X4,k6_numbers)|k6_xcmplx_0(k4_power(k18_complex1(k3_power(k6_numbers,X4)),X2),X1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(k4_power(k18_complex1(k3_power(k6_numbers,X4)),X2))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(X4)|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2088, negated_conjecture, (k1_lp_space(esk2_0,X1)=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|k4_power(k6_numbers,X1)!=k4_power(k18_complex1(k4_power(k6_numbers,X2)),X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(esk6_3(k1_lp_space(esk2_0,X2),X3,k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2337, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_seq_1(X2,k2_nat_1(esk12_2(X2,X1),np__1)),k1_funct_1(X1,esk12_2(X2,X1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(X2,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v3_valued_0(X2))).#trainneg
cnf(i_0_2990, plain, (k2_zfmisc_1(X1,X2)=k6_numbers|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2992, plain, (k2_zfmisc_1(X1,X2)=k6_numbers|~m1_subset_1(X2,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2926, plain, (v1_valued_0(k6_numbers))).#trainneg
cnf(i_0_1563, negated_conjecture, (k2_series_1(X1)=esk2_0|k1_funct_1(X1,k2_nat_1(esk12_2(X1,esk2_0),np__1))!=k1_seq_1(esk2_0,k2_nat_1(esk12_2(X1,esk2_0),np__1))|k1_funct_1(X1,k6_numbers)!=k1_seq_1(esk2_0,k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m2_subset_1(esk12_2(X1,esk2_0),k1_numbers,k5_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_405, plain, (k2_xcmplx_0(k4_xcmplx_0(X2),X1)=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2974, negated_conjecture, (~m1_subset_1(esk2_0,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2975, negated_conjecture, (~m1_subset_1(esk1_0,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2928, plain, (v3_valued_0(k6_numbers))).#trainneg
cnf(i_0_406, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2522, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_seq_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k4_power(k6_numbers,X1),k4_power(k6_numbers,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2255, plain, (r1_xxreal_0(k6_numbers,k6_xcmplx_0(X1,X2))|~v2_xxreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3119, plain, (k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k6_numbers)=k4_power(k6_numbers,X2)|r1_xxreal_0(X1,k6_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2288, plain, (r1_xxreal_0(k6_numbers,k7_xcmplx_0(X1,X2))|~v2_xxreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2061, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k3_power(k6_numbers,X2),k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k3_power(k6_numbers,X2),k1_numbers)|~m1_subset_1(k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k1_numbers))).#trainneg
cnf(i_0_2688, negated_conjecture, (k1_seq_1(esk2_0,X2)=k3_power(k6_numbers,X1)|r1_xxreal_0(X1,k6_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3123, negated_conjecture, (k4_power(k3_power(k6_numbers,X1),X2)=k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k6_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(k3_power(k6_numbers,X1))|~r1_xxreal_0(k6_numbers,k3_power(k6_numbers,X1)))).#trainneg
cnf(i_0_2704, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X4,k6_numbers)|k2_xcmplx_0(k4_power(k18_complex1(k3_power(k6_numbers,X4)),X1),k1_seq_1(k1_lp_space(esk2_0,X2),X3))!=k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X4)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2724, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X4,k6_numbers)|k9_real_1(k4_power(k18_complex1(k3_power(k6_numbers,X4)),X2),X1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k9_real_1(k4_power(k18_complex1(k3_power(k6_numbers,X4)),X2),X1),k1_numbers)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(X4)|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2984, plain, (X1=k6_numbers|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2986, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2360, negated_conjecture, (v1_xcmplx_0(k1_seq_2(esk2_0))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1147, negated_conjecture, (k2_series_1(X1)=esk2_0|k2_xcmplx_0(k6_numbers,k1_seq_1(X1,k2_nat_1(esk12_2(X1,esk2_0),np__1)))!=k1_funct_1(esk2_0,k2_nat_1(esk12_2(X1,esk2_0),np__1))|k1_funct_1(X1,k6_numbers)!=k1_seq_1(esk2_0,k6_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m2_subset_1(esk12_2(X1,esk2_0),k1_numbers,k5_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_409, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(np__1))=k2_xcmplx_0(X1,np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2981, plain, (~m1_subset_1(k5_numbers,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2982, plain, (~m1_subset_1(np__1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2936, plain, (v5_relat_1(k6_numbers,X1))).#trainneg
cnf(i_0_2935, plain, (v2_valued_0(k6_numbers)|~v2_membered(X1))).#trainneg
cnf(i_0_410, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_411, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X1,esk12_2(X2,X1)),k1_funct_1(X2,k2_nat_1(np__1,esk12_2(X2,X1))))!=k1_funct_1(X1,k2_nat_1(np__1,esk12_2(X2,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v7_ordinal1(esk12_2(X2,X1)))).#trainneg
cnf(i_0_3137, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_seq_1(X2,k2_nat_1(esk12_2(X2,X1),np__1)),k1_seq_1(X1,esk12_2(X2,X1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(X2,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v3_valued_0(X2)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_3226, plain, (k4_power(k6_numbers,X2)=k1_seq_1(k1_lp_space(esk2_0,X2),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3296, negated_conjecture, (k3_power(k6_numbers,X2)=k6_numbers|r1_xxreal_0(X2,k6_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3145, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k3_power(k6_numbers,X2),k1_funct_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2186, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k4_power(k6_numbers,X2),k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k4_power(k6_numbers,X2),k1_numbers)|~m1_subset_1(k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k1_numbers))).#trainneg
cnf(i_0_3144, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k4_power(k6_numbers,X2),k1_funct_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2530, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_seq_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k4_power(k6_numbers,X1),k3_power(k6_numbers,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2059, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k3_power(k6_numbers,X2),k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(k3_power(k6_numbers,X2))|~v1_xcmplx_0(k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1))))).#trainneg
cnf(i_0_2987, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_3387, plain, (v2_valued_0(k6_numbers))).#trainneg
cnf(i_0_3352, plain, (esk15_1(k1_zfmisc_1(k6_numbers))=k6_numbers)).#trainneg
cnf(i_0_1098, negated_conjecture, (k2_series_1(X1)=esk2_0|k2_xcmplx_0(k6_numbers,k1_funct_1(X1,k2_nat_1(esk12_2(X1,esk2_0),np__1)))!=k1_funct_1(esk2_0,k2_nat_1(esk12_2(X1,esk2_0),np__1))|k1_funct_1(X1,k6_numbers)!=k1_seq_1(esk2_0,k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m2_subset_1(esk12_2(X1,esk2_0),k1_numbers,k5_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2603, plain, (v3_valued_0(X1)|~v3_membered(k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2983, plain, (~m1_subset_1(k1_numbers,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_412, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X1,esk12_2(X2,X1)),k1_funct_1(X2,k2_nat_1(np__1,esk12_2(X2,X1))))!=k1_funct_1(X1,k2_nat_1(np__1,esk12_2(X2,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(esk12_2(X2,X1),k5_numbers))).#trainneg
cnf(i_0_415, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X1,esk12_2(X2,X1)),k1_funct_1(X2,k2_xcmplx_0(esk12_2(X2,X1),np__1)))!=k1_funct_1(X1,k2_xcmplx_0(esk12_2(X2,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(esk12_2(X2,X1),k5_numbers))).#trainneg
cnf(i_0_2518, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k2_xcmplx_0(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k4_power(k6_numbers,X1),k4_power(k6_numbers,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_xcmplx_0(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k5_numbers))).#trainneg
cnf(i_0_3225, plain, (k1_seq_1(k1_lp_space(k1_lp_space(k1_lp_space(esk2_0,X1),X2),X3),k6_numbers)=k3_power(k4_power(k6_numbers,X2),X3)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(k4_power(k6_numbers,X2))|~v1_funct_2(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2))|~r1_xxreal_0(k6_numbers,k4_power(k6_numbers,X2))|~m1_subset_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_xreal_0(X1)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3258, plain, (k1_seq_1(esk2_0,X2)=k1_seq_1(esk2_0,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2083, negated_conjecture, (X1=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|k1_seq_1(X1,esk6_3(k1_lp_space(esk2_0,X2),X3,X1))!=k4_power(k18_complex1(k1_seq_1(k1_lp_space(esk2_0,X2),X4)),X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2512, negated_conjecture, (k1_lp_space(esk2_0,X1)=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|k4_power(k6_numbers,X1)!=k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(esk6_3(k1_lp_space(esk2_0,X2),X3,k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2147, negated_conjecture, (k1_lp_space(k1_lp_space(esk2_0,X1),X2)=k1_lp_space(k1_lp_space(esk2_0,X3),X4)|k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2)!=k4_power(k18_complex1(k3_power(k6_numbers,X3)),X4)|~v1_funct_2(k1_lp_space(esk2_0,X3),k5_numbers,k1_numbers)|~v1_funct_2(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X3))|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2))|~m1_subset_1(k1_lp_space(esk2_0,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X4,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~m2_subset_1(esk6_3(k1_lp_space(esk2_0,X3),X4,k1_lp_space(k1_lp_space(esk2_0,X1),X2)),k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2169, negated_conjecture, (k3_power(k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2),X4)=k1_seq_1(k1_lp_space(k1_lp_space(k1_lp_space(esk2_0,X1),X2),X4),X3)|~v1_xreal_0(k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2))|~r1_xxreal_0(k6_numbers,k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2))|~m1_subset_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X4,k1_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3459, negated_conjecture, (k3_power(k6_numbers,X1)=k6_numbers|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3507, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k4_power(k6_numbers,X1),k3_power(k6_numbers,X2))!=k3_power(k6_numbers,X1)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2361, negated_conjecture, (v1_xxreal_0(k1_seq_2(esk2_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1190, negated_conjecture, (k2_series_1(X1)=esk2_0|k1_funct_1(X1,k2_nat_1(esk12_2(X1,esk2_0),np__1))!=k1_funct_1(esk2_0,k2_nat_1(esk12_2(X1,esk2_0),np__1))|k1_funct_1(X1,k6_numbers)!=k1_seq_1(esk2_0,k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m2_subset_1(esk12_2(X1,esk2_0),k1_numbers,k5_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2604, plain, (v1_valued_0(X1)|~v1_membered(k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2605, plain, (v2_valued_0(X1)|~v2_membered(k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2322, plain, (v1_xreal_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_419, negated_conjecture, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|~r2_hidden(X1,esk2_0))).#trainneg
cnf(i_0_423, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3428, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k4_power(k6_numbers,X2),k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(X1)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3429, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k3_power(k6_numbers,X2),k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(X1)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2513, negated_conjecture, (k1_lp_space(esk2_0,X1)=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|k3_power(k6_numbers,X1)!=k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(esk6_3(k1_lp_space(esk2_0,X2),X3,k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2168, negated_conjecture, (k4_power(k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2),X4)=k1_seq_1(k1_lp_space(k1_lp_space(k1_lp_space(esk2_0,X1),X2),X4),X3)|~v1_xreal_0(k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2))|~r1_xxreal_0(k6_numbers,k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2))|~m1_subset_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X4,k1_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2148, negated_conjecture, (X1=k1_lp_space(k1_lp_space(k1_lp_space(esk2_0,X2),X3),X4)|k1_seq_1(X1,esk6_3(k1_lp_space(k1_lp_space(esk2_0,X2),X3),X4,X1))!=k4_power(k18_complex1(k4_power(k18_complex1(k3_power(k6_numbers,X2)),X3)),X4)|~v1_funct_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v1_funct_1(X1)|~m1_subset_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X4,k1_numbers)|~m2_subset_1(esk6_3(k1_lp_space(k1_lp_space(esk2_0,X2),X3),X4,X1),k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2517, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(np__1,esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1))))!=k2_xcmplx_0(k4_power(k6_numbers,X1),k4_power(k6_numbers,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(np__1,esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1))),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k5_numbers))).#trainneg
cnf(i_0_2610, plain, (v1_funct_2(X1,X2,k6_numbers)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2362, negated_conjecture, (v3_ordinal1(k1_seq_2(esk2_0))|~v3_ordinal1(k1_numbers))).#trainneg
cnf(i_0_424, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|X2!=k1_seq_2(esk2_0)|~v2_comseq_2(esk2_0)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k6_numbers,X2)))|~r1_xxreal_0(esk3_3(esk2_0,X2,X1),X3))).#trainneg
cnf(i_0_2606, plain, (v5_relat_1(X1,k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2976, plain, (~m1_subset_1(k4_xcmplx_0(np__1),k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2979, negated_conjecture, (~m1_subset_1(k2_zfmisc_1(k5_numbers,k1_numbers),k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2462, negated_conjecture, (v1_xxreal_0(k1_seq_1(esk2_0,X1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_702, negated_conjecture, (v3_xxreal_0(esk1_0)|v2_xxreal_0(X1)|r1_xxreal_0(X1,esk1_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_432, plain, (v1_xreal_0(k3_series_1(X1))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_433, plain, (v7_ordinal1(k3_series_1(X1))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_3120, plain, (k4_power(k18_complex1(k1_seq_1(esk2_0,k6_numbers)),X2)=k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2146, negated_conjecture, (k1_lp_space(k1_lp_space(esk2_0,X1),X2)=k1_lp_space(k1_lp_space(esk2_0,X3),X4)|k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2)!=k4_power(k18_complex1(k4_power(k6_numbers,X3)),X4)|~v1_funct_2(k1_lp_space(esk2_0,X3),k5_numbers,k1_numbers)|~v1_funct_2(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X3))|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2))|~m1_subset_1(k1_lp_space(esk2_0,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X4,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~m2_subset_1(esk6_3(k1_lp_space(esk2_0,X3),X4,k1_lp_space(k1_lp_space(esk2_0,X1),X2)),k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3304, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X4,k6_numbers)|k2_xcmplx_0(k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k3_power(k6_numbers,X4))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_3788, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X2),X3),k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2298, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k6_numbers|k2_xcmplx_0(k6_numbers,k4_power(k6_numbers,X1))!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~v1_xreal_0(k4_power(k6_numbers,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k4_power(k6_numbers,X1),k5_numbers))).#trainneg
cnf(i_0_434, plain, (v1_xcmplx_0(k3_series_1(X1))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_3724, negated_conjecture, (~v3_xxreal_0(k1_seq_2(esk2_0))|~v1_xboole_0(k1_seq_2(esk2_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_3725, negated_conjecture, (~v3_xxreal_0(k1_seq_2(esk2_0))|~v2_xxreal_0(k1_seq_2(esk2_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_716, plain, (v3_xxreal_0(np__1)|v2_xxreal_0(X1)|r1_xxreal_0(X1,np__1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_435, plain, (v1_xxreal_0(k3_series_1(X1))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_436, plain, (v3_ordinal1(k3_series_1(X1))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_3718, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|k2_xcmplx_0(k4_power(k6_numbers,X1),k3_power(k6_numbers,X2))!=k3_power(k6_numbers,X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_xcmplx_0(np__1,esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1))),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v7_ordinal1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1))))).#trainneg
cnf(i_0_3802, plain, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k2_xcmplx_0(k1_seq_1(esk2_0,k6_numbers),k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3796, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k3_power(k6_numbers,X2),k3_power(k6_numbers,X1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3951, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X3,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|k2_xcmplx_0(k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k3_power(k6_numbers,X3))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1))).#trainneg
cnf(i_0_3783, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k4_power(k6_numbers,X2),k3_power(k6_numbers,X1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3795, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k3_power(k6_numbers,X2),k4_power(k6_numbers,X1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3790, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k2_xcmplx_0(k1_seq_1(esk2_0,X3),k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3782, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k4_power(k6_numbers,X2),k4_power(k6_numbers,X1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3884, plain, (X1=k1_lp_space(k1_lp_space(esk2_0,X3),X4)|r1_xxreal_0(X2,k6_numbers)|k1_seq_1(X1,esk6_3(k1_lp_space(esk2_0,X3),X4,X1))!=k4_power(k18_complex1(k4_power(k18_complex1(k3_power(k6_numbers,X2)),X3)),X4)|~m2_subset_1(esk6_3(k1_lp_space(esk2_0,X3),X4,X1),k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X4,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2985, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2299, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k6_numbers|k2_xcmplx_0(k6_numbers,k4_power(k6_numbers,X1))!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~v1_xreal_0(k4_power(k6_numbers,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k4_power(k6_numbers,X1),k1_numbers))).#trainneg
cnf(i_0_437, plain, (v3_membered(k3_series_1(X1))|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_3742, plain, (v1_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers))|~v3_membered(k6_numbers))).#trainneg
cnf(i_0_785, plain, (v3_xxreal_0(k6_numbers)|v2_xxreal_0(X1)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2609, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_439, plain, (v1_membered(k3_series_1(X1))|~v1_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_3723, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|k2_xcmplx_0(k4_power(k6_numbers,X1),k3_power(k6_numbers,X2))!=k3_power(k6_numbers,X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k5_numbers))).#trainneg
cnf(i_0_2357, negated_conjecture, (m1_subset_1(k1_seq_2(esk2_0),k1_numbers))).#trainneg
cnf(i_0_2540, negated_conjecture, (k1_lp_space(esk2_0,X1)=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|k3_power(k6_numbers,X1)!=k4_power(k18_complex1(k3_power(k6_numbers,X2)),X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2720, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|r1_xxreal_0(esk4_2(k1_lp_space(esk2_0,X2),X1),k18_complex1(k6_xcmplx_0(k4_power(k18_complex1(k3_power(k6_numbers,X4)),X2),X1)))|r1_xxreal_0(X4,k6_numbers)|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k4_power(k18_complex1(k3_power(k6_numbers,X4)),X2),k1_numbers)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(X4)|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3596, plain, (k1_seq_1(esk2_0,X1)=k1_seq_1(esk2_0,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1296, plain, (X1=k1_seq_2(X2)|k9_real_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1)!=k6_numbers|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),k1_numbers))).#trainneg
cnf(i_0_3111, negated_conjecture, (k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k6_numbers)=k1_seq_1(k1_lp_space(esk2_0,X2),X3)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4209, plain, (X1=k1_seq_2(X2)|k6_xcmplx_0(k1_seq_1(X2,esk5_3(X2,X1,X3)),X1)!=k6_numbers|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(k1_seq_1(X2,esk5_3(X2,X1,X3)),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2989, negated_conjecture, (r1_xxreal_0(X1,esk1_0)|r1_xxreal_0(esk1_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2945, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2972, negated_conjecture, (v1_xxreal_0(k1_seq_2(esk2_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_440, plain, (v2_membered(k3_series_1(X1))|~v2_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_441, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|~r2_hidden(X1,k3_series_1(X2))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_446, plain, (v3_valued_0(esk15_1(k1_zfmisc_1(X1)))|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_4286, negated_conjecture, (k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X3),X1),k6_numbers)=k3_power(k6_numbers,X1)|r1_xxreal_0(X3,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X3)|~v1_funct_2(k1_lp_space(esk2_0,X3),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X3))|~m1_subset_1(k1_lp_space(esk2_0,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4020, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|k2_xcmplx_0(k4_power(k6_numbers,X1),k3_power(k6_numbers,X2))!=k3_power(k6_numbers,X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v7_ordinal1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)))|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4204, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|k9_real_1(k3_power(k6_numbers,X2),X1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k3_power(k6_numbers,X2),k1_numbers)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4203, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|k9_real_1(k4_power(k6_numbers,X2),X1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k4_power(k6_numbers,X2),k1_numbers)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4291, negated_conjecture, (k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2)=k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X4),X2),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X4,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X4)|~v1_funct_2(k1_lp_space(esk2_0,X4),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X4))|~m1_subset_1(k1_lp_space(esk2_0,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_755, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1548, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~r1_xxreal_0(k6_numbers,X1))).#trainneg
cnf(i_0_3185, plain, (m1_subset_1(k1_seq_1(k6_numbers,X1),k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_450, plain, (v1_xreal_0(k4_power(X1,X2))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4303, plain, (k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k6_numbers)=k1_seq_1(esk2_0,X3)|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4150, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|k2_xcmplx_0(k4_power(k6_numbers,X1),k3_power(k6_numbers,X2))!=k3_power(k6_numbers,X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k5_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2129, negated_conjecture, (X1=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|k1_seq_1(X1,esk6_3(k1_lp_space(esk2_0,X2),X3,X1))!=k3_power(k18_complex1(k3_power(k6_numbers,X2)),X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k18_complex1(k3_power(k6_numbers,X2)),k1_numbers))).#trainneg
cnf(i_0_3108, negated_conjecture, (k3_power(k18_complex1(k3_power(k6_numbers,X1)),X2)=k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k6_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k18_complex1(k3_power(k6_numbers,X1)),k1_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1146, negated_conjecture, (X1=k2_series_1(esk2_0)|k1_funct_1(X1,k2_nat_1(esk12_2(esk2_0,X1),np__1))!=k2_xcmplx_0(k1_seq_1(X1,esk12_2(esk2_0,X1)),k6_numbers)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(esk2_0,k6_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1)|~m2_subset_1(k2_nat_1(esk12_2(esk2_0,X1),np__1),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2087, negated_conjecture, (X1=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|k1_seq_1(X1,esk6_3(k1_lp_space(esk2_0,X2),X3,X1))!=k3_power(k18_complex1(k4_power(k6_numbers,X2)),X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k18_complex1(k4_power(k6_numbers,X2)),k1_numbers))).#trainneg
cnf(i_0_4316, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))!=k2_xcmplx_0(k6_numbers,k1_funct_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4317, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))!=k2_xcmplx_0(k6_numbers,k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4318, negated_conjecture, (k1_seq_1(k1_lp_space(esk2_0,X1),X2)=k6_numbers|r1_xxreal_0(X1,k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_756, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_451, plain, (v7_ordinal1(k4_power(X1,X2))|~v6_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4408, plain, (v1_xcmplx_0(k1_seq_1(k6_numbers,X1))|~v1_membered(k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_4342, negated_conjecture, (~v3_xxreal_0(k1_seq_2(esk2_0))|~v1_xboole_0(k1_seq_2(esk2_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_4343, negated_conjecture, (~v3_xxreal_0(k1_seq_2(esk2_0))|~v2_xxreal_0(k1_seq_2(esk2_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_4406, plain, (v1_xreal_0(k1_seq_1(k6_numbers,X1))|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_4341, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_452, plain, (v1_xcmplx_0(k4_power(X1,X2))|~v1_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_453, plain, (v1_xxreal_0(k4_power(X1,X2))|~v2_membered(k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4329, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|k6_xcmplx_0(k3_power(k6_numbers,X2),X1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k3_power(k6_numbers,X2),k1_numbers)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4440, plain, (v1_xreal_0(k3_power(X1,X2))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4328, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|k6_xcmplx_0(k4_power(k6_numbers,X2),X1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k4_power(k6_numbers,X2),k1_numbers)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4420, plain, (r1_xxreal_0(k4_power(X1,X2),k6_numbers)|r1_xxreal_0(k6_numbers,k4_power(X1,X2))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4356, negated_conjecture, (k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k6_numbers)=k3_power(k6_numbers,X2)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4568, plain, (X1=k1_lp_space(esk2_0,X3)|r1_xxreal_0(X2,k6_numbers)|k1_seq_1(X1,esk6_3(esk2_0,X3,X1))!=k3_power(k18_complex1(k4_power(k6_numbers,X2)),X3)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k18_complex1(k4_power(k6_numbers,X2)),k1_numbers)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_757, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_454, plain, (v3_ordinal1(k4_power(X1,X2))|~v3_ordinal1(k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1191, negated_conjecture, (v1_xreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1967, negated_conjecture, (v1_xcmplx_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4407, plain, (v7_ordinal1(k1_seq_1(k6_numbers,X1))|~v6_membered(k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_4409, plain, (v1_xxreal_0(k1_seq_1(k6_numbers,X1))|~v2_membered(k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_455, plain, (m1_subset_1(k3_power(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_460, plain, (r1_xxreal_0(X1,k6_numbers)|X2!=k1_seq_2(X3)|~v2_comseq_2(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(X3,k5_numbers,k1_numbers)|~v1_funct_1(X3)|~r1_xxreal_0(X1,k18_complex1(k6_xcmplx_0(k1_seq_1(X3,X4),X2)))|~r1_xxreal_0(esk3_3(X3,X2,X1),X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X3,X4),k1_numbers))).#trainneg
cnf(i_0_4655, negated_conjecture, (k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2)=k6_numbers|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4597, plain, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k1_seq_1(esk2_0,k6_numbers),k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4594, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k3_power(k6_numbers,X2),k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4646, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|k2_xcmplx_0(k4_power(k6_numbers,X1),k4_power(k6_numbers,X2))!=k6_numbers|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4593, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k4_power(k6_numbers,X2),k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4637, plain, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|k2_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X1),esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1))),k6_numbers)!=k6_numbers|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4647, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|k2_xcmplx_0(k4_power(k6_numbers,X1),k3_power(k6_numbers,X2))!=k6_numbers|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4405, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4769, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k6_numbers|k2_xcmplx_0(k4_power(k6_numbers,X1),k6_numbers)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k4_power(k6_numbers,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(k4_power(k6_numbers,X1)))).#trainneg
cnf(i_0_463, plain, (v1_xreal_0(k9_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1968, negated_conjecture, (v1_xxreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_497, negated_conjecture, (v2_valued_0(esk2_0)|~v2_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_4677, plain, (v1_xcmplx_0(k1_seq_1(k6_numbers,X1))|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_4695, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_464, plain, (v7_ordinal1(k9_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_465, plain, (v1_xcmplx_0(k9_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4630, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k9_real_1(k6_numbers,X1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(X2)|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4620, negated_conjecture, (X1=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|r1_xxreal_0(X2,k6_numbers)|k1_seq_1(X1,esk6_3(k1_lp_space(esk2_0,X2),X3,X1))!=k3_power(k6_numbers,X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_xreal_0(X2)|~m2_subset_1(esk6_3(k1_lp_space(esk2_0,X2),X3,X1),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4588, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k6_numbers,k4_power(k6_numbers,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4603, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(X2)|r1_xxreal_0(X1,k6_numbers)|k2_xcmplx_0(k1_seq_1(X2,k2_nat_1(esk12_2(X2,k1_lp_space(esk2_0,X1)),np__1)),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(X2,k1_lp_space(esk2_0,X1)),np__1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(X2)|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(X2,k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4590, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|k2_xcmplx_0(k6_numbers,k1_seq_1(k1_lp_space(esk2_0,X2),X3))!=k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4659, plain, (v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1612, plain, (esk15_1(k1_zfmisc_1(X1))=k6_numbers|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3924, negated_conjecture, (r1_xxreal_0(k4_xcmplx_0(np__1),esk1_0))).#trainneg
cnf(i_0_4766, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k4_xcmplx_0(np__1)|k2_xcmplx_0(k4_power(k6_numbers,X1),np__1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k4_power(k6_numbers,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_power(k6_numbers,X1)))).#trainneg
cnf(i_0_466, plain, (v1_xxreal_0(k9_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5005, negated_conjecture, (v2_valued_0(esk2_0))).#trainneg
cnf(i_0_4998, negated_conjecture, (~v3_xxreal_0(esk2_0)|~v2_xxreal_0(esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5098, plain, (m1_subset_1(k6_numbers,k1_zfmisc_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5097, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_467, plain, (v3_ordinal1(k9_real_1(X1,X2))|~v3_ordinal1(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_468, plain, (m1_subset_1(k6_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4892, negated_conjecture, (k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2)=k6_numbers|r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4942, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|k1_seq_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k4_power(k6_numbers,X2),k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4916, plain, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|k1_seq_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k1_seq_1(esk2_0,k6_numbers),k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1)))).#trainneg
cnf(i_0_4924, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|k1_seq_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k3_power(k6_numbers,X2),k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_5059, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(X2)|r1_xxreal_0(X1,k6_numbers)|k2_xcmplx_0(k1_seq_1(X2,k2_nat_1(np__1,esk12_2(X2,k1_lp_space(esk2_0,X1)))),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(np__1,esk12_2(X2,k1_lp_space(esk2_0,X1))))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v1_relat_1(X2)|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v3_valued_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(X2,k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(X2)|~v7_ordinal1(esk12_2(X2,k1_lp_space(esk2_0,X1))))).#trainneg
cnf(i_0_4946, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|k2_xcmplx_0(k4_power(k6_numbers,X1),k6_numbers)!=k6_numbers|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4595, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|k2_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X2),X3),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4587, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k1_seq_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))!=k2_xcmplx_0(k6_numbers,k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1))).#trainneg
cnf(i_0_5115, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k6_numbers)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2292, plain, (k2_zfmisc_1(X1,esk15_1(k1_zfmisc_1(X2)))=k6_numbers|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4001, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),np__1))).#trainneg
cnf(i_0_4974, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k6_numbers|k2_xcmplx_0(k6_numbers,k4_power(k6_numbers,X1))!=k6_numbers|~v1_xcmplx_0(k4_power(k6_numbers,X1))|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k4_power(k6_numbers,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_469, plain, (v1_xreal_0(k7_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1870, plain, (r1_xxreal_0(k5_numbers,k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2347, negated_conjecture, (v1_xxreal_0(k1_seq_1(esk2_0,k6_numbers))|~v2_valued_0(esk2_0))).#trainneg
cnf(i_0_5110, plain, (v3_membered(k6_numbers)|~v3_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5112, plain, (v1_membered(k6_numbers)|~v1_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_470, plain, (v7_ordinal1(k7_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_471, plain, (v1_xcmplx_0(k7_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5061, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(X2)|r1_xxreal_0(X1,k6_numbers)|k2_xcmplx_0(k1_seq_1(X2,k2_nat_1(esk12_2(X2,k1_lp_space(esk2_0,X1)),np__1)),k6_numbers)!=k1_seq_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(X2,k1_lp_space(esk2_0,X1)),np__1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v1_relat_1(X2)|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v3_valued_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(X2,k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(X2))).#trainneg
cnf(i_0_5369, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk2_0,k6_numbers),k1_seq_1(esk2_0,k6_numbers)))).#trainneg
cnf(i_0_5427, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X3,k6_numbers)|k2_xcmplx_0(k4_power(k18_complex1(k3_power(k6_numbers,X3)),X2),k6_numbers)!=k1_seq_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X3)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_5370, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_seq_1(esk2_0,k6_numbers))|r1_xxreal_0(k1_seq_1(esk2_0,k6_numbers),k6_numbers))).#trainneg
cnf(i_0_5184, negated_conjecture, (X1=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X3,k6_numbers)|k1_seq_1(X1,esk6_3(k1_lp_space(esk2_0,X2),X3,X1))!=k6_numbers|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_4773, plain, (r1_xxreal_0(k6_numbers,k4_power(X1,X2))|~v2_xxreal_0(k4_power(X1,X2))|~v1_xreal_0(k4_power(X1,X2))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4589, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k6_numbers,k3_power(k6_numbers,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5353, plain, (r1_xxreal_0(k7_real_1(X1,X2),k6_numbers)|r1_xxreal_0(k6_numbers,k7_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5147, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_472, plain, (v1_xxreal_0(k7_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5367, negated_conjecture, (~v3_xxreal_0(k1_seq_1(esk2_0,k6_numbers))|~v1_xboole_0(k1_seq_1(esk2_0,k6_numbers)))).#trainneg
cnf(i_0_5368, negated_conjecture, (~v3_xxreal_0(k1_seq_1(esk2_0,k6_numbers))|~v2_xxreal_0(k1_seq_1(esk2_0,k6_numbers)))).#trainneg
cnf(i_0_5113, plain, (v2_membered(k6_numbers)|~v2_membered(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5386, plain, (v1_membered(k6_numbers)|~v1_xboole_0(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_473, plain, (v3_ordinal1(k7_real_1(X1,X2))|~v3_ordinal1(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_475, plain, (m1_subset_1(k7_real_1(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5056, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(X2)|r1_xxreal_0(X1,k6_numbers)|k2_xcmplx_0(k1_seq_1(X2,k2_nat_1(np__1,esk12_2(X2,k1_lp_space(esk2_0,X1)))),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(np__1,esk12_2(X2,k1_lp_space(esk2_0,X1))))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v1_relat_1(X2)|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v3_valued_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(X2,k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(X2)|~m1_subset_1(esk12_2(X2,k1_lp_space(esk2_0,X1)),k5_numbers))).#trainneg
cnf(i_0_5058, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(X2)|r1_xxreal_0(X1,k6_numbers)|k2_xcmplx_0(k1_seq_1(X2,k2_xcmplx_0(np__1,esk12_2(X2,k1_lp_space(esk2_0,X1)))),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X1),k2_xcmplx_0(np__1,esk12_2(X2,k1_lp_space(esk2_0,X1))))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v1_relat_1(X2)|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v3_valued_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(X2,k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(X2)|~v7_ordinal1(esk12_2(X2,k1_lp_space(esk2_0,X1))))).#trainneg
cnf(i_0_4581, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k2_xcmplx_0(k6_numbers,k1_funct_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)))!=k1_seq_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_4615, negated_conjecture, (k1_lp_space(esk2_0,X1)=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|r1_xxreal_0(X1,k6_numbers)|k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),X4)!=k6_numbers|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(esk6_3(k1_lp_space(esk2_0,X2),X3,k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4600, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(X2)|r1_xxreal_0(X1,k6_numbers)|k2_xcmplx_0(k6_numbers,k1_funct_1(X2,k2_xcmplx_0(np__1,esk12_2(X2,k1_lp_space(esk2_0,X1)))))!=k1_funct_1(k1_lp_space(esk2_0,X1),k2_xcmplx_0(np__1,esk12_2(X2,k1_lp_space(esk2_0,X1))))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v7_ordinal1(esk12_2(X2,k1_lp_space(esk2_0,X1)))|~v1_funct_1(X2)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(X2,k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4792, negated_conjecture, (k1_seq_1(k1_lp_space(k1_lp_space(k1_lp_space(esk2_0,X1),X2),X3),k6_numbers)=k3_power(k3_power(k6_numbers,X2),X3)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(k3_power(k6_numbers,X2))|~v1_funct_2(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2))|~r1_xxreal_0(k6_numbers,k3_power(k6_numbers,X2))|~m1_subset_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_xreal_0(X1)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4947, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|k2_xcmplx_0(k3_power(k6_numbers,X1),k6_numbers)!=k6_numbers|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5179, negated_conjecture, (k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),X3)=k6_numbers|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_5149, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5551, plain, (v2_membered(k6_numbers)|~v1_xboole_0(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_477, plain, (v1_xboole_0(esk15_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_480, plain, (v1_xboole_0(esk15_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5525, plain, (r1_xxreal_0(k6_numbers,k2_xcmplx_0(X2,X1))|r1_xxreal_0(k2_xcmplx_0(X2,X1),k6_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_5706, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X3,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|k1_seq_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k6_numbers,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_partfun1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v1_relat_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_5481, plain, (X1=k1_lp_space(esk2_0,X3)|r1_xxreal_0(X3,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|k1_seq_1(X1,esk6_3(esk2_0,X3,X1))!=k6_numbers|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_5479, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_seq_1(esk2_0,k6_numbers))|~v2_xxreal_0(k1_seq_1(esk2_0,k6_numbers)))).#trainneg
cnf(i_0_5527, plain, (r1_xxreal_0(k6_numbers,k2_xcmplx_0(X1,X2))|r1_xxreal_0(k2_xcmplx_0(X1,X2),k6_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4958, plain, (k2_series_1(k1_lp_space(esk2_0,X2))=esk2_0|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)!=k1_seq_1(esk2_0,k6_numbers)|k2_xcmplx_0(k4_power(k6_numbers,X1),k3_power(k6_numbers,X2))!=k6_numbers|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),esk2_0),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),esk2_0),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2294, plain, (k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3))=k6_numbers|~v1_xboole_0(X3))).#trainneg
cnf(i_0_483, plain, (v3_valued_0(esk15_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v3_membered(X2))).#trainneg
cnf(i_0_4410, plain, (v3_ordinal1(k1_seq_1(k6_numbers,X1))|~v3_ordinal1(k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_4830, plain, (v1_xxreal_0(k1_seq_1(k6_numbers,X1))|~v6_membered(k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_486, plain, (v1_valued_0(esk15_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_membered(X2))).#trainneg
cnf(i_0_488, plain, (v1_valued_0(k3_series_1(X1))|~v1_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4531, negated_conjecture, (X1=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|k1_seq_1(X1,esk6_3(k1_lp_space(esk2_0,X2),X3,X1))!=k3_power(k3_power(k6_numbers,X2),X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_power(k6_numbers,X2),k1_numbers)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k3_power(k6_numbers,X2))|~r1_xxreal_0(k6_numbers,k3_power(k6_numbers,X2)))).#trainneg
cnf(i_0_5183, negated_conjecture, (k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k6_numbers)=k6_numbers|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4566, negated_conjecture, (X1=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|k1_seq_1(X1,esk6_3(k1_lp_space(esk2_0,X2),X3,X1))!=k3_power(k4_power(k6_numbers,X2),X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k4_power(k6_numbers,X2),k1_numbers)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k4_power(k6_numbers,X2))|~r1_xxreal_0(k6_numbers,k4_power(k6_numbers,X2)))).#trainneg
cnf(i_0_3222, plain, (k4_power(k18_complex1(k4_power(k6_numbers,X2)),X3)=k1_seq_1(k1_lp_space(k1_lp_space(k1_lp_space(esk2_0,X1),X2),X3),k6_numbers)|r1_xxreal_0(X1,k6_numbers)|~v1_funct_2(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2))|~m1_subset_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_xreal_0(X1)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2145, negated_conjecture, (k1_lp_space(k1_lp_space(esk2_0,X1),X2)=k1_lp_space(esk2_0,X3)|k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2)!=k4_power(k6_numbers,X3)|~m2_subset_1(esk6_3(esk2_0,X3,k1_lp_space(k1_lp_space(esk2_0,X1),X2)),k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2))|~m1_subset_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2136, plain, (r1_xxreal_0(k6_numbers,k18_complex1(X1))|r1_xxreal_0(k18_complex1(X1),k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2328, plain, (k2_zfmisc_1(X1,k2_zfmisc_1(X2,X3))=k6_numbers|~v1_xboole_0(X2))).#trainneg
cnf(i_0_489, plain, (v2_valued_0(esk15_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v2_membered(X2))).#trainneg
cnf(i_0_4999, negated_conjecture, (r1_xxreal_0(esk2_0,esk2_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_725, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3775, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_491, plain, (v2_valued_0(k3_series_1(X1))|~v2_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_492, plain, (v1_valued_0(esk15_1(k1_zfmisc_1(X1)))|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5969, plain, (r1_xxreal_0(k6_numbers,k18_complex1(X1))|~v2_xxreal_0(k18_complex1(X1))|~v1_xreal_0(k18_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5917, plain, (X1=k1_lp_space(esk2_0,X3)|r1_xxreal_0(X2,k6_numbers)|k1_seq_1(X1,esk6_3(esk2_0,X3,X1))!=k3_power(k4_power(k6_numbers,X2),X3)|~v1_xreal_0(k4_power(k6_numbers,X2))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~r1_xxreal_0(k6_numbers,k4_power(k6_numbers,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_5890, plain, (X1=k1_lp_space(esk2_0,X3)|r1_xxreal_0(X2,k6_numbers)|k1_seq_1(X1,esk6_3(esk2_0,X3,X1))!=k3_power(k3_power(k6_numbers,X2),X3)|~v1_xreal_0(k3_power(k6_numbers,X2))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~r1_xxreal_0(k6_numbers,k3_power(k6_numbers,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_5965, negated_conjecture, (k1_lp_space(k1_lp_space(esk2_0,X1),X2)=k1_lp_space(esk2_0,X3)|k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2)!=k4_power(k6_numbers,X3)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_funct_2(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2))|~m1_subset_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_5763, plain, (r1_xxreal_0(k6_numbers,k2_xcmplx_0(X1,X2))|~v2_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_5950, negated_conjecture, (X1=k1_lp_space(esk2_0,X2)|r1_xxreal_0(X3,k6_numbers)|r1_xxreal_0(X4,k6_numbers)|k1_seq_1(X1,esk6_3(esk2_0,X2,X1))!=k1_seq_1(k1_lp_space(k1_lp_space(k1_lp_space(esk2_0,X4),X3),X2),k6_numbers)|~v1_xreal_0(X3)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X4)|~v1_funct_2(k1_lp_space(esk2_0,X4),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X4))|~m1_subset_1(k1_lp_space(esk2_0,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_5812, plain, (r1_xxreal_0(k6_numbers,k2_xcmplx_0(X1,X2))|~v2_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_494, plain, (v1_valued_0(k3_series_1(X1))|~v1_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_496, plain, (v2_valued_0(esk15_1(k1_zfmisc_1(X1)))|~v2_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_498, plain, (v2_valued_0(k3_series_1(X1))|~v2_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_5949, negated_conjecture, (X1=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|r1_xxreal_0(X4,k6_numbers)|k1_seq_1(X1,esk6_3(k1_lp_space(esk2_0,X2),X3,X1))!=k1_seq_1(k1_lp_space(k1_lp_space(k1_lp_space(esk2_0,X4),X2),X3),k6_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X4)|~v1_funct_2(k1_lp_space(esk2_0,X4),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X4))|~m1_subset_1(k1_lp_space(esk2_0,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2547, negated_conjecture, (k1_lp_space(esk2_0,X1)=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|k3_power(k6_numbers,X1)!=k4_power(k18_complex1(k4_power(k6_numbers,X2)),X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5069, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k6_numbers,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4206, negated_conjecture, (X1=k1_seq_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|k9_real_1(k4_power(k18_complex1(k3_power(k6_numbers,X2)),X3),X1)!=k6_numbers|~v2_comseq_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v1_xreal_0(X1)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~m1_subset_1(k4_power(k18_complex1(k3_power(k6_numbers,X2)),X3),k1_numbers)|~m1_subset_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk5_3(k1_lp_space(k1_lp_space(esk2_0,X2),X3),X1,X4),k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3150, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X1,esk12_2(X2,X1)),k1_seq_1(X2,k2_nat_1(esk12_2(X2,X1),np__1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(X2,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v3_valued_0(X2)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk12_2(X2,X1)),k1_numbers)|~m1_subset_1(k1_seq_1(X2,k2_nat_1(esk12_2(X2,X1),np__1)),k1_numbers))).#trainneg
cnf(i_0_3134, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k6_numbers|r1_xxreal_0(X2,k6_numbers)|k2_xcmplx_0(k4_power(k18_complex1(k3_power(k6_numbers,X2)),X1),k6_numbers)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~v1_xreal_0(k4_power(k18_complex1(k3_power(k6_numbers,X2)),X1))|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(k4_power(k18_complex1(k3_power(k6_numbers,X2)),X1)))).#trainneg
cnf(i_0_864, plain, (v1_valued_0(k2_series_1(X1))|~v1_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_5125, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k4_xcmplx_0(np__1)|k2_xcmplx_0(np__1,k4_power(k6_numbers,X1))!=k6_numbers|~v1_xcmplx_0(k4_power(k6_numbers,X1))|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k4_power(k6_numbers,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_504, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v2_comseq_2(esk2_0)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k6_numbers,esk7_1(esk2_0))))|~r1_xxreal_0(esk8_2(esk2_0,X1),X2))).#trainneg
cnf(i_0_2218, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_3776, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_505, plain, (r1_xxreal_0(X1,k6_numbers)|~v2_comseq_2(X2)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~r1_xxreal_0(X1,k18_complex1(k6_xcmplx_0(k1_seq_1(X2,X3),esk7_1(X2))))|~r1_xxreal_0(esk8_2(X2,X1),X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(esk7_1(X2))|~m1_subset_1(k1_seq_1(X2,X3),k1_numbers))).#trainneg
cnf(i_0_507, plain, (X1=k6_numbers|r1_xxreal_0(X2,k6_numbers)|X1!=k3_power(X3,X2)|X3!=k6_numbers|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_3310, plain, (k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k6_numbers)=k4_power(k1_seq_1(esk2_0,k6_numbers),X2)|r1_xxreal_0(X1,k6_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~r1_xxreal_0(k6_numbers,k1_seq_1(esk2_0,k6_numbers))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6122, negated_conjecture, (k1_lp_space(k1_lp_space(esk2_0,X1),X2)=k1_lp_space(esk2_0,X3)|k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2)!=k3_power(k6_numbers,X3)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6252, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k1_funct_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k3_power(k6_numbers,X2))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k1_numbers)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6251, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k1_funct_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k4_power(k6_numbers,X2))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k1_numbers)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6127, plain, (k1_lp_space(k1_lp_space(esk2_0,X1),X2)=k1_lp_space(esk2_0,X3)|r1_xxreal_0(X1,k6_numbers)|k4_power(k6_numbers,X2)!=k4_power(k6_numbers,X3)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6129, plain, (k1_lp_space(k1_lp_space(esk2_0,X1),X2)=k1_lp_space(esk2_0,X3)|r1_xxreal_0(X1,k6_numbers)|k4_power(k18_complex1(k1_seq_1(esk2_0,k6_numbers)),X2)!=k4_power(k6_numbers,X3)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6123, negated_conjecture, (k1_lp_space(k1_lp_space(esk2_0,X1),X2)=k1_lp_space(esk2_0,X3)|k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2)!=k1_seq_1(k1_lp_space(esk2_0,X3),X4)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6349, negated_conjecture, (X1=k6_numbers|r1_xxreal_0(X2,k6_numbers)|X1!=k1_seq_1(esk2_0,X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6194, plain, (v1_valued_0(k2_series_1(X1))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_508, plain, (v3_xxreal_0(esk7_1(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,esk7_1(X1))|~v1_xreal_0(X2)|~v2_comseq_2(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2386, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|k4_power(k1_seq_2(esk2_0),k7_xcmplx_0(np__1,esk1_0))!=k6_numbers|~v1_series_1(esk2_0))).#trainneg
cnf(i_0_2965, plain, (r1_xxreal_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_6043, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2108, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_509, plain, (m1_subset_1(k1_seq_1(k2_zfmisc_1(X1,X2),X3),k1_numbers)|~v1_relat_1(k2_zfmisc_1(X1,X2))|~v1_funct_1(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_515, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X1,esk12_2(X2,X1)),k1_seq_1(X2,k2_nat_1(esk12_2(X2,X1),np__1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(X2,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v3_valued_0(X2))).#trainneg
cnf(i_0_1155, plain, (m1_subset_1(k6_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5642, plain, (k2_series_1(k1_lp_space(esk2_0,X2))=esk2_0|r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)!=k1_seq_1(esk2_0,k6_numbers)|k2_xcmplx_0(k3_power(k6_numbers,X1),k6_numbers)!=k6_numbers|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),esk2_0),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),esk2_0),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2)))).#trainneg
cnf(i_0_5953, plain, (k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2)=k1_seq_1(k1_lp_space(k1_lp_space(k1_lp_space(esk2_0,X3),X1),X2),k6_numbers)|r1_xxreal_0(X3,k6_numbers)|~v1_xreal_0(X3)|~v1_funct_2(k1_lp_space(esk2_0,X3),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X3))|~m1_subset_1(k1_lp_space(esk2_0,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4599, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(X2)|r1_xxreal_0(X1,k6_numbers)|k2_xcmplx_0(k6_numbers,k1_funct_1(X2,k2_nat_1(esk12_2(X2,k1_lp_space(esk2_0,X1)),np__1)))!=k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(X2,k1_lp_space(esk2_0,X1)),np__1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_funct_1(X2)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(X2,k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4601, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(X2)|r1_xxreal_0(X1,k6_numbers)|k2_xcmplx_0(k6_numbers,k1_funct_1(X2,k2_nat_1(np__1,esk12_2(X2,k1_lp_space(esk2_0,X1)))))!=k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(np__1,esk12_2(X2,k1_lp_space(esk2_0,X1))))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v7_ordinal1(esk12_2(X2,k1_lp_space(esk2_0,X1)))|~v1_funct_1(X2)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(X2,k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6238, plain, (X1=k1_seq_2(k1_lp_space(esk2_0,X3))|r1_xxreal_0(X2,k6_numbers)|k9_real_1(k4_power(k18_complex1(k3_power(k6_numbers,X2)),X3),X1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X3))|~v1_xreal_0(X1)|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X3),X1,X4),k1_numbers,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(k4_power(k18_complex1(k3_power(k6_numbers,X2)),X3),k1_numbers)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4791, negated_conjecture, (k1_seq_1(k1_lp_space(k1_lp_space(k1_lp_space(esk2_0,X1),X2),X3),k6_numbers)=k4_power(k3_power(k6_numbers,X2),X3)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(k3_power(k6_numbers,X2))|~v1_funct_2(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2))|~r1_xxreal_0(k6_numbers,k3_power(k6_numbers,X2))|~m1_subset_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_xreal_0(X1)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_865, plain, (v2_valued_0(k2_series_1(X1))|~v2_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1602, plain, (r1_xxreal_0(esk15_1(k5_numbers),np__1)|r1_xxreal_0(np__1,esk15_1(k5_numbers)))).#trainneg
cnf(i_0_520, plain, (v3_xxreal_0(k1_funct_1(X1,X2))|v2_xxreal_0(X3)|r1_xxreal_0(X3,k1_funct_1(X1,X2))|~v1_xreal_0(X3)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5381, plain, (v3_membered(k6_numbers)|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3726, negated_conjecture, (r1_xxreal_0(k1_seq_2(esk2_0),k1_seq_2(esk2_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_521, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_522, plain, (v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_6481, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k1_funct_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k3_power(k6_numbers,X2))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6482, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))!=k2_xcmplx_0(k1_funct_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k6_numbers)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6480, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k1_funct_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k4_power(k6_numbers,X2))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6532, plain, (k4_power(k18_complex1(k4_power(k6_numbers,X1)),X2)=k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2)|r1_xxreal_0(X3,k6_numbers)|~v1_xreal_0(X3)|~v1_funct_2(k1_lp_space(esk2_0,X3),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X3))|~m1_subset_1(k1_lp_space(esk2_0,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4179, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|r1_xxreal_0(esk4_2(k1_lp_space(esk2_0,X2),X1),k18_complex1(k6_xcmplx_0(k4_power(k18_complex1(k3_power(k6_numbers,X3)),X2),X1)))|r1_xxreal_0(X3,k6_numbers)|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k4_power(k18_complex1(k3_power(k6_numbers,X3)),X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6641, plain, (v3_membered(k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_1732, plain, (v7_ordinal1(k1_seq_2(X1))|~v6_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_527, plain, (r1_xxreal_0(X1,k1_funct_1(X2,X3))|r1_xxreal_0(k1_funct_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v2_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_5557, plain, (v1_membered(k6_numbers)|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4344, negated_conjecture, (r1_xxreal_0(k1_seq_2(esk2_0),k1_seq_2(esk2_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2327, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2944, plain, (v2_valued_0(X1)|~v2_membered(X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_528, plain, (v1_xxreal_0(k1_seq_1(X1,X2))|~v2_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_529, plain, (v1_xreal_0(k6_xcmplx_0(X2,X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2534, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k2_xcmplx_0(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k3_power(k6_numbers,X1),k4_power(k6_numbers,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_xcmplx_0(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k5_numbers))).#trainneg
cnf(i_0_2423, plain, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k1_funct_1(X1,k2_xcmplx_0(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))!=k2_xcmplx_0(k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k6_numbers)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(k2_xcmplx_0(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(esk12_2(k1_lp_space(esk2_0,X2),X1),k5_numbers))).#trainneg
cnf(i_0_4596, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|k2_xcmplx_0(k1_seq_1(esk2_0,X3),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1709, negated_conjecture, (k1_seq_1(k1_lp_space(esk2_0,X1),X3)=k1_seq_1(k1_lp_space(esk2_0,X1),X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6937, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k1_funct_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k1_seq_1(k1_lp_space(esk2_0,X2),X3))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6784, plain, (v2_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_531, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),np__1)=k6_xcmplx_0(k4_xcmplx_0(np__1),X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6766, plain, (v1_membered(k6_numbers)|~m1_subset_1(k6_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1852, plain, (m2_subset_1(k4_xcmplx_0(np__1),k1_numbers,k5_numbers)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_532, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),X2)=k6_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_534, plain, (k6_xcmplx_0(np__1,k4_xcmplx_0(X1))=k6_xcmplx_0(X1,k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_6703, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k2_xcmplx_0(k1_funct_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k6_numbers)!=k1_seq_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_6769, plain, (v1_membered(k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_1733, plain, (v1_xcmplx_0(k1_seq_2(X1))|~v1_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1675, plain, (r1_xxreal_0(esk15_1(k5_numbers),k6_numbers)|r1_xxreal_0(k6_numbers,esk15_1(k5_numbers)))).#trainneg
cnf(i_0_535, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5731, plain, (v2_membered(k6_numbers)|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2973, negated_conjecture, (m2_subset_1(k1_seq_2(esk2_0),k1_numbers,k5_numbers)|~v1_xcmplx_0(k1_seq_2(esk2_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_6304, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|k3_power(k1_seq_2(esk2_0),k10_real_1(np__1,esk1_0))!=k6_numbers|~v1_series_1(esk2_0)|~m1_subset_1(k10_real_1(np__1,esk1_0),k1_numbers))).#trainneg
cnf(i_0_536, plain, (m2_subset_1(k2_nat_1(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_6958, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k2_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X1),X3),k6_numbers)!=k4_power(k6_numbers,X1)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6939, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X2),X3),k1_funct_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6935, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|k6_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X2),X4),X1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_seq_1(k1_lp_space(esk2_0,X2),X4),k1_numbers)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1734, plain, (v1_xxreal_0(k1_seq_2(X1))|~v2_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_3375, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k4_xcmplx_0(np__1)|k2_xcmplx_0(k4_power(k6_numbers,X1),np__1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~v1_xreal_0(k4_power(k6_numbers,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(k4_power(k6_numbers,X1)))).#trainneg
cnf(i_0_539, plain, (k2_xcmplx_0(k4_xcmplx_0(X2),k4_xcmplx_0(X1))=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7263, plain, (v2_membered(k6_numbers)|~m1_subset_1(k6_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_540, plain, (v1_xreal_0(k4_xcmplx_0(k2_xcmplx_0(X1,X2)))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7292, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|k2_xcmplx_0(k4_power(k6_numbers,X1),k4_power(k6_numbers,X2))!=k6_numbers|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~v7_ordinal1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1))))).#trainneg
cnf(i_0_7295, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|k2_xcmplx_0(k3_power(k6_numbers,X1),k6_numbers)!=k6_numbers|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v7_ordinal1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1))))).#trainneg
cnf(i_0_7293, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|k2_xcmplx_0(k4_power(k6_numbers,X1),k3_power(k6_numbers,X2))!=k6_numbers|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~v7_ordinal1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1))))).#trainneg
cnf(i_0_7294, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|k2_xcmplx_0(k4_power(k6_numbers,X1),k6_numbers)!=k6_numbers|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v7_ordinal1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1))))).#trainneg
cnf(i_0_7266, plain, (v2_membered(k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_1735, plain, (v3_ordinal1(k1_seq_2(X1))|~v3_ordinal1(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1828, plain, (r1_xxreal_0(np__1,esk15_1(k1_numbers))|r1_xxreal_0(esk15_1(k1_numbers),np__1))).#trainneg
cnf(i_0_3989, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k6_numbers|k2_xcmplx_0(k6_numbers,k3_power(k6_numbers,X1))!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_power(k6_numbers,X1),k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_542, plain, (k6_xcmplx_0(k4_xcmplx_0(X1),X2)=k4_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7267, negated_conjecture, (m1_subset_1(k1_seq_2(esk2_0),k5_numbers)|~v1_xcmplx_0(k1_seq_2(esk2_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_7209, plain, (r1_xxreal_0(k6_numbers,esk15_1(k5_numbers))|~v2_xxreal_0(esk15_1(k5_numbers))|~v1_xreal_0(esk15_1(k5_numbers)))).#trainneg
cnf(i_0_1721, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk2_0,X1),np__1)|r1_xxreal_0(np__1,k1_seq_1(esk2_0,X1)))).#trainneg
cnf(i_0_544, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(np__1)))=k2_xcmplx_0(k4_xcmplx_0(X1),np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7098, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X4,k6_numbers)|k2_xcmplx_0(k1_funct_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k4_power(k18_complex1(k3_power(k6_numbers,X4)),X2))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X4))).#trainneg
cnf(i_0_547, plain, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(np__1),X1))=k2_xcmplx_0(np__1,k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1722, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk2_0,X1),k6_numbers)|r1_xxreal_0(k6_numbers,k1_seq_1(esk2_0,X1)))).#trainneg
cnf(i_0_548, plain, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X1),X2))=k2_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_550, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|v2_xxreal_0(k4_xcmplx_0(k2_xcmplx_0(X1,X2)))|~v2_xxreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7569, plain, (k4_xcmplx_0(k6_xcmplx_0(X1,X2))=k2_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_7475, plain, (k4_xcmplx_0(k2_xcmplx_0(k4_xcmplx_0(X1),X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7469, plain, (k4_xcmplx_0(k2_xcmplx_0(X1,k4_xcmplx_0(X2)))=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2)))).#trainneg
cnf(i_0_7633, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_seq_1(esk2_0,X1))|~v2_xxreal_0(k1_seq_1(esk2_0,X1)))).#trainneg
cnf(i_0_7392, plain, (k4_xcmplx_0(k2_xcmplx_0(X2,k4_xcmplx_0(X1)))=k2_xcmplx_0(X1,k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6463, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X4)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_7741, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),np__1)=k6_numbers)).#trainneg
cnf(i_0_7477, negated_conjecture, (m1_subset_1(k1_seq_2(esk2_0),k5_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_7361, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k4_xcmplx_0(np__1)|k2_xcmplx_0(np__1,k4_power(k6_numbers,X1))!=k6_numbers|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~v1_xreal_0(k4_power(k6_numbers,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k4_power(k6_numbers,X1),k5_numbers))).#trainneg
cnf(i_0_552, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k6_xcmplx_0(X1,X2))|~v2_xxreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4831, plain, (m2_subset_1(k1_seq_1(k6_numbers,X1),k1_numbers,k5_numbers)|~v1_xcmplx_0(k1_seq_1(k6_numbers,X1))|~v6_membered(k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_6447, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|k3_power(k1_seq_2(esk2_0),k7_xcmplx_0(np__1,esk1_0))!=k6_numbers|~v1_series_1(esk2_0)|~m1_subset_1(k7_xcmplx_0(np__1,esk1_0),k1_numbers))).#trainneg
cnf(i_0_1993, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk2_0,X1),esk1_0)|r1_xxreal_0(esk1_0,k1_seq_1(esk2_0,X1)))).#trainneg
cnf(i_0_7476, plain, (k4_xcmplx_0(k2_xcmplx_0(k6_numbers,X1))=k6_xcmplx_0(k6_numbers,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6718, plain, (k4_power(k18_complex1(k4_power(k6_numbers,X1)),X2)=k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2)|r1_xxreal_0(X3,k6_numbers)|~v1_xreal_0(X3)|~v1_funct_2(k1_lp_space(esk2_0,X3),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X3))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_7279, plain, (m1_subset_1(k2_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_7800, plain, (k4_xcmplx_0(k6_xcmplx_0(X2,X1))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6464, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X4)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1899, plain, (v1_xxreal_0(k4_series_1(X1))|~v6_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_1909, plain, (r1_xxreal_0(k6_numbers,esk15_1(k1_numbers))|r1_xxreal_0(esk15_1(k1_numbers),k6_numbers))).#trainneg
cnf(i_0_7362, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k4_xcmplx_0(np__1)|k2_xcmplx_0(np__1,k4_power(k6_numbers,X1))!=k6_numbers|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~v1_xreal_0(k4_power(k6_numbers,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k4_power(k6_numbers,X1),k1_numbers))).#trainneg
cnf(i_0_2949, plain, (v1_funct_2(X1,k6_numbers,X2)|~v1_partfun1(X1,k6_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_1871, plain, (r1_xxreal_0(k6_numbers,k5_numbers)|r1_xxreal_0(k5_numbers,k6_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1872, plain, (r1_xxreal_0(np__1,k5_numbers)|r1_xxreal_0(k5_numbers,np__1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_972, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_557, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|v2_xxreal_0(k4_xcmplx_0(k2_xcmplx_0(X2,X1)))|~v2_xxreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_559, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|v2_xxreal_0(k6_xcmplx_0(X2,X1))|~v2_xxreal_0(X2)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_5069, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k6_numbers|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(k1_lp_space(esk2_0,X2)))).#trainneg
cnf(i_0_8156, plain, (k4_xcmplx_0(k6_xcmplx_0(X2,k4_xcmplx_0(X1)))=k6_xcmplx_0(k4_xcmplx_0(X2),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_8097, plain, (k6_xcmplx_0(k6_numbers,X1)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8099, plain, (k4_power(k18_complex1(k4_power(k6_numbers,X1)),X2)=k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2)|r1_xxreal_0(X3,k6_numbers)|~v1_xreal_0(X3)|~v1_funct_1(k1_lp_space(esk2_0,X3))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_7986, plain, (k4_xcmplx_0(k6_xcmplx_0(X1,X2))=k2_xcmplx_0(X2,k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_8088, plain, (k6_xcmplx_0(k6_numbers,k6_numbers)=k6_numbers)).#trainneg
cnf(i_0_5706, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X3,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|k1_seq_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k1_lp_space(esk2_0,X1)),np__1))!=k6_numbers|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k6_numbers)|~v1_partfun1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_830, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_564, plain, (X1=k3_power(X2,X3)|r1_xxreal_0(X2,k6_numbers)|X1!=k9_prepower(X2,X3)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_6928, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|k9_real_1(k1_seq_1(k1_lp_space(esk2_0,X2),X4),X1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(k9_real_1(k1_seq_1(k1_lp_space(esk2_0,X2),X4),X1))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6927, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|r1_xxreal_0(esk4_2(k1_lp_space(esk2_0,X2),X1),k18_complex1(k6_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X2),X4),X1)))|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_seq_1(k1_lp_space(esk2_0,X2),X4),k1_numbers)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6934, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|k6_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X2),X4),X1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(k1_seq_1(k1_lp_space(esk2_0,X2),X4))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_8150, plain, (k6_xcmplx_0(np__1,k6_numbers)=np__1)).#trainneg
cnf(i_0_6959, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k2_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X1),X3),k6_numbers)!=k3_power(k6_numbers,X1)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8016, plain, (k2_xcmplx_0(np__1,k4_xcmplx_0(np__1))=k6_numbers)).#trainneg
cnf(i_0_565, plain, (X1=k9_prepower(X2,X3)|r1_xxreal_0(X2,k6_numbers)|X1!=k3_power(X2,X3)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_575, plain, (m1_subset_1(k1_seq_1(k2_series_1(X1),X2),k1_numbers)|~v1_relat_1(k2_series_1(X1))|~v1_funct_1(k2_series_1(X1))|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_585, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_8245, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|k1_seq_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k6_numbers|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(k1_lp_space(esk2_0,X2)))).#trainneg
cnf(i_0_8322, plain, (k4_power(k18_complex1(k4_power(k6_numbers,X1)),X2)=k4_power(k18_complex1(k3_power(k6_numbers,X1)),X2)|r1_xxreal_0(X3,k6_numbers)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_8398, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X3,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k6_numbers)|~v1_partfun1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m2_subset_1(esk12_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v7_ordinal1(esk12_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k1_lp_space(esk2_0,X1))))).#trainneg
cnf(i_0_8397, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X3,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k6_numbers)|~v1_partfun1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~m2_subset_1(esk12_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(esk12_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k1_lp_space(esk2_0,X1)),k5_numbers))).#trainneg
cnf(i_0_6964, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|k1_seq_1(k1_lp_space(esk2_0,X1),X3)!=k2_xcmplx_0(k3_power(k6_numbers,X2),k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6962, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_seq_1(k1_lp_space(esk2_0,X1),X3)!=k2_xcmplx_0(k4_power(k6_numbers,X1),k3_power(k6_numbers,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6779, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_586, plain, (v1_xboole_0(X1)|m1_subset_1(esk16_2(X2,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1873, negated_conjecture, (r1_xxreal_0(esk1_0,k5_numbers)|r1_xxreal_0(k5_numbers,esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8180, plain, (r1_xxreal_0(k6_numbers,esk15_1(k1_numbers))|~v2_xxreal_0(esk15_1(k1_numbers))|~v1_xreal_0(esk15_1(k1_numbers)))).#trainneg
cnf(i_0_6044, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_587, plain, (v1_xboole_0(X1)|m1_subset_1(esk16_2(X2,X1),X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_589, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k4_xcmplx_0(X3)))=k6_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k4_xcmplx_0(X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8565, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k5_numbers))).#trainneg
cnf(i_0_8566, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v7_ordinal1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1))))).#trainneg
cnf(i_0_8602, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X3,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k6_numbers)|~v1_partfun1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(esk12_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k1_lp_space(esk2_0,X1)),k5_numbers)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xcmplx_0(esk12_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k1_lp_space(esk2_0,X1))))).#trainneg
cnf(i_0_8601, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X3,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k6_numbers)|~v1_partfun1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(esk12_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k1_lp_space(esk2_0,X1)),k5_numbers)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_8578, plain, (k4_power(k18_complex1(k4_power(k6_numbers,X1)),k7_xcmplx_0(X2,X3))=k4_power(k18_complex1(k3_power(k6_numbers,X1)),k7_xcmplx_0(X2,X3))|r1_xxreal_0(X4,k6_numbers)|~v1_xreal_0(X4)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8588, plain, (k4_power(k18_complex1(k4_power(k6_numbers,X1)),k10_real_1(X2,X3))=k4_power(k18_complex1(k3_power(k6_numbers,X1)),k10_real_1(X2,X3))|r1_xxreal_0(X4,k6_numbers)|~v1_xreal_0(X4)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_590, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k6_numbers))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k6_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_591, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_592, plain, (v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8800, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k5_numbers))).#trainneg
cnf(i_0_8562, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k5_numbers))).#trainneg
cnf(i_0_6930, negated_conjecture, (X1=k1_seq_2(k1_lp_space(esk2_0,X2))|k9_real_1(k1_seq_1(k1_lp_space(esk2_0,X2),X4),X1)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k9_real_1(k1_seq_1(k1_lp_space(esk2_0,X2),X4),X1),k1_numbers)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(esk5_3(k1_lp_space(esk2_0,X2),X1,X3),k1_numbers,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6923, negated_conjecture, (X1=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|k1_seq_1(X1,esk6_3(k1_lp_space(esk2_0,X2),X3,X1))!=k3_power(k18_complex1(k1_seq_1(k1_lp_space(esk2_0,X2),X4)),X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k18_complex1(k1_seq_1(k1_lp_space(esk2_0,X2),X4)),k1_numbers)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m2_subset_1(esk6_3(k1_lp_space(esk2_0,X2),X3,X1),k1_numbers,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6961, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_seq_1(k1_lp_space(esk2_0,X1),X3)!=k2_xcmplx_0(k4_power(k6_numbers,X1),k4_power(k6_numbers,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2548, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2094, plain, (v1_xxreal_0(k2_seq_2(X1))|~v6_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_8489, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),k6_numbers)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_593, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v2_xxreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8081, plain, (k6_xcmplx_0(k6_numbers,k4_xcmplx_0(np__1))=np__1|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_594, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X4)=k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7962, plain, (k2_xcmplx_0(X2,k4_xcmplx_0(X1))=k2_xcmplx_0(k4_xcmplx_0(X1),X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_9128, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_seq_1(k1_lp_space(esk2_0,X1),X3)!=k2_xcmplx_0(k4_power(k6_numbers,X1),k4_power(k6_numbers,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v7_ordinal1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1))))).#trainneg
cnf(i_0_9103, negated_conjecture, (X1=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|k1_seq_1(X1,esk6_3(k1_lp_space(esk2_0,X2),X3,X1))!=k3_power(k1_seq_1(k1_lp_space(esk2_0,X2),X4),X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(k1_lp_space(esk2_0,X2),X4),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k1_seq_1(k1_lp_space(esk2_0,X2),X4))|~r1_xxreal_0(k6_numbers,k1_seq_1(k1_lp_space(esk2_0,X2),X4)))).#trainneg
cnf(i_0_7979, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=k2_xcmplx_0(X1,k6_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_6987, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X1),X4),k1_seq_1(k1_lp_space(esk2_0,X2),X3))!=k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4968, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2295, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k6_numbers|k2_xcmplx_0(k3_power(k6_numbers,X1),k6_numbers)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~v1_xreal_0(k3_power(k6_numbers,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_596, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1)))|~v2_xxreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_9148, plain, (m2_subset_1(k4_xcmplx_0(np__1),k1_numbers,k5_numbers)|~m1_subset_1(k4_xcmplx_0(np__1),k5_numbers))).#trainneg
cnf(i_0_598, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k2_xcmplx_0(k4_xcmplx_0(X2),X3))=k2_xcmplx_0(k4_xcmplx_0(k2_xcmplx_0(X1,X2)),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_600, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(k4_xcmplx_0(X2),X3))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_9536, plain, (k1_seq_2(k1_lp_space(esk2_0,X1))=k6_numbers|r1_xxreal_0(X1,k6_numbers)|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_9127, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_seq_1(k1_lp_space(esk2_0,X1),X3)!=k2_xcmplx_0(k4_power(k6_numbers,X1),k4_power(k6_numbers,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k5_numbers))).#trainneg
cnf(i_0_8504, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|k2_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X1),X3),k6_numbers)!=k3_power(k6_numbers,X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X2)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X1,k1_numbers)|~v7_ordinal1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1))))).#trainneg
cnf(i_0_8428, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k5_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_8491, plain, (k2_xcmplx_0(k6_numbers,k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_601, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(k6_numbers,X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9155, plain, (m2_subset_1(k4_xcmplx_0(np__1),k1_numbers,k5_numbers)|~v7_ordinal1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_6075, negated_conjecture, (r1_xxreal_0(X1,k1_seq_1(esk2_0,X2))|r1_xxreal_0(k1_seq_1(esk2_0,X2),X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6076, plain, (r1_xxreal_0(X1,k1_seq_1(k6_numbers,X2))|r1_xxreal_0(k1_seq_1(k6_numbers,X2),X1)|~v1_xreal_0(X1)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_602, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9784, plain, (k1_seq_2(k1_lp_space(esk2_0,X1))=k6_numbers|r1_xxreal_0(X1,k6_numbers)|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_9782, plain, (k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k6_xcmplx_0(X3,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_9827, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X4,k6_numbers)|k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X4),X1),k6_numbers)!=k2_xcmplx_0(k4_power(k6_numbers,X1),k4_power(k6_numbers,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k5_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X4)|~v1_funct_2(k1_lp_space(esk2_0,X4),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X4))|~m1_subset_1(k1_lp_space(esk2_0,X4),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_9783, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,k4_xcmplx_0(X2)))=k2_xcmplx_0(k6_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_8182, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k4_xcmplx_0(np__1)|k2_xcmplx_0(np__1,k1_seq_1(k1_lp_space(esk2_0,X1),X3))!=k6_numbers|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~v1_xreal_0(k1_seq_1(k1_lp_space(esk2_0,X1),X3))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_7280, plain, (m1_subset_1(k2_nat_1(X1,X2),k1_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_6687, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k2_xcmplx_0(k1_funct_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k1_seq_1(esk2_0,X3))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9928, plain, (m2_subset_1(k4_xcmplx_0(np__1),k1_numbers,k5_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_4130, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k6_numbers|k2_xcmplx_0(k6_numbers,k3_power(k6_numbers,X1))!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~v1_xreal_0(k3_power(k6_numbers,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_604, plain, (v1_partfun1(esk15_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1)|v1_xboole_0(X2)|~v1_funct_2(esk15_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2))).#trainneg
cnf(i_0_607, plain, (k4_prepower(X1,X3)=X2|X2!=k3_power(X1,X3)|~epred1_3(X1,X3,X2)|~v1_int_1(X3))).#trainneg
cnf(i_0_609, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(esk15_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_funct_2(esk15_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2)|~v1_funct_1(esk15_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_8579, plain, (k4_power(k18_complex1(k4_power(k6_numbers,X1)),k7_xcmplx_0(X2,X3))=k4_power(k18_complex1(k3_power(k6_numbers,X1)),k7_xcmplx_0(X2,X3))|r1_xxreal_0(X4,k6_numbers)|~v1_xreal_0(X4)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_8580, plain, (k4_power(k18_complex1(k4_power(k6_numbers,X1)),k3_power(X2,X3))=k4_power(k18_complex1(k3_power(k6_numbers,X1)),k3_power(X2,X3))|r1_xxreal_0(X4,k6_numbers)|~v1_xreal_0(X4)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6965, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|k1_seq_1(k1_lp_space(esk2_0,X1),X3)!=k2_xcmplx_0(k4_power(k6_numbers,X2),k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7310, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|k2_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X1),X3),k6_numbers)!=k4_power(k6_numbers,X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X2)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X1,k1_numbers)|~v7_ordinal1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1))))).#trainneg
cnf(i_0_6985, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X1),X3),k4_power(k6_numbers,X2))!=k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7747, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),k6_numbers)=k4_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7429, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|k2_xcmplx_0(k4_power(k6_numbers,X1),k4_power(k6_numbers,X2))!=k6_numbers|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k5_numbers))).#trainneg
cnf(i_0_6311, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_612, plain, (v3_xxreal_0(esk9_2(X1,X2))|v2_xxreal_0(X3)|r1_xxreal_0(X3,esk9_2(X1,X2))|v2_comseq_2(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_10277, plain, (v1_partfun1(k6_numbers,k6_numbers)|v1_xboole_0(X1)|~v1_funct_2(k6_numbers,k6_numbers,X1))).#trainneg
cnf(i_0_613, plain, (v1_xreal_0(k1_lp_space(X1,X2))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_614, plain, (v7_ordinal1(k1_lp_space(X1,X2))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2332, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X2,k2_nat_1(esk12_2(X2,X1),np__1)),k1_seq_1(X1,esk12_2(X2,X1)))!=k1_funct_1(X1,k2_nat_1(esk12_2(X2,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_4270, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X3,k6_numbers)|k2_xcmplx_0(k1_seq_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X3),X2),k6_numbers))!=k1_funct_1(X1,k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~v1_xreal_0(X3)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X3),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X3))|~m1_subset_1(k1_lp_space(esk2_0,X3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6700, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k1_funct_1(X1,k2_xcmplx_0(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1))!=k2_xcmplx_0(k1_funct_1(X1,esk12_2(k1_lp_space(esk2_0,X2),X1)),k6_numbers)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X2)|~m2_subset_1(k2_xcmplx_0(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(esk12_2(k1_lp_space(esk2_0,X2),X1),k5_numbers))).#trainneg
cnf(i_0_5035, negated_conjecture, (k1_lp_space(k1_lp_space(esk2_0,X1),X2)=esk2_0|r1_xxreal_0(X1,k6_numbers)|k3_power(k6_numbers,X2)!=k6_numbers|~v1_xreal_0(X1)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(esk6_3(k1_lp_space(esk2_0,X1),X2,esk2_0),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4252, negated_conjecture, (k1_lp_space(esk2_0,X1)=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|r1_xxreal_0(X5,k6_numbers)|k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X5),X1),k6_numbers)!=k1_seq_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),X4)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X5)|~m2_subset_1(esk6_3(k1_lp_space(esk2_0,X2),X3,k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X5),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X5))|~m1_subset_1(k1_lp_space(esk2_0,X5),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8430, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_615, plain, (v1_xcmplx_0(k1_lp_space(X1,X2))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_616, plain, (v1_xxreal_0(k1_lp_space(X1,X2))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_617, plain, (v3_ordinal1(k1_lp_space(X1,X2))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6133, negated_conjecture, (k1_lp_space(k1_lp_space(esk2_0,X1),X2)=k1_lp_space(esk2_0,X3)|k4_power(k3_power(k6_numbers,X1),X2)!=k4_power(k6_numbers,X3)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(k3_power(k6_numbers,X1))|~r1_xxreal_0(k6_numbers,k3_power(k6_numbers,X1)))).#trainneg
cnf(i_0_10546, negated_conjecture, (k1_lp_space(k1_lp_space(esk2_0,X2),X1)=k1_lp_space(esk2_0,X1)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X2)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4605, negated_conjecture, (k1_lp_space(esk2_0,X1)=k1_lp_space(esk2_0,X2)|r1_xxreal_0(X1,k6_numbers)|k4_power(k6_numbers,X2)!=k6_numbers|~m2_subset_1(esk6_3(esk2_0,X2,k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7309, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|k2_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X1),X3),k6_numbers)!=k4_power(k6_numbers,X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X2)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k5_numbers))).#trainneg
cnf(i_0_1156, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6748, plain, (v1_xxreal_0(k1_seq_2(X1))|~v6_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_4739, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k6_numbers|k2_xcmplx_0(k3_power(k6_numbers,X1),k6_numbers)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_power(k6_numbers,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(k3_power(k6_numbers,X1)))).#trainneg
cnf(i_0_618, plain, (v3_membered(k1_lp_space(X1,X2))|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1718, negated_conjecture, (v3_xxreal_0(k1_seq_1(esk2_0,X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k1_seq_1(esk2_0,X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_620, plain, (v1_membered(k1_lp_space(X1,X2))|~v1_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1713, plain, (k1_seq_1(k1_lp_space(X1,X3),X2)=k6_numbers|r1_xxreal_0(X3,k6_numbers)|k1_seq_1(X1,X2)!=k6_numbers|~v1_xreal_0(X3)|~v1_xreal_0(k1_seq_1(X1,X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~r1_xxreal_0(k6_numbers,k1_seq_1(X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_10857, negated_conjecture, (X1=k1_lp_space(k1_lp_space(esk2_0,X2),X3)|r1_xxreal_0(X3,k6_numbers)|k1_seq_1(X1,esk6_3(k1_lp_space(esk2_0,X2),X3,X1))!=k6_numbers|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|k1_seq_1(k1_lp_space(esk2_0,X2),X4)!=k6_numbers|~v1_xreal_0(k1_seq_1(k1_lp_space(esk2_0,X2),X4))|~v1_xreal_0(X3)|~r1_xxreal_0(k6_numbers,k1_seq_1(k1_lp_space(esk2_0,X2),X4)))).#trainneg
cnf(i_0_10712, negated_conjecture, (v3_xxreal_0(k1_seq_1(esk2_0,X1))|r1_xxreal_0(k6_numbers,k1_seq_1(esk2_0,X1)))).#trainneg
cnf(i_0_2506, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_seq_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k3_power(k6_numbers,X1),k3_power(k6_numbers,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2991, plain, (X1=X2|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_621, plain, (v2_membered(k1_lp_space(X1,X2))|~v2_membered(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3183, plain, (k6_xcmplx_0(X1,k6_numbers)=k2_xcmplx_0(k6_numbers,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4411, plain, (v3_xxreal_0(k1_seq_1(k6_numbers,X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k1_seq_1(k6_numbers,X1))|~v1_xreal_0(X2)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_622, plain, (v3_valued_0(k1_lp_space(X1,X2))|~v1_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~v3_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_623, plain, (v1_valued_0(k1_lp_space(X1,X2))|~v1_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_5203, negated_conjecture, (k4_power(k3_power(k6_numbers,X1),X2)=k6_numbers|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(k3_power(k6_numbers,X1))|~r1_xxreal_0(k6_numbers,k3_power(k6_numbers,X1)))).#trainneg
cnf(i_0_3224, plain, (k1_seq_1(k1_lp_space(k1_lp_space(k1_lp_space(esk2_0,X1),X2),X3),k6_numbers)=k4_power(k4_power(k6_numbers,X2),X3)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(k4_power(k6_numbers,X2))|~v1_funct_2(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2))|~r1_xxreal_0(k6_numbers,k4_power(k6_numbers,X2))|~m1_subset_1(k1_lp_space(k1_lp_space(esk2_0,X1),X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_xreal_0(X1)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_8987, plain, (k2_xcmplx_0(k6_numbers,k2_xcmplx_0(X1,X2))=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_8184, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k4_xcmplx_0(np__1)|r1_xxreal_0(X1,k6_numbers)|k2_xcmplx_0(np__1,k1_seq_1(esk2_0,X3))!=k6_numbers|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1094, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_seq_1(X1,esk12_2(X2,X1)),k1_funct_1(X2,k2_xcmplx_0(esk12_2(X2,X1),np__1)))!=k1_funct_1(X1,k2_xcmplx_0(esk12_2(X2,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(esk12_2(X2,X1),k5_numbers))).#trainneg
cnf(i_0_1044, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_873, plain, (v1_xreal_0(esk5_3(X1,X2,X3))|X2=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_2194, negated_conjecture, (r1_xxreal_0(esk1_0,esk15_1(k1_numbers))|r1_xxreal_0(esk15_1(k1_numbers),esk1_0))).#trainneg
cnf(i_0_624, plain, (v2_valued_0(k1_lp_space(X1,X2))|~v2_valued_0(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_10881, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_seq_1(esk2_0,k6_numbers))|~v1_xboole_0(k1_seq_1(esk2_0,k6_numbers)))).#trainneg
cnf(i_0_5142, plain, (m1_subset_1(k6_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_625, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k1_numbers))|~r2_hidden(X1,k1_lp_space(X2,X3))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_10941, plain, (k2_xcmplx_0(k6_numbers,X1)=k2_xcmplx_0(X1,k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10897, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_seq_1(k1_lp_space(esk2_0,X1),X3)!=k2_xcmplx_0(k3_power(k6_numbers,X1),k3_power(k6_numbers,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_11041, plain, (k2_xcmplx_0(k6_numbers,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1)))=k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_10898, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|k2_xcmplx_0(k3_power(k6_numbers,X1),k3_power(k6_numbers,X2))!=k6_numbers|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7993, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X2)|~m1_subset_1(X4,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2995, plain, (v3_valued_0(esk15_1(k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))))).#trainneg
cnf(i_0_5000, negated_conjecture, (r1_xxreal_0(k6_numbers,esk2_0)|r1_xxreal_0(esk2_0,k6_numbers)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5143, plain, (m1_subset_1(k6_numbers,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_9962, negated_conjecture, (r1_xxreal_0(k1_seq_1(esk2_0,X1),k1_seq_1(k6_numbers,X2))|r1_xxreal_0(k1_seq_1(k6_numbers,X2),k1_seq_1(esk2_0,X1))|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_630, plain, (v2_valued_0(k1_lp_space(X1,X2))|~v2_membered(k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_631, plain, (v5_relat_1(k1_lp_space(X1,X2),k1_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_11210, plain, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(X2,k6_numbers)|k2_xcmplx_0(k6_numbers,k1_seq_1(k1_lp_space(esk2_0,X1),esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1))))!=k6_numbers|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X1),esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)))))).#trainneg
cnf(i_0_11401, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|k2_xcmplx_0(k3_power(k6_numbers,X1),k3_power(k6_numbers,X2))!=k6_numbers|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~v7_ordinal1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1))))).#trainneg
cnf(i_0_11400, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|k2_xcmplx_0(k3_power(k6_numbers,X1),k3_power(k6_numbers,X2))!=k6_numbers|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k5_numbers))).#trainneg
cnf(i_0_2312, negated_conjecture, (r1_xxreal_0(esk15_1(k5_numbers),esk1_0)|r1_xxreal_0(esk1_0,esk15_1(k5_numbers)))).#trainneg
cnf(i_0_640, plain, (v3_xxreal_0(esk4_2(X1,X2))|v2_xxreal_0(X3)|r1_xxreal_0(X3,esk4_2(X1,X2))|X2=k1_seq_2(X1)|~v1_xreal_0(X3)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_11176, plain, (v1_funct_2(k6_numbers,X1,X2)|~v1_partfun1(k6_numbers,X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_11147, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_seq_1(esk2_0,k6_numbers))|~m1_subset_1(k1_seq_1(esk2_0,k6_numbers),k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_5001, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)|r1_xxreal_0(esk2_0,np__1)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4412, plain, (r1_xxreal_0(k6_numbers,k1_seq_1(k6_numbers,X1))|r1_xxreal_0(k1_seq_1(k6_numbers,X1),k6_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_4413, plain, (r1_xxreal_0(np__1,k1_seq_1(k6_numbers,X1))|r1_xxreal_0(k1_seq_1(k6_numbers,X1),np__1)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_645, plain, (v3_xxreal_0(X1)|v2_comseq_2(X2)|~v3_xxreal_0(esk10_3(X2,X3,X1))|~v1_xreal_0(esk10_3(X2,X3,X1))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_646, plain, (v3_xxreal_0(X1)|v2_xxreal_0(esk10_3(X2,X3,X1))|v1_xboole_0(esk10_3(X2,X3,X1))|v2_comseq_2(X2)|~v1_xreal_0(esk10_3(X2,X3,X1))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_11291, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k1_seq_1(k1_lp_space(esk2_0,X1),X3)!=k2_xcmplx_0(k3_power(k6_numbers,X1),k3_power(k6_numbers,X2))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v7_ordinal1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1))))).#trainneg
cnf(i_0_11391, plain, (k2_series_1(k1_lp_space(esk2_0,X2))=esk2_0|r1_xxreal_0(X1,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)!=k1_seq_1(esk2_0,k6_numbers)|k2_xcmplx_0(k3_power(k6_numbers,X1),k3_power(k6_numbers,X2))!=k6_numbers|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),esk2_0),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),esk2_0),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6943, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X1,k6_numbers)|k2_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X2),X3),k6_numbers)!=k1_seq_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3135, plain, (k1_lp_space(k1_lp_space(esk2_0,X2),X3)=esk2_0|r1_xxreal_0(X1,k6_numbers)|k4_power(k6_numbers,X1)!=k4_power(k18_complex1(k4_power(k6_numbers,X2)),X3)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6990, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k4_power(k6_numbers,X2),k1_seq_1(k1_lp_space(esk2_0,X1),X3))!=k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1357, plain, (k2_xcmplx_0(k1_funct_1(X3,k2_nat_1(X2,np__1)),k1_funct_1(X1,X2))=k1_funct_1(X1,k2_nat_1(X2,np__1))|X1!=k2_series_1(X3)|~v1_valued_0(X3)|~v1_valued_0(X1)|~v1_partfun1(X3,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X3)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,X2),k5_numbers)|~m1_subset_1(k1_funct_1(X3,k2_nat_1(X2,np__1)),k5_numbers))).#trainneg
cnf(i_0_8166, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_11406, plain, (v1_valued_0(esk15_1(k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))))).#trainneg
cnf(i_0_2774, negated_conjecture, (r1_xxreal_0(k1_seq_2(esk2_0),np__1)|r1_xxreal_0(np__1,k1_seq_2(esk2_0)))).#trainneg
cnf(i_0_647, plain, (v3_xxreal_0(X1)|v2_xxreal_0(esk10_3(X2,X3,X1))|v1_xboole_0(X1)|v2_comseq_2(X2)|~v1_xreal_0(esk10_3(X2,X3,X1))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_11439, plain, (v1_funct_2(k6_numbers,X1,X2)|~v1_partfun1(k6_numbers,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5002, negated_conjecture, (r1_xxreal_0(esk1_0,esk2_0)|r1_xxreal_0(esk2_0,esk1_0)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4414, negated_conjecture, (r1_xxreal_0(esk1_0,k1_seq_1(k6_numbers,X1))|r1_xxreal_0(k1_seq_1(k6_numbers,X1),esk1_0)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_648, plain, (v2_xxreal_0(esk10_3(X1,X2,X3))|v2_comseq_2(X1)|~v2_xxreal_0(X3)|~v1_xreal_0(esk10_3(X1,X2,X3))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_651, plain, (k4_power(k18_complex1(k1_seq_1(X2,X1)),X3)=k1_seq_1(esk15_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))),X1)|esk15_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))!=k1_lp_space(X2,X3)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_2(esk15_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))),k5_numbers,k1_numbers)|~v1_funct_1(X2)|~v1_funct_1(esk15_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_11483, plain, (v2_comseq_2(esk2_0)|r1_xxreal_0(esk9_2(esk2_0,X2),k18_complex1(k9_real_1(k4_power(k6_numbers,X1),X2)))|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_11482, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),X2),k9_real_1(k4_power(k6_numbers,X1),X2))|~v1_xreal_0(X2)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(k9_real_1(k4_power(k6_numbers,X1),X2))|~r1_xxreal_0(k6_numbers,k9_real_1(k4_power(k6_numbers,X1),X2)))).#trainneg
cnf(i_0_11474, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),X2),k18_complex1(k9_real_1(k3_power(k6_numbers,X1),X2)))|~v1_xreal_0(X2)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11480, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),X2),k6_numbers)|~v1_xreal_0(X2)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|k9_real_1(k4_power(k6_numbers,X1),X2)!=k6_numbers|~v1_xreal_0(k9_real_1(k4_power(k6_numbers,X1),X2)))).#trainneg
cnf(i_0_8633, plain, (k2_zfmisc_1(X1,X2)=k2_zfmisc_1(X3,X4)|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_1613, plain, (v1_xreal_0(esk15_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_652, negated_conjecture, (k4_power(k18_complex1(k1_seq_1(X2,X1)),X3)=k1_seq_1(esk2_0,X1)|k1_lp_space(X2,X3)!=esk2_0|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_11440, plain, (v1_partfun1(k6_numbers,X1)|v1_xboole_0(X2)|~v1_funct_2(k6_numbers,X1,X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_11540, negated_conjecture, (~v3_xxreal_0(k1_seq_1(esk2_0,k6_numbers))|~m1_subset_1(k1_seq_1(esk2_0,k6_numbers),k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_657, plain, (k2_xcmplx_0(k1_funct_1(X3,k2_nat_1(X2,np__1)),k1_funct_1(X1,X2))=k1_funct_1(X1,k2_nat_1(X2,np__1))|~v1_xcmplx_0(k1_funct_1(X3,k2_nat_1(X2,np__1)))|~v1_xcmplx_0(k1_funct_1(X1,X2))|X1!=k2_series_1(X3)|~v1_valued_0(X3)|~v1_valued_0(X1)|~v1_partfun1(X3,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X3)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1))).#trainneg
cnf(i_0_658, plain, (v1_xreal_0(k1_funct_1(X1,k2_nat_1(X2,np__1)))|~v1_xreal_0(k1_funct_1(X3,k2_nat_1(X2,np__1)))|~v1_xreal_0(k1_funct_1(X1,X2))|X1!=k2_series_1(X3)|~v1_valued_0(X3)|~v1_valued_0(X1)|~v1_partfun1(X3,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X3)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1))).#trainneg
cnf(i_0_5429, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|r1_xxreal_0(X1,k6_numbers)|k1_seq_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k4_power(k18_complex1(k3_power(k6_numbers,X2)),X3),k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_partfun1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v1_relat_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~v1_xreal_0(X1)|~m2_subset_1(esk12_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~v1_funct_1(k1_lp_space(k1_lp_space(esk2_0,X2),X3))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(k1_lp_space(esk2_0,X2),X3),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X2),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~m1_subset_1(k1_lp_space(esk2_0,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_11633, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k3_power(k6_numbers,X2),k1_seq_1(k1_lp_space(esk2_0,X1),X3))!=k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9523, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|k2_xcmplx_0(k4_power(k6_numbers,X1),k1_seq_1(k1_lp_space(esk2_0,X2),X4))!=k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11768, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),X2),k9_real_1(k3_power(k6_numbers,X1),X2))|~v1_xreal_0(k9_real_1(k3_power(k6_numbers,X1),X2))|~v1_xreal_0(X2)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~r1_xxreal_0(k6_numbers,k9_real_1(k3_power(k6_numbers,X1),X2))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11773, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),X2),k6_xcmplx_0(k4_power(k6_numbers,X1),X2))|~v1_xreal_0(k6_xcmplx_0(k4_power(k6_numbers,X1),X2))|~v1_xreal_0(X2)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~r1_xxreal_0(k6_numbers,k6_xcmplx_0(k4_power(k6_numbers,X1),X2))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k4_power(k6_numbers,X1),k1_numbers))).#trainneg
cnf(i_0_11771, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),X2),k9_real_1(k1_seq_1(esk2_0,X3),X2))|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(k9_real_1(k1_seq_1(esk2_0,X3),X2))|~v1_xreal_0(X2)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~r1_xxreal_0(k6_numbers,k9_real_1(k1_seq_1(esk2_0,X3),X2))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_11914, negated_conjecture, (k4_power(k1_seq_1(X1,X2),X3)=k1_seq_1(esk2_0,X2)|k1_lp_space(X1,X3)!=esk2_0|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_xxreal_0(k6_numbers,k1_seq_1(X1,X2)))).#trainneg
cnf(i_0_5343, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k6_numbers|k2_xcmplx_0(k6_numbers,k3_power(k6_numbers,X1))!=k6_numbers|~v1_xcmplx_0(k3_power(k6_numbers,X1))|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k3_power(k6_numbers,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_664, plain, (k2_xcmplx_0(k1_funct_1(X1,X2),k1_funct_1(X3,k2_nat_1(np__1,X2)))=k1_funct_1(X1,k2_nat_1(np__1,X2))|X1!=k2_series_1(X3)|~v1_valued_0(X3)|~v1_valued_0(X1)|~v1_partfun1(X3,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X3)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_659, plain, (v3_xxreal_0(k1_funct_1(X1,X2))|v2_xxreal_0(k1_funct_1(X1,k2_nat_1(X2,np__1)))|~v2_xxreal_0(k1_funct_1(X3,k2_nat_1(X2,np__1)))|~v1_xreal_0(k1_funct_1(X1,X2))|~v1_xreal_0(k1_funct_1(X3,k2_nat_1(X2,np__1)))|X1!=k2_series_1(X3)|~v1_valued_0(X3)|~v1_valued_0(X1)|~v1_partfun1(X3,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X3)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1))).#trainneg
cnf(i_0_660, plain, (k2_xcmplx_0(k1_funct_1(X1,X2),k2_xcmplx_0(k1_funct_1(X3,k2_nat_1(X2,np__1)),X4))=k2_xcmplx_0(k1_funct_1(X1,k2_nat_1(X2,np__1)),X4)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k1_funct_1(X3,k2_nat_1(X2,np__1)))|~v1_xcmplx_0(k1_funct_1(X1,X2))|X1!=k2_series_1(X3)|~v1_valued_0(X3)|~v1_valued_0(X1)|~v1_partfun1(X3,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X3)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1))).#trainneg
cnf(i_0_12023, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))!=k2_xcmplx_0(k4_power(k6_numbers,X1),k6_numbers)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_12002, plain, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k2_xcmplx_0(k1_seq_1(esk2_0,k6_numbers),k1_seq_1(k1_lp_space(esk2_0,X1),X3))!=k1_funct_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_11772, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),X2),k9_real_1(k6_numbers,X2))|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(k9_real_1(k6_numbers,X2))|~v1_xreal_0(X2)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~r1_xxreal_0(k6_numbers,k9_real_1(k6_numbers,X2))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_11755, plain, (v2_comseq_2(esk2_0)|r1_xxreal_0(esk9_2(esk2_0,X1),k18_complex1(k9_real_1(k3_power(k6_numbers,X2),X1)))|r1_xxreal_0(X2,k6_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_12260, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),k6_numbers),k18_complex1(k1_seq_1(k1_lp_space(esk2_0,X1),X2)))|~v1_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X1),X2))|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(k1_lp_space(esk2_0,X1),X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12275, plain, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k2_xcmplx_0(k1_seq_1(esk2_0,k6_numbers),k1_seq_1(k1_lp_space(esk2_0,X1),X3))!=k1_seq_1(k1_lp_space(esk2_0,X1),k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1))|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11879, negated_conjecture, (k4_power(k6_numbers,X2)=k1_seq_1(esk2_0,X1)|k1_lp_space(esk2_0,X2)!=esk2_0|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_2775, negated_conjecture, (r1_xxreal_0(k1_seq_2(esk2_0),k6_numbers)|r1_xxreal_0(k6_numbers,k1_seq_2(esk2_0)))).#trainneg
cnf(i_0_665, plain, (k2_xcmplx_0(k1_funct_1(X1,X2),k1_funct_1(X3,k2_xcmplx_0(X2,np__1)))=k1_funct_1(X1,k2_xcmplx_0(X2,np__1))|X1!=k2_series_1(X3)|~v1_valued_0(X3)|~v1_valued_0(X1)|~v1_partfun1(X3,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X3)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_743, plain, (X1=k6_numbers|r1_xxreal_0(k6_numbers,X1)|k4_xcmplx_0(X1)!=k6_numbers|~v1_xreal_0(X1))).#trainneg
cnf(i_0_662, plain, (v3_xxreal_0(k1_funct_1(X1,k2_nat_1(X2,np__1)))|v2_xxreal_0(k1_funct_1(X3,k2_nat_1(X2,np__1)))|~v2_xxreal_0(k1_funct_1(X3,X2))|~v1_xreal_0(k1_funct_1(X1,k2_nat_1(X2,np__1)))|~v1_xreal_0(k1_funct_1(X3,X2))|X3!=k2_series_1(X1)|~v1_valued_0(X1)|~v1_valued_0(X3)|~v1_partfun1(X1,k5_numbers)|~v1_partfun1(X3,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_12307, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(X1,k6_numbers)|~v2_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),k6_numbers))|~v1_xreal_0(esk9_2(k1_lp_space(esk2_0,X1),k6_numbers))|~v1_xreal_0(X1)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_12337, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k2_xcmplx_0(k1_seq_1(esk2_0,k6_numbers),k1_seq_1(k1_lp_space(esk2_0,X1),X3))!=k3_power(k6_numbers,X1)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12322, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),k6_numbers),k18_complex1(k4_power(k18_complex1(k3_power(k6_numbers,X3)),X1)))|r1_xxreal_0(X3,k6_numbers)|~v1_xcmplx_0(k4_power(k18_complex1(k3_power(k6_numbers,X3)),X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k4_power(k18_complex1(k3_power(k6_numbers,X3)),X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_12336, negated_conjecture, (k1_lp_space(esk2_0,X1)=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k2_xcmplx_0(k1_seq_1(esk2_0,k6_numbers),k1_seq_1(k1_lp_space(esk2_0,X1),X3))!=k4_power(k6_numbers,X1)|k1_funct_1(k1_lp_space(esk2_0,X1),k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(k1_lp_space(esk2_0,X1),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X1),k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(k1_lp_space(esk2_0,X1))|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v3_valued_0(k1_lp_space(esk2_0,X1))|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),np__1),k1_numbers,k5_numbers)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),k1_lp_space(esk2_0,X1)),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12420, negated_conjecture, (k4_power(k6_numbers,X2)=k6_numbers|~m2_subset_1(X1,k1_numbers,k5_numbers)|k1_lp_space(esk2_0,X2)!=esk2_0|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1414, plain, (r1_xxreal_0(esk15_1(X1),esk15_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_12576, plain, (k4_xcmplx_0(np__1)=k6_numbers|np__1!=k6_numbers)).#trainneg
cnf(i_0_12480, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k6_numbers|k2_xcmplx_0(k1_seq_1(esk2_0,X3),k6_numbers)!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|k1_lp_space(esk2_0,X1)!=esk2_0|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_666, plain, (k2_xcmplx_0(k1_funct_1(X1,X2),k1_funct_1(X3,k2_nat_1(np__1,X2)))=k1_funct_1(X1,k2_nat_1(np__1,X2))|X1!=k2_series_1(X3)|~v1_valued_0(X3)|~v1_valued_0(X1)|~v1_partfun1(X3,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X3)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_838, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_12508, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_seq_2(esk2_0))|~v2_xxreal_0(k1_seq_2(esk2_0)))).#trainneg
cnf(i_0_667, plain, (k2_xcmplx_0(k1_funct_1(X1,X2),k1_seq_1(X3,k2_nat_1(X2,np__1)))=k1_funct_1(X1,k2_nat_1(X2,np__1))|X1!=k2_series_1(X3)|~v1_valued_0(X3)|~v1_valued_0(X1)|~v1_partfun1(X3,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X3)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v3_valued_0(X3))).#trainneg
cnf(i_0_668, plain, (k2_xcmplx_0(k1_seq_1(X1,X2),k1_funct_1(X3,k2_nat_1(X2,np__1)))=k1_funct_1(X1,k2_nat_1(X2,np__1))|X1!=k2_series_1(X3)|~v1_valued_0(X3)|~v1_valued_0(X1)|~v1_partfun1(X3,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X3)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_12667, plain, (np__1!=k6_numbers)).#trainneg
cnf(i_0_12856, negated_conjecture, (k1_funct_1(X1,k2_nat_1(X2,np__1))=k2_xcmplx_0(k1_funct_1(X1,X2),k6_numbers)|X1!=esk2_0|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X1)|~m2_subset_1(k2_nat_1(X2,np__1),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12076, negated_conjecture, (k3_power(k1_seq_1(X1,X2),X3)=k1_seq_1(esk2_0,X2)|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|k1_lp_space(X1,X3)!=esk2_0|~v1_xreal_0(k1_seq_1(X1,X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~r1_xxreal_0(k6_numbers,k1_seq_1(X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_12053, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),X2),k6_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X1),X3),X2))|~v1_xreal_0(k6_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X1),X3),X2))|~v1_xreal_0(X2)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~r1_xxreal_0(k6_numbers,k6_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X1),X3),X2))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(k1_lp_space(esk2_0,X1),X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12644, negated_conjecture, (k4_power(k6_numbers,X1)=k6_numbers|k1_lp_space(esk2_0,X1)!=esk2_0|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2612, plain, (k2_zfmisc_1(esk15_1(k1_zfmisc_1(X1)),X2)=k6_numbers|~v1_xboole_0(X1))).#trainneg
cnf(i_0_12483, negated_conjecture, (k1_seq_2(k1_lp_space(esk2_0,X1))=k6_numbers|k2_xcmplx_0(k6_numbers,k1_seq_1(esk2_0,X3))!=k6_numbers|~v2_comseq_2(k1_lp_space(esk2_0,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|k1_lp_space(esk2_0,X1)!=esk2_0|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_669, plain, (v3_xxreal_0(X1)|X3=k1_seq_2(X2)|~v3_xxreal_0(esk5_3(X2,X3,X1))|~v1_xreal_0(esk5_3(X2,X3,X1))|~v1_xreal_0(X1)|~v2_comseq_2(X2)|~v1_xreal_0(X3)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_839, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k5_numbers))).#trainneg
cnf(i_0_670, plain, (v3_xxreal_0(X1)|v2_xxreal_0(esk5_3(X2,X3,X1))|v1_xboole_0(esk5_3(X2,X3,X1))|X3=k1_seq_2(X2)|~v1_xreal_0(esk5_3(X2,X3,X1))|~v1_xreal_0(X1)|~v2_comseq_2(X2)|~v1_xreal_0(X3)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_671, plain, (v3_xxreal_0(X1)|v2_xxreal_0(esk5_3(X2,X3,X1))|v1_xboole_0(X1)|X3=k1_seq_2(X2)|~v1_xreal_0(esk5_3(X2,X3,X1))|~v1_xreal_0(X1)|~v2_comseq_2(X2)|~v1_xreal_0(X3)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_13053, negated_conjecture, (k1_funct_1(X1,k2_xcmplx_0(X2,np__1))=k2_xcmplx_0(k1_funct_1(X1,X2),k6_numbers)|X1!=esk2_0|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m2_subset_1(k2_xcmplx_0(X2,np__1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_13019, negated_conjecture, (k2_xcmplx_0(k1_funct_1(X1,X2),k6_numbers)=k1_seq_1(X1,k2_nat_1(X2,np__1))|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1)|X1!=esk2_0|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~m2_subset_1(k2_nat_1(X2,np__1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13043, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X2)|~m2_subset_1(k2_nat_1(esk12_2(k1_lp_space(esk2_0,X2),X1),np__1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|X1!=esk2_0|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),X1),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13054, negated_conjecture, (k1_funct_1(X1,k2_xcmplx_0(np__1,X2))=k2_xcmplx_0(k1_funct_1(X1,X2),k6_numbers)|X1!=esk2_0|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m2_subset_1(k2_xcmplx_0(np__1,X2),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13121, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),k6_numbers),k1_seq_1(k1_lp_space(esk2_0,X1),X2))|~v1_xreal_0(k1_seq_1(k1_lp_space(esk2_0,X1),X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~r1_xxreal_0(k6_numbers,k1_seq_1(k1_lp_space(esk2_0,X1),X2))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(k1_lp_space(esk2_0,X1),X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X1),X2)))).#trainneg
cnf(i_0_13055, negated_conjecture, (k1_funct_1(X1,k2_nat_1(np__1,X2))=k2_xcmplx_0(k1_funct_1(X1,X2),k6_numbers)|X1!=esk2_0|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m2_subset_1(k2_nat_1(np__1,X2),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13047, negated_conjecture, (k2_xcmplx_0(k1_funct_1(X1,X2),k2_xcmplx_0(k1_funct_1(X3,X2),k6_numbers))=k1_funct_1(X1,k2_nat_1(X2,np__1))|X1!=k2_series_1(X3)|~v1_valued_0(X3)|~v1_valued_0(X1)|~v1_partfun1(X3,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X3)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)|X3!=esk2_0|~m2_subset_1(k2_nat_1(X2,np__1),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13052, negated_conjecture, (k1_funct_1(X1,k2_nat_1(np__1,X2))=k2_xcmplx_0(k1_funct_1(X1,X2),k6_numbers)|X1!=esk2_0|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m2_subset_1(k2_nat_1(np__1,X2),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2019, negated_conjecture, (v1_series_1(esk2_0))).#trainneg
cnf(i_0_11500, negated_conjecture, (k1_seq_2(esk2_0)=k6_numbers)).#trainneg
cnf(i_0_2614, plain, (k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)=k6_numbers|~v1_xboole_0(X2))).#trainneg
cnf(i_0_672, plain, (v2_xxreal_0(esk5_3(X1,X2,X3))|X2=k1_seq_2(X1)|~v2_xxreal_0(X3)|~v1_xreal_0(esk5_3(X1,X2,X3))|~v1_xreal_0(X3)|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_6458, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_675, plain, (v3_xxreal_0(esk9_2(X1,X2))|v2_comseq_2(X1)|~v3_xxreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk10_3(X1,X2,X3)),X2)))|~v1_xreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk10_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk9_2(X1,X2))|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_676, plain, (v3_xxreal_0(esk9_2(X1,X2))|v2_xxreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk10_3(X1,X2,X3)),X2)))|v1_xboole_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk10_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk10_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk9_2(X1,X2))|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_975, negated_conjecture, (X1=k1_seq_2(esk2_0)|r1_xxreal_0(esk4_2(esk2_0,X1),k18_complex1(k9_real_1(k6_numbers,X1)))|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_924, negated_conjecture, (X1=k1_seq_2(esk2_0)|r1_xxreal_0(esk4_2(esk2_0,X1),k18_complex1(k6_xcmplx_0(k6_numbers,X1)))|~v1_xreal_0(X1)|~m2_subset_1(esk5_3(esk2_0,X1,X2),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_957, negated_conjecture, (X1=k1_seq_2(esk2_0)|r1_xxreal_0(esk4_2(esk2_0,X1),k9_real_1(k6_numbers,X1))|~v1_xreal_0(k9_real_1(k6_numbers,X1))|~v1_xreal_0(X1)|~m2_subset_1(esk5_3(esk2_0,X1,X2),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(k6_numbers,k9_real_1(k6_numbers,X1)))).#trainneg
cnf(i_0_13345, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|X1!=esk2_0|~v1_valued_0(X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X2)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),X1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~v7_ordinal1(esk12_2(k1_lp_space(esk2_0,X2),X1)))).#trainneg
cnf(i_0_13390, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),k6_numbers),k4_power(k6_numbers,X1))|~v1_xreal_0(k4_power(k6_numbers,X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~r1_xxreal_0(k6_numbers,k4_power(k6_numbers,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_13388, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),k6_numbers),k3_power(k6_numbers,X1))|~v1_xreal_0(k3_power(k6_numbers,X1))|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~r1_xxreal_0(k6_numbers,k3_power(k6_numbers,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5144, plain, (m1_subset_1(k6_numbers,k1_zfmisc_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k6_numbers)))).#trainneg
cnf(i_0_2615, plain, (k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3)=k6_numbers|~v1_xboole_0(X1))).#trainneg
cnf(i_0_504, negated_conjecture, (r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(X1,k18_complex1(k9_real_1(k6_numbers,esk7_1(esk2_0))))|~r1_xxreal_0(esk8_2(esk2_0,X1),X2))).#trainneg
cnf(i_0_677, plain, (v3_xxreal_0(esk9_2(X1,X2))|v2_xxreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk10_3(X1,X2,X3)),X2)))|v1_xboole_0(esk9_2(X1,X2))|v2_comseq_2(X1)|~v1_xreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk10_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk9_2(X1,X2))|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_7126, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))=k6_numbers|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_678, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk10_3(X1,X2,X3)),X2)))|v2_comseq_2(X1)|~v2_xxreal_0(esk9_2(X1,X2))|~v1_xreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk10_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk9_2(X1,X2))|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_13375, negated_conjecture, (k1_funct_1(X1,k2_xcmplx_0(X2,np__1))=k2_xcmplx_0(k1_funct_1(X1,X2),k6_numbers)|X1!=esk2_0|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(X2)|~m2_subset_1(k2_xcmplx_0(X2,np__1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_12990, negated_conjecture, (k2_series_1(X1)=esk2_0|k1_funct_1(X1,k2_nat_1(esk12_2(X1,esk2_0),np__1))!=k2_xcmplx_0(k1_funct_1(esk2_0,esk12_2(X1,esk2_0)),k6_numbers)|k1_funct_1(X1,k6_numbers)!=k1_seq_1(esk2_0,k6_numbers)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~m2_subset_1(esk12_2(X1,esk2_0),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~m2_subset_1(k2_nat_1(esk12_2(X1,esk2_0),np__1),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13344, negated_conjecture, (X1=k2_series_1(k1_lp_space(esk2_0,X2))|r1_xxreal_0(X2,k6_numbers)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(k1_lp_space(esk2_0,X2),k6_numbers)|X1!=esk2_0|~v1_valued_0(X1)|~v1_partfun1(k1_lp_space(esk2_0,X2),k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(k1_lp_space(esk2_0,X2),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(k1_lp_space(esk2_0,X2))|~v1_relat_1(X1)|~v3_valued_0(k1_lp_space(esk2_0,X2))|~v1_xreal_0(X2)|~m2_subset_1(esk12_2(k1_lp_space(esk2_0,X2),X1),k1_numbers,k5_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X2))|~v1_funct_1(X1)|~m1_subset_1(esk12_2(k1_lp_space(esk2_0,X2),X1),k5_numbers))).#trainneg
cnf(i_0_13184, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),k6_numbers),k6_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_numbers)|k1_lp_space(esk2_0,X1)!=esk2_0)).#trainneg
cnf(i_0_13325, negated_conjecture, (k1_seq_1(X1,k2_nat_1(X2,np__1))=k1_funct_1(X1,X2)|~v1_xcmplx_0(k1_funct_1(X1,X2))|X1!=esk2_0|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m2_subset_1(k2_nat_1(X2,np__1),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_2023, negated_conjecture, (k4_series_1(esk2_0)=k6_numbers)).#trainneg
cnf(i_0_13995, negated_conjecture, (k1_seq_1(k1_lp_space(X1,X3),k2_nat_1(X2,np__1))=k4_power(k1_funct_1(X1,X2),X3)|~v1_xreal_0(k1_funct_1(X1,X2))|~m2_subset_1(k2_nat_1(X2,np__1),k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~r1_xxreal_0(k6_numbers,k1_funct_1(X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|X1!=esk2_0|~v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13996, negated_conjecture, (k1_seq_1(k1_lp_space(X1,X3),k2_nat_1(X2,np__1))=k3_power(k1_funct_1(X1,X2),X3)|~v1_xreal_0(k1_funct_1(X1,X2))|~m2_subset_1(k2_nat_1(X2,np__1),k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~r1_xxreal_0(k6_numbers,k1_funct_1(X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|X1!=esk2_0|~v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2766, plain, (r1_xxreal_0(esk15_1(X1),esk15_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_2017, negated_conjecture, (k2_seq_2(esk2_0)=k6_numbers)).#trainneg
cnf(i_0_682, plain, (v2_comseq_2(X1)|r1_xxreal_0(esk9_2(X1,X2),k18_complex1(k6_xcmplx_0(k1_seq_1(X1,esk10_3(X1,X2,X3)),X2)))|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X1,esk10_3(X1,X2,X3)),k1_numbers))).#trainneg
cnf(i_0_683, plain, (v2_comseq_2(X1)|r1_xxreal_0(esk9_2(X1,X2),k6_numbers)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|k9_real_1(k1_seq_1(X1,esk10_3(X1,X2,X3)),X2)!=k6_numbers|~v1_xreal_0(k9_real_1(k1_seq_1(X1,esk10_3(X1,X2,X3)),X2)))).#trainneg
cnf(i_0_684, plain, (v3_xxreal_0(esk4_2(X1,X2))|X2=k1_seq_2(X1)|~v3_xxreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk5_3(X1,X2,X3)),X2)))|~v1_xreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk5_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk4_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_13650, negated_conjecture, (X1=k6_numbers|r1_xxreal_0(esk4_2(esk2_0,X1),k18_complex1(k9_real_1(k6_numbers,X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_13870, negated_conjecture, (k1_funct_1(X1,k2_xcmplx_0(np__1,X2))=k2_xcmplx_0(k1_funct_1(X1,X2),k6_numbers)|X1!=esk2_0|~v1_xcmplx_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v7_ordinal1(X2)|~m2_subset_1(k2_xcmplx_0(np__1,X2),k1_numbers,k5_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_14289, negated_conjecture, (X1=k6_numbers|r1_xxreal_0(esk4_2(esk2_0,X1),k18_complex1(k6_xcmplx_0(k6_numbers,X1)))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_14103, negated_conjecture, (k1_seq_1(k1_lp_space(X1,X2),k2_nat_1(np__1,X3))=k4_power(k1_funct_1(X1,X3),X2)|X1!=esk2_0|~v1_xcmplx_0(k1_funct_1(X1,X3))|~v1_partfun1(X1,k5_numbers)|~v3_valued_0(X1)|~m2_subset_1(k2_nat_1(np__1,X3),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(X1)|~r1_xxreal_0(k6_numbers,k1_funct_1(X1,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_14296, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|~v2_xxreal_0(k10_real_1(np__1,esk1_0))|~v1_xreal_0(k10_real_1(np__1,esk1_0)))).#trainneg
cnf(i_0_14223, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),X2),k18_complex1(k6_xcmplx_0(k4_power(k18_complex1(k3_power(k6_numbers,X4)),X1),X2)))|r1_xxreal_0(X4,k6_numbers)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k4_power(k18_complex1(k3_power(k6_numbers,X4)),X1),k1_numbers)|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(X4)|~m2_subset_1(esk10_3(k1_lp_space(esk2_0,X1),X2,X3),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14220, negated_conjecture, (v2_comseq_2(k1_lp_space(esk2_0,X1))|r1_xxreal_0(esk9_2(k1_lp_space(esk2_0,X1),X2),k18_complex1(k6_xcmplx_0(k1_seq_1(k1_lp_space(esk2_0,X1),X4),X2)))|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(k1_lp_space(esk2_0,X1),k5_numbers,k1_numbers)|~v1_funct_1(k1_lp_space(esk2_0,X1))|~m1_subset_1(k1_seq_1(k1_lp_space(esk2_0,X1),X4),k1_numbers)|~m1_subset_1(k1_lp_space(esk2_0,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m2_subset_1(X4,k1_numbers,k5_numbers)|~m2_subset_1(esk10_3(k1_lp_space(esk2_0,X1),X2,X3),k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14189, negated_conjecture, (k1_seq_1(k1_lp_space(X1,X2),k2_nat_1(np__1,X3))=k3_power(k1_funct_1(X1,X3),X2)|X1!=esk2_0|~v1_xcmplx_0(k1_funct_1(X1,X3))|~v1_partfun1(X1,k5_numbers)|~v3_valued_0(X1)|~m2_subset_1(k2_nat_1(np__1,X3),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(X1)|~r1_xxreal_0(k6_numbers,k1_funct_1(X1,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_14100, negated_conjecture, (k1_seq_1(k1_lp_space(X1,X2),k2_nat_1(np__1,X3))=k4_power(k1_funct_1(X1,X3),X2)|X1!=esk2_0|~v1_xcmplx_0(k1_funct_1(X1,X3))|~v1_partfun1(X1,k5_numbers)|~v3_valued_0(X1)|~m2_subset_1(k2_nat_1(np__1,X3),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(X1)|~r1_xxreal_0(k6_numbers,k1_funct_1(X1,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_5150, negated_conjecture, (r1_xxreal_0(esk1_0,X1)|r1_xxreal_0(X1,esk1_0)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1029, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2386, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|k4_power(k1_seq_2(esk2_0),k7_xcmplx_0(np__1,esk1_0))!=k6_numbers)).#trainneg
cnf(i_0_687, plain, (v2_xxreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk5_3(X1,X2,X3)),X2)))|X2=k1_seq_2(X1)|~v2_xxreal_0(esk4_2(X1,X2))|~v1_xreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk5_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk4_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_685, plain, (v3_xxreal_0(esk4_2(X1,X2))|v2_xxreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk5_3(X1,X2,X3)),X2)))|v1_xboole_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk5_3(X1,X2,X3)),X2)))|X2=k1_seq_2(X1)|~v1_xreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk5_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk4_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_686, plain, (v3_xxreal_0(esk4_2(X1,X2))|v2_xxreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk5_3(X1,X2,X3)),X2)))|v1_xboole_0(esk4_2(X1,X2))|X2=k1_seq_2(X1)|~v1_xreal_0(k18_complex1(k9_real_1(k1_seq_1(X1,esk5_3(X1,X2,X3)),X2)))|~v1_xreal_0(esk4_2(X1,X2))|~v2_comseq_2(X1)|~v1_xreal_0(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_14423, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|r1_xxreal_0(k10_real_1(np__1,esk1_0),esk1_0)|~v1_xreal_0(k10_real_1(np__1,esk1_0)))).#trainneg
cnf(i_0_14420, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|~v2_xxreal_0(k7_xcmplx_0(np__1,esk1_0))|~v1_xreal_0(k7_xcmplx_0(np__1,esk1_0)))).#trainneg
cnf(i_0_14416, negated_conjecture, (k4_power(k1_funct_1(X1,X3),X2)=k1_seq_1(k1_lp_space(X1,X2),k2_xcmplx_0(np__1,X3))|X1!=esk2_0|~v1_xcmplx_0(k1_funct_1(X1,X3))|~v1_partfun1(X1,k5_numbers)|~v3_valued_0(X1)|~v7_ordinal1(X3)|~m2_subset_1(k2_xcmplx_0(np__1,X3),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(X1)|~r1_xxreal_0(k6_numbers,k1_funct_1(X1,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_14528, negated_conjecture, (k1_seq_1(k1_lp_space(X1,X2),k2_xcmplx_0(np__1,X3))=k3_power(k1_funct_1(X1,X3),X2)|X1!=esk2_0|~v1_xcmplx_0(k1_funct_1(X1,X3))|~v1_partfun1(X1,k5_numbers)|~v3_valued_0(X1)|~v7_ordinal1(X3)|~m2_subset_1(k2_xcmplx_0(np__1,X3),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(X1)|~r1_xxreal_0(k6_numbers,k1_funct_1(X1,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4996, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6304, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|k3_power(k1_seq_2(esk2_0),k10_real_1(np__1,esk1_0))!=k6_numbers|~m1_subset_1(k10_real_1(np__1,esk1_0),k1_numbers))).#trainneg
cnf(i_0_695, negated_conjecture, (k1_seq_1(k3_series_1(k1_lp_space(esk2_0,X1)),esk11_2(X1,esk2_0))=k1_seq_1(k1_lp_space(esk2_0,X1),esk11_2(X1,esk2_0))|k1_seq_1(k3_series_1(k1_lp_space(esk2_0,X1)),X2)=k1_seq_1(k1_lp_space(esk2_0,X1),X2)|k1_seq_1(k3_series_1(k1_lp_space(esk2_0,X1)),k6_numbers)!=k1_seq_1(k1_lp_space(esk2_0,X1),k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_689, plain, (r1_xxreal_0(esk4_2(X1,esk7_1(X1)),k6_numbers)|esk7_1(X1)=k1_seq_2(X1)|~v2_comseq_2(X1)|~v1_xreal_0(esk4_2(X1,esk7_1(X1)))|~m2_subset_1(esk5_3(X1,esk7_1(X1),X2),k1_numbers,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~r1_xxreal_0(esk8_2(X1,esk4_2(X1,esk7_1(X1))),esk5_3(X1,esk7_1(X1),X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_xreal_0(esk7_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_694, plain, (k1_seq_1(k3_series_1(k1_lp_space(esk15_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))),X1)),esk11_2(X1,esk15_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))=k1_seq_1(k1_lp_space(esk15_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))),X1),esk11_2(X1,esk15_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))|k1_seq_1(k3_series_1(k1_lp_space(esk15_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))),X1)),X2)=k1_seq_1(k1_lp_space(esk15_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))),X1),X2)|k1_seq_1(k3_series_1(k1_lp_space(esk15_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))),X1)),k6_numbers)!=k1_seq_1(k1_lp_space(esk15_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))),X1),k6_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_2(esk15_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))),k5_numbers,k1_numbers)|~v1_funct_1(esk15_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14186, negated_conjecture, (k1_seq_1(k1_lp_space(X1,X2),k2_nat_1(np__1,X3))=k3_power(k1_funct_1(X1,X3),X2)|X1!=esk2_0|~v1_xcmplx_0(k1_funct_1(X1,X3))|~v1_partfun1(X1,k5_numbers)|~v3_valued_0(X1)|~m2_subset_1(k2_nat_1(np__1,X3),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(X1)|~r1_xxreal_0(k6_numbers,k1_funct_1(X1,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_14187, negated_conjecture, (k1_seq_1(k1_lp_space(X1,X2),k2_xcmplx_0(X3,np__1))=k3_power(k1_funct_1(X1,X3),X2)|X1!=esk2_0|~v1_xcmplx_0(k1_funct_1(X1,X3))|~v1_partfun1(X1,k5_numbers)|~v3_valued_0(X1)|~m2_subset_1(k2_xcmplx_0(X3,np__1),k1_numbers,k5_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_funct_1(X1)|~r1_xxreal_0(k6_numbers,k1_funct_1(X1,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_14422, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|r1_xxreal_0(k10_real_1(np__1,esk1_0),np__1)|~v1_xreal_0(k10_real_1(np__1,esk1_0)))).#trainneg
cnf(i_0_14768, negated_conjecture, (k1_seq_1(k1_lp_space(X1,X3),k2_xcmplx_0(np__1,X2))=k4_power(k1_seq_1(X1,X2),X3)|X1!=esk2_0|~v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_partfun1(X1,k5_numbers)|~v3_valued_0(X1)|~v7_ordinal1(X2)|~m2_subset_1(k2_xcmplx_0(np__1,X2),k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X1)|~r1_xxreal_0(k6_numbers,k1_seq_1(X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_relat_1(X1))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 154
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 230
# Removed in clause preprocessing      : 7
# Initial clauses in saturation        : 223
# Processed clauses                    : 2147
# ...of these trivial                  : 74
# ...subsumed                          : 727
# ...remaining for further processing  : 1346
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 78
# Backward-rewritten                   : 87
# Generated clauses                    : 14925
# ...of the previous two non-trivial   : 13746
# Contextual simplify-reflections      : 295
# Paramodulations                      : 14911
# Factorizations                       : 0
# Equation resolutions                 : 14
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1180
#    Positive orientable unit clauses  : 107
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 23
#    Non-unit-clauses                  : 1050
# Current number of unprocessed clauses: 11553
# ...number of literals in the above   : 128763
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1511
# Proof object given clauses           : 170
# Proof search given clauses           : 1346
# Clause-clause subsumption calls (NU) : 1613140
# Rec. Clause-clause subsumption calls : 50620
# Non-unit clause-clause subsumptions  : 1007
# Unit Clause-clause subsumption calls : 11269
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 41
# BW rewrite match successes           : 23
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 652093

# -------------------------------------------------
# User time                : 6.487 s
# System time              : 0.222 s
# Total time               : 6.709 s
# Maximum resident set size: 78260 pages
