2020-05-31 07:03:30.150199: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_15
2020-05-31 07:03:30.247072: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-05-31 07:03:30.415176: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-05-31 07:03:30.996104: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-05-31 07:03:32.667571: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 2517392 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_15' loaded (query=256; context=512; weigths=real; len_mult=0.000000).
# Preprocessing time       : 3.683 s
#TF# Flushing cached clauses: 0 unprocessed
#TF# Cached clauses flushed: 164 unprocessed

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t11_waybel_6, conjecture, ![X1]:(((((v3_orders_2(X1)&v4_orders_2(X1))&v5_orders_2(X1))&v2_lattice3(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v2_waybel_6(X2,X1)<=>![X3]:(((~(v1_xboole_0(X3))&v1_finset_1(X3))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))=>(X2=k2_yellow_0(X1,X3)=>r2_hidden(X2,X3)))))), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', t11_waybel_6)).
fof(dt_k7_domain_1, axiom, ![X1, X2, X3]:(((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))&m1_subset_1(X3,X1))=>m1_subset_1(k7_domain_1(X1,X2,X3),k1_zfmisc_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', dt_k7_domain_1)).
fof(cc2_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v2_lattice3(X1)=>~(v2_struct_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', cc2_lattice3)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', dt_l1_orders_2)).
fof(redefinition_k7_domain_1, axiom, ![X1, X2, X3]:(((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))&m1_subset_1(X3,X1))=>k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', redefinition_k7_domain_1)).
fof(fc2_finset_1, axiom, ![X1, X2]:v1_finset_1(k2_tarski(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', fc2_finset_1)).
fof(fc3_xboole_0, axiom, ![X1, X2]:~(v1_xboole_0(k2_tarski(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', fc3_xboole_0)).
fof(t40_yellow_0, axiom, ![X1]:(((((v3_orders_2(X1)&v4_orders_2(X1))&v5_orders_2(X1))&v2_lattice3(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k12_lattice3(X1,X2,X3)))), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', t40_yellow_0)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', d2_tarski)).
fof(redefinition_k12_lattice3, axiom, ![X1, X2, X3]:(((((v5_orders_2(X1)&v2_lattice3(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', redefinition_k12_lattice3)).
fof(d2_waybel_6, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v2_waybel_6(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~(((X2=k11_lattice3(X1,X3,X4)&X3!=X2)&X4!=X2))))))), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', d2_waybel_6)).
fof(t4_yellow_2, axiom, ![X1]:((((v4_orders_2(X1)&v5_orders_2(X1))&v2_lattice3(X1))&l1_orders_2(X1))=>![X2, X3]:((r2_yellow_0(X1,X2)&r2_yellow_0(X1,X3))=>(r2_yellow_0(X1,k2_xboole_0(X2,X3))&k2_yellow_0(X1,k2_xboole_0(X2,X3))=k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))))), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', t4_yellow_2)).
fof(commutativity_k12_lattice3, axiom, ![X1, X2, X3]:(((((v5_orders_2(X1)&v2_lattice3(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>k12_lattice3(X1,X2,X3)=k12_lattice3(X1,X3,X2)), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', commutativity_k12_lattice3)).
fof(dt_k2_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', dt_k2_yellow_0)).
fof(s2_finset_1__e5_25_1__waybel_6, axiom, ![X1, X2, X3]:(((((((((v3_orders_2(X1)&v4_orders_2(X1))&v5_orders_2(X1))&v2_lattice3(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))&~(v1_xboole_0(X3)))&v1_finset_1(X3))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))=>(((v1_finset_1(X3)&(X2=k2_yellow_0(X1,k1_xboole_0)=>(v1_xboole_0(k1_xboole_0)|r2_hidden(X2,k1_xboole_0))))&![X4, X5]:(((r2_hidden(X4,X3)&r1_tarski(X5,X3))&(X2=k2_yellow_0(X1,X5)=>(v1_xboole_0(X5)|r2_hidden(X2,X5))))=>(X2=k2_yellow_0(X1,k2_xboole_0(X5,k1_tarski(X4)))=>(v1_xboole_0(k2_xboole_0(X5,k1_tarski(X4)))|r2_hidden(X2,k2_xboole_0(X5,k1_tarski(X4)))))))=>(X2=k2_yellow_0(X1,X3)=>(v1_xboole_0(X3)|r2_hidden(X2,X3))))), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', s2_finset_1__e5_25_1__waybel_6)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', t1_xboole_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', t3_subset)).
fof(dt_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', dt_k6_domain_1)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', redefinition_k6_domain_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', fc1_xboole_0)).
fof(t55_yellow_0, axiom, ![X1]:(((((~(v2_struct_0(X1))&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&l1_orders_2(X1))=>(v2_lattice3(X1)<=>![X2]:(((~(v1_xboole_0(X2))&v1_finset_1(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>r2_yellow_0(X1,X2)))), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', t55_yellow_0)).
fof(reflexivity_r1_tarski, axiom, ![X1, X2]:r1_tarski(X1,X1), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', reflexivity_r1_tarski)).
fof(t39_yellow_0, axiom, ![X1]:((((~(v2_struct_0(X1))&v3_orders_2(X1))&v5_orders_2(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2))), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', t39_yellow_0)).
fof(fc1_finset_1, axiom, ![X1]:v1_finset_1(k1_tarski(X1)), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', fc1_finset_1)).
fof(fc2_xboole_0, axiom, ![X1]:~(v1_xboole_0(k1_tarski(X1))), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', fc2_xboole_0)).
fof(cc2_finset_1, axiom, ![X1]:(v1_finset_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_finset_1(X2))), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', cc2_finset_1)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', t4_subset)).
fof(d3_xboole_0, axiom, ![X1, X2, X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2)))), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', d3_xboole_0)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', t1_subset)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', t6_boole)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', t1_boole)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', commutativity_k2_xboole_0)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('/scratch/mptp/allpred/preds__0.005/t11_waybel_6', d1_tarski)).
fof(c_0_34, negated_conjecture, ~(![X1]:(((((v3_orders_2(X1)&v4_orders_2(X1))&v5_orders_2(X1))&v2_lattice3(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v2_waybel_6(X2,X1)<=>![X3]:(((~(v1_xboole_0(X3))&v1_finset_1(X3))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))=>(X2=k2_yellow_0(X1,X3)=>r2_hidden(X2,X3))))))), inference(assume_negation,[status(cth)],[t11_waybel_6])).
fof(c_0_35, negated_conjecture, ![X9]:(((((v3_orders_2(esk1_0)&v4_orders_2(esk1_0))&v5_orders_2(esk1_0))&v2_lattice3(esk1_0))&l1_orders_2(esk1_0))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(((((~v1_xboole_0(esk3_0)|~v2_waybel_6(esk2_0,esk1_0))&(v1_finset_1(esk3_0)|~v2_waybel_6(esk2_0,esk1_0)))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_waybel_6(esk2_0,esk1_0)))&((esk2_0=k2_yellow_0(esk1_0,esk3_0)|~v2_waybel_6(esk2_0,esk1_0))&(~r2_hidden(esk2_0,esk3_0)|~v2_waybel_6(esk2_0,esk1_0))))&(v2_waybel_6(esk2_0,esk1_0)|(v1_xboole_0(X9)|~v1_finset_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(u1_struct_0(esk1_0)))|(esk2_0!=k2_yellow_0(esk1_0,X9)|r2_hidden(esk2_0,X9))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_34])])])])])])).
fof(c_0_36, plain, ![X112, X113, X114]:(v1_xboole_0(X112)|~m1_subset_1(X113,X112)|~m1_subset_1(X114,X112)|m1_subset_1(k7_domain_1(X112,X113,X114),k1_zfmisc_1(X112))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_domain_1])])])).
fof(c_0_37, plain, ![X64]:(~l1_orders_2(X64)|(~v2_lattice3(X64)|~v2_struct_0(X64))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_lattice3])])])).
fof(c_0_38, plain, ![X135]:(v2_struct_0(X135)|~l1_struct_0(X135)|~v1_xboole_0(u1_struct_0(X135))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_39, plain, ![X141]:(~l1_orders_2(X141)|l1_struct_0(X141)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
cnf(c_0_40, negated_conjecture, (v2_waybel_6(esk2_0,esk1_0)|v1_xboole_0(X1)|r2_hidden(esk2_0,X1)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|esk2_0!=k2_yellow_0(esk1_0,X1)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41, plain, (v1_xboole_0(X1)|m1_subset_1(k7_domain_1(X1,X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_36])).
fof(c_0_42, plain, ![X99, X100, X101]:(v1_xboole_0(X99)|~m1_subset_1(X100,X99)|~m1_subset_1(X101,X99)|k7_domain_1(X99,X100,X101)=k2_tarski(X100,X101)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k7_domain_1])])])).
fof(c_0_43, plain, ![X166, X167]:v1_finset_1(k2_tarski(X166,X167)), inference(variable_rename,[status(thm)],[fc2_finset_1])).
fof(c_0_44, plain, ![X122, X123]:~v1_xboole_0(k2_tarski(X122,X123)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc3_xboole_0])])).
cnf(c_0_45, plain, (~l1_orders_2(X1)|~v2_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_46, negated_conjecture, (v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_47, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_48, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_49, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_50, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(esk2_0,X1)|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,X1)!=esk2_0|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[c_0_40])).
cnf(c_0_51, plain, (v1_xboole_0(X1)|m1_subset_1(k7_domain_1(X1,X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_41])).
cnf(c_0_52, plain, (v1_xboole_0(X1)|k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_53, plain, (v1_finset_1(k2_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_54, plain, (~v1_xboole_0(k2_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_55, plain, (~v2_lattice3(X1)|~l1_orders_2(X1)|~v2_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_45])).
cnf(c_0_56, negated_conjecture, (v2_lattice3(esk1_0)), inference(evalgc,[status(thm)],[c_0_46])).
cnf(c_0_57, negated_conjecture, (l1_orders_2(esk1_0)), inference(evalgc,[status(thm)],[c_0_47])).
cnf(c_0_58, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_48])).
cnf(c_0_59, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_49])).
fof(c_0_60, plain, ![X61, X62, X63]:(~v3_orders_2(X61)|~v4_orders_2(X61)|~v5_orders_2(X61)|~v2_lattice3(X61)|~l1_orders_2(X61)|(~m1_subset_1(X62,u1_struct_0(X61))|(~m1_subset_1(X63,u1_struct_0(X61))|k2_yellow_0(X61,k7_domain_1(u1_struct_0(X61),X62,X63))=k12_lattice3(X61,X62,X63)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t40_yellow_0])])])).
fof(c_0_61, plain, ![X40, X41, X42, X43, X44, X45, X46, X47]:(((~r2_hidden(X43,X42)|(X43=X40|X43=X41)|X42!=k2_tarski(X40,X41))&((X44!=X40|r2_hidden(X44,X42)|X42!=k2_tarski(X40,X41))&(X44!=X41|r2_hidden(X44,X42)|X42!=k2_tarski(X40,X41))))&(((esk10_3(X45,X46,X47)!=X45|~r2_hidden(esk10_3(X45,X46,X47),X47)|X47=k2_tarski(X45,X46))&(esk10_3(X45,X46,X47)!=X46|~r2_hidden(esk10_3(X45,X46,X47),X47)|X47=k2_tarski(X45,X46)))&(r2_hidden(esk10_3(X45,X46,X47),X47)|(esk10_3(X45,X46,X47)=X45|esk10_3(X45,X46,X47)=X46)|X47=k2_tarski(X45,X46)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
cnf(c_0_62, negated_conjecture, (r2_hidden(esk2_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))|v1_xboole_0(k7_domain_1(u1_struct_0(esk1_0),X1,X2))|v2_waybel_6(esk2_0,esk1_0)|v1_xboole_0(u1_struct_0(esk1_0))|k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))!=esk2_0|~v1_finset_1(k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51])])).
cnf(c_0_63, plain, (k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)|v1_xboole_0(X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_52])).
cnf(c_0_64, plain, (v1_finset_1(k2_tarski(X1,X2))), inference(evalgc,[status(thm)],[c_0_53])).
cnf(c_0_65, plain, (~v1_xboole_0(k2_tarski(X1,X2))), inference(evalgc,[status(thm)],[c_0_54])).
cnf(c_0_66, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_56]), c_0_57])])])).
cnf(c_0_67, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_59])])).
cnf(c_0_68, plain, (k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k12_lattice3(X1,X2,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_60])).
fof(c_0_69, plain, ![X82, X83, X84]:(~v5_orders_2(X82)|~v2_lattice3(X82)|~l1_orders_2(X82)|~m1_subset_1(X83,u1_struct_0(X82))|~m1_subset_1(X84,u1_struct_0(X82))|k12_lattice3(X82,X83,X84)=k11_lattice3(X82,X83,X84)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k12_lattice3])])).
fof(c_0_70, plain, ![X24, X25, X26, X27]:((~v2_waybel_6(X25,X24)|(~m1_subset_1(X26,u1_struct_0(X24))|(~m1_subset_1(X27,u1_struct_0(X24))|(X25!=k11_lattice3(X24,X26,X27)|X26=X25|X27=X25)))|~m1_subset_1(X25,u1_struct_0(X24))|(v2_struct_0(X24)|~l1_orders_2(X24)))&((m1_subset_1(esk7_2(X24,X25),u1_struct_0(X24))|v2_waybel_6(X25,X24)|~m1_subset_1(X25,u1_struct_0(X24))|(v2_struct_0(X24)|~l1_orders_2(X24)))&((m1_subset_1(esk8_2(X24,X25),u1_struct_0(X24))|v2_waybel_6(X25,X24)|~m1_subset_1(X25,u1_struct_0(X24))|(v2_struct_0(X24)|~l1_orders_2(X24)))&(((X25=k11_lattice3(X24,esk7_2(X24,X25),esk8_2(X24,X25))|v2_waybel_6(X25,X24)|~m1_subset_1(X25,u1_struct_0(X24))|(v2_struct_0(X24)|~l1_orders_2(X24)))&(esk7_2(X24,X25)!=X25|v2_waybel_6(X25,X24)|~m1_subset_1(X25,u1_struct_0(X24))|(v2_struct_0(X24)|~l1_orders_2(X24))))&(esk8_2(X24,X25)!=X25|v2_waybel_6(X25,X24)|~m1_subset_1(X25,u1_struct_0(X24))|(v2_struct_0(X24)|~l1_orders_2(X24))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_waybel_6])])])])])])).
cnf(c_0_71, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_72, negated_conjecture, (r2_hidden(esk2_0,k2_tarski(X1,X2))|v1_xboole_0(u1_struct_0(esk1_0))|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,k2_tarski(X1,X2))!=esk2_0|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63]), c_0_64])]), c_0_65])])).
cnf(c_0_73, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_66, c_0_67]), c_0_57])])])).
cnf(c_0_74, plain, (k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k12_lattice3(X1,X2,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_68])).
cnf(c_0_75, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_76, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_77, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_78, plain, (k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_79, plain, (X1=k11_lattice3(X2,esk7_2(X2,X1),esk8_2(X2,X1))|v2_waybel_6(X1,X2)|v2_struct_0(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_80, plain, (X1=X4|X1=X3|X2!=k2_tarski(X3,X4)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_71])).
cnf(c_0_81, negated_conjecture, (r2_hidden(esk2_0,k2_tarski(X1,X2))|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,k2_tarski(X1,X2))!=esk2_0|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_82, plain, (k2_yellow_0(X1,k2_tarski(X2,X3))=k12_lattice3(X1,X2,X3)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_63])])).
cnf(c_0_83, negated_conjecture, (v5_orders_2(esk1_0)), inference(evalgc,[status(thm)],[c_0_75])).
cnf(c_0_84, negated_conjecture, (v4_orders_2(esk1_0)), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_85, negated_conjecture, (v3_orders_2(esk1_0)), inference(evalgc,[status(thm)],[c_0_77])).
cnf(c_0_86, plain, (k11_lattice3(X1,X2,X3)=k12_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_78])).
cnf(c_0_87, plain, (k11_lattice3(X2,esk7_2(X2,X1),esk8_2(X2,X1))=X1|v2_struct_0(X2)|v2_waybel_6(X1,X2)|~l1_orders_2(X2)|~m1_subset_1(X1,u1_struct_0(X2))), inference(evalgc,[status(thm)],[c_0_79])).
cnf(c_0_88, plain, (m1_subset_1(esk7_2(X1,X2),u1_struct_0(X1))|v2_waybel_6(X2,X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_89, plain, (m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|v2_waybel_6(X2,X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_90, plain, (X1=X2|X1=X3|~r2_hidden(X1,k2_tarski(X2,X3))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_80])])).
cnf(c_0_91, negated_conjecture, (r2_hidden(esk2_0,k2_tarski(X1,X2))|v2_waybel_6(esk2_0,esk1_0)|k12_lattice3(esk1_0,X1,X2)!=esk2_0|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_57]), c_0_56]), c_0_83]), c_0_84]), c_0_85])]), c_0_73])])).
cnf(c_0_92, plain, (k12_lattice3(X1,esk7_2(X1,X2),esk8_2(X1,X2))=X2|v2_struct_0(X1)|v2_waybel_6(X2,X1)|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk7_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_87])])).
cnf(c_0_93, plain, (v2_struct_0(X1)|v2_waybel_6(X2,X1)|m1_subset_1(esk7_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_88])).
cnf(c_0_94, plain, (v2_struct_0(X1)|v2_waybel_6(X2,X1)|m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_95, negated_conjecture, (esk2_0=X1|esk2_0=X2|v2_waybel_6(esk2_0,esk1_0)|k12_lattice3(esk1_0,X2,X1)!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91])])).
cnf(c_0_96, plain, (k12_lattice3(X1,esk7_2(X1,X2),esk8_2(X1,X2))=X2|v2_waybel_6(X2,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_92, c_0_93]), c_0_94]), c_0_55])).
cnf(c_0_97, negated_conjecture, (esk7_2(esk1_0,X1)=esk2_0|esk8_2(esk1_0,X1)=esk2_0|v2_waybel_6(esk2_0,esk1_0)|v2_waybel_6(X1,esk1_0)|X1!=esk2_0|~m1_subset_1(esk8_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(esk7_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96]), c_0_57]), c_0_56]), c_0_83])])])).
fof(c_0_98, plain, ![X37, X38, X39]:((r2_yellow_0(X37,k2_xboole_0(X38,X39))|(~r2_yellow_0(X37,X38)|~r2_yellow_0(X37,X39))|(~v4_orders_2(X37)|~v5_orders_2(X37)|~v2_lattice3(X37)|~l1_orders_2(X37)))&(k2_yellow_0(X37,k2_xboole_0(X38,X39))=k12_lattice3(X37,k2_yellow_0(X37,X38),k2_yellow_0(X37,X39))|(~r2_yellow_0(X37,X38)|~r2_yellow_0(X37,X39))|(~v4_orders_2(X37)|~v5_orders_2(X37)|~v2_lattice3(X37)|~l1_orders_2(X37)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_yellow_2])])])])).
fof(c_0_99, plain, ![X70, X71, X72]:(~v5_orders_2(X70)|~v2_lattice3(X70)|~l1_orders_2(X70)|~m1_subset_1(X71,u1_struct_0(X70))|~m1_subset_1(X72,u1_struct_0(X70))|k12_lattice3(X70,X71,X72)=k12_lattice3(X70,X72,X71)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k12_lattice3])])).
cnf(c_0_100, plain, (v2_waybel_6(X2,X1)|v2_struct_0(X1)|esk8_2(X1,X2)!=X2|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_101, negated_conjecture, (esk8_2(esk1_0,X1)=esk2_0|esk7_2(esk1_0,X1)=esk2_0|v2_waybel_6(esk2_0,esk1_0)|v2_waybel_6(X1,esk1_0)|X1!=esk2_0|~m1_subset_1(esk7_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_94]), c_0_57])]), c_0_66])])).
cnf(c_0_102, plain, (X3=X1|X4=X1|v2_struct_0(X2)|~v2_waybel_6(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|X1!=k11_lattice3(X2,X3,X4)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_103, plain, (k2_yellow_0(X1,k2_xboole_0(X2,X3))=k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,X3)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_104, plain, (k12_lattice3(X1,X2,X3)=k12_lattice3(X1,X3,X2)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_99])).
fof(c_0_105, plain, ![X78, X79]:(~l1_orders_2(X78)|m1_subset_1(k2_yellow_0(X78,X79),u1_struct_0(X78))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_yellow_0])])).
cnf(c_0_106, plain, (v2_waybel_6(X2,X1)|v2_struct_0(X1)|esk7_2(X1,X2)!=X2|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_107, plain, (v2_struct_0(X1)|v2_waybel_6(X2,X1)|esk8_2(X1,X2)!=X2|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_100])).
cnf(c_0_108, negated_conjecture, (esk7_2(esk1_0,X1)=esk2_0|esk8_2(esk1_0,X1)=esk2_0|v2_waybel_6(esk2_0,esk1_0)|v2_waybel_6(X1,esk1_0)|X1!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_93]), c_0_57])]), c_0_66])])).
cnf(c_0_109, plain, (X1=X4|X1=X3|v2_struct_0(X2)|X1!=k11_lattice3(X2,X3,X4)|~l1_orders_2(X2)|~v2_waybel_6(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))), inference(evalgc,[status(thm)],[c_0_102])).
cnf(c_0_110, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(X2,X3))|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_103])).
cnf(c_0_111, plain, (k12_lattice3(X1,X2,X3)=k12_lattice3(X1,X3,X2)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_104])).
cnf(c_0_112, plain, (m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_105])).
fof(c_0_113, plain, ![X19, X20, X21]:(((((r2_hidden(esk5_3(X19,X20,X21),X21)|(X20=k2_yellow_0(X19,k1_xboole_0)|~v1_finset_1(X21))|(X20!=k2_yellow_0(X19,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21)))|(~v3_orders_2(X19)|~v4_orders_2(X19)|~v5_orders_2(X19)|~v2_lattice3(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,u1_struct_0(X19))|v1_xboole_0(X21)|~v1_finset_1(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X19)))))&(r1_tarski(esk6_3(X19,X20,X21),X21)|(X20=k2_yellow_0(X19,k1_xboole_0)|~v1_finset_1(X21))|(X20!=k2_yellow_0(X19,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21)))|(~v3_orders_2(X19)|~v4_orders_2(X19)|~v5_orders_2(X19)|~v2_lattice3(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,u1_struct_0(X19))|v1_xboole_0(X21)|~v1_finset_1(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X19))))))&(X20!=k2_yellow_0(X19,esk6_3(X19,X20,X21))|(v1_xboole_0(esk6_3(X19,X20,X21))|r2_hidden(X20,esk6_3(X19,X20,X21)))|(X20=k2_yellow_0(X19,k1_xboole_0)|~v1_finset_1(X21))|(X20!=k2_yellow_0(X19,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21)))|(~v3_orders_2(X19)|~v4_orders_2(X19)|~v5_orders_2(X19)|~v2_lattice3(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,u1_struct_0(X19))|v1_xboole_0(X21)|~v1_finset_1(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X19))))))&((X20=k2_yellow_0(X19,k2_xboole_0(esk6_3(X19,X20,X21),k1_tarski(esk5_3(X19,X20,X21))))|(X20=k2_yellow_0(X19,k1_xboole_0)|~v1_finset_1(X21))|(X20!=k2_yellow_0(X19,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21)))|(~v3_orders_2(X19)|~v4_orders_2(X19)|~v5_orders_2(X19)|~v2_lattice3(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,u1_struct_0(X19))|v1_xboole_0(X21)|~v1_finset_1(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X19)))))&((~v1_xboole_0(k2_xboole_0(esk6_3(X19,X20,X21),k1_tarski(esk5_3(X19,X20,X21))))|(X20=k2_yellow_0(X19,k1_xboole_0)|~v1_finset_1(X21))|(X20!=k2_yellow_0(X19,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21)))|(~v3_orders_2(X19)|~v4_orders_2(X19)|~v5_orders_2(X19)|~v2_lattice3(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,u1_struct_0(X19))|v1_xboole_0(X21)|~v1_finset_1(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X19)))))&(~r2_hidden(X20,k2_xboole_0(esk6_3(X19,X20,X21),k1_tarski(esk5_3(X19,X20,X21))))|(X20=k2_yellow_0(X19,k1_xboole_0)|~v1_finset_1(X21))|(X20!=k2_yellow_0(X19,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21)))|(~v3_orders_2(X19)|~v4_orders_2(X19)|~v5_orders_2(X19)|~v2_lattice3(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,u1_struct_0(X19))|v1_xboole_0(X21)|~v1_finset_1(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X19))))))))&(((((r2_hidden(esk5_3(X19,X20,X21),X21)|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X21))|(X20!=k2_yellow_0(X19,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21)))|(~v3_orders_2(X19)|~v4_orders_2(X19)|~v5_orders_2(X19)|~v2_lattice3(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,u1_struct_0(X19))|v1_xboole_0(X21)|~v1_finset_1(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X19)))))&(r1_tarski(esk6_3(X19,X20,X21),X21)|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X21))|(X20!=k2_yellow_0(X19,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21)))|(~v3_orders_2(X19)|~v4_orders_2(X19)|~v5_orders_2(X19)|~v2_lattice3(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,u1_struct_0(X19))|v1_xboole_0(X21)|~v1_finset_1(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X19))))))&(X20!=k2_yellow_0(X19,esk6_3(X19,X20,X21))|(v1_xboole_0(esk6_3(X19,X20,X21))|r2_hidden(X20,esk6_3(X19,X20,X21)))|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X21))|(X20!=k2_yellow_0(X19,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21)))|(~v3_orders_2(X19)|~v4_orders_2(X19)|~v5_orders_2(X19)|~v2_lattice3(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,u1_struct_0(X19))|v1_xboole_0(X21)|~v1_finset_1(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X19))))))&((X20=k2_yellow_0(X19,k2_xboole_0(esk6_3(X19,X20,X21),k1_tarski(esk5_3(X19,X20,X21))))|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X21))|(X20!=k2_yellow_0(X19,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21)))|(~v3_orders_2(X19)|~v4_orders_2(X19)|~v5_orders_2(X19)|~v2_lattice3(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,u1_struct_0(X19))|v1_xboole_0(X21)|~v1_finset_1(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X19)))))&((~v1_xboole_0(k2_xboole_0(esk6_3(X19,X20,X21),k1_tarski(esk5_3(X19,X20,X21))))|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X21))|(X20!=k2_yellow_0(X19,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21)))|(~v3_orders_2(X19)|~v4_orders_2(X19)|~v5_orders_2(X19)|~v2_lattice3(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,u1_struct_0(X19))|v1_xboole_0(X21)|~v1_finset_1(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X19)))))&(~r2_hidden(X20,k2_xboole_0(esk6_3(X19,X20,X21),k1_tarski(esk5_3(X19,X20,X21))))|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X21))|(X20!=k2_yellow_0(X19,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21)))|(~v3_orders_2(X19)|~v4_orders_2(X19)|~v5_orders_2(X19)|~v2_lattice3(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,u1_struct_0(X19))|v1_xboole_0(X21)|~v1_finset_1(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X19))))))))&((((r2_hidden(esk5_3(X19,X20,X21),X21)|(~r2_hidden(X20,k1_xboole_0)|~v1_finset_1(X21))|(X20!=k2_yellow_0(X19,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21)))|(~v3_orders_2(X19)|~v4_orders_2(X19)|~v5_orders_2(X19)|~v2_lattice3(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,u1_struct_0(X19))|v1_xboole_0(X21)|~v1_finset_1(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X19)))))&(r1_tarski(esk6_3(X19,X20,X21),X21)|(~r2_hidden(X20,k1_xboole_0)|~v1_finset_1(X21))|(X20!=k2_yellow_0(X19,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21)))|(~v3_orders_2(X19)|~v4_orders_2(X19)|~v5_orders_2(X19)|~v2_lattice3(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,u1_struct_0(X19))|v1_xboole_0(X21)|~v1_finset_1(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X19))))))&(X20!=k2_yellow_0(X19,esk6_3(X19,X20,X21))|(v1_xboole_0(esk6_3(X19,X20,X21))|r2_hidden(X20,esk6_3(X19,X20,X21)))|(~r2_hidden(X20,k1_xboole_0)|~v1_finset_1(X21))|(X20!=k2_yellow_0(X19,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21)))|(~v3_orders_2(X19)|~v4_orders_2(X19)|~v5_orders_2(X19)|~v2_lattice3(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,u1_struct_0(X19))|v1_xboole_0(X21)|~v1_finset_1(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X19))))))&((X20=k2_yellow_0(X19,k2_xboole_0(esk6_3(X19,X20,X21),k1_tarski(esk5_3(X19,X20,X21))))|(~r2_hidden(X20,k1_xboole_0)|~v1_finset_1(X21))|(X20!=k2_yellow_0(X19,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21)))|(~v3_orders_2(X19)|~v4_orders_2(X19)|~v5_orders_2(X19)|~v2_lattice3(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,u1_struct_0(X19))|v1_xboole_0(X21)|~v1_finset_1(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X19)))))&((~v1_xboole_0(k2_xboole_0(esk6_3(X19,X20,X21),k1_tarski(esk5_3(X19,X20,X21))))|(~r2_hidden(X20,k1_xboole_0)|~v1_finset_1(X21))|(X20!=k2_yellow_0(X19,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21)))|(~v3_orders_2(X19)|~v4_orders_2(X19)|~v5_orders_2(X19)|~v2_lattice3(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,u1_struct_0(X19))|v1_xboole_0(X21)|~v1_finset_1(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X19)))))&(~r2_hidden(X20,k2_xboole_0(esk6_3(X19,X20,X21),k1_tarski(esk5_3(X19,X20,X21))))|(~r2_hidden(X20,k1_xboole_0)|~v1_finset_1(X21))|(X20!=k2_yellow_0(X19,X21)|(v1_xboole_0(X21)|r2_hidden(X20,X21)))|(~v3_orders_2(X19)|~v4_orders_2(X19)|~v5_orders_2(X19)|~v2_lattice3(X19)|~l1_orders_2(X19)|~m1_subset_1(X20,u1_struct_0(X19))|v1_xboole_0(X21)|~v1_finset_1(X21)|~m1_subset_1(X21,k1_zfmisc_1(u1_struct_0(X19)))))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[s2_finset_1__e5_25_1__waybel_6])])])])])).
fof(c_0_114, plain, ![X54, X55, X56]:(~r1_tarski(X54,X55)|~r1_tarski(X55,X56)|r1_tarski(X54,X56)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_115, plain, ![X80, X81]:((~m1_subset_1(X80,k1_zfmisc_1(X81))|r1_tarski(X80,X81))&(~r1_tarski(X80,X81)|m1_subset_1(X80,k1_zfmisc_1(X81)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_116, plain, (v2_struct_0(X1)|v2_waybel_6(X2,X1)|esk7_2(X1,X2)!=X2|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_117, negated_conjecture, (v2_waybel_6(X1,esk1_0)|esk7_2(esk1_0,X1)=esk2_0|v2_waybel_6(esk2_0,esk1_0)|esk2_0!=X1|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108]), c_0_57])]), c_0_66])])).
cnf(c_0_118, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_119, plain, ![X102, X103]:(v1_xboole_0(X102)|~m1_subset_1(X103,X102)|m1_subset_1(k6_domain_1(X102,X103),k1_zfmisc_1(X102))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_domain_1])])])).
fof(c_0_120, plain, ![X92, X93]:(v1_xboole_0(X92)|~m1_subset_1(X93,X92)|k6_domain_1(X92,X93)=k1_tarski(X93)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_121, plain, (X1=X2|X1=X3|v2_struct_0(X4)|X1!=k12_lattice3(X4,X2,X3)|~v2_waybel_6(X1,X4)|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))|~m1_subset_1(X1,u1_struct_0(X4))|~l1_orders_2(X4)|~v2_lattice3(X4)|~v5_orders_2(X4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_86])])).
cnf(c_0_122, plain, (k12_lattice3(X1,k2_yellow_0(X1,X3),k2_yellow_0(X1,X2))=k2_yellow_0(X1,k2_xboole_0(X2,X3))|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111])])).
cnf(c_0_123, plain, (m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_112])).
cnf(c_0_124, plain, (X1=k2_yellow_0(X2,k2_xboole_0(esk6_3(X2,X1,X3),k1_tarski(esk5_3(X2,X1,X3))))|v1_xboole_0(X3)|r2_hidden(X1,X3)|v1_xboole_0(X3)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|X1!=k2_yellow_0(X2,X3)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v2_lattice3(X2)|~l1_orders_2(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_125, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_114])).
cnf(c_0_126, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_115])).
cnf(c_0_127, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_waybel_6(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_128, negated_conjecture, (v2_waybel_6(X1,esk1_0)|v2_waybel_6(esk2_0,esk1_0)|esk2_0!=X1|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117]), c_0_57])]), c_0_66])])).
cnf(c_0_129, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_118])).
cnf(c_0_130, plain, (v1_xboole_0(X1)|m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_131, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_132, plain, (X1=X2|X1=X3|X1!=k12_lattice3(X4,X3,X2)|~v2_waybel_6(X1,X4)|~m1_subset_1(X2,u1_struct_0(X4))|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X1,u1_struct_0(X4))|~l1_orders_2(X4)|~v2_lattice3(X4)|~v5_orders_2(X4)), inference(csr,[status(thm)],[c_0_121, c_0_55])).
cnf(c_0_133, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(X3,X2))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_122, c_0_123]), c_0_123])).
cnf(c_0_134, plain, (X1=k2_yellow_0(X2,k2_xboole_0(esk6_3(X2,X1,X3),k1_tarski(esk5_3(X2,X1,X3))))|v1_xboole_0(X3)|r2_hidden(X1,X3)|X1!=k2_yellow_0(X2,X3)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v2_lattice3(X2)|~l1_orders_2(X2)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(cn,[status(thm)],[c_0_124])).
cnf(c_0_135, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
fof(c_0_136, plain, ![X51, X52]:((~v2_lattice3(X51)|(v1_xboole_0(X52)|~v1_finset_1(X52)|~m1_subset_1(X52,k1_zfmisc_1(u1_struct_0(X51)))|r2_yellow_0(X51,X52))|(v2_struct_0(X51)|~v3_orders_2(X51)|~v4_orders_2(X51)|~v5_orders_2(X51)|~l1_orders_2(X51)))&((((~v1_xboole_0(esk11_1(X51))|v2_lattice3(X51)|(v2_struct_0(X51)|~v3_orders_2(X51)|~v4_orders_2(X51)|~v5_orders_2(X51)|~l1_orders_2(X51)))&(v1_finset_1(esk11_1(X51))|v2_lattice3(X51)|(v2_struct_0(X51)|~v3_orders_2(X51)|~v4_orders_2(X51)|~v5_orders_2(X51)|~l1_orders_2(X51))))&(m1_subset_1(esk11_1(X51),k1_zfmisc_1(u1_struct_0(X51)))|v2_lattice3(X51)|(v2_struct_0(X51)|~v3_orders_2(X51)|~v4_orders_2(X51)|~v5_orders_2(X51)|~l1_orders_2(X51))))&(~r2_yellow_0(X51,esk11_1(X51))|v2_lattice3(X51)|(v2_struct_0(X51)|~v3_orders_2(X51)|~v4_orders_2(X51)|~v5_orders_2(X51)|~l1_orders_2(X51))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t55_yellow_0])])])])])])).
cnf(c_0_137, plain, (r1_tarski(X1,X3)|~r1_tarski(X2,X3)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_125])).
cnf(c_0_138, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_126])).
cnf(c_0_139, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_waybel_6(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_127])).
cnf(c_0_140, negated_conjecture, (v2_waybel_6(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129])])).
cnf(c_0_141, plain, (v1_xboole_0(X1)|m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_130])).
cnf(c_0_142, plain, (k6_domain_1(X1,X2)=k1_tarski(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_131])).
fof(c_0_143, plain, ![X98]:r1_tarski(X98,X98), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r1_tarski])])).
fof(c_0_144, plain, ![X49, X50]:((k1_yellow_0(X49,k6_domain_1(u1_struct_0(X49),X50))=X50|~m1_subset_1(X50,u1_struct_0(X49))|(v2_struct_0(X49)|~v3_orders_2(X49)|~v5_orders_2(X49)|~l1_orders_2(X49)))&(k2_yellow_0(X49,k6_domain_1(u1_struct_0(X49),X50))=X50|~m1_subset_1(X50,u1_struct_0(X49))|(v2_struct_0(X49)|~v3_orders_2(X49)|~v5_orders_2(X49)|~l1_orders_2(X49)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t39_yellow_0])])])])])).
cnf(c_0_145, plain, (X1=k2_yellow_0(X2,X3)|X1=k2_yellow_0(X2,X4)|X1!=k2_yellow_0(X2,k2_xboole_0(X3,X4))|~v2_waybel_6(X1,X2)|~m1_subset_1(k2_yellow_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(k2_yellow_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~r2_yellow_0(X2,X4)|~r2_yellow_0(X2,X3)|~v4_orders_2(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_133])])).
cnf(c_0_146, plain, (k2_yellow_0(X2,k2_xboole_0(esk6_3(X2,X1,X3),k1_tarski(esk5_3(X2,X1,X3))))=X1|v1_xboole_0(X3)|r2_hidden(X1,X3)|X1!=k2_yellow_0(X2,X3)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v2_lattice3(X2)|~l1_orders_2(X2)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(evalgc,[status(thm)],[c_0_134])).
cnf(c_0_147, plain, (v1_xboole_0(k1_xboole_0)), inference(evalgc,[status(thm)],[c_0_135])).
cnf(c_0_148, plain, (v1_xboole_0(X2)|r2_yellow_0(X1,X2)|v2_struct_0(X1)|~v2_lattice3(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_136])).
fof(c_0_149, plain, ![X151]:v1_finset_1(k1_tarski(X151)), inference(variable_rename,[status(thm)],[fc1_finset_1])).
fof(c_0_150, plain, ![X108]:~v1_xboole_0(k1_tarski(X108)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_xboole_0])])).
cnf(c_0_151, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_115])).
cnf(c_0_152, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_138])])).
cnf(c_0_153, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_139, c_0_140])])])).
cnf(c_0_154, plain, (v1_xboole_0(X1)|m1_subset_1(k1_tarski(X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_142])])).
cnf(c_0_155, plain, (r1_tarski(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_143])).
cnf(c_0_156, negated_conjecture, (~v1_xboole_0(esk3_0)|~v2_waybel_6(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_157, plain, (k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_144])).
cnf(c_0_158, plain, (X1=k2_yellow_0(X2,X3)|X1=k2_yellow_0(X2,X4)|X1!=k2_yellow_0(X2,k2_xboole_0(X4,X3))|~r2_yellow_0(X2,X3)|~r2_yellow_0(X2,X4)|~v2_waybel_6(X1,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_145, c_0_123]), c_0_123])).
cnf(c_0_159, plain, (k2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X2,X3),k1_tarski(esk5_3(X1,X2,X3))))=X2|r2_hidden(X2,X3)|v1_xboole_0(X3)|X2!=k2_yellow_0(X1,X3)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_146, c_0_147])])).
cnf(c_0_160, plain, (v1_xboole_0(X2)|v2_struct_0(X1)|r2_yellow_0(X1,X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[c_0_148])).
cnf(c_0_161, plain, (v1_finset_1(k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_149])).
cnf(c_0_162, plain, (~v1_xboole_0(k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_150])).
cnf(c_0_163, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_151])).
cnf(c_0_164, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~r1_tarski(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_153])])).
cnf(c_0_165, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~r1_tarski(X1,k1_tarski(X3))|~m1_subset_1(X3,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_154])])).
cnf(c_0_166, plain, (r1_tarski(X1,X1)), inference(evalgc,[status(thm)],[c_0_155])).
cnf(c_0_167, negated_conjecture, (~v1_xboole_0(esk3_0)|~v2_waybel_6(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_156])).
cnf(c_0_168, plain, (k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2|v2_struct_0(X1)|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_157])).
cnf(c_0_169, plain, (k2_yellow_0(X2,k1_tarski(esk5_3(X2,X1,X3)))=X1|k2_yellow_0(X2,esk6_3(X2,X1,X3))=X1|r2_hidden(X1,X3)|v1_xboole_0(X3)|~r2_yellow_0(X2,k1_tarski(esk5_3(X2,X1,X3)))|~r2_yellow_0(X2,esk6_3(X2,X1,X3))|~v2_waybel_6(X1,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|X1!=k2_yellow_0(X2,X3)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v3_orders_2(X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159])])])).
cnf(c_0_170, plain, (r2_yellow_0(X1,X2)|v1_xboole_0(X2)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_160, c_0_55])).
cnf(c_0_171, plain, (v1_finset_1(k1_tarski(X1))), inference(evalgc,[status(thm)],[c_0_161])).
cnf(c_0_172, plain, (~v1_xboole_0(k1_tarski(X1))), inference(evalgc,[status(thm)],[c_0_162])).
cnf(c_0_173, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_tarski(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_164])])).
cnf(c_0_174, plain, (r1_tarski(k1_tarski(X1),X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_166])])).
cnf(c_0_175, negated_conjecture, (~v1_xboole_0(esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_167, c_0_140])])])).
cnf(c_0_176, plain, (r1_tarski(esk6_3(X1,X2,X3),X3)|v1_xboole_0(X3)|r2_hidden(X2,X3)|v1_xboole_0(X3)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|X2!=k2_yellow_0(X1,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_177, plain, (k2_yellow_0(X1,k1_tarski(X2))=X2|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_142])])).
cnf(c_0_178, plain, (k2_yellow_0(X1,k1_tarski(esk5_3(X1,X2,X3)))=X2|k2_yellow_0(X1,esk6_3(X1,X2,X3))=X2|r2_hidden(X2,X3)|v1_xboole_0(X3)|X2!=k2_yellow_0(X1,X3)|~r2_yellow_0(X1,esk6_3(X1,X2,X3))|~v1_finset_1(X3)|~v2_waybel_6(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(esk5_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_170]), c_0_171])]), c_0_172])])).
cnf(c_0_179, negated_conjecture, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_174]), c_0_175])])).
cnf(c_0_180, plain, (v1_xboole_0(X3)|r2_hidden(X2,X3)|r1_tarski(esk6_3(X1,X2,X3),X3)|X2!=k2_yellow_0(X1,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(cn,[status(thm)],[c_0_176])).
cnf(c_0_181, negated_conjecture, (v1_finset_1(esk3_0)|~v2_waybel_6(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_182, plain, (k2_yellow_0(X1,k1_tarski(X2))=X2|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_177, c_0_67])).
cnf(c_0_183, negated_conjecture, (k2_yellow_0(esk1_0,k1_tarski(esk5_3(esk1_0,X1,X2)))=X1|k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))=X1|r2_hidden(X1,X2)|v1_xboole_0(X2)|X1!=k2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))|~v1_finset_1(X2)|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(esk5_3(esk1_0,X1,X2),esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_179]), c_0_57]), c_0_56]), c_0_83]), c_0_84]), c_0_85])])])).
cnf(c_0_184, plain, (v1_xboole_0(X3)|r2_hidden(X2,X3)|r1_tarski(esk6_3(X1,X2,X3),X3)|X2!=k2_yellow_0(X1,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[c_0_180])).
cnf(c_0_185, negated_conjecture, (v1_finset_1(esk3_0)|~v2_waybel_6(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_181])).
cnf(c_0_186, negated_conjecture, (esk2_0=k2_yellow_0(esk1_0,esk3_0)|~v2_waybel_6(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_187, plain, ![X125, X126]:(~v1_finset_1(X125)|(~m1_subset_1(X126,k1_zfmisc_1(X125))|v1_finset_1(X126))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_finset_1])])])).
cnf(c_0_188, plain, (v1_xboole_0(esk6_3(X2,X1,X3))|r2_hidden(X1,esk6_3(X2,X1,X3))|v1_xboole_0(X3)|r2_hidden(X1,X3)|v1_xboole_0(X3)|X1!=k2_yellow_0(X2,esk6_3(X2,X1,X3))|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|X1!=k2_yellow_0(X2,X3)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v2_lattice3(X2)|~l1_orders_2(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_189, negated_conjecture, (esk5_3(esk1_0,X1,X2)=X1|k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))=X1|r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(esk5_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|X1!=k2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))|~v1_finset_1(X2)|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(esk5_3(esk1_0,X1,X2),esk3_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_183]), c_0_57]), c_0_83]), c_0_85])]), c_0_66])])).
cnf(c_0_190, plain, (r1_tarski(esk6_3(X1,X2,X3),X3)|r2_hidden(X2,X3)|v1_xboole_0(X3)|X2!=k2_yellow_0(X1,X3)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_184, c_0_147])])).
cnf(c_0_191, negated_conjecture, (v1_finset_1(esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_185, c_0_140])])])).
cnf(c_0_192, negated_conjecture, (k2_yellow_0(esk1_0,esk3_0)=esk2_0|~v2_waybel_6(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_186])).
cnf(c_0_193, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_187])).
fof(c_0_194, plain, ![X67, X68, X69]:(~r2_hidden(X67,X68)|~m1_subset_1(X68,k1_zfmisc_1(X69))|m1_subset_1(X67,X69)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_195, plain, (r2_hidden(esk5_3(X1,X2,X3),X3)|v1_xboole_0(X3)|r2_hidden(X2,X3)|v1_xboole_0(X3)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|X2!=k2_yellow_0(X1,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_196, plain, (v1_xboole_0(X3)|r2_hidden(X1,X3)|v1_xboole_0(X3)|~r2_hidden(X1,k2_xboole_0(esk6_3(X2,X1,X3),k1_tarski(esk5_3(X2,X1,X3))))|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|X1!=k2_yellow_0(X2,X3)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v2_lattice3(X2)|~l1_orders_2(X2)|~m1_subset_1(X1,u1_struct_0(X2))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(split_conjunct,[status(thm)],[c_0_113])).
fof(c_0_197, plain, ![X10, X11, X12, X13, X14, X15, X16, X17]:(((~r2_hidden(X13,X12)|(r2_hidden(X13,X10)|r2_hidden(X13,X11))|X12!=k2_xboole_0(X10,X11))&((~r2_hidden(X14,X10)|r2_hidden(X14,X12)|X12!=k2_xboole_0(X10,X11))&(~r2_hidden(X14,X11)|r2_hidden(X14,X12)|X12!=k2_xboole_0(X10,X11))))&(((~r2_hidden(esk4_3(X15,X16,X17),X15)|~r2_hidden(esk4_3(X15,X16,X17),X17)|X17=k2_xboole_0(X15,X16))&(~r2_hidden(esk4_3(X15,X16,X17),X16)|~r2_hidden(esk4_3(X15,X16,X17),X17)|X17=k2_xboole_0(X15,X16)))&(r2_hidden(esk4_3(X15,X16,X17),X17)|(r2_hidden(esk4_3(X15,X16,X17),X15)|r2_hidden(esk4_3(X15,X16,X17),X16))|X17=k2_xboole_0(X15,X16)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])).
cnf(c_0_198, plain, (v1_xboole_0(X3)|r2_hidden(X1,X3)|v1_xboole_0(esk6_3(X2,X1,X3))|r2_hidden(X1,esk6_3(X2,X1,X3))|X1!=k2_yellow_0(X2,X3)|X1!=k2_yellow_0(X2,esk6_3(X2,X1,X3))|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v2_lattice3(X2)|~l1_orders_2(X2)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(cn,[status(thm)],[c_0_188])).
cnf(c_0_199, negated_conjecture, (k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))=X1|esk5_3(esk1_0,X1,X2)=X1|r2_hidden(X1,X2)|v1_xboole_0(X2)|v1_xboole_0(esk6_3(esk1_0,X1,X2))|X1!=k2_yellow_0(esk1_0,X2)|~v1_finset_1(X2)|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(esk5_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(esk5_3(esk1_0,X1,X2),esk3_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v1_finset_1(esk6_3(esk1_0,X1,X2))|~m1_subset_1(esk6_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_170]), c_0_57]), c_0_56]), c_0_83]), c_0_84]), c_0_85])])])).
cnf(c_0_200, negated_conjecture, (m1_subset_1(esk6_3(X1,X2,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|r2_hidden(X2,esk3_0)|X2!=k2_yellow_0(X1,esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_190]), c_0_191])]), c_0_175])])).
cnf(c_0_201, negated_conjecture, (k2_yellow_0(esk1_0,esk3_0)=esk2_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_192, c_0_140])])])).
cnf(c_0_202, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_193])).
cnf(c_0_203, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_zfmisc_1(X3))|r2_hidden(X2,X3)|v1_xboole_0(X3)|X2!=k2_yellow_0(X1,X3)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_190])])).
cnf(c_0_204, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_194])).
cnf(c_0_205, plain, (v1_xboole_0(X3)|r2_hidden(X2,X3)|r2_hidden(esk5_3(X1,X2,X3),X3)|X2!=k2_yellow_0(X1,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(cn,[status(thm)],[c_0_195])).
cnf(c_0_206, plain, (v1_xboole_0(X3)|r2_hidden(X1,X3)|X1!=k2_yellow_0(X2,X3)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v2_lattice3(X2)|~l1_orders_2(X2)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,k2_xboole_0(esk6_3(X2,X1,X3),k1_tarski(esk5_3(X2,X1,X3))))), inference(cn,[status(thm)],[c_0_196])).
cnf(c_0_207, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_197])).
cnf(c_0_208, plain, (v1_xboole_0(X3)|r2_hidden(X1,X3)|v1_xboole_0(esk6_3(X2,X1,X3))|r2_hidden(X1,esk6_3(X2,X1,X3))|X1!=k2_yellow_0(X2,X3)|k2_yellow_0(X2,esk6_3(X2,X1,X3))!=X1|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v2_lattice3(X2)|~l1_orders_2(X2)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))), inference(evalgc,[status(thm)],[c_0_198])).
cnf(c_0_209, negated_conjecture, (k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,esk3_0))=X1|esk5_3(esk1_0,X1,esk3_0)=X1|r2_hidden(X1,esk3_0)|v1_xboole_0(esk6_3(esk1_0,X1,esk3_0))|X1!=esk2_0|~v1_finset_1(esk6_3(esk1_0,X1,esk3_0))|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(esk5_3(esk1_0,X1,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk5_3(esk1_0,X1,esk3_0),esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_200]), c_0_201]), c_0_191]), c_0_153]), c_0_57]), c_0_56]), c_0_83]), c_0_84]), c_0_85])]), c_0_175])])).
cnf(c_0_210, plain, (v1_finset_1(esk6_3(X1,X2,X3))|r2_hidden(X2,X3)|v1_xboole_0(X3)|~v1_finset_1(X3)|X2!=k2_yellow_0(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_203])])).
cnf(c_0_211, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_204])).
cnf(c_0_212, plain, (v1_xboole_0(X3)|r2_hidden(X2,X3)|r2_hidden(esk5_3(X1,X2,X3),X3)|X2!=k2_yellow_0(X1,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[c_0_205])).
cnf(c_0_213, plain, (v1_xboole_0(X3)|r2_hidden(X1,X3)|X1!=k2_yellow_0(X2,X3)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v2_lattice3(X2)|~l1_orders_2(X2)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,k2_xboole_0(esk6_3(X2,X1,X3),k1_tarski(esk5_3(X2,X1,X3))))), inference(evalgc,[status(thm)],[c_0_206])).
cnf(c_0_214, plain, (r2_hidden(X1,X3)|X3!=k2_xboole_0(X2,X4)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_207])).
cnf(c_0_215, plain, (r2_hidden(X1,esk6_3(X2,X1,X3))|r2_hidden(X1,X3)|v1_xboole_0(esk6_3(X2,X1,X3))|v1_xboole_0(X3)|k2_yellow_0(X2,esk6_3(X2,X1,X3))!=X1|X1!=k2_yellow_0(X2,X3)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_208, c_0_147])])).
cnf(c_0_216, negated_conjecture, (k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,esk3_0))=X1|esk5_3(esk1_0,X1,esk3_0)=X1|r2_hidden(X1,esk3_0)|v1_xboole_0(esk6_3(esk1_0,X1,esk3_0))|X1!=esk2_0|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(esk5_3(esk1_0,X1,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk5_3(esk1_0,X1,esk3_0),esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_210]), c_0_201]), c_0_191]), c_0_153]), c_0_57]), c_0_56]), c_0_83]), c_0_84]), c_0_85])]), c_0_175])])).
cnf(c_0_217, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_153])])).
cnf(c_0_218, plain, (r2_hidden(esk5_3(X1,X2,X3),X3)|r2_hidden(X2,X3)|v1_xboole_0(X3)|X2!=k2_yellow_0(X1,X3)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_212, c_0_147])])).
fof(c_0_219, plain, ![X59, X60]:(~r2_hidden(X59,X60)|m1_subset_1(X59,X60)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_220, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|X1!=k2_yellow_0(X3,X2)|~r2_hidden(X1,k2_xboole_0(esk6_3(X3,X1,X2),k1_tarski(esk5_3(X3,X1,X2))))|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)|~v2_lattice3(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_213, c_0_147])])).
cnf(c_0_221, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_214])])).
cnf(c_0_222, negated_conjecture, (r2_hidden(X1,esk6_3(esk1_0,X1,esk3_0))|r2_hidden(X1,esk3_0)|v1_xboole_0(esk6_3(esk1_0,X1,esk3_0))|esk5_3(esk1_0,X1,esk3_0)=X1|X1!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(esk5_3(esk1_0,X1,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk5_3(esk1_0,X1,esk3_0),esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_216]), c_0_201]), c_0_191]), c_0_153]), c_0_57]), c_0_56]), c_0_83]), c_0_84]), c_0_85])]), c_0_175])])).
cnf(c_0_223, negated_conjecture, (m1_subset_1(esk5_3(X1,X2,esk3_0),u1_struct_0(esk1_0))|r2_hidden(X2,esk3_0)|X2!=k2_yellow_0(X1,esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_218]), c_0_191])]), c_0_175])])).
cnf(c_0_224, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_219])).
fof(c_0_225, plain, ![X107]:(~v1_xboole_0(X107)|X107=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_226, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|X1!=k2_yellow_0(X3,X2)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)|~v2_lattice3(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)|~r2_hidden(X1,esk6_3(X3,X1,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_221])])).
cnf(c_0_227, negated_conjecture, (esk5_3(esk1_0,X1,esk3_0)=X1|r2_hidden(X1,esk6_3(esk1_0,X1,esk3_0))|r2_hidden(X1,esk3_0)|v1_xboole_0(esk6_3(esk1_0,X1,esk3_0))|X1!=esk2_0|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(esk5_3(esk1_0,X1,esk3_0),esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_223]), c_0_201]), c_0_153]), c_0_57]), c_0_56]), c_0_83]), c_0_84]), c_0_85])])])).
cnf(c_0_228, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_224])).
cnf(c_0_229, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_225])).
cnf(c_0_230, negated_conjecture, (r2_hidden(X1,esk3_0)|esk5_3(esk1_0,X1,esk3_0)=X1|v1_xboole_0(esk6_3(esk1_0,X1,esk3_0))|X1!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(esk5_3(esk1_0,X1,esk3_0),esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_227]), c_0_201]), c_0_191]), c_0_153]), c_0_57]), c_0_56]), c_0_83]), c_0_84]), c_0_85])]), c_0_175])])).
cnf(c_0_231, plain, (m1_subset_1(esk5_3(X1,X2,X3),X3)|r2_hidden(X2,X3)|v1_xboole_0(X3)|X2!=k2_yellow_0(X1,X3)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_218])])).
cnf(c_0_232, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_229])).
cnf(c_0_233, negated_conjecture, (esk5_3(esk1_0,X1,esk3_0)=X1|r2_hidden(X1,esk3_0)|v1_xboole_0(esk6_3(esk1_0,X1,esk3_0))|X1!=esk2_0|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_231]), c_0_201]), c_0_191]), c_0_153]), c_0_57]), c_0_56]), c_0_83]), c_0_84]), c_0_85])]), c_0_175])])).
cnf(c_0_234, negated_conjecture, (~r2_hidden(esk2_0,esk3_0)|~v2_waybel_6(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_35])).
fof(c_0_235, plain, ![X118]:k2_xboole_0(X118,k1_xboole_0)=X118, inference(variable_rename,[status(thm)],[t1_boole])).
fof(c_0_236, plain, ![X90, X91]:k2_xboole_0(X90,X91)=k2_xboole_0(X91,X90), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_237, negated_conjecture, (esk6_3(esk1_0,X1,esk3_0)=k1_xboole_0|esk5_3(esk1_0,X1,esk3_0)=X1|r2_hidden(X1,esk3_0)|X1!=esk2_0|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_233])])).
cnf(c_0_238, negated_conjecture, (~v2_waybel_6(esk2_0,esk1_0)|~r2_hidden(esk2_0,esk3_0)), inference(evalgc,[status(thm)],[c_0_234])).
cnf(c_0_239, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_235])).
cnf(c_0_240, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_236])).
cnf(c_0_241, negated_conjecture, (r2_hidden(X1,esk3_0)|esk6_3(esk1_0,X1,esk3_0)=k1_xboole_0|X1!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v2_waybel_6(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_237]), c_0_201]), c_0_191]), c_0_153]), c_0_57]), c_0_56]), c_0_83]), c_0_84]), c_0_85])]), c_0_175])])).
cnf(c_0_242, negated_conjecture, (~r2_hidden(esk2_0,esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_238, c_0_140])])])).
cnf(c_0_243, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(evalgc,[status(thm)],[c_0_239])).
cnf(c_0_244, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(evalgc,[status(thm)],[c_0_240])).
fof(c_0_245, plain, ![X30, X31, X32, X33, X34, X35]:(((~r2_hidden(X32,X31)|X32=X30|X31!=k1_tarski(X30))&(X33!=X30|r2_hidden(X33,X31)|X31!=k1_tarski(X30)))&((~r2_hidden(esk9_2(X34,X35),X35)|esk9_2(X34,X35)!=X34|X35=k1_tarski(X34))&(r2_hidden(esk9_2(X34,X35),X35)|esk9_2(X34,X35)=X34|X35=k1_tarski(X34)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_246, negated_conjecture, (esk6_3(esk1_0,esk2_0,esk3_0)=k1_xboole_0), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_140]), c_0_129])]), c_0_242])])).
cnf(c_0_247, plain, (k2_xboole_0(k1_xboole_0,X1)=X1), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_244])])).
cnf(c_0_248, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_245])).
cnf(c_0_249, negated_conjecture, (k2_yellow_0(esk1_0,k1_tarski(esk5_3(esk1_0,esk2_0,esk3_0)))=esk2_0), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_246]), c_0_247]), c_0_201]), c_0_191]), c_0_153]), c_0_129]), c_0_57]), c_0_56]), c_0_83]), c_0_84]), c_0_85])]), c_0_242]), c_0_175])])).
cnf(c_0_250, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(evalgc,[status(thm)],[c_0_248])).
cnf(c_0_251, negated_conjecture, (esk5_3(esk1_0,esk2_0,esk3_0)=esk2_0|~m1_subset_1(esk5_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_249]), c_0_57]), c_0_83]), c_0_85])]), c_0_66])])).
cnf(c_0_252, plain, (r2_hidden(X1,X2)|X2!=k1_tarski(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_250])])).
cnf(c_0_253, negated_conjecture, (~r2_hidden(esk2_0,k1_tarski(esk5_3(esk1_0,esk2_0,esk3_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_246]), c_0_201]), c_0_247]), c_0_191]), c_0_153]), c_0_129]), c_0_57]), c_0_56]), c_0_83]), c_0_84]), c_0_85])]), c_0_242]), c_0_175])])).
cnf(c_0_254, negated_conjecture, (esk5_3(esk1_0,esk2_0,esk3_0)=esk2_0), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_223]), c_0_201]), c_0_153]), c_0_129]), c_0_57]), c_0_56]), c_0_83]), c_0_84]), c_0_85])]), c_0_242])])).
cnf(c_0_255, plain, (r2_hidden(X1,k1_tarski(X1))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_252])])).
cnf(c_0_256, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_253, c_0_254]), c_0_255])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 257
# Proof object clause steps            : 189
# Proof object formula steps           : 68
# Proof object conjectures             : 68
# Proof object clause conjectures      : 65
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 55
# Proof object initial formulas used   : 34
# Proof object generating inferences   : 55
# Proof object simplifying inferences  : 210
# Training examples: 116 positive, 567 negative
# Training: Positive examples begin
cnf(c_0_85, negated_conjecture, (v3_orders_2(esk1_0))).# trainpos
cnf(c_0_84, negated_conjecture, (v4_orders_2(esk1_0))).# trainpos
cnf(c_0_83, negated_conjecture, (v5_orders_2(esk1_0))).# trainpos
cnf(c_0_56, negated_conjecture, (v2_lattice3(esk1_0))).# trainpos
cnf(c_0_57, negated_conjecture, (l1_orders_2(esk1_0))).# trainpos
cnf(c_0_147, plain, (v1_xboole_0(k1_xboole_0))).# trainpos
cnf(c_0_167, negated_conjecture, (~v1_xboole_0(esk3_0)|~v2_waybel_6(esk2_0,esk1_0))).# trainpos
cnf(c_0_129, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_238, negated_conjecture, (~v2_waybel_6(esk2_0,esk1_0)|~r2_hidden(esk2_0,esk3_0))).# trainpos
cnf(c_0_59, plain, (l1_struct_0(X1)|~l1_orders_2(X1))).# trainpos
cnf(c_0_232, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_171, plain, (v1_finset_1(k1_tarski(X1)))).# trainpos
cnf(c_0_243, plain, (k2_xboole_0(X1,k1_xboole_0)=X1)).# trainpos
cnf(c_0_166, plain, (r1_tarski(X1,X1))).# trainpos
cnf(c_0_172, plain, (~v1_xboole_0(k1_tarski(X1)))).# trainpos
cnf(c_0_64, plain, (v1_finset_1(k2_tarski(X1,X2)))).# trainpos
cnf(c_0_185, negated_conjecture, (v1_finset_1(esk3_0)|~v2_waybel_6(esk2_0,esk1_0))).# trainpos
cnf(c_0_192, negated_conjecture, (k2_yellow_0(esk1_0,esk3_0)=esk2_0|~v2_waybel_6(esk2_0,esk1_0))).# trainpos
cnf(c_0_139, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_waybel_6(esk2_0,esk1_0))).# trainpos
cnf(c_0_55, plain, (~v2_lattice3(X1)|~l1_orders_2(X1)|~v2_struct_0(X1))).# trainpos
cnf(c_0_65, plain, (~v1_xboole_0(k2_tarski(X1,X2)))).# trainpos
cnf(c_0_244, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1))).# trainpos
cnf(c_0_228, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_250, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2))).# trainpos
cnf(c_0_50, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(esk2_0,X1)|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,X1)!=esk2_0|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_163, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_58, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).# trainpos
cnf(c_0_138, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_123, plain, (m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1))).# trainpos
cnf(c_0_202, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_211, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_142, plain, (k6_domain_1(X1,X2)=k1_tarski(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_80, plain, (X1=X4|X1=X3|X2!=k2_tarski(X3,X4)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_137, plain, (r1_tarski(X1,X3)|~r1_tarski(X2,X3)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_214, plain, (r2_hidden(X1,X3)|X3!=k2_xboole_0(X2,X4)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_141, plain, (v1_xboole_0(X1)|m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_160, plain, (v1_xboole_0(X2)|v2_struct_0(X1)|r2_yellow_0(X1,X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_63, plain, (k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)|v1_xboole_0(X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_116, plain, (v2_struct_0(X1)|v2_waybel_6(X2,X1)|esk7_2(X1,X2)!=X2|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_107, plain, (v2_struct_0(X1)|v2_waybel_6(X2,X1)|esk8_2(X1,X2)!=X2|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_51, plain, (v1_xboole_0(X1)|m1_subset_1(k7_domain_1(X1,X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_93, plain, (v2_struct_0(X1)|v2_waybel_6(X2,X1)|m1_subset_1(esk7_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_94, plain, (v2_struct_0(X1)|v2_waybel_6(X2,X1)|m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_168, plain, (k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2|v2_struct_0(X1)|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_86, plain, (k11_lattice3(X1,X2,X3)=k12_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_110, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(X2,X3))|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2))).# trainpos
cnf(c_0_87, plain, (k11_lattice3(X2,esk7_2(X2,X1),esk8_2(X2,X1))=X1|v2_struct_0(X2)|v2_waybel_6(X1,X2)|~l1_orders_2(X2)|~m1_subset_1(X1,u1_struct_0(X2)))).# trainpos
cnf(c_0_111, plain, (k12_lattice3(X1,X2,X3)=k12_lattice3(X1,X3,X2)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_109, plain, (X1=X4|X1=X3|v2_struct_0(X2)|X1!=k11_lattice3(X2,X3,X4)|~l1_orders_2(X2)|~v2_waybel_6(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2)))).# trainpos
cnf(c_0_74, plain, (k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k12_lattice3(X1,X2,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_212, plain, (v1_xboole_0(X3)|r2_hidden(X2,X3)|r2_hidden(esk5_3(X1,X2,X3),X3)|X2!=k2_yellow_0(X1,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_184, plain, (v1_xboole_0(X3)|r2_hidden(X2,X3)|r1_tarski(esk6_3(X1,X2,X3),X3)|X2!=k2_yellow_0(X1,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_146, plain, (k2_yellow_0(X2,k2_xboole_0(esk6_3(X2,X1,X3),k1_tarski(esk5_3(X2,X1,X3))))=X1|v1_xboole_0(X3)|r2_hidden(X1,X3)|X1!=k2_yellow_0(X2,X3)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v2_lattice3(X2)|~l1_orders_2(X2)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))).# trainpos
cnf(c_0_213, plain, (v1_xboole_0(X3)|r2_hidden(X1,X3)|X1!=k2_yellow_0(X2,X3)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v2_lattice3(X2)|~l1_orders_2(X2)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~r2_hidden(X1,k2_xboole_0(esk6_3(X2,X1,X3),k1_tarski(esk5_3(X2,X1,X3)))))).# trainpos
cnf(c_0_208, plain, (v1_xboole_0(X3)|r2_hidden(X1,X3)|v1_xboole_0(esk6_3(X2,X1,X3))|r2_hidden(X1,esk6_3(X2,X1,X3))|X1!=k2_yellow_0(X2,X3)|k2_yellow_0(X2,esk6_3(X2,X1,X3))!=X1|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~v2_lattice3(X2)|~l1_orders_2(X2)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2))))).# trainpos
cnf(c_0_67, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1))).# trainpos
cnf(c_0_152, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_66, negated_conjecture, (~v2_struct_0(esk1_0))).# trainpos
cnf(c_0_252, plain, (r2_hidden(X1,X2)|X2!=k1_tarski(X1))).# trainpos
cnf(c_0_62, negated_conjecture, (r2_hidden(esk2_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))|v1_xboole_0(k7_domain_1(u1_struct_0(esk1_0),X1,X2))|v2_waybel_6(esk2_0,esk1_0)|v1_xboole_0(u1_struct_0(esk1_0))|k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))!=esk2_0|~v1_finset_1(k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_247, plain, (k2_xboole_0(k1_xboole_0,X1)=X1)).# trainpos
cnf(c_0_221, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_90, plain, (X1=X2|X1=X3|~r2_hidden(X1,k2_tarski(X2,X3)))).# trainpos
cnf(c_0_92, plain, (k12_lattice3(X1,esk7_2(X1,X2),esk8_2(X1,X2))=X2|v2_struct_0(X1)|v2_waybel_6(X2,X1)|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk7_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_121, plain, (X1=X2|X1=X3|v2_struct_0(X4)|X1!=k12_lattice3(X4,X2,X3)|~v2_waybel_6(X1,X4)|~m1_subset_1(X3,u1_struct_0(X4))|~m1_subset_1(X2,u1_struct_0(X4))|~m1_subset_1(X1,u1_struct_0(X4))|~l1_orders_2(X4)|~v2_lattice3(X4)|~v5_orders_2(X4))).# trainpos
cnf(c_0_177, plain, (k2_yellow_0(X1,k1_tarski(X2))=X2|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_122, plain, (k12_lattice3(X1,k2_yellow_0(X1,X3),k2_yellow_0(X1,X2))=k2_yellow_0(X1,k2_xboole_0(X2,X3))|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1)))).# trainpos
cnf(c_0_154, plain, (v1_xboole_0(X1)|m1_subset_1(k1_tarski(X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_203, plain, (m1_subset_1(esk6_3(X1,X2,X3),k1_zfmisc_1(X3))|r2_hidden(X2,X3)|v1_xboole_0(X3)|X2!=k2_yellow_0(X1,X3)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_73, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_210, plain, (v1_finset_1(esk6_3(X1,X2,X3))|r2_hidden(X2,X3)|v1_xboole_0(X3)|~v1_finset_1(X3)|X2!=k2_yellow_0(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_255, plain, (r2_hidden(X1,k1_tarski(X1)))).# trainpos
cnf(c_0_226, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|X1!=k2_yellow_0(X3,X2)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)|~v2_lattice3(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)|~r2_hidden(X1,esk6_3(X3,X1,X2)))).# trainpos
cnf(c_0_72, negated_conjecture, (r2_hidden(esk2_0,k2_tarski(X1,X2))|v1_xboole_0(u1_struct_0(esk1_0))|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,k2_tarski(X1,X2))!=esk2_0|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_145, plain, (X1=k2_yellow_0(X2,X3)|X1=k2_yellow_0(X2,X4)|X1!=k2_yellow_0(X2,k2_xboole_0(X3,X4))|~v2_waybel_6(X1,X2)|~m1_subset_1(k2_yellow_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(k2_yellow_0(X2,X4),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~r2_yellow_0(X2,X4)|~r2_yellow_0(X2,X3)|~v4_orders_2(X2))).# trainpos
cnf(c_0_165, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~r1_tarski(X1,k1_tarski(X3))|~m1_subset_1(X3,X2))).# trainpos
cnf(c_0_169, plain, (k2_yellow_0(X2,k1_tarski(esk5_3(X2,X1,X3)))=X1|k2_yellow_0(X2,esk6_3(X2,X1,X3))=X1|r2_hidden(X1,X3)|v1_xboole_0(X3)|~r2_yellow_0(X2,k1_tarski(esk5_3(X2,X1,X3)))|~r2_yellow_0(X2,esk6_3(X2,X1,X3))|~v2_waybel_6(X1,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|X1!=k2_yellow_0(X2,X3)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X2)))|~v3_orders_2(X2))).# trainpos
cnf(c_0_174, plain, (r1_tarski(k1_tarski(X1),X2)|v1_xboole_0(X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_178, plain, (k2_yellow_0(X1,k1_tarski(esk5_3(X1,X2,X3)))=X2|k2_yellow_0(X1,esk6_3(X1,X2,X3))=X2|r2_hidden(X2,X3)|v1_xboole_0(X3)|X2!=k2_yellow_0(X1,X3)|~r2_yellow_0(X1,esk6_3(X1,X2,X3))|~v1_finset_1(X3)|~v2_waybel_6(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(esk5_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_82, plain, (k2_yellow_0(X1,k2_tarski(X2,X3))=k12_lattice3(X1,X2,X3)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_231, plain, (m1_subset_1(esk5_3(X1,X2,X3),X3)|r2_hidden(X2,X3)|v1_xboole_0(X3)|X2!=k2_yellow_0(X1,X3)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_91, negated_conjecture, (r2_hidden(esk2_0,k2_tarski(X1,X2))|v2_waybel_6(esk2_0,esk1_0)|k12_lattice3(esk1_0,X1,X2)!=esk2_0|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_95, negated_conjecture, (esk2_0=X1|esk2_0=X2|v2_waybel_6(esk2_0,esk1_0)|k12_lattice3(esk1_0,X2,X1)!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_97, negated_conjecture, (esk7_2(esk1_0,X1)=esk2_0|esk8_2(esk1_0,X1)=esk2_0|v2_waybel_6(esk2_0,esk1_0)|v2_waybel_6(X1,esk1_0)|X1!=esk2_0|~m1_subset_1(esk8_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(esk7_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_101, negated_conjecture, (esk8_2(esk1_0,X1)=esk2_0|esk7_2(esk1_0,X1)=esk2_0|v2_waybel_6(esk2_0,esk1_0)|v2_waybel_6(X1,esk1_0)|X1!=esk2_0|~m1_subset_1(esk7_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_108, negated_conjecture, (esk7_2(esk1_0,X1)=esk2_0|esk8_2(esk1_0,X1)=esk2_0|v2_waybel_6(esk2_0,esk1_0)|v2_waybel_6(X1,esk1_0)|X1!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_117, negated_conjecture, (v2_waybel_6(X1,esk1_0)|esk7_2(esk1_0,X1)=esk2_0|v2_waybel_6(esk2_0,esk1_0)|esk2_0!=X1|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_128, negated_conjecture, (v2_waybel_6(X1,esk1_0)|v2_waybel_6(esk2_0,esk1_0)|esk2_0!=X1|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_140, negated_conjecture, (v2_waybel_6(esk2_0,esk1_0))).# trainpos
cnf(c_0_201, negated_conjecture, (k2_yellow_0(esk1_0,esk3_0)=esk2_0)).# trainpos
cnf(c_0_242, negated_conjecture, (~r2_hidden(esk2_0,esk3_0))).# trainpos
cnf(c_0_175, negated_conjecture, (~v1_xboole_0(esk3_0))).# trainpos
cnf(c_0_191, negated_conjecture, (v1_finset_1(esk3_0))).# trainpos
cnf(c_0_153, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_164, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~r1_tarski(X1,esk3_0))).# trainpos
cnf(c_0_217, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk3_0))).# trainpos
cnf(c_0_173, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_tarski(X1,esk3_0))).# trainpos
cnf(c_0_223, negated_conjecture, (m1_subset_1(esk5_3(X1,X2,esk3_0),u1_struct_0(esk1_0))|r2_hidden(X2,esk3_0)|X2!=k2_yellow_0(X1,esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_200, negated_conjecture, (m1_subset_1(esk6_3(X1,X2,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|r2_hidden(X2,esk3_0)|X2!=k2_yellow_0(X1,esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_179, negated_conjecture, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,esk3_0))).# trainpos
cnf(c_0_183, negated_conjecture, (k2_yellow_0(esk1_0,k1_tarski(esk5_3(esk1_0,X1,X2)))=X1|k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))=X1|r2_hidden(X1,X2)|v1_xboole_0(X2)|X1!=k2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))|~v1_finset_1(X2)|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(esk5_3(esk1_0,X1,X2),esk3_0))).# trainpos
cnf(c_0_189, negated_conjecture, (esk5_3(esk1_0,X1,X2)=X1|k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))=X1|r2_hidden(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(esk5_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|X1!=k2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))|~v1_finset_1(X2)|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(esk5_3(esk1_0,X1,X2),esk3_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_199, negated_conjecture, (k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))=X1|esk5_3(esk1_0,X1,X2)=X1|r2_hidden(X1,X2)|v1_xboole_0(X2)|v1_xboole_0(esk6_3(esk1_0,X1,X2))|X1!=k2_yellow_0(esk1_0,X2)|~v1_finset_1(X2)|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(esk5_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(esk5_3(esk1_0,X1,X2),esk3_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v1_finset_1(esk6_3(esk1_0,X1,X2))|~m1_subset_1(esk6_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_209, negated_conjecture, (k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,esk3_0))=X1|esk5_3(esk1_0,X1,esk3_0)=X1|r2_hidden(X1,esk3_0)|v1_xboole_0(esk6_3(esk1_0,X1,esk3_0))|X1!=esk2_0|~v1_finset_1(esk6_3(esk1_0,X1,esk3_0))|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(esk5_3(esk1_0,X1,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk5_3(esk1_0,X1,esk3_0),esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_216, negated_conjecture, (k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,esk3_0))=X1|esk5_3(esk1_0,X1,esk3_0)=X1|r2_hidden(X1,esk3_0)|v1_xboole_0(esk6_3(esk1_0,X1,esk3_0))|X1!=esk2_0|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(esk5_3(esk1_0,X1,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk5_3(esk1_0,X1,esk3_0),esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_222, negated_conjecture, (r2_hidden(X1,esk6_3(esk1_0,X1,esk3_0))|r2_hidden(X1,esk3_0)|v1_xboole_0(esk6_3(esk1_0,X1,esk3_0))|esk5_3(esk1_0,X1,esk3_0)=X1|X1!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(esk5_3(esk1_0,X1,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk5_3(esk1_0,X1,esk3_0),esk3_0))).# trainpos
cnf(c_0_227, negated_conjecture, (esk5_3(esk1_0,X1,esk3_0)=X1|r2_hidden(X1,esk6_3(esk1_0,X1,esk3_0))|r2_hidden(X1,esk3_0)|v1_xboole_0(esk6_3(esk1_0,X1,esk3_0))|X1!=esk2_0|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(esk5_3(esk1_0,X1,esk3_0),esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_230, negated_conjecture, (r2_hidden(X1,esk3_0)|esk5_3(esk1_0,X1,esk3_0)=X1|v1_xboole_0(esk6_3(esk1_0,X1,esk3_0))|X1!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(esk5_3(esk1_0,X1,esk3_0),esk3_0))).# trainpos
cnf(c_0_233, negated_conjecture, (esk5_3(esk1_0,X1,esk3_0)=X1|r2_hidden(X1,esk3_0)|v1_xboole_0(esk6_3(esk1_0,X1,esk3_0))|X1!=esk2_0|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_237, negated_conjecture, (esk6_3(esk1_0,X1,esk3_0)=k1_xboole_0|esk5_3(esk1_0,X1,esk3_0)=X1|r2_hidden(X1,esk3_0)|X1!=esk2_0|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_241, negated_conjecture, (r2_hidden(X1,esk3_0)|esk6_3(esk1_0,X1,esk3_0)=k1_xboole_0|X1!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v2_waybel_6(X1,esk1_0))).# trainpos
cnf(c_0_246, negated_conjecture, (esk6_3(esk1_0,esk2_0,esk3_0)=k1_xboole_0)).# trainpos
cnf(c_0_253, negated_conjecture, (~r2_hidden(esk2_0,k1_tarski(esk5_3(esk1_0,esk2_0,esk3_0))))).# trainpos
cnf(c_0_249, negated_conjecture, (k2_yellow_0(esk1_0,k1_tarski(esk5_3(esk1_0,esk2_0,esk3_0)))=esk2_0)).# trainpos
cnf(c_0_251, negated_conjecture, (esk5_3(esk1_0,esk2_0,esk3_0)=esk2_0|~m1_subset_1(esk5_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_254, negated_conjecture, (esk5_3(esk1_0,esk2_0,esk3_0)=esk2_0)).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_121, plain, (o_0_0_xboole_0=k1_xboole_0)).#trainneg
cnf(i_0_138, plain, (k5_numbers=k4_ordinal1)).#trainneg
cnf(i_0_122, plain, (v1_xboole_0(esk14_0))).#trainneg
cnf(i_0_161, plain, (v1_xboole_0(esk18_0))).#trainneg
cnf(i_0_134, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_171, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_160, plain, (v1_xxreal_0(esk18_0))).#trainneg
cnf(i_0_133, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_172, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_145, plain, (~v1_finset_1(k4_ordinal1))).#trainneg
cnf(i_0_140, plain, (v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_123, plain, (v1_xboole_0(esk15_1(X1)))).#trainneg
cnf(i_0_131, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_132, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_93, plain, (k2_xboole_0(X1,X1)=X1)).#trainneg
cnf(i_0_107, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_135, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_80, plain, (m1_subset_1(esk12_1(X1),X1))).#trainneg
cnf(i_0_147, plain, (~v1_subset_1(esk16_1(X1),X1))).#trainneg
cnf(i_0_153, plain, (v1_finset_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_151, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_67, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_154, plain, (v2_orders_2(X1)|~v3_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_137, plain, (v1_yellow_0(X1)|~l1_orders_2(X1)|~v3_yellow_0(X1))).#trainneg
cnf(i_0_162, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_136, plain, (v2_yellow_0(X1)|~l1_orders_2(X1)|~v3_yellow_0(X1))).#trainneg
cnf(i_0_124, plain, (m1_subset_1(esk15_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_148, plain, (m1_subset_1(esk16_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_76, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1))).#trainneg
cnf(i_0_165, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_142, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_xboole_0(X2,X1)))).#trainneg
cnf(i_0_143, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_xboole_0(X1,X2)))).#trainneg
cnf(i_0_141, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_96, plain, (~l1_orders_2(X1)|~v2_struct_0(X1)|~v1_lattice3(X1))).#trainneg
cnf(i_0_46, plain, (X1=X3|X2!=k1_tarski(X3)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_128, plain, (v2_struct_0(X1)|v2_yellow_0(X1)|~l1_orders_2(X1)|~v9_waybel_1(X1))).#trainneg
cnf(i_0_44, plain, (X2=k1_tarski(X1)|esk9_2(X1,X2)!=X1|~r2_hidden(esk9_2(X1,X2),X2))).#trainneg
cnf(i_0_104, plain, (v3_orders_2(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v9_waybel_1(X1))).#trainneg
cnf(i_0_98, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_164, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_101, plain, (v2_struct_0(X1)|v1_lattice3(X1)|~l1_orders_2(X1)|~v9_waybel_1(X1))).#trainneg
cnf(i_0_113, plain, (v2_struct_0(X1)|v1_lattice3(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_116, plain, (v3_orders_2(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_103, plain, (v4_orders_2(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v9_waybel_1(X1))).#trainneg
cnf(i_0_139, plain, (v1_finset_1(X2)|~v5_finset_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_52, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2))).#trainneg
cnf(i_0_53, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4))).#trainneg
cnf(i_0_71, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_63, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_95, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_152, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_146, plain, (v1_finset_1(k2_xboole_0(X1,X2))|~v1_finset_1(X2)|~v1_finset_1(X1))).#trainneg
cnf(i_0_115, plain, (v4_orders_2(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_16, plain, (r2_hidden(X1,X3)|X3!=k2_xboole_0(X4,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_126, plain, (v2_struct_0(X1)|v9_waybel_1(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_50, plain, (X3=k2_tarski(X1,X2)|esk10_3(X1,X2,X3)!=X2|~r2_hidden(esk10_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_102, plain, (v5_orders_2(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v9_waybel_1(X1))).#trainneg
cnf(i_0_114, plain, (v5_orders_2(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_158, plain, (v2_struct_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v1_xboole_0(esk17_1(X1)))).#trainneg
cnf(i_0_100, plain, (v2_lattice3(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v9_waybel_1(X1))).#trainneg
cnf(i_0_111, plain, (v2_struct_0(X1)|v3_yellow_0(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_149, plain, (v2_struct_0(X1)|v2_waybel_1(X1)|~l1_orders_2(X1)|~v9_waybel_1(X1))).#trainneg
cnf(i_0_112, plain, (v2_lattice3(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_110, plain, (v2_struct_0(X1)|v2_waybel_1(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_109, plain, (v2_struct_0(X1)|v10_waybel_1(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_157, plain, (v2_struct_0(X1)|v1_finset_1(esk17_1(X1))|~v3_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_43, plain, (X2=k1_tarski(X1)|esk9_2(X1,X2)=X1|r2_hidden(esk9_2(X1,X2),X2))).#trainneg
cnf(i_0_156, plain, (v2_struct_0(X1)|v1_waybel_0(esk17_1(X1),X1)|~v3_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_155, plain, (v2_struct_0(X1)|v2_waybel_0(esk17_1(X1),X1)|~v3_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_167, plain, (v2_struct_0(X1)|v24_waybel_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v3_waybel_3(X1))).#trainneg
cnf(i_0_166, plain, (v2_struct_0(X1)|v2_waybel_3(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v3_waybel_3(X1))).#trainneg
cnf(i_0_60, plain, (v2_lattice3(X1)|v2_struct_0(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~v1_xboole_0(esk11_1(X1)))).#trainneg
cnf(i_0_18, plain, (r2_hidden(X1,X4)|r2_hidden(X1,X3)|X2!=k2_xboole_0(X3,X4)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_59, plain, (v2_lattice3(X1)|v2_struct_0(X1)|v1_finset_1(esk11_1(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_51, plain, (X3=k2_tarski(X1,X2)|esk10_3(X1,X2,X3)!=X1|~r2_hidden(esk10_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_91, plain, (v2_struct_0(X1)|v2_waybel_6(esk13_1(X1),X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v2_yellow_0(X1))).#trainneg
cnf(i_0_119, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_57, plain, (v2_lattice3(X1)|v2_struct_0(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~r2_yellow_0(X1,esk11_1(X1)))).#trainneg
cnf(i_0_159, plain, (v2_struct_0(X1)|m1_subset_1(esk17_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_118, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_92, plain, (v2_struct_0(X1)|m1_subset_1(esk13_1(X1),u1_struct_0(X1))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v2_yellow_0(X1))).#trainneg
cnf(i_0_58, plain, (v2_lattice3(X1)|v2_struct_0(X1)|m1_subset_1(esk11_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_170, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m2_subset_1(esk19_2(X1,X2),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_14, plain, (X3=k2_xboole_0(X1,X2)|~r2_hidden(esk4_3(X1,X2,X3),X3)|~r2_hidden(esk4_3(X1,X2,X3),X2))).#trainneg
cnf(i_0_15, plain, (X3=k2_xboole_0(X1,X2)|~r2_hidden(esk4_3(X1,X2,X3),X3)|~r2_hidden(esk4_3(X1,X2,X3),X1))).#trainneg
cnf(i_0_88, plain, (k7_domain_1(X1,X2,X3)=k7_domain_1(X1,X3,X2)|v1_xboole_0(X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_108, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_77, plain, (m1_subset_1(k11_lattice3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_49, plain, (X3=k2_tarski(X1,X2)|esk10_3(X1,X2,X3)=X2|esk10_3(X1,X2,X3)=X1|r2_hidden(esk10_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_48, plain, (r2_yellow_0(X1,k2_xboole_0(X2,X3))|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2))).#trainneg
cnf(i_0_70, plain, (m1_subset_1(k12_lattice3(X1,X2,X3),u1_struct_0(X1))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_26, plain, (v1_xboole_0(X3)|r2_hidden(X2,X3)|X2!=k2_yellow_0(X1,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_xboole_0(k2_xboole_0(esk6_3(X1,X2,X3),k1_tarski(esk5_3(X1,X2,X3)))))).#trainneg
cnf(i_0_82, plain, (v2_struct_0(X2)|v1_waybel_0(X1,X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~v6_orders_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_81, plain, (v2_struct_0(X2)|v2_waybel_0(X1,X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~v6_orders_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_56, plain, (k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2|v2_struct_0(X1)|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_13, plain, (X3=k2_xboole_0(X1,X2)|r2_hidden(esk4_3(X1,X2,X3),X3)|r2_hidden(esk4_3(X1,X2,X3),X2)|r2_hidden(esk4_3(X1,X2,X3),X1))).#trainneg
cnf(i_0_201, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_175, plain, (esk18_0=k1_xboole_0)).#trainneg
cnf(i_0_176, plain, (esk14_0=k1_xboole_0)).#trainneg
cnf(i_0_298, negated_conjecture, (r2_hidden(esk2_0,esk17_1(esk1_0))|v1_xboole_0(esk17_1(esk1_0))|v2_waybel_6(esk2_0,esk1_0)|v2_struct_0(esk1_0)|k2_yellow_0(esk1_0,esk17_1(esk1_0))!=esk2_0|~v1_finset_1(esk17_1(esk1_0)))).#trainneg
cnf(i_0_53, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X1,X3))).#trainneg
cnf(i_0_247, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_242, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk16_1(X2)))).#trainneg
cnf(i_0_52, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X3,X1))).#trainneg
cnf(i_0_177, plain, (esk15_1(X1)=k1_xboole_0)).#trainneg
cnf(i_0_216, plain, (v1_finset_1(np__1)|~v5_finset_1(k4_ordinal1))).#trainneg
cnf(i_0_229, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_258, plain, (v2_yellow_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_207, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_276, negated_conjecture, (r2_hidden(esk2_0,k6_domain_1(u1_struct_0(esk1_0),X1))|v1_xboole_0(k6_domain_1(u1_struct_0(esk1_0),X1))|v2_waybel_6(esk2_0,esk1_0)|v1_xboole_0(u1_struct_0(esk1_0))|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1))!=esk2_0|~v1_finset_1(k6_domain_1(u1_struct_0(esk1_0),X1))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_385, plain, (r1_tarski(X1,X2)|r2_hidden(X4,X2)|v1_xboole_0(X2)|~r1_tarski(X1,esk6_3(X3,X4,X2))|X4!=k2_yellow_0(X3,X2)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_orders_2(X3)|~v2_lattice3(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3))).#trainneg
cnf(i_0_283, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k2_xboole_0(X2,X3)))).#trainneg
cnf(i_0_383, plain, (esk5_3(X1,X2,X3)=X4|r2_hidden(X2,X3)|v1_xboole_0(X3)|X3!=k1_tarski(X4)|X2!=k2_yellow_0(X1,X3)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_246, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_198, negated_conjecture, (r2_hidden(esk2_0,esk16_1(u1_struct_0(esk1_0)))|v1_xboole_0(esk16_1(u1_struct_0(esk1_0)))|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,esk16_1(u1_struct_0(esk1_0)))!=esk2_0|~v1_finset_1(esk16_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_233, plain, (~v1_finset_1(k1_numbers))).#trainneg
cnf(i_0_187, plain, (v1_xxreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_217, plain, (v1_finset_1(np__1)|~v5_finset_1(k1_numbers))).#trainneg
cnf(i_0_215, plain, (~v5_finset_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_237, plain, (v1_relat_1(k4_ordinal1)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_200, negated_conjecture, (r2_hidden(esk2_0,esk12_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|v1_xboole_0(esk12_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,esk12_1(k1_zfmisc_1(u1_struct_0(esk1_0))))!=esk2_0|~v1_finset_1(esk12_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_297, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r2_hidden(X1,esk17_1(X2))|~l1_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_225, plain, (X1=X2|v1_xboole_0(X3)|X3!=k1_tarski(X2)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_226, plain, (v1_xboole_0(esk16_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_312, plain, (k7_domain_1(X1,X3,X2)=k2_tarski(X2,X3)|v1_xboole_0(X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_244, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_230, plain, (v1_finset_1(esk16_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_214, negated_conjecture, (v1_finset_1(esk2_0)|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_203, plain, (~r2_hidden(X1,esk16_1(X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_211, plain, (v1_finset_1(esk16_1(X1))|~v5_finset_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_245, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k4_ordinal1))).#trainneg
cnf(i_0_205, plain, (~r2_hidden(X1,esk12_1(k1_zfmisc_1(X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_213, plain, (v1_finset_1(esk12_1(X1))|~v5_finset_1(X1))).#trainneg
cnf(i_0_220, plain, (v1_finset_1(k2_yellow_0(X1,X2))|~v5_finset_1(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_221, plain, (v1_finset_1(k1_yellow_0(X1,X2))|~v5_finset_1(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_376, plain, (k12_lattice3(X1,k12_lattice3(X1,X2,X3),k2_yellow_0(X1,X4))=k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X2,X3),X4))|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_387, plain, (k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(esk6_3(X1,X2,X3),k1_tarski(esk5_3(X1,X2,X3))),X4))=k12_lattice3(X1,X2,k2_yellow_0(X1,X4))|r2_hidden(X2,X3)|v1_xboole_0(X3)|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X2,X3),k1_tarski(esk5_3(X1,X2,X3))))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|X2!=k2_yellow_0(X1,X3)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_377, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k12_lattice3(X1,X3,X4))=k2_yellow_0(X1,k2_xboole_0(X2,k7_domain_1(u1_struct_0(X1),X3,X4)))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_389, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|v1_xboole_0(k2_xboole_0(esk6_3(X3,X1,X2),k1_tarski(esk5_3(X3,X1,X2))))|X1!=k2_yellow_0(X3,X2)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)|~v2_lattice3(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)|~m1_subset_1(X1,k2_xboole_0(esk6_3(X3,X1,X2),k1_tarski(esk5_3(X3,X1,X2)))))).#trainneg
cnf(i_0_379, plain, (k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X2))=k12_lattice3(X1,X2,X3)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_281, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,k6_domain_1(X2,X3))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_228, plain, (v1_xboole_0(esk12_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_234, plain, (v1_relat_1(esk16_1(X1))|~v1_relat_1(X1))).#trainneg
cnf(i_0_273, plain, (esk9_2(X1,X2)=X3|esk9_2(X1,X2)=X1|X2=k1_tarski(X1)|X2!=k1_tarski(X3))).#trainneg
cnf(i_0_222, plain, (X1=k1_tarski(X2)|v1_xboole_0(X1)|esk9_2(X2,X1)!=X2|~m1_subset_1(esk9_2(X2,X1),X1))).#trainneg
cnf(i_0_232, plain, (v1_finset_1(esk12_1(k1_zfmisc_1(X1)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_236, plain, (v1_relat_1(esk12_1(k1_zfmisc_1(X1)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_326, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_255, plain, (X1=k2_tarski(X2,X3)|v1_xboole_0(X1)|esk10_3(X2,X3,X1)!=X3|~m1_subset_1(esk10_3(X2,X3,X1),X1))).#trainneg
cnf(i_0_346, plain, (k2_yellow_0(X1,k2_xboole_0(X2,k6_domain_1(u1_struct_0(X1),X3)))=k12_lattice3(X1,k2_yellow_0(X1,X2),X3)|v2_struct_0(X1)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X3))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_160, plain, (v1_xxreal_0(k1_xboole_0))).#trainneg
cnf(i_0_124, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_272, plain, (m1_subset_1(esk9_2(X1,X2),X2)|esk9_2(X1,X2)=X1|X2=k1_tarski(X1))).#trainneg
cnf(i_0_402, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k4_ordinal1))).#trainneg
cnf(i_0_391, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))|v1_xboole_0(u1_struct_0(esk1_0))|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,u1_struct_0(esk1_0))!=esk2_0|~v1_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_390, plain, (v1_finset_1(X1)|~v5_finset_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_275, plain, (v1_finset_1(k6_domain_1(X1,X2))|v1_xboole_0(X1)|~v5_finset_1(k1_zfmisc_1(X1))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_279, plain, (v1_finset_1(k6_domain_1(X1,X2))|v1_xboole_0(X1)|~v1_finset_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_407, negated_conjecture, (v2_struct_0(esk1_0)|r2_hidden(esk2_0,esk17_1(esk1_0))|v1_xboole_0(esk17_1(esk1_0))|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,esk17_1(esk1_0))!=esk2_0)).#trainneg
cnf(i_0_404, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~r1_tarski(X1,k6_domain_1(X2,X3))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_409, plain, (r2_hidden(X1,k2_tarski(X1,X2)))).#trainneg
cnf(i_0_403, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~r1_tarski(X1,k7_domain_1(X2,X3,X4))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_463, plain, (r1_tarski(esk6_3(X1,X2,esk6_3(X3,X4,X5)),X5)|r2_hidden(X4,X5)|v1_xboole_0(X5)|r2_hidden(X2,esk6_3(X3,X4,X5))|v1_xboole_0(esk6_3(X3,X4,X5))|X4!=k2_yellow_0(X3,X5)|~v1_finset_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_orders_2(X3)|~v2_lattice3(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)|X2!=k2_yellow_0(X1,esk6_3(X3,X4,X5))|~v1_finset_1(esk6_3(X3,X4,X5))|~m1_subset_1(esk6_3(X3,X4,X5),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_553, negated_conjecture, (r2_hidden(esk2_0,esk16_1(u1_struct_0(esk1_0)))|v1_xboole_0(esk16_1(u1_struct_0(esk1_0)))|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,esk16_1(u1_struct_0(esk1_0)))!=esk2_0|~v1_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_524, plain, (np__1=X1|k1_tarski(X1)!=k4_ordinal1)).#trainneg
cnf(i_0_280, plain, (v1_relat_1(k6_domain_1(X1,X2))|v1_xboole_0(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_539, plain, (esk16_1(X1)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_561, negated_conjecture, (r2_hidden(esk2_0,esk16_1(u1_struct_0(esk1_0)))|v1_xboole_0(esk16_1(u1_struct_0(esk1_0)))|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,esk16_1(u1_struct_0(esk1_0)))!=esk2_0|~v5_finset_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_675, negated_conjecture, (r2_hidden(esk2_0,esk12_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|v1_xboole_0(esk12_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,esk12_1(k1_zfmisc_1(u1_struct_0(esk1_0))))!=esk2_0|~v1_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_641, plain, (v1_relat_1(k1_tarski(X1))|v1_xboole_0(X2)|~v1_relat_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_290, plain, (X1=k2_tarski(X2,X3)|v1_xboole_0(X1)|esk10_3(X2,X3,X1)!=X2|~m1_subset_1(esk10_3(X2,X3,X1),X1))).#trainneg
cnf(i_0_479, plain, (esk5_3(X1,X2,k1_tarski(X3))=X3|r2_hidden(X2,k1_tarski(X3))|X2!=k2_yellow_0(X1,k1_tarski(X3))|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_696, plain, (k2_yellow_0(X1,k2_xboole_0(k6_domain_1(u1_struct_0(X1),X3),X2))=k12_lattice3(X1,k2_yellow_0(X1,X2),X3)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X3))|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_477, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|v1_xboole_0(k2_xboole_0(X3,X2))|~m1_subset_1(X1,k2_xboole_0(X3,X2)))).#trainneg
cnf(i_0_485, plain, (~v1_xboole_0(k2_xboole_0(X2,X3))|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_488, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|X1!=k2_yellow_0(X3,X2)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,u1_struct_0(X3))|~l1_orders_2(X3)|~v2_lattice3(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3)|~r2_hidden(X1,k1_tarski(esk5_3(X3,X1,X2))))).#trainneg
cnf(i_0_576, negated_conjecture, (r2_hidden(esk2_0,esk12_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|v1_xboole_0(esk12_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,esk12_1(k1_zfmisc_1(u1_struct_0(esk1_0))))!=esk2_0|~v5_finset_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_525, plain, (np__1=X1|k1_tarski(X1)!=k1_numbers)).#trainneg
cnf(i_0_296, plain, (v1_relat_1(esk17_1(X1))|v2_struct_0(X1)|~v1_relat_1(u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_522, negated_conjecture, (esk2_0=X1|v1_xboole_0(u1_struct_0(esk1_0))|k1_tarski(X1)!=u1_struct_0(esk1_0))).#trainneg
cnf(i_0_452, plain, (~v1_xboole_0(k2_xboole_0(X2,X3))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_568, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_397, plain, (r1_tarski(X1,u1_struct_0(X2))|v2_struct_0(X2)|~r1_tarski(X1,esk17_1(X2))|~l1_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_637, negated_conjecture, (r2_hidden(esk2_0,k1_tarski(X1))|v2_waybel_6(esk2_0,esk1_0)|v1_xboole_0(u1_struct_0(esk1_0))|k2_yellow_0(esk1_0,k1_tarski(X1))!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_606, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(X3,k1_tarski(X2)))|v2_struct_0(X1)|~r2_yellow_0(X1,k1_tarski(X2))|~r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_604, plain, (k12_lattice3(X1,k12_lattice3(X1,X2,X3),k2_yellow_0(X1,X4))=k2_yellow_0(X1,k2_xboole_0(X4,k7_domain_1(u1_struct_0(X1),X2,X3)))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,X4)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_642, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,k1_tarski(X3))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_431, plain, (r2_hidden(X1,k2_tarski(X2,X1)))).#trainneg
cnf(i_0_523, plain, (k4_ordinal1=X1|k1_tarski(X1)!=k1_zfmisc_1(k1_numbers))).#trainneg
cnf(i_0_306, plain, (v1_relat_1(esk11_1(X1))|v2_struct_0(X1)|v2_lattice3(X1)|~v1_relat_1(u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_398, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,esk16_1(X2)))).#trainneg
cnf(i_0_673, plain, (esk12_1(k1_zfmisc_1(X1))=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_453, plain, (m1_subset_1(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_486, plain, (m1_subset_1(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_307, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|v2_lattice3(X2)|~r2_hidden(X1,esk11_1(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_309, plain, (v1_xboole_0(X1)|m1_subset_1(esk19_2(X2,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_855, negated_conjecture, (r2_hidden(esk2_0,k1_tarski(X1))|v2_waybel_6(esk2_0,esk1_0)|X1!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_749, plain, (r1_tarski(k7_domain_1(X1,X2,X3),X1)|v1_xboole_0(X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_931, plain, (r1_tarski(esk16_1(X1),X1))).#trainneg
cnf(i_0_737, plain, (r1_tarski(k6_domain_1(X1,X2),X1)|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_779, plain, (v1_relat_1(k1_tarski(np__1))|~v1_relat_1(k4_ordinal1))).#trainneg
cnf(i_0_699, plain, (k1_xboole_0=X1|k1_zfmisc_1(X2)!=k1_tarski(X1))).#trainneg
cnf(i_0_310, plain, (X1=k2_xboole_0(X2,X3)|v1_xboole_0(X1)|~r2_hidden(esk4_3(X2,X3,X1),X3)|~m1_subset_1(esk4_3(X2,X3,X1),X1))).#trainneg
cnf(i_0_704, plain, (v1_finset_1(k1_xboole_0)|~v1_finset_1(X1))).#trainneg
cnf(i_0_780, plain, (v1_relat_1(k1_tarski(np__1))|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_777, negated_conjecture, (v1_relat_1(k1_tarski(esk2_0))|~v1_relat_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_311, plain, (X1=k2_xboole_0(X2,X3)|v1_xboole_0(X1)|~r2_hidden(esk4_3(X2,X3,X1),X2)|~m1_subset_1(esk4_3(X2,X3,X1),X1))).#trainneg
cnf(i_0_314, plain, (v1_finset_1(k7_domain_1(X1,X2,X3))|v1_xboole_0(X1)|~v5_finset_1(k1_zfmisc_1(X1))|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_898, plain, (k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X4,X5),k7_domain_1(u1_struct_0(X1),X2,X3)))=k12_lattice3(X1,k12_lattice3(X1,X2,X3),k12_lattice3(X1,X4,X5))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X4,X5))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_878, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(k1_tarski(X2),X3))|~r2_yellow_0(X1,k1_tarski(X2))|~r2_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_902, plain, (k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(esk6_3(X1,X4,X5),k1_tarski(esk5_3(X1,X4,X5))),k7_domain_1(u1_struct_0(X1),X2,X3)))=k12_lattice3(X1,k12_lattice3(X1,X2,X3),X4)|r2_hidden(X4,X5)|v1_xboole_0(X5)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X4,X5),k1_tarski(esk5_3(X1,X4,X5))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|X4!=k2_yellow_0(X1,X5)|~v1_finset_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_732, plain, (m1_subset_1(X1,k1_tarski(X1)))).#trainneg
cnf(i_0_896, plain, (X1=k2_yellow_0(X2,X3)|X1=k12_lattice3(X2,X4,X5)|X1!=k2_yellow_0(X2,k2_xboole_0(X3,k7_domain_1(u1_struct_0(X2),X4,X5)))|~v2_waybel_6(X1,X2)|~m1_subset_1(k2_yellow_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(k12_lattice3(X2,X4,X5),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~r2_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X4,X5))|~r2_yellow_0(X2,X3)|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_706, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_778, plain, (v1_relat_1(k1_tarski(k4_ordinal1))|~v1_relat_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_532, plain, (X1=X2|k1_zfmisc_1(X1)!=k1_tarski(X2))).#trainneg
cnf(i_0_318, plain, (v1_finset_1(k7_domain_1(X1,X2,X3))|v1_xboole_0(X1)|~v1_finset_1(X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_705, plain, (v1_relat_1(k1_xboole_0)|~v1_relat_1(X1))).#trainneg
cnf(i_0_723, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))|~v1_xboole_0(esk3_0)|k2_yellow_0(esk1_0,u1_struct_0(esk1_0))!=esk2_0|~v1_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_725, negated_conjecture, (r2_hidden(esk2_0,u1_struct_0(esk1_0))|~r2_hidden(esk2_0,esk3_0)|k2_yellow_0(esk1_0,u1_struct_0(esk1_0))!=esk2_0|~v1_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_319, plain, (v1_relat_1(k7_domain_1(X1,X2,X3))|v1_xboole_0(X1)|~v1_relat_1(X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_320, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,k7_domain_1(X2,X3,X4))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_849, plain, (v2_struct_0(X1)|r1_tarski(esk17_1(X1),u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_900, plain, (k2_yellow_0(X1,k2_xboole_0(k6_domain_1(u1_struct_0(X1),X4),k7_domain_1(u1_struct_0(X1),X2,X3)))=k12_lattice3(X1,k12_lattice3(X1,X2,X3),X4)|v2_struct_0(X1)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_947, plain, (v1_finset_1(X1)|~v5_finset_1(k2_xboole_0(X2,X3))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_722, negated_conjecture, (v1_finset_1(esk3_0)|r2_hidden(esk2_0,u1_struct_0(esk1_0))|k2_yellow_0(esk1_0,u1_struct_0(esk1_0))!=esk2_0|~v1_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_733, plain, (X1=X2|k1_tarski(X1)!=k1_tarski(X2))).#trainneg
cnf(i_0_321, plain, (v1_xboole_0(X1)|m1_subset_1(k2_tarski(X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_724, negated_conjecture, (k2_yellow_0(esk1_0,esk3_0)=esk2_0|r2_hidden(esk2_0,u1_struct_0(esk1_0))|k2_yellow_0(esk1_0,u1_struct_0(esk1_0))!=esk2_0|~v1_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_726, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|r2_hidden(esk2_0,u1_struct_0(esk1_0))|k2_yellow_0(esk1_0,u1_struct_0(esk1_0))!=esk2_0|~v1_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_700, plain, (v1_finset_1(k1_xboole_0)|~v5_finset_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_324, plain, (v1_xboole_0(X1)|m1_subset_1(esk19_2(X2,X1),X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_327, plain, (v1_finset_1(k11_lattice3(X1,X2,X3))|~v5_finset_1(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_414, negated_conjecture, (r2_hidden(esk2_0,k7_domain_1(u1_struct_0(esk1_0),X2,X1))|v1_xboole_0(k7_domain_1(u1_struct_0(esk1_0),X2,X1))|v1_xboole_0(u1_struct_0(esk1_0))|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,X1))!=esk2_0|~v1_finset_1(k7_domain_1(u1_struct_0(esk1_0),X2,X1))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_991, negated_conjecture, (r2_hidden(esk2_0,X1)|v1_xboole_0(X1)|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(X2,X1)!=esk2_0|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(esk2_0,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|esk5_3(X2,esk2_0,X1)!=esk2_0|~m1_subset_1(esk5_3(X2,esk2_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1220, plain, (v1_finset_1(X1)|~v5_finset_1(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_1133, plain, (v1_finset_1(X1)|~v5_finset_1(k1_tarski(X1)))).#trainneg
cnf(i_0_1039, plain, (v1_finset_1(k1_xboole_0))).#trainneg
cnf(i_0_329, plain, (m1_subset_1(esk10_3(X1,X2,X3),X3)|esk10_3(X1,X2,X3)=X1|esk10_3(X1,X2,X3)=X2|X3=k2_tarski(X1,X2))).#trainneg
cnf(i_0_330, plain, (esk10_3(X1,X2,X3)=X4|esk10_3(X1,X2,X3)=X1|esk10_3(X1,X2,X3)=X2|X3=k2_tarski(X1,X2)|X3!=k1_tarski(X4))).#trainneg
cnf(i_0_333, plain, (v1_finset_1(esk7_2(X1,X2))|v2_struct_0(X1)|v2_waybel_6(X2,X1)|~v5_finset_1(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1073, plain, (k12_lattice3(X1,k12_lattice3(X1,X2,X3),k12_lattice3(X1,X2,X3))=k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_989, negated_conjecture, (m1_subset_1(esk2_0,k1_tarski(X1))|v2_waybel_6(esk2_0,esk1_0)|X1!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1221, plain, (v1_finset_1(X1)|~v5_finset_1(k2_xboole_0(X3,X2))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_516, plain, (esk16_1(X1)=X2|k1_zfmisc_1(X1)!=k1_tarski(X2))).#trainneg
cnf(i_0_338, plain, (v1_finset_1(k12_lattice3(X1,X2,X3))|~v5_finset_1(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1))).#trainneg
cnf(i_0_1015, plain, (r1_tarski(k1_xboole_0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_339, plain, (v1_finset_1(esk8_2(X1,X2))|v2_struct_0(X1)|v2_waybel_6(X2,X1)|~v5_finset_1(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_343, plain, (k1_yellow_0(X1,k1_tarski(X2))=X2|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1242, negated_conjecture, (r2_hidden(esk2_0,esk19_2(k1_zfmisc_1(u1_struct_0(esk1_0)),X1))|v1_xboole_0(esk19_2(k1_zfmisc_1(u1_struct_0(esk1_0)),X1))|v2_waybel_6(esk2_0,esk1_0)|v1_xboole_0(X1)|k2_yellow_0(esk1_0,esk19_2(k1_zfmisc_1(u1_struct_0(esk1_0)),X1))!=esk2_0|~v1_finset_1(esk19_2(k1_zfmisc_1(u1_struct_0(esk1_0)),X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_600, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k12_lattice3(X1,X3,X4))=k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X3,X4),X2))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_990, negated_conjecture, (esk2_0=X1|v2_waybel_6(esk2_0,esk1_0)|k1_tarski(X2)!=k1_tarski(X1)|X2!=esk2_0|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_465, plain, (r1_tarski(X1,X2)|r2_hidden(X3,X2)|v1_xboole_0(X2)|X3!=k2_yellow_0(X4,X2)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X4)))|~m1_subset_1(X3,u1_struct_0(X4))|~l1_orders_2(X4)|~v2_lattice3(X4)|~v5_orders_2(X4)|~v4_orders_2(X4)|~v3_orders_2(X4)|~m1_subset_1(X1,k1_zfmisc_1(esk6_3(X4,X3,X2))))).#trainneg
cnf(i_0_707, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_xboole_0))).#trainneg
cnf(i_0_347, plain, (k2_yellow_0(X1,k2_xboole_0(k6_domain_1(u1_struct_0(X1),X2),X3))=k12_lattice3(X1,X2,k2_yellow_0(X1,X3))|v2_struct_0(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_880, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,k1_tarski(X2)))=k2_yellow_0(X1,k1_tarski(X2))|~r2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_602, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),X3)=k2_yellow_0(X1,k2_xboole_0(k1_tarski(X3),X2))|v2_struct_0(X1)|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,k1_tarski(X3))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_503, plain, (X1=X2|X1=X3|~m1_subset_1(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_864, plain, (k12_lattice3(X1,esk5_3(X1,X2,X3),k2_yellow_0(X1,esk6_3(X1,X2,X3)))=X2|r2_hidden(X2,X3)|v1_xboole_0(X3)|X2!=k2_yellow_0(X1,X3)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k1_tarski(esk5_3(X1,X2,X3)))|~r2_yellow_0(X1,esk6_3(X1,X2,X3))|~m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(X1)))).#trainneg
cnf(i_0_1103, plain, (k2_yellow_0(X1,k2_xboole_0(k1_tarski(X2),k7_domain_1(u1_struct_0(X1),X3,X4)))=k12_lattice3(X1,X2,k12_lattice3(X1,X3,X4))|~r2_yellow_0(X1,k1_tarski(X2))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1330, negated_conjecture, (v2_waybel_6(esk2_0,esk1_0)|m1_subset_1(esk2_0,k1_tarski(k12_lattice3(esk1_0,X1,X2)))|k12_lattice3(esk1_0,X1,X2)!=esk2_0|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1322, negated_conjecture, (v2_waybel_6(esk2_0,esk1_0)|m1_subset_1(esk2_0,k1_tarski(esk13_1(esk1_0)))|esk13_1(esk1_0)!=esk2_0|~v2_yellow_0(esk1_0))).#trainneg
cnf(i_0_741, plain, (X1=X2|k2_tarski(X1,X3)!=k1_tarski(X2))).#trainneg
cnf(i_0_1417, plain, (r1_tarski(k1_xboole_0,X1))).#trainneg
cnf(i_0_1323, negated_conjecture, (v2_waybel_6(esk2_0,esk1_0)|m1_subset_1(esk2_0,k1_tarski(esk12_1(u1_struct_0(esk1_0))))|esk12_1(u1_struct_0(esk1_0))!=esk2_0)).#trainneg
cnf(i_0_787, plain, (v1_relat_1(k1_tarski(k1_xboole_0))|~v1_relat_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_354, plain, (X4=X5|r2_hidden(esk4_3(X5,X5,X4),X5)|r2_hidden(esk4_3(X5,X5,X4),X4))).#trainneg
cnf(i_0_1331, negated_conjecture, (v2_waybel_6(esk2_0,esk1_0)|m1_subset_1(esk2_0,k1_tarski(k11_lattice3(esk1_0,X1,X2)))|k11_lattice3(esk1_0,X1,X2)!=esk2_0|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1299, plain, (X1=k12_lattice3(X2,X3,X4)|X1!=k2_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))|~v2_waybel_6(X1,X2)|~m1_subset_1(k12_lattice3(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~r2_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_1332, negated_conjecture, (v2_waybel_6(esk2_0,esk1_0)|m1_subset_1(esk2_0,k1_tarski(k2_yellow_0(esk1_0,X1)))|k2_yellow_0(esk1_0,X1)!=esk2_0)).#trainneg
cnf(i_0_1400, plain, (r1_tarski(esk16_1(esk6_3(X1,X2,X3)),X3)|r2_hidden(X2,X3)|v1_xboole_0(X3)|X2!=k2_yellow_0(X1,X3)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1326, negated_conjecture, (v2_waybel_6(esk2_0,esk1_0)|m1_subset_1(esk2_0,k1_tarski(esk8_2(esk1_0,X1)))|v2_waybel_6(X1,esk1_0)|esk8_2(esk1_0,X1)!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1371, plain, (X1=k12_lattice3(X2,X3,X4)|X1=k2_yellow_0(X2,X5)|X1!=k2_yellow_0(X2,k2_xboole_0(k7_domain_1(u1_struct_0(X2),X3,X4),X5))|~v2_waybel_6(X1,X2)|~m1_subset_1(k12_lattice3(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(k2_yellow_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~r2_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))|~r2_yellow_0(X2,X5)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_1435, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4)))=k12_lattice3(X1,k12_lattice3(X1,X3,X4),X2)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1300, plain, (k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),esk7_2(X1,X2),esk8_2(X1,X2)))=k12_lattice3(X1,X2,X2)|v2_waybel_6(X2,X1)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),esk7_2(X1,X2),esk8_2(X1,X2)))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk7_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1219, plain, (v1_finset_1(X1)|~v5_finset_1(X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_1271, plain, (v1_finset_1(X1)|~v5_finset_1(k2_tarski(X1,X2)))).#trainneg
cnf(i_0_845, plain, (v1_finset_1(X1)|~v5_finset_1(k1_numbers)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_925, plain, (X1=X2|k2_tarski(X3,X1)!=k1_tarski(X2))).#trainneg
cnf(i_0_355, plain, (k2_xboole_0(X6,X7)=X6|r2_hidden(esk4_3(X6,X7,X6),X7)|r2_hidden(esk4_3(X6,X7,X6),X6))).#trainneg
cnf(i_0_702, plain, (~r2_hidden(X1,k1_xboole_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_760, plain, (~v1_subset_1(k1_xboole_0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_357, plain, (k2_xboole_0(X11,X10)=X10|r2_hidden(esk4_3(X11,X10,X10),X11)|r2_hidden(esk4_3(X11,X10,X10),X10))).#trainneg
cnf(i_0_1492, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),X3)=k2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(X3)))|~r2_yellow_0(X1,k1_tarski(X3))|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1272, plain, (v1_finset_1(X1)|~v5_finset_1(k2_tarski(X2,X1)))).#trainneg
cnf(i_0_1566, negated_conjecture, (v1_finset_1(esk2_0)|v2_waybel_6(esk2_0,esk1_0)|~v5_finset_1(k1_tarski(esk13_1(esk1_0)))|esk13_1(esk1_0)!=esk2_0|~v2_yellow_0(esk1_0))).#trainneg
cnf(i_0_1580, negated_conjecture, (v1_finset_1(esk2_0)|v2_waybel_6(esk2_0,esk1_0)|~v5_finset_1(k1_tarski(esk12_1(u1_struct_0(esk1_0))))|esk12_1(u1_struct_0(esk1_0))!=esk2_0)).#trainneg
cnf(i_0_361, plain, (m1_subset_1(esk4_3(X1,X2,X3),X3)|X3=k2_xboole_0(X1,X2)|r2_hidden(esk4_3(X1,X2,X3),X1)|r2_hidden(esk4_3(X1,X2,X3),X2))).#trainneg
cnf(i_0_1681, plain, (k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X4,X5),k7_domain_1(u1_struct_0(X1),X2,X3)))=k12_lattice3(X1,k12_lattice3(X1,X4,X5),k12_lattice3(X1,X2,X3))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X4,X5))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),k12_lattice3(X1,X2,X3)))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k12_lattice3(X1,X2,X3),u1_struct_0(X1)))).#trainneg
cnf(i_0_1702, plain, (k12_lattice3(X1,X2,k12_lattice3(X1,X3,X4))=k12_lattice3(X1,k12_lattice3(X1,X3,X4),X2)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1327, negated_conjecture, (v2_waybel_6(esk2_0,esk1_0)|m1_subset_1(esk2_0,k1_tarski(esk7_2(esk1_0,X1)))|v2_waybel_6(X1,esk1_0)|esk7_2(esk1_0,X1)!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1513, plain, (esk12_1(k2_tarski(X1,X2))=X1|esk12_1(k2_tarski(X1,X2))=X2)).#trainneg
cnf(i_0_1321, negated_conjecture, (v2_waybel_6(esk2_0,esk1_0)|m1_subset_1(esk2_0,k1_tarski(esk19_2(u1_struct_0(esk1_0),X1)))|v1_xboole_0(X1)|esk19_2(u1_struct_0(esk1_0),X1)!=esk2_0|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1361, plain, (k2_yellow_0(X1,k1_tarski(esk5_3(X1,X2,X3)))=X2|k2_yellow_0(X1,esk6_3(X1,X2,X3))=X2|r2_hidden(X2,X3)|v1_xboole_0(X3)|v1_xboole_0(u1_struct_0(X1))|X2!=k2_yellow_0(X1,X3)|~r2_yellow_0(X1,esk6_3(X1,X2,X3))|~v1_finset_1(X3)|~v2_waybel_6(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(X1)))).#trainneg
cnf(i_0_362, plain, (esk4_3(X1,X2,X3)=X4|X3=k2_xboole_0(X1,X2)|r2_hidden(esk4_3(X1,X2,X3),X1)|r2_hidden(esk4_3(X1,X2,X3),X2)|X3!=k1_tarski(X4))).#trainneg
cnf(i_0_721, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,X2)|~r1_tarski(X2,k4_ordinal1))).#trainneg
cnf(i_0_1174, plain, (v1_relat_1(k2_tarski(X2,X3))|v1_xboole_0(X1)|~v1_relat_1(X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_366, plain, (m1_subset_1(esk4_3(X1,X2,X3),X2)|X3=k2_xboole_0(X1,X2)|r2_hidden(esk4_3(X1,X2,X3),X1)|r2_hidden(esk4_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_367, plain, (esk4_3(X1,X2,X3)=X4|X3=k2_xboole_0(X1,X2)|r2_hidden(esk4_3(X1,X2,X3),X1)|r2_hidden(esk4_3(X1,X2,X3),X3)|X2!=k1_tarski(X4))).#trainneg
cnf(i_0_1329, negated_conjecture, (v2_waybel_6(esk2_0,esk1_0)|m1_subset_1(esk2_0,k1_tarski(k1_yellow_0(esk1_0,X1)))|k1_yellow_0(esk1_0,X1)!=esk2_0)).#trainneg
cnf(i_0_1499, plain, (k2_yellow_0(X1,k2_xboole_0(k1_tarski(X4),k2_xboole_0(esk6_3(X1,X2,X3),k1_tarski(esk5_3(X1,X2,X3)))))=k12_lattice3(X1,X2,X4)|r2_hidden(X2,X3)|v1_xboole_0(X3)|~r2_yellow_0(X1,k1_tarski(X4))|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X2,X3),k1_tarski(esk5_3(X1,X2,X3))))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|X2!=k2_yellow_0(X1,X3)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2133, plain, (X1=k2_yellow_0(X2,k2_xboole_0(esk6_3(X2,X3,X4),k1_tarski(esk5_3(X2,X3,X4))))|X1=k2_yellow_0(X2,k1_tarski(X5))|r2_hidden(X3,X4)|v1_xboole_0(X4)|X1!=k12_lattice3(X2,X3,X5)|~r2_yellow_0(X2,k2_xboole_0(esk6_3(X2,X3,X4),k1_tarski(esk5_3(X2,X3,X4))))|~r2_yellow_0(X2,k1_tarski(X5))|~v2_waybel_6(X1,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|X3!=k2_yellow_0(X2,X4)|~v1_finset_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~v3_orders_2(X2))).#trainneg
cnf(i_0_1932, negated_conjecture, (v2_waybel_6(esk2_0,esk1_0)|m1_subset_1(esk2_0,k1_tarski(k12_lattice3(esk1_0,k12_lattice3(esk1_0,X1,X2),k12_lattice3(esk1_0,X3,X4))))|k12_lattice3(esk1_0,k12_lattice3(esk1_0,X1,X2),k12_lattice3(esk1_0,X3,X4))!=esk2_0|~r2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),k12_lattice3(esk1_0,X3,X4)))|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X3,X4))|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2027, negated_conjecture, (v1_relat_1(k2_tarski(X1,esk2_0))|v2_waybel_6(esk2_0,esk1_0)|~v1_relat_1(k1_tarski(esk13_1(esk1_0)))|~m1_subset_1(X1,k1_tarski(esk13_1(esk1_0)))|esk13_1(esk1_0)!=esk2_0|~v2_yellow_0(esk1_0))).#trainneg
cnf(i_0_369, plain, (m1_subset_1(esk4_3(X1,X2,X3),X1)|X3=k2_xboole_0(X1,X2)|r2_hidden(esk4_3(X1,X2,X3),X2)|r2_hidden(esk4_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_370, plain, (esk4_3(X1,X2,X3)=X4|X3=k2_xboole_0(X1,X2)|r2_hidden(esk4_3(X1,X2,X3),X2)|r2_hidden(esk4_3(X1,X2,X3),X3)|X1!=k1_tarski(X4))).#trainneg
cnf(i_0_371, plain, (k11_lattice3(X1,X2,X3)=X2|k11_lattice3(X1,X2,X3)=X3|v2_struct_0(X1)|~v2_waybel_6(k11_lattice3(X1,X2,X3),X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k11_lattice3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1949, plain, (k12_lattice3(X1,k12_lattice3(X1,X2,X3),k12_lattice3(X1,X4,X5))=k12_lattice3(X1,k12_lattice3(X1,X4,X5),k12_lattice3(X1,X2,X3))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X4,X5))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),k12_lattice3(X1,X4,X5))))).#trainneg
cnf(i_0_1948, plain, (X1=k12_lattice3(X2,X3,X4)|X1=k2_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X5,X6))|X1!=k12_lattice3(X2,k12_lattice3(X2,X3,X4),k12_lattice3(X2,X5,X6))|~r2_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))|~r2_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X5,X6))|~v2_waybel_6(X1,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_yellow_0(X2,k6_domain_1(u1_struct_0(X2),k12_lattice3(X2,X5,X6)))|~m1_subset_1(X6,u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2)))).#trainneg
cnf(i_0_1405, plain, (r1_tarski(k7_domain_1(esk6_3(X1,X2,X3),X4,X5),X3)|r2_hidden(X2,X3)|v1_xboole_0(X3)|v1_xboole_0(esk6_3(X1,X2,X3))|X2!=k2_yellow_0(X1,X3)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X5,esk6_3(X1,X2,X3))|~m1_subset_1(X4,esk6_3(X1,X2,X3)))).#trainneg
cnf(i_0_2131, plain, (k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(k1_tarski(X4),k2_xboole_0(esk6_3(X1,X5,X6),k1_tarski(esk5_3(X1,X5,X6)))),k7_domain_1(u1_struct_0(X1),X2,X3)))=k12_lattice3(X1,k12_lattice3(X1,X2,X3),k12_lattice3(X1,X5,X4))|r2_hidden(X5,X6)|v1_xboole_0(X6)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,k2_xboole_0(k1_tarski(X4),k2_xboole_0(esk6_3(X1,X5,X6),k1_tarski(esk5_3(X1,X5,X6)))))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|X5!=k2_yellow_0(X1,X6)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X5,X6),k1_tarski(esk5_3(X1,X5,X6))))|~r2_yellow_0(X1,k1_tarski(X4))|~v1_finset_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1)))).#trainneg
cnf(i_0_1235, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~r1_tarski(X1,k2_tarski(X3,X4))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_1694, plain, (k2_yellow_0(X1,k2_xboole_0(k1_tarski(X2),k7_domain_1(u1_struct_0(X1),X3,X4)))=k12_lattice3(X1,k12_lattice3(X1,X3,X4),X2)|~r2_yellow_0(X1,k1_tarski(X2))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1652, negated_conjecture, (v2_waybel_6(esk2_0,esk1_0)|m1_subset_1(esk2_0,k1_tarski(k12_lattice3(esk1_0,X1,k2_yellow_0(esk1_0,X3))))|r2_hidden(X1,X2)|v1_xboole_0(X2)|k12_lattice3(esk1_0,X1,k2_yellow_0(esk1_0,X3))!=esk2_0|X1!=k2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,k2_xboole_0(esk6_3(esk1_0,X1,X2),k1_tarski(esk5_3(esk1_0,X1,X2))))|~r2_yellow_0(esk1_0,X3)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2028, negated_conjecture, (v1_relat_1(k2_tarski(X1,esk2_0))|v2_waybel_6(esk2_0,esk1_0)|~v1_relat_1(k1_tarski(esk12_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_tarski(esk12_1(u1_struct_0(esk1_0))))|esk12_1(u1_struct_0(esk1_0))!=esk2_0)).#trainneg
cnf(i_0_1637, negated_conjecture, (v1_finset_1(esk2_0)|v2_waybel_6(esk2_0,esk1_0)|~v5_finset_1(k1_tarski(k2_yellow_0(esk1_0,X1)))|k2_yellow_0(esk1_0,X1)!=esk2_0)).#trainneg
cnf(i_0_2114, negated_conjecture, (v1_finset_1(esk2_0)|v2_waybel_6(esk2_0,esk1_0)|~v5_finset_1(k1_tarski(k1_yellow_0(esk1_0,X1)))|k1_yellow_0(esk1_0,X1)!=esk2_0)).#trainneg
cnf(i_0_2237, plain, (k12_lattice3(X1,X2,X3)=X3|k12_lattice3(X1,X2,X3)=X2|v2_struct_0(X1)|~v2_waybel_6(k12_lattice3(X1,X2,X3),X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1))).#trainneg
cnf(i_0_2330, plain, (k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(k2_xboole_0(esk6_3(X1,X3,X4),k1_tarski(esk5_3(X1,X3,X4))),k1_tarski(X2)),k7_domain_1(u1_struct_0(X1),X5,X6)))=k12_lattice3(X1,k12_lattice3(X1,X5,X6),k12_lattice3(X1,X3,X2))|r2_hidden(X3,X4)|v1_xboole_0(X4)|X3!=k2_yellow_0(X1,X4)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X3,X4),k1_tarski(esk5_3(X1,X3,X4))))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X5,X6))|~r2_yellow_0(X1,k1_tarski(X2))|~v1_finset_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1071, plain, (k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X4,X5),k7_domain_1(u1_struct_0(X1),X2,X3)))=k12_lattice3(X1,k12_lattice3(X1,X4,X5),k12_lattice3(X1,X2,X3))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X4,X5))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1306, plain, (k12_lattice3(X1,k12_lattice3(X1,X3,X2),k12_lattice3(X1,X3,X2))=k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1411, plain, (r1_tarski(esk16_1(k1_xboole_0),X1))).#trainneg
cnf(i_0_1565, negated_conjecture, (esk2_0=X1|v2_waybel_6(esk2_0,esk1_0)|k1_tarski(esk13_1(esk1_0))!=k1_tarski(X1)|esk13_1(esk1_0)!=esk2_0|~v2_yellow_0(esk1_0))).#trainneg
cnf(i_0_400, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,esk12_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_388, plain, (k2_yellow_0(X1,k2_xboole_0(X2,k2_xboole_0(esk6_3(X1,X3,X4),k1_tarski(esk5_3(X1,X3,X4)))))=k12_lattice3(X1,k2_yellow_0(X1,X2),X3)|r2_hidden(X3,X4)|v1_xboole_0(X4)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X3,X4),k1_tarski(esk5_3(X1,X3,X4))))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|X3!=k2_yellow_0(X1,X4)|~v1_finset_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_401, plain, (r1_tarski(X1,u1_struct_0(X2))|v2_struct_0(X2)|v2_lattice3(X2)|~r1_tarski(X1,esk11_1(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_2304, plain, (r1_tarski(X1,X2)|r2_hidden(X4,X2)|v1_xboole_0(esk6_3(X3,X4,X2))|v1_xboole_0(X2)|~r1_tarski(X1,k7_domain_1(esk6_3(X3,X4,X2),X5,X6))|X4!=k2_yellow_0(X3,X2)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X6,esk6_3(X3,X4,X2))|~m1_subset_1(X5,esk6_3(X3,X4,X2))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_orders_2(X3)|~v2_lattice3(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3))).#trainneg
cnf(i_0_876, plain, (X1=k2_yellow_0(X2,X3)|X1=X4|X1!=k2_yellow_0(X2,k2_xboole_0(X3,k1_tarski(X4)))|~v2_waybel_6(X1,X2)|~m1_subset_1(k2_yellow_0(X2,X3),u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~r2_yellow_0(X2,k1_tarski(X4))|~r2_yellow_0(X2,X3)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_1107, plain, (k2_yellow_0(X1,k2_xboole_0(k1_tarski(X2),k2_xboole_0(esk6_3(X1,X3,X4),k1_tarski(esk5_3(X1,X3,X4)))))=k12_lattice3(X1,X2,X3)|r2_hidden(X3,X4)|v1_xboole_0(X4)|~r2_yellow_0(X1,k1_tarski(X2))|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X3,X4),k1_tarski(esk5_3(X1,X3,X4))))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|X3!=k2_yellow_0(X1,X4)|~v1_finset_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1697, plain, (X1=k2_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))|X1=X5|X1!=k12_lattice3(X2,k12_lattice3(X2,X3,X4),X5)|~v2_waybel_6(X1,X2)|~m1_subset_1(k2_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4)),u1_struct_0(X2))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~r2_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X3,X4))|~r2_yellow_0(X2,k6_domain_1(u1_struct_0(X2),X5))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_2440, negated_conjecture, (r2_hidden(esk2_0,k2_tarski(X2,X1))|v2_waybel_6(esk2_0,esk1_0)|k12_lattice3(esk1_0,X1,X2)!=esk2_0|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_740, plain, (m1_subset_1(X1,k2_tarski(X1,X2)))).#trainneg
cnf(i_0_2578, plain, (m1_subset_1(esk16_1(k1_xboole_0),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_417, plain, (r2_hidden(esk10_3(X1,X2,k1_xboole_0),X3)|esk10_3(X1,X2,k1_xboole_0)=X1|esk10_3(X1,X2,k1_xboole_0)=X2|k2_tarski(X1,X2)=k1_xboole_0)).#trainneg
cnf(i_0_1270, negated_conjecture, (v1_finset_1(esk2_0)|v2_waybel_6(esk2_0,esk1_0)|~v5_finset_1(k1_tarski(X1))|X1!=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1419, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_1600, plain, (v1_finset_1(esk4_3(X1,X1,X2))|X2=X1|r2_hidden(esk4_3(X1,X1,X2),X2)|~v5_finset_1(X1))).#trainneg
cnf(i_0_418, plain, (r2_hidden(esk9_2(X1,k1_xboole_0),X2)|esk9_2(X1,k1_xboole_0)=X1|k1_tarski(X1)=k1_xboole_0)).#trainneg
cnf(i_0_420, plain, (r2_hidden(esk4_3(X1,X2,k1_xboole_0),X3)|k2_xboole_0(X1,X2)=k1_xboole_0|r2_hidden(esk4_3(X1,X2,k1_xboole_0),X1)|r2_hidden(esk4_3(X1,X2,k1_xboole_0),X2))).#trainneg
cnf(i_0_2484, negated_conjecture, (v2_waybel_6(esk2_0,esk1_0)|m1_subset_1(esk2_0,k1_tarski(k12_lattice3(esk1_0,k12_lattice3(esk1_0,X1,X2),k12_lattice3(esk1_0,X3,X4))))|k12_lattice3(esk1_0,k12_lattice3(esk1_0,X1,X2),k12_lattice3(esk1_0,X3,X4))!=esk2_0|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X3,X4))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2502, plain, (k12_lattice3(X1,k12_lattice3(X1,X2,X3),k12_lattice3(X1,X4,X5))=k12_lattice3(X1,k12_lattice3(X1,X4,X5),k12_lattice3(X1,X2,X3))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X4,X5))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2756, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_2716, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk16_1(k1_xboole_0)))).#trainneg
cnf(i_0_2711, negated_conjecture, (r2_hidden(esk2_0,esk16_1(k1_xboole_0))|v1_xboole_0(esk16_1(k1_xboole_0))|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,esk16_1(k1_xboole_0))!=esk2_0|~v1_finset_1(esk16_1(k1_xboole_0)))).#trainneg
cnf(i_0_2706, plain, (esk16_1(k1_xboole_0)=X1|k1_zfmisc_1(X2)!=k1_tarski(X1))).#trainneg
cnf(i_0_421, plain, (r2_hidden(esk4_3(X1,k1_xboole_0,X2),X3)|X2=X1|r2_hidden(esk4_3(X1,k1_xboole_0,X2),X1)|r2_hidden(esk4_3(X1,k1_xboole_0,X2),X2))).#trainneg
cnf(i_0_2713, plain, (v1_xboole_0(esk16_1(k1_xboole_0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2714, plain, (v1_finset_1(esk16_1(k1_xboole_0))|~v1_finset_1(X1))).#trainneg
cnf(i_0_2715, plain, (v1_relat_1(esk16_1(k1_xboole_0))|~v1_relat_1(X1))).#trainneg
cnf(i_0_422, plain, (r2_hidden(esk4_3(k1_xboole_0,X1,X2),X3)|X2=k2_xboole_0(k1_xboole_0,X1)|r2_hidden(esk4_3(k1_xboole_0,X1,X2),X1)|r2_hidden(esk4_3(k1_xboole_0,X1,X2),X2))).#trainneg
cnf(i_0_424, plain, (m1_subset_1(esk10_3(X1,X2,esk16_1(X3)),X3)|esk10_3(X1,X2,esk16_1(X3))=X1|esk10_3(X1,X2,esk16_1(X3))=X2|esk16_1(X3)=k2_tarski(X1,X2))).#trainneg
cnf(i_0_2949, plain, (k12_lattice3(X1,X2,k12_lattice3(X1,X3,X4))=k12_lattice3(X1,k12_lattice3(X1,X3,X4),X2)|v2_waybel_6(X2,X1)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),esk7_2(X1,X2),esk8_2(X1,X2)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk7_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2856, negated_conjecture, (esk5_3(esk1_0,X1,X2)=esk2_0|k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))=esk2_0|v2_waybel_6(esk2_0,esk1_0)|r2_hidden(X1,X2)|v1_xboole_0(X2)|X1!=esk2_0|~m1_subset_1(k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2)),u1_struct_0(esk1_0))|~m1_subset_1(esk5_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|X1!=k2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,k1_tarski(esk5_3(esk1_0,X1,X2)))|~r2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_789, plain, (v1_relat_1(k1_tarski(X1))|~v1_relat_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_425, plain, (m1_subset_1(esk9_2(X1,esk16_1(X2)),X2)|esk9_2(X1,esk16_1(X2))=X1|esk16_1(X2)=k1_tarski(X1))).#trainneg
cnf(i_0_3145, plain, (v1_xboole_0(esk16_1(k1_xboole_0)))).#trainneg
cnf(i_0_3160, plain, (v1_finset_1(esk16_1(k1_xboole_0)))).#trainneg
cnf(i_0_426, plain, (m1_subset_1(esk5_3(X1,X2,esk16_1(X3)),X3)|r2_hidden(X2,esk16_1(X3))|v1_xboole_0(esk16_1(X3))|X2!=k2_yellow_0(X1,esk16_1(X3))|~v1_finset_1(esk16_1(X3))|~m1_subset_1(esk16_1(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_427, plain, (m1_subset_1(esk4_3(X1,X2,esk16_1(X3)),X3)|esk16_1(X3)=k2_xboole_0(X1,X2)|r2_hidden(esk4_3(X1,X2,esk16_1(X3)),X1)|r2_hidden(esk4_3(X1,X2,esk16_1(X3)),X2))).#trainneg
cnf(i_0_3275, negated_conjecture, (k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))=esk2_0|esk5_3(esk1_0,X1,X2)=esk2_0|r2_hidden(X1,X2)|v1_xboole_0(X2)|v2_waybel_6(esk2_0,esk1_0)|X1!=k2_yellow_0(esk1_0,X2)|X1!=esk2_0|~r2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))|~v1_finset_1(X2)|~m1_subset_1(k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2)),u1_struct_0(esk1_0))|~m1_subset_1(esk5_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(esk5_3(esk1_0,X1,X2)),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2862, negated_conjecture, (k2_yellow_0(esk1_0,X1)=esk2_0|k2_yellow_0(esk1_0,X2)=esk2_0|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,k2_xboole_0(X2,X1))!=esk2_0|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,X1)|~r2_yellow_0(esk1_0,X2))).#trainneg
cnf(i_0_2857, negated_conjecture, (k12_lattice3(esk1_0,X1,X2)=esk2_0|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))!=esk2_0|~m1_subset_1(k12_lattice3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2595, negated_conjecture, (v2_waybel_6(esk2_0,esk1_0)|m1_subset_1(esk2_0,k1_tarski(k12_lattice3(esk1_0,k2_yellow_0(esk1_0,X1),X2)))|r2_hidden(X2,X3)|v1_xboole_0(X3)|k12_lattice3(esk1_0,k2_yellow_0(esk1_0,X1),X2)!=esk2_0|X2!=k2_yellow_0(esk1_0,X3)|~r2_yellow_0(esk1_0,k2_xboole_0(esk6_3(esk1_0,X2,X3),k1_tarski(esk5_3(esk1_0,X2,X3))))|~r2_yellow_0(esk1_0,X1)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2604, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k12_lattice3(X1,k2_yellow_0(X1,X3),X4))=k2_yellow_0(X1,k2_xboole_0(X2,k2_xboole_0(X3,k2_xboole_0(esk6_3(X1,X4,X5),k1_tarski(esk5_3(X1,X4,X5))))))|r2_hidden(X4,X5)|v1_xboole_0(X5)|~r2_yellow_0(X1,k2_xboole_0(X3,k2_xboole_0(esk6_3(X1,X4,X5),k1_tarski(esk5_3(X1,X4,X5)))))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|X4!=k2_yellow_0(X1,X5)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X4,X5),k1_tarski(esk5_3(X1,X4,X5))))|~r2_yellow_0(X1,X3)|~v1_finset_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_2858, negated_conjecture, (k12_lattice3(esk1_0,X1,X2)=esk2_0|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,X1))!=esk2_0|~m1_subset_1(k12_lattice3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_720, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k4_ordinal1))).#trainneg
cnf(i_0_2712, plain, (~r2_hidden(X1,esk16_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_428, plain, (m1_subset_1(esk4_3(X1,esk16_1(X2),X3),X2)|X3=k2_xboole_0(X1,esk16_1(X2))|r2_hidden(esk4_3(X1,esk16_1(X2),X3),X1)|r2_hidden(esk4_3(X1,esk16_1(X2),X3),X3))).#trainneg
cnf(i_0_3298, plain, (esk16_1(k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_3541, plain, (m1_subset_1(esk16_1(k4_ordinal1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1620, plain, (v1_finset_1(esk4_3(X1,X1,X2))|X2=X1|r2_hidden(esk4_3(X1,X1,X2),X1)|~v5_finset_1(X2))).#trainneg
cnf(i_0_429, plain, (m1_subset_1(esk4_3(esk16_1(X1),X2,X3),X1)|X3=k2_xboole_0(esk16_1(X1),X2)|r2_hidden(esk4_3(esk16_1(X1),X2,X3),X2)|r2_hidden(esk4_3(esk16_1(X1),X2,X3),X3))).#trainneg
cnf(i_0_430, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk16_1(X2))|~m1_subset_1(X1,esk16_1(X2)))).#trainneg
cnf(i_0_3388, negated_conjecture, (k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))=esk2_0|esk5_3(esk1_0,X1,X2)=esk2_0|r2_hidden(X1,X2)|v1_xboole_0(X2)|v2_waybel_6(esk2_0,esk1_0)|X1!=k2_yellow_0(esk1_0,X2)|X1!=esk2_0|~r2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))|~v1_finset_1(X2)|~m1_subset_1(k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2)),u1_struct_0(esk1_0))|~m1_subset_1(esk5_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3398, negated_conjecture, (k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1))=esk2_0|k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,X3))=esk2_0|v2_waybel_6(esk2_0,esk1_0)|k12_lattice3(esk1_0,k12_lattice3(esk1_0,X2,X3),X1)!=esk2_0|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,X3))|~r2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1))|~m1_subset_1(k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1)),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,X3)),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3396, negated_conjecture, (k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))=esk2_0|k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X3,X4))=esk2_0|v2_waybel_6(esk2_0,esk1_0)|k12_lattice3(esk1_0,k12_lattice3(esk1_0,X1,X2),k12_lattice3(esk1_0,X3,X4))!=esk2_0|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X3,X4))|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~m1_subset_1(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2)),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X3,X4)),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3395, negated_conjecture, (k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))=esk2_0|k2_yellow_0(esk1_0,k1_tarski(esk5_3(esk1_0,X1,X2)))=esk2_0|v2_waybel_6(esk2_0,esk1_0)|r2_hidden(X1,X2)|v1_xboole_0(X2)|X1!=esk2_0|~r2_yellow_0(esk1_0,k1_tarski(esk5_3(esk1_0,X1,X2)))|~r2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))|~m1_subset_1(k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2)),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(esk5_3(esk1_0,X1,X2))),u1_struct_0(esk1_0))|X1!=k2_yellow_0(esk1_0,X2)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3550, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_3278, negated_conjecture, (v1_relat_1(k2_tarski(X1,esk2_0))|v2_waybel_6(esk2_0,esk1_0)|esk13_1(esk1_0)!=esk2_0|~v2_yellow_0(esk1_0)|~m1_subset_1(X1,k1_tarski(esk13_1(esk1_0)))|~v1_relat_1(k1_zfmisc_1(esk13_1(esk1_0))))).#trainneg
cnf(i_0_447, plain, (k2_xboole_0(X1,X2)=k2_tarski(X3,X4)|esk10_3(X3,X4,k2_xboole_0(X1,X2))!=X4|~r2_hidden(esk10_3(X3,X4,k2_xboole_0(X1,X2)),X1))).#trainneg
cnf(i_0_2004, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,esk16_1(k4_ordinal1)))).#trainneg
cnf(i_0_448, plain, (k2_xboole_0(X1,X2)=k2_tarski(X3,X4)|esk10_3(X3,X4,k2_xboole_0(X1,X2))!=X3|~r2_hidden(esk10_3(X3,X4,k2_xboole_0(X1,X2)),X1))).#trainneg
cnf(i_0_449, plain, (k2_xboole_0(X1,X2)=k1_tarski(X3)|esk9_2(X3,k2_xboole_0(X1,X2))!=X3|~r2_hidden(esk9_2(X3,k2_xboole_0(X1,X2)),X1))).#trainneg
cnf(i_0_3771, negated_conjecture, (k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))=esk2_0|esk5_3(esk1_0,X1,X2)=esk2_0|r2_hidden(X1,X2)|v1_xboole_0(X2)|v2_waybel_6(esk2_0,esk1_0)|v1_xboole_0(esk6_3(esk1_0,X1,X2))|X1!=k2_yellow_0(esk1_0,X2)|X1!=esk2_0|~v1_finset_1(X2)|~m1_subset_1(k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2)),u1_struct_0(esk1_0))|~m1_subset_1(esk5_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v1_finset_1(esk6_3(esk1_0,X1,X2))|~m1_subset_1(esk6_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3816, negated_conjecture, (k2_yellow_0(esk1_0,k1_tarski(esk5_3(esk1_0,X1,X2)))=esk2_0|k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))=esk2_0|r2_hidden(X1,X2)|v1_xboole_0(X2)|v2_waybel_6(esk2_0,esk1_0)|X1!=k2_yellow_0(esk1_0,X2)|X1!=esk2_0|~r2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))|~v1_finset_1(X2)|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(esk5_3(esk1_0,X1,X2))),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2)),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(esk5_3(esk1_0,X1,X2)),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3410, negated_conjecture, (k2_yellow_0(esk1_0,k2_xboole_0(esk6_3(esk1_0,X1,X2),k1_tarski(esk5_3(esk1_0,X1,X2))))=esk2_0|k2_yellow_0(esk1_0,X3)=esk2_0|v2_waybel_6(esk2_0,esk1_0)|r2_hidden(X1,X2)|v1_xboole_0(X2)|k12_lattice3(esk1_0,X1,k2_yellow_0(esk1_0,X3))!=esk2_0|~r2_yellow_0(esk1_0,X3)|~r2_yellow_0(esk1_0,k2_xboole_0(esk6_3(esk1_0,X1,X2),k1_tarski(esk5_3(esk1_0,X1,X2))))|~m1_subset_1(k2_yellow_0(esk1_0,k2_xboole_0(esk6_3(esk1_0,X1,X2),k1_tarski(esk5_3(esk1_0,X1,X2)))),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,X3),u1_struct_0(esk1_0))|X1!=k2_yellow_0(esk1_0,X2)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3417, negated_conjecture, (k12_lattice3(esk1_0,X1,X2)=esk2_0|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,k2_tarski(X2,X1))!=esk2_0|~r2_yellow_0(esk1_0,k2_tarski(X2,X1))|~m1_subset_1(k12_lattice3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3419, negated_conjecture, (k12_lattice3(esk1_0,X1,X2)=esk2_0|v2_waybel_6(esk2_0,esk1_0)|v1_xboole_0(k7_domain_1(u1_struct_0(esk1_0),X1,X2))|k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))!=esk2_0|~m1_subset_1(k12_lattice3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v1_finset_1(k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~m1_subset_1(k7_domain_1(u1_struct_0(esk1_0),X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3402, negated_conjecture, (k2_yellow_0(esk1_0,k1_tarski(X1))=esk2_0|k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,X3))=esk2_0|v2_waybel_6(esk2_0,esk1_0)|k12_lattice3(esk1_0,k12_lattice3(esk1_0,X2,X3),X1)!=esk2_0|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,X3))|~r2_yellow_0(esk1_0,k1_tarski(X1))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(X1)),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,X3)),u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3409, negated_conjecture, (k2_yellow_0(esk1_0,k2_xboole_0(esk6_3(esk1_0,X1,X2),k1_tarski(esk5_3(esk1_0,X1,X2))))=esk2_0|k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X3,X4))=esk2_0|v2_waybel_6(esk2_0,esk1_0)|r2_hidden(X1,X2)|v1_xboole_0(X2)|k12_lattice3(esk1_0,k12_lattice3(esk1_0,X3,X4),X1)!=esk2_0|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X3,X4))|~r2_yellow_0(esk1_0,k2_xboole_0(esk6_3(esk1_0,X1,X2),k1_tarski(esk5_3(esk1_0,X1,X2))))|~m1_subset_1(k2_yellow_0(esk1_0,k2_xboole_0(esk6_3(esk1_0,X1,X2),k1_tarski(esk5_3(esk1_0,X1,X2)))),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X3,X4)),u1_struct_0(esk1_0))|X1!=k2_yellow_0(esk1_0,X2)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3538, negated_conjecture, (k12_lattice3(esk1_0,X1,X2)=esk2_0|v2_waybel_6(esk2_0,esk1_0)|v1_xboole_0(k7_domain_1(u1_struct_0(esk1_0),X2,X1))|k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,X1))!=esk2_0|~m1_subset_1(k12_lattice3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~v1_finset_1(k7_domain_1(u1_struct_0(esk1_0),X2,X1))|~m1_subset_1(k7_domain_1(u1_struct_0(esk1_0),X2,X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3400, negated_conjecture, (k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1))=esk2_0|k2_yellow_0(esk1_0,X2)=esk2_0|v2_waybel_6(esk2_0,esk1_0)|k12_lattice3(esk1_0,X1,k2_yellow_0(esk1_0,X2))!=esk2_0|~r2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1))|~m1_subset_1(k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1)),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3401, negated_conjecture, (k2_yellow_0(esk1_0,k1_tarski(X1))=esk2_0|k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,X3))=esk2_0|v2_waybel_6(esk2_0,esk1_0)|k12_lattice3(esk1_0,X1,k12_lattice3(esk1_0,X2,X3))!=esk2_0|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,X3))|~r2_yellow_0(esk1_0,k1_tarski(X1))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(X1)),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,X3)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_450, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X3,X4)|~r2_hidden(esk4_3(X3,X4,k2_xboole_0(X1,X2)),X4)|~r2_hidden(esk4_3(X3,X4,k2_xboole_0(X1,X2)),X1))).#trainneg
cnf(i_0_3654, plain, (~v1_subset_1(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_2757, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_451, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X3,X4)|~r2_hidden(esk4_3(X3,X4,k2_xboole_0(X1,X2)),X3)|~r2_hidden(esk4_3(X3,X4,k2_xboole_0(X1,X2)),X1))).#trainneg
cnf(i_0_454, plain, (X1=X2|k2_xboole_0(X3,X4)!=k1_tarski(X2)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_3661, plain, (~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_3415, negated_conjecture, (k12_lattice3(esk1_0,X1,X2)=esk2_0|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,k2_tarski(X1,X2))!=esk2_0|~r2_yellow_0(esk1_0,k2_tarski(X1,X2))|~m1_subset_1(k12_lattice3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2704, plain, (m1_subset_1(X1,k2_tarski(X2,X1)))).#trainneg
cnf(i_0_3404, negated_conjecture, (k2_yellow_0(esk1_0,k1_tarski(X1))=esk2_0|k2_yellow_0(esk1_0,k2_xboole_0(esk6_3(esk1_0,X2,X3),k1_tarski(esk5_3(esk1_0,X2,X3))))=esk2_0|v2_waybel_6(esk2_0,esk1_0)|r2_hidden(X2,X3)|v1_xboole_0(X3)|k12_lattice3(esk1_0,X2,X1)!=esk2_0|~r2_yellow_0(esk1_0,k2_xboole_0(esk6_3(esk1_0,X2,X3),k1_tarski(esk5_3(esk1_0,X2,X3))))|~r2_yellow_0(esk1_0,k1_tarski(X1))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(X1)),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k2_xboole_0(esk6_3(esk1_0,X2,X3),k1_tarski(esk5_3(esk1_0,X2,X3)))),u1_struct_0(esk1_0))|X2!=k2_yellow_0(esk1_0,X3)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2526, plain, (k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X2))=k12_lattice3(X1,X2,X3)|~v2_waybel_6(k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X2)),X1)|~m1_subset_1(k12_lattice3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3411, negated_conjecture, (k2_yellow_0(esk1_0,X1)=esk2_0|k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X2))=esk2_0|v2_waybel_6(esk2_0,esk1_0)|k12_lattice3(esk1_0,k2_yellow_0(esk1_0,X1),X2)!=esk2_0|~r2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X2))|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X2)),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3407, negated_conjecture, (k2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X1),k2_xboole_0(esk6_3(esk1_0,X2,X3),k1_tarski(esk5_3(esk1_0,X2,X3)))))=esk2_0|k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X4,X5))=esk2_0|v2_waybel_6(esk2_0,esk1_0)|r2_hidden(X2,X3)|v1_xboole_0(X3)|k12_lattice3(esk1_0,k12_lattice3(esk1_0,X4,X5),k12_lattice3(esk1_0,X2,X1))!=esk2_0|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X4,X5))|~r2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X1),k2_xboole_0(esk6_3(esk1_0,X2,X3),k1_tarski(esk5_3(esk1_0,X2,X3)))))|~m1_subset_1(k2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X1),k2_xboole_0(esk6_3(esk1_0,X2,X3),k1_tarski(esk5_3(esk1_0,X2,X3))))),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X4,X5)),u1_struct_0(esk1_0))|X2!=k2_yellow_0(esk1_0,X3)|~r2_yellow_0(esk1_0,k2_xboole_0(esk6_3(esk1_0,X2,X3),k1_tarski(esk5_3(esk1_0,X2,X3))))|~r2_yellow_0(esk1_0,k1_tarski(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X5,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2436, negated_conjecture, (m1_subset_1(esk2_0,k2_tarski(X1,X2))|v2_waybel_6(esk2_0,esk1_0)|k12_lattice3(esk1_0,X1,X2)!=esk2_0|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3675, plain, (v1_relat_1(esk16_1(k4_ordinal1))|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_3668, plain, (esk16_1(k4_ordinal1)=X1|k1_tarski(X1)!=k1_zfmisc_1(k1_numbers))).#trainneg
cnf(i_0_471, plain, (r2_hidden(esk10_3(X1,X2,k2_xboole_0(X3,X4)),X4)|r2_hidden(esk10_3(X1,X2,k2_xboole_0(X3,X4)),X3)|esk10_3(X1,X2,k2_xboole_0(X3,X4))=X1|esk10_3(X1,X2,k2_xboole_0(X3,X4))=X2|k2_xboole_0(X3,X4)=k2_tarski(X1,X2))).#trainneg
cnf(i_0_472, plain, (r2_hidden(esk9_2(X1,k2_xboole_0(X2,X3)),X3)|r2_hidden(esk9_2(X1,k2_xboole_0(X2,X3)),X2)|esk9_2(X1,k2_xboole_0(X2,X3))=X1|k2_xboole_0(X2,X3)=k1_tarski(X1))).#trainneg
cnf(i_0_473, plain, (r2_hidden(esk5_3(X1,X2,k2_xboole_0(X3,X4)),X4)|r2_hidden(esk5_3(X1,X2,k2_xboole_0(X3,X4)),X3)|r2_hidden(X2,k2_xboole_0(X3,X4))|v1_xboole_0(k2_xboole_0(X3,X4))|X2!=k2_yellow_0(X1,k2_xboole_0(X3,X4))|~v1_finset_1(k2_xboole_0(X3,X4))|~m1_subset_1(k2_xboole_0(X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3876, negated_conjecture, (k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))=esk2_0|esk5_3(esk1_0,X1,X2)=esk2_0|r2_hidden(X1,X2)|v1_xboole_0(esk6_3(esk1_0,X1,X2))|v1_xboole_0(X2)|v2_waybel_6(esk2_0,esk1_0)|X1!=k2_yellow_0(esk1_0,X2)|X1!=esk2_0|~v1_finset_1(esk6_3(esk1_0,X1,X2))|~v1_finset_1(X2)|~m1_subset_1(esk6_3(esk1_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk5_3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3891, negated_conjecture, (k12_lattice3(esk1_0,X1,X2)=esk2_0|v1_xboole_0(k7_domain_1(u1_struct_0(esk1_0),X1,X2))|v2_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))!=esk2_0|~v1_finset_1(k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~m1_subset_1(k12_lattice3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3880, negated_conjecture, (k2_yellow_0(esk1_0,k1_tarski(esk5_3(esk1_0,X1,X2)))=esk2_0|k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))=esk2_0|r2_hidden(X1,X2)|v1_xboole_0(X2)|v2_waybel_6(esk2_0,esk1_0)|X1!=k2_yellow_0(esk1_0,X2)|X1!=esk2_0|~r2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2))|~v1_finset_1(X2)|~m1_subset_1(k2_yellow_0(esk1_0,esk6_3(esk1_0,X1,X2)),u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(esk5_3(esk1_0,X1,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2417, plain, (esk5_3(X1,X2,k2_tarski(X3,X4))=X3|esk5_3(X1,X2,k2_tarski(X3,X4))=X4|r2_hidden(X2,k2_tarski(X3,X4))|X2!=k2_yellow_0(X1,k2_tarski(X3,X4))|~m1_subset_1(k2_tarski(X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3405, negated_conjecture, (k2_yellow_0(esk1_0,k1_tarski(X1))=esk2_0|k2_yellow_0(esk1_0,X2)=esk2_0|v2_waybel_6(esk2_0,esk1_0)|k12_lattice3(esk1_0,X1,k2_yellow_0(esk1_0,X2))!=esk2_0|~r2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,k1_tarski(X1))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(X1)),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3849, plain, (r1_tarski(esk16_1(k4_ordinal1),k1_numbers))).#trainneg
cnf(i_0_474, plain, (r2_hidden(esk4_3(X1,X2,k2_xboole_0(X3,X4)),X4)|r2_hidden(esk4_3(X1,X2,k2_xboole_0(X3,X4)),X3)|k2_xboole_0(X3,X4)=k2_xboole_0(X1,X2)|r2_hidden(esk4_3(X1,X2,k2_xboole_0(X3,X4)),X1)|r2_hidden(esk4_3(X1,X2,k2_xboole_0(X3,X4)),X2))).#trainneg
cnf(i_0_3670, plain, (v1_relat_1(k1_tarski(esk16_1(k4_ordinal1)))|~v1_relat_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3841, plain, (r1_tarski(esk16_1(esk16_1(k4_ordinal1)),k1_numbers))).#trainneg
cnf(i_0_475, plain, (r2_hidden(esk4_3(X1,k2_xboole_0(X2,X3),X4),X3)|r2_hidden(esk4_3(X1,k2_xboole_0(X2,X3),X4),X2)|X4=k2_xboole_0(X1,k2_xboole_0(X2,X3))|r2_hidden(esk4_3(X1,k2_xboole_0(X2,X3),X4),X1)|r2_hidden(esk4_3(X1,k2_xboole_0(X2,X3),X4),X4))).#trainneg
cnf(i_0_476, plain, (r2_hidden(esk4_3(k2_xboole_0(X1,X2),X3,X4),X2)|r2_hidden(esk4_3(k2_xboole_0(X1,X2),X3,X4),X1)|X4=k2_xboole_0(k2_xboole_0(X1,X2),X3)|r2_hidden(esk4_3(k2_xboole_0(X1,X2),X3,X4),X3)|r2_hidden(esk4_3(k2_xboole_0(X1,X2),X3,X4),X4))).#trainneg
cnf(i_0_3464, plain, (k2_yellow_0(X1,k2_xboole_0(X2,k2_xboole_0(X3,k2_xboole_0(esk6_3(X1,X4,X6),k1_tarski(esk5_3(X1,X4,X6))))))=k2_yellow_0(X1,k2_xboole_0(X2,k2_xboole_0(X3,k2_xboole_0(esk6_3(X1,X4,X5),k1_tarski(esk5_3(X1,X4,X5))))))|r2_hidden(X4,X5)|v1_xboole_0(X5)|r2_hidden(X4,X6)|v1_xboole_0(X6)|X4!=k2_yellow_0(X1,X5)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X4,X5),k1_tarski(esk5_3(X1,X4,X5))))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,X3)|~v1_finset_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|X4!=k2_yellow_0(X1,X6)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X4,X6),k1_tarski(esk5_3(X1,X4,X6))))|~v1_finset_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_2979, plain, (k12_lattice3(X1,k12_lattice3(X1,X2,X3),k12_lattice3(X1,X5,X4))=k12_lattice3(X1,k12_lattice3(X1,X5,X4),k12_lattice3(X1,X2,X3))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X4,X5))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3676, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk16_1(k4_ordinal1)))).#trainneg
cnf(i_0_480, plain, (k2_xboole_0(X1,X2)=k2_tarski(X3,X4)|esk10_3(X3,X4,k2_xboole_0(X1,X2))!=X4|~r2_hidden(esk10_3(X3,X4,k2_xboole_0(X1,X2)),X2))).#trainneg
cnf(i_0_4392, plain, (m1_subset_1(esk16_1(esk16_1(k4_ordinal1)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2593, plain, (r1_tarski(esk12_1(k1_zfmisc_1(X1)),X1))).#trainneg
cnf(i_0_481, plain, (k2_xboole_0(X1,X2)=k2_tarski(X3,X4)|esk10_3(X3,X4,k2_xboole_0(X1,X2))!=X3|~r2_hidden(esk10_3(X3,X4,k2_xboole_0(X1,X2)),X2))).#trainneg
cnf(i_0_482, plain, (k2_xboole_0(X1,X2)=k1_tarski(X3)|esk9_2(X3,k2_xboole_0(X1,X2))!=X3|~r2_hidden(esk9_2(X3,k2_xboole_0(X1,X2)),X2))).#trainneg
cnf(i_0_4768, plain, (k12_lattice3(X1,k12_lattice3(X1,X2,X3),X4)=k12_lattice3(X1,X4,k12_lattice3(X1,X2,X3))|v2_waybel_6(X4,X1)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),esk8_2(X1,X4),esk7_2(X1,X4)))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~m1_subset_1(esk7_2(X1,X4),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_3828, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_3827, plain, (v1_relat_1(X1)|~v1_relat_1(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_4004, plain, (X1=X2|k2_xboole_0(k1_tarski(X1),X3)!=k1_tarski(X2))).#trainneg
cnf(i_0_483, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X3,X4)|~r2_hidden(esk4_3(X3,X4,k2_xboole_0(X1,X2)),X4)|~r2_hidden(esk4_3(X3,X4,k2_xboole_0(X1,X2)),X2))).#trainneg
cnf(i_0_484, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X3,X4)|~r2_hidden(esk4_3(X3,X4,k2_xboole_0(X1,X2)),X3)|~r2_hidden(esk4_3(X3,X4,k2_xboole_0(X1,X2)),X2))).#trainneg
cnf(i_0_487, plain, (X1=X2|k2_xboole_0(X3,X4)!=k1_tarski(X2)|~r2_hidden(X1,X4))).#trainneg
cnf(i_0_2446, plain, (k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X2,X3),X4))=k12_lattice3(X1,X2,X3)|k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X2,X3),X4))=k2_yellow_0(X1,X4)|~v2_waybel_6(k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X2,X3),X4)),X1)|~m1_subset_1(k2_yellow_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(k12_lattice3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_4840, negated_conjecture, (v1_xboole_0(esk3_0)|~m1_subset_1(esk2_0,esk3_0))).#trainneg
cnf(i_0_4896, negated_conjecture, (v1_relat_1(k1_tarski(esk3_0))|~v1_relat_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4894, negated_conjecture, (esk3_0=X1|k1_tarski(X1)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_497, plain, (esk10_3(X1,X2,k2_tarski(X3,X4))=X4|esk10_3(X1,X2,k2_tarski(X3,X4))=X3|esk10_3(X1,X2,k2_tarski(X3,X4))=X1|esk10_3(X1,X2,k2_tarski(X3,X4))=X2|k2_tarski(X3,X4)=k2_tarski(X1,X2))).#trainneg
cnf(i_0_4901, negated_conjecture, (v1_relat_1(esk3_0)|~v1_relat_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_986, plain, (v1_finset_1(esk19_2(X1,X2))|v1_xboole_0(X2)|~v5_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1251, plain, (v1_finset_1(esk19_2(X1,X2))|v1_xboole_0(X2)|~v5_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_498, plain, (esk9_2(X1,k2_tarski(X2,X3))=X3|esk9_2(X1,k2_tarski(X2,X3))=X2|esk9_2(X1,k2_tarski(X2,X3))=X1|k2_tarski(X2,X3)=k1_tarski(X1))).#trainneg
cnf(i_0_500, plain, (esk4_3(X1,X2,k2_tarski(X3,X4))=X4|esk4_3(X1,X2,k2_tarski(X3,X4))=X3|k2_tarski(X3,X4)=k2_xboole_0(X1,X2)|r2_hidden(esk4_3(X1,X2,k2_tarski(X3,X4)),X1)|r2_hidden(esk4_3(X1,X2,k2_tarski(X3,X4)),X2))).#trainneg
cnf(i_0_2532, plain, (X1=k12_lattice3(X2,X3,X4)|X1!=k2_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X4,X3))|~v2_waybel_6(X1,X2)|~m1_subset_1(k12_lattice3(X2,X3,X4),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~r2_yellow_0(X2,k7_domain_1(u1_struct_0(X2),X4,X3))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_2447, plain, (k2_yellow_0(X1,k2_xboole_0(X2,X3))=k2_yellow_0(X1,X2)|k2_yellow_0(X1,k2_xboole_0(X2,X3))=k2_yellow_0(X1,X3)|~v2_waybel_6(k2_yellow_0(X1,k2_xboole_0(X2,X3)),X1)|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2)|~v4_orders_2(X1))).#trainneg
cnf(i_0_2602, plain, (k12_lattice3(X1,k12_lattice3(X1,k2_yellow_0(X1,X2),X3),k2_yellow_0(X1,X5))=k2_yellow_0(X1,k2_xboole_0(X5,k2_xboole_0(X2,k2_xboole_0(esk6_3(X1,X3,X4),k1_tarski(esk5_3(X1,X3,X4))))))|r2_hidden(X3,X4)|v1_xboole_0(X4)|~r2_yellow_0(X1,k2_xboole_0(X2,k2_xboole_0(esk6_3(X1,X3,X4),k1_tarski(esk5_3(X1,X3,X4)))))|~r2_yellow_0(X1,X5)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|X3!=k2_yellow_0(X1,X4)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X3,X4),k1_tarski(esk5_3(X1,X3,X4))))|~r2_yellow_0(X1,X2)|~v1_finset_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_1008, plain, (r1_tarski(esk16_1(esk16_1(X1)),X1))).#trainneg
cnf(i_0_4921, plain, (v1_relat_1(esk16_1(esk16_1(k4_ordinal1)))|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_4914, plain, (esk16_1(esk16_1(k4_ordinal1))=X1|k1_tarski(X1)!=k1_zfmisc_1(k1_numbers))).#trainneg
cnf(i_0_501, plain, (esk4_3(X1,k2_tarski(X2,X3),X4)=X3|esk4_3(X1,k2_tarski(X2,X3),X4)=X2|X4=k2_xboole_0(X1,k2_tarski(X2,X3))|r2_hidden(esk4_3(X1,k2_tarski(X2,X3),X4),X1)|r2_hidden(esk4_3(X1,k2_tarski(X2,X3),X4),X4))).#trainneg
cnf(i_0_1729, plain, (v1_finset_1(esk4_3(X1,X2,X1))|k2_xboole_0(X1,X2)=X1|r2_hidden(esk4_3(X1,X2,X1),X1)|~v5_finset_1(X2))).#trainneg
cnf(i_0_1748, plain, (v1_finset_1(esk4_3(X1,X2,X1))|k2_xboole_0(X1,X2)=X1|r2_hidden(esk4_3(X1,X2,X1),X2)|~v5_finset_1(X1))).#trainneg
cnf(i_0_502, plain, (esk4_3(k2_tarski(X1,X2),X3,X4)=X2|esk4_3(k2_tarski(X1,X2),X3,X4)=X1|X4=k2_xboole_0(k2_tarski(X1,X2),X3)|r2_hidden(esk4_3(k2_tarski(X1,X2),X3,X4),X3)|r2_hidden(esk4_3(k2_tarski(X1,X2),X3,X4),X4))).#trainneg
cnf(i_0_507, plain, (v2_struct_0(X1)|m1_subset_1(esk10_3(X2,X3,esk17_1(X1)),u1_struct_0(X1))|esk10_3(X2,X3,esk17_1(X1))=X2|esk10_3(X2,X3,esk17_1(X1))=X3|esk17_1(X1)=k2_tarski(X2,X3)|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5479, plain, (k12_lattice3(X1,k12_lattice3(X1,X3,k2_yellow_0(X1,X2)),k2_yellow_0(X1,X4))=k2_yellow_0(X1,k2_xboole_0(X4,k2_xboole_0(X2,k2_xboole_0(esk6_3(X1,X3,X5),k1_tarski(esk5_3(X1,X3,X5))))))|r2_hidden(X3,X5)|v1_xboole_0(X5)|X3!=k2_yellow_0(X1,X5)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X3,X5),k1_tarski(esk5_3(X1,X3,X5))))|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,X2)|~v1_finset_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_5717, plain, (k12_lattice3(X1,k12_lattice3(X1,X2,k2_yellow_0(X1,X3)),k2_yellow_0(X1,X4))=k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(X3,k2_xboole_0(esk6_3(X1,X2,X5),k1_tarski(esk5_3(X1,X2,X5)))),X4))|r2_hidden(X2,X5)|v1_xboole_0(X5)|X2!=k2_yellow_0(X1,X5)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X2,X5),k1_tarski(esk5_3(X1,X2,X5))))|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,X3)|~v1_finset_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5336, plain, (k12_lattice3(X1,k12_lattice3(X1,X2,X3),k12_lattice3(X1,X4,X5))=k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X4,X5))|k12_lattice3(X1,k12_lattice3(X1,X2,X3),k12_lattice3(X1,X4,X5))=k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X4,X5))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~v2_waybel_6(k12_lattice3(X1,k12_lattice3(X1,X2,X3),k12_lattice3(X1,X4,X5)),X1)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_4922, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk16_1(esk16_1(k4_ordinal1))))).#trainneg
cnf(i_0_4940, plain, (r1_tarski(esk12_1(k1_zfmisc_1(k1_xboole_0)),X1))).#trainneg
cnf(i_0_5797, negated_conjecture, (k1_tarski(X1)=X2|k1_tarski(X2)!=k1_zfmisc_1(u1_struct_0(esk1_0))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_508, plain, (v2_struct_0(X1)|m1_subset_1(esk9_2(X2,esk17_1(X1)),u1_struct_0(X1))|esk9_2(X2,esk17_1(X1))=X2|esk17_1(X1)=k1_tarski(X2)|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_4916, plain, (v1_relat_1(k1_tarski(esk16_1(esk16_1(k4_ordinal1))))|~v1_relat_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1783, plain, (v1_finset_1(esk4_3(X1,X2,X2))|k2_xboole_0(X1,X2)=X2|r2_hidden(esk4_3(X1,X2,X2),X2)|~v5_finset_1(X1))).#trainneg
cnf(i_0_1803, plain, (v1_finset_1(esk4_3(X1,X2,X2))|k2_xboole_0(X1,X2)=X2|r2_hidden(esk4_3(X1,X2,X2),X1)|~v5_finset_1(X2))).#trainneg
cnf(i_0_509, plain, (v2_struct_0(X1)|m1_subset_1(esk5_3(X2,X3,esk17_1(X1)),u1_struct_0(X1))|r2_hidden(X3,esk17_1(X1))|v1_xboole_0(esk17_1(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)|X3!=k2_yellow_0(X2,esk17_1(X1))|~v1_finset_1(esk17_1(X1))|~m1_subset_1(esk17_1(X1),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_510, plain, (v2_struct_0(X1)|m1_subset_1(esk4_3(X2,X3,esk17_1(X1)),u1_struct_0(X1))|esk17_1(X1)=k2_xboole_0(X2,X3)|r2_hidden(esk4_3(X2,X3,esk17_1(X1)),X2)|r2_hidden(esk4_3(X2,X3,esk17_1(X1)),X3)|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5872, plain, (k12_lattice3(X1,k12_lattice3(X1,X2,k2_yellow_0(X1,X3)),k2_yellow_0(X1,X4))=k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(k2_xboole_0(esk6_3(X1,X2,X5),k1_tarski(esk5_3(X1,X2,X5))),X3),X4))|r2_hidden(X2,X5)|v1_xboole_0(X5)|X2!=k2_yellow_0(X1,X5)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X2,X5),k1_tarski(esk5_3(X1,X2,X5))))|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,X3)|~v1_finset_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5844, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k12_lattice3(X1,X3,k2_yellow_0(X1,X4)))=k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(X4,k2_xboole_0(esk6_3(X1,X3,X5),k1_tarski(esk5_3(X1,X3,X5)))),X2))|r2_hidden(X3,X5)|v1_xboole_0(X5)|~m1_subset_1(k12_lattice3(X1,X3,k2_yellow_0(X1,X4)),u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|X3!=k2_yellow_0(X1,X5)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X3,X5),k1_tarski(esk5_3(X1,X3,X5))))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,X4)|~v1_finset_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5867, plain, (k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(X3,k2_xboole_0(esk6_3(X1,X2,X6),k1_tarski(esk5_3(X1,X2,X6)))),X4))=k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(X3,k2_xboole_0(esk6_3(X1,X2,X5),k1_tarski(esk5_3(X1,X2,X5)))),X4))|r2_hidden(X2,X5)|v1_xboole_0(X5)|r2_hidden(X2,X6)|v1_xboole_0(X6)|X2!=k2_yellow_0(X1,X5)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X2,X5),k1_tarski(esk5_3(X1,X2,X5))))|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,X3)|~v1_finset_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|X2!=k2_yellow_0(X1,X6)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X2,X6),k1_tarski(esk5_3(X1,X2,X6))))|~v1_finset_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_3025, plain, (m1_subset_1(esk12_1(k1_zfmisc_1(k1_xboole_0)),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_511, plain, (v2_struct_0(X1)|m1_subset_1(esk4_3(X2,esk17_1(X1),X3),u1_struct_0(X1))|X3=k2_xboole_0(X2,esk17_1(X1))|r2_hidden(esk4_3(X2,esk17_1(X1),X3),X2)|r2_hidden(esk4_3(X2,esk17_1(X1),X3),X3)|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_788, plain, (v1_relat_1(k1_tarski(k2_yellow_0(X1,X2)))|v1_xboole_0(u1_struct_0(X1))|~v1_relat_1(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_1003, plain, (r1_tarski(k2_tarski(X2,X3),X1)|v1_xboole_0(X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_512, plain, (v2_struct_0(X1)|m1_subset_1(esk4_3(esk17_1(X1),X2,X3),u1_struct_0(X1))|X3=k2_xboole_0(esk17_1(X1),X2)|r2_hidden(esk4_3(esk17_1(X1),X2,X3),X2)|r2_hidden(esk4_3(esk17_1(X1),X2,X3),X3)|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_513, plain, (v2_struct_0(X1)|m1_subset_1(X2,u1_struct_0(X1))|v1_xboole_0(esk17_1(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,esk17_1(X1)))).#trainneg
cnf(i_0_6307, plain, (k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(X4,k2_xboole_0(esk6_3(X1,X3,X5),k1_tarski(esk5_3(X1,X3,X5)))),X2))=k2_yellow_0(X1,X2)|k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(X4,k2_xboole_0(esk6_3(X1,X3,X5),k1_tarski(esk5_3(X1,X3,X5)))),X2))=k12_lattice3(X1,X3,k2_yellow_0(X1,X4))|r2_hidden(X3,X5)|v1_xboole_0(X5)|~v2_waybel_6(k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(X4,k2_xboole_0(esk6_3(X1,X3,X5),k1_tarski(esk5_3(X1,X3,X5)))),X2)),X1)|~m1_subset_1(k12_lattice3(X1,X3,k2_yellow_0(X1,X4)),u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|X3!=k2_yellow_0(X1,X5)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X3,X5),k1_tarski(esk5_3(X1,X3,X5))))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,X4)|~v1_finset_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6207, plain, (k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(k2_xboole_0(esk6_3(X1,X2,X5),k1_tarski(esk5_3(X1,X2,X5))),k6_domain_1(u1_struct_0(X1),X3)),X4))=k12_lattice3(X1,k12_lattice3(X1,X2,X3),k2_yellow_0(X1,X4))|r2_hidden(X2,X5)|v1_xboole_0(X5)|v2_struct_0(X1)|X2!=k2_yellow_0(X1,X5)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X2,X5),k1_tarski(esk5_3(X1,X2,X5))))|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X3))|~v1_finset_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_6495, plain, (k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(k2_xboole_0(esk6_3(X1,X2,X3),k1_tarski(esk5_3(X1,X2,X3))),k2_xboole_0(esk6_3(X1,X2,X4),k1_tarski(esk5_3(X1,X2,X4)))),X5))=k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(esk6_3(X1,X2,X3),k1_tarski(esk5_3(X1,X2,X3))),X5))|r2_hidden(X2,X4)|r2_hidden(X2,X3)|v1_xboole_0(X4)|v1_xboole_0(X3)|X2!=k2_yellow_0(X1,X3)|X2!=k2_yellow_0(X1,X4)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X2,X3),k1_tarski(esk5_3(X1,X2,X3))))|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X2,X4),k1_tarski(esk5_3(X1,X2,X4))))|~r2_yellow_0(X1,X5)|~v1_finset_1(X3)|~v1_finset_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5887, plain, (k12_lattice3(X1,k12_lattice3(X1,X2,k2_yellow_0(X1,X3)),k12_lattice3(X1,k12_lattice3(X1,X6,X7),k12_lattice3(X1,X4,X5)))=k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(X3,k2_xboole_0(esk6_3(X1,X2,X8),k1_tarski(esk5_3(X1,X2,X8)))),k2_xboole_0(k7_domain_1(u1_struct_0(X1),X4,X5),k7_domain_1(u1_struct_0(X1),X6,X7))))|r2_hidden(X2,X8)|v1_xboole_0(X8)|X2!=k2_yellow_0(X1,X8)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X2,X8),k1_tarski(esk5_3(X1,X2,X8))))|~r2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X4,X5),k7_domain_1(u1_struct_0(X1),X6,X7)))|~r2_yellow_0(X1,X3)|~v1_finset_1(X8)|~m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X6,X7))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X4,X5))|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_6334, plain, (k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(k1_tarski(esk5_3(X1,X4,X5)),k2_xboole_0(esk6_3(X1,X3,X6),k1_tarski(esk5_3(X1,X3,X6)))),X2))=k12_lattice3(X1,k2_yellow_0(X1,X2),k12_lattice3(X1,X3,X4))|r2_hidden(X3,X6)|v1_xboole_0(X6)|k2_yellow_0(X1,esk6_3(X1,X4,X5))=X4|r2_hidden(X4,X5)|v1_xboole_0(X5)|X3!=k2_yellow_0(X1,X6)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X3,X6),k1_tarski(esk5_3(X1,X3,X6))))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,k1_tarski(esk5_3(X1,X4,X5)))|~v1_finset_1(X6)|~m1_subset_1(k12_lattice3(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|X4!=k2_yellow_0(X1,X5)|~r2_yellow_0(X1,esk6_3(X1,X4,X5))|~v1_finset_1(X5)|~v2_waybel_6(X4,X1)|~m1_subset_1(esk5_3(X1,X4,X5),u1_struct_0(X1))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_6711, plain, (X1=k2_yellow_0(X2,X3)|X1=k2_yellow_0(X2,k2_xboole_0(k2_xboole_0(esk6_3(X2,X4,X5),k1_tarski(esk5_3(X2,X4,X5))),k2_xboole_0(esk6_3(X2,X4,X6),k1_tarski(esk5_3(X2,X4,X6)))))|r2_hidden(X4,X5)|r2_hidden(X4,X6)|v1_xboole_0(X5)|v1_xboole_0(X6)|X1!=k2_yellow_0(X2,k2_xboole_0(k2_xboole_0(esk6_3(X2,X4,X5),k1_tarski(esk5_3(X2,X4,X5))),X3))|~r2_yellow_0(X2,X3)|~r2_yellow_0(X2,k2_xboole_0(k2_xboole_0(esk6_3(X2,X4,X5),k1_tarski(esk5_3(X2,X4,X5))),k2_xboole_0(esk6_3(X2,X4,X6),k1_tarski(esk5_3(X2,X4,X6)))))|~v2_waybel_6(X1,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|X4!=k2_yellow_0(X2,X5)|X4!=k2_yellow_0(X2,X6)|~r2_yellow_0(X2,k2_xboole_0(esk6_3(X2,X4,X5),k1_tarski(esk5_3(X2,X4,X5))))|~r2_yellow_0(X2,k2_xboole_0(esk6_3(X2,X4,X6),k1_tarski(esk5_3(X2,X4,X6))))|~v1_finset_1(X5)|~v1_finset_1(X6)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,u1_struct_0(X2))|~v3_orders_2(X2))).#trainneg
cnf(i_0_6723, plain, (k2_yellow_0(X1,k2_xboole_0(X5,k2_xboole_0(k2_xboole_0(esk6_3(X1,X2,X3),k1_tarski(esk5_3(X1,X2,X3))),k2_xboole_0(esk6_3(X1,X2,X4),k1_tarski(esk5_3(X1,X2,X4))))))=k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(esk6_3(X1,X2,X3),k1_tarski(esk5_3(X1,X2,X3))),X5))|r2_hidden(X2,X3)|r2_hidden(X2,X4)|v1_xboole_0(X3)|v1_xboole_0(X4)|X2!=k2_yellow_0(X1,X3)|X2!=k2_yellow_0(X1,X4)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X2,X3),k1_tarski(esk5_3(X1,X2,X3))))|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X2,X4),k1_tarski(esk5_3(X1,X2,X4))))|~r2_yellow_0(X1,X5)|~v1_finset_1(X3)|~v1_finset_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6520, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_4933, plain, (m1_subset_1(esk12_1(k1_zfmisc_1(k4_ordinal1)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6509, plain, (esk12_1(k1_zfmisc_1(k1_xboole_0))=X1|k1_zfmisc_1(X2)!=k1_tarski(X1))).#trainneg
cnf(i_0_518, plain, (esk12_1(X1)=X2|v1_xboole_0(X1)|X1!=k1_tarski(X2))).#trainneg
cnf(i_0_3941, plain, (r1_tarski(X1,X2)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k1_xboole_0))|~r1_tarski(X1,k4_ordinal1))).#trainneg
cnf(i_0_4928, plain, (r1_tarski(esk12_1(k1_zfmisc_1(esk16_1(k4_ordinal1))),k1_numbers))).#trainneg
cnf(i_0_5484, negated_conjecture, (m1_subset_1(esk16_1(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_520, plain, (esk8_2(X1,X2)=X3|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|v2_waybel_6(X2,X1)|u1_struct_0(X1)!=k1_tarski(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_521, plain, (esk7_2(X1,X2)=X3|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|v2_waybel_6(X2,X1)|u1_struct_0(X1)!=k1_tarski(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_5866, plain, (k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(X3,k2_xboole_0(esk6_3(X1,X2,X6),k1_tarski(esk5_3(X1,X2,X6)))),X4))=k2_yellow_0(X1,k2_xboole_0(X4,k2_xboole_0(X3,k2_xboole_0(esk6_3(X1,X2,X5),k1_tarski(esk5_3(X1,X2,X5))))))|r2_hidden(X2,X5)|v1_xboole_0(X5)|r2_hidden(X2,X6)|v1_xboole_0(X6)|X2!=k2_yellow_0(X1,X5)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X2,X5),k1_tarski(esk5_3(X1,X2,X5))))|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,X3)|~v1_finset_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|X2!=k2_yellow_0(X1,X6)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X2,X6),k1_tarski(esk5_3(X1,X2,X6))))|~v1_finset_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_7123, plain, (esk12_1(k1_tarski(X1))=X1)).#trainneg
cnf(i_0_7263, plain, (k2_yellow_0(X1,k2_xboole_0(X6,k2_xboole_0(k1_tarski(esk5_3(X1,X2,X3)),k2_xboole_0(esk6_3(X1,X4,X7),k1_tarski(esk5_3(X1,X4,X7))))))=k12_lattice3(X1,k2_yellow_0(X1,X6),k12_lattice3(X1,X4,X2))|k2_yellow_0(X1,esk6_3(X1,X2,X3))=X2|r2_hidden(X2,X3)|r2_hidden(X4,X5)|v1_xboole_0(X3)|v1_xboole_0(X5)|r2_hidden(X4,X7)|v1_xboole_0(X7)|X4!=k2_yellow_0(X1,X5)|X2!=k2_yellow_0(X1,X3)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X4,X5),k1_tarski(esk5_3(X1,X4,X5))))|~r2_yellow_0(X1,k1_tarski(esk5_3(X1,X2,X3)))|~r2_yellow_0(X1,esk6_3(X1,X2,X3))|~r2_yellow_0(X1,X6)|~v1_finset_1(X5)|~v1_finset_1(X3)|~v2_waybel_6(X2,X1)|~m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|X4!=k2_yellow_0(X1,X7)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X4,X7),k1_tarski(esk5_3(X1,X4,X7))))|~v1_finset_1(X7)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_7151, negated_conjecture, (v1_finset_1(esk16_1(esk3_0))|~v1_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7145, negated_conjecture, (esk16_1(esk3_0)=X1|k1_tarski(X1)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_527, plain, (k1_yellow_0(X1,X2)=X3|v1_xboole_0(u1_struct_0(X1))|u1_struct_0(X1)!=k1_tarski(X3)|~l1_orders_2(X1))).#trainneg
cnf(i_0_7137, plain, (r1_tarski(k4_ordinal1,X1)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_7152, negated_conjecture, (v1_relat_1(esk16_1(esk3_0))|~v1_relat_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7146, negated_conjecture, (v1_finset_1(esk16_1(esk3_0))|~v5_finset_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_529, plain, (k12_lattice3(X1,X2,X3)=X4|v1_xboole_0(u1_struct_0(X1))|u1_struct_0(X1)!=k1_tarski(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1))).#trainneg
cnf(i_0_530, plain, (k11_lattice3(X1,X2,X3)=X4|v1_xboole_0(u1_struct_0(X1))|u1_struct_0(X1)!=k1_tarski(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_6978, plain, (X1=k2_yellow_0(X2,k2_xboole_0(k2_xboole_0(esk6_3(X2,X3,X4),k1_tarski(esk5_3(X2,X3,X4))),k2_xboole_0(esk6_3(X2,X3,X5),k1_tarski(esk5_3(X2,X3,X5)))))|X1=k2_yellow_0(X2,X6)|r2_hidden(X3,X5)|r2_hidden(X3,X4)|v1_xboole_0(X5)|v1_xboole_0(X4)|X1!=k2_yellow_0(X2,k2_xboole_0(X6,k2_xboole_0(esk6_3(X2,X3,X4),k1_tarski(esk5_3(X2,X3,X4)))))|X3!=k2_yellow_0(X2,X4)|X3!=k2_yellow_0(X2,X5)|~r2_yellow_0(X2,k2_xboole_0(esk6_3(X2,X3,X4),k1_tarski(esk5_3(X2,X3,X4))))|~r2_yellow_0(X2,k2_xboole_0(esk6_3(X2,X3,X5),k1_tarski(esk5_3(X2,X3,X5))))|~r2_yellow_0(X2,X6)|~v1_finset_1(X4)|~v1_finset_1(X5)|~v2_waybel_6(X1,X2)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_7252, plain, (X1=k2_yellow_0(X2,X3)|X1=k2_yellow_0(X2,k2_xboole_0(X4,k2_xboole_0(esk6_3(X2,X5,X6),k1_tarski(esk5_3(X2,X5,X6)))))|r2_hidden(X5,X6)|r2_hidden(X5,X7)|v1_xboole_0(X6)|v1_xboole_0(X7)|X1!=k2_yellow_0(X2,k2_xboole_0(X3,k2_xboole_0(X4,k2_xboole_0(esk6_3(X2,X5,X7),k1_tarski(esk5_3(X2,X5,X7))))))|~r2_yellow_0(X2,X3)|~r2_yellow_0(X2,k2_xboole_0(X4,k2_xboole_0(esk6_3(X2,X5,X6),k1_tarski(esk5_3(X2,X5,X6)))))|~v2_waybel_6(X1,X2)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|X5!=k2_yellow_0(X2,X7)|X5!=k2_yellow_0(X2,X6)|~r2_yellow_0(X2,k2_xboole_0(esk6_3(X2,X5,X7),k1_tarski(esk5_3(X2,X5,X7))))|~r2_yellow_0(X2,k2_xboole_0(esk6_3(X2,X5,X6),k1_tarski(esk5_3(X2,X5,X6))))|~r2_yellow_0(X2,X4)|~v1_finset_1(X7)|~v1_finset_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X5,u1_struct_0(X2))|~v3_orders_2(X2))).#trainneg
cnf(i_0_3475, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k12_lattice3(X1,k2_yellow_0(X1,X3),X4))=k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(X3,k2_xboole_0(esk6_3(X1,X4,X5),k1_tarski(esk5_3(X1,X4,X5)))),X2))|r2_hidden(X4,X5)|v1_xboole_0(X5)|X4!=k2_yellow_0(X1,X5)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X4,X5),k1_tarski(esk5_3(X1,X4,X5))))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,X3)|~v1_finset_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_4971, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk12_1(k1_zfmisc_1(k4_ordinal1))))).#trainneg
cnf(i_0_7111, plain, (esk12_1(k1_zfmisc_1(k4_ordinal1))=X1|k1_tarski(X1)!=k1_zfmisc_1(k1_numbers))).#trainneg
cnf(i_0_531, plain, (k2_yellow_0(X1,X2)=X3|v1_xboole_0(u1_struct_0(X1))|u1_struct_0(X1)!=k1_tarski(X3)|~l1_orders_2(X1))).#trainneg
cnf(i_0_6517, plain, (v1_xboole_0(esk12_1(k1_zfmisc_1(k1_xboole_0)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6518, plain, (v1_finset_1(esk12_1(k1_zfmisc_1(k1_xboole_0)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_541, plain, (k2_yellow_0(X1,k2_tarski(X3,X2))=k12_lattice3(X1,X2,X3)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_546, plain, (m1_subset_1(esk10_3(X1,X2,esk12_1(k1_zfmisc_1(X3))),X3)|esk10_3(X1,X2,esk12_1(k1_zfmisc_1(X3)))=X1|esk10_3(X1,X2,esk12_1(k1_zfmisc_1(X3)))=X2|esk12_1(k1_zfmisc_1(X3))=k2_tarski(X1,X2))).#trainneg
cnf(i_0_7473, plain, (k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(k1_tarski(esk5_3(X1,X3,X4)),k2_xboole_0(esk6_3(X1,X5,X6),k1_tarski(esk5_3(X1,X5,X6)))),X2))=k12_lattice3(X1,k2_yellow_0(X1,X2),k12_lattice3(X1,X3,X5))|r2_hidden(X5,X6)|v1_xboole_0(X6)|k2_yellow_0(X1,esk6_3(X1,X3,X4))=X3|r2_hidden(X3,X4)|v1_xboole_0(X4)|X5!=k2_yellow_0(X1,X6)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X5,X6),k1_tarski(esk5_3(X1,X5,X6))))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,k1_tarski(esk5_3(X1,X3,X4)))|~v1_finset_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|X3!=k2_yellow_0(X1,X4)|~r2_yellow_0(X1,esk6_3(X1,X3,X4))|~v1_finset_1(X4)|~v2_waybel_6(X3,X1)|~m1_subset_1(esk5_3(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_6159, plain, (k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(k2_xboole_0(esk6_3(X1,X2,X6),k1_tarski(esk5_3(X1,X2,X6))),X3),X4))=k2_yellow_0(X1,k2_xboole_0(X4,k2_xboole_0(X3,k2_xboole_0(esk6_3(X1,X2,X5),k1_tarski(esk5_3(X1,X2,X5))))))|r2_hidden(X2,X5)|v1_xboole_0(X5)|r2_hidden(X2,X6)|v1_xboole_0(X6)|X2!=k2_yellow_0(X1,X5)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X2,X5),k1_tarski(esk5_3(X1,X2,X5))))|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,X3)|~v1_finset_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|X2!=k2_yellow_0(X1,X6)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X2,X6),k1_tarski(esk5_3(X1,X2,X6))))|~v1_finset_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_4904, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_4762, plain, (k12_lattice3(X1,k12_lattice3(X1,X3,X2),k12_lattice3(X1,X4,X5))=k12_lattice3(X1,k12_lattice3(X1,X4,X5),k12_lattice3(X1,X3,X2))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X5,X4))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_7621, plain, (v1_xboole_0(esk12_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_7634, negated_conjecture, (v1_finset_1(esk12_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_547, plain, (m1_subset_1(esk9_2(X1,esk12_1(k1_zfmisc_1(X2))),X2)|esk9_2(X1,esk12_1(k1_zfmisc_1(X2)))=X1|esk12_1(k1_zfmisc_1(X2))=k1_tarski(X1))).#trainneg
cnf(i_0_7118, plain, (v1_relat_1(esk12_1(k1_zfmisc_1(k4_ordinal1)))|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_6519, plain, (v1_relat_1(esk12_1(k1_zfmisc_1(k1_xboole_0)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_548, plain, (m1_subset_1(esk5_3(X1,X2,esk12_1(k1_zfmisc_1(X3))),X3)|r2_hidden(X2,esk12_1(k1_zfmisc_1(X3)))|v1_xboole_0(esk12_1(k1_zfmisc_1(X3)))|X2!=k2_yellow_0(X1,esk12_1(k1_zfmisc_1(X3)))|~v1_finset_1(esk12_1(k1_zfmisc_1(X3)))|~m1_subset_1(esk12_1(k1_zfmisc_1(X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_549, plain, (m1_subset_1(esk4_3(X1,X2,esk12_1(k1_zfmisc_1(X3))),X3)|esk12_1(k1_zfmisc_1(X3))=k2_xboole_0(X1,X2)|r2_hidden(esk4_3(X1,X2,esk12_1(k1_zfmisc_1(X3))),X1)|r2_hidden(esk4_3(X1,X2,esk12_1(k1_zfmisc_1(X3))),X2))).#trainneg
cnf(i_0_5828, plain, (k12_lattice3(X1,k12_lattice3(X1,X2,X3),k12_lattice3(X1,k12_lattice3(X1,X5,k2_yellow_0(X1,X4)),k2_yellow_0(X1,X7)))=k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(k2_xboole_0(X4,k2_xboole_0(esk6_3(X1,X5,X6),k1_tarski(esk5_3(X1,X5,X6)))),X7),k7_domain_1(u1_struct_0(X1),X2,X3)))|r2_hidden(X5,X6)|v1_xboole_0(X6)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,k2_xboole_0(k2_xboole_0(X4,k2_xboole_0(esk6_3(X1,X5,X6),k1_tarski(esk5_3(X1,X5,X6)))),X7))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|X5!=k2_yellow_0(X1,X6)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X5,X6),k1_tarski(esk5_3(X1,X5,X6))))|~r2_yellow_0(X1,X7)|~r2_yellow_0(X1,X4)|~v1_finset_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,u1_struct_0(X1)))).#trainneg
cnf(i_0_1007, plain, (v2_struct_0(X1)|r1_tarski(esk16_1(esk17_1(X1)),u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1969, plain, (esk12_1(k2_tarski(X1,X1))=X1)).#trainneg
cnf(i_0_5511, plain, (m1_subset_1(esk16_1(esk16_1(X1)),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_8490, plain, (esk16_1(esk16_1(X1))=X2|k1_zfmisc_1(X1)!=k1_tarski(X2))).#trainneg
cnf(i_0_550, plain, (m1_subset_1(esk4_3(X1,esk12_1(k1_zfmisc_1(X2)),X3),X2)|X3=k2_xboole_0(X1,esk12_1(k1_zfmisc_1(X2)))|r2_hidden(esk4_3(X1,esk12_1(k1_zfmisc_1(X2)),X3),X1)|r2_hidden(esk4_3(X1,esk12_1(k1_zfmisc_1(X2)),X3),X3))).#trainneg
cnf(i_0_551, plain, (m1_subset_1(esk4_3(esk12_1(k1_zfmisc_1(X1)),X2,X3),X1)|X3=k2_xboole_0(esk12_1(k1_zfmisc_1(X1)),X2)|r2_hidden(esk4_3(esk12_1(k1_zfmisc_1(X1)),X2,X3),X2)|r2_hidden(esk4_3(esk12_1(k1_zfmisc_1(X1)),X2,X3),X3))).#trainneg
cnf(i_0_552, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk12_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,esk12_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_3504, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k12_lattice3(X1,X4,k2_yellow_0(X1,X3)))=k2_yellow_0(X1,k2_xboole_0(X2,k2_xboole_0(X3,k2_xboole_0(esk6_3(X1,X4,X5),k1_tarski(esk5_3(X1,X4,X5))))))|r2_hidden(X4,X5)|v1_xboole_0(X5)|X4!=k2_yellow_0(X1,X5)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X4,X5),k1_tarski(esk5_3(X1,X4,X5))))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,X3)|~v1_finset_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1)))).#trainneg
cnf(i_0_8804, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k12_lattice3(X1,X3,k2_yellow_0(X1,X4)))=k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(X4,k2_xboole_0(esk6_3(X1,X3,X5),k1_tarski(esk5_3(X1,X3,X5)))),X2))|r2_hidden(X3,X5)|v1_xboole_0(X5)|X3!=k2_yellow_0(X1,X5)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X3,X5),k1_tarski(esk5_3(X1,X3,X5))))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,X4)|~v1_finset_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_8802, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k12_lattice3(X1,X3,k2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X3,X4),k1_tarski(esk5_3(X1,X3,X4))))))=k2_yellow_0(X1,k2_xboole_0(X2,k2_xboole_0(esk6_3(X1,X3,X4),k1_tarski(esk5_3(X1,X3,X4)))))|r2_hidden(X3,X4)|v1_xboole_0(X4)|X3!=k2_yellow_0(X1,X4)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X3,X4),k1_tarski(esk5_3(X1,X3,X4))))|~r2_yellow_0(X1,X2)|~v1_finset_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_8501, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk16_1(esk16_1(X2))))).#trainneg
cnf(i_0_8498, plain, (v1_xboole_0(esk16_1(esk16_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8091, plain, (esk12_1(k1_zfmisc_1(k1_xboole_0))=k1_xboole_0)).#trainneg
cnf(i_0_4992, plain, (X1=X2|k2_xboole_0(X3,k1_tarski(X1))!=k1_tarski(X2))).#trainneg
cnf(i_0_562, plain, (m1_subset_1(esk10_3(X1,X2,k4_ordinal1),k1_numbers)|esk10_3(X1,X2,k4_ordinal1)=X1|esk10_3(X1,X2,k4_ordinal1)=X2|k2_tarski(X1,X2)=k4_ordinal1)).#trainneg
cnf(i_0_7147, negated_conjecture, (v1_relat_1(k1_tarski(esk16_1(esk3_0)))|~v1_relat_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7113, plain, (v1_relat_1(k1_tarski(esk12_1(k1_zfmisc_1(k4_ordinal1))))|~v1_relat_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_563, plain, (m1_subset_1(esk9_2(X1,k4_ordinal1),k1_numbers)|esk9_2(X1,k4_ordinal1)=X1|k1_tarski(X1)=k4_ordinal1)).#trainneg
cnf(i_0_565, plain, (m1_subset_1(esk4_3(X1,X2,k4_ordinal1),k1_numbers)|k2_xboole_0(X1,X2)=k4_ordinal1|r2_hidden(esk4_3(X1,X2,k4_ordinal1),X1)|r2_hidden(esk4_3(X1,X2,k4_ordinal1),X2))).#trainneg
cnf(i_0_8800, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k12_lattice3(X1,X3,k2_yellow_0(X1,X4)))=k2_yellow_0(X1,k2_xboole_0(X2,k2_xboole_0(k2_xboole_0(esk6_3(X1,X3,X5),k1_tarski(esk5_3(X1,X3,X5))),X4)))|r2_hidden(X3,X5)|v1_xboole_0(X5)|X3!=k2_yellow_0(X1,X5)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X3,X5),k1_tarski(esk5_3(X1,X3,X5))))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,X4)|~v1_finset_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_9534, plain, (X1=k12_lattice3(X2,X3,k2_yellow_0(X2,X4))|X1=k2_yellow_0(X2,X5)|r2_hidden(X3,X6)|v1_xboole_0(X6)|X1!=k2_yellow_0(X2,k2_xboole_0(X5,k2_xboole_0(k2_xboole_0(esk6_3(X2,X3,X6),k1_tarski(esk5_3(X2,X3,X6))),X4)))|~v2_waybel_6(X1,X2)|~m1_subset_1(k12_lattice3(X2,X3,k2_yellow_0(X2,X4)),u1_struct_0(X2))|~m1_subset_1(k2_yellow_0(X2,X5),u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|X3!=k2_yellow_0(X2,X6)|~r2_yellow_0(X2,k2_xboole_0(esk6_3(X2,X3,X6),k1_tarski(esk5_3(X2,X3,X6))))|~r2_yellow_0(X2,X5)|~r2_yellow_0(X2,X4)|~v1_finset_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,u1_struct_0(X2))|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_7202, plain, (k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(X3,k2_xboole_0(esk6_3(X1,X4,X6),k1_tarski(esk5_3(X1,X4,X6)))),X2))=k2_yellow_0(X1,k2_xboole_0(X3,k2_xboole_0(esk6_3(X1,X4,X5),k1_tarski(esk5_3(X1,X4,X5)))))|k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(X3,k2_xboole_0(esk6_3(X1,X4,X6),k1_tarski(esk5_3(X1,X4,X6)))),X2))=k2_yellow_0(X1,X2)|r2_hidden(X4,X6)|r2_hidden(X4,X5)|v1_xboole_0(X6)|v1_xboole_0(X5)|~r2_yellow_0(X1,k2_xboole_0(X3,k2_xboole_0(esk6_3(X1,X4,X5),k1_tarski(esk5_3(X1,X4,X5)))))|~r2_yellow_0(X1,X2)|~v2_waybel_6(k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(X3,k2_xboole_0(esk6_3(X1,X4,X6),k1_tarski(esk5_3(X1,X4,X6)))),X2)),X1)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|X4!=k2_yellow_0(X1,X5)|X4!=k2_yellow_0(X1,X6)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X4,X5),k1_tarski(esk5_3(X1,X4,X5))))|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X4,X6),k1_tarski(esk5_3(X1,X4,X6))))|~r2_yellow_0(X1,X3)|~v1_finset_1(X5)|~v1_finset_1(X6)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_9559, plain, (k2_yellow_0(X1,k2_xboole_0(X2,k2_xboole_0(k2_xboole_0(esk6_3(X1,X3,X6),k1_tarski(esk5_3(X1,X3,X6))),k2_xboole_0(esk6_3(X1,X4,X5),k1_tarski(esk5_3(X1,X4,X5))))))=k12_lattice3(X1,k2_yellow_0(X1,X2),k12_lattice3(X1,X3,X4))|r2_hidden(X3,X6)|v1_xboole_0(X6)|r2_hidden(X4,X5)|v1_xboole_0(X5)|X3!=k2_yellow_0(X1,X6)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X3,X6),k1_tarski(esk5_3(X1,X3,X6))))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X4,X5),k1_tarski(esk5_3(X1,X4,X5))))|~v1_finset_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|X4!=k2_yellow_0(X1,X5)|~v1_finset_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_566, plain, (m1_subset_1(esk4_3(X1,k4_ordinal1,X2),k1_numbers)|X2=k2_xboole_0(X1,k4_ordinal1)|r2_hidden(esk4_3(X1,k4_ordinal1,X2),X1)|r2_hidden(esk4_3(X1,k4_ordinal1,X2),X2))).#trainneg
cnf(i_0_567, plain, (m1_subset_1(esk4_3(k4_ordinal1,X1,X2),k1_numbers)|X2=k2_xboole_0(k4_ordinal1,X1)|r2_hidden(esk4_3(k4_ordinal1,X1,X2),X1)|r2_hidden(esk4_3(k4_ordinal1,X1,X2),X2))).#trainneg
cnf(i_0_9693, negated_conjecture, (~m1_subset_1(esk2_0,k1_tarski(esk5_3(esk1_0,esk2_0,esk3_0))))).#trainneg
cnf(i_0_9667, negated_conjecture, (k2_yellow_0(esk1_0,k2_xboole_0(k2_xboole_0(k1_tarski(esk5_3(esk1_0,esk2_0,esk3_0)),k2_xboole_0(esk6_3(esk1_0,esk2_0,X1),k1_tarski(esk5_3(esk1_0,esk2_0,X1)))),X2))=k2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(esk5_3(esk1_0,esk2_0,esk3_0)),X2))|r2_hidden(esk2_0,X1)|v1_xboole_0(X1)|k2_yellow_0(esk1_0,X1)!=esk2_0|~r2_yellow_0(esk1_0,k1_tarski(esk5_3(esk1_0,esk2_0,esk3_0)))|~r2_yellow_0(esk1_0,k2_xboole_0(esk6_3(esk1_0,esk2_0,X1),k1_tarski(esk5_3(esk1_0,esk2_0,X1))))|~r2_yellow_0(esk1_0,X2)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9722, plain, (X1=k12_lattice3(X2,X3,k2_yellow_0(X2,X4))|X1=k2_yellow_0(X2,X5)|r2_hidden(X3,X6)|v1_xboole_0(X6)|X1!=k2_yellow_0(X2,k2_xboole_0(X5,k2_xboole_0(X4,k2_xboole_0(esk6_3(X2,X3,X6),k1_tarski(esk5_3(X2,X3,X6))))))|X3!=k2_yellow_0(X2,X6)|~r2_yellow_0(X2,k2_xboole_0(esk6_3(X2,X3,X6),k1_tarski(esk5_3(X2,X3,X6))))|~r2_yellow_0(X2,X5)|~r2_yellow_0(X2,X4)|~v1_finset_1(X6)|~v2_waybel_6(X1,X2)|~m1_subset_1(k12_lattice3(X2,X3,k2_yellow_0(X2,X4)),u1_struct_0(X2))|~m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_9724, plain, (X1=k12_lattice3(X2,X3,k2_yellow_0(X2,k2_xboole_0(esk6_3(X2,X3,X4),k1_tarski(esk5_3(X2,X3,X4)))))|X1=k2_yellow_0(X2,X5)|r2_hidden(X3,X4)|v1_xboole_0(X4)|X1!=k2_yellow_0(X2,k2_xboole_0(X5,k2_xboole_0(esk6_3(X2,X3,X4),k1_tarski(esk5_3(X2,X3,X4)))))|X3!=k2_yellow_0(X2,X4)|~r2_yellow_0(X2,k2_xboole_0(esk6_3(X2,X3,X4),k1_tarski(esk5_3(X2,X3,X4))))|~r2_yellow_0(X2,X5)|~v1_finset_1(X4)|~v2_waybel_6(X1,X2)|~m1_subset_1(k12_lattice3(X2,X3,k2_yellow_0(X2,k2_xboole_0(esk6_3(X2,X3,X4),k1_tarski(esk5_3(X2,X3,X4))))),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_9705, negated_conjecture, (k2_yellow_0(esk1_0,k2_xboole_0(X1,k2_xboole_0(k1_tarski(esk5_3(esk1_0,esk2_0,esk3_0)),k2_xboole_0(esk6_3(esk1_0,esk2_0,X2),k1_tarski(esk5_3(esk1_0,esk2_0,X2))))))=k2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(esk5_3(esk1_0,esk2_0,esk3_0)),X1))|r2_hidden(esk2_0,X2)|v1_xboole_0(X2)|k2_yellow_0(esk1_0,X2)!=esk2_0|~r2_yellow_0(esk1_0,k1_tarski(esk5_3(esk1_0,esk2_0,esk3_0)))|~r2_yellow_0(esk1_0,k2_xboole_0(esk6_3(esk1_0,esk2_0,X2),k1_tarski(esk5_3(esk1_0,esk2_0,X2))))|~r2_yellow_0(esk1_0,X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9681, negated_conjecture, (k2_yellow_0(esk1_0,k2_xboole_0(k2_xboole_0(k2_xboole_0(esk6_3(esk1_0,esk2_0,X1),k1_tarski(esk5_3(esk1_0,esk2_0,X1))),k1_tarski(esk5_3(esk1_0,esk2_0,esk3_0))),X2))=k2_yellow_0(esk1_0,k2_xboole_0(k2_xboole_0(esk6_3(esk1_0,esk2_0,X1),k1_tarski(esk5_3(esk1_0,esk2_0,X1))),X2))|r2_hidden(esk2_0,X1)|v1_xboole_0(X1)|k2_yellow_0(esk1_0,X1)!=esk2_0|~r2_yellow_0(esk1_0,k2_xboole_0(esk6_3(esk1_0,esk2_0,X1),k1_tarski(esk5_3(esk1_0,esk2_0,X1))))|~r2_yellow_0(esk1_0,k1_tarski(esk5_3(esk1_0,esk2_0,esk3_0)))|~r2_yellow_0(esk1_0,X2)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_8499, plain, (v1_finset_1(esk16_1(esk16_1(X1)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_588, plain, (k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X2,X3),k1_tarski(X4)))=k12_lattice3(X1,k12_lattice3(X1,X2,X3),X4)|v2_struct_0(X1)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,k1_tarski(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_2012, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_9679, negated_conjecture, (k2_yellow_0(esk1_0,k2_xboole_0(k2_xboole_0(X1,k2_xboole_0(esk6_3(esk1_0,esk2_0,X2),k1_tarski(esk5_3(esk1_0,esk2_0,X2)))),X3))=k2_yellow_0(esk1_0,k2_xboole_0(k2_xboole_0(X1,k1_tarski(esk5_3(esk1_0,esk2_0,esk3_0))),X3))|r2_hidden(esk2_0,X2)|v1_xboole_0(X2)|k2_yellow_0(esk1_0,X2)!=esk2_0|~r2_yellow_0(esk1_0,k1_tarski(esk5_3(esk1_0,esk2_0,esk3_0)))|~r2_yellow_0(esk1_0,k2_xboole_0(esk6_3(esk1_0,esk2_0,X2),k1_tarski(esk5_3(esk1_0,esk2_0,X2))))|~r2_yellow_0(esk1_0,X3)|~r2_yellow_0(esk1_0,X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9689, negated_conjecture, (X1=k2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(esk5_3(esk1_0,esk2_0,esk3_0)),k2_xboole_0(esk6_3(esk1_0,esk2_0,X2),k1_tarski(esk5_3(esk1_0,esk2_0,X2)))))|X1=k2_yellow_0(esk1_0,X3)|r2_hidden(esk2_0,X2)|v1_xboole_0(X2)|X1!=k2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(esk5_3(esk1_0,esk2_0,esk3_0)),X3))|k2_yellow_0(esk1_0,X2)!=esk2_0|~r2_yellow_0(esk1_0,k1_tarski(esk5_3(esk1_0,esk2_0,esk3_0)))|~r2_yellow_0(esk1_0,k2_xboole_0(esk6_3(esk1_0,esk2_0,X2),k1_tarski(esk5_3(esk1_0,esk2_0,X2))))|~r2_yellow_0(esk1_0,X3)|~v1_finset_1(X2)|~v2_waybel_6(X1,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6210, plain, (k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(k2_xboole_0(esk6_3(X1,X2,X5),k1_tarski(esk5_3(X1,X2,X5))),k1_tarski(X3)),X4))=k12_lattice3(X1,k12_lattice3(X1,X2,X3),k2_yellow_0(X1,X4))|r2_hidden(X2,X5)|v1_xboole_0(X5)|v2_struct_0(X1)|X2!=k2_yellow_0(X1,X5)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X2,X5),k1_tarski(esk5_3(X1,X2,X5))))|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,k1_tarski(X3))|~v1_finset_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1706, plain, (k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),esk7_2(X1,X2),esk8_2(X1,X2)))=k12_lattice3(X1,X2,X2)|v2_waybel_6(X2,X1)|v1_xboole_0(k7_domain_1(u1_struct_0(X1),esk7_2(X1,X2),esk8_2(X1,X2)))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk7_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(k7_domain_1(u1_struct_0(X1),esk7_2(X1,X2),esk8_2(X1,X2)))|~m1_subset_1(k7_domain_1(u1_struct_0(X1),esk7_2(X1,X2),esk8_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_5810, plain, (k12_lattice3(X1,k12_lattice3(X1,k12_lattice3(X1,X3,k2_yellow_0(X1,X2)),k2_yellow_0(X1,X5)),k2_yellow_0(X1,X6))=k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(k2_xboole_0(X2,k2_xboole_0(esk6_3(X1,X3,X4),k1_tarski(esk5_3(X1,X3,X4)))),X5),X6))|r2_hidden(X3,X4)|v1_xboole_0(X4)|~r2_yellow_0(X1,X6)|~r2_yellow_0(X1,k2_xboole_0(k2_xboole_0(X2,k2_xboole_0(esk6_3(X1,X3,X4),k1_tarski(esk5_3(X1,X3,X4)))),X5))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|X3!=k2_yellow_0(X1,X4)|~r2_yellow_0(X1,k2_xboole_0(esk6_3(X1,X3,X4),k1_tarski(esk5_3(X1,X3,X4))))|~r2_yellow_0(X1,X5)|~r2_yellow_0(X1,X2)|~v1_finset_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 85
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 173
# Removed in clause preprocessing      : 9
# Initial clauses in saturation        : 164
# Processed clauses                    : 1210
# ...of these trivial                  : 20
# ...subsumed                          : 507
# ...remaining for further processing  : 683
# Other redundant clauses eliminated   : 53
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 23
# Backward-rewritten                   : 109
# Generated clauses                    : 10675
# ...of the previous two non-trivial   : 10152
# Contextual simplify-reflections      : 62
# Paramodulations                      : 10470
# Factorizations                       : 120
# Equation resolutions                 : 85
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 548
#    Positive orientable unit clauses  : 58
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 16
#    Non-unit-clauses                  : 472
# Current number of unprocessed clauses: 8955
# ...number of literals in the above   : 130450
# Current number of archived formulas  : 0
# Current number of archived clauses   : 815
# Proof object given clauses           : 116
# Proof search given clauses           : 683
# Clause-clause subsumption calls (NU) : 346945
# Rec. Clause-clause subsumption calls : 17886
# Non-unit clause-clause subsumptions  : 527
# Unit Clause-clause subsumption calls : 10988
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 58
# BW rewrite match successes           : 22
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 535137

# -------------------------------------------------
# User time                : 37.125 s
# System time              : 1.125 s
# Total time               : 38.250 s
# Maximum resident set size: 572500 pages
37.12user 1.24system 0:38.75elapsed 99%CPU (0avgtext+0avgdata 572500maxresident)k
0inputs+464outputs (0major+403614minor)pagefaults 0swaps
