# Preprocessing time       : 0.028 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(dt_l2_struct_0, axiom, ![X1]:(l2_struct_0(X1)=>l1_struct_0(X1)), file('/home/mptp/big1/gnn_pred_train_big_knn512_all.probs/preds__1/t13_matrix_9', dt_l2_struct_0)).
fof(dt_l2_algstr_0, axiom, ![X1]:(l2_algstr_0(X1)=>(l2_struct_0(X1)&l1_algstr_0(X1))), file('/home/mptp/big1/gnn_pred_train_big_knn512_all.probs/preds__1/t13_matrix_9', dt_l2_algstr_0)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('/home/mptp/big1/gnn_pred_train_big_knn512_all.probs/preds__1/t13_matrix_9', fc2_struct_0)).
fof(t12_matrix_7, axiom, ![X1]:((~(v2_struct_0(X1))&(~(v6_struct_0(X1))&(v13_algstr_0(X1)&(v33_algstr_0(X1)&(v3_group_1(X1)&(v5_group_1(X1)&(v2_rlvect_1(X1)&(v3_rlvect_1(X1)&(v4_rlvect_1(X1)&(v4_vectsp_1(X1)&(v5_vectsp_1(X1)&l6_algstr_0(X1))))))))))))=>![X2]:(m1_matrix_1(X2,u1_struct_0(X1),np__2,np__2)=>k12_matrix_3(np__2,X1,X2)=k5_algstr_0(X1,k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),X2,np__1,np__1),k3_matrix_1(u1_struct_0(X1),X2,np__2,np__2)),k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),X2,np__1,np__2),k3_matrix_1(u1_struct_0(X1),X2,np__2,np__1))))), file('/home/mptp/big1/gnn_pred_train_big_knn512_all.probs/preds__1/t13_matrix_9', t12_matrix_7)).
fof(t6_matrix_2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,X1)=>![X3]:(m1_subset_1(X3,X1)=>![X4]:(m1_subset_1(X4,X1)=>![X5]:(m1_subset_1(X5,X1)=>(k3_matrix_1(X1,k6_matrix_2(X1,X2,X3,X4,X5),np__1,np__1)=X2&(k3_matrix_1(X1,k6_matrix_2(X1,X2,X3,X4,X5),np__1,np__2)=X3&(k3_matrix_1(X1,k6_matrix_2(X1,X2,X3,X4,X5),np__2,np__1)=X4&k3_matrix_1(X1,k6_matrix_2(X1,X2,X3,X4,X5),np__2,np__2)=X5)))))))), file('/home/mptp/big1/gnn_pred_train_big_knn512_all.probs/preds__1/t13_matrix_9', t6_matrix_2)).
fof(dt_l6_algstr_0, axiom, ![X1]:(l6_algstr_0(X1)=>(l2_algstr_0(X1)&l5_algstr_0(X1))), file('/home/mptp/big1/gnn_pred_train_big_knn512_all.probs/preds__1/t13_matrix_9', dt_l6_algstr_0)).
fof(dt_k6_matrix_2, axiom, ![X1, X2, X3, X4, X5]:((~(v1_xboole_0(X1))&(m1_subset_1(X2,X1)&(m1_subset_1(X3,X1)&(m1_subset_1(X4,X1)&m1_subset_1(X5,X1)))))=>m1_matrix_1(k6_matrix_2(X1,X2,X3,X4,X5),X1,np__2,np__2)), file('/home/mptp/big1/gnn_pred_train_big_knn512_all.probs/preds__1/t13_matrix_9', dt_k6_matrix_2)).
fof(t13_matrix_9, conjecture, ![X1]:((~(v2_struct_0(X1))&(~(v6_struct_0(X1))&(v13_algstr_0(X1)&(v33_algstr_0(X1)&(v3_group_1(X1)&(v5_group_1(X1)&(v2_rlvect_1(X1)&(v3_rlvect_1(X1)&(v4_rlvect_1(X1)&(v4_vectsp_1(X1)&(v5_vectsp_1(X1)&l6_algstr_0(X1))))))))))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5))=k5_algstr_0(X1,k8_group_1(X1,X2,X5),k8_group_1(X1,X3,X4))))))), file('/home/mptp/big1/gnn_pred_train_big_knn512_all.probs/preds__1/t13_matrix_9', t13_matrix_9)).
fof(c_0_8, plain, ![X38]:(~l2_struct_0(X38)|l1_struct_0(X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l2_struct_0])])).
fof(c_0_9, plain, ![X16]:((l2_struct_0(X16)|~l2_algstr_0(X16))&(l1_algstr_0(X16)|~l2_algstr_0(X16))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l2_algstr_0])])])).
fof(c_0_10, plain, ![X15]:(v2_struct_0(X15)|~l1_struct_0(X15)|~v1_xboole_0(u1_struct_0(X15))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_11, plain, (l1_struct_0(X1)|~l2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_8])).
cnf(c_0_12, plain, (l2_struct_0(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_9])).
fof(c_0_13, plain, ![X71, X72]:(v2_struct_0(X71)|(v6_struct_0(X71)|(~v13_algstr_0(X71)|(~v33_algstr_0(X71)|(~v3_group_1(X71)|(~v5_group_1(X71)|(~v2_rlvect_1(X71)|(~v3_rlvect_1(X71)|(~v4_rlvect_1(X71)|(~v4_vectsp_1(X71)|(~v5_vectsp_1(X71)|~l6_algstr_0(X71)))))))))))|(~m1_matrix_1(X72,u1_struct_0(X71),np__2,np__2)|k12_matrix_3(np__2,X71,X72)=k5_algstr_0(X71,k8_group_1(X71,k3_matrix_1(u1_struct_0(X71),X72,np__1,np__1),k3_matrix_1(u1_struct_0(X71),X72,np__2,np__2)),k8_group_1(X71,k3_matrix_1(u1_struct_0(X71),X72,np__1,np__2),k3_matrix_1(u1_struct_0(X71),X72,np__2,np__1))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t12_matrix_7])])])])).
fof(c_0_14, plain, ![X58, X59, X60, X61, X62]:((k3_matrix_1(X58,k6_matrix_2(X58,X59,X60,X61,X62),np__1,np__1)=X59|~m1_subset_1(X62,X58)|~m1_subset_1(X61,X58)|~m1_subset_1(X60,X58)|~m1_subset_1(X59,X58)|v1_xboole_0(X58))&((k3_matrix_1(X58,k6_matrix_2(X58,X59,X60,X61,X62),np__1,np__2)=X60|~m1_subset_1(X62,X58)|~m1_subset_1(X61,X58)|~m1_subset_1(X60,X58)|~m1_subset_1(X59,X58)|v1_xboole_0(X58))&((k3_matrix_1(X58,k6_matrix_2(X58,X59,X60,X61,X62),np__2,np__1)=X61|~m1_subset_1(X62,X58)|~m1_subset_1(X61,X58)|~m1_subset_1(X60,X58)|~m1_subset_1(X59,X58)|v1_xboole_0(X58))&(k3_matrix_1(X58,k6_matrix_2(X58,X59,X60,X61,X62),np__2,np__2)=X62|~m1_subset_1(X62,X58)|~m1_subset_1(X61,X58)|~m1_subset_1(X60,X58)|~m1_subset_1(X59,X58)|v1_xboole_0(X58))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t6_matrix_2])])])])])).
cnf(c_0_15, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_10])).
cnf(c_0_16, plain, (l1_struct_0(X1)|~l2_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_11])).
cnf(c_0_17, plain, (l2_struct_0(X1)|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_12])).
fof(c_0_18, plain, ![X13]:((l2_algstr_0(X13)|~l6_algstr_0(X13))&(l5_algstr_0(X13)|~l6_algstr_0(X13))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l6_algstr_0])])])).
cnf(c_0_19, plain, (v2_struct_0(X1)|v6_struct_0(X1)|k12_matrix_3(np__2,X1,X2)=k5_algstr_0(X1,k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),X2,np__1,np__1),k3_matrix_1(u1_struct_0(X1),X2,np__2,np__2)),k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),X2,np__1,np__2),k3_matrix_1(u1_struct_0(X1),X2,np__2,np__1)))|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~m1_matrix_1(X2,u1_struct_0(X1),np__2,np__2)), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_20, plain, (k3_matrix_1(X1,k6_matrix_2(X1,X2,X3,X4,X5),np__2,np__1)=X4|v1_xboole_0(X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_21, plain, ![X22, X23, X24, X25, X26]:(v1_xboole_0(X22)|(~m1_subset_1(X23,X22)|(~m1_subset_1(X24,X22)|(~m1_subset_1(X25,X22)|~m1_subset_1(X26,X22))))|m1_matrix_1(k6_matrix_2(X22,X23,X24,X25,X26),X22,np__2,np__2)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_matrix_2])])])).
cnf(c_0_22, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_15])).
cnf(c_0_23, plain, (l1_struct_0(X1)|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_16, c_0_17])])).
cnf(c_0_24, plain, (l2_algstr_0(X1)|~l6_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_18])).
cnf(c_0_25, plain, (k5_algstr_0(X1,k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),X2,np__1,np__1),k3_matrix_1(u1_struct_0(X1),X2,np__2,np__2)),k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),X2,np__1,np__2),k3_matrix_1(u1_struct_0(X1),X2,np__2,np__1)))=k12_matrix_3(np__2,X1,X2)|v2_struct_0(X1)|v6_struct_0(X1)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~m1_matrix_1(X2,u1_struct_0(X1),np__2,np__2)), inference(evalgc,[status(thm)],[c_0_19])).
cnf(c_0_26, plain, (k3_matrix_1(X1,k6_matrix_2(X1,X2,X3,X4,X5),np__2,np__1)=X4|v1_xboole_0(X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_20])).
cnf(c_0_27, plain, (v1_xboole_0(X1)|m1_matrix_1(k6_matrix_2(X1,X2,X3,X4,X5),X1,np__2,np__2)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X5,X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_28, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23])])).
cnf(c_0_29, plain, (l2_algstr_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_24])).
cnf(c_0_30, plain, (k5_algstr_0(X1,k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__1,np__1),k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__2,np__2)),k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__1,np__2),X4))=k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5))|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_matrix_1(k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),u1_struct_0(X1),np__2,np__2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_25, c_0_26])])).
cnf(c_0_31, plain, (v1_xboole_0(X1)|m1_matrix_1(k6_matrix_2(X1,X2,X3,X4,X5),X1,np__2,np__2)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_27])).
cnf(c_0_32, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_29])])).
cnf(c_0_33, plain, (k3_matrix_1(X1,k6_matrix_2(X1,X2,X3,X4,X5),np__1,np__1)=X2|v1_xboole_0(X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_34, plain, (k5_algstr_0(X1,k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__1,np__1),k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__2,np__2)),k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__1,np__2),X4))=k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_30, c_0_31]), c_0_32])).
cnf(c_0_35, plain, (k3_matrix_1(X1,k6_matrix_2(X1,X2,X3,X4,X5),np__1,np__1)=X2|v1_xboole_0(X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_33])).
cnf(c_0_36, plain, (k5_algstr_0(X1,k8_group_1(X1,X2,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__2,np__2)),k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__1,np__2),X4))=k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5))|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35])])).
cnf(c_0_37, plain, (k3_matrix_1(X1,k6_matrix_2(X1,X2,X3,X4,X5),np__1,np__2)=X3|v1_xboole_0(X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_38, plain, (k5_algstr_0(X1,k8_group_1(X1,X2,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__2,np__2)),k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__1,np__2),X4))=k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_36, c_0_32])).
cnf(c_0_39, plain, (k3_matrix_1(X1,k6_matrix_2(X1,X2,X3,X4,X5),np__1,np__2)=X3|v1_xboole_0(X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_37])).
fof(c_0_40, negated_conjecture, ~(![X1]:((~(v2_struct_0(X1))&(~(v6_struct_0(X1))&(v13_algstr_0(X1)&(v33_algstr_0(X1)&(v3_group_1(X1)&(v5_group_1(X1)&(v2_rlvect_1(X1)&(v3_rlvect_1(X1)&(v4_rlvect_1(X1)&(v4_vectsp_1(X1)&(v5_vectsp_1(X1)&l6_algstr_0(X1))))))))))))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5))=k5_algstr_0(X1,k8_group_1(X1,X2,X5),k8_group_1(X1,X3,X4)))))))), inference(assume_negation,[status(cth)],[t13_matrix_9])).
cnf(c_0_41, plain, (k5_algstr_0(X1,k8_group_1(X1,X2,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__2,np__2)),k8_group_1(X1,X3,X4))=k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5))|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_38, c_0_39])])).
cnf(c_0_42, plain, (k3_matrix_1(X1,k6_matrix_2(X1,X2,X3,X4,X5),np__2,np__2)=X5|v1_xboole_0(X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_14])).
fof(c_0_43, negated_conjecture, ((~v2_struct_0(esk1_0)&(~v6_struct_0(esk1_0)&(v13_algstr_0(esk1_0)&(v33_algstr_0(esk1_0)&(v3_group_1(esk1_0)&(v5_group_1(esk1_0)&(v2_rlvect_1(esk1_0)&(v3_rlvect_1(esk1_0)&(v4_rlvect_1(esk1_0)&(v4_vectsp_1(esk1_0)&(v5_vectsp_1(esk1_0)&l6_algstr_0(esk1_0))))))))))))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(m1_subset_1(esk4_0,u1_struct_0(esk1_0))&(m1_subset_1(esk5_0,u1_struct_0(esk1_0))&k12_matrix_3(np__2,esk1_0,k6_matrix_2(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0,esk5_0))!=k5_algstr_0(esk1_0,k8_group_1(esk1_0,esk2_0,esk5_0),k8_group_1(esk1_0,esk3_0,esk4_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_40])])])])).
cnf(c_0_44, plain, (k5_algstr_0(X1,k8_group_1(X1,X2,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__2,np__2)),k8_group_1(X1,X3,X4))=k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_41, c_0_32])).
cnf(c_0_45, plain, (k3_matrix_1(X1,k6_matrix_2(X1,X2,X3,X4,X5),np__2,np__2)=X5|v1_xboole_0(X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_42])).
cnf(c_0_46, negated_conjecture, (k12_matrix_3(np__2,esk1_0,k6_matrix_2(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0,esk5_0))!=k5_algstr_0(esk1_0,k8_group_1(esk1_0,esk2_0,esk5_0),k8_group_1(esk1_0,esk3_0,esk4_0))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_47, plain, (k5_algstr_0(X1,k8_group_1(X1,X2,X5),k8_group_1(X1,X3,X4))=k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5))|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_44, c_0_45])])).
cnf(c_0_48, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_49, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_50, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_51, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_52, negated_conjecture, (l6_algstr_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_53, negated_conjecture, (v5_vectsp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_54, negated_conjecture, (v4_vectsp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_55, negated_conjecture, (v4_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_56, negated_conjecture, (v3_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_57, negated_conjecture, (v2_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_58, negated_conjecture, (v5_group_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_59, negated_conjecture, (v3_group_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_60, negated_conjecture, (v33_algstr_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_61, negated_conjecture, (v13_algstr_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_62, negated_conjecture, (~v6_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_63, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_64, negated_conjecture, (k12_matrix_3(np__2,esk1_0,k6_matrix_2(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0,esk5_0))!=k5_algstr_0(esk1_0,k8_group_1(esk1_0,esk2_0,esk5_0),k8_group_1(esk1_0,esk3_0,esk4_0))), inference(evalgc,[status(thm)],[c_0_46])).
cnf(c_0_65, plain, (k5_algstr_0(X1,k8_group_1(X1,X2,X3),k8_group_1(X1,X4,X5))=k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X4,X5,X3))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)), inference(csr,[status(thm)],[c_0_47, c_0_32])).
cnf(c_0_66, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_48])).
cnf(c_0_67, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_49])).
cnf(c_0_68, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_50])).
cnf(c_0_69, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_51])).
cnf(c_0_70, negated_conjecture, (l6_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_52])).
cnf(c_0_71, negated_conjecture, (v5_vectsp_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_53])).
cnf(c_0_72, negated_conjecture, (v4_vectsp_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_54])).
cnf(c_0_73, negated_conjecture, (v4_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_55])).
cnf(c_0_74, negated_conjecture, (v3_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_56])).
cnf(c_0_75, negated_conjecture, (v2_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_57])).
cnf(c_0_76, negated_conjecture, (v5_group_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_58])).
cnf(c_0_77, negated_conjecture, (v3_group_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_59])).
cnf(c_0_78, negated_conjecture, (v33_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_60])).
cnf(c_0_79, negated_conjecture, (v13_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_61])).
cnf(c_0_80, negated_conjecture, (~v6_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_62])).
cnf(c_0_81, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_63])).
cnf(c_0_82, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_65]), c_0_66]), c_0_67]), c_0_68]), c_0_69]), c_0_70]), c_0_71]), c_0_72]), c_0_73]), c_0_74]), c_0_75]), c_0_76]), c_0_77]), c_0_78]), c_0_79])]), c_0_80]), c_0_81]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 83
# Proof object clause steps            : 66
# Proof object formula steps           : 17
# Proof object conjectures             : 38
# Proof object clause conjectures      : 35
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 27
# Proof object initial formulas used   : 8
# Proof object generating inferences   : 8
# Proof object simplifying inferences  : 22
# Training examples: 34 positive, 146 negative
# Training: Positive examples begin
cnf(c_0_81, negated_conjecture, (~v2_struct_0(esk1_0))).# trainpos
cnf(c_0_80, negated_conjecture, (~v6_struct_0(esk1_0))).# trainpos
cnf(c_0_64, negated_conjecture, (k12_matrix_3(np__2,esk1_0,k6_matrix_2(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0,esk5_0))!=k5_algstr_0(esk1_0,k8_group_1(esk1_0,esk2_0,esk5_0),k8_group_1(esk1_0,esk3_0,esk4_0)))).# trainpos
cnf(c_0_79, negated_conjecture, (v13_algstr_0(esk1_0))).# trainpos
cnf(c_0_78, negated_conjecture, (v33_algstr_0(esk1_0))).# trainpos
cnf(c_0_77, negated_conjecture, (v3_group_1(esk1_0))).# trainpos
cnf(c_0_76, negated_conjecture, (v5_group_1(esk1_0))).# trainpos
cnf(c_0_75, negated_conjecture, (v2_rlvect_1(esk1_0))).# trainpos
cnf(c_0_74, negated_conjecture, (v3_rlvect_1(esk1_0))).# trainpos
cnf(c_0_73, negated_conjecture, (v4_rlvect_1(esk1_0))).# trainpos
cnf(c_0_72, negated_conjecture, (v4_vectsp_1(esk1_0))).# trainpos
cnf(c_0_71, negated_conjecture, (v5_vectsp_1(esk1_0))).# trainpos
cnf(c_0_70, negated_conjecture, (l6_algstr_0(esk1_0))).# trainpos
cnf(c_0_29, plain, (l2_algstr_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_69, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_68, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_22, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).# trainpos
cnf(c_0_67, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_66, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_16, plain, (l1_struct_0(X1)|~l2_struct_0(X1))).# trainpos
cnf(c_0_17, plain, (l2_struct_0(X1)|~l2_algstr_0(X1))).# trainpos
cnf(c_0_31, plain, (v1_xboole_0(X1)|m1_matrix_1(k6_matrix_2(X1,X2,X3,X4,X5),X1,np__2,np__2)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_45, plain, (k3_matrix_1(X1,k6_matrix_2(X1,X2,X3,X4,X5),np__2,np__2)=X5|v1_xboole_0(X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_26, plain, (k3_matrix_1(X1,k6_matrix_2(X1,X2,X3,X4,X5),np__2,np__1)=X4|v1_xboole_0(X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_39, plain, (k3_matrix_1(X1,k6_matrix_2(X1,X2,X3,X4,X5),np__1,np__2)=X3|v1_xboole_0(X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_35, plain, (k3_matrix_1(X1,k6_matrix_2(X1,X2,X3,X4,X5),np__1,np__1)=X2|v1_xboole_0(X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_25, plain, (k5_algstr_0(X1,k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),X2,np__1,np__1),k3_matrix_1(u1_struct_0(X1),X2,np__2,np__2)),k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),X2,np__1,np__2),k3_matrix_1(u1_struct_0(X1),X2,np__2,np__1)))=k12_matrix_3(np__2,X1,X2)|v2_struct_0(X1)|v6_struct_0(X1)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~m1_matrix_1(X2,u1_struct_0(X1),np__2,np__2))).# trainpos
cnf(c_0_23, plain, (l1_struct_0(X1)|~l2_algstr_0(X1))).# trainpos
cnf(c_0_28, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1))).# trainpos
cnf(c_0_32, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l6_algstr_0(X1))).# trainpos
cnf(c_0_30, plain, (k5_algstr_0(X1,k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__1,np__1),k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__2,np__2)),k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__1,np__2),X4))=k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5))|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_matrix_1(k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),u1_struct_0(X1),np__2,np__2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_36, plain, (k5_algstr_0(X1,k8_group_1(X1,X2,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__2,np__2)),k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__1,np__2),X4))=k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5))|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_41, plain, (k5_algstr_0(X1,k8_group_1(X1,X2,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__2,np__2)),k8_group_1(X1,X3,X4))=k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5))|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).# trainpos
cnf(c_0_47, plain, (k5_algstr_0(X1,k8_group_1(X1,X2,X5),k8_group_1(X1,X3,X4))=k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5))|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_73, plain, (k4_ordinal1=k5_numbers)).#trainneg
cnf(i_0_24, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_87, plain, (v3_group_1(X1)|k6_algstr_0(X1,k6_algstr_0(X1,esk7_1(X1),esk8_1(X1)),esk9_1(X1))!=k6_algstr_0(X1,esk7_1(X1),k6_algstr_0(X1,esk8_1(X1),esk9_1(X1)))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_25, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_49, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_50, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_77, plain, (v1_xboole_0(np__0))).#trainneg
cnf(i_0_20, plain, (l5_algstr_0(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_66, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_22, plain, (m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_23, plain, (m1_subset_1(np__2,k5_numbers))).#trainneg
cnf(i_0_47, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_48, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_19, plain, (l3_struct_0(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_26, plain, (l4_struct_0(X1)|~l5_algstr_0(X1))).#trainneg
cnf(i_0_27, plain, (l4_algstr_0(X1)|~l5_algstr_0(X1))).#trainneg
cnf(i_0_41, plain, (l3_struct_0(X1)|~l4_struct_0(X1))).#trainneg
cnf(i_0_74, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_62, plain, (v2_struct_0(X1)|v1_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_46, plain, (l1_struct_0(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_18, plain, (l3_algstr_0(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_79, plain, (l2_algstr_0(X1)|~l1_rlvect_1(X1))).#trainneg
cnf(i_0_45, plain, (l1_struct_0(X1)|~l1_algstr_0(X1))).#trainneg
cnf(i_0_42, plain, (l2_struct_0(X1)|~l4_struct_0(X1))).#trainneg
cnf(i_0_29, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_68, plain, (m1_subset_1(esk6_1(X1),X1))).#trainneg
cnf(i_0_61, plain, (v2_struct_0(X1)|v2_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_84, plain, (v2_struct_0(X1)|v3_vectsp_1(X1)|~v4_vectsp_1(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_31, plain, (v1_finseq_1(k10_finseq_1(X1,X2)))).#trainneg
cnf(i_0_83, plain, (v2_struct_0(X1)|v6_vectsp_1(X1)|~v4_vectsp_1(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_53, plain, (m2_finseq_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_60, plain, (v1_relat_1(k10_finseq_1(X1,X2)))).#trainneg
cnf(i_0_54, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_59, plain, (v1_funct_1(k10_finseq_1(X1,X2)))).#trainneg
cnf(i_0_80, plain, (m1_subset_1(k5_struct_0(X1),u1_struct_0(X1))|~l3_struct_0(X1))).#trainneg
cnf(i_0_86, plain, (m1_subset_1(k5_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_65, plain, (v1_xboole_0(X2)|v1_matrix_1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4))).#trainneg
cnf(i_0_90, plain, (v3_group_1(X1)|m1_subset_1(esk7_1(X1),u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_89, plain, (v3_group_1(X1)|m1_subset_1(esk8_1(X1),u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_39, plain, (k8_group_1(X1,X2,X3)=k8_group_1(X1,X3,X2)|v2_struct_0(X1)|~v5_group_1(X1)|~l6_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_88, plain, (v3_group_1(X1)|m1_subset_1(esk9_1(X1),u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_35, plain, (k3_finseq_1(X1)=np__2|X1!=k10_finseq_1(X2,X3)|~v1_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_33, plain, (k1_funct_1(X1,np__2)=X2|X1!=k10_finseq_1(X3,X2)|~v1_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_92, plain, (k3_rlvect_1(X1,X2,X3)=k1_algstr_0(X1,X2,X3)|~v2_rlvect_1(X1)|~l1_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_34, plain, (k1_funct_1(X1,np__1)=X2|X1!=k10_finseq_1(X2,X3)|~v1_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_75, plain, (m1_subset_1(k4_algstr_0(X1,X2),u1_struct_0(X1))|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_44, plain, (v2_struct_0(X1)|m1_subset_1(k8_group_1(X1,X2,X3),u1_struct_0(X1))|~v5_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_32, plain, (X1=k10_finseq_1(X2,X3)|k3_finseq_1(X1)!=np__2|k1_funct_1(X1,np__2)!=X3|k1_funct_1(X1,np__1)!=X2|~v1_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_43, plain, (k8_group_1(X1,X2,X3)=k8_group_1(X1,X3,X2)|v2_struct_0(X1)|~v5_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_69, plain, (v1_matrix_1(k10_finseq_1(k10_finseq_1(X1,X2),k10_finseq_1(X3,X4))))).#trainneg
cnf(i_0_37, plain, (k6_algstr_0(X1,X2,X3)=k8_group_1(X1,X2,X3)|v2_struct_0(X1)|~v5_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_78, plain, (m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~l3_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_64, plain, (v1_xboole_0(X2)|m2_finseq_1(X1,k3_finseq_2(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4))).#trainneg
cnf(i_0_52, plain, (m1_subset_1(k3_matrix_1(X2,X1,X3,X4),X2)|~v1_matrix_1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_finseq_1(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_81, plain, (k1_algstr_0(X1,X2,k4_algstr_0(X1,X3))=k5_algstr_0(X1,X2,X3)|~l2_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_76, plain, (v2_struct_0(X2)|v6_struct_0(X2)|m1_subset_1(k12_matrix_3(X1,X2,X3),u1_struct_0(X2))|~v13_algstr_0(X2)|~v33_algstr_0(X2)|~v3_group_1(X2)|~v5_group_1(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v7_ordinal1(X1)|~m1_matrix_1(X3,u1_struct_0(X2),X1,X1))).#trainneg
cnf(i_0_93, plain, (k5_algstr_0(X1,k8_group_1(X1,X2,X4),k8_group_1(X1,X3,X4))=k8_group_1(X1,k5_algstr_0(X1,X2,X3),X4)|v2_struct_0(X1)|v6_struct_0(X1)|~v13_algstr_0(X1)|~v33_algstr_0(X1)|~v3_group_1(X1)|~v5_group_1(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_91, plain, (k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),X4)=k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4))|~v3_group_1(X1)|~l3_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_94, plain, (X5=k3_matrix_1(X2,X3,X4,X6)|X5!=k1_funct_1(X1,X6)|X1!=k1_funct_1(X3,X4)|~v1_matrix_1(X3)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~m1_subset_1(X5,X2)|~m2_finseq_1(X1,X2)|~m2_finseq_1(X3,k3_finseq_2(X2))|~r2_hidden(k4_tarski(X4,X6),k2_matrix_1(X3)))).#trainneg
cnf(i_0_97, plain, (m2_finseq_1(esk10_5(X1,X2,X3,X4,X5),X1)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X5,X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X2)))).#trainneg
cnf(i_0_38, plain, (k6_matrix_2(X1,X2,X3,X4,X5)=k10_finseq_1(k10_finseq_1(X2,X3),k10_finseq_1(X4,X5))|v1_xboole_0(X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_96, plain, (esk10_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X5,X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X2)))).#trainneg
cnf(i_0_95, plain, (k1_funct_1(esk10_5(X2,X3,X4,X5,X1),X5)=X1|X1!=k3_matrix_1(X2,X3,X4,X5)|~v1_matrix_1(X3)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,X2)|~m2_finseq_1(X3,k3_finseq_2(X2))|~r2_hidden(k4_tarski(X4,X5),k2_matrix_1(X3)))).#trainneg
cnf(i_0_98, plain, (k1_xboole_0=np__0)).#trainneg
cnf(i_0_99, plain, (l4_algstr_0(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_101, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l3_algstr_0(X1))).#trainneg
cnf(i_0_107, plain, (k3_finseq_1(k10_finseq_1(X1,X2))=np__2)).#trainneg
cnf(i_0_121, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|k5_algstr_0(esk1_0,k8_group_1(esk1_0,esk2_0,esk5_0),k8_group_1(esk1_0,esk3_0,esk4_0))!=k12_matrix_3(np__2,esk1_0,k10_finseq_1(k10_finseq_1(esk2_0,esk3_0),k10_finseq_1(esk4_0,esk5_0))))).#trainneg
cnf(i_0_66, plain, (X1=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_100, plain, (l3_struct_0(X1)|~l5_algstr_0(X1))).#trainneg
cnf(i_0_103, plain, (l1_struct_0(X1)|~l4_struct_0(X1))).#trainneg
cnf(i_0_132, plain, (l3_algstr_0(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_135, negated_conjecture, (l3_algstr_0(esk1_0))).#trainneg
cnf(i_0_108, plain, (k1_funct_1(k10_finseq_1(X1,X2),np__2)=X2)).#trainneg
cnf(i_0_109, plain, (k1_funct_1(k10_finseq_1(X1,X2),np__1)=X1)).#trainneg
cnf(i_0_139, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_111, plain, (k10_finseq_1(X2,k1_funct_1(X1,np__2))=X1|k1_funct_1(X1,np__1)!=X2|k3_finseq_1(X1)!=np__2|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_110, plain, (v1_matrix_1(k6_matrix_2(X1,X2,X3,X4,X5))|v1_xboole_0(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_133, plain, (l1_struct_0(X1)|~l5_algstr_0(X1))).#trainneg
cnf(i_0_143, plain, (l1_struct_0(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_116, plain, (m1_subset_1(k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5)),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~v7_ordinal1(np__2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_117, plain, (m1_subset_1(k8_group_1(X1,k5_algstr_0(X1,X2,X4),X3),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k8_group_1(X1,X4,X3),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_114, plain, (m2_finseq_1(k6_matrix_2(X1,X2,X3,X4,X5),k3_finseq_2(X1))|v1_xboole_0(X1)|~v7_ordinal1(np__2)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_115, plain, (m1_subset_1(k3_matrix_1(X1,X2,X3,X4),X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_118, plain, (m1_subset_1(k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4)),u1_struct_0(X1))|~l3_algstr_0(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_group_1(X1))).#trainneg
cnf(i_0_113, negated_conjecture, (k6_algstr_0(esk1_0,X1,X2)=k8_group_1(esk1_0,X1,X2)|~l3_algstr_0(esk1_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_127, plain, (m1_subset_1(k12_matrix_3(np__2,X1,X2),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1)|~m1_subset_1(k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),X2,np__1,np__2),k3_matrix_1(u1_struct_0(X1),X2,np__2,np__1)),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),X2,np__1,np__1),k3_matrix_1(u1_struct_0(X1),X2,np__2,np__2)),u1_struct_0(X1))|~m1_matrix_1(X2,u1_struct_0(X1),np__2,np__2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_119, plain, (k6_algstr_0(X1,k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4)),X5)=k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),k6_algstr_0(X1,X4,X5))|~l3_algstr_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~v3_group_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_122, plain, (m1_matrix_1(k10_finseq_1(k10_finseq_1(X2,X3),k10_finseq_1(X4,X5)),X1,np__2,np__2)|v1_xboole_0(X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_155, negated_conjecture, (m1_subset_1(k8_group_1(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_177, negated_conjecture, (m1_subset_1(k12_matrix_3(np__2,esk1_0,X1),u1_struct_0(esk1_0))|~m1_matrix_1(X1,u1_struct_0(esk1_0),np__2,np__2)|~m1_subset_1(k8_group_1(esk1_0,k3_matrix_1(u1_struct_0(esk1_0),X1,np__1,np__2),k3_matrix_1(u1_struct_0(esk1_0),X1,np__2,np__1)),u1_struct_0(esk1_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),X1,np__2,np__2),u1_struct_0(esk1_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),X1,np__1,np__1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_147, plain, (m1_subset_1(k8_group_1(X1,k8_group_1(X1,k5_algstr_0(X1,X2,X4),X3),X5),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X4,X3),X5),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X2,X3),X5),u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X4,X3),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X2,X3),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_183, plain, (m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k5_algstr_0(X1,X2,X4),X3),X5),X6),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X4,X3),X5),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X2,X3),X5),u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X4,X3),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X2,X3),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_123, plain, (k3_matrix_1(X1,k10_finseq_1(k10_finseq_1(X2,X3),k10_finseq_1(X4,X5)),np__2,np__2)=X5|v1_xboole_0(X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_124, plain, (k3_matrix_1(X1,k10_finseq_1(k10_finseq_1(X2,X3),k10_finseq_1(X4,X5)),np__2,np__1)=X4|v1_xboole_0(X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_176, plain, (m2_finseq_1(k10_finseq_1(k10_finseq_1(X1,X2),k10_finseq_1(X3,X4)),k3_finseq_2(X5))|v1_xboole_0(X5)|~v7_ordinal1(np__2)|~m1_subset_1(X4,X5)|~m1_subset_1(X3,X5)|~m1_subset_1(X2,X5)|~m1_subset_1(X1,X5))).#trainneg
cnf(i_0_180, negated_conjecture, (m1_subset_1(k12_matrix_3(np__2,esk1_0,X1),u1_struct_0(esk1_0))|~m1_matrix_1(X1,u1_struct_0(esk1_0),np__2,np__2)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),X1,np__2,np__2),u1_struct_0(esk1_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),X1,np__1,np__1),u1_struct_0(esk1_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),X1,np__2,np__1),u1_struct_0(esk1_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),X1,np__1,np__2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_178, negated_conjecture, (m1_subset_1(k12_matrix_3(np__2,esk1_0,k6_matrix_2(u1_struct_0(esk1_0),X1,X2,X3,X4)),u1_struct_0(esk1_0))|~m1_matrix_1(k6_matrix_2(u1_struct_0(esk1_0),X1,X2,X3,X4),u1_struct_0(esk1_0),np__2,np__2)|~m1_subset_1(k8_group_1(esk1_0,k3_matrix_1(u1_struct_0(esk1_0),k6_matrix_2(u1_struct_0(esk1_0),X1,X2,X3,X4),np__1,np__2),X3),u1_struct_0(esk1_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),k6_matrix_2(u1_struct_0(esk1_0),X1,X2,X3,X4),np__2,np__2),u1_struct_0(esk1_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),k6_matrix_2(u1_struct_0(esk1_0),X1,X2,X3,X4),np__1,np__1),u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_125, plain, (k3_matrix_1(X1,k10_finseq_1(k10_finseq_1(X2,X3),k10_finseq_1(X4,X5)),np__1,np__2)=X3|v1_xboole_0(X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_149, plain, (m1_subset_1(k3_matrix_1(X1,k6_matrix_2(X1,X2,X3,X4,X5),X6,X7),X1)|v1_xboole_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v1_matrix_1(k6_matrix_2(X1,X2,X3,X4,X5))|~v7_ordinal1(np__2)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_158, negated_conjecture, (m1_subset_1(k6_algstr_0(esk1_0,X1,k8_group_1(esk1_0,X2,X3)),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_185, plain, (m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k5_algstr_0(X1,X2,X4),X3),X5),X6),X7),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X4,X3),X5),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X2,X3),X5),u1_struct_0(X1))|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X4,X3),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X2,X3),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_206, plain, (m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k5_algstr_0(X1,X2,X4),X3),X5),X6),X7),X8),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X4,X3),X5),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X2,X3),X5),u1_struct_0(X1))|~m1_subset_1(X8,u1_struct_0(X1))|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X4,X3),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X2,X3),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_126, plain, (k3_matrix_1(X1,k10_finseq_1(k10_finseq_1(X2,X3),k10_finseq_1(X4,X5)),np__1,np__1)=X2|v1_xboole_0(X1)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_161, plain, (m1_subset_1(k12_matrix_3(np__2,X1,X2),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_matrix_1(X2,u1_struct_0(X1),np__2,np__2)|~m1_subset_1(k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),X2,np__1,np__2),k3_matrix_1(u1_struct_0(X1),X2,np__2,np__1)),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~l3_algstr_0(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(X1),X2,np__2,np__2),u1_struct_0(X1))|~m1_subset_1(k3_matrix_1(u1_struct_0(X1),X2,np__1,np__1),u1_struct_0(X1)))).#trainneg
cnf(i_0_150, plain, (m1_subset_1(k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,k6_algstr_0(X1,X4,X5))),u1_struct_0(X1))|~l3_algstr_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_group_1(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_179, negated_conjecture, (m1_subset_1(k12_matrix_3(np__2,esk1_0,k6_matrix_2(u1_struct_0(esk1_0),X1,X2,X3,X4)),u1_struct_0(esk1_0))|~m1_matrix_1(k6_matrix_2(u1_struct_0(esk1_0),X1,X2,X3,X4),u1_struct_0(esk1_0),np__2,np__2)|~m1_subset_1(k8_group_1(esk1_0,X2,k3_matrix_1(u1_struct_0(esk1_0),k6_matrix_2(u1_struct_0(esk1_0),X1,X2,X3,X4),np__2,np__1)),u1_struct_0(esk1_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),k6_matrix_2(u1_struct_0(esk1_0),X1,X2,X3,X4),np__2,np__2),u1_struct_0(esk1_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),k6_matrix_2(u1_struct_0(esk1_0),X1,X2,X3,X4),np__1,np__1),u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_190, plain, (m1_subset_1(k3_matrix_1(X1,k10_finseq_1(k10_finseq_1(X2,X3),k10_finseq_1(X4,X5)),X6,X7),X1)|v1_xboole_0(X1)|~v7_ordinal1(X7)|~v7_ordinal1(X6)|~v7_ordinal1(np__2)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_208, plain, (m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k5_algstr_0(X1,X2,X4),X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X4,X3),X5),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X2,X3),X5),u1_struct_0(X1))|~m1_subset_1(X9,u1_struct_0(X1))|~m1_subset_1(X8,u1_struct_0(X1))|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X4,X3),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X2,X3),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_254, plain, (m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k5_algstr_0(X1,X2,X4),X3),X5),X6),X7),X8),X9),X10),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X4,X3),X5),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X2,X3),X5),u1_struct_0(X1))|~m1_subset_1(X10,u1_struct_0(X1))|~m1_subset_1(X9,u1_struct_0(X1))|~m1_subset_1(X8,u1_struct_0(X1))|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X4,X3),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X2,X3),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_129, plain, (k5_algstr_0(X1,k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__1,np__1),k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__2,np__2)),k8_group_1(X1,X3,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__2,np__1)))=k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5))|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_matrix_1(k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),u1_struct_0(X1),np__2,np__2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_130, plain, (k5_algstr_0(X1,k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__1,np__1),X5),k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__1,np__2),k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__2,np__1)))=k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5))|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_matrix_1(k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),u1_struct_0(X1),np__2,np__2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_228, plain, (m1_subset_1(k12_matrix_3(np__2,X1,X2),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_matrix_1(X2,u1_struct_0(X1),np__2,np__2)|~m1_subset_1(k3_matrix_1(u1_struct_0(X1),X2,np__2,np__2),u1_struct_0(X1))|~m1_subset_1(k3_matrix_1(u1_struct_0(X1),X2,np__1,np__1),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~l3_algstr_0(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(X1),X2,np__2,np__1),u1_struct_0(X1))|~m1_subset_1(k3_matrix_1(u1_struct_0(X1),X2,np__1,np__2),u1_struct_0(X1)))).#trainneg
cnf(i_0_106, negated_conjecture, (k8_group_1(esk1_0,X1,X2)=k8_group_1(esk1_0,X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_163, plain, (m1_subset_1(k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),k6_algstr_0(X1,X4,X5)),u1_struct_0(X1))|~l3_algstr_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4)),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_group_1(X1))).#trainneg
cnf(i_0_169, plain, (k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),k6_algstr_0(X1,X4,X5))=k6_algstr_0(X1,X2,k6_algstr_0(X1,k6_algstr_0(X1,X3,X4),X5))|~l3_algstr_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X3,X4),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_group_1(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_258, plain, (m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k5_algstr_0(X1,X2,X4),X3),X5),X6),X7),X8),X9),X10),X11),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X4,X3),X5),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X2,X3),X5),u1_struct_0(X1))|~m1_subset_1(X11,u1_struct_0(X1))|~m1_subset_1(X10,u1_struct_0(X1))|~m1_subset_1(X9,u1_struct_0(X1))|~m1_subset_1(X8,u1_struct_0(X1))|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X4,X3),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X2,X3),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_412, plain, (m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k5_algstr_0(X1,X2,X4),X3),X5),X6),X7),X8),X9),X10),X11),X12),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X4,X3),X5),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X2,X3),X5),u1_struct_0(X1))|~m1_subset_1(X12,u1_struct_0(X1))|~m1_subset_1(X11,u1_struct_0(X1))|~m1_subset_1(X10,u1_struct_0(X1))|~m1_subset_1(X9,u1_struct_0(X1))|~m1_subset_1(X8,u1_struct_0(X1))|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X4,X3),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X2,X3),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_131, plain, (k5_algstr_0(X1,k8_group_1(X1,X2,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__2,np__2)),k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__1,np__2),k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__2,np__1)))=k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5))|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_matrix_1(k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),u1_struct_0(X1),np__2,np__2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_145, plain, (m1_subset_1(k12_matrix_3(np__2,X1,k10_finseq_1(k10_finseq_1(X2,X3),k10_finseq_1(X4,X5))),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~v7_ordinal1(np__2)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_400, plain, (m1_subset_1(k6_algstr_0(X1,X2,k6_algstr_0(X1,k6_algstr_0(X1,X3,X4),X5)),u1_struct_0(X1))|~l3_algstr_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_group_1(X1))).#trainneg
cnf(i_0_164, plain, (k6_algstr_0(X1,k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),k6_algstr_0(X1,X4,X5)),X6)=k6_algstr_0(X1,k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4)),k6_algstr_0(X1,X5,X6))|~l3_algstr_0(X1)|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4)),u1_struct_0(X1))|~v3_group_1(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_398, plain, (k6_algstr_0(X1,X2,k6_algstr_0(X1,k6_algstr_0(X1,X3,X4),X5))=k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,k6_algstr_0(X1,X4,X5)))|~l3_algstr_0(X1)|~m1_subset_1(k6_algstr_0(X1,X4,X5),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_group_1(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_434, plain, (m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k5_algstr_0(X1,X2,X4),X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X4,X3),X5),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X2,X3),X5),u1_struct_0(X1))|~m1_subset_1(X13,u1_struct_0(X1))|~m1_subset_1(X12,u1_struct_0(X1))|~m1_subset_1(X11,u1_struct_0(X1))|~m1_subset_1(X10,u1_struct_0(X1))|~m1_subset_1(X9,u1_struct_0(X1))|~m1_subset_1(X8,u1_struct_0(X1))|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X4,X3),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X2,X3),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_561, plain, (m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k5_algstr_0(X1,X2,X4),X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X4,X3),X5),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X2,X3),X5),u1_struct_0(X1))|~m1_subset_1(X14,u1_struct_0(X1))|~m1_subset_1(X13,u1_struct_0(X1))|~m1_subset_1(X12,u1_struct_0(X1))|~m1_subset_1(X11,u1_struct_0(X1))|~m1_subset_1(X10,u1_struct_0(X1))|~m1_subset_1(X9,u1_struct_0(X1))|~m1_subset_1(X8,u1_struct_0(X1))|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X4,X3),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X2,X3),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_146, plain, (m1_subset_1(k8_group_1(X1,k12_matrix_3(np__2,X1,X2),X3),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),X2,np__1,np__2),k3_matrix_1(u1_struct_0(X1),X2,np__2,np__1)),X3),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),X2,np__1,np__1),k3_matrix_1(u1_struct_0(X1),X2,np__2,np__2)),X3),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),X2,np__1,np__2),k3_matrix_1(u1_struct_0(X1),X2,np__2,np__1)),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),X2,np__1,np__1),k3_matrix_1(u1_struct_0(X1),X2,np__2,np__2)),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_matrix_1(X2,u1_struct_0(X1),np__2,np__2))).#trainneg
cnf(i_0_159, plain, (m1_subset_1(k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5)),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_matrix_1(k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),u1_struct_0(X1),np__2,np__2)|~m1_subset_1(k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__1,np__2),k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__2,np__1)),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__1,np__1),X5),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_519, plain, (k6_algstr_0(X1,k6_algstr_0(X1,X2,k6_algstr_0(X1,k6_algstr_0(X1,X3,X4),X5)),X6)=k6_algstr_0(X1,k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4)),k6_algstr_0(X1,X5,X6))|~l3_algstr_0(X1)|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_group_1(X1))).#trainneg
cnf(i_0_588, plain, (m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k5_algstr_0(X1,X2,X4),X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X4,X3),X5),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X2,X3),X5),u1_struct_0(X1))|~m1_subset_1(X15,u1_struct_0(X1))|~m1_subset_1(X14,u1_struct_0(X1))|~m1_subset_1(X13,u1_struct_0(X1))|~m1_subset_1(X12,u1_struct_0(X1))|~m1_subset_1(X11,u1_struct_0(X1))|~m1_subset_1(X10,u1_struct_0(X1))|~m1_subset_1(X9,u1_struct_0(X1))|~m1_subset_1(X8,u1_struct_0(X1))|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X4,X3),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X2,X3),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_684, plain, (m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k5_algstr_0(X1,X2,X4),X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),X16),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),X16),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),X16),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X4,X3),X5),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X2,X3),X5),u1_struct_0(X1))|~m1_subset_1(X16,u1_struct_0(X1))|~m1_subset_1(X15,u1_struct_0(X1))|~m1_subset_1(X14,u1_struct_0(X1))|~m1_subset_1(X13,u1_struct_0(X1))|~m1_subset_1(X12,u1_struct_0(X1))|~m1_subset_1(X11,u1_struct_0(X1))|~m1_subset_1(X10,u1_struct_0(X1))|~m1_subset_1(X9,u1_struct_0(X1))|~m1_subset_1(X8,u1_struct_0(X1))|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X4,X3),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X2,X3),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_160, plain, (m1_subset_1(k12_matrix_3(np__2,X1,k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5)),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_matrix_1(k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),u1_struct_0(X1),np__2,np__2)|~m1_subset_1(k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__1,np__2),k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__2,np__1)),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X2,k3_matrix_1(u1_struct_0(X1),k6_matrix_2(u1_struct_0(X1),X2,X3,X4,X5),np__2,np__2)),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_166, plain, (k6_algstr_0(X1,k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),k6_algstr_0(X1,X4,k6_algstr_0(X1,X5,X6))),X7)=k6_algstr_0(X1,k6_algstr_0(X1,k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4)),X5),k6_algstr_0(X1,X6,X7))|~l3_algstr_0(X1)|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4)),u1_struct_0(X1))|~v3_group_1(X1)|~m1_subset_1(k6_algstr_0(X1,X5,X6),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_620, negated_conjecture, (m1_subset_1(k8_group_1(esk1_0,k12_matrix_3(np__2,esk1_0,X1),X2),u1_struct_0(esk1_0))|~m1_matrix_1(X1,u1_struct_0(esk1_0),np__2,np__2)|~m1_subset_1(k8_group_1(esk1_0,k8_group_1(esk1_0,k3_matrix_1(u1_struct_0(esk1_0),X1,np__1,np__2),k3_matrix_1(u1_struct_0(esk1_0),X1,np__2,np__1)),X2),u1_struct_0(esk1_0))|~m1_subset_1(k8_group_1(esk1_0,k3_matrix_1(u1_struct_0(esk1_0),X1,np__1,np__2),k3_matrix_1(u1_struct_0(esk1_0),X1,np__2,np__1)),u1_struct_0(esk1_0))|~m1_subset_1(k8_group_1(esk1_0,k3_matrix_1(u1_struct_0(esk1_0),X1,np__1,np__1),k3_matrix_1(u1_struct_0(esk1_0),X1,np__2,np__2)),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_167, plain, (m1_subset_1(k6_algstr_0(X1,X2,k6_algstr_0(X1,k6_algstr_0(X1,X3,X4),k6_algstr_0(X1,X5,X6))),u1_struct_0(X1))|~l3_algstr_0(X1)|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X3,k6_algstr_0(X1,X4,X5)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_group_1(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_670, plain, (k6_algstr_0(X1,k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,k6_algstr_0(X1,X4,X5))),X6)=k6_algstr_0(X1,k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4)),k6_algstr_0(X1,X5,X6))|~l3_algstr_0(X1)|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_group_1(X1))).#trainneg
cnf(i_0_822, negated_conjecture, (m1_subset_1(k8_group_1(esk1_0,k12_matrix_3(np__2,esk1_0,X1),X2),u1_struct_0(esk1_0))|~m1_matrix_1(X1,u1_struct_0(esk1_0),np__2,np__2)|~m1_subset_1(k8_group_1(esk1_0,k3_matrix_1(u1_struct_0(esk1_0),X1,np__1,np__2),k3_matrix_1(u1_struct_0(esk1_0),X1,np__2,np__1)),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),X1,np__2,np__2),u1_struct_0(esk1_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),X1,np__1,np__1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_715, plain, (m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k5_algstr_0(X1,X2,X4),X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),X16),X17),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),X16),X17),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),X16),X17),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),X16),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),X16),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X4,X3),X5),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X2,X3),X5),u1_struct_0(X1))|~m1_subset_1(X17,u1_struct_0(X1))|~m1_subset_1(X16,u1_struct_0(X1))|~m1_subset_1(X15,u1_struct_0(X1))|~m1_subset_1(X14,u1_struct_0(X1))|~m1_subset_1(X13,u1_struct_0(X1))|~m1_subset_1(X12,u1_struct_0(X1))|~m1_subset_1(X11,u1_struct_0(X1))|~m1_subset_1(X10,u1_struct_0(X1))|~m1_subset_1(X9,u1_struct_0(X1))|~m1_subset_1(X8,u1_struct_0(X1))|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X4,X3),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X2,X3),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_168, plain, (k6_algstr_0(X1,k6_algstr_0(X1,X2,k6_algstr_0(X1,k6_algstr_0(X1,X3,X4),k6_algstr_0(X1,X5,X6))),X7)=k6_algstr_0(X1,k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,k6_algstr_0(X1,X4,X5))),k6_algstr_0(X1,X6,X7))|~l3_algstr_0(X1)|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X3,k6_algstr_0(X1,X4,X5)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_group_1(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_846, plain, (m1_subset_1(k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,k6_algstr_0(X1,k6_algstr_0(X1,X4,X5),X6))),u1_struct_0(X1))|~l3_algstr_0(X1)|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_group_1(X1))).#trainneg
cnf(i_0_931, negated_conjecture, (m1_subset_1(k8_group_1(esk1_0,k12_matrix_3(np__2,esk1_0,X1),X2),u1_struct_0(esk1_0))|~m1_matrix_1(X1,u1_struct_0(esk1_0),np__2,np__2)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),X1,np__2,np__2),u1_struct_0(esk1_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),X1,np__1,np__1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),X1,np__2,np__1),u1_struct_0(esk1_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),X1,np__1,np__2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_173, plain, (k6_algstr_0(X1,k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,k6_algstr_0(X1,X4,X5))),X6)=k6_algstr_0(X1,k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),X4),k6_algstr_0(X1,X5,X6))|~l3_algstr_0(X1)|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(k6_algstr_0(X1,X2,X3),u1_struct_0(X1))|~v3_group_1(X1)|~m1_subset_1(k6_algstr_0(X1,X4,X5),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_938, plain, (m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k5_algstr_0(X1,X2,X4),X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),X16),X17),X18),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),X16),X17),X18),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),X16),X17),X18),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),X16),X17),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),X16),X17),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),X16),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),X16),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),X15),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),X14),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),X13),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),X12),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),X12),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),X11),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),X11),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),X10),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),X10),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),X9),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),X8),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),X7),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X4,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,k8_group_1(X1,X2,X3),X5),X6),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X4,X3),X5),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k8_group_1(X1,X2,X3),X5),u1_struct_0(X1))|~m1_subset_1(X18,u1_struct_0(X1))|~m1_subset_1(X17,u1_struct_0(X1))|~m1_subset_1(X16,u1_struct_0(X1))|~m1_subset_1(X15,u1_struct_0(X1))|~m1_subset_1(X14,u1_struct_0(X1))|~m1_subset_1(X13,u1_struct_0(X1))|~m1_subset_1(X12,u1_struct_0(X1))|~m1_subset_1(X11,u1_struct_0(X1))|~m1_subset_1(X10,u1_struct_0(X1))|~m1_subset_1(X9,u1_struct_0(X1))|~m1_subset_1(X8,u1_struct_0(X1))|~m1_subset_1(X7,u1_struct_0(X1))|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X4,X3),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,X2,X3),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_186, plain, (k5_algstr_0(X1,k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k10_finseq_1(k10_finseq_1(X2,X3),k10_finseq_1(X4,X5)),np__1,np__1),X5),k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k10_finseq_1(k10_finseq_1(X2,X3),k10_finseq_1(X4,X5)),np__1,np__2),k3_matrix_1(u1_struct_0(X1),k10_finseq_1(k10_finseq_1(X2,X3),k10_finseq_1(X4,X5)),np__2,np__1)))=k12_matrix_3(np__2,X1,k10_finseq_1(k10_finseq_1(X2,X3),k10_finseq_1(X4,X5)))|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_matrix_1(k10_finseq_1(k10_finseq_1(X2,X3),k10_finseq_1(X4,X5)),u1_struct_0(X1),np__2,np__2)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_187, plain, (m1_subset_1(k12_matrix_3(np__2,X1,k10_finseq_1(k10_finseq_1(X2,X3),k10_finseq_1(X4,X5))),u1_struct_0(X1))|v6_struct_0(X1)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_matrix_1(k10_finseq_1(k10_finseq_1(X2,X3),k10_finseq_1(X4,X5)),u1_struct_0(X1),np__2,np__2)|~m1_subset_1(k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k10_finseq_1(k10_finseq_1(X2,X3),k10_finseq_1(X4,X5)),np__1,np__2),k3_matrix_1(u1_struct_0(X1),k10_finseq_1(k10_finseq_1(X2,X3),k10_finseq_1(X4,X5)),np__2,np__1)),u1_struct_0(X1))|~m1_subset_1(k8_group_1(X1,k3_matrix_1(u1_struct_0(X1),k10_finseq_1(k10_finseq_1(X2,X3),k10_finseq_1(X4,X5)),np__1,np__1),X5),u1_struct_0(X1))|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v2_rlvect_1(X1)|~v5_group_1(X1)|~v3_group_1(X1)|~v33_algstr_0(X1)|~v13_algstr_0(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_1068, plain, (m1_subset_1(k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,k6_algstr_0(X1,X4,k6_algstr_0(X1,X5,X6)))),u1_struct_0(X1))|~l3_algstr_0(X1)|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_group_1(X1))).#trainneg
cnf(i_0_1152, plain, (k6_algstr_0(X1,k6_algstr_0(X1,k6_algstr_0(X1,X2,X3),X4),k6_algstr_0(X1,X5,X6))=k6_algstr_0(X1,k6_algstr_0(X1,X2,k6_algstr_0(X1,X3,X4)),k6_algstr_0(X1,X5,X6))|~l3_algstr_0(X1)|~m1_subset_1(X6,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_group_1(X1))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 47
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 97
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 94
# Processed clauses                    : 226
# ...of these trivial                  : 4
# ...subsumed                          : 42
# ...remaining for further processing  : 180
# Other redundant clauses eliminated   : 2
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 0
# Backward-rewritten                   : 1
# Generated clauses                    : 1367
# ...of the previous two non-trivial   : 1320
# Contextual simplify-reflections      : 40
# Paramodulations                      : 1361
# Factorizations                       : 0
# Equation resolutions                 : 6
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 179
#    Positive orientable unit clauses  : 34
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 141
# Current number of unprocessed clauses: 1188
# ...number of literals in the above   : 23480
# Current number of archived formulas  : 0
# Current number of archived clauses   : 182
# Proof object given clauses           : 34
# Proof search given clauses           : 180
# Clause-clause subsumption calls (NU) : 53388
# Rec. Clause-clause subsumption calls : 1259
# Non-unit clause-clause subsumptions  : 82
# Unit Clause-clause subsumption calls : 40
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 1
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 233932

# -------------------------------------------------
# User time                : 0.191 s
# System time              : 0.004 s
# Total time               : 0.195 s
# Maximum resident set size: 7976 pages
