2020-06-08 19:00:27.169511: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_49
2020-06-08 19:00:27.249718: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-06-08 19:00:27.379414: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-06-08 19:00:28.114877: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-06-08 19:00:30.623672: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 3454170 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_49' loaded (query=128; context=512; weigths=real; len_mult=0.000000).
# Preprocessing time       : 4.441 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t17_rsspace, axiom, ![X1]:((v1_funct_1(X1)&(v1_funct_2(X1,k5_numbers,k1_numbers)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))))=>((![X2]:(m2_subset_1(X2,k1_numbers,k5_numbers)=>r1_xxreal_0(k6_numbers,k1_seq_1(X1,X2)))&(v1_series_1(X1)&k4_series_1(X1)=k6_numbers))=>![X2]:(m2_subset_1(X2,k1_numbers,k5_numbers)=>k1_seq_1(X1,X2)=k6_numbers))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', t17_rsspace)).
fof(redefinition_k1_seq_1, axiom, ![X1, X2]:((v1_relat_1(X1)&(v1_funct_1(X1)&v3_valued_0(X1)))=>k1_seq_1(X1,X2)=k1_funct_1(X1,X2)), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', redefinition_k1_seq_1)).
fof(t46_complex1, axiom, ![X1]:(v1_xcmplx_0(X1)=>r1_xxreal_0(k6_numbers,k17_complex1(X1))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', t46_complex1)).
fof(t18_valued_1, axiom, ![X1]:((v1_relat_1(X1)&(v1_funct_1(X1)&v1_valued_0(X1)))=>![X2]:k1_funct_1(k54_valued_1(X1),X2)=k17_complex1(k1_funct_1(X1,X2))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', t18_valued_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', cc1_relset_1)).
fof(fc61_valued_0, axiom, ![X1, X2]:((v1_relat_1(X1)&(v1_funct_1(X1)&v1_valued_0(X1)))=>v1_xcmplx_0(k1_funct_1(X1,X2))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', fc61_valued_0)).
fof(dt_k54_valued_1, axiom, ![X1]:((v1_relat_1(X1)&(v1_funct_1(X1)&v1_valued_0(X1)))=>(v1_relat_1(k54_valued_1(X1))&(v1_funct_1(k54_valued_1(X1))&v3_valued_0(k54_valued_1(X1))))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', dt_k54_valued_1)).
fof(fc7_comseq_3, axiom, ![X1]:((v1_funct_1(X1)&(v1_funct_2(X1,k5_numbers,k2_numbers)&(v2_comseq_3(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers))))))=>(v1_funct_1(k54_valued_1(X1))&(v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)&v1_series_1(k54_valued_1(X1))))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', fc7_comseq_3)).
fof(t13_csspace2, conjecture, ![X1]:((v1_funct_1(X1)&(v1_funct_2(X1,k5_numbers,k2_numbers)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))))=>((v2_comseq_3(X1)&k4_series_1(k55_valued_1(k5_numbers,k2_numbers,X1))=k6_numbers)=>![X2]:(m2_subset_1(X2,k1_numbers,k5_numbers)=>k3_funct_2(k5_numbers,k2_numbers,X1,X2)=k5_complex1))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', t13_csspace2)).
fof(redefinition_k55_valued_1, axiom, ![X1, X2, X3]:((v1_membered(X2)&(v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>k55_valued_1(X1,X2,X3)=k54_valued_1(X3)), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', redefinition_k55_valued_1)).
fof(cc16_valued_0, axiom, ![X1, X2]:(v1_membered(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_valued_0(X3))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', cc16_valued_0)).
fof(dt_k55_valued_1, axiom, ![X1, X2, X3]:((v1_membered(X2)&(v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))=>(v1_funct_1(k55_valued_1(X1,X2,X3))&m1_subset_1(k55_valued_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers))))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', dt_k55_valued_1)).
fof(fc1_membered, axiom, v1_membered(k2_numbers), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', fc1_membered)).
fof(t45_complex1, axiom, ![X1]:(v1_xcmplx_0(X1)=>(k17_complex1(X1)=k6_numbers=>X1=k6_numbers)), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', t45_complex1)).
fof(cc18_valued_0, axiom, ![X1, X2]:(v3_membered(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v3_valued_0(X3))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', cc18_valued_0)).
fof(redefinition_k5_complex1, axiom, k5_complex1=k1_xboole_0, file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', redefinition_k5_complex1)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', redefinition_k6_numbers)).
fof(fc6_ordinal1, axiom, (~(v1_xboole_0(k4_ordinal1))&v3_ordinal1(k4_ordinal1)), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', fc6_ordinal1)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&(~(v1_xboole_0(X2))&m1_subset_1(X2,k1_zfmisc_1(X1))))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', redefinition_m2_subset_1)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', cc1_subset_1)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', fc3_membered)).
fof(redefinition_k3_funct_2, axiom, ![X1, X2, X3, X4]:((~(v1_xboole_0(X1))&((v1_funct_1(X3)&(v1_funct_2(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))&m1_subset_1(X4,X1)))=>k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4)), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', redefinition_k3_funct_2)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', redefinition_k5_numbers)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t13_csspace2', dt_k5_numbers)).
fof(c_0_24, plain, ![X38, X40]:((m2_subset_1(esk4_1(X38),k1_numbers,k5_numbers)|(~v1_series_1(X38)|k4_series_1(X38)!=k6_numbers)|(~m2_subset_1(X40,k1_numbers,k5_numbers)|k1_seq_1(X38,X40)=k6_numbers)|(~v1_funct_1(X38)|(~v1_funct_2(X38,k5_numbers,k1_numbers)|~m1_subset_1(X38,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))&(~r1_xxreal_0(k6_numbers,k1_seq_1(X38,esk4_1(X38)))|(~v1_series_1(X38)|k4_series_1(X38)!=k6_numbers)|(~m2_subset_1(X40,k1_numbers,k5_numbers)|k1_seq_1(X38,X40)=k6_numbers)|(~v1_funct_1(X38)|(~v1_funct_2(X38,k5_numbers,k1_numbers)|~m1_subset_1(X38,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t17_rsspace])])])])])).
fof(c_0_25, plain, ![X41, X42]:(~v1_relat_1(X41)|(~v1_funct_1(X41)|~v3_valued_0(X41))|k1_seq_1(X41,X42)=k1_funct_1(X41,X42)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_seq_1])])).
fof(c_0_26, plain, ![X18]:(~v1_xcmplx_0(X18)|r1_xxreal_0(k6_numbers,k17_complex1(X18))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t46_complex1])])).
fof(c_0_27, plain, ![X22, X23]:(~v1_relat_1(X22)|(~v1_funct_1(X22)|~v1_valued_0(X22))|k1_funct_1(k54_valued_1(X22),X23)=k17_complex1(k1_funct_1(X22,X23))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t18_valued_1])])])).
cnf(c_0_28, plain, (k1_seq_1(X1,X2)=k6_numbers|~r1_xxreal_0(k6_numbers,k1_seq_1(X1,esk4_1(X1)))|~v1_series_1(X1)|k4_series_1(X1)!=k6_numbers|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_24])).
cnf(c_0_29, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_25])).
fof(c_0_30, plain, ![X19, X20, X21]:(~m1_subset_1(X21,k1_zfmisc_1(k2_zfmisc_1(X19,X20)))|v1_relat_1(X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_31, plain, (r1_xxreal_0(k6_numbers,k17_complex1(X1))|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_26])).
cnf(c_0_32, plain, (k1_funct_1(k54_valued_1(X1),X2)=k17_complex1(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
fof(c_0_33, plain, ![X95, X96]:(~v1_relat_1(X95)|(~v1_funct_1(X95)|~v1_valued_0(X95))|v1_xcmplx_0(k1_funct_1(X95,X96))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc61_valued_0])])).
cnf(c_0_34, plain, (k1_seq_1(X1,X2)=k6_numbers|k4_series_1(X1)!=k6_numbers|~v1_funct_1(X1)|~v1_series_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(k6_numbers,k1_seq_1(X1,esk4_1(X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_28])).
cnf(c_0_35, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_29])).
cnf(c_0_36, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_37, plain, (r1_xxreal_0(k6_numbers,k17_complex1(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_31])).
cnf(c_0_38, plain, (k17_complex1(k1_funct_1(X1,X2))=k1_funct_1(k54_valued_1(X1),X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_32])).
cnf(c_0_39, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_40, plain, (k1_seq_1(X1,X2)=k6_numbers|k4_series_1(X1)!=k6_numbers|~v1_series_1(X1)|~r1_xxreal_0(k6_numbers,k1_funct_1(X1,esk4_1(X1)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35])])).
cnf(c_0_41, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_36])).
cnf(c_0_42, plain, (r1_xxreal_0(k6_numbers,k1_funct_1(k54_valued_1(X1),X2))|~v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_38])])).
cnf(c_0_43, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_39])).
fof(c_0_44, plain, ![X49]:((v1_relat_1(k54_valued_1(X49))|(~v1_relat_1(X49)|(~v1_funct_1(X49)|~v1_valued_0(X49))))&((v1_funct_1(k54_valued_1(X49))|(~v1_relat_1(X49)|(~v1_funct_1(X49)|~v1_valued_0(X49))))&(v3_valued_0(k54_valued_1(X49))|(~v1_relat_1(X49)|(~v1_funct_1(X49)|~v1_valued_0(X49)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k54_valued_1])])])).
cnf(c_0_45, plain, (k1_seq_1(X1,X2)=k6_numbers|k4_series_1(X1)!=k6_numbers|~v3_valued_0(X1)|~v1_series_1(X1)|~r1_xxreal_0(k6_numbers,k1_funct_1(X1,esk4_1(X1)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)), inference(csr,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_46, plain, (r1_xxreal_0(k6_numbers,k1_funct_1(k54_valued_1(X1),X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(csr,[status(thm)],[c_0_42, c_0_43])).
cnf(c_0_47, plain, (v1_funct_1(k54_valued_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_48, plain, (v3_valued_0(k54_valued_1(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
fof(c_0_49, plain, ![X195]:((v1_funct_1(k54_valued_1(X195))|(~v1_funct_1(X195)|(~v1_funct_2(X195,k5_numbers,k2_numbers)|(~v2_comseq_3(X195)|~m1_subset_1(X195,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))))))&((v1_funct_2(k54_valued_1(X195),k5_numbers,k1_numbers)|(~v1_funct_1(X195)|(~v1_funct_2(X195,k5_numbers,k2_numbers)|(~v2_comseq_3(X195)|~m1_subset_1(X195,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))))))&(v1_series_1(k54_valued_1(X195))|(~v1_funct_1(X195)|(~v1_funct_2(X195,k5_numbers,k2_numbers)|(~v2_comseq_3(X195)|~m1_subset_1(X195,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc7_comseq_3])])])).
cnf(c_0_50, plain, (k1_seq_1(k54_valued_1(X1),X2)=k6_numbers|k4_series_1(k54_valued_1(X1))!=k6_numbers|~v3_valued_0(k54_valued_1(X1))|~v1_series_1(k54_valued_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46])])).
cnf(c_0_51, plain, (v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_47])).
cnf(c_0_52, plain, (v3_valued_0(k54_valued_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_48])).
cnf(c_0_53, plain, (v1_series_1(k54_valued_1(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k2_numbers)|~v2_comseq_3(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_54, plain, (k1_seq_1(k54_valued_1(X1),X2)=k6_numbers|k4_series_1(k54_valued_1(X1))!=k6_numbers|~v1_series_1(k54_valued_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_50, c_0_51]), c_0_52])).
cnf(c_0_55, plain, (v1_series_1(k54_valued_1(X1))|~v1_funct_1(X1)|~v2_comseq_3(X1)|~v1_funct_2(X1,k5_numbers,k2_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))), inference(evalgc,[status(thm)],[c_0_53])).
cnf(c_0_56, plain, (v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k2_numbers)|~v2_comseq_3(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_57, negated_conjecture, ~(![X1]:((v1_funct_1(X1)&(v1_funct_2(X1,k5_numbers,k2_numbers)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))))=>((v2_comseq_3(X1)&k4_series_1(k55_valued_1(k5_numbers,k2_numbers,X1))=k6_numbers)=>![X2]:(m2_subset_1(X2,k1_numbers,k5_numbers)=>k3_funct_2(k5_numbers,k2_numbers,X1,X2)=k5_complex1)))), inference(assume_negation,[status(cth)],[t13_csspace2])).
cnf(c_0_58, plain, (k1_seq_1(k54_valued_1(X1),X2)=k6_numbers|k4_series_1(k54_valued_1(X1))!=k6_numbers|~v1_valued_0(X1)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v2_comseq_3(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_2(X1,k5_numbers,k2_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_55])])).
cnf(c_0_59, plain, (v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v2_comseq_3(X1)|~v1_funct_2(X1,k5_numbers,k2_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))), inference(evalgc,[status(thm)],[c_0_56])).
fof(c_0_60, negated_conjecture, ((v1_funct_1(esk1_0)&(v1_funct_2(esk1_0,k5_numbers,k2_numbers)&m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))))&((v2_comseq_3(esk1_0)&k4_series_1(k55_valued_1(k5_numbers,k2_numbers,esk1_0))=k6_numbers)&(m2_subset_1(esk2_0,k1_numbers,k5_numbers)&k3_funct_2(k5_numbers,k2_numbers,esk1_0,esk2_0)!=k5_complex1))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_57])])])).
fof(c_0_61, plain, ![X63, X64, X65]:(~v1_membered(X64)|(~v1_funct_1(X65)|~m1_subset_1(X65,k1_zfmisc_1(k2_zfmisc_1(X63,X64))))|k55_valued_1(X63,X64,X65)=k54_valued_1(X65)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k55_valued_1])])).
fof(c_0_62, plain, ![X32, X33, X34]:(~v1_membered(X33)|(~m1_subset_1(X34,k1_zfmisc_1(k2_zfmisc_1(X32,X33)))|v1_valued_0(X34))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc16_valued_0])])])).
fof(c_0_63, plain, ![X72, X73, X74]:((v1_funct_1(k55_valued_1(X72,X73,X74))|(~v1_membered(X73)|(~v1_funct_1(X74)|~m1_subset_1(X74,k1_zfmisc_1(k2_zfmisc_1(X72,X73))))))&(m1_subset_1(k55_valued_1(X72,X73,X74),k1_zfmisc_1(k2_zfmisc_1(X72,k1_numbers)))|(~v1_membered(X73)|(~v1_funct_1(X74)|~m1_subset_1(X74,k1_zfmisc_1(k2_zfmisc_1(X72,X73))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k55_valued_1])])])).
cnf(c_0_64, plain, (k1_seq_1(k54_valued_1(X1),X2)=k6_numbers|k4_series_1(k54_valued_1(X1))!=k6_numbers|~v1_valued_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v2_comseq_3(X1)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_2(X1,k5_numbers,k2_numbers)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_58, c_0_59]), c_0_41])).
cnf(c_0_65, plain, (v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k2_numbers)|~v2_comseq_3(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_66, negated_conjecture, (k4_series_1(k55_valued_1(k5_numbers,k2_numbers,esk1_0))=k6_numbers), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_67, plain, (k55_valued_1(X3,X1,X2)=k54_valued_1(X2)|~v1_membered(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_68, plain, (v1_membered(k2_numbers)), inference(split_conjunct,[status(thm)],[fc1_membered])).
cnf(c_0_69, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_70, negated_conjecture, (v1_funct_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_71, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_72, plain, (m1_subset_1(k55_valued_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_membered(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_63])).
fof(c_0_73, plain, ![X37]:(~v1_xcmplx_0(X37)|(k17_complex1(X37)!=k6_numbers|X37=k6_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t45_complex1])])).
cnf(c_0_74, plain, (k1_funct_1(k54_valued_1(X1),X2)=k6_numbers|~v3_valued_0(k54_valued_1(X1))|~v1_relat_1(k54_valued_1(X1))|~v1_funct_1(k54_valued_1(X1))|k4_series_1(k54_valued_1(X1))!=k6_numbers|~v1_valued_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v2_comseq_3(X1)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_2(X1,k5_numbers,k2_numbers)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_64])])).
cnf(c_0_75, plain, (v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1)|~v2_comseq_3(X1)|~v1_funct_2(X1,k5_numbers,k2_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_76, negated_conjecture, (k4_series_1(k55_valued_1(k5_numbers,k2_numbers,esk1_0))=k6_numbers), inference(evalgc,[status(thm)],[c_0_66])).
cnf(c_0_77, plain, (k55_valued_1(X3,X1,X2)=k54_valued_1(X2)|~v1_funct_1(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_67])).
cnf(c_0_78, plain, (v1_membered(k2_numbers)), inference(evalgc,[status(thm)],[c_0_68])).
cnf(c_0_79, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))), inference(evalgc,[status(thm)],[c_0_69])).
cnf(c_0_80, negated_conjecture, (v1_funct_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_70])).
cnf(c_0_81, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_71])).
cnf(c_0_82, negated_conjecture, (v2_comseq_3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_83, negated_conjecture, (v1_funct_2(esk1_0,k5_numbers,k2_numbers)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_84, plain, (m1_subset_1(k55_valued_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X3)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_72])).
fof(c_0_85, plain, ![X120, X121, X122]:(~v3_membered(X121)|(~m1_subset_1(X122,k1_zfmisc_1(k2_zfmisc_1(X120,X121)))|v3_valued_0(X122))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc18_valued_0])])])).
cnf(c_0_86, plain, (k5_complex1=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k5_complex1])).
cnf(c_0_87, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
fof(c_0_88, plain, (~v1_xboole_0(k4_ordinal1)&v3_ordinal1(k4_ordinal1)), inference(fof_simplification,[status(thm)],[fc6_ordinal1])).
fof(c_0_89, plain, ![X11, X12, X13]:((~m2_subset_1(X13,X11,X12)|m1_subset_1(X13,X12)|(v1_xboole_0(X11)|(v1_xboole_0(X12)|~m1_subset_1(X12,k1_zfmisc_1(X11)))))&(~m1_subset_1(X13,X12)|m2_subset_1(X13,X11,X12)|(v1_xboole_0(X11)|(v1_xboole_0(X12)|~m1_subset_1(X12,k1_zfmisc_1(X11)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])).
fof(c_0_90, plain, ![X69, X70]:(~v1_xboole_0(X69)|(~m1_subset_1(X70,k1_zfmisc_1(X69))|v1_xboole_0(X70))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
cnf(c_0_91, plain, (X1=k6_numbers|~v1_xcmplx_0(X1)|k17_complex1(X1)!=k6_numbers), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_92, plain, (k1_funct_1(k54_valued_1(X1),X2)=k6_numbers|k4_series_1(k54_valued_1(X1))!=k6_numbers|~v3_valued_0(k54_valued_1(X1))|~v1_valued_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v2_comseq_3(X1)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_2(X1,k5_numbers,k2_numbers)|~v1_funct_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_74, c_0_75]), c_0_41])).
cnf(c_0_93, negated_conjecture, (k4_series_1(k54_valued_1(esk1_0))=k6_numbers), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_78]), c_0_79]), c_0_80])])])).
cnf(c_0_94, negated_conjecture, (v1_valued_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_79]), c_0_78])])])).
cnf(c_0_95, negated_conjecture, (v2_comseq_3(esk1_0)), inference(evalgc,[status(thm)],[c_0_82])).
cnf(c_0_96, negated_conjecture, (v1_funct_2(esk1_0,k5_numbers,k2_numbers)), inference(evalgc,[status(thm)],[c_0_83])).
cnf(c_0_97, plain, (m1_subset_1(k54_valued_1(X3),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_77])])).
cnf(c_0_98, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_99, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
cnf(c_0_100, negated_conjecture, (k3_funct_2(k5_numbers,k2_numbers,esk1_0,esk2_0)!=k5_complex1), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_101, plain, (k5_complex1=k1_xboole_0), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_102, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_87])).
fof(c_0_103, plain, ![X26, X27, X28, X29]:(v1_xboole_0(X26)|(~v1_funct_1(X28)|(~v1_funct_2(X28,X26,X27)|~m1_subset_1(X28,k1_zfmisc_1(k2_zfmisc_1(X26,X27))))|~m1_subset_1(X29,X26))|k3_funct_2(X26,X27,X28,X29)=k1_funct_1(X28,X29)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_funct_2])])])).
cnf(c_0_104, plain, (~v1_xboole_0(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_105, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
cnf(c_0_106, plain, (m1_subset_1(X1,X3)|v1_xboole_0(X2)|v1_xboole_0(X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_89])).
cnf(c_0_107, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_108, plain, (X1=k6_numbers|k17_complex1(X1)!=k6_numbers|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_91])).
cnf(c_0_109, negated_conjecture, (k1_funct_1(k54_valued_1(esk1_0),X1)=k6_numbers|~v3_valued_0(k54_valued_1(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(esk1_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93]), c_0_94]), c_0_95]), c_0_79]), c_0_96]), c_0_80])])])).
cnf(c_0_110, negated_conjecture, (m1_subset_1(k54_valued_1(esk1_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_79]), c_0_78]), c_0_80])])])).
cnf(c_0_111, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_112, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_99])).
cnf(c_0_113, negated_conjecture, (k3_funct_2(k5_numbers,k2_numbers,esk1_0,esk2_0)!=k5_complex1), inference(evalgc,[status(thm)],[c_0_100])).
cnf(c_0_114, plain, (k5_complex1=k6_numbers), inference(rw,[status(thm)],[c_0_101, c_0_102])).
cnf(c_0_115, plain, (v1_xboole_0(X1)|k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|~v1_funct_1(X2)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_116, plain, (~v1_xboole_0(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_104])).
cnf(c_0_117, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_118, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_119, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_107])).
cnf(c_0_120, negated_conjecture, (m2_subset_1(esk2_0,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_121, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
cnf(c_0_122, plain, (k1_funct_1(X1,X2)=k6_numbers|k1_funct_1(k54_valued_1(X1),X2)!=k6_numbers|~v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_38])])).
cnf(c_0_123, negated_conjecture, (k1_funct_1(k54_valued_1(esk1_0),X1)=k6_numbers|~v3_valued_0(k54_valued_1(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_109, c_0_110])])])).
cnf(c_0_124, negated_conjecture, (v3_valued_0(k54_valued_1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_110]), c_0_112])])])).
cnf(c_0_125, negated_conjecture, (k3_funct_2(k5_numbers,k2_numbers,esk1_0,esk2_0)!=k6_numbers), inference(rw,[status(thm)],[c_0_113, c_0_114])).
cnf(c_0_126, plain, (k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))), inference(evalgc,[status(thm)],[c_0_115])).
cnf(c_0_127, plain, (~v1_xboole_0(k5_numbers)), inference(rw,[status(thm)],[c_0_116, c_0_117])).
cnf(c_0_128, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X1)|~m2_subset_1(X2,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(csr,[status(thm)],[c_0_118, c_0_119])).
cnf(c_0_129, negated_conjecture, (m2_subset_1(esk2_0,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_120])).
cnf(c_0_130, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_121])).
cnf(c_0_131, plain, (k1_funct_1(X1,X2)=k6_numbers|k1_funct_1(k54_valued_1(X1),X2)!=k6_numbers|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(csr,[status(thm)],[c_0_122, c_0_43])).
cnf(c_0_132, negated_conjecture, (k1_funct_1(k54_valued_1(esk1_0),X1)=k6_numbers|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_123, c_0_124])])])).
cnf(c_0_133, negated_conjecture, (v1_relat_1(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_79])])).
cnf(c_0_134, negated_conjecture, (k1_funct_1(esk1_0,esk2_0)!=k6_numbers|~m1_subset_1(esk2_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_79]), c_0_96]), c_0_80])]), c_0_127])])).
cnf(c_0_135, negated_conjecture, (m1_subset_1(esk2_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129]), c_0_130])]), c_0_127])])).
cnf(c_0_136, negated_conjecture, (k1_funct_1(esk1_0,X1)=k6_numbers|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_132]), c_0_94]), c_0_133]), c_0_80])])])).
cnf(c_0_137, negated_conjecture, (k1_funct_1(esk1_0,esk2_0)!=k6_numbers), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_134, c_0_135])])])).
cnf(c_0_138, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_129]), c_0_137]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 139
# Proof object clause steps            : 96
# Proof object formula steps           : 43
# Proof object conjectures             : 31
# Proof object clause conjectures      : 28
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 33
# Proof object initial formulas used   : 24
# Proof object generating inferences   : 17
# Proof object simplifying inferences  : 49
# Training examples: 52 positive, 401 negative
# Training: Positive examples begin
cnf(c_0_102, plain, (k1_xboole_0=k6_numbers)).# trainpos
cnf(c_0_101, plain, (k5_complex1=k1_xboole_0)).# trainpos
cnf(c_0_117, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_80, negated_conjecture, (v1_funct_1(esk1_0))).# trainpos
cnf(c_0_95, negated_conjecture, (v2_comseq_3(esk1_0))).# trainpos
cnf(c_0_113, negated_conjecture, (k3_funct_2(k5_numbers,k2_numbers,esk1_0,esk2_0)!=k5_complex1)).# trainpos
cnf(c_0_96, negated_conjecture, (v1_funct_2(esk1_0,k5_numbers,k2_numbers))).# trainpos
cnf(c_0_108, plain, (X1=k6_numbers|k17_complex1(X1)!=k6_numbers|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_78, plain, (v1_membered(k2_numbers))).# trainpos
cnf(c_0_112, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_116, plain, (~v1_xboole_0(k4_ordinal1))).# trainpos
cnf(c_0_129, negated_conjecture, (m2_subset_1(esk2_0,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_76, negated_conjecture, (k4_series_1(k55_valued_1(k5_numbers,k2_numbers,esk1_0))=k6_numbers)).# trainpos
cnf(c_0_130, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_79, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers))))).# trainpos
cnf(c_0_75, plain, (v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1)|~v2_comseq_3(X1)|~v1_funct_2(X1,k5_numbers,k2_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers))))).# trainpos
cnf(c_0_37, plain, (r1_xxreal_0(k6_numbers,k17_complex1(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_51, plain, (v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).# trainpos
cnf(c_0_119, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_52, plain, (v3_valued_0(k54_valued_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).# trainpos
cnf(c_0_41, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_55, plain, (v1_series_1(k54_valued_1(X1))|~v1_funct_1(X1)|~v2_comseq_3(X1)|~v1_funct_2(X1,k5_numbers,k2_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers))))).# trainpos
cnf(c_0_118, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_81, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_111, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_43, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).# trainpos
cnf(c_0_35, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_59, plain, (v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v2_comseq_3(X1)|~v1_funct_2(X1,k5_numbers,k2_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers))))).# trainpos
cnf(c_0_38, plain, (k17_complex1(k1_funct_1(X1,X2))=k1_funct_1(k54_valued_1(X1),X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).# trainpos
cnf(c_0_77, plain, (k55_valued_1(X3,X1,X2)=k54_valued_1(X2)|~v1_funct_1(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_34, plain, (k1_seq_1(X1,X2)=k6_numbers|k4_series_1(X1)!=k6_numbers|~v1_funct_1(X1)|~v1_series_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~r1_xxreal_0(k6_numbers,k1_seq_1(X1,esk4_1(X1)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_84, plain, (m1_subset_1(k55_valued_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X3)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_126, plain, (k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).# trainpos
cnf(c_0_133, negated_conjecture, (v1_relat_1(esk1_0))).# trainpos
cnf(c_0_134, negated_conjecture, (k1_funct_1(esk1_0,esk2_0)!=k6_numbers|~m1_subset_1(esk2_0,k5_numbers))).# trainpos
cnf(c_0_135, negated_conjecture, (m1_subset_1(esk2_0,k5_numbers))).# trainpos
cnf(c_0_93, negated_conjecture, (k4_series_1(k54_valued_1(esk1_0))=k6_numbers)).# trainpos
cnf(c_0_40, plain, (k1_seq_1(X1,X2)=k6_numbers|k4_series_1(X1)!=k6_numbers|~v1_series_1(X1)|~r1_xxreal_0(k6_numbers,k1_funct_1(X1,esk4_1(X1)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_122, plain, (k1_funct_1(X1,X2)=k6_numbers|k1_funct_1(k54_valued_1(X1),X2)!=k6_numbers|~v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).# trainpos
cnf(c_0_94, negated_conjecture, (v1_valued_0(esk1_0))).# trainpos
cnf(c_0_137, negated_conjecture, (k1_funct_1(esk1_0,esk2_0)!=k6_numbers)).# trainpos
cnf(c_0_42, plain, (r1_xxreal_0(k6_numbers,k1_funct_1(k54_valued_1(X1),X2))|~v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).# trainpos
cnf(c_0_50, plain, (k1_seq_1(k54_valued_1(X1),X2)=k6_numbers|k4_series_1(k54_valued_1(X1))!=k6_numbers|~v3_valued_0(k54_valued_1(X1))|~v1_series_1(k54_valued_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).# trainpos
cnf(c_0_58, plain, (k1_seq_1(k54_valued_1(X1),X2)=k6_numbers|k4_series_1(k54_valued_1(X1))!=k6_numbers|~v1_valued_0(X1)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v2_comseq_3(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_2(X1,k5_numbers,k2_numbers))).# trainpos
cnf(c_0_74, plain, (k1_funct_1(k54_valued_1(X1),X2)=k6_numbers|~v3_valued_0(k54_valued_1(X1))|~v1_relat_1(k54_valued_1(X1))|~v1_funct_1(k54_valued_1(X1))|k4_series_1(k54_valued_1(X1))!=k6_numbers|~v1_valued_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v2_comseq_3(X1)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_2(X1,k5_numbers,k2_numbers)|~v1_funct_1(X1))).# trainpos
cnf(c_0_109, negated_conjecture, (k1_funct_1(k54_valued_1(esk1_0),X1)=k6_numbers|~v3_valued_0(k54_valued_1(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(esk1_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_97, plain, (m1_subset_1(k54_valued_1(X3),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).# trainpos
cnf(c_0_110, negated_conjecture, (m1_subset_1(k54_valued_1(esk1_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).# trainpos
cnf(c_0_123, negated_conjecture, (k1_funct_1(k54_valued_1(esk1_0),X1)=k6_numbers|~v3_valued_0(k54_valued_1(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_124, negated_conjecture, (v3_valued_0(k54_valued_1(esk1_0)))).# trainpos
cnf(c_0_132, negated_conjecture, (k1_funct_1(k54_valued_1(esk1_0),X1)=k6_numbers|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_136, negated_conjecture, (k1_funct_1(esk1_0,X1)=k6_numbers|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_77, plain, (k4_xcmplx_0(np__0)=np__0)).#trainneg
cnf(i_0_144, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_73, plain, (v1_xboole_0(np__0))).#trainneg
cnf(i_0_84, plain, (v1_xboole_0(esk5_0))).#trainneg
cnf(i_0_177, plain, (v1_xboole_0(esk9_0))).#trainneg
cnf(i_0_81, plain, (v1_xreal_0(esk5_0))).#trainneg
cnf(i_0_18, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_83, plain, (v1_xcmplx_0(esk5_0))).#trainneg
cnf(i_0_64, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_30, plain, (~v1_xboole_0(k2_numbers))).#trainneg
cnf(i_0_16, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_130, plain, (m1_subset_1(k5_complex1,k2_numbers))).#trainneg
cnf(i_0_65, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_82, plain, (v1_xxreal_0(esk5_0))).#trainneg
cnf(i_0_176, plain, (v1_xxreal_0(esk9_0))).#trainneg
cnf(i_0_101, plain, (v6_membered(k4_ordinal1))).#trainneg
cnf(i_0_172, plain, (v1_xboole_0(esk8_1(X1)))).#trainneg
cnf(i_0_161, plain, (k2_xcmplx_0(np__0,np__0)=np__0)).#trainneg
cnf(i_0_134, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_63, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_62, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_41, plain, (m1_subset_1(np__0,k5_numbers))).#trainneg
cnf(i_0_40, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_178, plain, (r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_136, plain, (r1_xxreal_0(np__0,np__0))).#trainneg
cnf(i_0_10, plain, (X1=k2_series_1(X2)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|k2_xcmplx_0(k1_funct_1(X1,esk3_2(X2,X1)),k1_funct_1(X2,k2_nat_1(esk3_2(X2,X1),np__1)))!=k1_funct_1(X1,k2_nat_1(esk3_2(X2,X1),np__1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_36, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_42, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_168, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_125, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_171, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_79, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_92, plain, (v1_xcmplx_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_143, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_56, plain, (v1_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_173, plain, (m1_subset_1(esk8_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_72, plain, (k16_complex1(X1)=k17_complex1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_148, plain, (v1_funct_1(k3_sin_cos(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_106, plain, (v2_comseq_3(k27_sin_cos(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_146, plain, (v2_comseq_3(k3_sin_cos(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_116, plain, (v1_xreal_0(k16_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_114, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_61, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_105, plain, (r1_xxreal_0(k6_numbers,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_54, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_151, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_140, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_111, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_112, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_99, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k2_numbers))).#trainneg
cnf(i_0_160, plain, (~v1_xboole_0(X1)|~v2_xxreal_0(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_158, plain, (~v2_xxreal_0(X1)|~v1_xxreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_88, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_115, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_108, plain, (k6_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_51, plain, (k17_complex1(k17_complex1(X1))=k17_complex1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_93, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_80, plain, (k16_complex1(k16_complex1(X1))=k16_complex1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_69, plain, (m1_subset_1(k17_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_155, plain, (v3_valued_0(k2_zfmisc_1(X2,X1))|~v3_membered(X1))).#trainneg
cnf(i_0_141, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_147, plain, (v1_funct_2(k3_sin_cos(X1),k5_numbers,k2_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_139, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_102, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k2_numbers))).#trainneg
cnf(i_0_78, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_76, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_110, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_55, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_123, plain, (v2_comseq_3(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k2_numbers)|~v1_series_1(k55_valued_1(k5_numbers,k2_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers))))).#trainneg
cnf(i_0_174, plain, (v1_partfun1(X2,X1)|k1_relset_1(X1,X2)!=X1|~v1_relat_1(X2)|~v4_relat_1(X2,X1))).#trainneg
cnf(i_0_50, plain, (v1_relat_1(k54_valued_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_75, plain, (k2_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_129, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_14, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_71, plain, (v1_funct_1(k55_valued_1(X1,X2,X3))|~v1_funct_1(X3)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_53, plain, (k54_valued_1(k54_valued_1(X1))=k54_valued_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_138, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_164, plain, (v1_xcmplx_0(k6_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_47, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_46, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_104, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_157, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_120, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_89, plain, (k2_nat_1(X1,X2)=k2_nat_1(X2,X1)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_119, plain, (k3_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_142, plain, (k3_real_1(X1,X2)=k3_real_1(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_118, plain, (k7_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_154, plain, (v1_funct_1(k3_series_1(X1))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_156, plain, (m1_subset_1(k4_series_1(X1),k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_180, plain, (X3=X4|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_2(X4,X1,X2)|~v1_funct_2(X3,X1,X2)|~r2_funct_2(X1,X2,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_131, plain, (v2_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_128, plain, (k9_real_1(X1,X2)=k6_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_145, plain, (m1_subset_1(k2_seq_2(X1),k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_153, plain, (v1_funct_2(k3_series_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_163, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_121, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X2)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_150, plain, (k7_real_1(X1,X2)=k7_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_126, plain, (m1_subset_1(k3_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_162, plain, (m1_subset_1(k7_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_100, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_107, plain, (m2_subset_1(k2_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_74, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_175, plain, (k1_relset_1(X2,X1)=X2|~v1_relat_1(X1)|~v4_relat_1(X1,X2)|~v1_partfun1(X1,X2))).#trainneg
cnf(i_0_149, plain, (k8_real_1(k17_complex1(X1),k17_complex1(X2))=k17_complex1(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_94, plain, (k2_seq_2(k3_series_1(X1))=k4_series_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_152, plain, (m1_subset_1(k3_series_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_95, plain, (v7_valued_0(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~r1_xxreal_0(k8_nat_1(k1_numbers,X1,esk6_1(X1)),k8_nat_1(k1_numbers,X1,k2_nat_1(esk6_1(X1),esk7_1(X1)))))).#trainneg
cnf(i_0_135, plain, (k3_series_1(X1)=k2_series_1(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_124, plain, (v1_series_1(k55_valued_1(k5_numbers,k2_numbers,X1))|~v1_funct_1(X1)|~v2_comseq_3(X1)|~v1_funct_2(X1,k5_numbers,k2_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers))))).#trainneg
cnf(i_0_33, plain, (v1_funct_1(k2_series_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_170, plain, (v1_xboole_0(X2)|v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_169, plain, (v1_xboole_0(X1)|v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_28, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(k3_real_1(X1,np__1),X2)|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_25, plain, (r2_hidden(k6_xcmplx_0(X2,X1),k5_numbers)|~v1_int_1(X2)|~v1_int_1(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_133, plain, (k1_seq_2(X1)=k2_seq_2(X1)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_67, plain, (v1_xboole_0(X1)|v1_partfun1(X2,X3)|~v1_funct_2(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_35, plain, (v1_relat_1(k2_series_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_31, plain, (v1_valued_0(k2_series_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_34, plain, (v4_relat_1(k2_series_1(X1),k5_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_32, plain, (v1_partfun1(k2_series_1(X1),k5_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_58, plain, (v1_partfun1(k54_valued_1(X1),X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1)|~v4_relat_1(X1,X2)|~v1_partfun1(X1,X2))).#trainneg
cnf(i_0_127, plain, (m1_subset_1(k8_nat_1(X2,X1,X3),X2)|~v1_funct_1(X1)|~v7_ordinal1(X3)|~v1_funct_2(X1,k5_numbers,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_117, plain, (k1_seq_1(X1,X2)=k6_numbers|k4_series_1(k56_valued_1(k5_numbers,k1_numbers,X1))!=k6_numbers|~v1_funct_1(X1)|~v2_series_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_97, plain, (v7_valued_0(X1)|m2_subset_1(esk6_1(X1),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_17, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X3),k2_xcmplx_0(X2,X4))|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X3,X4)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_96, plain, (v7_valued_0(X1)|m2_subset_1(esk7_1(X1),k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_86, plain, (v1_xboole_0(X2)|v1_funct_1(k54_valued_1(X1))|~v1_funct_1(X1)|~v1_membered(X2)|~v1_funct_2(X1,X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_87, plain, (r1_xxreal_0(k8_nat_1(k1_numbers,X1,X2),k2_seq_2(X1))|~v1_funct_1(X1)|~v1_seq_2(X1)|~v7_valued_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_45, plain, (k55_valued_1(X3,X1,k55_valued_1(X3,X1,X2))=k55_valued_1(X3,X1,X2)|~v1_funct_1(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_13, plain, (k1_funct_1(X1,k6_numbers)=k1_funct_1(X2,k6_numbers)|X1!=k2_series_1(X2)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_85, plain, (v1_xboole_0(X3)|v1_partfun1(k54_valued_1(X1),X2)|~v1_funct_1(X1)|~v1_membered(X3)|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_137, plain, (k25_valued_1(X3,X1,X2,X4)=k24_valued_1(X2,X4)|~v1_funct_1(X2)|~v1_xcmplx_0(X4)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_52, plain, (k8_nat_1(X2,X1,X3)=k1_funct_1(X1,X3)|~v1_funct_1(X1)|~v7_ordinal1(X3)|~v1_funct_2(X1,k5_numbers,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_91, plain, (v1_xboole_0(X1)|m1_subset_1(k3_funct_2(X1,X3,X2,X4),X3)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_39, plain, (k1_seq_1(X1,X2)=k6_numbers|m2_subset_1(esk4_1(X1),k1_numbers,k5_numbers)|k4_series_1(X1)!=k6_numbers|~v1_funct_1(X1)|~v1_series_1(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_179, plain, (r2_funct_2(X3,X4,X1,X2)|X1!=X2|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,X3,X4)|~v1_funct_2(X1,X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_98, plain, (r1_xxreal_0(k8_nat_1(k1_numbers,X1,X2),k8_nat_1(k1_numbers,X1,k2_nat_1(X2,X3)))|~v1_funct_1(X1)|~v7_valued_0(X1)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers))))).#trainneg
cnf(i_0_11, plain, (X2=k2_series_1(X1)|m2_subset_1(esk3_2(X1,X2),k1_numbers,k5_numbers)|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers))).#trainneg
cnf(i_0_12, plain, (k2_xcmplx_0(k1_funct_1(X2,X1),k1_funct_1(X3,k2_nat_1(X1,np__1)))=k1_funct_1(X2,k2_nat_1(X1,np__1))|X2!=k2_series_1(X3)|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_valued_0(X3)|~v1_valued_0(X2)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v1_partfun1(X3,k5_numbers)|~v1_partfun1(X2,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_268, negated_conjecture, (v4_relat_1(esk1_0,k5_numbers))).#trainneg
cnf(i_0_361, plain, (k2_seq_2(k2_series_1(X1))=k4_series_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_181, plain, (esk9_0=k6_numbers)).#trainneg
cnf(i_0_183, plain, (esk5_0=k6_numbers)).#trainneg
cnf(i_0_184, plain, (np__0=k6_numbers)).#trainneg
cnf(i_0_182, plain, (esk8_1(X1)=k6_numbers)).#trainneg
cnf(i_0_270, negated_conjecture, (v5_relat_1(esk1_0,k2_numbers))).#trainneg
cnf(i_0_187, plain, (v1_membered(k1_numbers))).#trainneg
cnf(i_0_235, plain, (v1_xcmplx_0(k6_numbers))).#trainneg
cnf(i_0_246, plain, (v7_ordinal1(k6_numbers))).#trainneg
cnf(i_0_186, plain, (v1_xcmplx_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_188, plain, (v1_xcmplx_0(k16_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_256, plain, (v3_membered(k5_numbers))).#trainneg
cnf(i_0_391, negated_conjecture, (v1_partfun1(esk1_0,k5_numbers))).#trainneg
cnf(i_0_363, plain, (m1_subset_1(k2_series_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_362, plain, (v1_funct_2(k2_series_1(X1),k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_360, plain, (v1_funct_1(k2_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_440, plain, (v4_relat_1(k55_valued_1(X1,X2,X3),X1)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_279, plain, (k1_relset_1(X1,X2)=k1_relset_1(X3,X2)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v4_relat_1(X2,X3))).#trainneg
cnf(i_0_455, plain, (k1_funct_1(k2_series_1(X1),k6_numbers)=k1_funct_1(X1,k6_numbers)|~v1_valued_0(X1)|~v1_valued_0(k2_series_1(X1))|~v1_partfun1(X1,k5_numbers)|~v1_partfun1(k2_series_1(X1),k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(k2_series_1(X1),k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(k2_series_1(X1))|~v1_funct_1(X1)|~v1_funct_1(k2_series_1(X1)))).#trainneg
cnf(i_0_387, plain, (v2_comseq_3(X1)|~v1_series_1(k54_valued_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_2(X1,k5_numbers,k2_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_194, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_223, plain, (v1_xreal_0(np__1))).#trainneg
cnf(i_0_189, plain, (v1_xreal_0(k17_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_244, plain, (v7_ordinal1(np__1))).#trainneg
cnf(i_0_205, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_193, plain, (r1_xxreal_0(k6_numbers,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_196, negated_conjecture, (v1_funct_1(k54_valued_1(esk1_0)))).#trainneg
cnf(i_0_450, plain, (k54_valued_1(k55_valued_1(X1,X2,X3))=k55_valued_1(X1,X2,X3)|~v1_membered(X2)|~m1_subset_1(k55_valued_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(k55_valued_1(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_442, plain, (v1_partfun1(k55_valued_1(X1,X2,X3),X1)|~v1_funct_2(k55_valued_1(X1,X2,X3),X1,k1_numbers)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_441, plain, (v1_funct_2(k55_valued_1(X1,X2,X3),X1,k1_numbers)|~v1_partfun1(k55_valued_1(X1,X2,X3),X1)|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_480, plain, (v1_xboole_0(X1)|m1_subset_1(k1_funct_1(X3,X4),X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_378, plain, (k8_real_1(k17_complex1(X1),k1_funct_1(k54_valued_1(X2),X3))=k17_complex1(k3_xcmplx_0(X1,k1_funct_1(X2,X3)))|~v1_xcmplx_0(k1_funct_1(X2,X3))|~v1_xcmplx_0(X1)|~v1_valued_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_273, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k6_numbers))).#trainneg
cnf(i_0_439, plain, (v1_relat_1(k55_valued_1(X1,X2,X3))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_462, plain, (v7_valued_0(X1)|~r1_xxreal_0(k1_funct_1(X1,esk6_1(X1)),k8_nat_1(k1_numbers,X1,k2_nat_1(esk6_1(X1),esk7_1(X1))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v7_ordinal1(esk6_1(X1)))).#trainneg
cnf(i_0_225, plain, (v1_xreal_0(k6_numbers)|~v3_membered(k2_numbers))).#trainneg
cnf(i_0_206, plain, (r1_xxreal_0(k6_numbers,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_176, plain, (v1_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_197, plain, (v1_funct_1(k54_valued_1(k3_sin_cos(X1)))|~m1_subset_1(k3_sin_cos(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_2(k3_sin_cos(X1),k5_numbers,k2_numbers)|~v1_funct_1(k3_sin_cos(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_42, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_81, plain, (v1_xreal_0(k6_numbers))).#trainneg
cnf(i_0_198, plain, (v1_funct_1(k54_valued_1(k27_sin_cos(X1)))|~m1_subset_1(k27_sin_cos(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_2(k27_sin_cos(X1),k5_numbers,k2_numbers)|~v1_funct_1(k27_sin_cos(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_519, negated_conjecture, (v7_ordinal1(esk2_0))).#trainneg
cnf(i_0_517, negated_conjecture, (v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_516, plain, (v1_membered(k5_numbers))).#trainneg
cnf(i_0_496, plain, (r1_xxreal_0(k1_funct_1(X1,X2),k8_nat_1(k1_numbers,X1,k2_nat_1(X2,X3)))|~v7_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_538, plain, (k1_relset_1(X3,X2)=X1|~v1_partfun1(X2,X1)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v4_relat_1(X2,X3))).#trainneg
cnf(i_0_533, plain, (v3_valued_0(k2_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_518, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_575, plain, (r1_xxreal_0(k6_numbers,np__1))).#trainneg
cnf(i_0_77, plain, (k4_xcmplx_0(k6_numbers)=k6_numbers)).#trainneg
cnf(i_0_569, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_161, plain, (k2_xcmplx_0(k6_numbers,k6_numbers)=k6_numbers)).#trainneg
cnf(i_0_207, plain, (k17_complex1(k16_complex1(X1))=k16_complex1(X1)|~v1_xcmplx_0(k16_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_136, plain, (r1_xxreal_0(k6_numbers,k6_numbers))).#trainneg
cnf(i_0_40, plain, (m1_subset_1(k6_numbers,k1_numbers))).#trainneg
cnf(i_0_604, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(k6_numbers,X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_529, plain, (v1_partfun1(k2_series_1(X1),k5_numbers)|~v1_funct_2(k2_series_1(X1),k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_499, plain, (k2_xcmplx_0(k1_funct_1(X1,k2_nat_1(X2,np__1)),k1_funct_1(X3,X2))=k1_funct_1(X3,k2_nat_1(X2,np__1))|~v1_xcmplx_0(k1_funct_1(X3,X2))|~v1_xcmplx_0(k1_funct_1(X1,k2_nat_1(X2,np__1)))|X3!=k2_series_1(X1)|~v1_valued_0(X1)|~v1_valued_0(X3)|~v1_partfun1(X1,k5_numbers)|~v1_partfun1(X3,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X1)|~v1_relat_1(X3)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_492, plain, (r1_xxreal_0(k8_nat_1(k1_numbers,X1,X2),k8_nat_1(k1_numbers,X1,k2_nat_1(X3,X2)))|~v7_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_503, plain, (r1_xxreal_0(k2_xcmplx_0(X1,X2),k1_funct_1(X3,k2_nat_1(X4,np__1)))|~r1_xxreal_0(X2,k1_funct_1(X5,k2_nat_1(X4,np__1)))|~r1_xxreal_0(X1,k1_funct_1(X3,X4))|~v1_xreal_0(k1_funct_1(X5,k2_nat_1(X4,np__1)))|~v1_xreal_0(X2)|~v1_xreal_0(k1_funct_1(X3,X4))|~v1_xreal_0(X1)|X3!=k2_series_1(X5)|~v1_valued_0(X5)|~v1_valued_0(X3)|~v1_partfun1(X5,k5_numbers)|~v1_partfun1(X3,k5_numbers)|~v4_relat_1(X5,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v1_relat_1(X5)|~v1_relat_1(X3)|~m2_subset_1(X4,k1_numbers,k5_numbers)|~v1_funct_1(X5)|~v1_funct_1(X3))).#trainneg
cnf(i_0_292, plain, (v1_series_1(k54_valued_1(X1))|~v2_comseq_3(k54_valued_1(X1))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k2_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_353, plain, (v1_relat_1(k3_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_173, plain, (m1_subset_1(k6_numbers,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_307, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers))).#trainneg
cnf(i_0_179, plain, (r2_funct_2(X1,X2,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_210, plain, (k16_complex1(k17_complex1(X1))=k17_complex1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_331, negated_conjecture, (v3_valued_0(esk1_0)|~v3_membered(k2_numbers))).#trainneg
cnf(i_0_557, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_354, plain, (v4_relat_1(k3_series_1(X1),k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_493, plain, (r1_xxreal_0(k8_nat_1(k1_numbers,X1,X2),k8_nat_1(k1_numbers,X1,k2_nat_1(X3,X2)))|~v7_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_540, plain, (v1_partfun1(X1,X2)|k1_relset_1(X3,X1)!=X2|~v4_relat_1(X1,X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X3))).#trainneg
cnf(i_0_526, plain, (v4_relat_1(k2_series_1(X1),k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_464, plain, (m1_subset_1(k1_funct_1(X2,X3),X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_201, plain, (~v2_xxreal_0(X1)|~v1_xboole_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_383, plain, (k8_real_1(k1_funct_1(k54_valued_1(X1),X2),k17_complex1(X3))=k17_complex1(k3_xcmplx_0(k1_funct_1(X1,X2),X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_448, plain, (v1_series_1(k55_valued_1(k5_numbers,k2_numbers,X1))|~v2_comseq_3(k55_valued_1(k5_numbers,k2_numbers,X1))|~m1_subset_1(k55_valued_1(k5_numbers,k2_numbers,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_2(k55_valued_1(k5_numbers,k2_numbers,X1),k5_numbers,k2_numbers)|~v1_funct_1(k55_valued_1(k5_numbers,k2_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_289, plain, (k2_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_622, negated_conjecture, (r1_xxreal_0(k6_numbers,esk2_0))).#trainneg
cnf(i_0_716, plain, (v1_relat_1(k6_numbers))).#trainneg
cnf(i_0_717, plain, (v4_relat_1(k6_numbers,X1))).#trainneg
cnf(i_0_614, plain, (k1_seq_1(X1,X2)=k6_numbers|k4_series_1(X1)!=k6_numbers|~v1_series_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_seq_1(X1,esk4_1(X1)),k5_numbers))).#trainneg
cnf(i_0_213, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~v1_relat_1(k2_zfmisc_1(X1,X2))|~v3_membered(X2))).#trainneg
cnf(i_0_638, plain, (v3_xxreal_0(k6_numbers)|~v3_xxreal_0(np__1))).#trainneg
cnf(i_0_763, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_725, plain, (v5_relat_1(k6_numbers,X1))).#trainneg
cnf(i_0_723, plain, (v1_valued_0(k6_numbers)|~v1_membered(X1))).#trainneg
cnf(i_0_217, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_219, negated_conjecture, (v1_xreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers))))).#trainneg
cnf(i_0_634, negated_conjecture, (k1_relset_1(X1,esk1_0)=k5_numbers|~v4_relat_1(esk1_0,X1))).#trainneg
cnf(i_0_758, negated_conjecture, (m1_subset_1(k1_funct_1(X1,esk2_0),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(X1,k5_numbers,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_720, plain, (v1_xboole_0(X1)|v1_partfun1(k6_numbers,X2)|~v1_funct_2(k6_numbers,X2,X1))).#trainneg
cnf(i_0_629, plain, (r1_xxreal_0(k1_funct_1(X1,X2),k8_nat_1(k1_numbers,X1,k2_nat_1(X3,X2)))|~v7_valued_0(X1)|~v7_ordinal1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_718, plain, (v1_funct_2(k6_numbers,X1,X2)|~v1_partfun1(k6_numbers,X1))).#trainneg
cnf(i_0_762, plain, (m1_subset_1(k1_funct_1(X1,X2),X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X3)))|~v1_funct_2(X1,k5_numbers,X3)|~v1_funct_1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_630, plain, (r1_xxreal_0(k1_funct_1(X1,X2),k8_nat_1(k1_numbers,X1,k2_nat_1(X3,X2)))|~v7_valued_0(X1)|~v7_ordinal1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_724, plain, (v3_valued_0(k6_numbers)|~v3_membered(X1))).#trainneg
cnf(i_0_759, plain, (m1_subset_1(k1_funct_1(X1,np__1),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(X1,k5_numbers,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_615, plain, (k1_seq_1(X1,X2)=k6_numbers|k4_series_1(X1)!=k6_numbers|~v3_valued_0(X1)|~v1_series_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(k1_funct_1(X1,esk4_1(X1)),k5_numbers))).#trainneg
cnf(i_0_728, plain, (v3_membered(k6_numbers)|~v3_membered(X1))).#trainneg
cnf(i_0_227, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_237, plain, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_616, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_564, plain, (v1_valued_0(k2_zfmisc_1(X1,X2))|~v1_membered(X2))).#trainneg
cnf(i_0_229, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers))))).#trainneg
cnf(i_0_232, plain, (v1_xcmplx_0(k17_complex1(X1))|~v1_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_631, plain, (r1_xxreal_0(k1_funct_1(X1,X2),k8_nat_1(k1_numbers,X1,k2_xcmplx_0(X2,X3)))|~v7_valued_0(X1)|~v7_ordinal1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_676, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X3,X2))|~r1_xxreal_0(k6_numbers,X2)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_705, plain, (v1_relat_1(k2_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_635, plain, (k1_relset_1(X1,k2_series_1(X2))=k5_numbers|~v1_partfun1(k2_series_1(X2),k5_numbers)|~v4_relat_1(k2_series_1(X2),X1)|~v1_relat_1(k2_series_1(X2))|~v1_valued_0(X2)|~v1_partfun1(X2,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v1_relat_1(X2)|~v1_funct_1(X2))).#trainneg
cnf(i_0_879, plain, (k1_seq_1(X1,X2)=k6_numbers|k4_series_1(X1)!=k6_numbers|~v3_valued_0(X1)|~v1_series_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k5_numbers)))|~m1_subset_1(esk4_1(X1),k5_numbers)|~v1_funct_2(X1,k5_numbers,k5_numbers))).#trainneg
cnf(i_0_855, plain, (r1_xxreal_0(k1_funct_1(X1,X2),k8_nat_1(k1_numbers,X1,k2_nat_1(X2,X3)))|~v7_valued_0(X1)|~v7_ordinal1(X3)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_793, plain, (v1_valued_0(k6_numbers))).#trainneg
cnf(i_0_861, plain, (v3_valued_0(k6_numbers))).#trainneg
cnf(i_0_880, plain, (v3_membered(k6_numbers))).#trainneg
cnf(i_0_239, negated_conjecture, (v7_ordinal1(esk1_0)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers))))).#trainneg
cnf(i_0_252, negated_conjecture, (v1_xboole_0(esk1_0)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k2_numbers)))).#trainneg
cnf(i_0_785, negated_conjecture, (v3_xxreal_0(k6_numbers)|~v3_xxreal_0(esk2_0))).#trainneg
cnf(i_0_788, negated_conjecture, (v2_xxreal_0(esk2_0)|~v2_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_791, plain, (v1_valued_0(k2_zfmisc_1(X1,k1_numbers)))).#trainneg
cnf(i_0_792, plain, (v1_valued_0(k2_zfmisc_1(X1,k5_numbers)))).#trainneg
cnf(i_0_242, plain, (v7_ordinal1(k17_complex1(X1))|~v6_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_247, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_801, plain, (k1_seq_1(k54_valued_1(X1),X2)=k6_numbers|k4_series_1(k54_valued_1(X1))!=k6_numbers|~v1_valued_0(X1)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v2_comseq_3(k54_valued_1(X1))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k2_numbers))).#trainneg
cnf(i_0_789, plain, (k1_relset_1(X1,k6_numbers)=X2|~v1_partfun1(k6_numbers,X2))).#trainneg
cnf(i_0_899, plain, (r1_xxreal_0(k1_funct_1(X1,X2),k8_nat_1(k1_numbers,X1,X2))|~v7_valued_0(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X2,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_803, negated_conjecture, (k5_numbers=X1|~v1_partfun1(esk1_0,X1)|~v4_relat_1(esk1_0,X1))).#trainneg
cnf(i_0_830, plain, (r1_xxreal_0(k1_funct_1(X1,X2),k8_nat_1(k1_numbers,X1,k2_xcmplx_0(X3,X2)))|~v7_valued_0(X1)|~v7_ordinal1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_757, plain, (k1_relset_1(X1,k2_series_1(X2))=k5_numbers|~v1_partfun1(k2_series_1(X2),k5_numbers)|~v4_relat_1(k2_series_1(X2),X1)|~v1_relat_1(k2_series_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X2,k5_numbers,k1_numbers)|~v1_funct_1(X2))).#trainneg
cnf(i_0_772, plain, (v7_valued_0(X1)|~v7_ordinal1(esk6_1(X1))|~r1_xxreal_0(k1_funct_1(X1,esk6_1(X1)),k8_nat_1(k1_numbers,X1,k2_xcmplx_0(esk7_1(X1),esk6_1(X1))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(esk7_1(X1),k5_numbers))).#trainneg
cnf(i_0_554, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_786, negated_conjecture, (v3_xxreal_0(k6_numbers)|v2_xxreal_0(esk2_0)|v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_255, negated_conjecture, (v3_membered(esk1_0)|~v3_membered(k2_zfmisc_1(k5_numbers,k2_numbers)))).#trainneg
cnf(i_0_613, plain, (v2_xxreal_0(X1)|~v2_xxreal_0(k6_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(k6_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_457, negated_conjecture, (v1_partfun1(k54_valued_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_885, plain, (v1_valued_0(k2_zfmisc_1(X1,k2_numbers)))).#trainneg
cnf(i_0_250, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X1,esk3_2(X2,X1)),k1_funct_1(X2,k2_xcmplx_0(esk3_2(X2,X1),np__1)))!=k1_funct_1(X1,k2_xcmplx_0(esk3_2(X2,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(esk3_2(X2,X1),k5_numbers))).#trainneg
cnf(i_0_257, plain, (v1_valued_0(k54_valued_1(X1))|~v1_relat_1(k54_valued_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_805, negated_conjecture, (v1_partfun1(esk1_0,X1)|k5_numbers!=X1|~v4_relat_1(esk1_0,X1))).#trainneg
cnf(i_0_421, plain, (r1_xxreal_0(k2_xcmplx_0(X2,X1),k2_xcmplx_0(X3,X4))|~r1_xxreal_0(X2,X4)|~r1_xxreal_0(X1,X3)|~v1_xreal_0(X4)|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_392, plain, (v1_partfun1(k3_series_1(X1),k5_numbers)|~v1_funct_2(k3_series_1(X1),k5_numbers,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_505, plain, (k2_xcmplx_0(k1_funct_1(X1,X2),k1_funct_1(X3,k2_nat_1(np__1,X2)))=k1_funct_1(X1,k2_nat_1(np__1,X2))|X1!=k2_series_1(X3)|~v1_valued_0(X3)|~v1_valued_0(X1)|~v1_partfun1(X3,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X3)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_588, plain, (v1_funct_2(k54_valued_1(X3),X1,k1_numbers)|~v1_membered(X2)|~v1_partfun1(k54_valued_1(X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_502, plain, (r1_xxreal_0(k1_funct_1(X1,k2_nat_1(X2,np__1)),k2_xcmplx_0(X4,X5))|~r1_xxreal_0(k1_funct_1(X3,k2_nat_1(X2,np__1)),X5)|~r1_xxreal_0(k1_funct_1(X1,X2),X4)|~v1_xreal_0(X5)|~v1_xreal_0(k1_funct_1(X3,k2_nat_1(X2,np__1)))|~v1_xreal_0(X4)|~v1_xreal_0(k1_funct_1(X1,X2))|X1!=k2_series_1(X3)|~v1_valued_0(X3)|~v1_valued_0(X1)|~v1_partfun1(X3,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X3,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X3)|~v1_relat_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_funct_1(X3)|~v1_funct_1(X1))).#trainneg
cnf(i_0_494, plain, (r1_xxreal_0(k8_nat_1(k1_numbers,X1,X2),k8_nat_1(k1_numbers,X1,k2_xcmplx_0(X2,X3)))|~v7_valued_0(X1)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_966, plain, (r1_xxreal_0(k1_funct_1(X1,X2),k8_nat_1(k1_numbers,X1,X2))|~v7_valued_0(X1)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_940, plain, (v1_membered(k6_numbers))).#trainneg
cnf(i_0_939, plain, (v1_valued_0(k2_zfmisc_1(X1,k6_numbers)))).#trainneg
cnf(i_0_579, plain, (v2_xxreal_0(X1)|~v2_xxreal_0(k6_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(k6_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_941, negated_conjecture, (~v1_xboole_0(esk2_0)|~v2_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_1027, negated_conjecture, (v1_funct_2(k54_valued_1(esk1_0),k5_numbers,k1_numbers))).#trainneg
cnf(i_0_883, plain, (r1_xxreal_0(X1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_272, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X2,k2_nat_1(esk3_2(X2,X1),np__1)),k1_funct_1(X1,esk3_2(X2,X1)))!=k1_funct_1(X1,k2_nat_1(esk3_2(X2,X1),np__1))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_xcmplx_0(k1_funct_1(X1,esk3_2(X2,X1)))|~v1_xcmplx_0(k1_funct_1(X2,k2_nat_1(esk3_2(X2,X1),np__1))))).#trainneg
cnf(i_0_926, plain, (k1_seq_1(k54_valued_1(X1),X2)=k6_numbers|k4_series_1(k54_valued_1(X1))!=k6_numbers|~v1_valued_0(k54_valued_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v2_comseq_3(k54_valued_1(X1))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k2_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_555, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_760, plain, (m1_subset_1(k1_funct_1(X1,k6_numbers),X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(X1,k5_numbers,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_583, plain, (k54_valued_1(k54_valued_1(X3))=k54_valued_1(X3)|~v1_membered(X2)|~m1_subset_1(k54_valued_1(X3),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_610, plain, (v3_xxreal_0(k6_numbers)|~v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(k6_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_802, negated_conjecture, (k9_xtuple_0(esk1_0)=k5_numbers|~v4_relat_1(esk1_0,X1))).#trainneg
cnf(i_0_576, plain, (v3_xxreal_0(k6_numbers)|~v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(k6_numbers)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_558, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_566, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_280, plain, (v1_partfun1(X1,X2)|k9_xtuple_0(X1)!=X2|~v4_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_781, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_342, plain, (X1=k9_xtuple_0(X2)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_partfun1(X2,X1))).#trainneg
cnf(i_0_632, plain, (r1_xxreal_0(k1_funct_1(X1,X2),k1_funct_1(X1,k2_nat_1(X2,X3)))|~v7_valued_0(X1)|~v7_ordinal1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v7_ordinal1(k2_nat_1(X2,X3)))).#trainneg
cnf(i_0_449, plain, (v2_comseq_3(k55_valued_1(k5_numbers,k2_numbers,X1))|~v1_series_1(k55_valued_1(k5_numbers,k2_numbers,X1))|~m1_subset_1(k55_valued_1(k5_numbers,k2_numbers,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_2(k55_valued_1(k5_numbers,k2_numbers,X1),k5_numbers,k2_numbers)|~v1_funct_1(k55_valued_1(k5_numbers,k2_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_1(X1))).#trainneg
cnf(i_0_556, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_719, plain, (k1_seq_1(k6_numbers,X1)=k6_numbers|m2_subset_1(esk4_1(k6_numbers),k1_numbers,k5_numbers)|k4_series_1(k6_numbers)!=k6_numbers|~v1_series_1(k6_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k6_numbers,k5_numbers,k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_286, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X1,esk3_2(X2,X1)),k1_funct_1(X2,k2_nat_1(np__1,esk3_2(X2,X1))))!=k1_funct_1(X1,k2_nat_1(np__1,esk3_2(X2,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v7_ordinal1(esk3_2(X2,X1)))).#trainneg
cnf(i_0_281, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),X2)=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1173, plain, (X1=k9_xtuple_0(k6_numbers)|~v1_partfun1(k6_numbers,X1))).#trainneg
cnf(i_0_1076, negated_conjecture, (r1_xxreal_0(esk2_0,esk2_0))).#trainneg
cnf(i_0_1114, negated_conjecture, (k9_xtuple_0(esk1_0)=k5_numbers)).#trainneg
cnf(i_0_1080, plain, (r1_xxreal_0(X1,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1074, plain, (k1_funct_1(k54_valued_1(X1),X2)=k6_numbers|k4_series_1(k54_valued_1(X1))!=k6_numbers|~v3_valued_0(k54_valued_1(X1))|~v1_valued_0(k54_valued_1(X1))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v2_comseq_3(k54_valued_1(X1))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_2(k54_valued_1(X1),k5_numbers,k2_numbers)|~v1_funct_1(k54_valued_1(X1))|~v1_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1171, plain, (k9_xtuple_0(k2_series_1(X1))=k5_numbers|~v1_partfun1(k2_series_1(X1),k5_numbers)|~v1_relat_1(k2_series_1(X1))|~v1_valued_0(X1)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1172, plain, (k9_xtuple_0(k2_series_1(X1))=k5_numbers|~v1_partfun1(k2_series_1(X1),k5_numbers)|~v1_relat_1(k2_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1174, plain, (k9_xtuple_0(k55_valued_1(X1,X2,X3))=X1|~v1_partfun1(k55_valued_1(X1,X2,X3),X1)|~v1_relat_1(k55_valued_1(X1,X2,X3))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1081, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_284, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1122, plain, (v1_partfun1(k6_numbers,X1)|k9_xtuple_0(k6_numbers)!=X1)).#trainneg
cnf(i_0_313, plain, (v3_xxreal_0(k6_numbers)|~v3_xxreal_0(k17_complex1(X1))|~v1_xreal_0(k17_complex1(X1))|~v1_xreal_0(k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_611, plain, (v3_xxreal_0(k6_numbers)|v2_xxreal_0(X1)|v1_xboole_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(k6_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_385, plain, (r2_hidden(X1,k5_numbers)|~v1_int_1(X1)|~v1_int_1(k6_numbers)|~r1_xxreal_0(k6_numbers,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_288, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X1,esk3_2(X2,X1)),k1_funct_1(X2,k2_nat_1(np__1,esk3_2(X2,X1))))!=k1_funct_1(X1,k2_nat_1(np__1,esk3_2(X2,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v7_ordinal1(np__1)|~m1_subset_1(esk3_2(X2,X1),k5_numbers))).#trainneg
cnf(i_0_1189, plain, (v2_comseq_3(k55_valued_1(k5_numbers,k2_numbers,X1))|~m1_subset_1(k55_valued_1(k5_numbers,k2_numbers,X1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_2(k55_valued_1(k5_numbers,k2_numbers,X1),k5_numbers,k2_numbers)|~v1_funct_1(k55_valued_1(k5_numbers,k2_numbers,X1))|~v1_funct_1(X1)|~v2_comseq_3(X1)|~v1_funct_2(X1,k5_numbers,k2_numbers))).#trainneg
cnf(i_0_404, plain, (v1_xcmplx_0(k8_nat_1(X1,X2,X3))|~v1_membered(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~v1_funct_2(X2,k5_numbers,X1)|~v1_funct_1(X2))).#trainneg
cnf(i_0_293, plain, (v1_xreal_0(k4_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_300, plain, (v1_xreal_0(k2_seq_2(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_560, plain, (k1_seq_1(k2_zfmisc_1(k5_numbers,k1_numbers),X1)=k6_numbers|m2_subset_1(esk4_1(k2_zfmisc_1(k5_numbers,k1_numbers)),k1_numbers,k5_numbers)|k4_series_1(k2_zfmisc_1(k5_numbers,k1_numbers))!=k6_numbers|~v1_series_1(k2_zfmisc_1(k5_numbers,k1_numbers))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_funct_2(k2_zfmisc_1(k5_numbers,k1_numbers),k5_numbers,k1_numbers)|~v1_funct_1(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_1083, plain, (X1=k2_series_1(X2)|k2_xcmplx_0(k1_funct_1(X2,k2_nat_1(np__1,esk3_2(X2,X1))),k1_funct_1(X1,esk3_2(X2,X1)))!=k1_funct_1(X1,k2_nat_1(np__1,esk3_2(X2,X1)))|k1_funct_1(X1,k6_numbers)!=k1_funct_1(X2,k6_numbers)|~v1_valued_0(X2)|~v1_valued_0(X1)|~v1_partfun1(X2,k5_numbers)|~v1_partfun1(X1,k5_numbers)|~v4_relat_1(X2,k5_numbers)|~v4_relat_1(X1,k5_numbers)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(esk3_2(X2,X1),k5_numbers))).#trainneg
cnf(i_0_290, plain, (k3_real_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_644, plain, (k6_xcmplx_0(X1,k6_numbers)=k2_xcmplx_0(X1,k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_590, negated_conjecture, (m1_subset_1(k1_funct_1(esk1_0,X1),k2_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_310, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_294, plain, (v1_xcmplx_0(k4_series_1(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_291, plain, (k3_real_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_295, plain, (v7_ordinal1(k4_series_1(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_356, plain, (v1_valued_0(k3_series_1(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_823, negated_conjecture, (v7_ordinal1(k1_funct_1(X1,esk2_0))|~v6_membered(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~v1_funct_2(X1,k5_numbers,X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_782, plain, (m2_subset_1(k2_xcmplx_0(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_637, plain, (v1_valued_0(k2_series_1(X1))|~v1_relat_1(k2_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_584, plain, (k54_valued_1(k55_valued_1(X1,k1_numbers,X2))=k55_valued_1(X1,k1_numbers,X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1340, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1217, negated_conjecture, (v3_xxreal_0(esk2_0)|v2_xxreal_0(esk2_0)|v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_299, plain, (v2_xxreal_0(k17_complex1(X1))|~v2_xxreal_0(k6_numbers)|~v1_xreal_0(k17_complex1(X1))|~v1_xreal_0(k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1259, plain, (m1_subset_1(X1,k5_numbers)|~v1_int_1(k6_numbers)|~v1_int_1(X1)|~v1_xcmplx_0(X1)|~r1_xxreal_0(k6_numbers,X1))).#trainneg
cnf(i_0_1245, plain, (v1_partfun1(k6_numbers,k9_xtuple_0(k6_numbers)))).#trainneg
cnf(i_0_1321, negated_conjecture, (v7_ordinal1(k1_funct_1(esk1_0,esk2_0))|~v6_membered(k2_numbers))).#trainneg
cnf(i_0_1238, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X1)|v1_xboole_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_301, plain, (v1_xcmplx_0(k2_seq_2(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_302, plain, (v7_ordinal1(k2_seq_2(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1233, plain, (k9_xtuple_0(k55_valued_1(X1,X2,X3))=X1|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3)|~v1_funct_2(k55_valued_1(X1,X2,X3),X1,k1_numbers))).#trainneg
cnf(i_0_1286, plain, (k3_real_1(X2,X1)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1285, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1373, plain, (k1_seq_1(k55_valued_1(X1,k1_numbers,X2),X3)=k6_numbers|k4_series_1(k55_valued_1(X1,k1_numbers,X2))!=k6_numbers|~v1_series_1(k55_valued_1(X1,k1_numbers,X2))|~v1_valued_0(k55_valued_1(X1,k1_numbers,X2))|~v1_relat_1(k55_valued_1(X1,k1_numbers,X2))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(k55_valued_1(X1,k1_numbers,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(k55_valued_1(X1,k1_numbers,X2),k5_numbers,k1_numbers)|~v1_funct_1(k55_valued_1(X1,k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_1228, negated_conjecture, (k1_funct_1(k54_valued_1(esk1_0),X1)=k6_numbers|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v2_comseq_3(k54_valued_1(esk1_0))|~m1_subset_1(k54_valued_1(esk1_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k54_valued_1(esk1_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_2(k54_valued_1(esk1_0),k5_numbers,k2_numbers))).#trainneg
cnf(i_0_1333, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_357, plain, (v3_valued_0(k3_series_1(X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1))).#trainneg
cnf(i_0_1370, plain, (k1_funct_1(k55_valued_1(X1,k1_numbers,X2),X3)=k6_numbers|k4_series_1(k55_valued_1(X1,k1_numbers,X2))!=k6_numbers|~v1_valued_0(k55_valued_1(X1,k1_numbers,X2))|~v1_relat_1(k55_valued_1(X1,k1_numbers,X2))|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v2_comseq_3(k55_valued_1(X1,k1_numbers,X2))|~m1_subset_1(k55_valued_1(X1,k1_numbers,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(k55_valued_1(X1,k1_numbers,X2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k2_numbers)))|~v1_funct_2(k55_valued_1(X1,k1_numbers,X2),k5_numbers,k2_numbers)|~v1_funct_1(k55_valued_1(X1,k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_314, plain, (k7_real_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1421, negated_conjecture, (v4_relat_1(k54_valued_1(esk1_0),k5_numbers))).#trainneg
cnf(i_0_1420, negated_conjecture, (v1_relat_1(k54_valued_1(esk1_0)))).#trainneg
cnf(i_0_1429, negated_conjecture, (v1_valued_0(k54_valued_1(esk1_0)))).#trainneg
cnf(i_0_1294, negated_conjecture, (v1_xboole_0(esk1_0)|m1_subset_1(X1,k2_zfmisc_1(k5_numbers,k2_numbers))|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_315, plain, (k7_real_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_316, plain, (v1_xreal_0(k3_real_1(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1413, plain, (k9_xtuple_0(k54_valued_1(X3))=X1|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(k54_valued_1(X3),X1,k1_numbers)|~v1_funct_1(X3))).#trainneg
cnf(i_0_1392, plain, (k1_relset_1(X1,k6_numbers)=k9_xtuple_0(k6_numbers))).#trainneg
cnf(i_0_1431, negated_conjecture, (v5_relat_1(k54_valued_1(esk1_0),k1_numbers))).#trainneg
cnf(i_0_317, plain, (v1_xcmplx_0(k3_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1246, plain, (v1_partfun1(k6_numbers,X1)|k1_relset_1(X2,k6_numbers)!=X1)).#trainneg
cnf(i_0_1417, negated_conjecture, (v1_xboole_0(k54_valued_1(esk1_0))|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_1418, negated_conjecture, (v3_membered(k54_valued_1(esk1_0))|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_318, plain, (v7_ordinal1(k3_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1521, plain, (k7_real_1(X2,X1)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1382, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1461, plain, (r1_xxreal_0(k1_funct_1(X1,X2),k8_nat_1(k1_numbers,X1,k2_xcmplx_0(X2,X3)))|~v7_valued_0(X1)|~v7_ordinal1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~m1_subset_1(X3,k5_numbers)|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_344, plain, (k8_real_1(k17_complex1(X1),k17_complex1(X2))=k17_complex1(k3_xcmplx_0(X1,k17_complex1(X2)))|~v1_xcmplx_0(k17_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_463, plain, (v7_valued_0(X1)|~r1_xxreal_0(k8_nat_1(k1_numbers,X1,esk6_1(X1)),k1_funct_1(X1,k2_nat_1(esk6_1(X1),esk7_1(X1))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,k1_numbers)))|~v1_funct_2(X1,k5_numbers,k1_numbers)|~v1_funct_1(X1)|~v7_ordinal1(k2_nat_1(esk6_1(X1),esk7_1(X1))))).#trainneg
cnf(i_0_1391, plain, (v1_funct_2(k6_numbers,k9_xtuple_0(k6_numbers),X1))).#trainneg
cnf(i_0_1380, plain, (k54_valued_1(k54_valued_1(X2))=k54_valued_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v1_funct_1(X2))).#trainneg
cnf(i_0_320, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1658, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1427, negated_conjecture, (v1_funct_1(k54_valued_1(k54_valued_1(esk1_0))))).#trainneg
cnf(i_0_1580, negated_conjecture, (v1_xcmplx_0(k54_valued_1(esk1_0))|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_445, plain, (v3_valued_0(k55_valued_1(X1,X2,X3))|~v1_membered(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X3))).#trainneg
cnf(i_0_675, plain, (r1_xxreal_0(X1,X2)|~r1_xxreal_0(k6_numbers,X2)|~r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_321, plain, (m1_subset_1(k3_real_1(X2,X1),k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_322, plain, (v1_xreal_0(k7_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1599, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1623, plain, (k17_complex1(k3_xcmplx_0(X1,k17_complex1(X2)))=k17_complex1(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1581, negated_conjecture, (k54_valued_1(esk1_0)=k6_numbers|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_1554, plain, (k9_xtuple_0(k54_valued_1(k6_numbers))=X1|~v1_membered(X2)|~v1_funct_2(k54_valued_1(k6_numbers),X1,k1_numbers)|~v1_funct_1(k6_numbers))).#trainneg
cnf(i_0_323, plain, (v1_xcmplx_0(k7_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1583, negated_conjecture, (v1_membered(k54_valued_1(esk1_0))|~v3_membered(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_1774, negated_conjecture, (v1_funct_1(k6_numbers)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_1779, negated_conjecture, (v1_series_1(k6_numbers)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_1806, negated_conjecture, (k4_series_1(k6_numbers)=k6_numbers|~v1_xboole_0(k2_zfmisc_1(k5_numbers,k1_numbers)))).#trainneg
cnf(i_0_324, plain, (v7_ordinal1(k7_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 129
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 180
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 177
# Processed clauses                    : 625
# ...of these trivial                  : 22
# ...subsumed                          : 150
# ...remaining for further processing  : 453
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 5
# Backward-rewritten                   : 23
# Generated clauses                    : 1661
# ...of the previous two non-trivial   : 1488
# Contextual simplify-reflections      : 80
# Paramodulations                      : 1654
# Factorizations                       : 0
# Equation resolutions                 : 7
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 424
#    Positive orientable unit clauses  : 85
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 7
#    Non-unit-clauses                  : 331
# Current number of unprocessed clauses: 995
# ...number of literals in the above   : 8050
# Current number of archived formulas  : 0
# Current number of archived clauses   : 481
# Proof object given clauses           : 52
# Proof search given clauses           : 453
# Clause-clause subsumption calls (NU) : 59796
# Rec. Clause-clause subsumption calls : 8516
# Non-unit clause-clause subsumptions  : 232
# Unit Clause-clause subsumption calls : 3947
# Rewrite failures with RHS unbound    : 3
# BW rewrite match attempts            : 30
# BW rewrite match successes           : 14
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 56503

# -------------------------------------------------
# User time                : 8.843 s
# System time              : 0.561 s
# Total time               : 9.404 s
# Maximum resident set size: 352068 pages
8.84user 0.62system 0:09.50elapsed 99%CPU (0avgtext+0avgdata 352660maxresident)k
0inputs+208outputs (0major+130309minor)pagefaults 0swaps
