2020-05-31 17:20:33.496842: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_69
2020-05-31 17:20:33.583016: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-05-31 17:20:33.781632: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-05-31 17:20:34.577605: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-05-31 17:20:37.855180: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 4358313 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_69' loaded (query=128; context=256; weigths=real; len_mult=0.000000).
# Preprocessing time       : 5.639 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t9_fdiff_1, axiom, ![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(r2_fdiff_1(X2,X1)<=>(r1_tarski(X1,k1_relset_1(k1_numbers,X2))&![X3]:(m1_subset_1(X3,k1_numbers)=>(r2_hidden(X3,X1)=>r1_fdiff_1(X2,X3))))))), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', t9_fdiff_1)).
fof(t13_fdiff_2, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>((r1_fdiff_1(X2,X1)&r1_fdiff_1(X3,k1_seq_1(X2,X1)))=>(r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),X1)&k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),X1)=k8_real_1(k1_fdiff_1(X3,k1_seq_1(X2,X1)),k1_fdiff_1(X2,X1))))))), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', t13_fdiff_2)).
fof(dt_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:((((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_funct_1(X6))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>(v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))&m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4))))), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', dt_k1_partfun1)).
fof(t64_sin_cos, axiom, ![X1]:(v1_xreal_0(X1)=>(r1_fdiff_1(k16_sin_cos,X1)&k1_fdiff_1(k16_sin_cos,X1)=k1_seq_1(k19_sin_cos,X1))), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', t64_sin_cos)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', t4_subset)).
fof(dt_k16_sin_cos, axiom, ((v1_funct_1(k16_sin_cos)&v1_funct_2(k16_sin_cos,k1_numbers,k1_numbers))&m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', dt_k16_sin_cos)).
fof(cc1_xreal_0, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>v1_xreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', cc1_xreal_0)).
fof(dt_k1_seq_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>m1_subset_1(k1_seq_1(X1,X2),k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', dt_k1_seq_1)).
fof(t13_fdiff_11, conjecture, ![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>((r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))&r1_tarski(X1,k2_rcomp_1(k1_real_1(np__1),np__1)))=>(r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),X1)&![X2]:(m1_subset_1(X2,k1_numbers)=>(r2_hidden(X2,X1)=>k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),X1),X2)=k10_real_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,X2)),k7_real_1(np__1,k5_square_1(X2)))))))), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', t13_fdiff_11)).
fof(d9_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>![X2]:(v1_xcmplx_0(X2)=>k7_xcmplx_0(X1,X2)=k3_xcmplx_0(X1,k5_xcmplx_0(X2)))), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', d9_xcmplx_0)).
fof(spc3_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k7_xcmplx_0(np__1,X1)=k5_xcmplx_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', spc3_arithm)).
fof(t81_sin_cos9, axiom, ![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>(r1_tarski(X1,k2_rcomp_1(k1_real_1(np__1),np__1))=>(r2_fdiff_1(k1_sin_cos9,X1)&![X2]:(m1_subset_1(X2,k1_numbers)=>(r2_hidden(X2,X1)=>k1_seq_1(k2_fdiff_1(k1_sin_cos9,X1),X2)=k10_real_1(np__1,k7_real_1(np__1,k5_square_1(X2)))))))), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', t81_sin_cos9)).
fof(redefinition_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:((((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_funct_1(X6))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>k1_partfun1(X1,X2,X3,X4,X5,X6)=k3_relat_1(X5,X6)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', redefinition_k1_partfun1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', cc1_relset_1)).
fof(cc18_valued_0, axiom, ![X1, X2]:(v3_membered(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v3_valued_0(X3))), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', cc18_valued_0)).
fof(commutativity_k3_xcmplx_0, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', commutativity_k3_xcmplx_0)).
fof(dt_k5_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>v1_xcmplx_0(k5_xcmplx_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', dt_k5_xcmplx_0)).
fof(dt_k1_sin_cos9, axiom, (v1_funct_1(k1_sin_cos9)&m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', dt_k1_sin_cos9)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', fc3_membered)).
fof(redefinition_k10_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', redefinition_k10_real_1)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', spc1_numerals)).
fof(cc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xcmplx_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', cc3_xreal_0)).
fof(involutiveness_k5_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>k5_xcmplx_0(k5_xcmplx_0(X1))=X1), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', involutiveness_k5_xcmplx_0)).
fof(d7_fdiff_1, axiom, ![X1]:((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>![X2]:(r2_fdiff_1(X1,X2)=>![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(X3=k2_fdiff_1(X1,X2)<=>(k1_relset_1(k1_numbers,X3)=X2&![X4]:(m1_subset_1(X4,k1_numbers)=>(r2_hidden(X4,X2)=>k1_seq_1(X3,X4)=k1_fdiff_1(X1,X4)))))))), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', d7_fdiff_1)).
fof(dt_k10_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>m1_subset_1(k10_real_1(X1,X2),k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', dt_k10_real_1)).
fof(cc1_xcmplx_0, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>v1_xcmplx_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', cc1_xcmplx_0)).
fof(dt_k2_fdiff_1, axiom, ![X1, X2]:((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(v1_funct_1(k2_fdiff_1(X1,X2))&m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', dt_k2_fdiff_1)).
fof(commutativity_k8_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k8_real_1(X1,X2)=k8_real_1(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', commutativity_k8_real_1)).
fof(redefinition_k8_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', redefinition_k8_real_1)).
fof(fc4_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>(v1_xcmplx_0(k5_xcmplx_0(X1))&v1_xreal_0(k5_xcmplx_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', fc4_xreal_0)).
fof(dt_k1_fdiff_1, axiom, ![X1, X2]:(((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))&v1_xreal_0(X2))=>m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', dt_k1_fdiff_1)).
fof(cc14_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xreal_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', cc14_membered)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', t3_subset)).
fof(cc17_membered, axiom, ![X1]:(v6_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v7_ordinal1(X2))), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', cc17_membered)).
fof(fc6_membered, axiom, v6_membered(k4_ordinal1), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', fc6_membered)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', redefinition_k5_numbers)).
fof(cc2_int_1, axiom, ![X1]:(v7_ordinal1(X1)=>v1_int_1(X1)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', cc2_int_1)).
fof(dt_k7_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>m1_subset_1(k7_real_1(X1,X2),k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', dt_k7_real_1)).
fof(cc8_membered, axiom, ![X1]:(m1_subset_1(X1,k1_zfmisc_1(k1_numbers))=>v3_membered(X1)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', cc8_membered)).
fof(cc2_xxreal_0, axiom, ![X1]:(v7_ordinal1(X1)=>v1_xxreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', cc2_xxreal_0)).
fof(fc3_int_1, axiom, ![X1]:(v1_int_1(X1)=>(v1_xcmplx_0(k4_xcmplx_0(X1))&v1_int_1(k4_xcmplx_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', fc3_int_1)).
fof(involutiveness_k4_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>k4_xcmplx_0(k4_xcmplx_0(X1))=X1), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', involutiveness_k4_xcmplx_0)).
fof(dt_k2_rcomp_1, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>m1_subset_1(k2_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', dt_k2_rcomp_1)).
fof(redefinition_k1_real_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>k1_real_1(X1)=k4_xcmplx_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', redefinition_k1_real_1)).
fof(cc3_int_1, axiom, ![X1]:(v1_int_1(X1)=>v1_xreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', cc3_int_1)).
fof(rqRealNeg__k4_xcmplx_0__rm1_r1, axiom, k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1, file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', rqRealNeg__k4_xcmplx_0__rm1_r1)).
fof(dt_k5_square_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>m1_subset_1(k5_square_1(X1),k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', dt_k5_square_1)).
fof(cc4_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xxreal_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t13_fdiff_11', cc4_xreal_0)).
fof(c_0_48, plain, ![X19, X20, X21]:(((r1_tarski(X19,k1_relset_1(k1_numbers,X20))|~r2_fdiff_1(X20,X19)|(~v1_funct_1(X20)|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X19)|~m1_subset_1(X19,k1_zfmisc_1(k1_numbers))))&(~m1_subset_1(X21,k1_numbers)|(~r2_hidden(X21,X19)|r1_fdiff_1(X20,X21))|~r2_fdiff_1(X20,X19)|(~v1_funct_1(X20)|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X19)|~m1_subset_1(X19,k1_zfmisc_1(k1_numbers)))))&((m1_subset_1(esk4_2(X19,X20),k1_numbers)|~r1_tarski(X19,k1_relset_1(k1_numbers,X20))|r2_fdiff_1(X20,X19)|(~v1_funct_1(X20)|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X19)|~m1_subset_1(X19,k1_zfmisc_1(k1_numbers))))&((r2_hidden(esk4_2(X19,X20),X19)|~r1_tarski(X19,k1_relset_1(k1_numbers,X20))|r2_fdiff_1(X20,X19)|(~v1_funct_1(X20)|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X19)|~m1_subset_1(X19,k1_zfmisc_1(k1_numbers))))&(~r1_fdiff_1(X20,esk4_2(X19,X20))|~r1_tarski(X19,k1_relset_1(k1_numbers,X20))|r2_fdiff_1(X20,X19)|(~v1_funct_1(X20)|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X19)|~m1_subset_1(X19,k1_zfmisc_1(k1_numbers))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_fdiff_1])])])])])).
fof(c_0_49, plain, ![X9, X10, X11]:((r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X10,X11),X9)|(~r1_fdiff_1(X10,X9)|~r1_fdiff_1(X11,k1_seq_1(X10,X9)))|(~v1_funct_1(X11)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v1_funct_1(X10)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~m1_subset_1(X9,k1_numbers))&(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X10,X11),X9)=k8_real_1(k1_fdiff_1(X11,k1_seq_1(X10,X9)),k1_fdiff_1(X10,X9))|(~r1_fdiff_1(X10,X9)|~r1_fdiff_1(X11,k1_seq_1(X10,X9)))|(~v1_funct_1(X11)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v1_funct_1(X10)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~m1_subset_1(X9,k1_numbers))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_fdiff_2])])])])).
cnf(c_0_50, plain, (r2_fdiff_1(X1,X2)|~r1_fdiff_1(X1,esk4_2(X2,X1))|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_51, plain, (r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~r1_fdiff_1(X1,X3)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_52, plain, ![X28, X29, X30, X31, X32, X33]:((v1_funct_1(k1_partfun1(X28,X29,X30,X31,X32,X33))|(~v1_funct_1(X32)|~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(X28,X29)))|~v1_funct_1(X33)|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X30,X31)))))&(m1_subset_1(k1_partfun1(X28,X29,X30,X31,X32,X33),k1_zfmisc_1(k2_zfmisc_1(X28,X31)))|(~v1_funct_1(X32)|~m1_subset_1(X32,k1_zfmisc_1(k2_zfmisc_1(X28,X29)))|~v1_funct_1(X33)|~m1_subset_1(X33,k1_zfmisc_1(k2_zfmisc_1(X30,X31)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_partfun1])])])).
cnf(c_0_53, plain, (r2_fdiff_1(X1,X2)|~v3_rcomp_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r1_fdiff_1(X1,esk4_2(X2,X1))|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_50])).
cnf(c_0_54, plain, (r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X3)|~m1_subset_1(X3,k1_numbers)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_51])).
cnf(c_0_55, plain, (m1_subset_1(esk4_2(X1,X2),k1_numbers)|r2_fdiff_1(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_56, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_57, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_52])).
fof(c_0_58, plain, ![X25]:((r1_fdiff_1(k16_sin_cos,X25)|~v1_xreal_0(X25))&(k1_fdiff_1(k16_sin_cos,X25)=k1_seq_1(k19_sin_cos,X25)|~v1_xreal_0(X25))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t64_sin_cos])])])).
fof(c_0_59, plain, ![X34, X35, X36]:(~r2_hidden(X34,X35)|~m1_subset_1(X35,k1_zfmisc_1(X36))|m1_subset_1(X34,X36)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_60, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))))|~r1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_54])])).
cnf(c_0_61, plain, (r2_fdiff_1(X2,X1)|m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_55])).
cnf(c_0_62, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_56])).
cnf(c_0_63, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_57])).
cnf(c_0_64, plain, (r1_fdiff_1(k16_sin_cos,X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_65, plain, (v1_funct_1(k16_sin_cos)), inference(split_conjunct,[status(thm)],[dt_k16_sin_cos])).
cnf(c_0_66, plain, (m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k16_sin_cos])).
fof(c_0_67, plain, ![X41]:(~m1_subset_1(X41,k1_numbers)|v1_xreal_0(X41)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_xreal_0])])).
cnf(c_0_68, plain, (r1_fdiff_1(X3,X1)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,X2)|~r2_fdiff_1(X3,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_69, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_70, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))))|~r1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_60, c_0_61]), c_0_62]), c_0_63])).
cnf(c_0_71, plain, (r1_fdiff_1(k16_sin_cos,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_64])).
cnf(c_0_72, plain, (v1_funct_1(k16_sin_cos)), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_73, plain, (m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_66])).
cnf(c_0_74, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_75, plain, (r1_fdiff_1(X3,X1)|~v3_rcomp_1(X2)|~v1_funct_1(X3)|~r2_fdiff_1(X3,X2)|~r2_hidden(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_68])).
cnf(c_0_76, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_69])).
cnf(c_0_77, plain, (r2_hidden(esk4_2(X1,X2),X1)|r2_fdiff_1(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_78, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~r1_fdiff_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_xreal_0(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_71]), c_0_72]), c_0_73])])])).
cnf(c_0_79, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_74])).
cnf(c_0_80, plain, (r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~r2_hidden(X2,X3)|~r2_fdiff_1(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)), inference(csr,[status(thm)],[c_0_75, c_0_76])).
cnf(c_0_81, plain, (r2_fdiff_1(X2,X1)|r2_hidden(esk4_2(X1,X2),X1)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_77])).
fof(c_0_82, plain, ![X94, X95]:(~v1_relat_1(X94)|~v1_funct_1(X94)|~v3_valued_0(X94)|m1_subset_1(k1_seq_1(X94,X95),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_seq_1])])).
fof(c_0_83, negated_conjecture, ~(![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>((r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))&r1_tarski(X1,k2_rcomp_1(k1_real_1(np__1),np__1)))=>(r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),X1)&![X2]:(m1_subset_1(X2,k1_numbers)=>(r2_hidden(X2,X1)=>k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),X1),X2)=k10_real_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,X2)),k7_real_1(np__1,k5_square_1(X2))))))))), inference(assume_negation,[status(cth)],[t13_fdiff_11])).
fof(c_0_84, plain, ![X154, X155]:(~v1_xcmplx_0(X154)|(~v1_xcmplx_0(X155)|k7_xcmplx_0(X154,X155)=k3_xcmplx_0(X154,k5_xcmplx_0(X155)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_xcmplx_0])])])).
fof(c_0_85, plain, ![X60]:(~v1_xcmplx_0(X60)|k7_xcmplx_0(np__1,X60)=k5_xcmplx_0(X60)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc3_arithm])])).
fof(c_0_86, plain, ![X17, X18]:((r2_fdiff_1(k1_sin_cos9,X17)|~r1_tarski(X17,k2_rcomp_1(k1_real_1(np__1),np__1))|(~v3_rcomp_1(X17)|~m1_subset_1(X17,k1_zfmisc_1(k1_numbers))))&(~m1_subset_1(X18,k1_numbers)|(~r2_hidden(X18,X17)|k1_seq_1(k2_fdiff_1(k1_sin_cos9,X17),X18)=k10_real_1(np__1,k7_real_1(np__1,k5_square_1(X18))))|~r1_tarski(X17,k2_rcomp_1(k1_real_1(np__1),np__1))|(~v3_rcomp_1(X17)|~m1_subset_1(X17,k1_zfmisc_1(k1_numbers))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t81_sin_cos9])])])])).
fof(c_0_87, plain, ![X45, X46, X47, X48, X49, X50]:(~v1_funct_1(X49)|~m1_subset_1(X49,k1_zfmisc_1(k2_zfmisc_1(X45,X46)))|~v1_funct_1(X50)|~m1_subset_1(X50,k1_zfmisc_1(k2_zfmisc_1(X47,X48)))|k1_partfun1(X45,X46,X47,X48,X49,X50)=k3_relat_1(X49,X50)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_partfun1])])).
cnf(c_0_88, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~r1_fdiff_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos))),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_78, c_0_79])])).
cnf(c_0_89, plain, (r1_fdiff_1(X1,esk4_2(X2,X3))|r2_fdiff_1(X3,X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_funct_1(X3)|~r1_tarski(X2,k1_relset_1(k1_numbers,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_81])])).
cnf(c_0_90, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_82])).
fof(c_0_91, plain, ![X55, X56, X57]:(~m1_subset_1(X57,k1_zfmisc_1(k2_zfmisc_1(X55,X56)))|v1_relat_1(X57)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
fof(c_0_92, plain, ![X86, X87, X88]:(~v3_membered(X87)|(~m1_subset_1(X88,k1_zfmisc_1(k2_zfmisc_1(X86,X87)))|v3_valued_0(X88))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc18_valued_0])])])).
fof(c_0_93, negated_conjecture, ((v3_rcomp_1(esk1_0)&m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers)))&((r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))&r1_tarski(esk1_0,k2_rcomp_1(k1_real_1(np__1),np__1)))&((m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0))&((r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0))&(k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)!=k10_real_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k7_real_1(np__1,k5_square_1(esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_83])])])])).
fof(c_0_94, plain, ![X107, X108]:(~v1_xcmplx_0(X107)|~v1_xcmplx_0(X108)|k3_xcmplx_0(X107,X108)=k3_xcmplx_0(X108,X107)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k3_xcmplx_0])])).
cnf(c_0_95, plain, (k7_xcmplx_0(X1,X2)=k3_xcmplx_0(X1,k5_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_96, plain, (k7_xcmplx_0(np__1,X1)=k5_xcmplx_0(X1)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
fof(c_0_97, plain, ![X89]:(~v1_xcmplx_0(X89)|v1_xcmplx_0(k5_xcmplx_0(X89))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_xcmplx_0])])).
cnf(c_0_98, plain, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,X2),X1)=k10_real_1(np__1,k7_real_1(np__1,k5_square_1(X1)))|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,X2)|~r1_tarski(X2,k2_rcomp_1(k1_real_1(np__1),np__1))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_99, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_100, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)))|~m1_subset_1(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos))|~r2_fdiff_1(X1,X2)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_89])])).
cnf(c_0_101, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_102, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_103, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_104, plain, (m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k1_sin_cos9])).
cnf(c_0_105, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
cnf(c_0_106, plain, (r2_fdiff_1(k1_sin_cos9,X1)|~r1_tarski(X1,k2_rcomp_1(k1_real_1(np__1),np__1))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_107, negated_conjecture, (r1_tarski(esk1_0,k2_rcomp_1(k1_real_1(np__1),np__1))), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_108, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_109, negated_conjecture, (v3_rcomp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_110, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_111, plain, (k3_xcmplx_0(X1,k5_xcmplx_0(X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_95])).
cnf(c_0_112, plain, (k5_xcmplx_0(X1)=k7_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_96])).
cnf(c_0_113, plain, (v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_97])).
fof(c_0_114, plain, ![X26, X27]:(~m1_subset_1(X26,k1_numbers)|~v1_xreal_0(X27)|k10_real_1(X26,X27)=k7_xcmplx_0(X26,X27)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k10_real_1])])).
cnf(c_0_115, plain, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,X2),X1)=k10_real_1(np__1,k7_real_1(np__1,k5_square_1(X1)))|~v3_rcomp_1(X2)|~r2_hidden(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r1_tarski(X2,k2_rcomp_1(k1_real_1(np__1),np__1))), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_116, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_117, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_99])).
cnf(c_0_118, plain, (v1_funct_1(k1_sin_cos9)), inference(split_conjunct,[status(thm)],[dt_k1_sin_cos9])).
cnf(c_0_119, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v3_valued_0(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101])])).
cnf(c_0_120, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_102])).
cnf(c_0_121, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_122, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_103])).
cnf(c_0_123, plain, (m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_104])).
cnf(c_0_124, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_125, plain, (r2_fdiff_1(k1_sin_cos9,X1)|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k2_rcomp_1(k1_real_1(np__1),np__1))), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_126, negated_conjecture, (r1_tarski(esk1_0,k2_rcomp_1(k1_real_1(np__1),np__1))), inference(evalgc,[status(thm)],[c_0_107])).
cnf(c_0_127, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_108])).
cnf(c_0_128, negated_conjecture, (v3_rcomp_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_109])).
cnf(c_0_129, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_110])).
cnf(c_0_130, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112])])).
cnf(c_0_131, plain, (v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_113])).
cnf(c_0_132, plain, (k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_114])).
cnf(c_0_133, plain, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,X1),X2)=k10_real_1(np__1,k7_real_1(np__1,k5_square_1(X2)))|~r2_hidden(X2,X1)|~r1_tarski(X1,k2_rcomp_1(k1_real_1(np__1),np__1))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)), inference(csr,[status(thm)],[c_0_115, c_0_76])).
cnf(c_0_134, plain, (m1_subset_1(np__1,k1_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_135, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)), inference(evalgc,[status(thm)],[c_0_116])).
cnf(c_0_136, plain, (k1_partfun1(X1,X2,X3,X4,X5,X6)=k1_partfun1(X7,X8,X9,X10,X5,X6)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_117])])).
cnf(c_0_137, plain, (v1_funct_1(k1_sin_cos9)), inference(evalgc,[status(thm)],[c_0_118])).
cnf(c_0_138, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~v3_valued_0(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)), inference(csr,[status(thm)],[c_0_119, c_0_120])).
cnf(c_0_139, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))), inference(evalgc,[status(thm)],[c_0_121])).
cnf(c_0_140, plain, (v3_valued_0(k1_sin_cos9)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_123]), c_0_124])])])).
cnf(c_0_141, negated_conjecture, (r2_fdiff_1(k1_sin_cos9,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_127]), c_0_128])])])).
cnf(c_0_142, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,X1),X2)=k7_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_130])])).
cnf(c_0_143, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_112])])).
cnf(c_0_144, plain, (k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_132])).
cnf(c_0_145, negated_conjecture, (k10_real_1(np__1,k7_real_1(np__1,k5_square_1(X1)))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r2_hidden(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_126]), c_0_127]), c_0_128])])])).
cnf(c_0_146, plain, (m1_subset_1(np__1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_134])).
fof(c_0_147, plain, ![X98]:(~v1_xreal_0(X98)|v1_xcmplx_0(X98)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_xreal_0])])).
cnf(c_0_148, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_136]), c_0_72]), c_0_137]), c_0_73]), c_0_123])])])).
cnf(c_0_149, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_140]), c_0_137]), c_0_141]), c_0_123]), c_0_127]), c_0_128])])])).
cnf(c_0_150, plain, (m1_subset_1(k1_partfun1(X7,X8,X9,X10,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_136])])).
fof(c_0_151, plain, ![X78]:(~v1_xcmplx_0(X78)|k5_xcmplx_0(k5_xcmplx_0(X78))=X78), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k5_xcmplx_0])])).
cnf(c_0_152, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)!=k10_real_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k7_real_1(np__1,k5_square_1(esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_153, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,X1),X2)=k7_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_142, c_0_143])).
cnf(c_0_154, negated_conjecture, (k7_xcmplx_0(np__1,k7_real_1(np__1,k5_square_1(X1)))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~v1_xreal_0(k7_real_1(np__1,k5_square_1(X1)))|~r2_hidden(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_145]), c_0_146])])])).
cnf(c_0_155, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_147])).
cnf(c_0_156, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_149]), c_0_73]), c_0_123])])])).
fof(c_0_157, plain, ![X12, X13, X14, X15]:(((k1_relset_1(k1_numbers,X14)=X13|X14!=k2_fdiff_1(X12,X13)|(~v1_funct_1(X14)|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X12,X13)|(~v1_funct_1(X12)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(~m1_subset_1(X15,k1_numbers)|(~r2_hidden(X15,X13)|k1_seq_1(X14,X15)=k1_fdiff_1(X12,X15))|X14!=k2_fdiff_1(X12,X13)|(~v1_funct_1(X14)|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X12,X13)|(~v1_funct_1(X12)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))))&((m1_subset_1(esk3_3(X12,X13,X14),k1_numbers)|k1_relset_1(k1_numbers,X14)!=X13|X14=k2_fdiff_1(X12,X13)|(~v1_funct_1(X14)|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X12,X13)|(~v1_funct_1(X12)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&((r2_hidden(esk3_3(X12,X13,X14),X13)|k1_relset_1(k1_numbers,X14)!=X13|X14=k2_fdiff_1(X12,X13)|(~v1_funct_1(X14)|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X12,X13)|(~v1_funct_1(X12)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(k1_seq_1(X14,esk3_3(X12,X13,X14))!=k1_fdiff_1(X12,esk3_3(X12,X13,X14))|k1_relset_1(k1_numbers,X14)!=X13|X14=k2_fdiff_1(X12,X13)|(~v1_funct_1(X14)|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X12,X13)|(~v1_funct_1(X12)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_fdiff_1])])])])])).
cnf(c_0_158, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_159, plain, (r2_fdiff_1(k1_partfun1(X3,X4,X5,X6,X1,k16_sin_cos),X2)|~v3_valued_0(X1)|~v1_funct_1(k1_partfun1(X3,X4,X5,X6,X1,k16_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k16_sin_cos)))|~m1_subset_1(k1_partfun1(X3,X4,X5,X6,X1,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_136]), c_0_72]), c_0_73])])])).
cnf(c_0_160, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X6,k1_numbers)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_73]), c_0_72])])])).
cnf(c_0_161, plain, (k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)=k8_real_1(k1_fdiff_1(X2,k1_seq_1(X1,X3)),k1_fdiff_1(X1,X3))|~r1_fdiff_1(X1,X3)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_162, plain, (k1_fdiff_1(k16_sin_cos,X1)=k1_seq_1(k19_sin_cos,X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
fof(c_0_163, plain, ![X42, X43]:(~m1_subset_1(X42,k1_numbers)|~v1_xreal_0(X43)|m1_subset_1(k10_real_1(X42,X43),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k10_real_1])])).
cnf(c_0_164, plain, (k5_xcmplx_0(k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_151])).
cnf(c_0_165, negated_conjecture, (k10_real_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k7_real_1(np__1,k5_square_1(esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)), inference(evalgc,[status(thm)],[c_0_152])).
cnf(c_0_166, negated_conjecture, (k7_xcmplx_0(X2,k7_real_1(np__1,k5_square_1(X1)))=k3_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k7_real_1(np__1,k5_square_1(X1)))|~v1_xreal_0(k7_real_1(np__1,k5_square_1(X1)))|~r2_hidden(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_154])])).
cnf(c_0_167, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_155])).
cnf(c_0_168, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_62]), c_0_72]), c_0_137]), c_0_73]), c_0_123])])])).
fof(c_0_169, plain, ![X93]:(~m1_subset_1(X93,k1_numbers)|v1_xcmplx_0(X93)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_xcmplx_0])])).
cnf(c_0_170, plain, (k1_seq_1(X3,X1)=k1_fdiff_1(X4,X1)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,X2)|X3!=k2_fdiff_1(X4,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X4,X2)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_157])).
fof(c_0_171, plain, ![X23, X24]:((v1_funct_1(k2_fdiff_1(X23,X24))|(~v1_funct_1(X23)|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(m1_subset_1(k2_fdiff_1(X23,X24),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|(~v1_funct_1(X23)|~m1_subset_1(X23,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_fdiff_1])])])).
cnf(c_0_172, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)), inference(evalgc,[status(thm)],[c_0_158])).
cnf(c_0_173, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),X6)|~v3_valued_0(X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos))|~v1_funct_1(X5)|~r2_fdiff_1(X5,X6)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)), inference(csr,[status(thm)],[c_0_159, c_0_160])).
fof(c_0_174, plain, ![X51, X52]:(~m1_subset_1(X51,k1_numbers)|~v1_xreal_0(X52)|k8_real_1(X51,X52)=k8_real_1(X52,X51)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k8_real_1])])).
cnf(c_0_175, plain, (k8_real_1(k1_fdiff_1(X2,k1_seq_1(X1,X3)),k1_fdiff_1(X1,X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X3)|~m1_subset_1(X3,k1_numbers)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_161])).
cnf(c_0_176, plain, (k1_fdiff_1(k16_sin_cos,X1)=k1_seq_1(k19_sin_cos,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_162])).
cnf(c_0_177, plain, (m1_subset_1(k10_real_1(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_163])).
cnf(c_0_178, plain, (k5_xcmplx_0(k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_164])).
cnf(c_0_179, negated_conjecture, (k7_xcmplx_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k7_real_1(np__1,k5_square_1(esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_144])])).
cnf(c_0_180, negated_conjecture, (k7_xcmplx_0(X1,k7_real_1(np__1,k5_square_1(X2)))=k3_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X2),X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(k7_real_1(np__1,k5_square_1(X2)))|~r2_hidden(X2,esk1_0)), inference(csr,[status(thm)],[c_0_166, c_0_167])).
cnf(c_0_181, negated_conjecture, (r2_hidden(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_63]), c_0_72]), c_0_137]), c_0_73]), c_0_123])])])).
cnf(c_0_182, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_169])).
cnf(c_0_183, plain, (k1_fdiff_1(X4,X1)=k1_seq_1(X3,X1)|X3!=k2_fdiff_1(X4,X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~r2_fdiff_1(X4,X2)|~r2_hidden(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_170])).
cnf(c_0_184, plain, (m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_171])).
cnf(c_0_185, plain, (v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_171])).
cnf(c_0_186, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_136]), c_0_72]), c_0_137]), c_0_73]), c_0_123])])])).
cnf(c_0_187, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_139]), c_0_140]), c_0_137]), c_0_141]), c_0_123]), c_0_73]), c_0_127]), c_0_128])])])).
fof(c_0_188, plain, ![X39, X40]:(~m1_subset_1(X39,k1_numbers)|~v1_xreal_0(X40)|k8_real_1(X39,X40)=k3_xcmplx_0(X39,X40)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_real_1])])).
cnf(c_0_189, plain, (k8_real_1(X1,X2)=k8_real_1(X2,X1)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_190, plain, (k8_real_1(k1_seq_1(k19_sin_cos,k1_seq_1(X1,X2)),k1_fdiff_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~r1_fdiff_1(k16_sin_cos,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_176]), c_0_72]), c_0_73])])])).
cnf(c_0_191, plain, (m1_subset_1(k10_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_177])).
cnf(c_0_192, plain, (k7_xcmplx_0(np__1,k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_112])])).
fof(c_0_193, plain, ![X104]:((v1_xcmplx_0(k5_xcmplx_0(X104))|~v1_xreal_0(X104))&(v1_xreal_0(k5_xcmplx_0(X104))|~v1_xreal_0(X104))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc4_xreal_0])])])).
fof(c_0_194, plain, ![X83, X84]:(~v1_funct_1(X83)|~m1_subset_1(X83,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X84)|m1_subset_1(k1_fdiff_1(X83,X84),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_fdiff_1])])).
cnf(c_0_195, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk2_0),k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_numbers)|~v1_xcmplx_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_180]), c_0_181])])])).
cnf(c_0_196, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_182])).
cnf(c_0_197, plain, (k1_seq_1(k2_fdiff_1(X1,X3),X2)=k1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X1,X3))|~r2_hidden(X2,X3)|~r2_fdiff_1(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_183])])).
cnf(c_0_198, plain, (m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_184])).
cnf(c_0_199, plain, (v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_185])).
cnf(c_0_200, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_187]), c_0_73]), c_0_123])])])).
cnf(c_0_201, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_188])).
cnf(c_0_202, plain, (k8_real_1(X1,X2)=k8_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_189])).
cnf(c_0_203, plain, (k8_real_1(k1_seq_1(k19_sin_cos,k1_seq_1(X1,X2)),k1_fdiff_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[c_0_190, c_0_71])).
cnf(c_0_204, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_144])])).
cnf(c_0_205, plain, (k7_xcmplx_0(np__1,k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_192, c_0_131])).
cnf(c_0_206, plain, (v1_xreal_0(k5_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_193])).
fof(c_0_207, plain, ![X37, X38]:(~v3_membered(X37)|(~m1_subset_1(X38,X37)|v1_xreal_0(X38))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc14_membered])])])).
cnf(c_0_208, plain, (m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_194])).
fof(c_0_209, plain, ![X72, X73]:((~m1_subset_1(X72,k1_zfmisc_1(X73))|r1_tarski(X72,X73))&(~r1_tarski(X72,X73)|m1_subset_1(X72,k1_zfmisc_1(X73)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_210, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk2_0),k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_numbers)), inference(csr,[status(thm)],[c_0_195, c_0_196])).
cnf(c_0_211, plain, (k1_seq_1(k2_fdiff_1(X1,X2),X3)=k1_fdiff_1(X1,X3)|~v1_funct_1(X1)|~r2_hidden(X3,X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_197, c_0_198]), c_0_199])).
cnf(c_0_212, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_63]), c_0_72]), c_0_137]), c_0_73]), c_0_123])])])).
cnf(c_0_213, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_201])).
cnf(c_0_214, plain, (k8_real_1(k1_fdiff_1(X1,X2),k1_seq_1(k19_sin_cos,k1_seq_1(X1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~v1_xreal_0(k1_seq_1(k19_sin_cos,k1_seq_1(X1,X2)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_203])])).
cnf(c_0_215, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_205]), c_0_146])])])).
cnf(c_0_216, plain, (v1_xreal_0(k5_xcmplx_0(X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_206])).
cnf(c_0_217, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_207])).
cnf(c_0_218, plain, (m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_208])).
cnf(c_0_219, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_209])).
fof(c_0_220, plain, ![X68, X69]:(~v6_membered(X68)|(~m1_subset_1(X69,X68)|v7_ordinal1(X69))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc17_membered])])])).
cnf(c_0_221, plain, (v6_membered(k4_ordinal1)), inference(split_conjunct,[status(thm)],[fc6_membered])).
cnf(c_0_222, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
cnf(c_0_223, negated_conjecture, (k3_xcmplx_0(k1_fdiff_1(k1_sin_cos9,esk2_0),k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_211]), c_0_137]), c_0_181]), c_0_141]), c_0_123]), c_0_212])])])).
cnf(c_0_224, plain, (k3_xcmplx_0(k1_fdiff_1(X1,X2),k1_seq_1(k19_sin_cos,k1_seq_1(X1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~v1_xreal_0(k1_seq_1(k19_sin_cos,k1_seq_1(X1,X2)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_214])])).
cnf(c_0_225, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_216])])).
cnf(c_0_226, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_217])).
cnf(c_0_227, plain, (m1_subset_1(k1_seq_1(k19_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_176]), c_0_72]), c_0_73])])])).
cnf(c_0_228, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_219])).
fof(c_0_229, plain, ![X74]:(~v7_ordinal1(X74)|v1_int_1(X74)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_int_1])])).
cnf(c_0_230, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_220])).
cnf(c_0_231, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_232, plain, (v6_membered(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_221])).
cnf(c_0_233, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_222])).
cnf(c_0_234, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk2_0)|~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_numbers)|~v1_xreal_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)))|~v1_xreal_0(k1_seq_1(k1_sin_cos9,esk2_0))|~r1_fdiff_1(k1_sin_cos9,esk2_0)|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk2_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_224]), c_0_137]), c_0_123]), c_0_212])])])).
cnf(c_0_235, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_225, c_0_167])).
cnf(c_0_236, plain, (v1_xreal_0(k1_seq_1(k19_sin_cos,X1))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_227]), c_0_124])])])).
fof(c_0_237, plain, ![X112, X113]:(~m1_subset_1(X112,k1_numbers)|~v1_xreal_0(X113)|m1_subset_1(k7_real_1(X112,X113),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_real_1])])).
cnf(c_0_238, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k2_rcomp_1(k1_real_1(np__1),np__1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_126])])).
fof(c_0_239, plain, ![X118]:(~m1_subset_1(X118,k1_zfmisc_1(k1_numbers))|v3_membered(X118)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_membered])])).
fof(c_0_240, plain, ![X161]:(~v7_ordinal1(X161)|v1_xxreal_0(X161)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_xxreal_0])])).
fof(c_0_241, plain, ![X105]:((v1_xcmplx_0(k4_xcmplx_0(X105))|~v1_int_1(X105))&(v1_int_1(k4_xcmplx_0(X105))|~v1_int_1(X105))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc3_int_1])])])).
fof(c_0_242, plain, ![X90]:(~v1_xcmplx_0(X90)|k4_xcmplx_0(k4_xcmplx_0(X90))=X90), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k4_xcmplx_0])])).
cnf(c_0_243, plain, (v1_int_1(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_229])).
cnf(c_0_244, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_230])).
cnf(c_0_245, plain, (m1_subset_1(np__1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_231])).
cnf(c_0_246, plain, (v6_membered(k5_numbers)), inference(rw,[status(thm)],[c_0_232, c_0_233])).
cnf(c_0_247, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk2_0)|~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~v1_xreal_0(k1_seq_1(k1_sin_cos9,esk2_0))|~r1_fdiff_1(k1_sin_cos9,esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk2_0),k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_234, c_0_235]), c_0_236])).
cnf(c_0_248, plain, (m1_subset_1(k7_real_1(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_237])).
cnf(c_0_249, negated_conjecture, (m1_subset_1(X1,k2_rcomp_1(k1_real_1(np__1),np__1))|~r2_hidden(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_238])])).
cnf(c_0_250, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_239])).
fof(c_0_251, plain, ![X143, X144]:(~v1_xxreal_0(X143)|~v1_xxreal_0(X144)|m1_subset_1(k2_rcomp_1(X143,X144),k1_zfmisc_1(k1_numbers))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_rcomp_1])])).
cnf(c_0_252, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_240])).
cnf(c_0_253, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1)), inference(split_conjunct,[status(thm)],[c_0_241])).
cnf(c_0_254, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_242])).
cnf(c_0_255, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(X1)), inference(split_conjunct,[status(thm)],[c_0_241])).
cnf(c_0_256, plain, (v1_int_1(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_243])).
cnf(c_0_257, plain, (v7_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_245]), c_0_246])])])).
cnf(c_0_258, negated_conjecture, (~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~v1_xreal_0(k1_seq_1(k1_sin_cos9,esk2_0))|~r1_fdiff_1(k1_sin_cos9,esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk2_0),k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_211]), c_0_181]), c_0_212])])])).
cnf(c_0_259, plain, (m1_subset_1(k7_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_248])).
cnf(c_0_260, negated_conjecture, (v1_xreal_0(X1)|~v3_membered(k2_rcomp_1(k1_real_1(np__1),np__1))|~r2_hidden(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_249])])).
cnf(c_0_261, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_250])).
cnf(c_0_262, plain, (m1_subset_1(k2_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_251])).
cnf(c_0_263, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_252])).
fof(c_0_264, plain, ![X82]:(~m1_subset_1(X82,k1_numbers)|k1_real_1(X82)=k4_xcmplx_0(X82)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_real_1])])).
fof(c_0_265, plain, ![X101]:(~v1_int_1(X101)|v1_xreal_0(X101)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_int_1])])).
cnf(c_0_266, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1)), inference(evalgc,[status(thm)],[c_0_253])).
cnf(c_0_267, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_254])).
cnf(c_0_268, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1), inference(split_conjunct,[status(thm)],[rqRealNeg__k4_xcmplx_0__rm1_r1])).
cnf(c_0_269, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(X1)), inference(evalgc,[status(thm)],[c_0_255])).
cnf(c_0_270, plain, (v1_int_1(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_257])])).
cnf(c_0_271, negated_conjecture, (~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~v1_xreal_0(k1_seq_1(k1_sin_cos9,esk2_0))|~r1_fdiff_1(k1_sin_cos9,esk2_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk2_0),k1_numbers)), inference(csr,[status(thm)],[c_0_258, c_0_187])).
cnf(c_0_272, plain, (v1_xreal_0(k7_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_259]), c_0_124])])])).
fof(c_0_273, plain, ![X128]:(~m1_subset_1(X128,k1_numbers)|m1_subset_1(k5_square_1(X128),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_square_1])])).
cnf(c_0_274, negated_conjecture, (v1_xreal_0(X1)|~r2_hidden(X1,esk1_0)|~m1_subset_1(k2_rcomp_1(k1_real_1(np__1),np__1),k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_261])])).
cnf(c_0_275, plain, (m1_subset_1(k2_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_262])).
cnf(c_0_276, plain, (v1_xxreal_0(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_257])])).
cnf(c_0_277, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_264])).
fof(c_0_278, plain, ![X116]:(~v1_xreal_0(X116)|v1_xxreal_0(X116)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_xreal_0])])).
cnf(c_0_279, plain, (v1_xreal_0(X1)|~v1_int_1(X1)), inference(split_conjunct,[status(thm)],[c_0_265])).
cnf(c_0_280, plain, (v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_267])])).
cnf(c_0_281, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1), inference(evalgc,[status(thm)],[c_0_268])).
cnf(c_0_282, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_269, c_0_270])])).
cnf(c_0_283, negated_conjecture, (~v1_xreal_0(k1_seq_1(k1_sin_cos9,esk2_0))|~r1_fdiff_1(k1_sin_cos9,esk2_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk2_0),k1_numbers)|~v1_xreal_0(k5_square_1(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_272]), c_0_146])])])).
cnf(c_0_284, plain, (m1_subset_1(k5_square_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_273])).
cnf(c_0_285, negated_conjecture, (v1_xreal_0(X1)|~r2_hidden(X1,esk1_0)|~v1_xxreal_0(k1_real_1(np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_274, c_0_275]), c_0_276])])])).
cnf(c_0_286, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_277])).
cnf(c_0_287, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_278])).
cnf(c_0_288, plain, (v1_xreal_0(X1)|~v1_int_1(X1)), inference(evalgc,[status(thm)],[c_0_279])).
cnf(c_0_289, plain, (v1_int_1(k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_280, c_0_281]), c_0_270]), c_0_282])])])).
cnf(c_0_290, negated_conjecture, (~v1_xreal_0(k1_seq_1(k1_sin_cos9,esk2_0))|~v1_xreal_0(k5_square_1(esk2_0))|~r1_fdiff_1(k1_sin_cos9,esk2_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk2_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_62]), c_0_72]), c_0_137]), c_0_73]), c_0_123])])])).
cnf(c_0_291, plain, (m1_subset_1(k5_square_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_284])).
cnf(c_0_292, negated_conjecture, (v1_xreal_0(X1)|~v1_xxreal_0(k4_xcmplx_0(np__1))|~r2_hidden(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_286]), c_0_146])])])).
cnf(c_0_293, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_287])).
cnf(c_0_294, plain, (v1_xreal_0(k4_xcmplx_0(np__1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_289])])).
cnf(c_0_295, negated_conjecture, (~v1_xreal_0(k1_seq_1(k1_sin_cos9,esk2_0))|~v1_xreal_0(k5_square_1(esk2_0))|~r1_fdiff_1(k1_sin_cos9,esk2_0)|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk2_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_63]), c_0_72]), c_0_137]), c_0_73]), c_0_123])])])).
cnf(c_0_296, plain, (v1_xreal_0(k5_square_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_291]), c_0_124])])])).
cnf(c_0_297, negated_conjecture, (v1_xreal_0(X1)|~r2_hidden(X1,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_293]), c_0_294])])])).
cnf(c_0_298, negated_conjecture, (~v1_xreal_0(k1_seq_1(k1_sin_cos9,esk2_0))|~r1_fdiff_1(k1_sin_cos9,esk2_0)|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk2_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_295, c_0_296]), c_0_212])])])).
cnf(c_0_299, negated_conjecture, (v1_xreal_0(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_297, c_0_181])])).
cnf(c_0_300, negated_conjecture, (~v1_xreal_0(k1_seq_1(k1_sin_cos9,esk2_0))|~r1_fdiff_1(k1_sin_cos9,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_218]), c_0_299]), c_0_137]), c_0_123])])])).
cnf(c_0_301, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_101]), c_0_124])])])).
cnf(c_0_302, plain, (v1_relat_1(k1_sin_cos9)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_123])])).
cnf(c_0_303, negated_conjecture, (~r1_fdiff_1(k1_sin_cos9,esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_301]), c_0_140]), c_0_302]), c_0_137])])])).
cnf(c_0_304, negated_conjecture, (r1_fdiff_1(X1,esk2_0)|~v1_funct_1(X1)|~r2_fdiff_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_80, c_0_181]), c_0_127]), c_0_128])])])).
cnf(c_0_305, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_303, c_0_304]), c_0_137]), c_0_141]), c_0_123])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 306
# Proof object clause steps            : 216
# Proof object formula steps           : 90
# Proof object conjectures             : 55
# Proof object clause conjectures      : 52
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 66
# Proof object initial formulas used   : 48
# Proof object generating inferences   : 69
# Proof object simplifying inferences  : 154
# Training examples: 134 positive, 1091 negative
# Training: Positive examples begin
cnf(c_0_233, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_128, negated_conjecture, (v3_rcomp_1(esk1_0))).# trainpos
cnf(c_0_137, plain, (v1_funct_1(k1_sin_cos9))).# trainpos
cnf(c_0_72, plain, (v1_funct_1(k16_sin_cos))).# trainpos
cnf(c_0_124, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_127, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_126, negated_conjecture, (r1_tarski(esk1_0,k2_rcomp_1(k1_real_1(np__1),np__1)))).# trainpos
cnf(c_0_165, negated_conjecture, (k10_real_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k7_real_1(np__1,k5_square_1(esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0))).# trainpos
cnf(c_0_146, plain, (m1_subset_1(np__1,k1_numbers))).# trainpos
cnf(c_0_232, plain, (v6_membered(k4_ordinal1))).# trainpos
cnf(c_0_281, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1)).# trainpos
cnf(c_0_245, plain, (m1_subset_1(np__1,k5_numbers))).# trainpos
cnf(c_0_293, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_228, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_172, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0))).# trainpos
cnf(c_0_135, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0))).# trainpos
cnf(c_0_288, plain, (v1_xreal_0(X1)|~v1_int_1(X1))).# trainpos
cnf(c_0_263, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_167, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_256, plain, (v1_int_1(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_216, plain, (v1_xreal_0(k5_xcmplx_0(X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_76, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_125, plain, (r2_fdiff_1(k1_sin_cos9,X1)|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k2_rcomp_1(k1_real_1(np__1),np__1)))).# trainpos
cnf(c_0_71, plain, (r1_fdiff_1(k16_sin_cos,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_123, plain, (m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_73, plain, (m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_79, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_196, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_131, plain, (v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_269, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).# trainpos
cnf(c_0_266, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1))).# trainpos
cnf(c_0_178, plain, (k5_xcmplx_0(k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_267, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_112, plain, (k5_xcmplx_0(X1)=k7_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_261, plain, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_139, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))))).# trainpos
cnf(c_0_286, plain, (k1_real_1(X1)=k4_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_176, plain, (k1_fdiff_1(k16_sin_cos,X1)=k1_seq_1(k19_sin_cos,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_291, plain, (m1_subset_1(k5_square_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_226, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_244, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_120, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_61, plain, (r2_fdiff_1(X2,X1)|m1_subset_1(esk4_2(X1,X2),k1_numbers)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_129, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_144, plain, (k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_259, plain, (m1_subset_1(k7_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_191, plain, (m1_subset_1(k10_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_122, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_213, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_101, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_275, plain, (m1_subset_1(k2_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_111, plain, (k3_xcmplx_0(X1,k5_xcmplx_0(X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_202, plain, (k8_real_1(X1,X2)=k8_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_199, plain, (v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_218, plain, (m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_198, plain, (m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_53, plain, (r2_fdiff_1(X1,X2)|~v3_rcomp_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r1_fdiff_1(X1,esk4_2(X2,X1))|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_75, plain, (r1_fdiff_1(X3,X1)|~v3_rcomp_1(X2)|~v1_funct_1(X3)|~r2_fdiff_1(X3,X2)|~r2_hidden(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_63, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_117, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_62, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_115, plain, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,X2),X1)=k10_real_1(np__1,k7_real_1(np__1,k5_square_1(X1)))|~v3_rcomp_1(X2)|~r2_hidden(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r1_tarski(X2,k2_rcomp_1(k1_real_1(np__1),np__1)))).# trainpos
cnf(c_0_81, plain, (r2_fdiff_1(X2,X1)|r2_hidden(esk4_2(X1,X2),X1)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_183, plain, (k1_fdiff_1(X4,X1)=k1_seq_1(X3,X1)|X3!=k2_fdiff_1(X4,X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~r2_fdiff_1(X4,X2)|~r2_hidden(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_54, plain, (r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X3)|~m1_subset_1(X3,k1_numbers)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_175, plain, (k8_real_1(k1_fdiff_1(X2,k1_seq_1(X1,X3)),k1_fdiff_1(X1,X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X3)|~m1_subset_1(X3,k1_numbers)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_238, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k2_rcomp_1(k1_real_1(np__1),np__1))))).# trainpos
cnf(c_0_141, negated_conjecture, (r2_fdiff_1(k1_sin_cos9,esk1_0))).# trainpos
cnf(c_0_257, plain, (v7_ordinal1(np__1))).# trainpos
cnf(c_0_89, plain, (r1_fdiff_1(X1,esk4_2(X2,X3))|r2_fdiff_1(X3,X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_funct_1(X3)|~r1_tarski(X2,k1_relset_1(k1_numbers,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_60, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))))|~r1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)),k1_numbers))).# trainpos
cnf(c_0_197, plain, (k1_seq_1(k2_fdiff_1(X1,X3),X2)=k1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X1,X3))|~r2_hidden(X2,X3)|~r2_fdiff_1(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_302, plain, (v1_relat_1(k1_sin_cos9))).# trainpos
cnf(c_0_249, negated_conjecture, (m1_subset_1(X1,k2_rcomp_1(k1_real_1(np__1),np__1))|~r2_hidden(X1,esk1_0))).# trainpos
cnf(c_0_145, negated_conjecture, (k10_real_1(np__1,k7_real_1(np__1,k5_square_1(X1)))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r2_hidden(X1,esk1_0))).# trainpos
cnf(c_0_136, plain, (k1_partfun1(X1,X2,X3,X4,X5,X6)=k1_partfun1(X7,X8,X9,X10,X5,X6)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).# trainpos
cnf(c_0_179, negated_conjecture, (k7_xcmplx_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k7_real_1(np__1,k5_square_1(esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_numbers))).# trainpos
cnf(c_0_140, plain, (v3_valued_0(k1_sin_cos9))).# trainpos
cnf(c_0_276, plain, (v1_xxreal_0(np__1))).# trainpos
cnf(c_0_148, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_78, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~r1_fdiff_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_xreal_0(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)))))).# trainpos
cnf(c_0_270, plain, (v1_int_1(np__1))).# trainpos
cnf(c_0_150, plain, (m1_subset_1(k1_partfun1(X7,X8,X9,X10,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).# trainpos
cnf(c_0_154, negated_conjecture, (k7_xcmplx_0(np__1,k7_real_1(np__1,k5_square_1(X1)))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~v1_xreal_0(k7_real_1(np__1,k5_square_1(X1)))|~r2_hidden(X1,esk1_0))).# trainpos
cnf(c_0_186, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_88, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~r1_fdiff_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos))),k1_numbers))).# trainpos
cnf(c_0_282, plain, (v1_xcmplx_0(k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_280, plain, (v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_192, plain, (k7_xcmplx_0(np__1,k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1)))).# trainpos
cnf(c_0_143, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_260, negated_conjecture, (v1_xreal_0(X1)|~v3_membered(k2_rcomp_1(k1_real_1(np__1),np__1))|~r2_hidden(X1,esk1_0))).# trainpos
cnf(c_0_160, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X6,k1_numbers)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_190, plain, (k8_real_1(k1_seq_1(k19_sin_cos,k1_seq_1(X1,X2)),k1_fdiff_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~r1_fdiff_1(k16_sin_cos,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2)))).# trainpos
cnf(c_0_296, plain, (v1_xreal_0(k5_square_1(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_100, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)))|~m1_subset_1(k1_seq_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos))),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos))|~r2_fdiff_1(X1,X2)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_274, negated_conjecture, (v1_xreal_0(X1)|~r2_hidden(X1,esk1_0)|~m1_subset_1(k2_rcomp_1(k1_real_1(np__1),np__1),k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_214, plain, (k8_real_1(k1_fdiff_1(X1,X2),k1_seq_1(k19_sin_cos,k1_seq_1(X1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~v1_xreal_0(k1_seq_1(k19_sin_cos,k1_seq_1(X1,X2)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_289, plain, (v1_int_1(k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_119, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v3_valued_0(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_294, plain, (v1_xreal_0(k4_xcmplx_0(np__1)))).# trainpos
cnf(c_0_285, negated_conjecture, (v1_xreal_0(X1)|~r2_hidden(X1,esk1_0)|~v1_xxreal_0(k1_real_1(np__1)))).# trainpos
cnf(c_0_224, plain, (k3_xcmplx_0(k1_fdiff_1(X1,X2),k1_seq_1(k19_sin_cos,k1_seq_1(X1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~v1_xreal_0(k1_seq_1(k19_sin_cos,k1_seq_1(X1,X2)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_149, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_159, plain, (r2_fdiff_1(k1_partfun1(X3,X4,X5,X6,X1,k16_sin_cos),X2)|~v3_valued_0(X1)|~v1_funct_1(k1_partfun1(X3,X4,X5,X6,X1,k16_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k16_sin_cos)))|~m1_subset_1(k1_partfun1(X3,X4,X5,X6,X1,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).# trainpos
cnf(c_0_292, negated_conjecture, (v1_xreal_0(X1)|~v1_xxreal_0(k4_xcmplx_0(np__1))|~r2_hidden(X1,esk1_0))).# trainpos
cnf(c_0_156, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_187, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))).# trainpos
cnf(c_0_297, negated_conjecture, (v1_xreal_0(X1)|~r2_hidden(X1,esk1_0))).# trainpos
cnf(c_0_168, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))).# trainpos
cnf(c_0_200, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))).# trainpos
cnf(c_0_181, negated_conjecture, (r2_hidden(esk2_0,esk1_0))).# trainpos
cnf(c_0_212, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers))).# trainpos
cnf(c_0_272, plain, (v1_xreal_0(k7_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_299, negated_conjecture, (v1_xreal_0(esk2_0))).# trainpos
cnf(c_0_304, negated_conjecture, (r1_fdiff_1(X1,esk2_0)|~v1_funct_1(X1)|~r2_fdiff_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_227, plain, (m1_subset_1(k1_seq_1(k19_sin_cos,X1),k1_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_204, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_236, plain, (v1_xreal_0(k1_seq_1(k19_sin_cos,X1))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_215, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_225, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_301, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).# trainpos
cnf(c_0_130, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_142, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,X1),X2)=k7_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_166, negated_conjecture, (k7_xcmplx_0(X2,k7_real_1(np__1,k5_square_1(X1)))=k3_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k7_real_1(np__1,k5_square_1(X1)))|~v1_xreal_0(k7_real_1(np__1,k5_square_1(X1)))|~r2_hidden(X1,esk1_0))).# trainpos
cnf(c_0_195, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk2_0),k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_numbers)|~v1_xcmplx_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0))))).# trainpos
cnf(c_0_223, negated_conjecture, (k3_xcmplx_0(k1_fdiff_1(k1_sin_cos9,esk2_0),k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_numbers))).# trainpos
cnf(c_0_234, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk2_0)|~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_numbers)|~v1_xreal_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)))|~v1_xreal_0(k1_seq_1(k1_sin_cos9,esk2_0))|~r1_fdiff_1(k1_sin_cos9,esk2_0)|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk2_0),k1_numbers))).# trainpos
cnf(c_0_258, negated_conjecture, (~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~v1_xreal_0(k1_seq_1(k1_sin_cos9,esk2_0))|~r1_fdiff_1(k1_sin_cos9,esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk2_0),k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_283, negated_conjecture, (~v1_xreal_0(k1_seq_1(k1_sin_cos9,esk2_0))|~r1_fdiff_1(k1_sin_cos9,esk2_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk2_0),k1_numbers)|~v1_xreal_0(k5_square_1(esk2_0)))).# trainpos
cnf(c_0_290, negated_conjecture, (~v1_xreal_0(k1_seq_1(k1_sin_cos9,esk2_0))|~v1_xreal_0(k5_square_1(esk2_0))|~r1_fdiff_1(k1_sin_cos9,esk2_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk2_0),k1_numbers))).# trainpos
cnf(c_0_295, negated_conjecture, (~v1_xreal_0(k1_seq_1(k1_sin_cos9,esk2_0))|~v1_xreal_0(k5_square_1(esk2_0))|~r1_fdiff_1(k1_sin_cos9,esk2_0)|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk2_0),k1_numbers))).# trainpos
cnf(c_0_298, negated_conjecture, (~v1_xreal_0(k1_seq_1(k1_sin_cos9,esk2_0))|~r1_fdiff_1(k1_sin_cos9,esk2_0)|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk2_0),k1_numbers))).# trainpos
cnf(c_0_300, negated_conjecture, (~v1_xreal_0(k1_seq_1(k1_sin_cos9,esk2_0))|~r1_fdiff_1(k1_sin_cos9,esk2_0))).# trainpos
cnf(c_0_303, negated_conjecture, (~r1_fdiff_1(k1_sin_cos9,esk2_0))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_160, plain, (k4_xcmplx_0(np__0)=np__0)).#trainneg
cnf(i_0_59, plain, (v1_funct_1(k19_sin_cos))).#trainneg
cnf(i_0_125, plain, (v1_xreal_0(esk5_0))).#trainneg
cnf(i_0_140, plain, (v1_xreal_0(esk6_0))).#trainneg
cnf(i_0_29, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_63, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_126, plain, (v1_xxreal_0(esk5_0))).#trainneg
cnf(i_0_142, plain, (v1_xxreal_0(esk6_0))).#trainneg
cnf(i_0_196, plain, (v1_xxreal_0(esk8_0))).#trainneg
cnf(i_0_28, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_144, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_193, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_121, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_127, plain, (v1_xcmplx_0(esk5_0))).#trainneg
cnf(i_0_141, plain, (v1_xcmplx_0(esk6_0))).#trainneg
cnf(i_0_53, plain, (v1_xboole_0(np__0))).#trainneg
cnf(i_0_128, plain, (v1_xboole_0(esk5_0))).#trainneg
cnf(i_0_197, plain, (v1_xboole_0(esk8_0))).#trainneg
cnf(i_0_139, plain, (v1_int_1(esk6_0))).#trainneg
cnf(i_0_198, plain, (~v1_xboole_0(np__2))).#trainneg
cnf(i_0_189, plain, (k4_xcmplx_0(k4_xcmplx_0(np__2))=np__2)).#trainneg
cnf(i_0_145, plain, (k7_xcmplx_0(np__1,np__1)=np__1)).#trainneg
cnf(i_0_205, plain, (k3_xcmplx_0(np__2,np__0)=np__0)).#trainneg
cnf(i_0_70, plain, (k2_xcmplx_0(np__1,np__1)=np__2)).#trainneg
cnf(i_0_130, plain, (k2_xcmplx_0(np__1,np__0)=np__1)).#trainneg
cnf(i_0_79, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_72, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_90, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_55, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_67, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_65, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_66, plain, (m1_subset_1(np__0,k5_numbers))).#trainneg
cnf(i_0_115, plain, (v1_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_105, plain, (k2_xcmplx_0(np__0,np__1)=np__1)).#trainneg
cnf(i_0_83, plain, (k6_xcmplx_0(np__1,np__1)=np__0)).#trainneg
cnf(i_0_170, plain, (k6_xcmplx_0(np__1,np__0)=np__1)).#trainneg
cnf(i_0_168, plain, (k6_xcmplx_0(np__2,np__1)=np__1)).#trainneg
cnf(i_0_202, plain, (k6_xcmplx_0(np__2,np__2)=np__0)).#trainneg
cnf(i_0_133, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_60, plain, (m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_61, plain, (m1_subset_1(np__2,k5_numbers))).#trainneg
cnf(i_0_143, plain, (m1_subset_1(esk6_0,k1_numbers))).#trainneg
cnf(i_0_82, plain, (k6_xcmplx_0(np__1,np__2)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_117, plain, (k6_xcmplx_0(np__0,np__1)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_186, plain, (k6_xcmplx_0(np__0,np__2)=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_159, plain, (k2_xcmplx_0(np__1,k4_xcmplx_0(np__1))=np__0)).#trainneg
cnf(i_0_187, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),np__1)=np__0)).#trainneg
cnf(i_0_88, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_122, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_62, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_36, plain, (v1_funct_2(k16_sin_cos,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_58, plain, (v1_funct_2(k19_sin_cos,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_120, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_131, plain, (v2_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_177, plain, (v3_membered(X1)|~v4_membered(X1))).#trainneg
cnf(i_0_169, plain, (m1_subset_1(esk7_1(X1),X1))).#trainneg
cnf(i_0_135, plain, (v1_relat_1(k3_relat_1(X1,X2)))).#trainneg
cnf(i_0_137, plain, (v1_xcmplx_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_190, plain, (v4_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_123, plain, (k7_xcmplx_0(np__1,k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_200, plain, (k2_xcmplx_0(np__1,k4_xcmplx_0(np__2))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_180, plain, (k2_xcmplx_0(np__0,k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_194, plain, (k2_xcmplx_0(k4_xcmplx_0(np__2),np__1)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_129, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_182, plain, (k7_xcmplx_0(k4_xcmplx_0(np__1),np__2)=k7_xcmplx_0(np__1,k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_84, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_99, plain, (v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_85, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_119, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_106, plain, (k7_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_108, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_188, plain, (k2_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))=k4_xcmplx_0(np__2))).#trainneg
cnf(i_0_57, plain, (m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_102, plain, (k3_square_1(X1)=k3_xcmplx_0(X1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_147, plain, (k3_square_1(X1)=k5_square_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_206, plain, (k7_xcmplx_0(k4_xcmplx_0(np__1),np__2)=k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_10, plain, (X1=k2_fdiff_1(X2,X3)|k1_relset_1(k1_numbers,X1)!=X3|k1_fdiff_1(X2,esk3_3(X2,X3,X1))!=k1_seq_1(X1,esk3_3(X2,X3,X1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_112, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_109, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_199, plain, (k6_xcmplx_0(np__1,k7_xcmplx_0(np__1,np__2))=k7_xcmplx_0(np__1,np__2))).#trainneg
cnf(i_0_56, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_45, plain, (m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_191, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_78, plain, (k1_real_1(k1_real_1(X1))=X1|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_166, plain, (v3_valued_0(k2_zfmisc_1(X2,X1))|~v3_membered(X1))).#trainneg
cnf(i_0_176, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_146, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_134, plain, (v1_int_1(X2)|~v5_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_52, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_95, plain, (k7_real_1(X1,X2)=k7_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_165, plain, (v1_membered(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_157, plain, (v3_rcomp_1(k4_xxreal_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_183, plain, (v2_membered(X2)|~v2_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_162, plain, (v3_valued_0(X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_50, plain, (k4_xxreal_1(X1,X2)=k2_rcomp_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_93, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_153, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_185, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_104, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_97, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_172, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_178, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_167, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_181, plain, (v1_xcmplx_0(k6_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_195, plain, (v1_int_1(k3_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_175, plain, (v1_int_1(k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_201, plain, (v1_int_1(k6_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_76, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_75, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_48, plain, (m1_subset_1(k8_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_203, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_118, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_158, plain, (v1_valued_0(X2)|~v1_relat_1(X1)|~v1_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_204, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_156, plain, (k2_xcmplx_0(X1,k4_xcmplx_0(X2))=k6_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_54, plain, (k7_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_161, plain, (v1_valued_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_132, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_136, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_valued_0(X1))).#trainneg
cnf(i_0_179, plain, (k7_xcmplx_0(k5_xcmplx_0(X1),k5_xcmplx_0(X2))=k7_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_116, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_69, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_173, plain, (k3_xcmplx_0(k5_xcmplx_0(X1),k5_xcmplx_0(X2))=k5_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_21, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_184, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xboole_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_152, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_14, plain, (k1_relset_1(k1_numbers,X1)=X2|X1!=k2_fdiff_1(X3,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_151, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_174, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_164, plain, (k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k7_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_171, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_163, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_148, plain, (k2_xcmplx_0(k3_xcmplx_0(X1,X3),k3_xcmplx_0(X2,X3))=k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12, plain, (X3=k2_fdiff_1(X1,X2)|m1_subset_1(esk3_3(X1,X2,X3),k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_11, plain, (X3=k2_fdiff_1(X1,X2)|r2_hidden(esk3_3(X1,X2,X3),X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_222, negated_conjecture, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_223, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_207, plain, (k1_xboole_0=esk8_0)).#trainneg
cnf(i_0_182, plain, (k7_xcmplx_0(k4_xcmplx_0(np__1),np__2)=k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_282, plain, (v1_xreal_0(np__2))).#trainneg
cnf(i_0_285, plain, (v1_xreal_0(np__0))).#trainneg
cnf(i_0_291, plain, (v1_xreal_0(np__1))).#trainneg
cnf(i_0_211, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_216, plain, (v3_membered(X1)|~v5_membered(X1))).#trainneg
cnf(i_0_272, plain, (r2_fdiff_1(k1_sin_cos9,X1)|~r1_tarski(X1,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_226, plain, (r2_fdiff_1(k1_sin_cos9,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_rcomp_1(k1_real_1(np__1),np__1))))).#trainneg
cnf(i_0_638, plain, (r2_fdiff_1(k16_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k16_sin_cos))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(esk4_2(X1,k16_sin_cos)))).#trainneg
cnf(i_0_650, plain, (k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2))=X2|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_264, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))))).#trainneg
cnf(i_0_271, negated_conjecture, (r1_tarski(esk1_0,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))).#trainneg
cnf(i_0_750, plain, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,X1),X2)=k10_real_1(np__1,k7_real_1(np__1,k5_square_1(X2)))|~r2_hidden(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_rcomp_1(k1_real_1(np__1),np__1))))).#trainneg
cnf(i_0_583, negated_conjecture, (k10_real_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k2_xcmplx_0(np__1,k5_square_1(esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~v1_xreal_0(k5_square_1(esk2_0)))).#trainneg
cnf(i_0_225, plain, (r2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1))|~m1_subset_1(k2_rcomp_1(k1_real_1(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k1_real_1(np__1),np__1)))).#trainneg
cnf(i_0_221, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk7_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_329, plain, (v7_ordinal1(np__2))).#trainneg
cnf(i_0_55, plain, (X1=esk8_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_509, plain, (v4_relat_1(k19_sin_cos,k1_numbers))).#trainneg
cnf(i_0_332, plain, (v7_ordinal1(np__0))).#trainneg
cnf(i_0_359, plain, (v1_relat_1(k19_sin_cos))).#trainneg
cnf(i_0_360, plain, (v1_relat_1(k16_sin_cos))).#trainneg
cnf(i_0_363, negated_conjecture, (v3_membered(esk1_0))).#trainneg
cnf(i_0_277, plain, (X1=k2_fdiff_1(k16_sin_cos,X2)|k1_seq_1(k19_sin_cos,esk3_3(k16_sin_cos,X2,X1))!=k1_seq_1(X1,esk3_3(k16_sin_cos,X2,X1))|k1_relset_1(k1_numbers,X1)!=X2|~v1_funct_1(X1)|~r2_fdiff_1(k16_sin_cos,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(esk3_3(k16_sin_cos,X2,X1)))).#trainneg
cnf(i_0_770, negated_conjecture, (m1_subset_1(esk4_2(esk1_0,X1),k1_numbers)|r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~r1_tarski(esk1_0,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_747, plain, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,X1),X2)=k10_real_1(np__1,k7_real_1(np__1,k5_square_1(X2)))|~r2_hidden(X2,X1)|~r1_tarski(X1,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_374, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|m1_subset_1(esk4_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)),k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_375, plain, (r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|m1_subset_1(esk4_2(k1_relset_1(k1_numbers,X1),X1),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,X1)))).#trainneg
cnf(i_0_274, plain, (k4_xcmplx_0(k1_real_1(X1))=X1|~m1_subset_1(k1_real_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_555, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_364, plain, (v3_membered(k5_numbers))).#trainneg
cnf(i_0_517, plain, (v3_valued_0(k19_sin_cos))).#trainneg
cnf(i_0_133, plain, (m1_subset_1(esk8_0,k5_numbers))).#trainneg
cnf(i_0_510, plain, (v4_relat_1(k16_sin_cos,k1_numbers))).#trainneg
cnf(i_0_518, plain, (v3_valued_0(k16_sin_cos))).#trainneg
cnf(i_0_793, plain, (v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_231, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k1_sin_cos9))).#trainneg
cnf(i_0_233, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k16_sin_cos))).#trainneg
cnf(i_0_276, plain, (k1_real_1(k4_xcmplx_0(X1))=X1|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_848, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))).#trainneg
cnf(i_0_819, plain, (r2_fdiff_1(k1_sin_cos9,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))).#trainneg
cnf(i_0_850, plain, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,X1),X2)=k10_real_1(np__1,k7_real_1(np__1,k5_square_1(X2)))|~r2_hidden(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_791, plain, (v1_xcmplx_0(np__2))).#trainneg
cnf(i_0_792, plain, (v1_xcmplx_0(np__0))).#trainneg
cnf(i_0_863, plain, (esk8_0=esk5_0)).#trainneg
cnf(i_0_854, negated_conjecture, (k7_xcmplx_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k2_xcmplx_0(np__1,k5_square_1(esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k5_square_1(esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~v1_xreal_0(k2_xcmplx_0(np__1,k5_square_1(esk2_0)))|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_511, plain, (v4_relat_1(k1_sin_cos9,k1_numbers))).#trainneg
cnf(i_0_369, negated_conjecture, (v1_membered(esk1_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_383, negated_conjecture, (v2_membered(esk1_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_943, plain, (m1_subset_1(k1_relset_1(k1_numbers,k16_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_886, plain, (r1_fdiff_1(k1_partfun1(X4,X5,X6,X7,X1,X2),X3)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~r1_fdiff_1(X1,X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_887, plain, (r2_fdiff_1(k1_partfun1(X4,X5,X6,X7,X1,X2),X3)|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X3,k1_partfun1(X4,X5,X6,X7,X1,X2))))|~r1_fdiff_1(X1,esk4_2(X3,k1_partfun1(X4,X5,X6,X7,X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(X4,X5,X6,X7,X1,X2)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_919, plain, (v1_xreal_0(esk4_2(k1_relset_1(k1_numbers,X1),X1))|r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,X1)))).#trainneg
cnf(i_0_876, negated_conjecture, (m1_subset_1(X1,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_890, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_sin_cos9,k16_sin_cos)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_860, plain, (v1_xxreal_0(np__2))).#trainneg
cnf(i_0_862, plain, (v1_int_1(np__2))).#trainneg
cnf(i_0_513, plain, (v5_relat_1(k19_sin_cos,k1_numbers))).#trainneg
cnf(i_0_868, plain, (v1_xxreal_0(np__0))).#trainneg
cnf(i_0_912, negated_conjecture, (v1_xreal_0(esk4_2(esk1_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_825, plain, (r2_fdiff_1(k16_sin_cos,X1)|~v1_xreal_0(esk4_2(X1,k16_sin_cos))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k16_sin_cos))))).#trainneg
cnf(i_0_842, negated_conjecture, (m1_subset_1(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_828, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4)|~r1_fdiff_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))|~v1_funct_1(X1)|~r1_tarski(X4,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X4)|~r1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))))))|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)))))|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)))),k1_numbers))).#trainneg
cnf(i_0_883, negated_conjecture, (r2_fdiff_1(X1,esk1_0)|m1_subset_1(esk4_2(esk1_0,X1),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_881, plain, (k2_fdiff_1(X1,X2)=k2_fdiff_1(k16_sin_cos,X3)|k1_fdiff_1(X1,esk3_3(k16_sin_cos,X3,k2_fdiff_1(X1,X2)))!=k1_seq_1(k19_sin_cos,esk3_3(k16_sin_cos,X3,k2_fdiff_1(X1,X2)))|k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2))!=X3|~v1_xreal_0(esk3_3(k16_sin_cos,X3,k2_fdiff_1(X1,X2)))|~v1_funct_1(k2_fdiff_1(X1,X2))|~r2_fdiff_1(k16_sin_cos,X3)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~r2_hidden(esk3_3(k16_sin_cos,X3,k2_fdiff_1(X1,X2)),X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_3(k16_sin_cos,X3,k2_fdiff_1(X1,X2)),k1_numbers))).#trainneg
cnf(i_0_879, negated_conjecture, (k10_real_1(np__1,k2_xcmplx_0(np__1,k5_square_1(X1)))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r2_hidden(X1,esk1_0)|~v1_xreal_0(k5_square_1(X1)))).#trainneg
cnf(i_0_870, plain, (v1_int_1(np__0))).#trainneg
cnf(i_0_242, plain, (v1_xcmplx_0(k4_xcmplx_0(esk6_0)))).#trainneg
cnf(i_0_956, plain, (k1_real_1(np__0)=np__0)).#trainneg
cnf(i_0_759, plain, (X3=k2_fdiff_1(X2,X1)|~v1_xboole_0(X1)|k1_relset_1(k1_numbers,X3)!=X1|~v1_funct_1(X3)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_243, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_514, plain, (v5_relat_1(k16_sin_cos,k1_numbers))).#trainneg
cnf(i_0_244, plain, (v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_845, plain, (k2_fdiff_1(X1,X2)=k2_fdiff_1(X3,X4)|k1_fdiff_1(X3,esk3_3(X3,X4,k2_fdiff_1(X1,X2)))!=k1_fdiff_1(X1,esk3_3(X3,X4,k2_fdiff_1(X1,X2)))|k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2))!=X4|~v1_funct_1(X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r2_fdiff_1(X3,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~r2_hidden(esk3_3(X3,X4,k2_fdiff_1(X1,X2)),X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_3(X3,X4,k2_fdiff_1(X1,X2)),k1_numbers))).#trainneg
cnf(i_0_824, plain, (r2_fdiff_1(k16_sin_cos,k1_relset_1(k1_numbers,k16_sin_cos))|~v1_xreal_0(esk4_2(k1_relset_1(k1_numbers,k16_sin_cos),k16_sin_cos))|~m1_subset_1(k1_relset_1(k1_numbers,k16_sin_cos),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,k16_sin_cos)))).#trainneg
cnf(i_0_882, negated_conjecture, (r2_fdiff_1(k2_fdiff_1(X1,X2),esk1_0)|m1_subset_1(esk4_2(esk1_0,k2_fdiff_1(X1,X2)),k1_numbers)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(esk1_0,X2)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1036, plain, (r2_fdiff_1(k16_sin_cos,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k16_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(esk4_2(X1,k16_sin_cos),k1_numbers))).#trainneg
cnf(i_0_1037, plain, (r2_fdiff_1(k16_sin_cos,k1_relset_1(k1_numbers,k16_sin_cos))|~v3_rcomp_1(k1_relset_1(k1_numbers,k16_sin_cos)))).#trainneg
cnf(i_0_996, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k16_sin_cos)))).#trainneg
cnf(i_0_133, plain, (m1_subset_1(esk5_0,k5_numbers))).#trainneg
cnf(i_0_207, plain, (k1_xboole_0=esk5_0)).#trainneg
cnf(i_0_55, plain, (X1=esk5_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_515, plain, (v5_relat_1(k1_sin_cos9,k1_numbers))).#trainneg
cnf(i_0_248, plain, (v1_xreal_0(X1)|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_251, plain, (v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1054, negated_conjecture, (v1_xreal_0(esk4_2(esk1_0,X1))|r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_1097, negated_conjecture, (v1_xreal_0(esk4_2(esk1_0,k2_fdiff_1(X1,X2)))|r2_fdiff_1(k2_fdiff_1(X1,X2),esk1_0)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r1_tarski(esk1_0,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_978, plain, (r2_fdiff_1(k1_partfun1(X3,X4,X5,X6,X1,k16_sin_cos),X2)|~v1_xreal_0(k1_seq_1(X1,esk4_2(X2,k1_partfun1(X3,X4,X5,X6,X1,k16_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X2,k1_partfun1(X3,X4,X5,X6,X1,k16_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k16_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_1018, plain, (v1_xreal_0(esk4_2(X2,k2_fdiff_1(X1,X2)))|r2_fdiff_1(k2_fdiff_1(X1,X2),X2)|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1027, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1038, negated_conjecture, (v1_xreal_0(X1)|~v3_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_980, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos),X3)|~v1_xreal_0(k1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos))))|~r1_fdiff_1(k2_fdiff_1(X1,X2),esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos)))|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_hidden(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos)),X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_370, plain, (v1_membered(k5_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_232, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k1_sin_cos9))).#trainneg
cnf(i_0_588, plain, (v1_valued_0(k19_sin_cos)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_589, plain, (v1_valued_0(k16_sin_cos)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_256, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k19_sin_cos))).#trainneg
cnf(i_0_1034, negated_conjecture, (v1_xreal_0(esk4_2(esk1_0,k1_partfun1(X1,X2,X3,X4,k1_sin_cos9,k16_sin_cos)))|r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_sin_cos9,k16_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k1_sin_cos9,k16_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k1_sin_cos9,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_917, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_sin_cos9,k16_sin_cos),esk1_0)|m1_subset_1(esk4_2(esk1_0,k1_partfun1(X1,X2,X3,X4,k1_sin_cos9,k16_sin_cos)),k1_numbers)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k1_sin_cos9,k16_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k1_sin_cos9,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_843, negated_conjecture, (~v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_844, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3),X4)|~r1_fdiff_1(X3,k1_fdiff_1(X1,esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3))))|~r1_fdiff_1(k2_fdiff_1(X1,X2),esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3)))|~v1_funct_1(X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(X4,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X4)|~v1_funct_1(X1)|~r2_hidden(esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3)),X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3)),k1_numbers))).#trainneg
cnf(i_0_1060, plain, (k2_fdiff_1(X1,X2)=k2_fdiff_1(k16_sin_cos,X2)|k1_fdiff_1(X1,esk3_3(k16_sin_cos,X2,k2_fdiff_1(X1,X2)))!=k1_seq_1(k19_sin_cos,esk3_3(k16_sin_cos,X2,k2_fdiff_1(X1,X2)))|k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2))!=X2|~v1_funct_1(X1)|~r2_fdiff_1(k16_sin_cos,X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(esk3_3(k16_sin_cos,X2,k2_fdiff_1(X1,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1031, plain, (v1_xcmplx_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_384, plain, (v2_membered(k5_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_234, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k16_sin_cos))).#trainneg
cnf(i_0_259, plain, (v1_xreal_0(k7_xcmplx_0(np__1,X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_590, plain, (v1_valued_0(k1_sin_cos9)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_260, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1087, negated_conjecture, (v1_xreal_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~v1_xreal_0(k7_real_1(np__1,k5_square_1(X1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_975, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~r2_fdiff_1(k3_relat_1(k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_635, plain, (v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1061, negated_conjecture, (k7_xcmplx_0(np__1,k2_xcmplx_0(np__1,k5_square_1(X1)))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~v1_xreal_0(k2_xcmplx_0(np__1,k5_square_1(X1)))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_273, plain, (m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_257, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k19_sin_cos))).#trainneg
cnf(i_0_262, plain, (k5_xcmplx_0(k7_xcmplx_0(np__1,X1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1118, plain, (esk5_0=np__0)).#trainneg
cnf(i_0_327, plain, (v7_ordinal1(esk6_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_263, plain, (k5_square_1(X1)=k3_xcmplx_0(X1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_267, plain, (v3_xxreal_0(k4_xcmplx_0(np__2))|~v1_xreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1139, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),X6)|~r1_fdiff_1(X5,esk4_2(X6,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos)))|~v1_funct_1(X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~m1_subset_1(k1_seq_1(X5,esk4_2(X6,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_1106, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k16_sin_cos),X3)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k16_sin_cos)))|~m1_subset_1(k1_seq_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k16_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k16_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k16_sin_cos)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1147, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos),X3)|~v1_xreal_0(k1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos))))|~v1_funct_1(X1)|~r2_hidden(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos)),X2)|~r2_fdiff_1(X1,X2)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos)))|~m1_subset_1(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos))|~r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1144, negated_conjecture, (v1_xreal_0(X1)|~r2_hidden(X1,esk1_0)|~m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1107, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,X6),k16_sin_cos),X7)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,X6),k16_sin_cos)))|~m1_subset_1(k1_seq_1(k1_partfun1(X1,X2,X3,X4,X5,X6),esk4_2(X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,X6),k16_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7)|~r1_fdiff_1(X6,k1_seq_1(X5,esk4_2(X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,X6),k16_sin_cos))))|~r1_fdiff_1(X5,esk4_2(X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,X6),k16_sin_cos)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(esk4_2(X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,X6),k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1141, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos)))|~m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1230, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1254, plain, (k5_xcmplx_0(np__1)=np__1)).#trainneg
cnf(i_0_1149, plain, (X2=k2_fdiff_1(X1,k1_sin_cos9)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|k1_relset_1(k1_numbers,X2)!=k1_sin_cos9|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_sin_cos9)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_268, plain, (v3_xxreal_0(k4_xcmplx_0(np__1))|~v1_xreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_389, negated_conjecture, (v3_valued_0(esk1_0)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_531, negated_conjecture, (v1_valued_0(esk1_0)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_269, plain, (v3_xxreal_0(k4_xcmplx_0(X1))|~v2_xxreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_270, plain, (v3_xxreal_0(np__0)|~v2_xxreal_0(np__0)|~v1_xreal_0(np__0))).#trainneg
cnf(i_0_1207, negated_conjecture, (v1_xreal_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k7_real_1(np__1,k5_square_1(X1)),k1_numbers))).#trainneg
cnf(i_0_1138, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos),X7)|~v1_xreal_0(k1_fdiff_1(X5,esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos))))|~r1_fdiff_1(k2_fdiff_1(X5,X6),esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)))|~v1_funct_1(k2_fdiff_1(X5,X6))|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_fdiff_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7)|~v1_funct_1(X5)|~r2_hidden(esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)),X6)|~r2_fdiff_1(X5,X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1142, negated_conjecture, (v1_xreal_0(X1)|~v3_membered(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k1_sin_cos9,k16_sin_cos)))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1209, negated_conjecture, (v1_xreal_0(X1)|~v3_membered(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_1179, plain, (r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X1,X2),X3),X4)|~r1_fdiff_1(X3,k1_fdiff_1(X1,esk4_2(X4,k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X1,X2),X3))))|~r1_fdiff_1(k2_fdiff_1(X1,X2),esk4_2(X4,k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X1,X2),X3)))|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_hidden(esk4_2(X4,k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X1,X2),X3)),X2)|~r2_fdiff_1(X1,X2)|~r1_tarski(X4,k1_relset_1(k1_numbers,k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X1,X2),X3)))|~m1_subset_1(esk4_2(X4,k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X1,X2),X3)),k1_numbers)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X4)|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1185, plain, (X2=k2_fdiff_1(X1,k16_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|k1_relset_1(k1_numbers,X2)!=k16_sin_cos|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k16_sin_cos)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_278, plain, (v1_xreal_0(esk7_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_817, plain, (v3_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_991, plain, (v3_membered(k1_relset_1(k1_numbers,k16_sin_cos)))).#trainneg
cnf(i_0_280, negated_conjecture, (v1_xreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_286, plain, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1220, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1274, negated_conjecture, (k10_real_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k7_real_1(np__1,k3_xcmplx_0(esk2_0,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(esk2_0,k1_numbers))).#trainneg
cnf(i_0_1318, negated_conjecture, (r1_tarski(esk1_0,k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos))))).#trainneg
cnf(i_0_1285, plain, (r2_fdiff_1(k1_partfun1(X3,X4,X5,X6,X1,k16_sin_cos),X2)|~v1_funct_1(k1_partfun1(X3,X4,X5,X6,X1,k16_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(X3,X4,X5,X6,X1,k16_sin_cos)))|~m1_subset_1(k1_seq_1(X1,esk4_2(X2,k1_partfun1(X3,X4,X5,X6,X1,k16_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(X3,X4,X5,X6,X1,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_1287, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos),X3)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(X1,X2))|~r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos)))|~m1_subset_1(k1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_hidden(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos)),X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1276, negated_conjecture, (k10_real_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k2_xcmplx_0(np__1,k3_xcmplx_0(esk2_0,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k3_xcmplx_0(esk2_0,esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(esk2_0,k1_numbers))).#trainneg
cnf(i_0_55, plain, (X1=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_207, plain, (k1_xboole_0=np__0)).#trainneg
cnf(i_0_863, plain, (esk8_0=np__0)).#trainneg
cnf(i_0_288, plain, (v1_xreal_0(k19_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_289, plain, (v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1239, plain, (v1_xcmplx_0(k5_square_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_292, plain, (v1_xreal_0(k16_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_293, plain, (v1_xreal_0(k1_sin_cos9)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1305, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos),X2)|~v1_xreal_0(k1_seq_1(k19_sin_cos,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos))))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos))|~r2_hidden(esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos)),X1)|~r2_fdiff_1(k2_fdiff_1(k16_sin_cos,X1),X2)|~r2_fdiff_1(k16_sin_cos,X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_xreal_0(esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos))))).#trainneg
cnf(i_0_1275, negated_conjecture, (k7_xcmplx_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k7_real_1(np__1,k3_xcmplx_0(esk2_0,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k7_real_1(np__1,k3_xcmplx_0(esk2_0,esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_numbers)|~m1_subset_1(esk2_0,k1_numbers))).#trainneg
cnf(i_0_1351, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos),X7)|~v1_xreal_0(k1_fdiff_1(X5,esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos))))|~v1_funct_1(X5)|~r2_hidden(esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)),X6)|~r2_fdiff_1(X5,X6)|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)))|~m1_subset_1(esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)),k1_numbers)|~m1_subset_1(k2_fdiff_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7)|~v1_funct_1(k2_fdiff_1(X5,X6))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos))|~r2_fdiff_1(k2_fdiff_1(X5,X6),X7)|~m1_subset_1(k2_fdiff_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1281, negated_conjecture, (k10_real_1(np__1,k7_real_1(np__1,k3_xcmplx_0(X1,X1)))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r2_hidden(X1,esk1_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1277, negated_conjecture, (k7_xcmplx_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k2_xcmplx_0(np__1,k3_xcmplx_0(esk2_0,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k2_xcmplx_0(np__1,k3_xcmplx_0(esk2_0,esk2_0)))|~v1_xreal_0(k3_xcmplx_0(esk2_0,esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_numbers)|~m1_subset_1(esk2_0,k1_numbers))).#trainneg
cnf(i_0_1343, negated_conjecture, (v1_xreal_0(X1)|~r2_hidden(X1,esk1_0)|~m1_subset_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1325, plain, (k7_xcmplx_0(X1,np__1)=k3_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1319, plain, (m1_subset_1(k4_xcmplx_0(np__2),k1_numbers))).#trainneg
cnf(i_0_1241, plain, (X2=k2_fdiff_1(X1,k19_sin_cos)|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|k1_relset_1(k1_numbers,X2)!=k19_sin_cos|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,k19_sin_cos)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_294, plain, (v1_xxreal_0(esk7_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_992, plain, (v1_membered(k1_relset_1(k1_numbers,k16_sin_cos))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_296, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_302, plain, (v1_xxreal_0(k5_numbers)|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1294, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos)))|~m1_subset_1(k1_seq_1(k16_sin_cos,esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos))),k1_numbers)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos))))).#trainneg
cnf(i_0_1280, plain, (v1_xreal_0(k3_xcmplx_0(X1,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1292, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k16_sin_cos),X7)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6))|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k16_sin_cos)))|~m1_subset_1(k1_seq_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k16_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7)|~r1_fdiff_1(X6,k1_seq_1(X5,esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k16_sin_cos))))|~r1_fdiff_1(X5,esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k16_sin_cos)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1312, plain, (r2_fdiff_1(k1_partfun1(X8,X9,X10,X11,k1_partfun1(X1,X2,X3,X4,X5,X6),k16_sin_cos),X7)|~r1_fdiff_1(X6,k1_seq_1(X5,esk4_2(X7,k1_partfun1(X8,X9,X10,X11,k1_partfun1(X1,X2,X3,X4,X5,X6),k16_sin_cos))))|~r1_fdiff_1(X5,esk4_2(X7,k1_partfun1(X8,X9,X10,X11,k1_partfun1(X1,X2,X3,X4,X5,X6),k16_sin_cos)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(X8,X9,X10,X11,k1_partfun1(X1,X2,X3,X4,X5,X6),k16_sin_cos)))|~m1_subset_1(k1_seq_1(k1_partfun1(X1,X2,X3,X4,X5,X6),esk4_2(X7,k1_partfun1(X8,X9,X10,X11,k1_partfun1(X1,X2,X3,X4,X5,X6),k16_sin_cos))),k1_numbers)|~m1_subset_1(esk4_2(X7,k1_partfun1(X8,X9,X10,X11,k1_partfun1(X1,X2,X3,X4,X5,X6),k16_sin_cos)),k1_numbers)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X10,X11)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X8,X9))))).#trainneg
cnf(i_0_1102, negated_conjecture, (r2_fdiff_1(k16_sin_cos,esk1_0)|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k16_sin_cos))))).#trainneg
cnf(i_0_1223, plain, (v1_xreal_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_772, negated_conjecture, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_1320, plain, (m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).#trainneg
cnf(i_0_942, plain, (k9_xtuple_0(k16_sin_cos)=k1_relset_1(k1_numbers,k16_sin_cos))).#trainneg
cnf(i_0_303, plain, (v1_xxreal_0(k5_square_1(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_390, plain, (v3_valued_0(k5_numbers)|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_312, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_318, plain, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1430, plain, (v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_304, plain, (v1_xxreal_0(k19_sin_cos)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_305, plain, (v1_xxreal_0(k1_real_1(X1))|~v2_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1455, negated_conjecture, (k10_real_1(np__1,k7_real_1(np__1,k3_square_1(X1)))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r2_hidden(X1,esk1_0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1496, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5)|~v1_xreal_0(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos)))|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5))).#trainneg
cnf(i_0_1421, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)))|~v3_valued_0(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos))))).#trainneg
cnf(i_0_775, plain, (m1_subset_1(X1,k1_numbers)|v1_xboole_0(k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_268, plain, (v3_xxreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_982, plain, (k9_xtuple_0(k1_sin_cos9)=k1_relset_1(k1_numbers,k1_sin_cos9))).#trainneg
cnf(i_0_308, plain, (v1_xxreal_0(k16_sin_cos)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_532, plain, (v1_valued_0(k5_numbers)|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_328, negated_conjecture, (v7_ordinal1(esk1_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_334, plain, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_993, plain, (v2_membered(k1_relset_1(k1_numbers,k16_sin_cos))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_309, plain, (v1_xxreal_0(k1_sin_cos9)|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_310, plain, (v1_xcmplx_0(esk7_1(X1))|~v1_membered(X1))).#trainneg
cnf(i_0_1463, negated_conjecture, (v1_xreal_0(X1)|~r2_hidden(X1,esk1_0)|~v1_xreal_0(k1_real_1(np__1)))).#trainneg
cnf(i_0_1416, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos),X3)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos))|~v1_funct_1(X1)|~r2_hidden(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos)),X2)|~r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|~r2_fdiff_1(X1,X2)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_xreal_0(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos))))).#trainneg
cnf(i_0_1411, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos),X7)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(X5,X6))|~r2_fdiff_1(k2_fdiff_1(X5,X6),X7)|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)))|~m1_subset_1(k1_fdiff_1(X5,esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_fdiff_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7)|~v1_funct_1(X5)|~r2_hidden(esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)),X6)|~r2_fdiff_1(X5,X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1480, plain, (k3_xcmplx_0(np__1,np__1)=np__1)).#trainneg
cnf(i_0_1448, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos),X7)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos))|~v1_funct_1(X5)|~r2_hidden(esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)),X6)|~r2_fdiff_1(k2_fdiff_1(X5,X6),X7)|~r2_fdiff_1(X5,X6)|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7)|~m1_subset_1(k1_fdiff_1(X5,esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_344, negated_conjecture, (v1_int_1(esk1_0)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1484, plain, (v1_xreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1488, plain, (v1_int_1(k4_xcmplx_0(np__2))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_365, plain, (v3_membered(k19_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_350, plain, (v1_int_1(k5_numbers)|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1338, plain, (v3_xxreal_0(np__2)|~v2_xxreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_335, plain, (v7_ordinal1(k5_square_1(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_320, plain, (v1_xcmplx_0(k19_sin_cos)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_324, plain, (v1_xcmplx_0(k16_sin_cos)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1422, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~v3_valued_0(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)))))).#trainneg
cnf(i_0_1447, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos),X6)|~v1_xreal_0(k1_seq_1(k19_sin_cos,esk4_2(X6,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos))))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos))|~r2_hidden(esk4_2(X6,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos)),X5)|~r2_fdiff_1(k2_fdiff_1(k16_sin_cos,X5),X6)|~r2_fdiff_1(k16_sin_cos,X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k16_sin_cos,X5),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~v1_xreal_0(esk4_2(X6,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos))))).#trainneg
cnf(i_0_1465, plain, (v1_xreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2))|~v1_xreal_0(k1_seq_1(k19_sin_cos,k1_seq_1(X1,X2)))|~v1_xreal_0(k1_fdiff_1(X1,X2))|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1440, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos),X2)|~v1_xreal_0(esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos))|~r2_hidden(esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos)),X1)|~r2_fdiff_1(k2_fdiff_1(k16_sin_cos,X1),X2)|~r2_fdiff_1(k16_sin_cos,X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(k1_seq_1(k19_sin_cos,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_1449, negated_conjecture, (k7_xcmplx_0(np__1,k7_real_1(np__1,k3_xcmplx_0(X1,X1)))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~v1_xreal_0(k7_real_1(np__1,k3_xcmplx_0(X1,X1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_1502, plain, (v1_xreal_0(k3_square_1(X1))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1528, plain, (v1_xreal_0(X1)|~m1_subset_1(k4_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_798, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1524, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_865, plain, (k9_xtuple_0(k19_sin_cos)=k1_relset_1(k1_numbers,k19_sin_cos)|~v1_relat_1(k19_sin_cos))).#trainneg
cnf(i_0_325, plain, (v1_xcmplx_0(k1_sin_cos9)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_366, plain, (v3_membered(k16_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1339, plain, (v3_xxreal_0(np__1)|~v2_xxreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1491, negated_conjecture, (v1_xxreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_337, plain, (v7_ordinal1(k1_real_1(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_351, plain, (v1_int_1(k5_square_1(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_326, plain, (v7_ordinal1(esk7_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_336, plain, (v7_ordinal1(k19_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1561, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_sin_cos9,k16_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k1_sin_cos9,k16_sin_cos))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1593, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos),X2)|~v1_xreal_0(esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(k2_fdiff_1(X1,X2),X2)|~r2_fdiff_1(X1,X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2))).#trainneg
cnf(i_0_1577, plain, (r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,X1,k16_sin_cos),k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,X1,k16_sin_cos)))|~v3_valued_0(X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,X1,k16_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,X1,k16_sin_cos)))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,X1,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,X1,k16_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,X1,k16_sin_cos)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_267, plain, (v3_xxreal_0(k4_xcmplx_0(np__2)))).#trainneg
cnf(i_0_1619, plain, (v1_int_1(k4_xcmplx_0(np__2))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_340, plain, (v7_ordinal1(k16_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_367, plain, (v3_membered(k1_sin_cos9)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1492, plain, (v1_xxreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1539, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1681, negated_conjecture, (v1_xxreal_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_341, plain, (v7_ordinal1(k1_sin_cos9)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_342, plain, (v1_int_1(esk7_1(X1))|~v5_membered(X1))).#trainneg
cnf(i_0_1615, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos),X7)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos))|~v1_funct_1(X5)|~r2_hidden(esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)),X6)|~r2_fdiff_1(k2_fdiff_1(X5,X6),X7)|~r2_fdiff_1(X5,X6)|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7)|~v1_xreal_0(esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos))))).#trainneg
cnf(i_0_1574, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos)))|~v1_xreal_0(esk4_2(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos)),k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos))))).#trainneg
cnf(i_0_1598, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos),X7)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos))|~v1_funct_1(X5)|~r2_hidden(esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)),X6)|~r2_fdiff_1(k2_fdiff_1(X5,X6),X7)|~r2_fdiff_1(X5,X6)|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)))|~m1_subset_1(esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)),k1_numbers)|~m1_subset_1(k2_fdiff_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7)|~v1_xreal_0(esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos))))).#trainneg
cnf(i_0_1639, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos),X3)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos))|~r2_hidden(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos)),X2)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),X3)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_xreal_0(k1_seq_1(k19_sin_cos,k1_seq_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos)))))|~v1_xreal_0(k1_seq_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos))),k1_numbers)|~m1_subset_1(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1563, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),X6)|~v3_valued_0(X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos))|~v1_funct_1(X5)|~r2_fdiff_1(X5,X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos)))))).#trainneg
cnf(i_0_1645, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos),X2)|~v1_xreal_0(esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos))|~r2_hidden(esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos)),X1)|~r2_fdiff_1(k2_fdiff_1(k16_sin_cos,X1),X2)|~r2_fdiff_1(k16_sin_cos,X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2))).#trainneg
cnf(i_0_799, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1485, plain, (v1_xxreal_0(k4_xcmplx_0(np__2))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_352, plain, (v1_int_1(k19_sin_cos)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_371, plain, (v1_membered(k19_sin_cos)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1540, plain, (v1_xcmplx_0(k5_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1616, negated_conjecture, (v1_int_1(esk1_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_353, plain, (v1_int_1(k1_real_1(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_356, plain, (v1_int_1(k16_sin_cos)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_357, plain, (v1_int_1(k1_sin_cos9)|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1575, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5)|~v1_xreal_0(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos)))))).#trainneg
cnf(i_0_1592, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos),X3)|v1_xboole_0(X2)|~v1_xreal_0(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|~r2_fdiff_1(X1,X2)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~m1_subset_1(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos)),X2))).#trainneg
cnf(i_0_1614, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos),X6)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos))|~r2_hidden(esk4_2(X6,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos)),X5)|~r2_fdiff_1(k2_fdiff_1(k16_sin_cos,X5),X6)|~r2_fdiff_1(k16_sin_cos,X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(k19_sin_cos,esk4_2(X6,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(k16_sin_cos,X5),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~v1_xreal_0(esk4_2(X6,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos))))).#trainneg
cnf(i_0_800, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1662, plain, (v1_xcmplx_0(k3_square_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1487, plain, (v7_ordinal1(k4_xcmplx_0(np__2))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_358, plain, (v1_relat_1(esk7_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_372, plain, (v1_membered(k16_sin_cos)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1621, plain, (v1_int_1(k5_numbers)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_373, plain, (v1_membered(k1_sin_cos9)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1268, plain, (v1_int_1(k5_square_1(X1))|~v1_int_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_362, plain, (v3_membered(esk7_1(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_368, plain, (v1_membered(esk7_1(k1_zfmisc_1(X1)))|~v1_membered(X1))).#trainneg
cnf(i_0_1234, plain, (k8_real_1(k1_fdiff_1(X1,X2),k1_seq_1(k19_sin_cos,k1_seq_1(X1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~v1_xreal_0(k1_fdiff_1(X1,X2))|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(X1,X2)),k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_752, plain, (m1_subset_1(esk4_2(X1,X2),X1)|r2_fdiff_1(X2,X1)|~v1_funct_1(X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_830, plain, (r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|m1_subset_1(esk4_2(X3,k2_fdiff_1(X1,X2)),k1_numbers)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(X3,X2)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_831, plain, (r1_tarski(X1,X3)|~v1_funct_1(k2_fdiff_1(X2,X3))|~r2_fdiff_1(k2_fdiff_1(X2,X3),X1)|~m1_subset_1(k2_fdiff_1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_801, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_807, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1533, plain, (v1_xxreal_0(k4_xcmplx_0(np__1))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_376, plain, (r2_fdiff_1(X1,X2)|m1_subset_1(esk4_2(X2,X1),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_385, plain, (v2_membered(k19_sin_cos)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_386, plain, (v2_membered(k16_sin_cos)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_387, plain, (v2_membered(k1_sin_cos9)|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1623, plain, (v1_xxreal_0(k5_square_1(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_377, plain, (k3_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_1820, plain, (v1_xreal_0(esk4_2(X1,X2))|r2_fdiff_1(X2,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1825, plain, (r2_fdiff_1(k16_sin_cos,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k16_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_802, plain, (v1_int_1(X1)|~v5_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1767, plain, (v1_membered(k19_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1799, plain, (v1_membered(k16_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1800, plain, (v1_membered(k1_sin_cos9)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1625, plain, (v1_int_1(k5_square_1(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_382, plain, (v2_membered(esk7_1(k1_zfmisc_1(X1)))|~v2_membered(X1))).#trainneg
cnf(i_0_1706, negated_conjecture, (v1_xreal_0(esk4_2(esk1_0,X1))|r2_fdiff_1(X1,esk1_0)|~v1_funct_1(X1)|~r1_tarski(esk1_0,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1815, plain, (v1_xreal_0(esk4_2(X1,k2_fdiff_1(X2,X3)))|r2_fdiff_1(k2_fdiff_1(X2,X3),X1)|~v1_funct_1(X2)|~r2_fdiff_1(X2,X3)|~r1_tarski(X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1791, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos),X6)|~v1_xreal_0(esk4_2(X6,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos))|~r2_hidden(esk4_2(X6,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos)),X5)|~r2_fdiff_1(k2_fdiff_1(k16_sin_cos,X5),X6)|~r2_fdiff_1(k16_sin_cos,X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k16_sin_cos,X5),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6))).#trainneg
cnf(i_0_1813, plain, (r2_fdiff_1(X1,X2)|m1_subset_1(esk4_2(X2,X1),X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_1754, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos),X3)|~v1_xreal_0(k1_seq_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos))|~v1_funct_1(X1)|~r2_hidden(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos)),X2)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),X3)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos)))),k1_numbers))).#trainneg
cnf(i_0_1781, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),esk1_0)|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos))))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1764, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos),X1)|~v1_xreal_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k16_sin_cos,X1),X1)|~r2_fdiff_1(k16_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_1751, plain, (r2_fdiff_1(k1_partfun1(X4,X5,X6,X7,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos),X3)|~v1_xreal_0(k1_seq_1(k19_sin_cos,k1_seq_1(X1,esk4_2(X3,k1_partfun1(X4,X5,X6,X7,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos)))))|~v1_xreal_0(k1_seq_1(X1,esk4_2(X3,k1_partfun1(X4,X5,X6,X7,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X3,k1_partfun1(X4,X5,X6,X7,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X4,X5,X6,X7,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos))|~v1_funct_1(X1)|~r2_hidden(esk4_2(X3,k1_partfun1(X4,X5,X6,X7,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos)),X2)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),X3)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(X4,X5,X6,X7,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X4,X5,X6,X7,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,esk4_2(X3,k1_partfun1(X4,X5,X6,X7,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2))|~m1_subset_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_1712, negated_conjecture, (r2_hidden(esk2_0,esk1_0)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k1_sin_cos9,k16_sin_cos)))).#trainneg
cnf(i_0_1780, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5)|~m1_subset_1(X5,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos))))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~m1_subset_1(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1765, plain, (v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1535, plain, (v7_ordinal1(k4_xcmplx_0(np__1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_388, plain, (v3_valued_0(esk7_1(k1_zfmisc_1(X1)))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1826, plain, (v2_membered(k19_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1827, plain, (v2_membered(k16_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1828, plain, (v2_membered(k1_sin_cos9)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1683, plain, (v1_xxreal_0(k1_real_1(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1685, plain, (v1_int_1(k1_real_1(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1711, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k1_sin_cos9,k16_sin_cos)))).#trainneg
cnf(i_0_1718, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos),X2)|~v1_xreal_0(esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos))|~v1_funct_1(X1)|~r2_fdiff_1(k2_fdiff_1(X1,X2),X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k16_sin_cos)))))).#trainneg
cnf(i_0_1735, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos),X6)|~v1_xreal_0(esk4_2(X6,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos))|~v1_funct_1(X5)|~r2_fdiff_1(k2_fdiff_1(X5,X6),X6)|~r2_fdiff_1(X5,X6)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6))).#trainneg
cnf(i_0_1748, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos),X6)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos))|~v1_funct_1(X5)|~r2_fdiff_1(k2_fdiff_1(X5,X6),X6)|~r2_fdiff_1(X5,X6)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)))|~m1_subset_1(esk4_2(X6,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)),k1_numbers)|~m1_subset_1(k2_fdiff_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1782, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1742, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos)))|~m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_rcomp_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos)))|~m1_subset_1(esk4_2(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos)),k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_1796, plain, (v1_xxreal_0(k4_xcmplx_0(np__2))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_218, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk7_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1708, negated_conjecture, (v1_xreal_0(X1)|v1_xboole_0(esk1_0)|~m1_subset_1(X1,esk1_0))).#trainneg
cnf(i_0_1936, plain, (v1_xxreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_395, plain, (v1_xxreal_0(k7_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1874, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos),X5)|~v1_xreal_0(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k16_sin_cos,X5),X5)|~r2_fdiff_1(k16_sin_cos,X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k16_sin_cos,X5),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5))).#trainneg
cnf(i_0_1852, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5)|~m1_subset_1(X5,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos))))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1792, plain, (v1_xcmplx_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1938, plain, (v1_xxreal_0(k4_xcmplx_0(np__1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_396, plain, (v1_xcmplx_0(k7_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1620, plain, (v3_membered(k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_397, plain, (v7_ordinal1(k7_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_398, plain, (v1_int_1(k7_real_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1916, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos),X1)|~v1_xreal_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k16_sin_cos,X1),X1)|~r2_fdiff_1(k16_sin_cos,X1)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k16_sin_cos,X1),k16_sin_cos)))))).#trainneg
cnf(i_0_1910, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos),X3)|~v1_xreal_0(k1_seq_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos))|~v1_funct_1(X1)|~r2_hidden(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos)),X2)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),X3)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2),k16_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3))).#trainneg
cnf(i_0_1856, plain, (v1_int_1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1256, plain, (k5_xcmplx_0(k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_630, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k2_fdiff_1(X2,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1680, plain, (v3_membered(k16_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_400, plain, (m1_subset_1(k7_real_1(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_401, plain, (v1_xreal_0(k10_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1924, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),k16_sin_cos),X7)|~v1_xreal_0(k1_seq_1(X5,esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),k16_sin_cos))))|~r1_fdiff_1(X5,esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),k16_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos))|~v1_funct_1(X5)|~r2_hidden(esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),k16_sin_cos)),X6)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),X7)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6)|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X5,esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),k16_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7)|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(X5,esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),k16_sin_cos)))),k1_numbers))).#trainneg
cnf(i_0_1808, plain, (r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|m1_subset_1(esk4_2(X3,k2_fdiff_1(X1,X2)),X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~r1_tarski(X3,X2)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1435, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1727, plain, (v1_xxreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_402, plain, (v1_xxreal_0(k10_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1985, negated_conjecture, (m1_subset_1(esk2_0,esk1_0))).#trainneg
cnf(i_0_1984, negated_conjecture, (~v1_xboole_0(esk1_0))).#trainneg
cnf(i_0_403, plain, (v1_xcmplx_0(k10_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_404, plain, (v7_ordinal1(k10_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2078, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),k16_sin_cos),X7)|~v1_xreal_0(k1_seq_1(X5,esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),k16_sin_cos))))|~r1_fdiff_1(X5,esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),k16_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos))|~v1_funct_1(X5)|~r2_hidden(esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),k16_sin_cos)),X6)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),X7)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6)|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X5,esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),k16_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7))).#trainneg
cnf(i_0_2010, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos),X5)|~v1_xreal_0(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k16_sin_cos,X5),X5)|~r2_fdiff_1(k16_sin_cos,X5)|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k16_sin_cos,X5),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k16_sin_cos,X5),k16_sin_cos)))))).#trainneg
cnf(i_0_1990, negated_conjecture, (v1_xxreal_0(esk2_0)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_1991, negated_conjecture, (v1_xcmplx_0(esk2_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_1995, plain, (X3=k2_fdiff_1(X2,esk7_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1)|k1_relset_1(k1_numbers,X3)!=esk7_1(k1_zfmisc_1(X1))|~v1_funct_1(X3)|~v1_funct_1(X2)|~r2_fdiff_1(X2,esk7_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_405, plain, (v1_int_1(k10_real_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1992, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1993, negated_conjecture, (v1_int_1(esk2_0)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_407, plain, (v3_rcomp_1(k2_rcomp_1(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_2002, plain, (v1_xreal_0(k7_real_1(X2,X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1812, plain, (r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|m1_subset_1(esk4_2(k1_relset_1(k1_numbers,X1),X1),k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,X1)))).#trainneg
cnf(i_0_1707, negated_conjecture, (v1_xreal_0(esk3_3(X1,esk1_0,X2))|X2=k2_fdiff_1(X1,esk1_0)|k1_relset_1(k1_numbers,X2)!=esk1_0|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_933, plain, (k4_xcmplx_0(k4_xcmplx_0(X1))=X1|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2036, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X4,X5,X6,X7,X1,k16_sin_cos),X2),k16_sin_cos),X3)|~v1_xreal_0(k1_seq_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X4,X5,X6,X7,X1,k16_sin_cos),X2),k16_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X4,X5,X6,X7,X1,k16_sin_cos),X2),k16_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X4,X5,X6,X7,X1,k16_sin_cos),X2),k16_sin_cos))|~v1_funct_1(k1_partfun1(X4,X5,X6,X7,X1,k16_sin_cos))|~v1_funct_1(X1)|~r2_hidden(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X4,X5,X6,X7,X1,k16_sin_cos),X2),k16_sin_cos)),X2)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X4,X5,X6,X7,X1,k16_sin_cos),X2),X3)|~r2_fdiff_1(k1_partfun1(X4,X5,X6,X7,X1,k16_sin_cos),X2)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X4,X5,X6,X7,X1,k16_sin_cos),X2),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X4,X5,X6,X7,X1,k16_sin_cos),X2),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X4,X5,X6,X7,X1,k16_sin_cos),X2),k16_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(X4,X5,X6,X7,X1,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_1997, negated_conjecture, (v1_xboole_0(esk1_0)|v1_xreal_0(esk7_1(esk1_0)))).#trainneg
cnf(i_0_2055, plain, (X4=k2_fdiff_1(X2,k2_fdiff_1(X1,X3))|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|k1_relset_1(k1_numbers,X4)!=k2_fdiff_1(X1,X3)|~v1_funct_1(X4)|~v1_funct_1(X2)|~r2_fdiff_1(X2,k2_fdiff_1(X1,X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_418, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1726, plain, (v3_membered(k1_sin_cos9)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_784, negated_conjecture, (v1_membered(esk1_0)|~v1_membered(k2_rcomp_1(k1_real_1(np__1),np__1)))).#trainneg
cnf(i_0_785, negated_conjecture, (v2_membered(esk1_0)|~v2_membered(k2_rcomp_1(k1_real_1(np__1),np__1)))).#trainneg
cnf(i_0_983, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_sin_cos9),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_419, plain, (v1_xreal_0(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)))|~v1_xreal_0(k4_xcmplx_0(np__2))|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_423, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2175, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos),X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos))|~r2_fdiff_1(k19_sin_cos,X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos))))).#trainneg
cnf(i_0_2174, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos))|~r2_fdiff_1(k19_sin_cos,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos))))).#trainneg
cnf(i_0_2193, plain, (r2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1))|~m1_subset_1(k2_rcomp_1(k1_real_1(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(k1_real_1(np__1)))).#trainneg
cnf(i_0_2158, negated_conjecture, (v1_xxreal_0(esk2_0))).#trainneg
cnf(i_0_2148, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X8,X9,X10,X11,X5,k16_sin_cos),X6),k16_sin_cos),X7)|~v1_xreal_0(k1_seq_1(X5,esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X8,X9,X10,X11,X5,k16_sin_cos),X6),k16_sin_cos))))|~r1_fdiff_1(X5,esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X8,X9,X10,X11,X5,k16_sin_cos),X6),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X8,X9,X10,X11,X5,k16_sin_cos),X6),k16_sin_cos))|~v1_funct_1(k1_partfun1(X8,X9,X10,X11,X5,k16_sin_cos))|~v1_funct_1(X5)|~r2_hidden(esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X8,X9,X10,X11,X5,k16_sin_cos),X6),k16_sin_cos)),X6)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X8,X9,X10,X11,X5,k16_sin_cos),X6),X7)|~r2_fdiff_1(k1_partfun1(X8,X9,X10,X11,X5,k16_sin_cos),X6)|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X8,X9,X10,X11,X5,k16_sin_cos),X6),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X8,X9,X10,X11,X5,k16_sin_cos),X6),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X5,esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X8,X9,X10,X11,X5,k16_sin_cos),X6),k16_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(k1_partfun1(X8,X9,X10,X11,X5,k16_sin_cos),X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k1_partfun1(X8,X9,X10,X11,X5,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X10,X11)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X8,X9))))).#trainneg
cnf(i_0_2262, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k1_sin_cos9)))).#trainneg
cnf(i_0_2133, negated_conjecture, (v1_xcmplx_0(esk2_0))).#trainneg
cnf(i_0_426, plain, (v1_xreal_0(k3_square_1(X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2129, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(esk1_0))).#trainneg
cnf(i_0_2130, negated_conjecture, (v1_int_1(esk2_0)|~v5_membered(esk1_0))).#trainneg
cnf(i_0_2168, negated_conjecture, (v1_int_1(esk2_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2232, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos),X6)|~v1_xreal_0(k1_seq_1(k16_sin_cos,esk4_2(X6,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos))))|~r1_fdiff_1(k16_sin_cos,esk4_2(X6,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos))|~r2_hidden(esk4_2(X6,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos)),X5)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),X6)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(k19_sin_cos,esk4_2(X6,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos))),k1_numbers)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~v1_xreal_0(esk4_2(X6,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos))))).#trainneg
cnf(i_0_2257, plain, (v3_membered(k1_relset_1(k1_numbers,k1_sin_cos9)))).#trainneg
cnf(i_0_2234, negated_conjecture, (v1_xcmplx_0(esk7_1(esk1_0)))).#trainneg
cnf(i_0_429, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1857, plain, (v1_membered(k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_963, negated_conjecture, (v1_membered(esk1_0)|~v1_membered(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))).#trainneg
cnf(i_0_964, negated_conjecture, (v2_membered(esk1_0)|~v2_membered(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))).#trainneg
cnf(i_0_2258, plain, (v1_membered(k1_relset_1(k1_numbers,k1_sin_cos9))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_2259, plain, (v2_membered(k1_relset_1(k1_numbers,k1_sin_cos9))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_924, plain, (r2_fdiff_1(k2_fdiff_1(X1,X2),X2)|m1_subset_1(esk4_2(X2,k2_fdiff_1(X1,X2)),k1_numbers)|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1221, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k1_sin_cos9,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1163, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,X1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1459, negated_conjecture, (k10_real_1(np__1,k2_xcmplx_0(np__1,k3_xcmplx_0(X1,X1)))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r2_hidden(X1,esk1_0)|~v1_xreal_0(k3_xcmplx_0(X1,X1)))).#trainneg
cnf(i_0_1968, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos),X6)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos))|~v1_funct_1(X5)|~r2_fdiff_1(k2_fdiff_1(X5,X6),X6)|~r2_fdiff_1(X5,X6)|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),k16_sin_cos)))))).#trainneg
cnf(i_0_1278, plain, (m1_subset_1(k3_xcmplx_0(X1,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1858, plain, (v1_membered(k16_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1374, negated_conjecture, (v1_xreal_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1375, plain, (v1_xreal_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2171, plain, (v1_xcmplx_0(k1_seq_1(k19_sin_cos,X1))|~v1_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_986, plain, (v1_xreal_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2285, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos),X1)|~r1_fdiff_1(k19_sin_cos,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos)))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos))))).#trainneg
cnf(i_0_2289, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos),k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos)))|~v1_xreal_0(esk4_2(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos)),k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos))|~r2_fdiff_1(k19_sin_cos,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos))))).#trainneg
cnf(i_0_2290, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos),X1)|~v1_xreal_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos))|~r2_fdiff_1(k19_sin_cos,X1)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k16_sin_cos)))))).#trainneg
cnf(i_0_2281, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos),k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)))|~v1_xreal_0(esk4_2(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)),k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos))|~r2_fdiff_1(k19_sin_cos,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos))))).#trainneg
cnf(i_0_2286, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos),X5)|~r1_fdiff_1(k19_sin_cos,esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)))|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos))))).#trainneg
cnf(i_0_2282, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos),X5)|~v1_xreal_0(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos))|~r2_fdiff_1(k19_sin_cos,X5)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)))))).#trainneg
cnf(i_0_2299, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos),X10)|~v1_xreal_0(k1_seq_1(k16_sin_cos,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos))))|~r1_fdiff_1(k16_sin_cos,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos)),X9)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),X10)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9)|~r1_tarski(X10,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(k19_sin_cos,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X10)|~v1_xreal_0(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos))))).#trainneg
cnf(i_0_2300, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos),X11)|~v1_xreal_0(k1_seq_1(X9,esk4_2(X11,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos))))|~r1_fdiff_1(X9,esk4_2(X11,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos))|~v1_funct_1(X9)|~r2_hidden(esk4_2(X11,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos)),X10)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),X11)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10)|~r1_tarski(X11,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X11,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X11)|~v1_xreal_0(esk4_2(X11,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos))))).#trainneg
cnf(i_0_2292, plain, (r2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2307, negated_conjecture, (v1_int_1(esk2_0)|~v6_membered(esk1_0))).#trainneg
cnf(i_0_441, plain, (v1_xcmplx_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1859, plain, (v1_membered(k1_sin_cos9)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1545, plain, (v1_xxreal_0(k19_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1585, plain, (v1_xxreal_0(k16_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2284, plain, (v1_xcmplx_0(k1_seq_1(k19_sin_cos,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2313, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos),X6)|~v1_xreal_0(esk4_2(X6,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos))|~r2_hidden(esk4_2(X6,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos)),X5)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),X6)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk4_2(X6,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_2381, negated_conjecture, (k7_xcmplx_0(np__1,k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))=k7_real_1(np__1,k5_square_1(X1))|~v1_xcmplx_0(k7_real_1(np__1,k5_square_1(X1)))|~v1_xreal_0(k7_real_1(np__1,k5_square_1(X1)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_2319, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_2318, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(k7_xcmplx_0(np__1,X1)))).#trainneg
cnf(i_0_487, plain, (v1_int_1(k4_xcmplx_0(np__2))|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1941, plain, (v2_membered(k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1586, plain, (v1_xxreal_0(k1_sin_cos9)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1627, plain, (v1_xcmplx_0(k19_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_987, plain, (v1_xxreal_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_989, plain, (v7_ordinal1(k1_relset_1(k1_numbers,k16_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_450, plain, (v1_xcmplx_0(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)))|~v1_xcmplx_0(k4_xcmplx_0(np__2))|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_456, plain, (v1_xcmplx_0(k3_square_1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1003, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k1_sin_cos9),k1_zfmisc_1(k2_zfmisc_1(X6,k1_numbers)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k1_sin_cos9,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_749, plain, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1)),X1)=k10_real_1(np__1,k7_real_1(np__1,k5_square_1(X1)))|~r2_hidden(X1,k2_rcomp_1(k1_real_1(np__1),np__1))|~m1_subset_1(k2_rcomp_1(k1_real_1(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k1_real_1(np__1),np__1)))).#trainneg
cnf(i_0_2481, plain, (m1_subset_1(k3_square_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2469, negated_conjecture, (m1_subset_1(esk7_1(esk1_0),k1_numbers))).#trainneg
cnf(i_0_2577, plain, (k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1))=k2_fdiff_1(k16_sin_cos,X1)|k10_real_1(np__1,k7_real_1(np__1,k5_square_1(esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1))))))!=k1_seq_1(k19_sin_cos,esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1))))|k1_relset_1(k1_numbers,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1)))!=X1|~v1_xreal_0(esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1))))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1)))|~r2_fdiff_1(k16_sin_cos,X1)|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_hidden(esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1))),k2_rcomp_1(k1_real_1(np__1),np__1))|~m1_subset_1(k2_rcomp_1(k1_real_1(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k1_real_1(np__1),np__1)))).#trainneg
cnf(i_0_1942, plain, (v2_membered(k16_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1628, plain, (v1_xcmplx_0(k16_sin_cos)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1679, plain, (v1_xcmplx_0(k1_sin_cos9)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2523, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos),X6)|~v1_xreal_0(esk4_2(X6,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos))|~r2_hidden(esk4_2(X6,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos)),X5)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),X6)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6))).#trainneg
cnf(i_0_2480, plain, (m1_subset_1(k3_square_1(X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2449, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos),X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos))|~r2_fdiff_1(k19_sin_cos,X5)|~m1_subset_1(X5,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos))))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~m1_subset_1(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_2450, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos))|~r2_fdiff_1(k19_sin_cos,esk1_0)|~m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos))))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2439, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos),k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos))|~r2_fdiff_1(k19_sin_cos,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)))|~m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_rcomp_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)))|~m1_subset_1(esk4_2(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)),k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_2440, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos),k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos))|~r2_fdiff_1(k19_sin_cos,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)))|~m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_rcomp_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2526, negated_conjecture, (k7_xcmplx_0(np__1,k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))=k7_real_1(np__1,k5_square_1(X1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k7_real_1(np__1,k5_square_1(X1)),k1_numbers))).#trainneg
cnf(i_0_1943, plain, (v2_membered(k1_sin_cos9)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1728, negated_conjecture, (v1_xcmplx_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1766, plain, (v1_int_1(k19_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_990, plain, (v1_int_1(k1_relset_1(k1_numbers,k16_sin_cos))|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2497, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2459, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos),X10)|~v1_xreal_0(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos)),X9)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),X10)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9)|~r1_tarski(X10,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X10)|~m1_subset_1(k1_seq_1(k16_sin_cos,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_2467, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos),X11)|~v1_xreal_0(esk4_2(X11,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos)))|~r1_fdiff_1(X9,esk4_2(X11,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos))|~v1_funct_1(X9)|~r2_hidden(esk4_2(X11,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos)),X10)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),X11)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10)|~r1_tarski(X11,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X11,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X11)|~m1_subset_1(k1_seq_1(X9,esk4_2(X11,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_2468, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos),X9),k16_sin_cos),X10)|~v1_xreal_0(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos),X9),k16_sin_cos)))|~r1_fdiff_1(k19_sin_cos,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos),X9),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos),X9),k16_sin_cos)),X9)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos),X9),X10)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos),X9)|~r1_tarski(X10,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X10))).#trainneg
cnf(i_0_2533, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_2578, plain, (k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1))=k2_fdiff_1(X1,X2)|k1_fdiff_1(X1,esk3_3(X1,X2,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1))))!=k10_real_1(np__1,k7_real_1(np__1,k5_square_1(esk3_3(X1,X2,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1))))))|k1_relset_1(k1_numbers,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1)))!=X2|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1)))|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_hidden(esk3_3(X1,X2,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1))),k2_rcomp_1(k1_real_1(np__1),np__1))|~m1_subset_1(k2_rcomp_1(k1_real_1(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k1_real_1(np__1),np__1)))).#trainneg
cnf(i_0_1768, plain, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1775, plain, (v1_int_1(k16_sin_cos)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2252, plain, (v1_xreal_0(k1_relset_1(k1_numbers,k1_sin_cos9))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2253, plain, (v1_xxreal_0(k1_relset_1(k1_numbers,k1_sin_cos9))|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2427, plain, (r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k16_sin_cos),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k16_sin_cos))|~r2_fdiff_1(k19_sin_cos,X1)|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k16_sin_cos))))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2576, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1)),k16_sin_cos),X5)|~v1_xreal_0(k10_real_1(np__1,k7_real_1(np__1,k5_square_1(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1)),k16_sin_cos))))))|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1)),esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1)),k16_sin_cos)))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1)))|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1)),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1)),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~r2_hidden(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1)),k16_sin_cos)),k2_rcomp_1(k1_real_1(np__1),np__1))|~m1_subset_1(k2_rcomp_1(k1_real_1(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k1_real_1(np__1),np__1)))).#trainneg
cnf(i_0_2586, plain, (k10_real_1(np__1,k7_real_1(np__1,k5_square_1(X1)))=k1_fdiff_1(k1_sin_cos9,X1)|~r2_hidden(X1,k2_rcomp_1(k1_real_1(np__1),np__1))|~r2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k1_real_1(np__1),np__1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k2_rcomp_1(k1_real_1(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k1_real_1(np__1),np__1)))).#trainneg
cnf(i_0_2731, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)=k1_fdiff_1(k1_sin_cos9,X1)|~r2_hidden(X1,esk1_0)|~r2_hidden(X1,k2_rcomp_1(k1_real_1(np__1),np__1))|~m1_subset_1(k2_rcomp_1(k1_real_1(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k1_real_1(np__1),np__1)))).#trainneg
cnf(i_0_2674, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos),X10)|~v1_xreal_0(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos)),X9)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),X10)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9)|~r1_tarski(X10,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X10))).#trainneg
cnf(i_0_2646, negated_conjecture, (k7_xcmplx_0(np__1,k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))=k7_real_1(np__1,k5_square_1(X1))|~r2_hidden(X1,esk1_0)|~v1_xreal_0(k5_square_1(X1)))).#trainneg
cnf(i_0_2596, negated_conjecture, (v1_xxreal_0(esk7_1(esk1_0))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_2600, plain, (k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))=k2_fdiff_1(k16_sin_cos,X1)|k10_real_1(np__1,k7_real_1(np__1,k5_square_1(esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))))!=k1_seq_1(k19_sin_cos,esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))|k1_relset_1(k1_numbers,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))!=X1|~v1_xreal_0(esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))|~r2_hidden(esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))),k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r2_fdiff_1(k16_sin_cos,X1)|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))).#trainneg
cnf(i_0_478, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_1776, plain, (v1_int_1(k1_sin_cos9)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2255, plain, (v7_ordinal1(k1_relset_1(k1_numbers,k1_sin_cos9))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_479, plain, (v1_xcmplx_0(k4_xcmplx_0(k3_xcmplx_0(X1,X2)))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_482, plain, (v1_int_1(k3_square_1(X1))|~v1_int_1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2684, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos),X11)|~v1_xreal_0(esk4_2(X11,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos))|~v1_funct_1(X9)|~r2_hidden(esk4_2(X11,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos)),X10)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),X11)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10)|~r1_tarski(X11,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X9,esk4_2(X11,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X11,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X11)|~r2_fdiff_1(X9,X11))).#trainneg
cnf(i_0_2612, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos),X5)|~v1_xreal_0(esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),X5)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5))).#trainneg
cnf(i_0_2681, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X9,X10),k16_sin_cos),X11),k16_sin_cos),X12)|~v1_xreal_0(esk4_2(X12,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X9,X10),k16_sin_cos),X11),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X9,X10),k16_sin_cos),X11),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X9,X10),k16_sin_cos))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X9,X10))|~r2_hidden(esk4_2(X12,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X9,X10),k16_sin_cos),X11),k16_sin_cos)),X11)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X9,X10),k16_sin_cos),X11),X12)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X9,X10),k16_sin_cos),X11)|~r1_tarski(X12,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X9,X10),k16_sin_cos),X11),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X9,X10),k16_sin_cos),X11),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X9,X10),esk4_2(X12,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X9,X10),k16_sin_cos),X11),k16_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X9,X10),k16_sin_cos),X11),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X9,X10),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X9,X10),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X12,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X12)|~r1_fdiff_1(X10,k1_seq_1(X9,esk4_2(X12,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X9,X10),k16_sin_cos),X11),k16_sin_cos))))|~r1_fdiff_1(X9,esk4_2(X12,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X9,X10),k16_sin_cos),X11),k16_sin_cos)))|~v1_funct_1(X10)|~v1_funct_1(X9)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X12,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X9,X10),k16_sin_cos),X11),k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_2598, negated_conjecture, (v7_ordinal1(esk7_1(esk1_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2735, plain, (k10_real_1(np__1,k2_xcmplx_0(np__1,k5_square_1(X1)))=k1_fdiff_1(k1_sin_cos9,X1)|~r2_hidden(X1,k2_rcomp_1(k1_real_1(np__1),np__1))|~m1_subset_1(k2_rcomp_1(k1_real_1(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k1_real_1(np__1),np__1))|~v1_xreal_0(k5_square_1(X1)))).#trainneg
cnf(i_0_838, negated_conjecture, (v1_membered(esk1_0)|~v1_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))))).#trainneg
cnf(i_0_2249, negated_conjecture, (v1_membered(esk1_0)|~v3_membered(k2_rcomp_1(k1_real_1(np__1),np__1)))).#trainneg
cnf(i_0_2256, plain, (v1_int_1(k1_relset_1(k1_numbers,k1_sin_cos9))|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2419, plain, (v1_xcmplx_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_485, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2724, plain, (k7_xcmplx_0(np__1,k7_real_1(np__1,k5_square_1(X1)))=k1_fdiff_1(k1_sin_cos9,X1)|~v1_xreal_0(k7_real_1(np__1,k5_square_1(X1)))|~r2_hidden(X1,k2_rcomp_1(k1_real_1(np__1),np__1))|~m1_subset_1(k2_rcomp_1(k1_real_1(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k1_real_1(np__1),np__1)))).#trainneg
cnf(i_0_2647, negated_conjecture, (k7_xcmplx_0(np__1,k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))=k7_real_1(np__1,k5_square_1(X1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers))).#trainneg
cnf(i_0_2691, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos),X9),k16_sin_cos),X10)|~v1_xreal_0(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos),X9),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos),X9),k16_sin_cos)),X9)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos),X9),X10)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos),X9)|~r1_tarski(X10,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k19_sin_cos,k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X10)|~r2_fdiff_1(k19_sin_cos,X10))).#trainneg
cnf(i_0_2491, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2745, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos),X9)|~v1_xreal_0(esk4_2(X9,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),X9)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9)|~r1_tarski(X9,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X9,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X9))).#trainneg
cnf(i_0_2758, negated_conjecture, (k7_real_1(np__1,k5_square_1(X1))=k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2812, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos),X5)|~v1_xreal_0(esk4_2(X5,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),X5)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k16_sin_cos,k16_sin_cos),X5),k16_sin_cos)))))).#trainneg
cnf(i_0_2737, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)=k1_fdiff_1(k1_sin_cos9,X1)|~r2_hidden(X1,k2_rcomp_1(k1_real_1(np__1),np__1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k2_rcomp_1(k1_real_1(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(k1_real_1(np__1)))).#trainneg
cnf(i_0_2751, negated_conjecture, (k7_xcmplx_0(np__1,k7_real_1(np__1,k5_square_1(X1)))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~v1_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_2599, negated_conjecture, (v1_int_1(esk7_1(esk1_0))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_2759, plain, (k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))=k2_fdiff_1(k16_sin_cos,X1)|k10_real_1(np__1,k2_xcmplx_0(np__1,k5_square_1(esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))))!=k1_seq_1(k19_sin_cos,esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))|k1_relset_1(k1_numbers,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))!=X1|~v1_xreal_0(esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))|~r2_hidden(esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))),k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r2_fdiff_1(k16_sin_cos,X1)|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~v1_xreal_0(k5_square_1(esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))))).#trainneg
cnf(i_0_486, plain, (v1_xcmplx_0(k4_xcmplx_0(k2_xcmplx_0(X1,X2)))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2251, negated_conjecture, (v2_membered(esk1_0)|~v3_membered(k2_rcomp_1(k1_real_1(np__1),np__1)))).#trainneg
cnf(i_0_2335, negated_conjecture, (v1_membered(esk1_0)|~v3_membered(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))).#trainneg
cnf(i_0_2804, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k16_sin_cos),X12)|~v1_xreal_0(esk4_2(X12,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(X9,X10))|~r2_hidden(esk4_2(X12,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k16_sin_cos)),X11)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),X12)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11)|~r2_fdiff_1(k2_fdiff_1(X9,X10),X12)|~r1_tarski(X12,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X9,esk4_2(X12,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k16_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_fdiff_1(X9,X10),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k2_fdiff_1(X9,X10),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X12,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X12)|~v1_funct_1(X9)|~r2_hidden(esk4_2(X12,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k16_sin_cos)),X10)|~r2_fdiff_1(X9,X10)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X12,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_2868, negated_conjecture, (k7_real_1(np__1,k5_square_1(X1))=k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2466, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k16_sin_cos),X12)|~v1_xreal_0(k1_fdiff_1(X9,esk4_2(X12,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k16_sin_cos))))|~v1_xreal_0(esk4_2(X12,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k16_sin_cos)))|~r1_fdiff_1(k2_fdiff_1(X9,X10),esk4_2(X12,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(X9,X10))|~r2_hidden(esk4_2(X12,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k16_sin_cos)),X11)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),X12)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11)|~r1_tarski(X12,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_fdiff_1(X9,X10),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k2_fdiff_1(X9,X10),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X12,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X12)|~v1_funct_1(X9)|~r2_hidden(esk4_2(X12,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k16_sin_cos)),X10)|~r2_fdiff_1(X9,X10)|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X12,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(X9,X10),k16_sin_cos),X11),k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_2494, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_789, negated_conjecture, (~v1_xboole_0(k2_rcomp_1(k1_real_1(np__1),np__1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_2734, plain, (k10_real_1(np__1,k7_real_1(np__1,k3_xcmplx_0(X1,X1)))=k1_fdiff_1(k1_sin_cos9,X1)|~r2_hidden(X1,k2_rcomp_1(k1_real_1(np__1),np__1))|~m1_subset_1(k2_rcomp_1(k1_real_1(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k1_real_1(np__1),np__1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2826, negated_conjecture, (v1_xxreal_0(esk7_1(esk1_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2760, plain, (k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))=k2_fdiff_1(k16_sin_cos,X1)|k7_xcmplx_0(np__1,k7_real_1(np__1,k5_square_1(esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))))!=k1_seq_1(k19_sin_cos,esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))|k1_relset_1(k1_numbers,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))!=X1|~v1_xreal_0(esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))|~r2_hidden(esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))),k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r2_fdiff_1(k16_sin_cos,X1)|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))))))).#trainneg
cnf(i_0_839, negated_conjecture, (v2_membered(esk1_0)|~v2_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))))).#trainneg
cnf(i_0_2336, negated_conjecture, (v2_membered(esk1_0)|~v3_membered(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))).#trainneg
cnf(i_0_2553, plain, (v1_xxreal_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2555, plain, (v1_int_1(k1_relset_1(k1_numbers,k16_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_498, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2999, negated_conjecture, (k10_real_1(np__1,k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers))).#trainneg
cnf(i_0_2923, negated_conjecture, (k10_real_1(np__1,k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r2_hidden(X1,esk1_0)|~v1_xreal_0(k5_square_1(X1)))).#trainneg
cnf(i_0_2991, negated_conjecture, (k10_real_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k5_square_1(esk2_0),k1_numbers))).#trainneg
cnf(i_0_2916, negated_conjecture, (k10_real_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~v1_xreal_0(k5_square_1(esk2_0)))).#trainneg
cnf(i_0_3000, negated_conjecture, (k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers))).#trainneg
cnf(i_0_2924, negated_conjecture, (k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))|~r2_hidden(X1,esk1_0)|~v1_xreal_0(k5_square_1(X1)))).#trainneg
cnf(i_0_2993, negated_conjecture, (v1_xreal_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)),k1_numbers)|~m1_subset_1(k5_square_1(X1),k1_numbers))).#trainneg
cnf(i_0_2902, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos),X9)|~v1_xreal_0(esk4_2(X9,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),X9)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9)|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X9,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X9)|~m1_subset_1(X9,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k16_sin_cos,k16_sin_cos),X9),k16_sin_cos)))))).#trainneg
cnf(i_0_2992, negated_conjecture, (k7_xcmplx_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_numbers)|~m1_subset_1(k5_square_1(esk2_0),k1_numbers))).#trainneg
cnf(i_0_2918, negated_conjecture, (v1_xreal_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)),k1_numbers)|~v1_xreal_0(k5_square_1(X1)))).#trainneg
cnf(i_0_2917, negated_conjecture, (k7_xcmplx_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_numbers)|~v1_xreal_0(k5_square_1(esk2_0)))).#trainneg
cnf(i_0_2828, negated_conjecture, (v1_int_1(esk7_1(esk1_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2990, negated_conjecture, (k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))=k2_fdiff_1(k16_sin_cos,X1)|k10_real_1(np__1,k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))))!=k1_seq_1(k19_sin_cos,esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))|k1_relset_1(k1_numbers,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))!=X1|~v1_xreal_0(esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))|~r2_hidden(esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))),k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r2_fdiff_1(k16_sin_cos,X1)|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r2_hidden(esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))),esk1_0)|~m1_subset_1(k5_square_1(esk3_3(k16_sin_cos,X1,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))),k1_numbers))).#trainneg
cnf(i_0_499, plain, (v1_xcmplx_0(k4_xcmplx_0(k6_xcmplx_0(X1,X2)))|~v1_int_1(X2)|~v1_int_1(X1))).#trainneg
cnf(i_0_2833, negated_conjecture, (v1_membered(esk1_0)|~v3_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))))).#trainneg
cnf(i_0_778, negated_conjecture, (v1_xreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k2_rcomp_1(k1_real_1(np__1),np__1))))).#trainneg
cnf(i_0_2709, plain, (v1_xcmplx_0(k1_relset_1(k1_numbers,k1_sin_cos9))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2778, plain, (v1_xxreal_0(k1_relset_1(k1_numbers,k1_sin_cos9))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2998, negated_conjecture, (v1_xreal_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers))).#trainneg
cnf(i_0_2922, negated_conjecture, (v1_xreal_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))|~r2_hidden(X1,esk1_0)|~v1_xreal_0(k5_square_1(X1)))).#trainneg
cnf(i_0_2919, negated_conjecture, (k7_xcmplx_0(np__1,k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))=k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1))|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))|~r2_hidden(X1,esk1_0)|~v1_xreal_0(k5_square_1(X1)))).#trainneg
cnf(i_0_2994, negated_conjecture, (k7_xcmplx_0(np__1,k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))=k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1))|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers))).#trainneg
cnf(i_0_2949, negated_conjecture, (k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~v1_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_2904, negated_conjecture, (k2_xcmplx_0(np__1,k5_square_1(X1))=k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_929, plain, (v1_int_1(X1)|~v1_int_1(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_806, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_3036, negated_conjecture, (v2_membered(esk1_0)|~v3_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos))))).#trainneg
cnf(i_0_779, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k1_zfmisc_1(k2_rcomp_1(k1_real_1(np__1),np__1))))).#trainneg
cnf(i_0_780, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k1_zfmisc_1(k2_rcomp_1(k1_real_1(np__1),np__1))))).#trainneg
cnf(i_0_2780, plain, (v1_int_1(k1_relset_1(k1_numbers,k1_sin_cos9))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_550, plain, (v3_membered(k2_rcomp_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_508, plain, (v4_relat_1(esk7_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_512, plain, (v5_relat_1(esk7_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_3078, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)=k1_fdiff_1(k1_sin_cos9,X1)|~v1_xcmplx_0(k1_fdiff_1(k1_sin_cos9,X1))|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers))).#trainneg
cnf(i_0_3180, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)=k1_fdiff_1(k1_sin_cos9,X1)|~v1_xcmplx_0(k1_fdiff_1(k1_sin_cos9,X1))|~v1_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_3130, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X1)|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))))|~r2_hidden(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~m1_subset_1(k5_square_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3091, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)=k1_fdiff_1(k1_sin_cos9,X1)|~v1_xcmplx_0(k1_fdiff_1(k1_sin_cos9,X1))|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_3136, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X1)|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))))|~v1_xreal_0(k5_square_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~r2_hidden(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0))).#trainneg
cnf(i_0_3137, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),X10)|~v1_xreal_0(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),X9)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),X10)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9)|~r1_tarski(X10,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X10)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))))|~v1_xreal_0(k5_square_1(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),esk1_0))).#trainneg
cnf(i_0_3094, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X1)|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~r2_hidden(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))),k1_numbers)|~m1_subset_1(k5_square_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3131, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),X10)|~v1_xreal_0(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),X9)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),X10)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9)|~r1_tarski(X10,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X10)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),esk1_0)|~m1_subset_1(k5_square_1(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3095, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),X10)|~v1_xreal_0(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),X9)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),X10)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9)|~r1_tarski(X10,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X10)|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),esk1_0)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))),k1_numbers)|~m1_subset_1(k5_square_1(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3106, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X1)|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(k5_square_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~r2_hidden(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))),k1_numbers))).#trainneg
cnf(i_0_809, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_953, plain, (k1_real_1(np__2)=k4_xcmplx_0(np__2)|~m1_subset_1(k4_xcmplx_0(np__2),k1_numbers))).#trainneg
cnf(i_0_516, plain, (v3_valued_0(esk7_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v3_membered(X2))).#trainneg
cnf(i_0_781, negated_conjecture, (v7_ordinal1(esk1_0)|~v6_membered(k1_zfmisc_1(k2_rcomp_1(k1_real_1(np__1),np__1))))).#trainneg
cnf(i_0_782, negated_conjecture, (v1_int_1(esk1_0)|~v5_membered(k1_zfmisc_1(k2_rcomp_1(k1_real_1(np__1),np__1))))).#trainneg
cnf(i_0_994, plain, (v3_valued_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_995, plain, (v1_valued_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_520, plain, (v1_xreal_0(k8_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_521, plain, (v1_xxreal_0(k8_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3107, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),X10)|~v1_xreal_0(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),X9)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),X10)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9)|~r1_tarski(X10,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X10)|~v1_xreal_0(k5_square_1(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),esk1_0)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))),k1_numbers))).#trainneg
cnf(i_0_3132, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X5)|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))))|~r2_hidden(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~m1_subset_1(k5_square_1(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3138, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X5)|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))))|~v1_xreal_0(k5_square_1(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~r2_hidden(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0))).#trainneg
cnf(i_0_3096, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X5)|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~r2_hidden(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))),k1_numbers)|~m1_subset_1(k5_square_1(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3108, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X5)|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(k5_square_1(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~r2_hidden(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))),k1_numbers))).#trainneg
cnf(i_0_954, plain, (k1_real_1(np__1)=k4_xcmplx_0(np__1)|~m1_subset_1(k4_xcmplx_0(np__1),k1_numbers))).#trainneg
cnf(i_0_522, plain, (v1_xcmplx_0(k8_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3211, negated_conjecture, (v2_membered(esk1_0)|~v1_xxreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_3212, negated_conjecture, (v1_membered(esk1_0)|~v1_xxreal_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_3200, plain, (v1_int_1(X1)|~m1_subset_1(X1,k1_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2170, plain, (v1_xxreal_0(k1_seq_1(k19_sin_cos,X1))|~v2_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_523, plain, (v7_ordinal1(k8_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_524, plain, (v1_int_1(k8_real_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3226, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)=k1_fdiff_1(k1_sin_cos9,X1)|~v1_xcmplx_0(k1_fdiff_1(k1_sin_cos9,X1))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0)|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,X1)))).#trainneg
cnf(i_0_3219, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)=k1_fdiff_1(k1_sin_cos9,X1)|~v1_xcmplx_0(k1_fdiff_1(k1_sin_cos9,X1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers)|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,X1)))).#trainneg
cnf(i_0_3221, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)=k1_fdiff_1(k1_sin_cos9,X1)|~v1_xcmplx_0(k1_fdiff_1(k1_sin_cos9,X1))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1),k1_numbers))).#trainneg
cnf(i_0_3173, negated_conjecture, (k7_xcmplx_0(np__1,k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))=k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1))|~v1_xcmplx_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))|~v1_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_3051, negated_conjecture, (m1_subset_1(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1),k1_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers))).#trainneg
cnf(i_0_3313, plain, (k1_real_1(k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_3201, plain, (v3_valued_0(k2_zfmisc_1(X1,k1_numbers)))).#trainneg
cnf(i_0_3260, plain, (k1_real_1(k4_xcmplx_0(np__2))=np__2)).#trainneg
cnf(i_0_2731, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)=k1_fdiff_1(k1_sin_cos9,X1)|~r2_hidden(X1,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))).#trainneg
cnf(i_0_525, plain, (v1_xboole_0(esk7_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1227, plain, (v1_int_1(k4_xcmplx_0(X1))|~v5_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2172, plain, (v7_ordinal1(k1_seq_1(k19_sin_cos,X1))|~v6_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2173, plain, (v1_int_1(k1_seq_1(k19_sin_cos,X1))|~v5_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_529, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_530, plain, (v1_valued_0(esk7_1(k1_zfmisc_1(X1)))|~v1_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3092, negated_conjecture, (v1_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)),k1_numbers)|~m1_subset_1(k5_square_1(X1),k1_numbers))).#trainneg
cnf(i_0_3058, negated_conjecture, (m1_subset_1(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1),k1_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_3068, negated_conjecture, (v1_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~v1_xcmplx_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers))).#trainneg
cnf(i_0_2737, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)=k1_fdiff_1(k1_sin_cos9,X1)|~r2_hidden(X1,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_533, plain, (v1_valued_0(k19_sin_cos)|~v1_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_2260, plain, (v3_valued_0(k1_relset_1(k1_numbers,k1_sin_cos9))|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_2261, plain, (v1_valued_0(k1_relset_1(k1_numbers,k1_sin_cos9))|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_534, plain, (v1_valued_0(k16_sin_cos)|~v1_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_535, plain, (v1_valued_0(k1_sin_cos9)|~v1_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_3390, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~r2_hidden(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~m1_subset_1(k5_square_1(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3366, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~v1_xreal_0(k5_square_1(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~r2_hidden(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0))).#trainneg
cnf(i_0_3384, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~r2_hidden(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~m1_subset_1(k5_square_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3387, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),X10)|~v1_xreal_0(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),X9)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),X10)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9)|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X10)|~r1_tarski(X10,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X10)|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),esk1_0)|~m1_subset_1(k5_square_1(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3407, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)=k1_fdiff_1(k1_sin_cos9,X1)|~v1_xcmplx_0(k1_fdiff_1(k1_sin_cos9,X1))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0)|~v3_valued_0(k2_fdiff_1(k1_sin_cos9,esk1_0))|~v1_relat_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0)))).#trainneg
cnf(i_0_3360, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))),k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~v1_xreal_0(k5_square_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~r2_hidden(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0))).#trainneg
cnf(i_0_3363, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),X10)|~v1_xreal_0(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),X9)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),X10)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9)|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X10)|~r1_tarski(X10,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X10)|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))))|~v1_xreal_0(k5_square_1(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),esk1_0))).#trainneg
cnf(i_0_3362, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),X10)|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))))|~v1_xreal_0(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),X9)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),X10)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9)|~r1_tarski(X10,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X10)|~v1_xreal_0(k5_square_1(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),esk1_0))).#trainneg
cnf(i_0_3361, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X1)|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(k5_square_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~r2_hidden(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0))).#trainneg
cnf(i_0_3331, negated_conjecture, (v2_membered(esk1_0))).#trainneg
cnf(i_0_3335, negated_conjecture, (v1_membered(esk1_0))).#trainneg
cnf(i_0_3487, plain, (v1_valued_0(k19_sin_cos))).#trainneg
cnf(i_0_2578, plain, (k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))=k2_fdiff_1(X1,X2)|k1_fdiff_1(X1,esk3_3(X1,X2,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))!=k10_real_1(np__1,k7_real_1(np__1,k5_square_1(esk3_3(X1,X2,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))))|k1_relset_1(k1_numbers,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))!=X2|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))|~v1_funct_1(X1)|~r2_hidden(esk3_3(X1,X2,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))),k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r2_fdiff_1(X1,X2)|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))).#trainneg
cnf(i_0_536, plain, (v1_xboole_0(esk7_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3488, plain, (v1_valued_0(k16_sin_cos))).#trainneg
cnf(i_0_957, negated_conjecture, (v1_xreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))).#trainneg
cnf(i_0_958, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k1_zfmisc_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))).#trainneg
cnf(i_0_541, plain, (v1_xxreal_0(k1_seq_1(X1,X2))|~v2_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3473, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)=k1_fdiff_1(k1_sin_cos9,X1)|~v1_xcmplx_0(k1_fdiff_1(k1_sin_cos9,X1))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)),k1_numbers)|~m1_subset_1(k5_square_1(X1),k1_numbers))).#trainneg
cnf(i_0_3385, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X1)|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~r2_hidden(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~m1_subset_1(k5_square_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3386, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),X10)|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))))|~v1_xreal_0(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),X9)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),X10)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9)|~r1_tarski(X10,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X10)|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),esk1_0)|~m1_subset_1(k5_square_1(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3391, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X5)|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~r2_hidden(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~m1_subset_1(k5_square_1(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3367, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X5)|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(k5_square_1(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~r2_hidden(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0))).#trainneg
cnf(i_0_2213, plain, (v1_int_1(X1)|~v1_int_1(k4_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_542, plain, (v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_959, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k1_zfmisc_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))).#trainneg
cnf(i_0_960, negated_conjecture, (v7_ordinal1(esk1_0)|~v6_membered(k1_zfmisc_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))).#trainneg
cnf(i_0_3451, plain, (v1_xxreal_0(k1_seq_1(k19_sin_cos,X1))|~v6_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3453, plain, (v1_int_1(k1_seq_1(k19_sin_cos,X1))|~v6_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_543, plain, (v7_ordinal1(k1_seq_1(X1,X2))|~v6_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_544, plain, (v1_int_1(k1_seq_1(X1,X2))|~v5_membered(k1_numbers)|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3397, negated_conjecture, (k8_real_1(k1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1),k1_seq_1(k19_sin_cos,k1_fdiff_1(k1_sin_cos9,X1)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X1)|~v1_xreal_0(k1_seq_1(k19_sin_cos,k1_fdiff_1(k1_sin_cos9,X1)))|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,X1))|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers))).#trainneg
cnf(i_0_3541, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos),X2)|~r1_fdiff_1(X1,esk4_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)))|~v1_funct_1(X1)|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k16_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X2)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3489, plain, (v1_valued_0(k1_sin_cos9))).#trainneg
cnf(i_0_545, plain, (v1_xreal_0(k2_rcomp_1(X1,X2))|~v3_membered(k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_961, negated_conjecture, (v1_int_1(esk1_0)|~v5_membered(k1_zfmisc_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))).#trainneg
cnf(i_0_546, plain, (v1_xxreal_0(k2_rcomp_1(X1,X2))|~v2_membered(k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_547, plain, (v1_xcmplx_0(k2_rcomp_1(X1,X2))|~v1_membered(k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_3520, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),X10)|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),esk1_0)|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),X9)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),X10)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9)|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X10)|~r1_tarski(X10,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X10)|~m1_subset_1(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_3546, plain, (v1_xreal_0(k1_fdiff_1(X1,X3))|~v3_valued_0(k2_fdiff_1(X1,X2))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~r2_hidden(X3,X2)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_3543, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),X6)|~r1_fdiff_1(X5,esk4_2(X6,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos)))|~v1_funct_1(X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~v3_valued_0(X5)|~v1_relat_1(X5))).#trainneg
cnf(i_0_3493, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k5_square_1(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3542, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos),X11)|~v1_xreal_0(esk4_2(X11,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos)))|~r1_fdiff_1(X9,esk4_2(X11,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos))|~v1_funct_1(X9)|~r2_hidden(esk4_2(X11,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos)),X10)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),X11)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10)|~r1_tarski(X11,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,X9,k16_sin_cos),X10),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X11,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X11)|~v3_valued_0(X9)|~v1_relat_1(X9))).#trainneg
cnf(i_0_789, negated_conjecture, (~v1_xboole_0(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_548, plain, (v7_ordinal1(k2_rcomp_1(X1,X2))|~v6_membered(k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1953, plain, (v1_int_1(k4_xcmplx_0(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1999, plain, (v1_xcmplx_0(k7_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2066, plain, (v1_xcmplx_0(k10_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_549, plain, (v1_int_1(k2_rcomp_1(X1,X2))|~v5_membered(k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_551, plain, (v1_membered(k2_rcomp_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_3498, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X5)|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~m1_subset_1(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_3539, plain, (v1_xcmplx_0(k1_seq_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_3557, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xcmplx_0(k1_fdiff_1(k1_sin_cos9,esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~v1_xreal_0(k5_square_1(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~r2_hidden(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))),k1_numbers))).#trainneg
cnf(i_0_3478, negated_conjecture, (m1_subset_1(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1),k1_numbers)|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0)|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,X1)))).#trainneg
cnf(i_0_3093, negated_conjecture, (m1_subset_1(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1),k1_numbers)|~r2_hidden(X1,esk1_0)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)),k1_numbers)|~m1_subset_1(k5_square_1(X1),k1_numbers))).#trainneg
cnf(i_0_3423, negated_conjecture, (m1_subset_1(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1),k1_numbers)|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers)|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,X1)))).#trainneg
cnf(i_0_2181, plain, (m1_subset_1(k4_xcmplx_0(k7_xcmplx_0(np__1,np__2)),k1_numbers))).#trainneg
cnf(i_0_552, plain, (v2_membered(k2_rcomp_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_2194, plain, (v1_xcmplx_0(k7_real_1(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_866, plain, (m1_subset_1(k1_relset_1(k1_numbers,k19_sin_cos),k1_zfmisc_1(k1_numbers))|~v1_relat_1(k19_sin_cos))).#trainneg
cnf(i_0_553, plain, (v3_valued_0(k2_rcomp_1(X1,X2))|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_554, plain, (v1_valued_0(k2_rcomp_1(X1,X2))|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_3636, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k16_sin_cos,k16_sin_cos),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k16_sin_cos,k16_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k16_sin_cos,k16_sin_cos))))).#trainneg
cnf(i_0_3634, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k16_sin_cos),X3)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X3)|~r1_fdiff_1(X2,k1_seq_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k16_sin_cos))))|~r1_fdiff_1(X1,esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k16_sin_cos)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_3651, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k16_sin_cos),X7)|~v3_valued_0(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6))|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7)|~r1_fdiff_1(X6,k1_seq_1(X5,esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k16_sin_cos))))|~r1_fdiff_1(X5,esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k16_sin_cos)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X7,k1_partfun1(X1,X2,X3,X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X5,X6),k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_3763, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_3758, plain, (v3_membered(k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_2576, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)),k16_sin_cos),X5)|~v1_xreal_0(k10_real_1(np__1,k7_real_1(np__1,k5_square_1(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)),k16_sin_cos))))))|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)),esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)),k16_sin_cos)))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))|~r2_hidden(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)),k16_sin_cos)),k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~v3_rcomp_1(X5))).#trainneg
cnf(i_0_3759, plain, (v1_membered(k1_relset_1(k1_numbers,k19_sin_cos))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_3538, negated_conjecture, (v1_xxreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))).#trainneg
cnf(i_0_3606, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v3_membered(k1_zfmisc_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))).#trainneg
cnf(i_0_3760, plain, (v2_membered(k1_relset_1(k1_numbers,k19_sin_cos))|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_3753, plain, (v1_xreal_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_558, plain, (k3_square_1(k5_xcmplx_0(X1))=k7_xcmplx_0(k5_xcmplx_0(X1),X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_559, plain, (k3_xcmplx_0(k5_xcmplx_0(X1),X2)=k7_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3658, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~m1_subset_1(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_3685, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_3652, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,X9,X10),k16_sin_cos),X11)|~v3_valued_0(k1_partfun1(X5,X6,X7,X8,X9,X10))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,X9,X10))|~r1_tarski(X11,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,X9,X10),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X5,X6,X7,X8,X9,X10),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_partfun1(X5,X6,X7,X8,X9,X10),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X11,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X11)|~r1_fdiff_1(X10,k1_seq_1(X9,esk4_2(X11,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,X9,X10),k16_sin_cos))))|~r1_fdiff_1(X9,esk4_2(X11,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,X9,X10),k16_sin_cos)))|~v1_funct_1(X10)|~v1_funct_1(X9)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(esk4_2(X11,k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,X9,X10),k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_3862, plain, (k7_xcmplx_0(X1,np__1)=k3_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3635, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,X6),k16_sin_cos),X7)|~v3_valued_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~r1_tarski(X7,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,X6),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X7,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X7)|~r1_fdiff_1(X6,k1_seq_1(X5,esk4_2(X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,X6),k16_sin_cos))))|~r1_fdiff_1(X5,esk4_2(X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,X6),k16_sin_cos)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(esk4_2(X7,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,X6),k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_3831, plain, (k3_square_1(np__1)=np__1)).#trainneg
cnf(i_0_1581, plain, (v1_xboole_0(k5_numbers)|m1_subset_1(esk7_1(k5_numbers),k1_numbers))).#trainneg
cnf(i_0_2586, plain, (k10_real_1(np__1,k7_real_1(np__1,k5_square_1(X1)))=k1_fdiff_1(k1_sin_cos9,X1)|~r2_hidden(X1,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))).#trainneg
cnf(i_0_2516, plain, (v1_xxreal_0(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3640, negated_conjecture, (v1_int_1(esk1_0)|~v6_membered(k1_zfmisc_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))))).#trainneg
cnf(i_0_3754, plain, (v1_xxreal_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v2_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3756, plain, (v7_ordinal1(k1_relset_1(k1_numbers,k19_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_561, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(k5_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3554, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),X10)|~v1_xreal_0(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),X9)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),X10)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9)|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X10)|~r1_tarski(X10,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X10)|~v1_xcmplx_0(k1_fdiff_1(k1_sin_cos9,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))))|~v1_xreal_0(k5_square_1(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),esk1_0)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))),k1_numbers))).#trainneg
cnf(i_0_3230, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X1)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))))|~v1_xreal_0(k5_square_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3083, negated_conjecture, (v1_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_3721, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~r2_hidden(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))),k1_numbers)|~m1_subset_1(k5_square_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3218, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)=k1_fdiff_1(k1_sin_cos9,X1)|~v1_xcmplx_0(k1_fdiff_1(k1_sin_cos9,X1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)),k1_numbers))).#trainneg
cnf(i_0_2517, plain, (v1_xxreal_0(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_749, plain, (k10_real_1(np__1,k7_real_1(np__1,k5_square_1(X1)))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,k2_rcomp_1(k4_xcmplx_0(np__1),np__1)),X1)|~r2_hidden(X1,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))).#trainneg
cnf(i_0_563, plain, (v1_int_1(k7_xcmplx_0(X1,X2))|~v1_int_1(k5_xcmplx_0(X2))|~v1_int_1(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3936, plain, (k5_square_1(np__1)=np__1)).#trainneg
cnf(i_0_1429, plain, (v1_xreal_0(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3757, plain, (v1_int_1(k1_relset_1(k1_numbers,k19_sin_cos))|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3808, plain, (v1_xcmplx_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_564, plain, (k7_xcmplx_0(X1,k5_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_3932, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k16_sin_cos),X6)|~v3_valued_0(k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos))|~r1_fdiff_1(X5,esk4_2(X6,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k16_sin_cos)))|~v1_funct_1(X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k16_sin_cos)))|~m1_subset_1(esk4_2(X6,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k16_sin_cos)),k1_numbers)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~v1_xreal_0(k1_seq_1(X5,esk4_2(X6,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k16_sin_cos)))))).#trainneg
cnf(i_0_4001, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),X10)|~v1_xreal_0(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),X9)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),X10)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9)|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X10)|~r1_tarski(X10,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k1_sin_cos9,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X10)|~v1_xcmplx_0(k1_fdiff_1(k1_sin_cos9,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),esk1_0)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))),k1_numbers)|~m1_subset_1(k5_square_1(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_3970, plain, (v1_xreal_0(k7_xcmplx_0(X1,np__1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3943, plain, (v1_xreal_0(esk7_1(k5_numbers))|v1_xboole_0(k5_numbers))).#trainneg
cnf(i_0_2735, plain, (k10_real_1(np__1,k2_xcmplx_0(np__1,k5_square_1(X1)))=k1_fdiff_1(k1_sin_cos9,X1)|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))).#trainneg
cnf(i_0_566, plain, (k6_xcmplx_0(k4_xcmplx_0(np__1),np__1)=k4_xcmplx_0(np__2)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_1436, plain, (v1_xreal_0(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1437, plain, (v1_xreal_0(k1_sin_cos9)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3964, plain, (v1_xxreal_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3966, plain, (v1_int_1(k1_relset_1(k1_numbers,k19_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_570, plain, (k2_xcmplx_0(k4_xcmplx_0(X1),X2)=k6_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4145, negated_conjecture, (k3_xcmplx_0(X1,k7_real_1(np__1,k5_square_1(X2)))=k7_xcmplx_0(X1,k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X2))|~v1_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X2))|~v1_xcmplx_0(X1)|~r2_hidden(X2,esk1_0)|~m1_subset_1(k5_square_1(X2),k1_numbers))).#trainneg
cnf(i_0_4144, negated_conjecture, (k3_xcmplx_0(X1,k7_real_1(np__1,k5_square_1(X2)))=k7_xcmplx_0(X1,k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X2))|~v1_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k5_square_1(X2))|~r2_hidden(X2,esk1_0))).#trainneg
cnf(i_0_4160, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k16_sin_cos),X6)|~v3_valued_0(k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos))|~r1_fdiff_1(X5,esk4_2(X6,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k16_sin_cos)))|~v1_funct_1(X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k16_sin_cos)))|~m1_subset_1(esk4_2(X6,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k16_sin_cos)),k1_numbers)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~m1_subset_1(k1_seq_1(X5,esk4_2(X6,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k16_sin_cos))),k1_numbers))).#trainneg
cnf(i_0_4176, plain, (m1_subset_1(k7_xcmplx_0(X1,np__1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4183, plain, (v1_xcmplx_0(esk7_1(k5_numbers))|v1_xboole_0(k5_numbers))).#trainneg
cnf(i_0_2724, plain, (k7_xcmplx_0(np__1,k7_real_1(np__1,k5_square_1(X1)))=k1_fdiff_1(k1_sin_cos9,X1)|~v1_xreal_0(k7_real_1(np__1,k5_square_1(X1)))|~r2_hidden(X1,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))).#trainneg
cnf(i_0_572, plain, (v1_xreal_0(k6_xcmplx_0(X1,X2))|~v1_xreal_0(k4_xcmplx_0(X2))|~v1_xreal_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2550, plain, (v1_xxreal_0(k1_sin_cos9)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2551, plain, (v1_xcmplx_0(k19_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2604, plain, (v1_xcmplx_0(k16_sin_cos)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3761, plain, (v3_valued_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v3_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_3762, plain, (v1_valued_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v1_valued_0(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_574, plain, (v1_int_1(k6_xcmplx_0(X1,X2))|~v1_int_1(k4_xcmplx_0(X2))|~v1_int_1(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4153, plain, (r2_fdiff_1(k1_partfun1(X7,X8,X9,X10,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k16_sin_cos),X6)|~v3_valued_0(k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos))|~v1_xreal_0(k1_seq_1(X5,esk4_2(X6,k1_partfun1(X7,X8,X9,X10,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k16_sin_cos))))|~r1_fdiff_1(X5,esk4_2(X6,k1_partfun1(X7,X8,X9,X10,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k16_sin_cos)))|~v1_funct_1(X5)|~r1_tarski(X6,k1_relset_1(k1_numbers,k1_partfun1(X7,X8,X9,X10,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k16_sin_cos)))|~m1_subset_1(esk4_2(X6,k1_partfun1(X7,X8,X9,X10,k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k16_sin_cos)),k1_numbers)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X6)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_4148, negated_conjecture, (k7_xcmplx_0(X1,k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X2))=k3_xcmplx_0(X1,k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X2)))|~v1_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X2)))|~v1_xreal_0(k5_square_1(X2))|~r2_hidden(X2,esk1_0))).#trainneg
cnf(i_0_3982, negated_conjecture, (r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X1)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X1,k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))))|~v1_xreal_0(k5_square_1(esk4_2(X1,k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X1,k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_4117, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(k7_xcmplx_0(np__1,X2))|~v1_xreal_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4137, negated_conjecture, (k3_xcmplx_0(X1,k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X2))=k7_xcmplx_0(X1,k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X2)))|~v1_xcmplx_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X2)))|~v1_xcmplx_0(X1)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X2)))|~v1_xreal_0(k5_square_1(X2))|~r2_hidden(X2,esk1_0))).#trainneg
cnf(i_0_1404, negated_conjecture, (m1_subset_1(esk1_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos)))))).#trainneg
cnf(i_0_3944, plain, (v1_xxreal_0(esk7_1(k5_numbers))|v1_xboole_0(k5_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_2734, plain, (k10_real_1(np__1,k7_real_1(np__1,k3_xcmplx_0(X1,X1)))=k1_fdiff_1(k1_sin_cos9,X1)|~r2_hidden(X1,k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~m1_subset_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))).#trainneg
cnf(i_0_575, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(np__2))=k2_xcmplx_0(X1,np__2)|~v1_xcmplx_0(k4_xcmplx_0(np__2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4298, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,np__1))|~v1_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2605, plain, (v1_xcmplx_0(k1_sin_cos9)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3537, negated_conjecture, (v1_xreal_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_576, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(np__1))=k2_xcmplx_0(X1,np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_577, plain, (k6_xcmplx_0(X1,k4_xcmplx_0(X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(k4_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4448, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X5)|~v1_xreal_0(k5_square_1(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))),k1_numbers))).#trainneg
cnf(i_0_4412, negated_conjecture, (k3_xcmplx_0(k7_real_1(np__1,k5_square_1(X1)),X2)=k7_xcmplx_0(X2,k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_4405, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1),X2)=k7_xcmplx_0(X2,k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))|~v1_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_1673, plain, (v1_xreal_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4520, negated_conjecture, (m1_subset_1(X1,k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos)))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_578, plain, (k6_xcmplx_0(X1,np__0)=k2_xcmplx_0(X1,np__0)|~v1_xcmplx_0(np__0)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4300, plain, (v1_int_1(k7_xcmplx_0(X1,np__1))|~v5_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4510, negated_conjecture, (v1_xreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos)))))).#trainneg
cnf(i_0_4511, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos)))))).#trainneg
cnf(i_0_988, plain, (v1_xcmplx_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2254, plain, (v1_xcmplx_0(k1_relset_1(k1_numbers,k1_sin_cos9))|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_579, plain, (k8_real_1(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4474, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k7_xcmplx_0(np__1,X2),k1_numbers))).#trainneg
cnf(i_0_4413, negated_conjecture, (k3_xcmplx_0(k7_real_1(np__1,k5_square_1(X1)),X2)=k7_xcmplx_0(X2,k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers))).#trainneg
cnf(i_0_4682, plain, (k8_real_1(X2,X1)=k3_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4696, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X1)|~v1_xreal_0(k5_square_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xcmplx_0(k1_fdiff_1(k1_sin_cos9,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))),k1_numbers))).#trainneg
cnf(i_0_4483, negated_conjecture, (k7_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1),k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))=k5_square_1(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~m1_subset_1(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1),k1_numbers)|~v1_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~v1_xreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1)))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_1372, plain, (v1_xcmplx_0(esk7_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_3946, plain, (v7_ordinal1(esk7_1(k5_numbers))|v1_xboole_0(k5_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_581, plain, (k8_real_1(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4669, plain, (v1_int_1(k7_xcmplx_0(X1,np__1))|~v1_xreal_0(X1)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_4512, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos)))))).#trainneg
cnf(i_0_4513, negated_conjecture, (v7_ordinal1(esk1_0)|~v6_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos)))))).#trainneg
cnf(i_0_582, plain, (m1_subset_1(k8_real_1(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_584, plain, (k7_real_1(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_4469, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(X1,k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X2)))|~v1_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X2))|~v1_xreal_0(k7_real_1(np__1,k5_square_1(X2)))|~v1_xreal_0(X1)|~r2_hidden(X2,esk1_0)|~m1_subset_1(k5_square_1(X2),k1_numbers))).#trainneg
cnf(i_0_4877, negated_conjecture, (k10_real_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k2_xcmplx_0(k5_square_1(esk2_0),np__1))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~v1_xreal_0(k5_square_1(esk2_0)))).#trainneg
cnf(i_0_4891, negated_conjecture, (k10_real_1(np__1,k2_xcmplx_0(k5_square_1(X1),np__1))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r2_hidden(X1,esk1_0)|~v1_xreal_0(k5_square_1(X1)))).#trainneg
cnf(i_0_4887, negated_conjecture, (v1_xreal_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~v1_xreal_0(k2_xcmplx_0(k5_square_1(X1),np__1))|~r2_hidden(X1,esk1_0)|~v1_xreal_0(k5_square_1(X1)))).#trainneg
cnf(i_0_4868, negated_conjecture, (k10_real_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k2_xcmplx_0(k3_xcmplx_0(esk2_0,esk2_0),np__1))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~v1_xreal_0(k3_xcmplx_0(esk2_0,esk2_0)))).#trainneg
cnf(i_0_4885, negated_conjecture, (k2_xcmplx_0(k5_square_1(X1),np__1)=k7_xcmplx_0(np__1,k1_fdiff_1(k1_sin_cos9,X1))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_2018, plain, (v1_xcmplx_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1693, plain, (v1_xreal_0(esk7_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_3947, plain, (v1_int_1(esk7_1(k5_numbers))|v1_xboole_0(k5_numbers)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_4830, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_numbers)|~v1_xcmplx_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk2_0))|~v1_xcmplx_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0))))).#trainneg
cnf(i_0_585, plain, (k7_real_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4514, negated_conjecture, (v1_int_1(esk1_0)|~v5_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos)))))).#trainneg
cnf(i_0_4673, negated_conjecture, (v1_xxreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos)))))).#trainneg
cnf(i_0_2552, plain, (v1_xreal_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2554, plain, (v1_xcmplx_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_586, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_587, plain, (v1_valued_0(esk7_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_membered(X2))).#trainneg
cnf(i_0_4797, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_5005, negated_conjecture, (k10_real_1(np__1,k2_xcmplx_0(k5_square_1(X1),np__1))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers))).#trainneg
cnf(i_0_4881, negated_conjecture, (v1_xreal_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k2_xcmplx_0(k5_square_1(X1),np__1),k1_numbers)|~v1_xreal_0(k5_square_1(X1)))).#trainneg
cnf(i_0_4791, plain, (k3_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_fdiff_1(X2,X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3)|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_xreal_0(k1_fdiff_1(X1,k1_seq_1(X2,X3)))|~v1_xreal_0(k1_fdiff_1(X2,X3)))).#trainneg
cnf(i_0_643, plain, (v1_xboole_0(X1)|~v1_xboole_0(k3_square_1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1694, plain, (v1_xxreal_0(esk7_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_4991, negated_conjecture, (k10_real_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k2_xcmplx_0(k5_square_1(esk2_0),np__1))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k5_square_1(esk2_0),k1_numbers))).#trainneg
cnf(i_0_591, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v3_valued_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_4849, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos)))))).#trainneg
cnf(i_0_5023, negated_conjecture, (v1_int_1(esk1_0)|~v6_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos)))))).#trainneg
cnf(i_0_2648, plain, (v1_xreal_0(k1_relset_1(k1_numbers,k16_sin_cos))|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2777, plain, (v1_xreal_0(k1_relset_1(k1_numbers,k1_sin_cos9))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_592, plain, (v1_xreal_0(k7_xcmplx_0(X2,X1))|~v1_xreal_0(k5_xcmplx_0(X2))|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4898, plain, (k7_real_1(X2,X1)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4833, plain, (k3_xcmplx_0(k1_fdiff_1(X2,X3),k1_fdiff_1(X1,k1_seq_1(X2,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3)|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~v1_xreal_0(k1_fdiff_1(X1,k1_seq_1(X2,X3)))|~m1_subset_1(k1_fdiff_1(X2,X3),k1_numbers))).#trainneg
cnf(i_0_5134, negated_conjecture, (k7_xcmplx_0(np__1,k2_xcmplx_0(k5_square_1(X1),np__1))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~v1_xreal_0(k2_xcmplx_0(k5_square_1(X1),np__1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers))).#trainneg
cnf(i_0_5243, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5143, negated_conjecture, (k10_real_1(np__1,k2_xcmplx_0(np__1,k5_square_1(X1)))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r2_hidden(X1,esk1_0)|~m1_subset_1(k5_square_1(X1),k1_numbers)|~v1_xcmplx_0(k5_square_1(X1)))).#trainneg
cnf(i_0_1075, plain, (v1_xreal_0(X1)|~v1_int_1(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1695, plain, (v1_xcmplx_0(esk7_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_4822, plain, (v1_xxreal_0(esk7_1(k5_numbers))|v1_xboole_0(k5_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_5187, negated_conjecture, (k10_real_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k2_xcmplx_0(np__1,k5_square_1(esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k5_square_1(esk2_0),k1_numbers)|~v1_xcmplx_0(k5_square_1(esk2_0)))).#trainneg
cnf(i_0_594, plain, (k7_xcmplx_0(k5_xcmplx_0(X1),X2)=k7_xcmplx_0(k5_xcmplx_0(X2),X1)|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_3806, negated_conjecture, (v1_xxreal_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3807, negated_conjecture, (v1_xcmplx_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2779, plain, (v1_xcmplx_0(k1_relset_1(k1_numbers,k1_sin_cos9))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2836, plain, (v1_xreal_0(k1_relset_1(k1_numbers,k1_sin_cos9))|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_595, plain, (k7_xcmplx_0(k5_xcmplx_0(X1),k7_xcmplx_0(np__1,X2))=k7_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_596, plain, (k7_xcmplx_0(X1,k5_xcmplx_0(X2))=k7_xcmplx_0(X2,k5_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5153, negated_conjecture, (v1_xreal_0(k1_fdiff_1(k1_sin_cos9,X1))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0)|~m1_subset_1(k2_xcmplx_0(k5_square_1(X1),np__1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5190, negated_conjecture, (k7_xcmplx_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k2_xcmplx_0(k5_square_1(esk2_0),np__1))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k5_square_1(esk2_0),k1_numbers)|~v1_xreal_0(k2_xcmplx_0(k5_square_1(esk2_0),np__1))|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_1696, plain, (v1_int_1(esk7_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_597, plain, (k7_xcmplx_0(k7_xcmplx_0(np__1,X1),k5_xcmplx_0(X2))=k7_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5508, negated_conjecture, (v1_xreal_0(k1_fdiff_1(k1_sin_cos9,np__1))|~r2_hidden(np__1,esk1_0))).#trainneg
cnf(i_0_4671, negated_conjecture, (v1_xreal_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5025, negated_conjecture, (v1_xxreal_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4929, negated_conjecture, (k7_xcmplx_0(np__1,k2_xcmplx_0(k5_square_1(X1),np__1))=k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~v1_xreal_0(k2_xcmplx_0(k5_square_1(X1),np__1))|~v1_xreal_0(k5_square_1(X1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_5360, negated_conjecture, (k7_xcmplx_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k2_xcmplx_0(np__1,k5_square_1(esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k5_square_1(esk2_0),k1_numbers)|~v1_xreal_0(k2_xcmplx_0(np__1,k5_square_1(esk2_0)))|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_4945, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),X10)|~v1_xreal_0(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~r1_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),X9)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),X10)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9)|~r1_tarski(X10,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X10)|~v1_xreal_0(k2_xcmplx_0(k5_square_1(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))),np__1))|~v1_xreal_0(k5_square_1(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),esk1_0))).#trainneg
cnf(i_0_1729, plain, (v1_xreal_0(esk7_1(X1))|~v5_membered(X1))).#trainneg
cnf(i_0_4824, plain, (v1_int_1(esk7_1(k5_numbers))|v1_xboole_0(k5_numbers)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_600, plain, (k3_square_1(k5_xcmplx_0(X1))=k5_xcmplx_0(k3_xcmplx_0(X1,X1))|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5193, negated_conjecture, (v1_xcmplx_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_965, negated_conjecture, (v3_valued_0(esk1_0)|~v3_valued_0(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~v1_relat_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))).#trainneg
cnf(i_0_3755, plain, (v1_xcmplx_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3963, plain, (v1_xreal_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_601, plain, (k3_xcmplx_0(k5_xcmplx_0(X1),k5_xcmplx_0(X2))=k5_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_5661, plain, (k3_square_1(k5_xcmplx_0(X1))=k5_xcmplx_0(k3_square_1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5648, plain, (k7_xcmplx_0(k5_xcmplx_0(X1),X1)=k5_xcmplx_0(k3_xcmplx_0(X1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5704, plain, (k5_xcmplx_0(k3_xcmplx_0(X2,X1))=k5_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5728, plain, (k5_xcmplx_0(k3_square_1(X1))=k7_xcmplx_0(k5_xcmplx_0(X1),X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5839, plain, (k5_xcmplx_0(k3_xcmplx_0(X1,np__1))=k5_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5729, plain, (k5_xcmplx_0(k3_square_1(X1))=k5_xcmplx_0(k3_xcmplx_0(X1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2304, plain, (v1_xcmplx_0(k3_square_1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_603, plain, (v1_xreal_0(k5_xcmplx_0(k3_xcmplx_0(X1,X2)))|~v1_xreal_0(k5_xcmplx_0(X2))|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5568, negated_conjecture, (m1_subset_1(k1_fdiff_1(k1_sin_cos9,np__1),k1_numbers)|~r2_hidden(np__1,esk1_0))).#trainneg
cnf(i_0_966, negated_conjecture, (v1_valued_0(esk1_0)|~v1_valued_0(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~v1_relat_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))).#trainneg
cnf(i_0_2951, negated_conjecture, (v1_xcmplx_0(k4_xcmplx_0(esk7_1(esk1_0)))|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_3965, plain, (v1_xcmplx_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_604, plain, (v1_xcmplx_0(k5_xcmplx_0(k3_xcmplx_0(X1,X2)))|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_605, plain, (v1_int_1(k5_xcmplx_0(k3_xcmplx_0(X1,X2)))|~v1_int_1(k5_xcmplx_0(X2))|~v1_int_1(k5_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6049, plain, (k5_xcmplx_0(k3_square_1(X1))=k5_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6110, plain, (v1_xcmplx_0(k5_xcmplx_0(np__0)))).#trainneg
cnf(i_0_606, plain, (k7_xcmplx_0(k5_xcmplx_0(X1),X2)=k5_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5567, negated_conjecture, (v1_xcmplx_0(k1_fdiff_1(k1_sin_cos9,np__1))|~r2_hidden(np__1,esk1_0))).#trainneg
cnf(i_0_3111, negated_conjecture, (v1_xcmplx_0(k4_xcmplx_0(esk7_1(esk1_0)))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_6108, negated_conjecture, (v1_valued_0(esk1_0)|~v1_relat_1(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~v3_valued_0(k2_rcomp_1(k4_xcmplx_0(np__1),np__1)))).#trainneg
cnf(i_0_4080, plain, (v1_xreal_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v5_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_607, plain, (k5_xcmplx_0(k3_xcmplx_0(X1,k5_xcmplx_0(X2)))=k3_xcmplx_0(k5_xcmplx_0(X1),X2)|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_608, plain, (k3_xcmplx_0(k5_xcmplx_0(X1),k7_xcmplx_0(np__1,X2))=k5_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6323, plain, (k5_xcmplx_0(k7_xcmplx_0(X1,X2))=k3_xcmplx_0(k5_xcmplx_0(X1),X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_6238, plain, (k7_xcmplx_0(k5_xcmplx_0(X2),X1)=k5_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_5943, plain, (k5_xcmplx_0(k5_xcmplx_0(X1))=k3_xcmplx_0(X1,np__1)|~v1_xcmplx_0(k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6058, plain, (k5_xcmplx_0(k5_square_1(X1))=k5_xcmplx_0(k3_xcmplx_0(X1,X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5922, plain, (k5_xcmplx_0(k3_square_1(X1))=k7_xcmplx_0(k7_xcmplx_0(np__1,X1),X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5662, plain, (k3_square_1(k5_xcmplx_0(X1))=k5_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6496, plain, (k5_xcmplx_0(k3_xcmplx_0(X2,k5_xcmplx_0(X1)))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1122, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_6215, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,np__0)))).#trainneg
cnf(i_0_3829, plain, (k3_square_1(k4_xcmplx_0(np__1))=k7_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_609, plain, (k5_xcmplx_0(k3_xcmplx_0(k5_xcmplx_0(X1),X2))=k3_xcmplx_0(X1,k5_xcmplx_0(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_832, negated_conjecture, (v1_xreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))))).#trainneg
cnf(i_0_833, negated_conjecture, (v1_xxreal_0(esk1_0)|~v2_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))))).#trainneg
cnf(i_0_610, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,X1),k5_xcmplx_0(X2))=k5_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_611, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~v1_funct_1(X2)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1))).#trainneg
cnf(i_0_6860, plain, (k5_xcmplx_0(k7_xcmplx_0(X1,X2))=k7_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_6853, plain, (k5_xcmplx_0(k5_xcmplx_0(X1))=k7_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1)))).#trainneg
cnf(i_0_7204, plain, (k7_xcmplx_0(np__1,k5_xcmplx_0(X1))=k7_xcmplx_0(X1,np__1)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5845, plain, (k5_xcmplx_0(k3_xcmplx_0(np__1,X1))=k5_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5944, plain, (k7_xcmplx_0(np__1,k3_xcmplx_0(X1,np__1))=k5_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6889, plain, (k5_square_1(k4_xcmplx_0(np__1))=k7_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_5059, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_numbers)|~v1_xreal_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)))|~v1_xreal_0(k1_seq_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk2_0)))).#trainneg
cnf(i_0_613, plain, (v1_xreal_0(k1_fdiff_1(X1,X2))|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_6104, negated_conjecture, (v1_xxreal_0(k1_fdiff_1(k1_sin_cos9,np__1))|~v2_membered(k1_numbers)|~r2_hidden(np__1,esk1_0))).#trainneg
cnf(i_0_6896, plain, (v1_xboole_0(k4_xcmplx_0(np__1))|~v1_xboole_0(k7_xcmplx_0(k4_xcmplx_0(np__1),k4_xcmplx_0(np__1))))).#trainneg
cnf(i_0_834, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v1_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))))).#trainneg
cnf(i_0_614, plain, (v1_xxreal_0(k1_fdiff_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_615, plain, (v1_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_5925, plain, (k5_xcmplx_0(k5_square_1(X1))=k7_xcmplx_0(k5_xcmplx_0(X1),X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7518, plain, (k7_xcmplx_0(k5_xcmplx_0(X1),X1)=k5_xcmplx_0(k3_xcmplx_0(X1,X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5606, plain, (v1_xcmplx_0(esk7_1(X1))|~v5_membered(X1))).#trainneg
cnf(i_0_6439, plain, (k3_xcmplx_0(np__1,k4_xcmplx_0(np__1))=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_616, plain, (v7_ordinal1(k1_fdiff_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_835, negated_conjecture, (v7_ordinal1(esk1_0)|~v6_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))))).#trainneg
cnf(i_0_836, negated_conjecture, (v1_int_1(esk1_0)|~v5_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))))).#trainneg
cnf(i_0_617, plain, (v1_int_1(k1_fdiff_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_619, plain, (v1_xreal_0(k2_fdiff_1(X1,X2))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_7439, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X5)|~r1_tarski(X5,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X5)|~v1_xreal_0(esk4_2(X5,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))).#trainneg
cnf(i_0_7441, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k5_square_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~v1_xreal_0(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))))).#trainneg
cnf(i_0_7440, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),X10)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),esk1_0)|~r2_hidden(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),X9)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),X10)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9)|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X10)|~r1_tarski(X10,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(X10,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X10)|~v1_xreal_0(esk4_2(X10,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))))).#trainneg
cnf(i_0_7500, plain, (k5_xcmplx_0(k7_xcmplx_0(k5_xcmplx_0(X1),X1))=k5_square_1(X1)|~v1_xcmplx_0(k5_square_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7689, plain, (k7_xcmplx_0(X1,k5_xcmplx_0(X1))=k5_square_1(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7633, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_7501, plain, (k7_xcmplx_0(k5_xcmplx_0(X1),X1)=k7_xcmplx_0(np__1,k5_square_1(X1))|~v1_xcmplx_0(k5_square_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1020, negated_conjecture, (v1_xxreal_0(X1)|~v2_membered(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_7727, plain, (k5_xcmplx_0(k5_xcmplx_0(k5_square_1(X1)))=k5_square_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7152, plain, (k7_xcmplx_0(k4_xcmplx_0(np__1),np__1)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_620, plain, (v1_xxreal_0(k2_fdiff_1(X1,X2))|~v2_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_6989, negated_conjecture, (v1_xxreal_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))))).#trainneg
cnf(i_0_7447, negated_conjecture, (v1_xcmplx_0(esk1_0)|~v3_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))))).#trainneg
cnf(i_0_621, plain, (v1_xcmplx_0(k2_fdiff_1(X1,X2))|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_622, plain, (v7_ordinal1(k2_fdiff_1(X1,X2))|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_7796, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk1_0)|~r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_7795, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(esk4_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0))).#trainneg
cnf(i_0_7208, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(X1,np__1))=k5_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1021, negated_conjecture, (v1_xcmplx_0(X1)|~v1_membered(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_7728, plain, (k5_xcmplx_0(k5_xcmplx_0(k3_xcmplx_0(X1,X1)))=k5_square_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_623, plain, (v1_int_1(k2_fdiff_1(X1,X2))|~v5_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_6106, negated_conjecture, (v7_ordinal1(k1_fdiff_1(k1_sin_cos9,np__1))|~v6_membered(k1_numbers)|~r2_hidden(np__1,esk1_0))).#trainneg
cnf(i_0_7621, negated_conjecture, (v1_int_1(esk1_0)|~v6_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos)))))).#trainneg
cnf(i_0_2265, plain, (v1_xreal_0(k7_xcmplx_0(np__1,np__2))|~v1_xcmplx_0(k7_xcmplx_0(np__1,np__2)))).#trainneg
cnf(i_0_2506, plain, (m1_subset_1(k1_relset_1(k1_numbers,k16_sin_cos),k1_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2710, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_sin_cos9),k1_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_624, plain, (v3_membered(k2_fdiff_1(X1,X2))|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_625, plain, (v1_membered(k2_fdiff_1(X1,X2))|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_8002, plain, (k7_xcmplx_0(np__1,k3_xcmplx_0(X1,np__1))=k5_xcmplx_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7957, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk1_0)|~r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5720, plain, (k5_xcmplx_0(k3_square_1(X1))=k5_square_1(k5_xcmplx_0(X1))|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7961, negated_conjecture, (r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk4_2(X1,k1_partfun1(X2,X3,X4,X5,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos)),esk1_0)|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(X1)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_6711, plain, (k7_xcmplx_0(k5_xcmplx_0(X1),X1)=k7_xcmplx_0(k7_xcmplx_0(np__1,X1),X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1022, negated_conjecture, (v7_ordinal1(X1)|~v6_membered(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_4975, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_fdiff_1(k1_sin_cos9,esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~v1_xcmplx_0(k1_fdiff_1(k1_sin_cos9,esk2_0))|~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0)|~m1_subset_1(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_numbers))).#trainneg
cnf(i_0_626, plain, (v2_membered(k2_fdiff_1(X1,X2))|~v2_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_6107, negated_conjecture, (v1_int_1(k1_fdiff_1(k1_sin_cos9,np__1))|~v5_membered(k1_numbers)|~r2_hidden(np__1,esk1_0))).#trainneg
cnf(i_0_3809, plain, (m1_subset_1(k1_relset_1(k1_numbers,k19_sin_cos),k1_numbers)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_627, plain, (v3_valued_0(k2_fdiff_1(X1,X2))|~v3_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_628, plain, (v1_valued_0(k2_fdiff_1(X1,X2))|~v1_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_relat_1(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_8126, plain, (k3_xcmplx_0(X1,k5_xcmplx_0(X2))=k7_xcmplx_0(X1,k3_xcmplx_0(X2,np__1))|~v1_xcmplx_0(k3_xcmplx_0(X2,np__1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_7203, plain, (k5_xcmplx_0(k7_xcmplx_0(X1,np__1))=k5_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1023, negated_conjecture, (v1_int_1(X1)|~v5_membered(k2_rcomp_1(k4_xcmplx_0(np__1),np__1))|~r2_hidden(X1,esk1_0))).#trainneg
cnf(i_0_7933, plain, (k3_xcmplx_0(k4_xcmplx_0(np__1),np__1)=k4_xcmplx_0(np__1))).#trainneg
cnf(i_0_629, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k2_fdiff_1(X2,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_931, plain, (v3_xxreal_0(k1_real_1(X1))|~v2_xxreal_0(X1)|~v1_xreal_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7688, plain, (k3_xcmplx_0(k5_xcmplx_0(k5_xcmplx_0(X1)),X1)=k5_square_1(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6447, plain, (k3_xcmplx_0(np__1,k7_xcmplx_0(np__1,X1))=k5_xcmplx_0(X1)|~v1_xcmplx_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8072, plain, (k5_xcmplx_0(k5_xcmplx_0(k3_square_1(X1)))=k5_square_1(X1)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7987, plain, (k3_xcmplx_0(X1,k5_xcmplx_0(X2))=k7_xcmplx_0(X1,k7_xcmplx_0(X2,np__1))|~v1_xcmplx_0(k7_xcmplx_0(X2,np__1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_8658, plain, (k7_xcmplx_0(k3_square_1(X1),np__1)=k5_square_1(X1)|~v1_xcmplx_0(k3_square_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8723, plain, (k5_xcmplx_0(k7_xcmplx_0(X1,k7_xcmplx_0(X2,np__1)))=k7_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k7_xcmplx_0(X2,np__1)))).#trainneg
cnf(i_0_7431, negated_conjecture, (k3_xcmplx_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)),k1_fdiff_1(k1_sin_cos9,esk2_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)))|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,esk2_0))|~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0))).#trainneg
cnf(i_0_633, plain, (v1_valued_0(k2_fdiff_1(X1,X2))|~v1_membered(k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4518, negated_conjecture, (v3_valued_0(esk1_0)|~v3_valued_0(k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos)))|~v1_relat_1(k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos))))).#trainneg
cnf(i_0_1155, plain, (v1_int_1(k1_real_1(X1))|~v1_int_1(X1)|~v1_xcmplx_0(k1_real_1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_634, plain, (v5_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_8724, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(X2,np__1))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k7_xcmplx_0(X2,np__1)))).#trainneg
cnf(i_0_8870, negated_conjecture, (k3_xcmplx_0(k1_fdiff_1(k1_sin_cos9,esk2_0),k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0),esk2_0)|~v1_xreal_0(k1_seq_1(k19_sin_cos,k1_seq_1(k1_sin_cos9,esk2_0)))|~v1_xreal_0(k7_real_1(np__1,k5_square_1(esk2_0)))|~v1_xreal_0(k1_fdiff_1(k1_sin_cos9,esk2_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k1_sin_cos9,k16_sin_cos),esk1_0))).#trainneg
cnf(i_0_8635, plain, (k7_xcmplx_0(np__1,k5_xcmplx_0(k3_square_1(X1)))=k5_square_1(X1)|~v1_xcmplx_0(k5_xcmplx_0(k3_square_1(X1)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8900, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X2,k7_xcmplx_0(X3,np__1)))=k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k7_xcmplx_0(X3,np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_7641, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),esk1_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos))|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos))|~v1_funct_1(k2_fdiff_1(k1_sin_cos9,esk1_0))|~r2_hidden(esk4_2(esk1_0,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)),X9)|~r2_fdiff_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),esk1_0)|~r2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9)|~r2_fdiff_1(k2_fdiff_1(k1_sin_cos9,esk1_0),esk1_0)|~r1_tarski(esk1_0,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k16_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_partfun1(X5,X6,X7,X8,k2_fdiff_1(k1_sin_cos9,esk1_0),k16_sin_cos),X9),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(k1_sin_cos9,esk1_0),k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_5730, plain, (k5_xcmplx_0(k3_square_1(k4_xcmplx_0(np__1)))=k3_square_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_637, plain, (v4_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_8094, negated_conjecture, (v1_xxreal_0(k1_fdiff_1(k1_sin_cos9,np__1))|~v6_membered(k1_numbers)|~r2_hidden(np__1,esk1_0))).#trainneg
cnf(i_0_4519, negated_conjecture, (v1_valued_0(esk1_0)|~v1_valued_0(k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos)))|~v1_relat_1(k1_relset_1(k1_numbers,k3_relat_1(k1_sin_cos9,k16_sin_cos))))).#trainneg
cnf(i_0_640, plain, (v1_xboole_0(k5_xcmplx_0(X1))|v1_xboole_0(k5_xcmplx_0(X2))|~v1_xboole_0(k5_xcmplx_0(k3_xcmplx_0(X1,X2)))|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_642, plain, (v1_xboole_0(X1)|v1_xboole_0(k5_xcmplx_0(X2))|~v1_xboole_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_9072, negated_conjecture, (~r1_fdiff_1(k1_sin_cos9,esk2_0)|~m1_subset_1(k1_seq_1(k1_sin_cos9,esk2_0),k1_numbers))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 157
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 206
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 204
# Processed clauses                    : 2042
# ...of these trivial                  : 102
# ...subsumed                          : 715
# ...remaining for further processing  : 1225
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 71
# Backward-rewritten                   : 74
# Generated clauses                    : 9027
# ...of the previous two non-trivial   : 8328
# Contextual simplify-reflections      : 228
# Paramodulations                      : 9022
# Factorizations                       : 0
# Equation resolutions                 : 5
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1080
#    Positive orientable unit clauses  : 162
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 913
# Current number of unprocessed clauses: 6379
# ...number of literals in the above   : 37717
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1370
# Proof object given clauses           : 134
# Proof search given clauses           : 1225
# Clause-clause subsumption calls (NU) : 525290
# Rec. Clause-clause subsumption calls : 52264
# Non-unit clause-clause subsumptions  : 996
# Unit Clause-clause subsumption calls : 6611
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 49
# BW rewrite match successes           : 18
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 276527

# -------------------------------------------------
# User time                : 27.829 s
# System time              : 0.689 s
# Total time               : 28.518 s
# Maximum resident set size: 366468 pages
27.83user 0.78system 0:29.50elapsed 96%CPU (0avgtext+0avgdata 366468maxresident)k
0inputs+856outputs (0major+192338minor)pagefaults 0swaps
