# ENIGMATIC: Connected to the TF server '127.0.0.1:8888'.
# ENIGMATIC: LightGBM model './big03/model.lgb' loaded with 17544 features 'C(l,x[c=3],r,h[b=4096],v,c[b=256],d,a):G:P'
# Preprocessing time       : 0.736 s
# ENIGMATIC: LightGBM parental guidance model #1 './froc_gnn_model2/model.lgb' loaded with 13963 features 'C(l,x[c=3],r,h[b=2048],v,c[b=256],d,a):G:F'
#TF#SERVER: Lgb model filter skipped 31 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=133, query=168, total=307).
#TF#SERVER: Sent 112674 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1208 bytes expected).
#TF#SERVER: Received 1208 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 88 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=303, query=311, total=620).
#TF#SERVER: Sent 273627 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2460 bytes expected).
#TF#SERVER: Received 2460 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 251 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=545, query=786, total=1337).
#TF#SERVER: Sent 636982 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5328 bytes expected).
#TF#SERVER: Received 5328 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 395 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=712, query=629, total=1347).
#TF#SERVER: Sent 588486 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5368 bytes expected).
#TF#SERVER: Received 5368 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 418 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=945, query=622, total=1573).
#TF#SERVER: Sent 795987 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6272 bytes expected).
#TF#SERVER: Received 6272 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 378 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1079, query=649, total=1734).
#TF#SERVER: Sent 859069 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6916 bytes expected).
#TF#SERVER: Received 6916 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 528 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1201, query=510, total=1717).
#TF#SERVER: Sent 825898 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6848 bytes expected).
#TF#SERVER: Received 6848 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 599 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1351, query=448, total=1805).
#TF#SERVER: Sent 817788 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7200 bytes expected).
#TF#SERVER: Received 7200 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 545 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1496, query=479, total=1981).
#TF#SERVER: Sent 906357 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7904 bytes expected).
#TF#SERVER: Received 7904 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 427 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1536, query=606, total=2148).
#TF#SERVER: Sent 919083 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8572 bytes expected).
#TF#SERVER: Received 8572 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 505 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1536, query=524, total=2066).
#TF#SERVER: Sent 897854 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8244 bytes expected).
#TF#SERVER: Received 8244 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 478 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1536, query=551, total=2093).
#TF#SERVER: Sent 938525 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8352 bytes expected).
#TF#SERVER: Received 8352 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 622 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1536, query=402, total=1944).
#TF#SERVER: Sent 887983 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7756 bytes expected).
#TF#SERVER: Received 7756 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 500 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1536, query=529, total=2071).
#TF#SERVER: Sent 968344 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8264 bytes expected).
#TF#SERVER: Received 8264 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 507 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1536, query=519, total=2061).
#TF#SERVER: Sent 955257 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8224 bytes expected).
#TF#SERVER: Received 8224 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 481 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1536, query=556, total=2098).
#TF#SERVER: Sent 912650 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8372 bytes expected).
#TF#SERVER: Received 8372 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 354 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1536, query=681, total=2223).
#TF#SERVER: Sent 1005073 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8872 bytes expected).
#TF#SERVER: Received 8872 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 341 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1536, query=689, total=2231).
#TF#SERVER: Sent 1068429 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8904 bytes expected).
#TF#SERVER: Received 8904 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 613 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1536, query=428, total=1970).
#TF#SERVER: Sent 933636 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7860 bytes expected).
#TF#SERVER: Received 7860 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 429 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1536, query=615, total=2157).
#TF#SERVER: Sent 1100396 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8608 bytes expected).
#TF#SERVER: Received 8608 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 387 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1536, query=651, total=2193).
#TF#SERVER: Sent 1114544 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8752 bytes expected).
#TF#SERVER: Received 8752 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 476 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1536, query=563, total=2105).
#TF#SERVER: Sent 1023725 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8400 bytes expected).
#TF#SERVER: Received 8400 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 458 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1536, query=567, total=2109).
#TF#SERVER: Sent 987478 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8416 bytes expected).
#TF#SERVER: Received 8416 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 575 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1536, query=452, total=1994).
#TF#SERVER: Sent 987997 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (7956 bytes expected).
#TF#SERVER: Received 7956 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 420 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1536, query=631, total=2173).
#TF#SERVER: Sent 1107133 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (8672 bytes expected).
#TF#SERVER: Received 8672 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 326 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1536, query=787, total=2329).
#TF#SERVER: Sent 1309349 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9296 bytes expected).
#TF#SERVER: Received 9296 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 330 clauses
#TF#SERVER: Sending data to the TF server (conj=6, context=1536, query=810, total=2352).
#TF#SERVER: Sent 1152717 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (9388 bytes expected).
#TF#SERVER: Received 9388 bytes from the TF server.

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t14_fdiff_8, conjecture, ![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>(r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))=>(r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)&![X2]:(m1_subset_1(X2,k1_numbers)=>(r2_hidden(X2,X1)=>k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1),X2)=k10_real_1(np__1,k8_real_1(X2,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,X2)))))))))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', t14_fdiff_8)).
fof(t11_funct_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v1_funct_1(X2))=>![X3]:((v1_relat_1(X3)&v1_funct_1(X3))=>(r2_hidden(X1,k9_xtuple_0(k3_relat_1(X3,X2)))<=>(r2_hidden(X1,k9_xtuple_0(X3))&r2_hidden(k1_funct_1(X3,X1),k9_xtuple_0(X2)))))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', t11_funct_1)).
fof(redefinition_k1_seq_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>k1_seq_1(X1,X2)=k1_funct_1(X1,X2)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', redefinition_k1_seq_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', cc2_relset_1)).
fof(dt_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:((((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_funct_1(X6))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>(v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))&m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4))))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', dt_k1_partfun1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', t3_subset)).
fof(redefinition_k1_partfun1, axiom, ![X1, X2, X3, X4, X5, X6]:((((v1_funct_1(X5)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&v1_funct_1(X6))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))=>k1_partfun1(X1,X2,X3,X4,X5,X6)=k3_relat_1(X5,X6)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', redefinition_k1_partfun1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>k1_relset_1(X1,X2)=k9_xtuple_0(X2)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', redefinition_k1_relset_1)).
fof(dt_k3_relat_1, axiom, ![X1, X2]:v1_relat_1(k3_relat_1(X1,X2)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', dt_k3_relat_1)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', t4_subset)).
fof(fraenkel_a_0_0_fdiff_8, axiom, ![X1]:(r2_hidden(X1,a_0_0_fdiff_8)<=>?[X2]:((m1_subset_1(X2,k1_numbers)&X1=X2)&~(r1_xxreal_0(X2,k6_numbers)))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', fraenkel_a_0_0_fdiff_8)).
fof(t25_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(v1_relat_1(X2)=>r1_tarski(k9_xtuple_0(k3_relat_1(X1,X2)),k9_xtuple_0(X1)))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', t25_relat_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', cc1_relset_1)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', t2_subset)).
fof(t5_subset, axiom, ![X1, X2, X3]:~(((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))&v1_xboole_0(X3))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', t5_subset)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', t1_subset)).
fof(t9_fdiff_1, axiom, ![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(r2_fdiff_1(X2,X1)<=>(r1_tarski(X1,k1_relset_1(k1_numbers,X2))&![X3]:(m1_subset_1(X3,k1_numbers)=>(r2_hidden(X3,X1)=>r1_fdiff_1(X2,X3))))))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', t9_fdiff_1)).
fof(t18_taylor_1, axiom, (((((((r2_relset_1(k1_numbers,k1_numbers,k3_taylor_1,k2_partfun2(k1_numbers,k1_numbers,k24_sin_cos))&v2_funct_1(k3_taylor_1))&k1_relset_1(k1_numbers,k3_taylor_1)=k3_limfunc1(k6_numbers))&k2_relset_1(k1_numbers,k3_taylor_1)=k1_numbers)&r2_fdiff_1(k3_taylor_1,k3_limfunc1(k6_numbers)))&![X1]:(m1_subset_1(X1,k1_numbers)=>(~(r1_xxreal_0(X1,k6_numbers))=>r1_fdiff_1(k3_taylor_1,X1))))&![X1]:(m2_subset_1(X1,k1_numbers,k3_limfunc1(k6_numbers))=>k1_fdiff_1(k3_taylor_1,X1)=k10_real_1(np__1,X1)))&![X1]:(m2_subset_1(X1,k1_numbers,k3_limfunc1(k6_numbers))=>~(r1_xxreal_0(k1_fdiff_1(k3_taylor_1,X1),k6_numbers)))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', t18_taylor_1)).
fof(dt_k29_sin_cos, axiom, (v1_funct_1(k29_sin_cos)&m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', dt_k29_sin_cos)).
fof(cc18_valued_0, axiom, ![X1, X2]:(v3_membered(X2)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v3_valued_0(X3))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', cc18_valued_0)).
fof(t13_fdiff_2, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>![X2]:((v1_funct_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>((r1_fdiff_1(X2,X1)&r1_fdiff_1(X3,k1_seq_1(X2,X1)))=>(r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),X1)&k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),X1)=k8_real_1(k1_fdiff_1(X3,k1_seq_1(X2,X1)),k1_fdiff_1(X2,X1))))))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', t13_fdiff_2)).
fof(t7_boole, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', t7_boole)).
fof(dt_k3_taylor_1, axiom, (v1_funct_1(k3_taylor_1)&m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', dt_k3_taylor_1)).
fof(l5_fdiff_8, axiom, k3_limfunc1(k6_numbers)=a_0_0_fdiff_8, file('/scratch/mptp/minsubpreds5/t14_fdiff_8', l5_fdiff_8)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', t1_xboole_1)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', fc3_membered)).
fof(dt_k1_seq_1, axiom, ![X1, X2]:(((v1_relat_1(X1)&v1_funct_1(X1))&v3_valued_0(X1))=>m1_subset_1(k1_seq_1(X1,X2),k1_numbers)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', dt_k1_seq_1)).
fof(t46_fdiff_7, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>(k1_seq_1(k19_sin_cos,X1)!=k6_numbers=>(r1_fdiff_1(k3_rfunct_1(k1_numbers,k1_numbers,k16_sin_cos,k19_sin_cos),X1)&k1_fdiff_1(k3_rfunct_1(k1_numbers,k1_numbers,k16_sin_cos,k19_sin_cos),X1)=k10_real_1(np__1,k5_square_1(k1_seq_1(k19_sin_cos,X1)))))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', t46_fdiff_7)).
fof(t1_fdiff_8, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>~((r2_hidden(X1,k1_relset_1(k1_numbers,k29_sin_cos))&k1_seq_1(k19_sin_cos,X1)=k6_numbers))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', t1_fdiff_8)).
fof(d26_sin_cos, axiom, k29_sin_cos=k3_rfunct_1(k1_numbers,k1_numbers,k16_sin_cos,k19_sin_cos), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', d26_sin_cos)).
fof(redefinition_k10_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', redefinition_k10_real_1)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', spc1_numerals)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', redefinition_m2_subset_1)).
fof(fc1_numbers, axiom, ~(v1_xboole_0(k1_numbers)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', fc1_numbers)).
fof(dt_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>m1_subset_1(k1_relset_1(X1,X2),k1_zfmisc_1(X1))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', dt_k1_relset_1)).
fof(cc14_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xreal_0(X2))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', cc14_membered)).
fof(reflexivity_r1_tarski, axiom, ![X1, X2]:r1_tarski(X1,X1), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', reflexivity_r1_tarski)).
fof(dt_k1_fdiff_1, axiom, ![X1, X2]:(((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))&v1_xreal_0(X2))=>m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', dt_k1_fdiff_1)).
fof(redefinition_k8_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', redefinition_k8_real_1)).
fof(commutativity_k8_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k8_real_1(X1,X2)=k8_real_1(X2,X1)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', commutativity_k8_real_1)).
fof(dt_k8_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>m1_subset_1(k8_real_1(X1,X2),k1_numbers)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', dt_k8_real_1)).
fof(dt_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)=>m1_subset_1(X3,X1))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', dt_m2_subset_1)).
fof(cc12_membered, axiom, ![X1]:(v1_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xcmplx_0(X2))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', cc12_membered)).
fof(cc21_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v3_membered(X2))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', cc21_membered)).
fof(involutiveness_k5_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>k5_xcmplx_0(k5_xcmplx_0(X1))=X1), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', involutiveness_k5_xcmplx_0)).
fof(spc3_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k7_xcmplx_0(np__1,X1)=k5_xcmplx_0(X1)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', spc3_arithm)).
fof(cc5_membered, axiom, ![X1]:(v3_membered(X1)=>v1_membered(X1)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', cc5_membered)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', dt_k5_numbers)).
fof(dt_k5_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>v1_xcmplx_0(k5_xcmplx_0(X1))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', dt_k5_xcmplx_0)).
fof(t78_xcmplx_1, axiom, ![X1]:(v1_xcmplx_0(X1)=>![X2]:(v1_xcmplx_0(X2)=>![X3]:(v1_xcmplx_0(X3)=>k7_xcmplx_0(X1,k3_xcmplx_0(X2,X3))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X3)))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', t78_xcmplx_1)).
fof(dt_k5_square_1, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>m1_subset_1(k5_square_1(X1),k1_numbers)), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', dt_k5_square_1)).
fof(dt_k19_sin_cos, axiom, ((v1_funct_1(k19_sin_cos)&v1_funct_2(k19_sin_cos,k1_numbers,k1_numbers))&m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', dt_k19_sin_cos)).
fof(d7_fdiff_1, axiom, ![X1]:((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>![X2]:(r2_fdiff_1(X1,X2)=>![X3]:((v1_funct_1(X3)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(X3=k2_fdiff_1(X1,X2)<=>(k1_relset_1(k1_numbers,X3)=X2&![X4]:(m1_subset_1(X4,k1_numbers)=>(r2_hidden(X4,X2)=>k1_seq_1(X3,X4)=k1_fdiff_1(X1,X4)))))))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', d7_fdiff_1)).
fof(dt_k2_fdiff_1, axiom, ![X1, X2]:((v1_funct_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))=>(v1_funct_1(k2_fdiff_1(X1,X2))&m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', dt_k2_fdiff_1)).
fof(d9_xcmplx_0, axiom, ![X1]:(v1_xcmplx_0(X1)=>![X2]:(v1_xcmplx_0(X2)=>k7_xcmplx_0(X1,X2)=k3_xcmplx_0(X1,k5_xcmplx_0(X2)))), file('/scratch/mptp/minsubpreds5/t14_fdiff_8', d9_xcmplx_0)).
fof(c_0_55, negated_conjecture, ~(![X1]:((v3_rcomp_1(X1)&m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))=>(r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))=>(r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)&![X2]:(m1_subset_1(X2,k1_numbers)=>(r2_hidden(X2,X1)=>k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1),X2)=k10_real_1(np__1,k8_real_1(X2,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,X2))))))))))), inference(assume_negation,[status(cth)],[t14_fdiff_8])).
fof(c_0_56, plain, ![X112, X113, X114]:(((r2_hidden(X112,k9_xtuple_0(X114))|~r2_hidden(X112,k9_xtuple_0(k3_relat_1(X114,X113)))|(~v1_relat_1(X114)|~v1_funct_1(X114))|(~v1_relat_1(X113)|~v1_funct_1(X113)))&(r2_hidden(k1_funct_1(X114,X112),k9_xtuple_0(X113))|~r2_hidden(X112,k9_xtuple_0(k3_relat_1(X114,X113)))|(~v1_relat_1(X114)|~v1_funct_1(X114))|(~v1_relat_1(X113)|~v1_funct_1(X113))))&(~r2_hidden(X112,k9_xtuple_0(X114))|~r2_hidden(k1_funct_1(X114,X112),k9_xtuple_0(X113))|r2_hidden(X112,k9_xtuple_0(k3_relat_1(X114,X113)))|(~v1_relat_1(X114)|~v1_funct_1(X114))|(~v1_relat_1(X113)|~v1_funct_1(X113)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t11_funct_1])])])])).
fof(c_0_57, plain, ![X91, X92]:(~v1_relat_1(X91)|~v1_funct_1(X91)|~v3_valued_0(X91)|k1_seq_1(X91,X92)=k1_funct_1(X91,X92)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_seq_1])])).
fof(c_0_58, plain, ![X24, X25, X26]:((v4_relat_1(X26,X24)|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(X24,X25))))&(v5_relat_1(X26,X25)|~m1_subset_1(X26,k1_zfmisc_1(k2_zfmisc_1(X24,X25))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_59, plain, ![X46, X47, X48, X49, X50, X51]:((v1_funct_1(k1_partfun1(X46,X47,X48,X49,X50,X51))|(~v1_funct_1(X50)|~m1_subset_1(X50,k1_zfmisc_1(k2_zfmisc_1(X46,X47)))|~v1_funct_1(X51)|~m1_subset_1(X51,k1_zfmisc_1(k2_zfmisc_1(X48,X49)))))&(m1_subset_1(k1_partfun1(X46,X47,X48,X49,X50,X51),k1_zfmisc_1(k2_zfmisc_1(X46,X49)))|(~v1_funct_1(X50)|~m1_subset_1(X50,k1_zfmisc_1(k2_zfmisc_1(X46,X47)))|~v1_funct_1(X51)|~m1_subset_1(X51,k1_zfmisc_1(k2_zfmisc_1(X48,X49)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_partfun1])])])).
fof(c_0_60, plain, ![X137, X138]:((~m1_subset_1(X137,k1_zfmisc_1(X138))|r1_tarski(X137,X138))&(~r1_tarski(X137,X138)|m1_subset_1(X137,k1_zfmisc_1(X138)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_61, negated_conjecture, ((v3_rcomp_1(esk3_0)&m1_subset_1(esk3_0,k1_zfmisc_1(k1_numbers)))&(r1_tarski(esk3_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))&((m1_subset_1(esk4_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0))&((r2_hidden(esk4_0,esk3_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0))&(k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)!=k10_real_1(np__1,k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_55])])])])).
cnf(c_0_62, plain, (r2_hidden(k1_funct_1(X1,X2),k9_xtuple_0(X3))|~r2_hidden(X2,k9_xtuple_0(k3_relat_1(X1,X3)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_63, plain, (k1_seq_1(X1,X2)=k1_funct_1(X1,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
fof(c_0_64, plain, ![X83, X84, X85, X86, X87, X88]:(~v1_funct_1(X87)|~m1_subset_1(X87,k1_zfmisc_1(k2_zfmisc_1(X83,X84)))|~v1_funct_1(X88)|~m1_subset_1(X88,k1_zfmisc_1(k2_zfmisc_1(X85,X86)))|k1_partfun1(X83,X84,X85,X86,X87,X88)=k3_relat_1(X87,X88)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_partfun1])])).
fof(c_0_65, plain, ![X89, X90]:(~v1_relat_1(X90)|~v4_relat_1(X90,X89)|k1_relset_1(X89,X90)=k9_xtuple_0(X90)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_66, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_67, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_59])).
fof(c_0_68, plain, ![X59, X60]:v1_relat_1(k3_relat_1(X59,X60)), inference(variable_rename,[status(thm)],[dt_k3_relat_1])).
fof(c_0_69, plain, ![X140, X141, X142]:(~r2_hidden(X140,X141)|~m1_subset_1(X141,k1_zfmisc_1(X142))|m1_subset_1(X140,X142)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_70, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_71, negated_conjecture, (r1_tarski(esk3_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))), inference(split_conjunct,[status(thm)],[c_0_61])).
fof(c_0_72, plain, ![X76, X78, X79]:((((m1_subset_1(esk2_1(X76),k1_numbers)|~r2_hidden(X76,a_0_0_fdiff_8))&(X76=esk2_1(X76)|~r2_hidden(X76,a_0_0_fdiff_8)))&(~r1_xxreal_0(esk2_1(X76),k6_numbers)|~r2_hidden(X76,a_0_0_fdiff_8)))&(~m1_subset_1(X79,k1_numbers)|X78!=X79|r1_xxreal_0(X79,k6_numbers)|r2_hidden(X78,a_0_0_fdiff_8))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fraenkel_a_0_0_fdiff_8])])])])])])])).
fof(c_0_73, plain, ![X129, X130]:(~v1_relat_1(X129)|(~v1_relat_1(X130)|r1_tarski(k9_xtuple_0(k3_relat_1(X129,X130)),k9_xtuple_0(X129)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t25_relat_1])])])).
cnf(c_0_74, plain, (r2_hidden(k1_funct_1(X1,X2),k9_xtuple_0(X3))|~v1_relat_1(X3)|~v1_relat_1(X1)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_hidden(X2,k9_xtuple_0(k3_relat_1(X1,X3)))), inference(evalgc,[status(thm)],[c_0_62])).
cnf(c_0_75, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[c_0_63])).
cnf(c_0_76, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))), inference(split_conjunct,[status(thm)],[c_0_64])).
fof(c_0_77, plain, ![X18, X19, X20]:(~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(X18,X19)))|v1_relat_1(X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_78, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
cnf(c_0_79, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_66])).
cnf(c_0_80, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_67])).
cnf(c_0_81, plain, (v1_relat_1(k3_relat_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_68])).
fof(c_0_82, plain, ![X131, X132]:(~m1_subset_1(X131,X132)|(v1_xboole_0(X132)|r2_hidden(X131,X132))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_83, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_84, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_70])).
cnf(c_0_85, negated_conjecture, (r1_tarski(esk3_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))), inference(evalgc,[status(thm)],[c_0_71])).
fof(c_0_86, plain, ![X143, X144, X145]:(~r2_hidden(X143,X144)|~m1_subset_1(X144,k1_zfmisc_1(X145))|~v1_xboole_0(X145)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_subset])])).
fof(c_0_87, plain, ![X124, X125]:(~r2_hidden(X124,X125)|m1_subset_1(X124,X125)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
fof(c_0_88, plain, ![X153, X154, X155]:(((r1_tarski(X153,k1_relset_1(k1_numbers,X154))|~r2_fdiff_1(X154,X153)|(~v1_funct_1(X154)|~m1_subset_1(X154,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X153)|~m1_subset_1(X153,k1_zfmisc_1(k1_numbers))))&(~m1_subset_1(X155,k1_numbers)|(~r2_hidden(X155,X153)|r1_fdiff_1(X154,X155))|~r2_fdiff_1(X154,X153)|(~v1_funct_1(X154)|~m1_subset_1(X154,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X153)|~m1_subset_1(X153,k1_zfmisc_1(k1_numbers)))))&((m1_subset_1(esk6_2(X153,X154),k1_numbers)|~r1_tarski(X153,k1_relset_1(k1_numbers,X154))|r2_fdiff_1(X154,X153)|(~v1_funct_1(X154)|~m1_subset_1(X154,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X153)|~m1_subset_1(X153,k1_zfmisc_1(k1_numbers))))&((r2_hidden(esk6_2(X153,X154),X153)|~r1_tarski(X153,k1_relset_1(k1_numbers,X154))|r2_fdiff_1(X154,X153)|(~v1_funct_1(X154)|~m1_subset_1(X154,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X153)|~m1_subset_1(X153,k1_zfmisc_1(k1_numbers))))&(~r1_fdiff_1(X154,esk6_2(X153,X154))|~r1_tarski(X153,k1_relset_1(k1_numbers,X154))|r2_fdiff_1(X154,X153)|(~v1_funct_1(X154)|~m1_subset_1(X154,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v3_rcomp_1(X153)|~m1_subset_1(X153,k1_zfmisc_1(k1_numbers))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t9_fdiff_1])])])])])).
cnf(c_0_89, plain, (~r1_xxreal_0(esk2_1(X1),k6_numbers)|~r2_hidden(X1,a_0_0_fdiff_8)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_90, plain, (X1=esk2_1(X1)|~r2_hidden(X1,a_0_0_fdiff_8)), inference(split_conjunct,[status(thm)],[c_0_72])).
fof(c_0_91, plain, ![X120, X121, X122]:(((((((r2_relset_1(k1_numbers,k1_numbers,k3_taylor_1,k2_partfun2(k1_numbers,k1_numbers,k24_sin_cos))&v2_funct_1(k3_taylor_1))&k1_relset_1(k1_numbers,k3_taylor_1)=k3_limfunc1(k6_numbers))&k2_relset_1(k1_numbers,k3_taylor_1)=k1_numbers)&r2_fdiff_1(k3_taylor_1,k3_limfunc1(k6_numbers)))&(~m1_subset_1(X120,k1_numbers)|(r1_xxreal_0(X120,k6_numbers)|r1_fdiff_1(k3_taylor_1,X120))))&(~m2_subset_1(X121,k1_numbers,k3_limfunc1(k6_numbers))|k1_fdiff_1(k3_taylor_1,X121)=k10_real_1(np__1,X121)))&(~m2_subset_1(X122,k1_numbers,k3_limfunc1(k6_numbers))|~r1_xxreal_0(k1_fdiff_1(k3_taylor_1,X122),k6_numbers))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t18_taylor_1])])])])).
cnf(c_0_92, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,X3)))|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_funct_1(X3)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_93, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(X1,X2)),k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_relat_1(X2)), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_94, plain, (r2_hidden(k1_seq_1(X1,X2),k9_xtuple_0(X3))|~r2_hidden(X2,k9_xtuple_0(k3_relat_1(X1,X3)))|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X1)|~v3_valued_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75])])).
cnf(c_0_95, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_96, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_97, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_78])).
cnf(c_0_98, plain, (v4_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6),X1)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_80])])).
cnf(c_0_99, plain, (v1_relat_1(k3_relat_1(X1,X2))), inference(evalgc,[status(thm)],[c_0_81])).
cnf(c_0_100, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_82])).
cnf(c_0_101, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_83])).
cnf(c_0_102, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_85])])).
cnf(c_0_103, plain, (~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_xboole_0(X3)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_104, plain, (m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k29_sin_cos])).
fof(c_0_105, plain, ![X15, X16, X17]:(~v3_membered(X16)|(~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(X15,X16)))|v3_valued_0(X17))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc18_valued_0])])])).
fof(c_0_106, plain, ![X115, X116, X117]:((r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X116,X117),X115)|(~r1_fdiff_1(X116,X115)|~r1_fdiff_1(X117,k1_seq_1(X116,X115)))|(~v1_funct_1(X117)|~m1_subset_1(X117,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v1_funct_1(X116)|~m1_subset_1(X116,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~m1_subset_1(X115,k1_numbers))&(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X116,X117),X115)=k8_real_1(k1_fdiff_1(X117,k1_seq_1(X116,X115)),k1_fdiff_1(X116,X115))|(~r1_fdiff_1(X116,X115)|~r1_fdiff_1(X117,k1_seq_1(X116,X115)))|(~v1_funct_1(X117)|~m1_subset_1(X117,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|(~v1_funct_1(X116)|~m1_subset_1(X116,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~m1_subset_1(X115,k1_numbers))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_fdiff_2])])])])).
cnf(c_0_107, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_108, plain, (r2_hidden(esk6_2(X1,X2),X1)|r2_fdiff_1(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_88])).
fof(c_0_109, plain, ![X151, X152]:(~r2_hidden(X151,X152)|~v1_xboole_0(X152)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
cnf(c_0_110, plain, (~r2_hidden(X1,a_0_0_fdiff_8)|~r1_xxreal_0(esk2_1(X1),k6_numbers)), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_111, plain, (esk2_1(X1)=X1|~r2_hidden(X1,a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_112, plain, (r1_xxreal_0(X1,k6_numbers)|r1_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_113, plain, (m1_subset_1(esk2_1(X1),k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_8)), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_114, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_92])).
cnf(c_0_115, plain, (m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k3_taylor_1])).
cnf(c_0_116, plain, (k1_relset_1(k1_numbers,k3_taylor_1)=k3_limfunc1(k6_numbers)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_117, plain, (k3_limfunc1(k6_numbers)=a_0_0_fdiff_8), inference(split_conjunct,[status(thm)],[l5_fdiff_8])).
fof(c_0_118, plain, ![X126, X127, X128]:(~r1_tarski(X126,X127)|~r1_tarski(X127,X128)|r1_tarski(X126,X128)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
cnf(c_0_119, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(X1,X2)),k9_xtuple_0(X1))|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[c_0_93])).
cnf(c_0_120, plain, (r2_hidden(k1_seq_1(X1,X2),k9_xtuple_0(X3))|~r2_hidden(X2,k9_xtuple_0(k1_partfun1(X4,X5,X6,X7,X1,X3)))|~v1_funct_1(X3)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X3)|~v1_relat_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95])])).
cnf(c_0_121, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(evalgc,[status(thm)],[c_0_96])).
cnf(c_0_122, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,X5,X6))=k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98])])).
cnf(c_0_123, plain, (v1_relat_1(k1_partfun1(X3,X4,X5,X6,X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_95])])).
cnf(c_0_124, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_100])).
cnf(c_0_125, negated_conjecture, (m1_subset_1(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102])])).
cnf(c_0_126, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_103])).
cnf(c_0_127, plain, (m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_104])).
cnf(c_0_128, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_105])).
cnf(c_0_129, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
fof(c_0_130, plain, ![X54, X55]:(~v1_relat_1(X54)|~v1_funct_1(X54)|~v3_valued_0(X54)|m1_subset_1(k1_seq_1(X54,X55),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_seq_1])])).
cnf(c_0_131, plain, (r2_fdiff_1(X1,X2)|~r1_fdiff_1(X1,esk6_2(X2,X1))|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_132, plain, (r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~r1_fdiff_1(X1,X3)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_106])).
fof(c_0_133, plain, ![X139]:((r1_fdiff_1(k3_rfunct_1(k1_numbers,k1_numbers,k16_sin_cos,k19_sin_cos),X139)|k1_seq_1(k19_sin_cos,X139)=k6_numbers|~m1_subset_1(X139,k1_numbers))&(k1_fdiff_1(k3_rfunct_1(k1_numbers,k1_numbers,k16_sin_cos,k19_sin_cos),X139)=k10_real_1(np__1,k5_square_1(k1_seq_1(k19_sin_cos,X139)))|k1_seq_1(k19_sin_cos,X139)=k6_numbers|~m1_subset_1(X139,k1_numbers))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t46_fdiff_7])])])).
cnf(c_0_134, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_107])).
cnf(c_0_135, plain, (r2_fdiff_1(X2,X1)|r2_hidden(esk6_2(X1,X2),X1)|~v1_funct_1(X2)|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_108])).
cnf(c_0_136, negated_conjecture, (v3_rcomp_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_137, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_138, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_109])).
cnf(c_0_139, plain, (~r1_xxreal_0(X1,k6_numbers)|~r2_hidden(X1,a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_111])])).
cnf(c_0_140, plain, (r1_xxreal_0(X1,k6_numbers)|r1_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_112])).
cnf(c_0_141, plain, (m1_subset_1(esk2_1(X1),k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[c_0_113])).
cnf(c_0_142, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(X1,k9_xtuple_0(k1_partfun1(X4,X5,X6,X7,X2,X3)))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_95])])).
cnf(c_0_143, plain, (m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_115])).
cnf(c_0_144, plain, (k3_limfunc1(k6_numbers)=k1_relset_1(k1_numbers,k3_taylor_1)), inference(evalgc,[status(thm)],[c_0_116])).
cnf(c_0_145, plain, (k3_limfunc1(k6_numbers)=a_0_0_fdiff_8), inference(evalgc,[status(thm)],[c_0_117])).
cnf(c_0_146, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_147, plain, (r1_tarski(k9_xtuple_0(k1_partfun1(X3,X4,X5,X6,X1,X2)),k9_xtuple_0(X1))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_95])])).
fof(c_0_148, plain, ![X123]:(~m1_subset_1(X123,k1_numbers)|(~r2_hidden(X123,k1_relset_1(k1_numbers,k29_sin_cos))|k1_seq_1(k19_sin_cos,X123)!=k6_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_fdiff_8])])).
cnf(c_0_149, plain, (r2_hidden(k1_seq_1(X1,X2),k9_xtuple_0(X3))|~r2_hidden(X2,k9_xtuple_0(k1_partfun1(X4,X5,X6,X7,X1,X3)))|~v1_funct_1(X3)|~v1_funct_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_120, c_0_121]), c_0_121])).
cnf(c_0_150, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,X5,X6))=k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(csr,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_151, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_hidden(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125])])).
cnf(c_0_152, negated_conjecture, (~v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_102])])).
cnf(c_0_153, plain, (v4_relat_1(k29_sin_cos,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_127])])).
cnf(c_0_154, plain, (v1_relat_1(k29_sin_cos)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_127])])).
cnf(c_0_155, plain, (v1_funct_1(k29_sin_cos)), inference(split_conjunct,[status(thm)],[dt_k29_sin_cos])).
cnf(c_0_156, plain, (v1_funct_1(k3_taylor_1)), inference(split_conjunct,[status(thm)],[dt_k3_taylor_1])).
cnf(c_0_157, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(evalgc,[status(thm)],[c_0_128])).
cnf(c_0_158, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_129])).
cnf(c_0_159, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v3_valued_0(X1)), inference(split_conjunct,[status(thm)],[c_0_130])).
cnf(c_0_160, plain, (r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r1_fdiff_1(X1,esk6_2(X2,X1))|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_131])).
cnf(c_0_161, plain, (r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X3)|~m1_subset_1(X3,k1_numbers)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_132])).
cnf(c_0_162, plain, (m1_subset_1(esk6_2(X1,X2),k1_numbers)|r2_fdiff_1(X2,X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_163, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_164, plain, (r1_fdiff_1(k3_rfunct_1(k1_numbers,k1_numbers,k16_sin_cos,k19_sin_cos),X1)|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_133])).
cnf(c_0_165, plain, (k29_sin_cos=k3_rfunct_1(k1_numbers,k1_numbers,k16_sin_cos,k19_sin_cos)), inference(split_conjunct,[status(thm)],[d26_sin_cos])).
cnf(c_0_166, plain, (m1_subset_1(esk6_2(X1,X2),X1)|r2_fdiff_1(X2,X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135])])).
cnf(c_0_167, negated_conjecture, (v3_rcomp_1(esk3_0)), inference(evalgc,[status(thm)],[c_0_136])).
cnf(c_0_168, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_169, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_138])).
cnf(c_0_170, plain, (r1_fdiff_1(k3_taylor_1,X1)|~r2_hidden(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140])])).
cnf(c_0_171, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_111])])).
cnf(c_0_172, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(X1,k9_xtuple_0(k1_partfun1(X3,X4,X5,X6,X2,X7)))|~v1_funct_1(X7)|~v1_funct_1(X2)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_142, c_0_121]), c_0_121])).
cnf(c_0_173, plain, (v4_relat_1(k3_taylor_1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_143])])).
cnf(c_0_174, plain, (k1_relset_1(k1_numbers,k3_taylor_1)=a_0_0_fdiff_8), inference(rw,[status(thm)],[c_0_144, c_0_145])).
cnf(c_0_175, plain, (v1_relat_1(k3_taylor_1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_143])])).
cnf(c_0_176, plain, (r1_tarski(X1,X3)|~r1_tarski(X2,X3)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_146])).
cnf(c_0_177, plain, (r1_tarski(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,X6)),k9_xtuple_0(X5))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_147, c_0_121]), c_0_121])).
cnf(c_0_178, plain, (~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k29_sin_cos))|k1_seq_1(k19_sin_cos,X1)!=k6_numbers), inference(split_conjunct,[status(thm)],[c_0_148])).
cnf(c_0_179, plain, (r2_hidden(k1_seq_1(X1,X2),k9_xtuple_0(X3))|~r2_hidden(X2,k1_relset_1(X4,k1_partfun1(X4,X5,X6,X7,X1,X3)))|~v1_funct_1(X3)|~v1_funct_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_150])])).
cnf(c_0_180, negated_conjecture, (r2_hidden(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,esk3_0)), inference(csr,[status(thm)],[c_0_151, c_0_152])).
cnf(c_0_181, plain, (k9_xtuple_0(k29_sin_cos)=k1_relset_1(k1_numbers,k29_sin_cos)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_153]), c_0_154])])])).
cnf(c_0_182, plain, (v1_funct_1(k29_sin_cos)), inference(evalgc,[status(thm)],[c_0_155])).
cnf(c_0_183, plain, (v1_funct_1(k3_taylor_1)), inference(evalgc,[status(thm)],[c_0_156])).
cnf(c_0_184, plain, (v3_valued_0(k3_taylor_1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_143]), c_0_158])])])).
cnf(c_0_185, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[c_0_159])).
cnf(c_0_186, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v3_rcomp_1(X3)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~r1_fdiff_1(X2,k1_seq_1(X1,esk6_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))))|~r1_fdiff_1(X1,esk6_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk6_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_161])])).
cnf(c_0_187, plain, (r2_fdiff_1(X2,X1)|m1_subset_1(esk6_2(X1,X2),k1_numbers)|~v1_funct_1(X2)|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_162])).
cnf(c_0_188, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[c_0_163])).
cnf(c_0_189, plain, (k1_seq_1(k19_sin_cos,X1)=k6_numbers|r1_fdiff_1(k3_rfunct_1(k1_numbers,k1_numbers,k16_sin_cos,k19_sin_cos),X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_164])).
cnf(c_0_190, plain, (k3_rfunct_1(k1_numbers,k1_numbers,k16_sin_cos,k19_sin_cos)=k29_sin_cos), inference(evalgc,[status(thm)],[c_0_165])).
cnf(c_0_191, plain, (r1_fdiff_1(X3,X1)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,X2)|~r2_fdiff_1(X3,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_192, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|m1_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_85]), c_0_167]), c_0_168])])])).
cnf(c_0_193, plain, (r2_fdiff_1(X2,X1)|~v1_xboole_0(X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_135])])).
cnf(c_0_194, plain, (r1_fdiff_1(k3_taylor_1,X1)|~r2_hidden(X1,a_0_0_fdiff_8)), inference(csr,[status(thm)],[c_0_170, c_0_171])).
cnf(c_0_195, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(X1,k1_relset_1(X3,k1_partfun1(X3,X4,X5,X6,X2,X7)))|~v1_funct_1(X7)|~v1_funct_1(X2)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_150])])).
cnf(c_0_196, plain, (k9_xtuple_0(k3_taylor_1)=a_0_0_fdiff_8), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_173]), c_0_174]), c_0_175])])])).
cnf(c_0_197, plain, (r1_tarski(X1,k9_xtuple_0(X2))|~r1_tarski(X1,k9_xtuple_0(k1_partfun1(X3,X4,X5,X6,X2,X7)))|~v1_funct_1(X7)|~v1_funct_1(X2)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_177])])).
cnf(c_0_198, plain, (k1_seq_1(k19_sin_cos,X1)!=k6_numbers|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k29_sin_cos))), inference(evalgc,[status(thm)],[c_0_178])).
cnf(c_0_199, negated_conjecture, (r2_hidden(k1_seq_1(k3_taylor_1,X1),k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_180]), c_0_181]), c_0_182]), c_0_183]), c_0_184]), c_0_127]), c_0_143])])])).
cnf(c_0_200, plain, (m1_subset_1(k1_seq_1(k3_taylor_1,X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_184]), c_0_183]), c_0_175])])])).
cnf(c_0_201, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v3_rcomp_1(X3)|~r1_fdiff_1(X2,k1_seq_1(X1,esk6_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))))|~r1_fdiff_1(X1,esk6_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_186, c_0_187]), c_0_80]), c_0_188])).
cnf(c_0_202, plain, (k1_seq_1(k19_sin_cos,X1)=k6_numbers|r1_fdiff_1(k29_sin_cos,X1)|~m1_subset_1(X1,k1_numbers)), inference(rw,[status(thm)],[c_0_189, c_0_190])).
cnf(c_0_203, plain, (r1_fdiff_1(X3,X1)|~v1_funct_1(X3)|~v3_rcomp_1(X2)|~r2_fdiff_1(X3,X2)|~r2_hidden(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_191])).
cnf(c_0_204, negated_conjecture, (v1_xboole_0(esk3_0)|r2_hidden(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),esk3_0)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_192])])).
cnf(c_0_205, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xboole_0(esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_85]), c_0_167]), c_0_168])])])).
cnf(c_0_206, plain, (r2_fdiff_1(k3_taylor_1,X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r2_hidden(esk6_2(X1,k3_taylor_1),a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_194]), c_0_174]), c_0_183]), c_0_143])])])).
cnf(c_0_207, negated_conjecture, (r2_hidden(X1,a_0_0_fdiff_8)|~r2_hidden(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_180]), c_0_196]), c_0_182]), c_0_183]), c_0_127]), c_0_143])])])).
cnf(c_0_208, plain, (r1_tarski(X1,k9_xtuple_0(X2))|~r1_tarski(X1,k1_relset_1(X3,k1_partfun1(X3,X4,X5,X6,X2,X7)))|~v1_funct_1(X7)|~v1_funct_1(X2)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_150])])).
cnf(c_0_209, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,X1))!=k6_numbers|~r2_hidden(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_199]), c_0_200])])])).
cnf(c_0_210, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k29_sin_cos),X2)|k1_seq_1(k19_sin_cos,k1_seq_1(X1,esk6_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k29_sin_cos))))=k6_numbers|~v3_rcomp_1(X2)|~r1_fdiff_1(X1,esk6_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k29_sin_cos)))|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k29_sin_cos)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(k1_seq_1(X1,esk6_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k29_sin_cos))),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_202]), c_0_182]), c_0_127])])])).
cnf(c_0_211, plain, (r1_fdiff_1(X1,X2)|~v3_rcomp_1(X3)|~r2_hidden(X2,X3)|~r2_fdiff_1(X1,X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))), inference(csr,[status(thm)],[c_0_203, c_0_101])).
cnf(c_0_212, negated_conjecture, (r2_hidden(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),esk3_0)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_204, c_0_205])).
cnf(c_0_213, negated_conjecture, (r2_fdiff_1(k3_taylor_1,X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r2_hidden(esk6_2(X1,k3_taylor_1),esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_207])])).
cnf(c_0_214, negated_conjecture, (r1_tarski(esk3_0,a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_85]), c_0_196]), c_0_182]), c_0_183]), c_0_127]), c_0_143])])])).
cnf(c_0_215, negated_conjecture, (m1_subset_1(esk4_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_216, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~r2_hidden(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),esk3_0)|~v3_rcomp_1(X1)|~r1_fdiff_1(k3_taylor_1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_210]), c_0_183]), c_0_200]), c_0_143])])])).
cnf(c_0_217, negated_conjecture, (r1_fdiff_1(X1,esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_212]), c_0_167]), c_0_168])])])).
cnf(c_0_218, negated_conjecture, (r2_fdiff_1(k3_taylor_1,esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_135]), c_0_167]), c_0_214]), c_0_168]), c_0_174]), c_0_214]), c_0_183]), c_0_143])])])).
cnf(c_0_219, plain, (k1_fdiff_1(k3_taylor_1,X1)=k10_real_1(np__1,X1)|~m2_subset_1(X1,k1_numbers,k3_limfunc1(k6_numbers))), inference(split_conjunct,[status(thm)],[c_0_91])).
fof(c_0_220, plain, ![X81, X82]:(~m1_subset_1(X81,k1_numbers)|~v1_xreal_0(X82)|k10_real_1(X81,X82)=k7_xcmplx_0(X81,X82)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k10_real_1])])).
cnf(c_0_221, negated_conjecture, (m1_subset_1(esk4_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)), inference(evalgc,[status(thm)],[c_0_215])).
cnf(c_0_222, plain, (k1_partfun1(X1,X2,X3,X4,X5,X6)=k1_partfun1(X7,X8,X9,X10,X5,X6)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_95])])).
cnf(c_0_223, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~r2_hidden(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_217]), c_0_167]), c_0_85]), c_0_168]), c_0_218]), c_0_183]), c_0_143])])])).
cnf(c_0_224, negated_conjecture, (r2_hidden(esk4_0,esk3_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_225, plain, (k10_real_1(np__1,X1)=k1_fdiff_1(k3_taylor_1,X1)|~m2_subset_1(X1,k1_numbers,k3_limfunc1(k6_numbers))), inference(evalgc,[status(thm)],[c_0_219])).
cnf(c_0_226, plain, (k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_220])).
cnf(c_0_227, plain, (m1_subset_1(np__1,k1_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
fof(c_0_228, plain, ![X96, X97, X98]:((~m2_subset_1(X98,X96,X97)|m1_subset_1(X98,X97)|(v1_xboole_0(X96)|v1_xboole_0(X97)|~m1_subset_1(X97,k1_zfmisc_1(X96))))&(~m1_subset_1(X98,X97)|m2_subset_1(X98,X96,X97)|(v1_xboole_0(X96)|v1_xboole_0(X97)|~m1_subset_1(X97,k1_zfmisc_1(X96))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])).
fof(c_0_229, plain, ~v1_xboole_0(k1_numbers), inference(fof_simplification,[status(thm)],[fc1_numbers])).
fof(c_0_230, plain, ![X52, X53]:(~v1_relat_1(X53)|~v4_relat_1(X53,X52)|m1_subset_1(k1_relset_1(X52,X53),k1_zfmisc_1(X52))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_relset_1])])).
cnf(c_0_231, negated_conjecture, (m1_subset_1(esk4_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),esk3_0)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_222]), c_0_182]), c_0_183]), c_0_127]), c_0_143])])])).
cnf(c_0_232, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_223, c_0_212])).
cnf(c_0_233, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r1_tarski(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_85])])).
cnf(c_0_234, negated_conjecture, (r2_hidden(esk4_0,esk3_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)), inference(evalgc,[status(thm)],[c_0_224])).
cnf(c_0_235, plain, (k10_real_1(np__1,X1)=k1_fdiff_1(k3_taylor_1,X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)), inference(rw,[status(thm)],[c_0_225, c_0_145])).
cnf(c_0_236, plain, (k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_226])).
cnf(c_0_237, plain, (m1_subset_1(np__1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_227])).
cnf(c_0_238, plain, (m2_subset_1(X1,X3,X2)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_228])).
cnf(c_0_239, plain, (~v1_xboole_0(k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_229])).
cnf(c_0_240, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_230])).
fof(c_0_241, plain, ![X11, X12]:(~v3_membered(X11)|(~m1_subset_1(X12,X11)|v1_xreal_0(X12))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc14_membered])])])).
cnf(c_0_242, negated_conjecture, (m1_subset_1(esk4_0,k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_232]), c_0_127]), c_0_143])])])).
cnf(c_0_243, negated_conjecture, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_233]), c_0_196]), c_0_182]), c_0_183]), c_0_127]), c_0_143])])])).
cnf(c_0_244, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_245, negated_conjecture, (r2_hidden(esk4_0,esk3_0)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),esk3_0)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_222]), c_0_182]), c_0_183]), c_0_127]), c_0_143])])])).
fof(c_0_246, plain, ![X103]:r1_tarski(X103,X103), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r1_tarski])])).
cnf(c_0_247, plain, (k1_fdiff_1(k3_taylor_1,X1)=k7_xcmplx_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_236]), c_0_237])])])).
cnf(c_0_248, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_238])).
cnf(c_0_249, plain, (~v1_xboole_0(k1_numbers)), inference(evalgc,[status(thm)],[c_0_239])).
cnf(c_0_250, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_240])).
cnf(c_0_251, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_241])).
cnf(c_0_252, negated_conjecture, (m1_subset_1(esk4_0,k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_80]), c_0_182]), c_0_183]), c_0_127]), c_0_143])])])).
cnf(c_0_253, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(a_0_0_fdiff_8))|~r1_tarski(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_243])])).
cnf(c_0_254, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_244])).
cnf(c_0_255, negated_conjecture, (r2_hidden(esk4_0,esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_232]), c_0_127]), c_0_143])])])).
cnf(c_0_256, plain, (r1_tarski(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_246])).
fof(c_0_257, plain, ![X44, X45]:(~v1_funct_1(X44)|~m1_subset_1(X44,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X45)|m1_subset_1(k1_fdiff_1(X44,X45),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_fdiff_1])])).
cnf(c_0_258, plain, (k1_fdiff_1(k3_taylor_1,X1)=k7_xcmplx_0(np__1,X1)|v1_xboole_0(a_0_0_fdiff_8)|~v1_xreal_0(X1)|~m1_subset_1(a_0_0_fdiff_8,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_248]), c_0_249])])).
cnf(c_0_259, plain, (m1_subset_1(a_0_0_fdiff_8,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_173]), c_0_174]), c_0_175])])])).
cnf(c_0_260, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_251])).
cnf(c_0_261, negated_conjecture, (m1_subset_1(esk4_0,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_188]), c_0_182]), c_0_183]), c_0_127]), c_0_143])])])).
cnf(c_0_262, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_254])])).
cnf(c_0_263, negated_conjecture, (r2_hidden(esk4_0,esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_80]), c_0_182]), c_0_183]), c_0_127]), c_0_143])])])).
cnf(c_0_264, plain, (r1_tarski(X1,X1)), inference(evalgc,[status(thm)],[c_0_256])).
fof(c_0_265, plain, ![X94, X95]:(~m1_subset_1(X94,k1_numbers)|~v1_xreal_0(X95)|k8_real_1(X94,X95)=k3_xcmplx_0(X94,X95)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_real_1])])).
fof(c_0_266, plain, ![X32, X33]:(~m1_subset_1(X32,k1_numbers)|~v1_xreal_0(X33)|k8_real_1(X32,X33)=k8_real_1(X33,X32)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k8_real_1])])).
fof(c_0_267, plain, ![X68, X69]:(~m1_subset_1(X68,k1_numbers)|~v1_xreal_0(X69)|m1_subset_1(k8_real_1(X68,X69),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_real_1])])).
cnf(c_0_268, plain, (m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_257])).
cnf(c_0_269, plain, (k1_fdiff_1(k3_taylor_1,X1)=k7_xcmplx_0(np__1,X1)|v1_xboole_0(a_0_0_fdiff_8)|~v1_xreal_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_258, c_0_259])])).
cnf(c_0_270, negated_conjecture, (v1_xreal_0(esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_261]), c_0_158])])])).
cnf(c_0_271, negated_conjecture, (m1_subset_1(X1,a_0_0_fdiff_8)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(esk3_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_262])])).
cnf(c_0_272, negated_conjecture, (r2_hidden(esk4_0,esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_188]), c_0_182]), c_0_183]), c_0_127]), c_0_143])])])).
cnf(c_0_273, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_264])])).
fof(c_0_274, plain, ![X70, X71, X72]:(v1_xboole_0(X70)|v1_xboole_0(X71)|~m1_subset_1(X71,k1_zfmisc_1(X70))|(~m2_subset_1(X72,X70,X71)|m1_subset_1(X72,X70))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_subset_1])])])])).
cnf(c_0_275, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)!=k10_real_1(np__1,k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_276, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_265])).
cnf(c_0_277, plain, (k8_real_1(X1,X2)=k8_real_1(X2,X1)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_266])).
cnf(c_0_278, plain, (m1_subset_1(k8_real_1(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_267])).
fof(c_0_279, plain, ![X9, X10]:(~v1_membered(X9)|(~m1_subset_1(X10,X9)|v1_xcmplx_0(X10))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc12_membered])])])).
fof(c_0_280, plain, ![X21, X22]:(~v3_membered(X21)|(~m1_subset_1(X22,k1_zfmisc_1(X21))|v3_membered(X22))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc21_membered])])])).
cnf(c_0_281, plain, (m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_268])).
cnf(c_0_282, negated_conjecture, (k1_fdiff_1(k3_taylor_1,esk4_0)=k7_xcmplx_0(np__1,esk4_0)|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(esk4_0,a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_269, c_0_270])])).
cnf(c_0_283, negated_conjecture, (m1_subset_1(esk4_0,a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_272]), c_0_273])])])).
fof(c_0_284, plain, ![X80]:(~v1_xcmplx_0(X80)|k5_xcmplx_0(k5_xcmplx_0(X80))=X80), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[involutiveness_k5_xcmplx_0])])).
fof(c_0_285, plain, ![X108]:(~v1_xcmplx_0(X108)|k7_xcmplx_0(np__1,X108)=k5_xcmplx_0(X108)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc3_arithm])])).
cnf(c_0_286, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m2_subset_1(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_274])).
cnf(c_0_287, negated_conjecture, (r2_hidden(esk4_0,esk3_0)|~v1_xboole_0(esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_205]), c_0_127]), c_0_143])])])).
cnf(c_0_288, negated_conjecture, (k10_real_1(np__1,k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)), inference(evalgc,[status(thm)],[c_0_275])).
cnf(c_0_289, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_276])).
cnf(c_0_290, plain, (k8_real_1(X1,X2)=k8_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_277])).
cnf(c_0_291, plain, (m1_subset_1(k8_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_278])).
cnf(c_0_292, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_279])).
cnf(c_0_293, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
fof(c_0_294, plain, ![X28]:(~v3_membered(X28)|v1_membered(X28)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_membered])])).
cnf(c_0_295, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_280])).
cnf(c_0_296, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
cnf(c_0_297, plain, (~m2_subset_1(X1,k1_numbers,k3_limfunc1(k6_numbers))|~r1_xxreal_0(k1_fdiff_1(k3_taylor_1,X1),k6_numbers)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_298, plain, (v1_xreal_0(k1_fdiff_1(X1,X2))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_281]), c_0_158])])])).
cnf(c_0_299, negated_conjecture, (k1_fdiff_1(k3_taylor_1,esk4_0)=k7_xcmplx_0(np__1,esk4_0)|v1_xboole_0(a_0_0_fdiff_8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_282, c_0_283])])).
cnf(c_0_300, negated_conjecture, (~v1_xboole_0(a_0_0_fdiff_8)|~r2_hidden(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_207])])).
cnf(c_0_301, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_8)|r2_hidden(esk4_0,a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_283])])).
cnf(c_0_302, plain, (k5_xcmplx_0(k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_284])).
cnf(c_0_303, plain, (k7_xcmplx_0(np__1,X1)=k5_xcmplx_0(X1)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_285])).
fof(c_0_304, plain, ![X67]:(~v1_xcmplx_0(X67)|v1_xcmplx_0(k5_xcmplx_0(X67))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_xcmplx_0])])).
cnf(c_0_305, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_286])).
cnf(c_0_306, negated_conjecture, (~v1_xboole_0(esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(csr,[status(thm)],[c_0_287, c_0_169])).
cnf(c_0_307, negated_conjecture, (k7_xcmplx_0(np__1,k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_236]), c_0_237])])])).
cnf(c_0_308, plain, (k8_real_1(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_290])])).
cnf(c_0_309, plain, (v1_xreal_0(k8_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_291]), c_0_158])])])).
fof(c_0_310, plain, ![X148, X149, X150]:(~v1_xcmplx_0(X148)|(~v1_xcmplx_0(X149)|(~v1_xcmplx_0(X150)|k7_xcmplx_0(X148,k3_xcmplx_0(X149,X150))=k7_xcmplx_0(k7_xcmplx_0(X148,X149),X150)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t78_xcmplx_1])])])).
cnf(c_0_311, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_292])).
cnf(c_0_312, plain, (m1_subset_1(np__1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_293])).
cnf(c_0_313, plain, (v1_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_294])).
cnf(c_0_314, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_295])).
cnf(c_0_315, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_296])).
cnf(c_0_316, plain, (k1_fdiff_1(k3_rfunct_1(k1_numbers,k1_numbers,k16_sin_cos,k19_sin_cos),X1)=k10_real_1(np__1,k5_square_1(k1_seq_1(k19_sin_cos,X1)))|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_133])).
fof(c_0_317, plain, ![X66]:(~m1_subset_1(X66,k1_numbers)|m1_subset_1(k5_square_1(X66),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_square_1])])).
cnf(c_0_318, plain, (~r1_xxreal_0(k1_fdiff_1(k3_taylor_1,X1),k6_numbers)|~m2_subset_1(X1,k1_numbers,k3_limfunc1(k6_numbers))), inference(evalgc,[status(thm)],[c_0_297])).
cnf(c_0_319, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(np__1,esk4_0))|v1_xboole_0(a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_299]), c_0_183]), c_0_270]), c_0_143])])])).
cnf(c_0_320, negated_conjecture, (r2_hidden(esk4_0,a_0_0_fdiff_8)|~r2_hidden(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_301])])).
cnf(c_0_321, plain, (k5_xcmplx_0(k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_302])).
cnf(c_0_322, plain, (k5_xcmplx_0(X1)=k7_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_303])).
cnf(c_0_323, plain, (v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_304])).
cnf(c_0_324, plain, (r1_xxreal_0(X1,k6_numbers)|r2_hidden(X2,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers)|X2!=X1), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_325, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_248])])).
cnf(c_0_326, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|~r1_tarski(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_233])])).
cnf(c_0_327, negated_conjecture, (~v1_xboole_0(esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_306, c_0_80]), c_0_182]), c_0_183]), c_0_127]), c_0_143])])])).
cnf(c_0_328, negated_conjecture, (k7_xcmplx_0(np__1,k3_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),esk4_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k3_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),esk4_0))|~v1_xreal_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_307, c_0_308])])).
cnf(c_0_329, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_289])])).
cnf(c_0_330, plain, (k7_xcmplx_0(X1,k3_xcmplx_0(X2,X3))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)), inference(split_conjunct,[status(thm)],[c_0_310])).
cnf(c_0_331, plain, (v1_xcmplx_0(np__1)|~v1_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_312])])).
cnf(c_0_332, plain, (v1_membered(X1)|~v3_membered(X1)), inference(evalgc,[status(thm)],[c_0_313])).
cnf(c_0_333, plain, (v3_membered(k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_314, c_0_315]), c_0_158])])])).
cnf(c_0_334, plain, (k1_seq_1(k19_sin_cos,X1)=k6_numbers|k10_real_1(np__1,k5_square_1(k1_seq_1(k19_sin_cos,X1)))=k1_fdiff_1(k3_rfunct_1(k1_numbers,k1_numbers,k16_sin_cos,k19_sin_cos),X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_316])).
cnf(c_0_335, plain, (m1_subset_1(k5_square_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_317])).
cnf(c_0_336, plain, (m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[dt_k19_sin_cos])).
cnf(c_0_337, plain, (~r1_xxreal_0(k1_fdiff_1(k3_taylor_1,X1),k6_numbers)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)), inference(rw,[status(thm)],[c_0_318, c_0_145])).
cnf(c_0_338, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,esk4_0))=k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0))|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_269, c_0_319])])).
cnf(c_0_339, negated_conjecture, (r2_hidden(esk4_0,a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_320, c_0_272])])).
cnf(c_0_340, plain, (k7_xcmplx_0(np__1,k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_321, c_0_322])])).
cnf(c_0_341, plain, (v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_323])).
cnf(c_0_342, plain, (r2_hidden(X2,a_0_0_fdiff_8)|r1_xxreal_0(X1,k6_numbers)|X1!=X2|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_324])).
cnf(c_0_343, negated_conjecture, (v1_xboole_0(X1)|v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|m1_subset_1(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(X2,X1)|~r1_tarski(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_326])])).
cnf(c_0_344, negated_conjecture, (~v1_xboole_0(esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_188]), c_0_182]), c_0_183]), c_0_127]), c_0_143])])])).
cnf(c_0_345, negated_conjecture, (k7_xcmplx_0(np__1,k3_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),esk4_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_328, c_0_329])])).
cnf(c_0_346, plain, (k7_xcmplx_0(X1,k3_xcmplx_0(X2,X3))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_330])).
cnf(c_0_347, plain, (v1_xcmplx_0(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_331, c_0_332]), c_0_333])])])).
cnf(c_0_348, plain, (k10_real_1(np__1,k5_square_1(k1_seq_1(k19_sin_cos,X1)))=k1_fdiff_1(k29_sin_cos,X1)|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~m1_subset_1(X1,k1_numbers)), inference(rw,[status(thm)],[c_0_334, c_0_190])).
cnf(c_0_349, plain, (m1_subset_1(k5_square_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_335])).
cnf(c_0_350, plain, (m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_336])).
cnf(c_0_351, plain, (v1_funct_1(k19_sin_cos)), inference(split_conjunct,[status(thm)],[dt_k19_sin_cos])).
cnf(c_0_352, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_8)|~r1_xxreal_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0)),k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_338])])).
cnf(c_0_353, negated_conjecture, (~v1_xboole_0(a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_339])])).
cnf(c_0_354, plain, (k7_xcmplx_0(np__1,k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_340, c_0_341])).
cnf(c_0_355, negated_conjecture, (v1_xcmplx_0(esk4_0)|~v1_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_261])])).
cnf(c_0_356, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_8)|~r1_xxreal_0(k7_xcmplx_0(np__1,esk4_0),k6_numbers)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_299])])).
cnf(c_0_357, plain, (r1_xxreal_0(X1,k6_numbers)|r2_hidden(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_342])])).
cnf(c_0_358, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|m1_subset_1(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_343, c_0_264]), c_0_344])])).
cnf(c_0_359, negated_conjecture, (m1_subset_1(esk4_0,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_272])])).
cnf(c_0_360, negated_conjecture, (k7_xcmplx_0(k7_xcmplx_0(np__1,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))),esk4_0)!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers)|~v1_xcmplx_0(esk4_0)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_345, c_0_346]), c_0_347])])])).
cnf(c_0_361, plain, (k7_xcmplx_0(np__1,k5_square_1(k1_seq_1(k19_sin_cos,X1)))=k1_fdiff_1(k29_sin_cos,X1)|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,X1)))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_348]), c_0_237])])])).
cnf(c_0_362, plain, (v1_xreal_0(k5_square_1(X1))|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_349]), c_0_158])])])).
cnf(c_0_363, plain, (v3_valued_0(k19_sin_cos)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_350]), c_0_158])])])).
cnf(c_0_364, plain, (v1_funct_1(k19_sin_cos)), inference(evalgc,[status(thm)],[c_0_351])).
cnf(c_0_365, plain, (v1_relat_1(k19_sin_cos)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_350])])).
fof(c_0_366, plain, ![X35, X36, X37, X38]:(((k1_relset_1(k1_numbers,X37)=X36|X37!=k2_fdiff_1(X35,X36)|(~v1_funct_1(X37)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X35,X36)|(~v1_funct_1(X35)|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(~m1_subset_1(X38,k1_numbers)|(~r2_hidden(X38,X36)|k1_seq_1(X37,X38)=k1_fdiff_1(X35,X38))|X37!=k2_fdiff_1(X35,X36)|(~v1_funct_1(X37)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X35,X36)|(~v1_funct_1(X35)|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))))&((m1_subset_1(esk1_3(X35,X36,X37),k1_numbers)|k1_relset_1(k1_numbers,X37)!=X36|X37=k2_fdiff_1(X35,X36)|(~v1_funct_1(X37)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X35,X36)|(~v1_funct_1(X35)|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&((r2_hidden(esk1_3(X35,X36,X37),X36)|k1_relset_1(k1_numbers,X37)!=X36|X37=k2_fdiff_1(X35,X36)|(~v1_funct_1(X37)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X35,X36)|(~v1_funct_1(X35)|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(k1_seq_1(X37,esk1_3(X35,X36,X37))!=k1_fdiff_1(X35,esk1_3(X35,X36,X37))|k1_relset_1(k1_numbers,X37)!=X36|X37=k2_fdiff_1(X35,X36)|(~v1_funct_1(X37)|~m1_subset_1(X37,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))|~r2_fdiff_1(X35,X36)|(~v1_funct_1(X35)|~m1_subset_1(X35,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d7_fdiff_1])])])])])).
cnf(c_0_367, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0)),k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)), inference(sr,[status(thm)],[c_0_352, c_0_353])).
cnf(c_0_368, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,X1))=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_354, c_0_322])])).
cnf(c_0_369, negated_conjecture, (v1_xcmplx_0(esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_355, c_0_332]), c_0_158])])])).
cnf(c_0_370, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_8)|r2_hidden(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_357])])).
cnf(c_0_371, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers)|v1_xboole_0(a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_281, c_0_299]), c_0_183]), c_0_270]), c_0_143])])])).
cnf(c_0_372, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|m1_subset_1(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_358, c_0_359])])).
cnf(c_0_373, negated_conjecture, (k7_xcmplx_0(k7_xcmplx_0(np__1,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))),esk4_0)!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_360, c_0_270])])])).
cnf(c_0_374, plain, (k7_xcmplx_0(np__1,k5_square_1(k1_seq_1(k19_sin_cos,X1)))=k1_fdiff_1(k29_sin_cos,X1)|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_362])])).
cnf(c_0_375, plain, (m1_subset_1(k1_seq_1(k19_sin_cos,X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_363]), c_0_364]), c_0_365])])])).
cnf(c_0_376, plain, (k1_seq_1(X3,X1)=k1_fdiff_1(X4,X1)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,X2)|X3!=k2_fdiff_1(X4,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X4,X2)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_366])).
fof(c_0_377, plain, ![X56, X57]:((v1_funct_1(k2_fdiff_1(X56,X57))|(~v1_funct_1(X56)|~m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))&(m1_subset_1(k2_fdiff_1(X56,X57),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|(~v1_funct_1(X56)|~m1_subset_1(X56,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_fdiff_1])])])).
cnf(c_0_378, plain, (k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)=k8_real_1(k1_fdiff_1(X2,k1_seq_1(X1,X3)),k1_fdiff_1(X1,X3))|~r1_fdiff_1(X1,X3)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_106])).
cnf(c_0_379, negated_conjecture, (~r1_xxreal_0(esk4_0,k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_368]), c_0_369])])])).
cnf(c_0_380, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_8)|r2_hidden(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8)), inference(csr,[status(thm)],[c_0_370, c_0_371])).
cnf(c_0_381, negated_conjecture, (m1_subset_1(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_372])])).
cnf(c_0_382, negated_conjecture, (k7_xcmplx_0(k7_xcmplx_0(np__1,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))),esk4_0)!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_373, c_0_369])])).
cnf(c_0_383, plain, (k7_xcmplx_0(np__1,k5_square_1(k1_seq_1(k19_sin_cos,X1)))=k1_fdiff_1(k29_sin_cos,X1)|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~m1_subset_1(X1,k1_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_374, c_0_375])])).
cnf(c_0_384, plain, (k1_fdiff_1(X4,X1)=k1_seq_1(X3,X1)|X3!=k2_fdiff_1(X4,X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~r2_fdiff_1(X4,X2)|~r2_hidden(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_376])).
cnf(c_0_385, plain, (m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_377])).
cnf(c_0_386, plain, (v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(split_conjunct,[status(thm)],[c_0_377])).
cnf(c_0_387, plain, (k8_real_1(k1_fdiff_1(X2,k1_seq_1(X1,X3)),k1_fdiff_1(X1,X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X3)|~m1_subset_1(X3,k1_numbers)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_378])).
cnf(c_0_388, negated_conjecture, (~r1_xxreal_0(esk4_0,k6_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_379, c_0_248]), c_0_259])]), c_0_353]), c_0_249])])).
cnf(c_0_389, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|v1_xboole_0(a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_380])])).
cnf(c_0_390, negated_conjecture, (m1_subset_1(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_381, c_0_272])])).
cnf(c_0_391, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))=k6_numbers|k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)!=k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_383]), c_0_200])])])).
cnf(c_0_392, plain, (v1_xcmplx_0(k5_square_1(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_349])])).
cnf(c_0_393, plain, (k1_seq_1(k2_fdiff_1(X1,X3),X2)=k1_fdiff_1(X1,X2)|~r2_hidden(X2,X3)|~r2_fdiff_1(X1,X3)|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X1,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_384])])).
cnf(c_0_394, plain, (m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_385])).
cnf(c_0_395, plain, (v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[c_0_386])).
fof(c_0_396, plain, ![X40, X41]:(~v1_xcmplx_0(X40)|(~v1_xcmplx_0(X41)|k7_xcmplx_0(X40,X41)=k3_xcmplx_0(X40,k5_xcmplx_0(X41)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d9_xcmplx_0])])])).
cnf(c_0_397, plain, (k3_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_fdiff_1(X2,X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3)|~v1_xreal_0(k1_fdiff_1(X2,X3))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_numbers)|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_387])])).
cnf(c_0_398, negated_conjecture, (k1_fdiff_1(k3_taylor_1,esk4_0)=k7_xcmplx_0(np__1,esk4_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_299, c_0_353])])).
cnf(c_0_399, negated_conjecture, (r1_fdiff_1(k3_taylor_1,esk4_0)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_388, c_0_140]), c_0_261])])])).
cnf(c_0_400, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8)), inference(sr,[status(thm)],[c_0_389, c_0_353])).
cnf(c_0_401, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_hidden(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_390])])).
cnf(c_0_402, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))=k6_numbers|k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)!=k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers)|~v1_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_391, c_0_392]), c_0_375])])])).
cnf(c_0_403, plain, (k1_seq_1(k2_fdiff_1(X1,X2),X3)=k1_fdiff_1(X1,X3)|~r2_hidden(X3,X2)|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_393, c_0_394]), c_0_395])).
cnf(c_0_404, plain, (k7_xcmplx_0(X1,X2)=k3_xcmplx_0(X1,k5_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_396])).
cnf(c_0_405, negated_conjecture, (k3_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k7_xcmplx_0(np__1,esk4_0))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~r1_fdiff_1(k3_taylor_1,esk4_0)|~v1_funct_1(X1)|~v1_xreal_0(k7_xcmplx_0(np__1,esk4_0))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_397, c_0_398]), c_0_183]), c_0_143]), c_0_261])])])).
cnf(c_0_406, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(np__1,esk4_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[c_0_319, c_0_353])])).
cnf(c_0_407, negated_conjecture, (r1_fdiff_1(k3_taylor_1,esk4_0)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_399, c_0_400])])).
cnf(c_0_408, negated_conjecture, (r2_hidden(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_401])])).
cnf(c_0_409, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))=k6_numbers|k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_membered(k1_numbers)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_403]), c_0_272]), c_0_261])])])).
cnf(c_0_410, plain, (k3_xcmplx_0(X1,k5_xcmplx_0(X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_404])).
cnf(c_0_411, negated_conjecture, (k3_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k7_xcmplx_0(np__1,esk4_0))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~r1_fdiff_1(k3_taylor_1,esk4_0)|~v1_funct_1(X1)|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_405, c_0_406])])).
cnf(c_0_412, negated_conjecture, (r1_fdiff_1(k3_taylor_1,esk4_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_407, c_0_248]), c_0_259]), c_0_283])]), c_0_353]), c_0_249])])).
cnf(c_0_413, negated_conjecture, (r2_hidden(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_408, c_0_272])])).
cnf(c_0_414, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))=k6_numbers|k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk4_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~v1_membered(k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers)), inference(csr,[status(thm)],[c_0_409, c_0_232])).
cnf(c_0_415, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_410, c_0_322])])).
cnf(c_0_416, negated_conjecture, (k3_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k7_xcmplx_0(np__1,esk4_0))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_411, c_0_412])])])).
cnf(c_0_417, negated_conjecture, (r2_hidden(k1_seq_1(k3_taylor_1,esk4_0),k1_relset_1(k1_numbers,k29_sin_cos))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_413]), c_0_181]), c_0_182]), c_0_183]), c_0_184]), c_0_127]), c_0_143])])])).
cnf(c_0_418, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))=k6_numbers|k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk4_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~v1_membered(k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_414, c_0_349]), c_0_375])])])).
cnf(c_0_419, negated_conjecture, (k7_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~v1_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)))|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_415, c_0_416]), c_0_369])])])).
cnf(c_0_420, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))!=k6_numbers), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_417]), c_0_200])])])).
cnf(c_0_421, negated_conjecture, (~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~v1_membered(k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))|~v1_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))|~m1_subset_1(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_418, c_0_419]), c_0_182]), c_0_127])]), c_0_420])])).
cnf(c_0_422, negated_conjecture, (~r1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~v1_membered(k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers)), inference(csr,[status(thm)],[c_0_421, c_0_311])).
cnf(c_0_423, plain, (v1_xreal_0(k1_seq_1(k3_taylor_1,X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_200]), c_0_158])])])).
cnf(c_0_424, negated_conjecture, (~r1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~v1_membered(k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_422, c_0_281]), c_0_182]), c_0_423]), c_0_127])])])).
cnf(c_0_425, negated_conjecture, (~r1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~v1_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_424, c_0_80]), c_0_182]), c_0_183]), c_0_127]), c_0_143])])])).
cnf(c_0_426, negated_conjecture, (~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~v1_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_425, c_0_202]), c_0_200])]), c_0_420])])).
cnf(c_0_427, negated_conjecture, (~v1_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_426, c_0_188]), c_0_182]), c_0_183]), c_0_127]), c_0_143])])])).
cnf(c_0_428, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_427, c_0_332]), c_0_158])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 429
# Proof object clause steps            : 326
# Proof object formula steps           : 103
# Proof object conjectures             : 113
# Proof object clause conjectures      : 110
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 79
# Proof object initial formulas used   : 55
# Proof object generating inferences   : 135
# Proof object simplifying inferences  : 271
# Training examples: 218 positive, 2918 negative
# Training: Positive examples begin
cnf(c_0_145, plain, (k3_limfunc1(k6_numbers)=a_0_0_fdiff_8)).# trainpos
cnf(c_0_158, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_182, plain, (v1_funct_1(k29_sin_cos))).# trainpos
cnf(c_0_364, plain, (v1_funct_1(k19_sin_cos))).# trainpos
cnf(c_0_183, plain, (v1_funct_1(k3_taylor_1))).# trainpos
cnf(c_0_134, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_249, plain, (~v1_xboole_0(k1_numbers))).# trainpos
cnf(c_0_167, negated_conjecture, (v3_rcomp_1(esk3_0))).# trainpos
cnf(c_0_237, plain, (m1_subset_1(np__1,k1_numbers))).# trainpos
cnf(c_0_312, plain, (m1_subset_1(np__1,k5_numbers))).# trainpos
cnf(c_0_264, plain, (r1_tarski(X1,X1))).# trainpos
cnf(c_0_144, plain, (k3_limfunc1(k6_numbers)=k1_relset_1(k1_numbers,k3_taylor_1))).# trainpos
cnf(c_0_315, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_168, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_99, plain, (v1_relat_1(k3_relat_1(X1,X2)))).# trainpos
cnf(c_0_332, plain, (v1_membered(X1)|~v3_membered(X1))).# trainpos
cnf(c_0_84, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_169, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_111, plain, (esk2_1(X1)=X1|~r2_hidden(X1,a_0_0_fdiff_8))).# trainpos
cnf(c_0_221, negated_conjecture, (m1_subset_1(esk4_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0))).# trainpos
cnf(c_0_234, negated_conjecture, (r2_hidden(esk4_0,esk3_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0))).# trainpos
cnf(c_0_127, plain, (m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_190, plain, (k3_rfunct_1(k1_numbers,k1_numbers,k16_sin_cos,k19_sin_cos)=k29_sin_cos)).# trainpos
cnf(c_0_341, plain, (v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_321, plain, (k5_xcmplx_0(k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_350, plain, (m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_143, plain, (m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_322, plain, (k5_xcmplx_0(X1)=k7_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_254, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_110, plain, (~r2_hidden(X1,a_0_0_fdiff_8)|~r1_xxreal_0(esk2_1(X1),k6_numbers))).# trainpos
cnf(c_0_288, negated_conjecture, (k10_real_1(np__1,k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0))).# trainpos
cnf(c_0_85, negated_conjecture, (r1_tarski(esk3_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))).# trainpos
cnf(c_0_318, plain, (~r1_xxreal_0(k1_fdiff_1(k3_taylor_1,X1),k6_numbers)|~m2_subset_1(X1,k1_numbers,k3_limfunc1(k6_numbers)))).# trainpos
cnf(c_0_198, plain, (k1_seq_1(k19_sin_cos,X1)!=k6_numbers|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k29_sin_cos)))).# trainpos
cnf(c_0_349, plain, (m1_subset_1(k5_square_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_141, plain, (m1_subset_1(esk2_1(X1),k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_8))).# trainpos
cnf(c_0_311, plain, (v1_xcmplx_0(X2)|~v1_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_260, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_121, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_124, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_314, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_79, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_291, plain, (m1_subset_1(k8_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_289, plain, (k8_real_1(X1,X2)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_225, plain, (k10_real_1(np__1,X1)=k1_fdiff_1(k3_taylor_1,X1)|~m2_subset_1(X1,k1_numbers,k3_limfunc1(k6_numbers)))).# trainpos
cnf(c_0_126, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_157, plain, (v3_valued_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).# trainpos
cnf(c_0_140, plain, (r1_xxreal_0(X1,k6_numbers)|r1_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_189, plain, (k1_seq_1(k19_sin_cos,X1)=k6_numbers|r1_fdiff_1(k3_rfunct_1(k1_numbers,k1_numbers,k16_sin_cos,k19_sin_cos),X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_395, plain, (v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_97, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).# trainpos
cnf(c_0_176, plain, (r1_tarski(X1,X3)|~r1_tarski(X2,X3)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_410, plain, (k3_xcmplx_0(X1,k5_xcmplx_0(X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_290, plain, (k8_real_1(X1,X2)=k8_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_236, plain, (k10_real_1(X1,X2)=k7_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_101, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_342, plain, (r2_hidden(X2,a_0_0_fdiff_8)|r1_xxreal_0(X1,k6_numbers)|X1!=X2|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_250, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).# trainpos
cnf(c_0_185, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1))).# trainpos
cnf(c_0_160, plain, (r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v3_rcomp_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r1_fdiff_1(X1,esk6_2(X2,X1))|~r1_tarski(X2,k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_75, plain, (k1_funct_1(X1,X2)=k1_seq_1(X1,X2)|~v1_relat_1(X1)|~v3_valued_0(X1)|~v1_funct_1(X1))).# trainpos
cnf(c_0_281, plain, (m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_119, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(X1,X2)),k9_xtuple_0(X1))|~v1_relat_1(X2)|~v1_relat_1(X1))).# trainpos
cnf(c_0_394, plain, (m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_334, plain, (k1_seq_1(k19_sin_cos,X1)=k6_numbers|k10_real_1(np__1,k5_square_1(k1_seq_1(k19_sin_cos,X1)))=k1_fdiff_1(k3_rfunct_1(k1_numbers,k1_numbers,k16_sin_cos,k19_sin_cos),X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_114, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,X3))))).# trainpos
cnf(c_0_203, plain, (r1_fdiff_1(X3,X1)|~v1_funct_1(X3)|~v3_rcomp_1(X2)|~r2_fdiff_1(X3,X2)|~r2_hidden(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_346, plain, (k7_xcmplx_0(X1,k3_xcmplx_0(X2,X3))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_248, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_305, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_74, plain, (r2_hidden(k1_funct_1(X1,X2),k9_xtuple_0(X3))|~v1_relat_1(X3)|~v1_relat_1(X1)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_hidden(X2,k9_xtuple_0(k3_relat_1(X1,X3))))).# trainpos
cnf(c_0_188, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_95, plain, (k1_partfun1(X2,X3,X5,X6,X1,X4)=k3_relat_1(X1,X4)|~v1_funct_1(X4)|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).# trainpos
cnf(c_0_187, plain, (r2_fdiff_1(X2,X1)|m1_subset_1(esk6_2(X1,X2),k1_numbers)|~v1_funct_1(X2)|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_135, plain, (r2_fdiff_1(X2,X1)|r2_hidden(esk6_2(X1,X2),X1)|~v1_funct_1(X2)|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_161, plain, (r1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X3)|~m1_subset_1(X3,k1_numbers)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_384, plain, (k1_fdiff_1(X4,X1)=k1_seq_1(X3,X1)|X3!=k2_fdiff_1(X4,X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~r2_fdiff_1(X4,X2)|~r2_hidden(X1,X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_80, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_387, plain, (k8_real_1(k1_fdiff_1(X2,k1_seq_1(X1,X3)),k1_fdiff_1(X1,X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_fdiff_1(X1,X3)|~m1_subset_1(X3,k1_numbers)|~r1_fdiff_1(X2,k1_seq_1(X1,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_307, negated_conjecture, (k7_xcmplx_0(np__1,k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))))).# trainpos
cnf(c_0_362, plain, (v1_xreal_0(k5_square_1(X1))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_186, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),X3)|~v3_rcomp_1(X3)|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~r1_fdiff_1(X2,k1_seq_1(X1,esk6_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2))))|~r1_fdiff_1(X1,esk6_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk6_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X2)),k1_numbers))).# trainpos
cnf(c_0_247, plain, (k1_fdiff_1(k3_taylor_1,X1)=k7_xcmplx_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_233, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r1_tarski(X1,esk3_0))).# trainpos
cnf(c_0_222, plain, (k1_partfun1(X1,X2,X3,X4,X5,X6)=k1_partfun1(X7,X8,X9,X10,X5,X6)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).# trainpos
cnf(c_0_357, plain, (r1_xxreal_0(X1,k6_numbers)|r2_hidden(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_175, plain, (v1_relat_1(k3_taylor_1))).# trainpos
cnf(c_0_365, plain, (v1_relat_1(k19_sin_cos))).# trainpos
cnf(c_0_154, plain, (v1_relat_1(k29_sin_cos))).# trainpos
cnf(c_0_333, plain, (v3_membered(k5_numbers))).# trainpos
cnf(c_0_397, plain, (k3_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_fdiff_1(X2,X3))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3)|~v1_xreal_0(k1_fdiff_1(X2,X3))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_numbers)|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).# trainpos
cnf(c_0_393, plain, (k1_seq_1(k2_fdiff_1(X1,X3),X2)=k1_fdiff_1(X1,X2)|~r2_hidden(X2,X3)|~r2_fdiff_1(X1,X3)|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X1,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_298, plain, (v1_xreal_0(k1_fdiff_1(X1,X2))|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_184, plain, (v3_valued_0(k3_taylor_1))).# trainpos
cnf(c_0_171, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_8))).# trainpos
cnf(c_0_363, plain, (v3_valued_0(k19_sin_cos))).# trainpos
cnf(c_0_331, plain, (v1_xcmplx_0(np__1)|~v1_membered(k5_numbers))).# trainpos
cnf(c_0_361, plain, (k7_xcmplx_0(np__1,k5_square_1(k1_seq_1(k19_sin_cos,X1)))=k1_fdiff_1(k29_sin_cos,X1)|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,X1)))|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_309, plain, (v1_xreal_0(k8_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_166, plain, (m1_subset_1(esk6_2(X1,X2),X1)|r2_fdiff_1(X2,X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_308, plain, (k8_real_1(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_102, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))))).# trainpos
cnf(c_0_173, plain, (v4_relat_1(k3_taylor_1,k1_numbers))).# trainpos
cnf(c_0_273, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_147, plain, (r1_tarski(k9_xtuple_0(k1_partfun1(X3,X4,X5,X6,X1,X2)),k9_xtuple_0(X1))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).# trainpos
cnf(c_0_94, plain, (r2_hidden(k1_seq_1(X1,X2),k9_xtuple_0(X3))|~r2_hidden(X2,k9_xtuple_0(k3_relat_1(X1,X3)))|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X1)|~v3_valued_0(X1))).# trainpos
cnf(c_0_153, plain, (v4_relat_1(k29_sin_cos,k1_numbers))).# trainpos
cnf(c_0_340, plain, (k7_xcmplx_0(np__1,k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1)))).# trainpos
cnf(c_0_193, plain, (r2_fdiff_1(X2,X1)|~v1_xboole_0(X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_139, plain, (~r1_xxreal_0(X1,k6_numbers)|~r2_hidden(X1,a_0_0_fdiff_8))).# trainpos
cnf(c_0_392, plain, (v1_xcmplx_0(k5_square_1(X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_142, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(X1,k9_xtuple_0(k1_partfun1(X4,X5,X6,X7,X2,X3)))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).# trainpos
cnf(c_0_98, plain, (v4_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6),X1)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_325, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,X2))).# trainpos
cnf(c_0_123, plain, (v1_relat_1(k1_partfun1(X3,X4,X5,X6,X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).# trainpos
cnf(c_0_415, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_347, plain, (v1_xcmplx_0(np__1))).# trainpos
cnf(c_0_196, plain, (k9_xtuple_0(k3_taylor_1)=a_0_0_fdiff_8)).# trainpos
cnf(c_0_210, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k29_sin_cos),X2)|k1_seq_1(k19_sin_cos,k1_seq_1(X1,esk6_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k29_sin_cos))))=k6_numbers|~v3_rcomp_1(X2)|~r1_fdiff_1(X1,esk6_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k29_sin_cos)))|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k29_sin_cos)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(k1_seq_1(X1,esk6_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k29_sin_cos))),k1_numbers))).# trainpos
cnf(c_0_122, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,X5,X6))=k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_181, plain, (k9_xtuple_0(k29_sin_cos)=k1_relset_1(k1_numbers,k29_sin_cos))).# trainpos
cnf(c_0_259, plain, (m1_subset_1(a_0_0_fdiff_8,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_170, plain, (r1_fdiff_1(k3_taylor_1,X1)|~r2_hidden(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_375, plain, (m1_subset_1(k1_seq_1(k19_sin_cos,X1),k1_numbers))).# trainpos
cnf(c_0_200, plain, (m1_subset_1(k1_seq_1(k3_taylor_1,X1),k1_numbers))).# trainpos
cnf(c_0_245, negated_conjecture, (r2_hidden(esk4_0,esk3_0)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),esk3_0)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_231, negated_conjecture, (m1_subset_1(esk4_0,k1_numbers)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),esk3_0)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).# trainpos
cnf(c_0_328, negated_conjecture, (k7_xcmplx_0(np__1,k3_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),esk4_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k3_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),esk4_0))|~v1_xreal_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers))).# trainpos
cnf(c_0_326, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|~r1_tarski(X1,esk3_0))).# trainpos
cnf(c_0_329, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_205, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xboole_0(esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_374, plain, (k7_xcmplx_0(np__1,k5_square_1(k1_seq_1(k19_sin_cos,X1)))=k1_fdiff_1(k29_sin_cos,X1)|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,X1),k1_numbers))).# trainpos
cnf(c_0_152, negated_conjecture, (~v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,esk3_0))).# trainpos
cnf(c_0_258, plain, (k1_fdiff_1(k3_taylor_1,X1)=k7_xcmplx_0(np__1,X1)|v1_xboole_0(a_0_0_fdiff_8)|~v1_xreal_0(X1)|~m1_subset_1(a_0_0_fdiff_8,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,a_0_0_fdiff_8))).# trainpos
cnf(c_0_192, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|m1_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_197, plain, (r1_tarski(X1,k9_xtuple_0(X2))|~r1_tarski(X1,k9_xtuple_0(k1_partfun1(X3,X4,X5,X6,X2,X7)))|~v1_funct_1(X7)|~v1_funct_1(X2)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).# trainpos
cnf(c_0_125, negated_conjecture, (m1_subset_1(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,esk3_0))).# trainpos
cnf(c_0_368, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,X1))=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_120, plain, (r2_hidden(k1_seq_1(X1,X2),k9_xtuple_0(X3))|~r2_hidden(X2,k9_xtuple_0(k1_partfun1(X4,X5,X6,X7,X1,X3)))|~v1_funct_1(X3)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X3)|~v1_relat_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).# trainpos
cnf(c_0_206, plain, (r2_fdiff_1(k3_taylor_1,X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r2_hidden(esk6_2(X1,k3_taylor_1),a_0_0_fdiff_8))).# trainpos
cnf(c_0_204, negated_conjecture, (v1_xboole_0(esk3_0)|r2_hidden(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),esk3_0)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_423, plain, (v1_xreal_0(k1_seq_1(k3_taylor_1,X1)))).# trainpos
cnf(c_0_208, plain, (r1_tarski(X1,k9_xtuple_0(X2))|~r1_tarski(X1,k1_relset_1(X3,k1_partfun1(X3,X4,X5,X6,X2,X7)))|~v1_funct_1(X7)|~v1_funct_1(X2)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).# trainpos
cnf(c_0_179, plain, (r2_hidden(k1_seq_1(X1,X2),k9_xtuple_0(X3))|~r2_hidden(X2,k1_relset_1(X4,k1_partfun1(X4,X5,X6,X7,X1,X3)))|~v1_funct_1(X3)|~v1_funct_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).# trainpos
cnf(c_0_151, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_hidden(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,esk3_0))).# trainpos
cnf(c_0_345, negated_conjecture, (k7_xcmplx_0(np__1,k3_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),esk4_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers))).# trainpos
cnf(c_0_287, negated_conjecture, (r2_hidden(esk4_0,esk3_0)|~v1_xboole_0(esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_195, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(X1,k1_relset_1(X3,k1_partfun1(X3,X4,X5,X6,X2,X7)))|~v1_funct_1(X7)|~v1_funct_1(X2)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).# trainpos
cnf(c_0_327, negated_conjecture, (~v1_xboole_0(esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))).# trainpos
cnf(c_0_214, negated_conjecture, (r1_tarski(esk3_0,a_0_0_fdiff_8))).# trainpos
cnf(c_0_243, negated_conjecture, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,esk3_0))).# trainpos
cnf(c_0_207, negated_conjecture, (r2_hidden(X1,a_0_0_fdiff_8)|~r2_hidden(X1,esk3_0))).# trainpos
cnf(c_0_199, negated_conjecture, (r2_hidden(k1_seq_1(k3_taylor_1,X1),k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X1,esk3_0))).# trainpos
cnf(c_0_217, negated_conjecture, (r1_fdiff_1(X1,esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_209, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,X1))!=k6_numbers|~r2_hidden(X1,esk3_0))).# trainpos
cnf(c_0_213, negated_conjecture, (r2_fdiff_1(k3_taylor_1,X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r2_hidden(esk6_2(X1,k3_taylor_1),esk3_0))).# trainpos
cnf(c_0_344, negated_conjecture, (~v1_xboole_0(esk3_0))).# trainpos
cnf(c_0_300, negated_conjecture, (~v1_xboole_0(a_0_0_fdiff_8)|~r2_hidden(X1,esk3_0))).# trainpos
cnf(c_0_253, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(a_0_0_fdiff_8))|~r1_tarski(X1,esk3_0))).# trainpos
cnf(c_0_218, negated_conjecture, (r2_fdiff_1(k3_taylor_1,esk3_0))).# trainpos
cnf(c_0_216, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~r2_hidden(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),esk3_0)|~v3_rcomp_1(X1)|~r1_fdiff_1(k3_taylor_1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_262, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).# trainpos
cnf(c_0_360, negated_conjecture, (k7_xcmplx_0(k7_xcmplx_0(np__1,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))),esk4_0)!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers)|~v1_xcmplx_0(esk4_0)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))).# trainpos
cnf(c_0_223, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~r2_hidden(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_271, negated_conjecture, (m1_subset_1(X1,a_0_0_fdiff_8)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(esk3_0)))).# trainpos
cnf(c_0_255, negated_conjecture, (r2_hidden(esk4_0,esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_242, negated_conjecture, (m1_subset_1(esk4_0,k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_263, negated_conjecture, (r2_hidden(esk4_0,esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))).# trainpos
cnf(c_0_252, negated_conjecture, (m1_subset_1(esk4_0,k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))).# trainpos
cnf(c_0_343, negated_conjecture, (v1_xboole_0(X1)|v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|m1_subset_1(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(X2,X1)|~r1_tarski(X1,esk3_0))).# trainpos
cnf(c_0_272, negated_conjecture, (r2_hidden(esk4_0,esk3_0))).# trainpos
cnf(c_0_261, negated_conjecture, (m1_subset_1(esk4_0,k1_numbers))).# trainpos
cnf(c_0_358, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|m1_subset_1(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(X1,esk3_0))).# trainpos
cnf(c_0_270, negated_conjecture, (v1_xreal_0(esk4_0))).# trainpos
cnf(c_0_355, negated_conjecture, (v1_xcmplx_0(esk4_0)|~v1_membered(k1_numbers))).# trainpos
cnf(c_0_283, negated_conjecture, (m1_subset_1(esk4_0,a_0_0_fdiff_8))).# trainpos
cnf(c_0_359, negated_conjecture, (m1_subset_1(esk4_0,esk3_0))).# trainpos
cnf(c_0_369, negated_conjecture, (v1_xcmplx_0(esk4_0))).# trainpos
cnf(c_0_372, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|m1_subset_1(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))).# trainpos
cnf(c_0_301, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_8)|r2_hidden(esk4_0,a_0_0_fdiff_8))).# trainpos
cnf(c_0_282, negated_conjecture, (k1_fdiff_1(k3_taylor_1,esk4_0)=k7_xcmplx_0(np__1,esk4_0)|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(esk4_0,a_0_0_fdiff_8))).# trainpos
cnf(c_0_373, negated_conjecture, (k7_xcmplx_0(k7_xcmplx_0(np__1,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))),esk4_0)!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers))).# trainpos
cnf(c_0_320, negated_conjecture, (r2_hidden(esk4_0,a_0_0_fdiff_8)|~r2_hidden(X1,esk3_0))).# trainpos
cnf(c_0_356, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_8)|~r1_xxreal_0(k7_xcmplx_0(np__1,esk4_0),k6_numbers)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).# trainpos
cnf(c_0_319, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(np__1,esk4_0))|v1_xboole_0(a_0_0_fdiff_8))).# trainpos
cnf(c_0_371, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers)|v1_xboole_0(a_0_0_fdiff_8))).# trainpos
cnf(c_0_391, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))=k6_numbers|k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)!=k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers))).# trainpos
cnf(c_0_381, negated_conjecture, (m1_subset_1(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,esk3_0))).# trainpos
cnf(c_0_339, negated_conjecture, (r2_hidden(esk4_0,a_0_0_fdiff_8))).# trainpos
cnf(c_0_390, negated_conjecture, (m1_subset_1(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))).# trainpos
cnf(c_0_402, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))=k6_numbers|k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)!=k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers)|~v1_membered(k1_numbers))).# trainpos
cnf(c_0_338, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,esk4_0))=k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0))|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).# trainpos
cnf(c_0_370, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_8)|r2_hidden(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers))).# trainpos
cnf(c_0_353, negated_conjecture, (~v1_xboole_0(a_0_0_fdiff_8))).# trainpos
cnf(c_0_352, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_8)|~r1_xxreal_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0)),k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).# trainpos
cnf(c_0_401, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_hidden(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))).# trainpos
cnf(c_0_409, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))=k6_numbers|k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_membered(k1_numbers)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_389, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|v1_xboole_0(a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).# trainpos
cnf(c_0_398, negated_conjecture, (k1_fdiff_1(k3_taylor_1,esk4_0)=k7_xcmplx_0(np__1,esk4_0))).# trainpos
cnf(c_0_406, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(np__1,esk4_0)))).# trainpos
cnf(c_0_379, negated_conjecture, (~r1_xxreal_0(esk4_0,k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).# trainpos
cnf(c_0_418, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))=k6_numbers|k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk4_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~v1_membered(k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_408, negated_conjecture, (r2_hidden(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,esk3_0))).# trainpos
cnf(c_0_413, negated_conjecture, (r2_hidden(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))).# trainpos
cnf(c_0_388, negated_conjecture, (~r1_xxreal_0(esk4_0,k6_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).# trainpos
cnf(c_0_405, negated_conjecture, (k3_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k7_xcmplx_0(np__1,esk4_0))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~r1_fdiff_1(k3_taylor_1,esk4_0)|~v1_funct_1(X1)|~v1_xreal_0(k7_xcmplx_0(np__1,esk4_0))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_417, negated_conjecture, (r2_hidden(k1_seq_1(k3_taylor_1,esk4_0),k1_relset_1(k1_numbers,k29_sin_cos)))).# trainpos
cnf(c_0_399, negated_conjecture, (r1_fdiff_1(k3_taylor_1,esk4_0)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).# trainpos
cnf(c_0_420, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))!=k6_numbers)).# trainpos
cnf(c_0_407, negated_conjecture, (r1_fdiff_1(k3_taylor_1,esk4_0)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).# trainpos
cnf(c_0_412, negated_conjecture, (r1_fdiff_1(k3_taylor_1,esk4_0))).# trainpos
cnf(c_0_416, negated_conjecture, (k3_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k7_xcmplx_0(np__1,esk4_0))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_419, negated_conjecture, (k7_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~v1_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)))|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_421, negated_conjecture, (~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~v1_membered(k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))|~v1_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))|~m1_subset_1(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers))).# trainpos
cnf(c_0_424, negated_conjecture, (~r1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~v1_membered(k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).# trainpos
cnf(c_0_425, negated_conjecture, (~r1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~v1_membered(k1_numbers))).# trainpos
cnf(c_0_426, negated_conjecture, (~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~v1_membered(k1_numbers))).# trainpos
cnf(c_0_427, negated_conjecture, (~v1_membered(k1_numbers))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_76, plain, (k1_xboole_0=k6_numbers)).#trainneg
cnf(i_0_74, plain, (k4_ordinal1=k5_numbers)).#trainneg
cnf(i_0_61, plain, (v6_membered(k4_ordinal1))).#trainneg
cnf(i_0_37, plain, (v1_funct_1(k24_sin_cos))).#trainneg
cnf(i_0_56, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_85, plain, (v1_xboole_0(np__0))).#trainneg
cnf(i_0_92, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_62, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_63, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_112, plain, (v2_funct_1(k3_taylor_1))).#trainneg
cnf(i_0_86, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_87, plain, (m1_subset_1(np__0,k5_numbers))).#trainneg
cnf(i_0_84, plain, (k4_xcmplx_0(k4_xcmplx_0(np__1))=np__1)).#trainneg
cnf(i_0_110, plain, (k2_relset_1(k1_numbers,k3_taylor_1)=k1_numbers)).#trainneg
cnf(i_0_128, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_109, plain, (r2_fdiff_1(k3_taylor_1,k3_limfunc1(k6_numbers)))).#trainneg
cnf(i_0_28, plain, (v1_funct_2(k19_sin_cos,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_36, plain, (v1_funct_2(k24_sin_cos,k1_numbers,k1_numbers))).#trainneg
cnf(i_0_8, plain, (v1_int_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_91, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_52, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_88, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_48, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_58, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_57, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_127, plain, (k7_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_35, plain, (m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_14, plain, (v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_18, plain, (k3_square_1(X1)=k3_xcmplx_0(X1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_75, plain, (k3_square_1(X1)=k5_square_1(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11, plain, (v3_valued_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,k1_numbers))).#trainneg
cnf(i_0_113, plain, (r2_relset_1(k1_numbers,k1_numbers,k3_taylor_1,k2_partfun2(k1_numbers,k1_numbers,k24_sin_cos)))).#trainneg
cnf(i_0_42, plain, (m1_subset_1(k3_limfunc1(X1),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_120, plain, (r2_hidden(esk5_2(X1,X2),X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_16, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_119, plain, (~r2_hidden(X3,X2)|~r2_hidden(X1,X2)|~r2_hidden(X1,esk5_2(X3,X2)))).#trainneg
cnf(i_0_1, plain, (v1_valued_0(X2)|~v1_relat_1(X1)|~v1_valued_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_9, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_26, plain, (m1_subset_1(k10_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_60, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_relat_1(X1)|~v1_valued_0(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_79, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_135, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~v3_rcomp_1(X1)|~r2_fdiff_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_94, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X2,X3))=k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_83, plain, (r2_relset_1(X2,X3,X1,X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_80, plain, (r2_relset_1(X3,X4,X1,X2)|X1!=X2|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_81, plain, (X3=X4|~r2_relset_1(X1,X2,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_45, plain, (v1_funct_1(k3_rfunct_1(X1,X2,X3,X4))|~v3_membered(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_95, plain, (r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,X3)))|~v1_relat_1(X3)|~v1_relat_1(X2)|~v1_funct_1(X3)|~v1_funct_1(X2)|~r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(k1_funct_1(X2,X1),k9_xtuple_0(X3)))).#trainneg
cnf(i_0_24, plain, (k1_relset_1(k1_numbers,X1)=X2|X1!=k2_fdiff_1(X3,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X3,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_44, plain, (m1_subset_1(k3_rfunct_1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(X1,k1_numbers)))|~v3_membered(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_22, plain, (X3=k2_fdiff_1(X1,X2)|m1_subset_1(esk1_3(X1,X2,X3),k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_20, plain, (X1=k2_fdiff_1(X2,X3)|k1_relset_1(k1_numbers,X1)!=X3|k1_fdiff_1(X2,esk1_3(X2,X3,X1))!=k1_seq_1(X1,esk1_3(X2,X3,X1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~r2_fdiff_1(X2,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_21, plain, (X3=k2_fdiff_1(X1,X2)|r2_hidden(esk1_3(X1,X2,X3),X2)|k1_relset_1(k1_numbers,X3)!=X2|~v1_funct_1(X3)|~v1_funct_1(X1)|~r2_fdiff_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_399, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~m2_subset_1(k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))),k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_395, negated_conjecture, (k10_real_1(np__1,k3_xcmplx_0(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~m1_subset_1(esk4_0,k1_numbers))).#trainneg
cnf(i_0_453, plain, (r2_fdiff_1(k29_sin_cos,X1)|k1_seq_1(k19_sin_cos,esk6_2(X1,k29_sin_cos))=k6_numbers|~v3_rcomp_1(X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(esk6_2(X1,k29_sin_cos),k1_numbers))).#trainneg
cnf(i_0_529, plain, (r1_fdiff_1(X1,esk6_2(X2,X3))|r2_fdiff_1(X3,X2)|~v3_rcomp_1(X2)|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r1_tarski(X2,k1_relset_1(k1_numbers,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_339, plain, (v1_xreal_0(np__1))).#trainneg
cnf(i_0_445, negated_conjecture, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_221, plain, (np__0=k6_numbers)).#trainneg
cnf(i_0_351, plain, (v7_ordinal1(np__1))).#trainneg
cnf(i_0_364, plain, (v1_relat_1(k24_sin_cos))).#trainneg
cnf(i_0_380, negated_conjecture, (v3_membered(esk3_0))).#trainneg
cnf(i_0_481, plain, (k1_fdiff_1(k3_taylor_1,k5_square_1(k1_seq_1(k19_sin_cos,X1)))=k1_fdiff_1(k29_sin_cos,X1)|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~m2_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,X1)),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_520, plain, (r2_fdiff_1(X1,X2)|m1_subset_1(esk6_2(X2,X1),k1_numbers)|~v3_rcomp_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_518, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|m1_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_517, plain, (r2_fdiff_1(k3_taylor_1,X1)|m1_subset_1(esk6_2(X1,k3_taylor_1),k1_numbers)|~v3_rcomp_1(X1)|~r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_427, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_574, plain, (k8_real_1(k1_fdiff_1(X1,X2),k1_fdiff_1(X3,k1_seq_1(X1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X3),X2)|~v1_xreal_0(k1_fdiff_1(X3,k1_seq_1(X1,X2)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~r1_fdiff_1(X3,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_573, plain, (k8_real_1(k1_fdiff_1(X2,X3),k1_fdiff_1(X1,k1_seq_1(X2,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X1),X3)|~v1_xreal_0(k1_fdiff_1(X2,X3))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(X2,X3)),k1_numbers)|~r1_fdiff_1(X1,k1_seq_1(X2,X3))|~r1_fdiff_1(X2,X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_80, plain, (r2_relset_1(X1,X2,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_418, plain, (v3_valued_0(k29_sin_cos))).#trainneg
cnf(i_0_416, plain, (v3_valued_0(k24_sin_cos))).#trainneg
cnf(i_0_326, plain, (v1_xcmplx_0(np__1)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_308, plain, (v1_xcmplx_0(np__1)|~v1_xcmplx_0(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_443, plain, (m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~v1_xreal_0(X1)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_373, plain, (r2_hidden(k5_square_1(X1),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_477, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~r1_tarski(X1,X3)|~v3_rcomp_1(X3)|~r2_fdiff_1(X2,X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_460, plain, (r1_tarski(X1,k9_xtuple_0(X2))|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(X2,X3)))|~v1_relat_1(X3)|~v1_relat_1(X2))).#trainneg
cnf(i_0_446, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_372, plain, (r2_hidden(np__1,k1_numbers))).#trainneg
cnf(i_0_371, plain, (r2_hidden(np__1,k5_numbers))).#trainneg
cnf(i_0_309, plain, (v1_xcmplx_0(np__1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_310, plain, (v1_int_1(np__1)|~v1_int_1(k4_xcmplx_0(np__1)))).#trainneg
cnf(i_0_321, negated_conjecture, (v1_xcmplx_0(esk3_0)|~v1_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_334, negated_conjecture, (v1_xreal_0(esk3_0)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_347, negated_conjecture, (v7_ordinal1(esk3_0)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_316, plain, (k5_square_1(X1)=k3_xcmplx_0(X1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_478, plain, (r1_tarski(X1,a_0_0_fdiff_8)|~v3_rcomp_1(X1)|~r2_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_530, plain, (k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2))=X2|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_476, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~v3_rcomp_1(X1)|~r2_fdiff_1(X2,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_455, plain, (r2_hidden(k1_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_435, plain, (m1_subset_1(k8_real_1(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_391, plain, (r2_hidden(k8_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_389, plain, (v4_relat_1(k19_sin_cos,k1_numbers))).#trainneg
cnf(i_0_442, plain, (v1_xreal_0(k10_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_448, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k3_limfunc1(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_421, plain, (v5_relat_1(k3_taylor_1,k1_numbers))).#trainneg
cnf(i_0_423, plain, (v5_relat_1(k19_sin_cos,k1_numbers))).#trainneg
cnf(i_0_424, plain, (v5_relat_1(k29_sin_cos,k1_numbers))).#trainneg
cnf(i_0_388, plain, (v4_relat_1(k24_sin_cos,k1_numbers))).#trainneg
cnf(i_0_422, plain, (v5_relat_1(k24_sin_cos,k1_numbers))).#trainneg
cnf(i_0_463, plain, (m1_subset_1(k6_numbers,k5_numbers))).#trainneg
cnf(i_0_497, plain, (m1_subset_1(k6_numbers,k1_numbers))).#trainneg
cnf(i_0_319, plain, (m1_subset_1(a_0_0_fdiff_8,k1_zfmisc_1(k1_numbers))|~v1_xreal_0(k6_numbers))).#trainneg
cnf(i_0_444, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_396, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_480, plain, (m1_subset_1(k1_fdiff_1(k29_sin_cos,X1),k1_numbers)|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,X1)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_516, plain, (v1_funct_1(k3_relat_1(X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_439, plain, (r2_hidden(k10_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_473, plain, (v4_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_474, plain, (v3_valued_0(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_367, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|r2_hidden(esk3_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_313, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_403, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_taylor_1))).#trainneg
cnf(i_0_406, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k19_sin_cos))).#trainneg
cnf(i_0_374, plain, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|r2_hidden(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_505, plain, (k2_partfun2(k1_numbers,k1_numbers,k24_sin_cos)=k3_taylor_1|~m1_subset_1(k2_partfun2(k1_numbers,k1_numbers,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_375, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|r2_hidden(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_377, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|r2_hidden(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_378, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|r2_hidden(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_419, plain, (r1_fdiff_1(k3_taylor_1,esk2_1(X1))|~r2_hidden(X1,a_0_0_fdiff_8)|~m1_subset_1(esk2_1(X1),k1_numbers))).#trainneg
cnf(i_0_570, plain, (m1_subset_1(k3_relat_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_472, plain, (v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_314, plain, (k5_xcmplx_0(k7_xcmplx_0(np__1,X1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_368, plain, (r2_hidden(esk2_1(X1),k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_522, plain, (r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,X3)))|~r2_hidden(k1_seq_1(X2,X1),k9_xtuple_0(X3))|~r2_hidden(X1,k9_xtuple_0(X2))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~v3_valued_0(X2))).#trainneg
cnf(i_0_493, plain, (k7_xcmplx_0(X1,k3_square_1(X2))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_541, plain, (v4_relat_1(k3_rfunct_1(X1,X2,X3,X4),X1)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v3_membered(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_353, plain, (v7_ordinal1(k5_square_1(X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_470, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k2_fdiff_1(X2,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_376, plain, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|r2_hidden(k3_limfunc1(X1),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_407, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k29_sin_cos))).#trainneg
cnf(i_0_405, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k24_sin_cos))).#trainneg
cnf(i_0_383, plain, (v3_membered(k3_limfunc1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_379, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|r2_hidden(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_471, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k2_fdiff_1(X2,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_447, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_taylor_1))).#trainneg
cnf(i_0_449, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k24_sin_cos))).#trainneg
cnf(i_0_450, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k19_sin_cos))).#trainneg
cnf(i_0_459, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(X1,X2)),k1_zfmisc_1(k9_xtuple_0(X1)))|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_507, plain, (m1_subset_1(k1_funct_1(X1,X2),k9_xtuple_0(X3))|~r2_hidden(X2,k9_xtuple_0(k3_relat_1(X1,X3)))|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X1))).#trainneg
cnf(i_0_568, plain, (v3_valued_0(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v3_membered(X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_456, plain, (v1_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_membered(k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_506, plain, (~v1_xboole_0(k9_xtuple_0(X1))|~r2_hidden(X3,k9_xtuple_0(k3_relat_1(X2,X1)))|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_491, plain, (k7_xcmplx_0(X1,k3_xcmplx_0(X3,X2))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_392, plain, (v1_xcmplx_0(k8_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_454, plain, (v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v1_valued_0(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_393, plain, (v7_ordinal1(k8_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_451, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k29_sin_cos))).#trainneg
cnf(i_0_440, plain, (v1_xcmplx_0(k10_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_441, plain, (v7_ordinal1(k10_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_420, plain, (r1_fdiff_1(k3_taylor_1,k1_fdiff_1(k3_taylor_1,X1))|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_526, plain, (r2_fdiff_1(X1,k1_relset_1(k1_numbers,k29_sin_cos))|k1_seq_1(k19_sin_cos,esk6_2(k1_relset_1(k1_numbers,k29_sin_cos),X1))!=k6_numbers|~m1_subset_1(esk6_2(k1_relset_1(k1_numbers,k29_sin_cos),X1),k1_numbers)|~v3_rcomp_1(k1_relset_1(k1_numbers,k29_sin_cos))|~r1_tarski(k1_relset_1(k1_numbers,k29_sin_cos),k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,k29_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_468, plain, (v3_membered(k2_fdiff_1(X1,X2))|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_475, plain, (v5_relat_1(k2_fdiff_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_464, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|r2_hidden(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_431, plain, (k7_xcmplx_0(X1,k5_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_554, plain, (X3=k2_fdiff_1(X2,X1)|~v1_xboole_0(X1)|k1_relset_1(k1_numbers,X3)!=X1|~r2_fdiff_1(X2,X1)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_540, plain, (v1_relat_1(k3_rfunct_1(X1,X2,X3,X4))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v3_membered(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_430, plain, (k3_square_1(k5_xcmplx_0(X1))=k7_xcmplx_0(k5_xcmplx_0(X1),X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_484, plain, (k3_xcmplx_0(k7_xcmplx_0(X2,X3),X1)=k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k7_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_428, plain, (k3_xcmplx_0(k5_xcmplx_0(X2),X1)=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_489, plain, (k7_xcmplx_0(k7_xcmplx_0(X1,X2),k5_xcmplx_0(X3))=k7_xcmplx_0(X1,k7_xcmplx_0(X2,X3))|~v1_xcmplx_0(k5_xcmplx_0(X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_486, plain, (k7_xcmplx_0(k3_xcmplx_0(k7_xcmplx_0(X1,X2),X1),X2)=k3_square_1(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_569, plain, (v5_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6),X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_457, plain, (v7_ordinal1(k1_fdiff_1(X1,X2))|~v6_membered(k1_numbers)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_537, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,k1_numbers))|~r2_hidden(X1,k3_rfunct_1(X2,X3,X4,X5))|~v1_funct_1(X5)|~v1_funct_1(X4)|~v3_membered(X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_487, plain, (k7_xcmplx_0(k3_xcmplx_0(X1,X2),np__1)=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_543, plain, (v5_relat_1(k3_rfunct_1(X1,X2,X3,X4),k1_numbers)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v3_membered(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_519, plain, (r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|m1_subset_1(esk6_2(k1_relset_1(k1_numbers,X1),X1),k1_numbers)|~v3_rcomp_1(k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_563, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,k1_partfun1(X2,X4,X5,X3,X6,X7))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X3)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X4))))).#trainneg
cnf(i_0_553, plain, (X2=k2_fdiff_1(X1,k1_relset_1(k1_numbers,k29_sin_cos))|k1_seq_1(k19_sin_cos,esk1_3(X1,k1_relset_1(k1_numbers,k29_sin_cos),X2))!=k6_numbers|~m1_subset_1(esk1_3(X1,k1_relset_1(k1_numbers,k29_sin_cos),X2),k1_numbers)|k1_relset_1(k1_numbers,X2)!=k1_relset_1(k1_numbers,k29_sin_cos)|~r2_fdiff_1(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_509, plain, (r2_hidden(k1_funct_1(X1,X2),k9_xtuple_0(X3))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_relat_1(X4)|~v1_relat_1(X3)|~r2_hidden(X2,k9_xtuple_0(k3_relat_1(X1,k3_relat_1(X3,X4))))|~v1_funct_1(k3_relat_1(X3,X4))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_508, plain, (r1_fdiff_1(X1,k1_funct_1(X2,X3))|~v3_rcomp_1(k9_xtuple_0(X4))|~r2_fdiff_1(X1,k9_xtuple_0(X4))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k9_xtuple_0(X4),k1_zfmisc_1(k1_numbers))|~r2_hidden(X3,k9_xtuple_0(k3_relat_1(X2,X4)))|~v1_funct_1(X4)|~v1_funct_1(X2)|~v1_relat_1(X4)|~v1_relat_1(X2))).#trainneg
cnf(i_0_548, plain, (v1_xreal_0(esk1_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|k1_relset_1(k1_numbers,X3)!=X2|~r2_fdiff_1(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_546, plain, (v1_xcmplx_0(esk1_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v1_membered(k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~r2_fdiff_1(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_488, plain, (k7_xcmplx_0(k3_xcmplx_0(X1,X2),k3_xcmplx_0(X3,X4))=k3_xcmplx_0(X1,k7_xcmplx_0(k7_xcmplx_0(X2,X3),X4))|~v1_xcmplx_0(k3_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_490, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(k3_xcmplx_0(X2,X3),X4))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),k7_xcmplx_0(X3,X4))|~v1_xcmplx_0(k7_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_545, plain, (r2_hidden(esk1_3(X1,X2,X3),k1_numbers)|X3=k2_fdiff_1(X1,X2)|k1_relset_1(k1_numbers,X3)!=X2|~r2_fdiff_1(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_555, plain, (m1_subset_1(esk1_3(X1,X2,X3),X2)|X3=k2_fdiff_1(X1,X2)|k1_relset_1(k1_numbers,X3)!=X2|~r2_fdiff_1(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_557, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|r2_hidden(k1_partfun1(X1,X3,X4,X2,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_547, plain, (v7_ordinal1(esk1_3(X1,X2,X3))|X3=k2_fdiff_1(X1,X2)|~v6_membered(k1_numbers)|k1_relset_1(k1_numbers,X3)!=X2|~r2_fdiff_1(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_556, plain, (r1_fdiff_1(X1,esk1_3(X2,X3,X4))|X4=k2_fdiff_1(X2,X3)|~v3_rcomp_1(X3)|~r2_fdiff_1(X1,X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|k1_relset_1(k1_numbers,X4)!=X3|~r2_fdiff_1(X2,X3)|~v1_funct_1(X4)|~v1_funct_1(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_525, plain, (r2_hidden(esk6_2(k9_xtuple_0(k3_relat_1(X1,X2)),X3),k9_xtuple_0(X1))|r2_fdiff_1(X3,k9_xtuple_0(k3_relat_1(X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v3_rcomp_1(k9_xtuple_0(k3_relat_1(X1,X2)))|~r1_tarski(k9_xtuple_0(k3_relat_1(X1,X2)),k1_relset_1(k1_numbers,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k9_xtuple_0(k3_relat_1(X1,X2)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_552, plain, (r2_hidden(esk1_3(X1,k9_xtuple_0(k3_relat_1(X2,X3)),X4),k9_xtuple_0(X2))|X4=k2_fdiff_1(X1,k9_xtuple_0(k3_relat_1(X2,X3)))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|k1_relset_1(k1_numbers,X4)!=k9_xtuple_0(k3_relat_1(X2,X3))|~r2_fdiff_1(X1,k9_xtuple_0(k3_relat_1(X2,X3)))|~v1_funct_1(X4)|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_615, plain, (v1_int_1(np__1))).#trainneg
cnf(i_0_815, plain, (v7_ordinal1(k6_numbers))).#trainneg
cnf(i_0_816, plain, (v1_xreal_0(k6_numbers))).#trainneg
cnf(i_0_814, plain, (v1_xcmplx_0(k6_numbers)|~v1_membered(k5_numbers))).#trainneg
cnf(i_0_818, plain, (v1_xcmplx_0(k6_numbers)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_596, plain, (m1_subset_1(k1_partfun1(X7,X8,X9,X10,X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X4)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_1037, plain, (v1_xcmplx_0(k1_seq_1(k19_sin_cos,X1)))).#trainneg
cnf(i_0_1033, plain, (v1_xcmplx_0(k1_seq_1(k3_taylor_1,X1)))).#trainneg
cnf(i_0_1038, plain, (v1_xcmplx_0(k1_seq_1(k29_sin_cos,X1)))).#trainneg
cnf(i_0_970, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_789, plain, (k9_xtuple_0(k19_sin_cos)=k1_relset_1(k1_numbers,k19_sin_cos))).#trainneg
cnf(i_0_940, plain, (v3_membered(a_0_0_fdiff_8)|~v1_xreal_0(k6_numbers))).#trainneg
cnf(i_0_813, plain, (r2_hidden(k6_numbers,k5_numbers))).#trainneg
cnf(i_0_817, plain, (r2_hidden(k6_numbers,k1_numbers))).#trainneg
cnf(i_0_759, negated_conjecture, (r1_tarski(esk3_0,a_0_0_fdiff_8)|~r2_fdiff_1(k3_taylor_1,esk3_0))).#trainneg
cnf(i_0_610, negated_conjecture, (m1_subset_1(esk6_2(esk3_0,X1),k1_numbers)|r2_fdiff_1(X1,esk3_0)|~r1_tarski(esk3_0,k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_635, negated_conjecture, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,esk3_0))).#trainneg
cnf(i_0_855, plain, (k9_xtuple_0(k2_fdiff_1(X1,X2))=k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_579, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k3_xcmplx_0(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~m2_subset_1(k3_xcmplx_0(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))),k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_788, plain, (m1_subset_1(k1_relset_1(k1_numbers,k19_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_576, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_8)|k1_fdiff_1(k3_taylor_1,k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~m1_subset_1(a_0_0_fdiff_8,k1_zfmisc_1(k1_numbers))|~m1_subset_1(k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))),a_0_0_fdiff_8))).#trainneg
cnf(i_0_900, plain, (k5_xcmplx_0(np__1)=np__1|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_811, plain, (k9_xtuple_0(k24_sin_cos)=k1_relset_1(k1_numbers,k24_sin_cos))).#trainneg
cnf(i_0_636, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k5_numbers))).#trainneg
cnf(i_0_698, negated_conjecture, (v1_xcmplx_0(esk3_0)|~v1_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))))).#trainneg
cnf(i_0_699, negated_conjecture, (v7_ordinal1(esk3_0)|~v6_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))))).#trainneg
cnf(i_0_700, negated_conjecture, (v1_xreal_0(esk3_0)|~v3_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))))).#trainneg
cnf(i_0_798, plain, (m1_subset_1(k1_relset_1(k1_numbers,k29_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_810, plain, (m1_subset_1(k1_relset_1(k1_numbers,k24_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_590, negated_conjecture, (r1_tarski(esk3_0,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1066, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_669, negated_conjecture, (k7_xcmplx_0(np__1,k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~m1_subset_1(esk4_0,k1_numbers))).#trainneg
cnf(i_0_905, plain, (r2_hidden(X1,k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_577, negated_conjecture, (k7_xcmplx_0(np__1,k3_xcmplx_0(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k3_xcmplx_0(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~m1_subset_1(esk4_0,k1_numbers))).#trainneg
cnf(i_0_1034, plain, (v1_xcmplx_0(k1_seq_1(k24_sin_cos,X1)))).#trainneg
cnf(i_0_662, plain, (m1_subset_1(k1_seq_1(k29_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_725, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_939, plain, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|r2_hidden(a_0_0_fdiff_8,k1_zfmisc_1(k1_numbers))|~v1_xreal_0(k6_numbers))).#trainneg
cnf(i_0_857, negated_conjecture, (k1_zfmisc_1(k1_numbers)=k6_numbers|r2_hidden(esk3_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_870, plain, (k1_zfmisc_1(k1_numbers)=k6_numbers|r2_hidden(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_686, negated_conjecture, (k10_real_1(np__1,k3_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),esk4_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers))).#trainneg
cnf(i_0_614, plain, (r2_hidden(k1_fdiff_1(k3_taylor_1,X1),a_0_0_fdiff_8)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_650, plain, (k3_xcmplx_0(k1_fdiff_1(X1,X2),k1_fdiff_1(X3,k1_seq_1(X1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,X3),X2)|~v1_xreal_0(k1_fdiff_1(X3,k1_seq_1(X1,X2)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~r1_fdiff_1(X3,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_634, plain, (v1_xreal_0(esk6_2(X1,k3_taylor_1))|r2_fdiff_1(k3_taylor_1,X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_586, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers)|~v3_rcomp_1(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,esk3_0))).#trainneg
cnf(i_0_688, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k3_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),esk4_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~m2_subset_1(k3_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),esk4_0),k1_numbers,a_0_0_fdiff_8)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers))).#trainneg
cnf(i_0_697, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|r2_hidden(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))))).#trainneg
cnf(i_0_1075, plain, (k2_fdiff_1(X1,X2)=k3_taylor_1|a_0_0_fdiff_8!=X2|~v1_xboole_0(X2)|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_580, plain, (k8_real_1(k1_fdiff_1(X1,k6_numbers),k1_fdiff_1(k19_sin_cos,esk6_2(X2,k29_sin_cos)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,X1),esk6_2(X2,k29_sin_cos))|r2_fdiff_1(k29_sin_cos,X2)|~r1_fdiff_1(X1,k6_numbers)|~r1_fdiff_1(k19_sin_cos,esk6_2(X2,k29_sin_cos))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk6_2(X2,k29_sin_cos),k1_numbers)|~v3_rcomp_1(X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_663, plain, (m1_subset_1(k1_seq_1(k24_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_968, negated_conjecture, (v1_xboole_0(esk3_0)|m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_871, plain, (k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))=k6_numbers|r2_hidden(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_872, plain, (k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))=k6_numbers|r2_hidden(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_873, plain, (k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))=k6_numbers|r2_hidden(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_941, plain, (k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))=k6_numbers|r2_hidden(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_715, plain, (r1_tarski(X1,k9_xtuple_0(X2))|~v1_relat_1(X3)|~v1_relat_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k9_xtuple_0(k3_relat_1(X2,X3)))))).#trainneg
cnf(i_0_585, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r1_tarski(X1,X2)|~r1_tarski(X2,esk3_0))).#trainneg
cnf(i_0_680, plain, (r2_fdiff_1(k3_taylor_1,X1)|m1_subset_1(esk6_2(X1,k3_taylor_1),X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_679, plain, (v1_xreal_0(k1_fdiff_1(k3_taylor_1,X1))|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_721, plain, (v1_xboole_0(k1_zfmisc_1(X1))|r2_hidden(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_825, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_581, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)),X4)|~v3_rcomp_1(X4)|~r1_fdiff_1(X1,esk6_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))))|~r1_tarski(X4,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))|~v1_funct_1(X1)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~r1_fdiff_1(X3,k1_seq_1(X2,k1_seq_1(X1,esk6_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3))))))|~r1_fdiff_1(X2,k1_seq_1(X1,esk6_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)))))|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X1,esk6_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,X3)))),k1_numbers))).#trainneg
cnf(i_0_946, plain, (v1_xcmplx_0(X1)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_taylor_1))).#trainneg
cnf(i_0_947, plain, (v7_ordinal1(X1)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_taylor_1))).#trainneg
cnf(i_0_948, plain, (v1_xreal_0(X1)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_taylor_1))).#trainneg
cnf(i_0_624, negated_conjecture, (v1_xreal_0(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_622, negated_conjecture, (v1_xcmplx_0(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_membered(k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_623, negated_conjecture, (v7_ordinal1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v6_membered(k1_numbers)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_621, negated_conjecture, (r2_hidden(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_684, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),X1)|~v3_rcomp_1(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,esk3_0))).#trainneg
cnf(i_0_620, plain, (v1_xreal_0(esk6_2(X1,X2))|r2_fdiff_1(X2,X1)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_792, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1198, plain, (r2_fdiff_1(k2_fdiff_1(X1,X2),X2)|m1_subset_1(esk6_2(X2,k2_fdiff_1(X1,X2)),k1_numbers)|~v3_rcomp_1(X2)|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_683, plain, (r2_fdiff_1(X1,X2)|m1_subset_1(esk6_2(X2,X1),X2)|~v3_rcomp_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_674, plain, (v1_xreal_0(k8_real_1(X2,X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_638, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k3_limfunc1(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_954, plain, (v1_xcmplx_0(X1)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k19_sin_cos))).#trainneg
cnf(i_0_955, plain, (v7_ordinal1(X1)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k19_sin_cos))).#trainneg
cnf(i_0_956, plain, (v1_xreal_0(X1)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k19_sin_cos))).#trainneg
cnf(i_0_965, plain, (~v1_xboole_0(k9_xtuple_0(X1))|~r2_hidden(X2,k9_xtuple_0(k3_relat_1(X1,X3)))|~v1_relat_1(X3)|~v1_relat_1(X1))).#trainneg
cnf(i_0_950, plain, (v1_xcmplx_0(X1)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k24_sin_cos))).#trainneg
cnf(i_0_951, plain, (v7_ordinal1(X1)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k24_sin_cos))).#trainneg
cnf(i_0_952, plain, (v1_xreal_0(X1)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k24_sin_cos))).#trainneg
cnf(i_0_591, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k29_sin_cos)))|~r1_tarski(X1,esk3_0)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_693, plain, (k3_xcmplx_0(X2,X1)=k3_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_640, plain, (r1_tarski(X1,k2_zfmisc_1(X2,X3))|~r1_tarski(X1,k1_partfun1(X2,X4,X5,X3,X6,X7))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X3)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X4))))).#trainneg
cnf(i_0_1073, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_749, plain, (k7_xcmplx_0(X1,k5_square_1(X2))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_858, plain, (r2_fdiff_1(k3_taylor_1,X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,a_0_0_fdiff_8)|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1048, plain, (v1_xcmplx_0(X1)|~v1_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k29_sin_cos))).#trainneg
cnf(i_0_1049, plain, (v7_ordinal1(X1)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k29_sin_cos))).#trainneg
cnf(i_0_1050, plain, (v1_xreal_0(X1)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k29_sin_cos))).#trainneg
cnf(i_0_637, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k3_taylor_1))).#trainneg
cnf(i_0_639, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k24_sin_cos))).#trainneg
cnf(i_0_643, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k19_sin_cos))).#trainneg
cnf(i_0_644, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k29_sin_cos))).#trainneg
cnf(i_0_891, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,X5,X6)),k1_zfmisc_1(X1))|~v1_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_613, plain, (r2_hidden(esk2_1(X1),a_0_0_fdiff_8)|~r2_hidden(X1,a_0_0_fdiff_8)|~m1_subset_1(esk2_1(X1),k1_numbers))).#trainneg
cnf(i_0_631, plain, (r2_hidden(esk6_2(X1,k3_taylor_1),k1_numbers)|r2_fdiff_1(k3_taylor_1,X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_862, plain, (r2_fdiff_1(X1,X2)|~v3_rcomp_1(X2)|~v1_xboole_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_768, plain, (v1_xcmplx_0(k8_real_1(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_616, plain, (k1_fdiff_1(k3_taylor_1,k5_square_1(k1_seq_1(k19_sin_cos,X1)))=k1_fdiff_1(k29_sin_cos,X1)|k1_seq_1(k19_sin_cos,X1)=k6_numbers|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(a_0_0_fdiff_8,k1_zfmisc_1(k1_numbers))|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,X1)),a_0_0_fdiff_8))).#trainneg
cnf(i_0_726, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_728, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_945, plain, (v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|r2_hidden(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_taylor_1))).#trainneg
cnf(i_0_949, plain, (v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|r2_hidden(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k24_sin_cos))).#trainneg
cnf(i_0_953, plain, (v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|r2_hidden(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k19_sin_cos))).#trainneg
cnf(i_0_1047, plain, (v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|r2_hidden(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k29_sin_cos))).#trainneg
cnf(i_0_854, plain, (m1_subset_1(k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_751, plain, (v1_xreal_0(k3_xcmplx_0(X1,X1))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_929, plain, (k1_relset_1(X1,k3_rfunct_1(X1,X2,X3,X4))=k9_xtuple_0(k3_rfunct_1(X1,X2,X3,X4))|~v1_relat_1(k3_rfunct_1(X1,X2,X3,X4))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v3_membered(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_656, plain, (k8_real_1(k1_fdiff_1(k19_sin_cos,esk6_2(X1,k29_sin_cos)),k1_fdiff_1(X2,k6_numbers))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,X2),esk6_2(X1,k29_sin_cos))|r2_fdiff_1(k29_sin_cos,X1)|~r1_fdiff_1(X2,k6_numbers)|~r1_fdiff_1(k19_sin_cos,esk6_2(X1,k29_sin_cos))|~v1_funct_1(X2)|~v1_xreal_0(k1_fdiff_1(k19_sin_cos,esk6_2(X1,k29_sin_cos)))|~m1_subset_1(k1_fdiff_1(X2,k6_numbers),k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk6_2(X1,k29_sin_cos),k1_numbers)|~v3_rcomp_1(X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_630, plain, (k3_xcmplx_0(k1_fdiff_1(X1,k6_numbers),k1_fdiff_1(k19_sin_cos,esk6_2(X2,k29_sin_cos)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,X1),esk6_2(X2,k29_sin_cos))|r2_fdiff_1(k29_sin_cos,X2)|~r1_fdiff_1(X1,k6_numbers)|~r1_fdiff_1(k19_sin_cos,esk6_2(X2,k29_sin_cos))|~v1_funct_1(X1)|~v1_xreal_0(k1_fdiff_1(k19_sin_cos,esk6_2(X2,k29_sin_cos)))|~m1_subset_1(k1_fdiff_1(X1,k6_numbers),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk6_2(X2,k29_sin_cos),k1_numbers)|~v3_rcomp_1(X2)|~r1_tarski(X2,k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1127, plain, (k3_xcmplx_0(k7_xcmplx_0(X1,X2),X3)=k7_xcmplx_0(X1,k7_xcmplx_0(X2,X3))|~v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_835, plain, (m1_subset_1(k3_xcmplx_0(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_617, plain, (r2_hidden(esk6_2(X1,X2),k1_numbers)|r2_fdiff_1(X2,X1)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_769, plain, (v7_ordinal1(k8_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_863, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v3_rcomp_1(X1)|~v1_xboole_0(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,esk3_0))).#trainneg
cnf(i_0_618, plain, (v1_xcmplx_0(esk6_2(X1,X2))|r2_fdiff_1(X2,X1)|~v1_membered(k1_numbers)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1000, plain, (k3_xcmplx_0(X1,np__1)=k7_xcmplx_0(X1,np__1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_767, plain, (r2_hidden(k8_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_967, negated_conjecture, (v1_xboole_0(esk3_0)|v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|m1_subset_1(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_1064, plain, (v1_xcmplx_0(k3_xcmplx_0(np__1,X1))|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_677, plain, (v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~v1_membered(k1_numbers)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_678, plain, (v7_ordinal1(k1_fdiff_1(k3_taylor_1,X1))|~v6_membered(k1_numbers)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_652, plain, (k8_real_1(k1_fdiff_1(k19_sin_cos,esk6_2(X1,k29_sin_cos)),k1_fdiff_1(X2,k6_numbers))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,X2),esk6_2(X1,k29_sin_cos))|r2_fdiff_1(k29_sin_cos,X1)|~r1_fdiff_1(X2,k6_numbers)|~r1_fdiff_1(k19_sin_cos,esk6_2(X1,k29_sin_cos))|~v1_funct_1(X2)|~v1_xreal_0(k1_fdiff_1(X2,k6_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k19_sin_cos,esk6_2(X1,k29_sin_cos)),k1_numbers)|~m1_subset_1(esk6_2(X1,k29_sin_cos),k1_numbers)|~v3_rcomp_1(X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_760, plain, (r1_tarski(X1,X3)|~v3_rcomp_1(X1)|~r2_fdiff_1(k2_fdiff_1(X2,X3),X1)|~v1_funct_1(k2_fdiff_1(X2,X3))|~m1_subset_1(k2_fdiff_1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r2_fdiff_1(X2,X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_641, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k2_fdiff_1(X2,X3))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_763, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1)))|~r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_784, plain, (r2_hidden(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_750, plain, (m1_subset_1(k3_xcmplx_0(X1,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1121, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,X1),X2)=k7_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_736, plain, (r1_fdiff_1(X1,np__1)|~v3_rcomp_1(k1_numbers)|~r2_fdiff_1(X1,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_619, plain, (v7_ordinal1(esk6_2(X1,X2))|r2_fdiff_1(X2,X1)|~v6_membered(k1_numbers)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_752, plain, (r2_hidden(k3_xcmplx_0(X1,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_821, plain, (r2_hidden(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_676, plain, (r2_hidden(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_837, plain, (m1_subset_1(k3_square_1(X1),k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_739, plain, (r1_fdiff_1(X1,np__1)|~v3_rcomp_1(k5_numbers)|~r2_fdiff_1(X1,k5_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_706, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos))))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_753, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,X6)),k1_zfmisc_1(k9_xtuple_0(X5)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_839, plain, (k1_seq_1(k19_sin_cos,X1)=k6_numbers|m1_subset_1(k1_fdiff_1(k29_sin_cos,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_926, plain, (k7_xcmplx_0(X1,k3_xcmplx_0(X2,X2))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_997, plain, (k5_square_1(k7_xcmplx_0(np__1,X1))=k7_xcmplx_0(k7_xcmplx_0(np__1,X1),X1)|~v1_xcmplx_0(k7_xcmplx_0(np__1,X1))|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_963, plain, (m1_subset_1(X1,k9_xtuple_0(X2))|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,X3)))|~v1_relat_1(X3)|~v1_relat_1(X2))).#trainneg
cnf(i_0_632, plain, (v1_xcmplx_0(esk6_2(X1,k3_taylor_1))|r2_fdiff_1(k3_taylor_1,X1)|~v1_membered(k1_numbers)|~v3_rcomp_1(X1)|~r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_822, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_823, plain, (v7_ordinal1(k7_xcmplx_0(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_829, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1083, plain, (m1_subset_1(k3_square_1(k5_xcmplx_0(X1)),k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_840, plain, (v1_funct_1(k3_relat_1(X1,k3_taylor_1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_841, plain, (v1_funct_1(k3_relat_1(X1,k24_sin_cos))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_845, plain, (v1_funct_1(k3_relat_1(X1,k19_sin_cos))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_846, plain, (v1_funct_1(k3_relat_1(X1,k29_sin_cos))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_928, plain, (m1_subset_1(k1_relset_1(X1,k3_rfunct_1(X1,X2,X3,X4)),k1_zfmisc_1(X1))|~v1_relat_1(k3_rfunct_1(X1,X2,X3,X4))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v3_membered(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_875, plain, (m1_subset_1(k3_relat_1(X1,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_985, plain, (m1_subset_1(k1_funct_1(X1,X2),k9_xtuple_0(X3))|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X1)|~r2_hidden(k1_seq_1(X1,X2),k9_xtuple_0(X3))|~r2_hidden(X2,k9_xtuple_0(X1))|~v3_valued_0(X1))).#trainneg
cnf(i_0_710, plain, (r1_fdiff_1(X1,k5_square_1(X2))|~v3_rcomp_1(k1_numbers)|~r2_fdiff_1(X1,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_numbers,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1119, plain, (k7_xcmplx_0(X2,k5_xcmplx_0(X1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_633, plain, (v7_ordinal1(esk6_2(X1,k3_taylor_1))|r2_fdiff_1(k3_taylor_1,X1)|~v6_membered(k1_numbers)|~v3_rcomp_1(X1)|~r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_830, plain, (v7_ordinal1(k3_xcmplx_0(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_883, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X1))|~v1_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_932, plain, (v7_ordinal1(k3_xcmplx_0(X1,X1))|~v6_membered(k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_907, plain, (~v1_xboole_0(k9_xtuple_0(k3_relat_1(X1,X2)))|~r2_hidden(k1_seq_1(X1,X3),k9_xtuple_0(X2))|~r2_hidden(X3,k9_xtuple_0(X1))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_850, plain, (r1_fdiff_1(X1,k10_real_1(X2,X3))|~v3_rcomp_1(k1_numbers)|~r2_fdiff_1(X1,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1087, plain, (k7_xcmplx_0(X1,k5_xcmplx_0(X1))=k3_square_1(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_756, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(X5,X6)),k9_xtuple_0(X5))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_881, plain, (m1_subset_1(k3_relat_1(X1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1166, plain, (v3_valued_0(k1_partfun1(X1,X2,X3,k1_numbers,X4,X5))|~v1_relat_1(k1_partfun1(X1,X2,X3,k1_numbers,X4,X5))|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1017, plain, (k7_xcmplx_0(k7_xcmplx_0(X1,X3),X2)=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1120, plain, (k7_xcmplx_0(X2,k7_xcmplx_0(np__1,X1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_880, plain, (m1_subset_1(k3_relat_1(X1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_744, plain, (k5_square_1(k5_xcmplx_0(X1))=k7_xcmplx_0(k5_xcmplx_0(X1),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_833, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(k5_xcmplx_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_642, plain, (r1_tarski(X1,k2_zfmisc_1(X2,k1_numbers))|~r1_tarski(X1,k3_rfunct_1(X2,X3,X4,X5))|~v1_funct_1(X5)|~v1_funct_1(X4)|~v3_membered(X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_972, plain, (v1_xboole_0(k3_limfunc1(X1))|m1_subset_1(X2,k1_numbers)|~m1_subset_1(X2,k3_limfunc1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1068, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(k5_xcmplx_0(X2))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1113, plain, (m1_subset_1(k7_xcmplx_0(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_961, plain, (v3_membered(k9_xtuple_0(k3_relat_1(X1,X2)))|~v3_membered(k9_xtuple_0(X1))|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_718, plain, (m1_subset_1(esk6_2(k5_numbers,X1),k1_numbers)|r2_fdiff_1(X1,k5_numbers)|~v3_rcomp_1(k5_numbers)|~r1_tarski(k5_numbers,k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_658, plain, (m1_subset_1(esk6_2(a_0_0_fdiff_8,X1),k1_numbers)|r2_fdiff_1(X1,a_0_0_fdiff_8)|~v3_rcomp_1(a_0_0_fdiff_8)|~r1_tarski(a_0_0_fdiff_8,k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(a_0_0_fdiff_8,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_762, plain, (r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|m1_subset_1(esk6_2(X3,k2_fdiff_1(X1,X2)),X3)|~v3_rcomp_1(X3)|~r1_tarski(X3,X2)|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_780, plain, (r1_fdiff_1(X1,k8_real_1(X2,X3))|~v3_rcomp_1(k1_numbers)|~r2_fdiff_1(X1,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_904, plain, (r1_fdiff_1(X1,esk2_1(X2))|~v3_rcomp_1(k1_numbers)|~r2_fdiff_1(X1,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_hidden(X2,a_0_0_fdiff_8))).#trainneg
cnf(i_0_804, plain, (k7_xcmplx_0(k3_xcmplx_0(X1,np__1),k5_xcmplx_0(X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1191, plain, (v3_valued_0(k3_rfunct_1(X1,X2,X3,X4))|~v1_relat_1(k3_rfunct_1(X1,X2,X3,X4))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v3_membered(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1089, plain, (k3_square_1(k7_xcmplx_0(np__1,X1))=k7_xcmplx_0(k7_xcmplx_0(np__1,X1),X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1036, plain, (v1_xcmplx_0(k1_seq_1(k2_fdiff_1(X1,X2),X3))|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_884, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,X7)))|~v1_funct_1(X7)|~v1_funct_1(X2)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_645, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3),X4)|~v3_rcomp_1(X4)|~r1_fdiff_1(X3,k1_fdiff_1(X1,esk6_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3))))|~r1_fdiff_1(k2_fdiff_1(X1,X2),esk6_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3)))|~r1_tarski(X4,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3)))|~v1_funct_1(X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_zfmisc_1(k1_numbers))|~r2_hidden(esk6_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3)),X2)|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk6_2(X4,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X3)),k1_numbers))).#trainneg
cnf(i_0_1104, plain, (k7_xcmplx_0(k3_xcmplx_0(X2,X1),np__1)=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_876, plain, (m1_subset_1(k3_relat_1(X1,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X2,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_938, plain, (r1_fdiff_1(X1,k3_limfunc1(X2))|v1_xboole_0(k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k1_zfmisc_1(k1_numbers))|~r2_fdiff_1(X1,k1_zfmisc_1(k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_861, plain, (r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|~v3_rcomp_1(k1_relset_1(k1_numbers,X1))|~v1_xboole_0(k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_924, plain, (k7_xcmplx_0(np__1,k3_square_1(k5_xcmplx_0(X1)))=k7_xcmplx_0(X1,k5_xcmplx_0(X1))|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1129, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(X2,k5_xcmplx_0(X3)))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X3))|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_893, plain, (v4_relat_1(k3_relat_1(X5,X6),X1)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_908, plain, (m1_subset_1(X1,k9_xtuple_0(k3_relat_1(X2,X3)))|~r2_hidden(k1_seq_1(X2,X1),k9_xtuple_0(X3))|~r2_hidden(X1,k9_xtuple_0(X2))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v3_valued_0(X2)|~v1_relat_1(X3)|~v1_relat_1(X2))).#trainneg
cnf(i_0_745, plain, (k7_xcmplx_0(k3_xcmplx_0(k7_xcmplx_0(X1,X2),X1),X2)=k5_square_1(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers))).#trainneg
cnf(i_0_1070, plain, (k7_xcmplx_0(k3_xcmplx_0(X1,X2),k5_xcmplx_0(X3))=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(k5_xcmplx_0(X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_1074, plain, (k2_fdiff_1(X2,k1_relset_1(k1_numbers,X1))=X1|~v1_xboole_0(k1_relset_1(k1_numbers,X1))|~r2_fdiff_1(X2,k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1109, plain, (k7_xcmplx_0(k5_xcmplx_0(X2),X1)=k7_xcmplx_0(k5_xcmplx_0(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_980, plain, (m1_subset_1(k1_funct_1(X1,X2),k9_xtuple_0(X3))|~r2_hidden(X2,k9_xtuple_0(k1_partfun1(X4,X5,X6,X7,X1,X3)))|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_987, plain, (v3_valued_0(k3_relat_1(X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v3_membered(X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1019, plain, (k7_xcmplx_0(k7_xcmplx_0(X1,k5_xcmplx_0(X3)),X2)=k7_xcmplx_0(X1,k7_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k5_xcmplx_0(X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_856, plain, (m1_subset_1(k1_seq_1(k2_fdiff_1(X1,X2),X3),k1_numbers)|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1199, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,k3_relat_1(X6,X7))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X3)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X4))))).#trainneg
cnf(i_0_646, plain, (k2_fdiff_1(X1,X2)=k2_fdiff_1(X3,X4)|k1_fdiff_1(X3,esk1_3(X3,X4,k2_fdiff_1(X1,X2)))!=k1_fdiff_1(X1,esk1_3(X3,X4,k2_fdiff_1(X1,X2)))|k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2))!=X4|~r2_fdiff_1(X3,X4)|~v1_funct_1(X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_hidden(esk1_3(X3,X4,k2_fdiff_1(X1,X2)),X2)|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk1_3(X3,X4,k2_fdiff_1(X1,X2)),k1_numbers))).#trainneg
cnf(i_0_1084, plain, (k7_xcmplx_0(k3_square_1(k5_xcmplx_0(X1)),X1)=k7_xcmplx_0(k5_xcmplx_0(X1),k3_square_1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1)))).#trainneg
cnf(i_0_1065, plain, (k7_xcmplx_0(X1,k5_xcmplx_0(X2))=k3_xcmplx_0(X1,k3_xcmplx_0(np__1,X2))|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1101, plain, (k7_xcmplx_0(k3_xcmplx_0(X2,np__1),k5_xcmplx_0(X1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_1079, plain, (k3_square_1(k5_xcmplx_0(k5_xcmplx_0(X1)))=k3_xcmplx_0(k5_xcmplx_0(k5_xcmplx_0(X1)),X1)|~v1_xcmplx_0(k5_xcmplx_0(k5_xcmplx_0(X1)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1)))).#trainneg
cnf(i_0_1062, plain, (r1_fdiff_1(X1,k2_fdiff_1(X2,X3))|v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_766, plain, (r1_fdiff_1(X1,k1_fdiff_1(X2,X3))|~v3_rcomp_1(k1_numbers)|~r2_fdiff_1(X1,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X2)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_957, plain, (v1_xboole_0(k1_zfmisc_1(k9_xtuple_0(X1)))|r2_hidden(k9_xtuple_0(k3_relat_1(X1,X2)),k1_zfmisc_1(k9_xtuple_0(X1)))|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1167, plain, (v5_relat_1(k3_relat_1(X5,X6),X4)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_713, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(X1,X2),X3)),k9_xtuple_0(X1))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X3))).#trainneg
cnf(i_0_1105, plain, (k7_xcmplx_0(k3_xcmplx_0(X3,X1),k5_xcmplx_0(X2))=k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1076, plain, (k2_fdiff_1(X1,X2)=k2_fdiff_1(X3,X2)|~v1_xboole_0(X2)|~r2_fdiff_1(X3,X2)|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X3)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1425, plain, (v3_membered(k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_487, plain, (k7_xcmplx_0(k3_xcmplx_0(X1,X2),np__1)=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2801, plain, (k3_square_1(np__1)=np__1)).#trainneg
cnf(i_0_1333, plain, (v1_int_1(k6_numbers))).#trainneg
cnf(i_0_1334, plain, (v1_xcmplx_0(k6_numbers))).#trainneg
cnf(i_0_1401, negated_conjecture, (r1_tarski(esk3_0,k1_numbers))).#trainneg
cnf(i_0_1449, plain, (r1_tarski(k5_numbers,k1_numbers))).#trainneg
cnf(i_0_1761, plain, (v1_xboole_0(X1)|r2_hidden(esk6_2(X1,X2),X1)|r2_fdiff_1(X2,X1)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1734, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1440, plain, (k3_square_1(np__1)=k7_xcmplx_0(np__1,np__1))).#trainneg
cnf(i_0_1498, plain, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~v3_rcomp_1(X2)|~r1_fdiff_1(X1,esk6_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1)))|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r2_hidden(k1_seq_1(X1,esk6_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1))),a_0_0_fdiff_8))).#trainneg
cnf(i_0_1462, plain, (v3_membered(k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_1490, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_526, plain, (r2_fdiff_1(X1,k1_relset_1(k1_numbers,k29_sin_cos))|k1_seq_1(k19_sin_cos,esk6_2(k1_relset_1(k1_numbers,k29_sin_cos),X1))!=k6_numbers|~v3_rcomp_1(k1_relset_1(k1_numbers,k29_sin_cos))|~r1_tarski(k1_relset_1(k1_numbers,k29_sin_cos),k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2516, plain, (k5_square_1(np__1)=k7_xcmplx_0(np__1,np__1))).#trainneg
cnf(i_0_1472, plain, (v3_membered(k1_relset_1(k1_numbers,k24_sin_cos)))).#trainneg
cnf(i_0_1933, negated_conjecture, (r2_fdiff_1(k3_taylor_1,esk3_0)|~r1_tarski(esk3_0,a_0_0_fdiff_8)|~v1_xboole_0(esk3_0))).#trainneg
cnf(i_0_1653, plain, (v1_xboole_0(X1)|r2_hidden(esk6_2(X1,k3_taylor_1),X1)|r2_fdiff_1(k3_taylor_1,X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1504, plain, (v1_xreal_0(k1_seq_1(k19_sin_cos,X1)))).#trainneg
cnf(i_0_3282, plain, (r2_fdiff_1(k2_fdiff_1(X1,X2),X2)|~v3_rcomp_1(X2)|~v1_xboole_0(X2)|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1418, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_647, plain, (k8_real_1(k1_fdiff_1(X1,k1_fdiff_1(X2,X4)),k1_fdiff_1(k2_fdiff_1(X2,X3),X4))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X2,X3),X1),X4)|~r1_fdiff_1(X1,k1_fdiff_1(X2,X4))|~r1_fdiff_1(k2_fdiff_1(X2,X3),X4)|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X2,X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_numbers)|~r2_hidden(X4,X3)|~r2_fdiff_1(X2,X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2596, plain, (v1_funct_1(k3_relat_1(k3_taylor_1,k3_taylor_1)))).#trainneg
cnf(i_0_2597, plain, (v1_funct_1(k3_relat_1(k24_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_2602, plain, (v1_funct_1(k3_relat_1(k19_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_2603, plain, (v1_funct_1(k3_relat_1(k29_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_2609, plain, (v1_funct_1(k3_relat_1(k3_taylor_1,k24_sin_cos)))).#trainneg
cnf(i_0_3906, plain, (k3_square_1(np__1)=k3_xcmplx_0(np__1,np__1))).#trainneg
cnf(i_0_1967, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(X2,k1_partfun1(X2,X3,X4,X5,X6,X7)))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1710, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),X1)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1654, plain, (v1_xcmplx_0(esk6_2(X1,k3_taylor_1))|r2_fdiff_1(k3_taylor_1,X1)|~v1_membered(X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1362, plain, (r1_tarski(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,X5,X6)),k9_xtuple_0(X5))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1489, plain, (k3_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1517, plain, (v1_xreal_0(k1_seq_1(k29_sin_cos,X1)))).#trainneg
cnf(i_0_648, plain, (k3_xcmplx_0(k1_fdiff_1(X1,k1_fdiff_1(X2,X4)),k1_fdiff_1(k2_fdiff_1(X2,X3),X4))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X2,X3),X1),X4)|~r1_fdiff_1(X1,k1_fdiff_1(X2,X4))|~r1_fdiff_1(k2_fdiff_1(X2,X3),X4)|~v1_funct_1(X1)|~v1_funct_1(k2_fdiff_1(X2,X3))|~v1_xreal_0(k1_fdiff_1(k2_fdiff_1(X2,X3),X4))|~m1_subset_1(k1_fdiff_1(X1,k1_fdiff_1(X2,X4)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X4,k1_numbers)|~r2_hidden(X4,X3)|~r2_fdiff_1(X2,X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2610, plain, (v1_funct_1(k3_relat_1(k24_sin_cos,k24_sin_cos)))).#trainneg
cnf(i_0_2615, plain, (v1_funct_1(k3_relat_1(k19_sin_cos,k24_sin_cos)))).#trainneg
cnf(i_0_2616, plain, (v1_funct_1(k3_relat_1(k29_sin_cos,k24_sin_cos)))).#trainneg
cnf(i_0_2622, plain, (v1_funct_1(k3_relat_1(k3_taylor_1,k19_sin_cos)))).#trainneg
cnf(i_0_2623, plain, (v1_funct_1(k3_relat_1(k24_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_2003, plain, (m1_subset_1(esk2_1(X1),a_0_0_fdiff_8)|~r2_hidden(X1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_1397, negated_conjecture, (r2_fdiff_1(X1,esk3_0)|m1_subset_1(esk6_2(esk3_0,X1),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_2274, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~r1_tarski(X1,esk3_0)|~r2_fdiff_1(X2,esk3_0)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2687, plain, (m1_subset_1(k1_funct_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1503, plain, (v7_ordinal1(k1_seq_1(k19_sin_cos,X1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1428, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_651, plain, (k8_real_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_fdiff_1(X4,k1_fdiff_1(X1,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X4),X3)|~r1_fdiff_1(X4,k1_fdiff_1(X1,X3))|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X4)|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_xreal_0(k1_fdiff_1(X4,k1_fdiff_1(X1,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_numbers)|~m1_subset_1(X3,k1_numbers)|~r2_hidden(X3,X2)|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2628, plain, (v1_funct_1(k3_relat_1(k19_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_2629, plain, (v1_funct_1(k3_relat_1(k29_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_2635, plain, (v1_funct_1(k3_relat_1(k3_taylor_1,k29_sin_cos)))).#trainneg
cnf(i_0_2636, plain, (v1_funct_1(k3_relat_1(k24_sin_cos,k29_sin_cos)))).#trainneg
cnf(i_0_2641, plain, (v1_funct_1(k3_relat_1(k19_sin_cos,k29_sin_cos)))).#trainneg
cnf(i_0_1711, negated_conjecture, (v1_xcmplx_0(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v1_membered(X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1509, plain, (v7_ordinal1(k1_seq_1(k3_taylor_1,X1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1465, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_2642, plain, (v1_funct_1(k3_relat_1(k29_sin_cos,k29_sin_cos)))).#trainneg
cnf(i_0_1393, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(a_0_0_fdiff_8))|~r2_fdiff_1(k3_taylor_1,esk3_0))).#trainneg
cnf(i_0_1793, negated_conjecture, (v1_xcmplx_0(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_membered(esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1364, plain, (k1_relset_1(X1,k1_partfun1(X7,X8,X9,X10,X5,X6))=k9_xtuple_0(k1_partfun1(X7,X8,X9,X10,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_1516, plain, (v7_ordinal1(k1_seq_1(k29_sin_cos,X1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_1475, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k24_sin_cos)))).#trainneg
cnf(i_0_655, plain, (k8_real_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_fdiff_1(X4,k1_fdiff_1(X1,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X4),X3)|~r1_fdiff_1(X4,k1_fdiff_1(X1,X3))|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X4)|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_xreal_0(k1_fdiff_1(k2_fdiff_1(X1,X2),X3))|~m1_subset_1(k1_fdiff_1(X4,k1_fdiff_1(X1,X3)),k1_numbers)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers)|~r2_hidden(X3,X2)|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1589, plain, (v1_xreal_0(k1_seq_1(k24_sin_cos,X1)))).#trainneg
cnf(i_0_1499, negated_conjecture, (k7_xcmplx_0(np__1,k3_xcmplx_0(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~m1_subset_1(esk4_0,k1_numbers))).#trainneg
cnf(i_0_1765, negated_conjecture, (k7_xcmplx_0(np__1,k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers))).#trainneg
cnf(i_0_1972, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,X6)),k1_zfmisc_1(X1))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1421, plain, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|r2_hidden(k1_relset_1(k1_numbers,k19_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2383, plain, (r2_hidden(X1,k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_682, plain, (r2_fdiff_1(X1,k1_relset_1(k1_numbers,X1))|m1_subset_1(esk6_2(k1_relset_1(k1_numbers,X1),X1),k1_relset_1(k1_numbers,X1))|~v3_rcomp_1(k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1501, plain, (r2_hidden(k1_seq_1(k19_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_1507, plain, (r2_hidden(k1_seq_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_1514, plain, (r2_hidden(k1_seq_1(k29_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_1549, negated_conjecture, (v1_xreal_0(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1482, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),esk3_0)|m1_subset_1(esk6_2(esk3_0,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos)),esk3_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1342, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X6,k1_numbers)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2080, plain, (m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X1,X2))).#trainneg
cnf(i_0_1458, plain, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|r2_hidden(k1_relset_1(k1_numbers,k29_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1429, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_1402, negated_conjecture, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_1450, plain, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1518, plain, (k1_zfmisc_1(k1_numbers)=k6_numbers|r2_hidden(a_0_0_fdiff_8,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_2673, plain, (m1_subset_1(k3_relat_1(k3_taylor_1,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2674, plain, (m1_subset_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2680, plain, (m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1546, negated_conjecture, (r2_hidden(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1860, plain, (r1_tarski(k1_partfun1(X1,X2,X3,X4,X5,X6),k2_zfmisc_1(X1,X4))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2731, plain, (v7_ordinal1(X1)|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1466, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_761, plain, (r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|m1_subset_1(esk6_2(X3,k2_fdiff_1(X1,X2)),k1_numbers)|~v3_rcomp_1(X3)|~r1_tarski(X3,X2)|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2282, plain, (r2_hidden(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1586, plain, (r2_hidden(k1_seq_1(k24_sin_cos,X1),k1_numbers))).#trainneg
cnf(i_0_2681, plain, (m1_subset_1(k3_relat_1(k29_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2840, plain, (m1_subset_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2841, plain, (m1_subset_1(k3_relat_1(k24_sin_cos,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2847, plain, (m1_subset_1(k3_relat_1(k19_sin_cos,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2848, plain, (m1_subset_1(k3_relat_1(k29_sin_cos,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1343, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X6,k1_numbers)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2092, plain, (v1_xreal_0(k3_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1950, plain, (r1_tarski(X1,k9_xtuple_0(X2))|~v1_funct_1(X7)|~v1_funct_1(X2)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k9_xtuple_0(k1_partfun1(X3,X4,X5,X6,X2,X7)))))).#trainneg
cnf(i_0_2315, plain, (m1_subset_1(k3_square_1(X1),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1337, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(X6,k1_numbers)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1476, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k24_sin_cos)))).#trainneg
cnf(i_0_794, plain, (r2_hidden(k1_seq_1(X1,esk6_2(k9_xtuple_0(k3_relat_1(X1,X2)),X3)),k9_xtuple_0(X2))|r2_fdiff_1(X3,k9_xtuple_0(k3_relat_1(X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v3_rcomp_1(k9_xtuple_0(k3_relat_1(X1,X2)))|~r1_tarski(k9_xtuple_0(k3_relat_1(X1,X2)),k1_relset_1(k1_numbers,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k9_xtuple_0(k3_relat_1(X1,X2)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1775, plain, (r1_tarski(k3_limfunc1(X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1728, plain, (v1_xreal_0(X1)|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1588, plain, (v7_ordinal1(k1_seq_1(k24_sin_cos,X1))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_3003, plain, (m1_subset_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3004, plain, (m1_subset_1(k3_relat_1(k24_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3010, plain, (m1_subset_1(k3_relat_1(k19_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3011, plain, (m1_subset_1(k3_relat_1(k29_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3270, plain, (m1_subset_1(k3_relat_1(k3_taylor_1,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2437, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,X5,X6)),k1_zfmisc_1(k9_xtuple_0(X5)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2264, plain, (r1_tarski(k2_fdiff_1(X1,X2),k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_2514, plain, (k7_xcmplx_0(X1,k5_xcmplx_0(X1))=k5_square_1(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1662, plain, (m1_subset_1(k5_xcmplx_0(X1),k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2011, plain, (r2_fdiff_1(k3_taylor_1,X1)|~v3_rcomp_1(X1)|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_797, plain, (r2_hidden(k1_seq_1(X1,k1_funct_1(X2,X3)),k9_xtuple_0(X4))|~v1_funct_1(X4)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X4)|~v1_relat_1(X1)|~r2_hidden(X3,k9_xtuple_0(k3_relat_1(X2,k3_relat_1(X1,X4))))|~v1_funct_1(k3_relat_1(X1,X4))|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1439, plain, (k3_xcmplx_0(np__1,X1)=k7_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3271, plain, (m1_subset_1(k3_relat_1(k24_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3277, plain, (m1_subset_1(k3_relat_1(k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3278, plain, (m1_subset_1(k3_relat_1(k29_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1395, negated_conjecture, (r2_fdiff_1(k3_taylor_1,esk3_0)|m1_subset_1(esk6_2(esk3_0,k3_taylor_1),k1_numbers)|~r1_tarski(esk3_0,a_0_0_fdiff_8))).#trainneg
cnf(i_0_1552, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k3_xcmplx_0(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~m2_subset_1(k3_xcmplx_0(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))),k1_numbers,a_0_0_fdiff_8)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers)|~v1_xcmplx_0(esk4_0)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))).#trainneg
cnf(i_0_1939, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(X1,k19_sin_cos))).#trainneg
cnf(i_0_1378, plain, (r1_tarski(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,X5)),k1_relset_1(k1_numbers,k19_sin_cos))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1776, plain, (r1_tarski(X1,k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k3_limfunc1(X2))))).#trainneg
cnf(i_0_859, plain, (r2_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v3_rcomp_1(X3)|~r1_tarski(X3,X2)|~v1_xboole_0(X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1468, plain, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|r2_hidden(k1_relset_1(k1_numbers,k24_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1638, plain, (m1_subset_1(esk6_2(k5_numbers,k3_taylor_1),k1_numbers)|r2_fdiff_1(k3_taylor_1,k5_numbers)|~v3_rcomp_1(k5_numbers)|~r1_tarski(k5_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_1781, plain, (k1_fdiff_1(k3_taylor_1,np__1)=k7_xcmplx_0(np__1,np__1)|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(np__1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_1782, plain, (k1_fdiff_1(k3_taylor_1,k6_numbers)=k7_xcmplx_0(np__1,k6_numbers)|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_1555, negated_conjecture, (m1_subset_1(esk4_0,k1_numbers)|~r2_fdiff_1(k3_relat_1(k3_taylor_1,k29_sin_cos),esk3_0)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1664, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2833, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(X1,k19_sin_cos)),k9_xtuple_0(X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1416, plain, (v1_xboole_0(a_0_0_fdiff_8)|m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_3713, plain, (k7_xcmplx_0(k3_square_1(X1),X2)=k7_xcmplx_0(X1,k7_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2375, plain, (r2_hidden(X1,k1_numbers)|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_887, plain, (r2_hidden(esk6_2(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,X6)),X7),k9_xtuple_0(X5))|r2_fdiff_1(X7,k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,X6)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_rcomp_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,X6)))|~r1_tarski(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,X6)),k1_relset_1(k1_numbers,X7))|~v1_funct_1(X7)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,X6)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1525, negated_conjecture, (k10_real_1(np__1,k3_xcmplx_0(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers)|~v1_xcmplx_0(esk4_0)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))).#trainneg
cnf(i_0_1420, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_8)|k1_fdiff_1(k3_taylor_1,k3_xcmplx_0(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~m1_subset_1(k3_xcmplx_0(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))),a_0_0_fdiff_8))).#trainneg
cnf(i_0_2044, plain, (m1_subset_1(k1_relset_1(X1,k2_zfmisc_1(X1,X2)),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1529, plain, (~v1_xboole_0(a_0_0_fdiff_8)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_1533, plain, (r2_hidden(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_1976, negated_conjecture, (v1_xreal_0(X1)|~v3_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_2413, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos))))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1375, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k19_sin_cos,X1)),k1_relset_1(k1_numbers,k19_sin_cos))|~v1_relat_1(X1))).#trainneg
cnf(i_0_1974, negated_conjecture, (v1_xcmplx_0(X1)|~v1_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_1966, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,k1_partfun1(X2,X3,X4,X5,X6,X7)))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2000, plain, (k7_xcmplx_0(np__1,k1_fdiff_1(k29_sin_cos,X1))=k5_square_1(k1_seq_1(k19_sin_cos,X1))|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,X1)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1935, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(X1,k3_taylor_1))).#trainneg
cnf(i_0_890, plain, (r2_hidden(k1_funct_1(X1,X2),k9_xtuple_0(X3))|~v1_funct_1(X8)|~v1_funct_1(X3)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~r2_hidden(X2,k9_xtuple_0(k3_relat_1(X1,k1_partfun1(X4,X5,X6,X7,X3,X8))))|~v1_funct_1(k1_partfun1(X4,X5,X6,X7,X3,X8))|~v1_funct_1(X1)|~v1_relat_1(k1_partfun1(X4,X5,X6,X7,X3,X8))|~v1_relat_1(X1))).#trainneg
cnf(i_0_1660, plain, (k1_zfmisc_1(X1)=k6_numbers|r2_hidden(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1432, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_8)|k1_fdiff_1(k3_taylor_1,k3_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),esk4_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~m1_subset_1(k3_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),esk4_0),a_0_0_fdiff_8)|~v1_xreal_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers))).#trainneg
cnf(i_0_1686, negated_conjecture, (v1_int_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~v6_membered(k1_numbers)|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1794, negated_conjecture, (v7_ordinal1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v6_membered(esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1338, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X6,k1_numbers)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2233, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|v1_xboole_0(esk3_0)|m1_subset_1(esk6_2(esk3_0,X1),k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_3109, plain, (v1_xboole_0(X1)|r2_hidden(esk6_2(X1,k2_fdiff_1(X2,X3)),X1)|r2_fdiff_1(k2_fdiff_1(X2,X3),X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,X3)|~r2_fdiff_1(X2,X3)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_906, plain, (r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,k3_relat_1(X3,X4))))|~r2_hidden(X1,k9_xtuple_0(X2))|~v1_funct_1(k3_relat_1(X3,X4))|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(k1_seq_1(X3,k1_funct_1(X2,X1)),k9_xtuple_0(X4))|~r2_hidden(k1_funct_1(X2,X1),k9_xtuple_0(X3))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v3_valued_0(X3)|~v1_relat_1(X4)|~v1_relat_1(X3))).#trainneg
cnf(i_0_4237, plain, (r1_tarski(a_0_0_fdiff_8,k1_numbers))).#trainneg
cnf(i_0_4723, plain, (v3_valued_0(k3_relat_1(k3_taylor_1,k3_taylor_1)))).#trainneg
cnf(i_0_4131, plain, (r1_fdiff_1(X1,esk6_2(X2,X3))|r2_fdiff_1(X3,X2)|~v3_rcomp_1(X2)|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X3))))).#trainneg
cnf(i_0_4136, negated_conjecture, (k7_xcmplx_0(np__1,k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xcmplx_0(k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))|~m1_subset_1(k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))),k1_numbers))).#trainneg
cnf(i_0_4692, plain, (r1_tarski(k1_relset_1(k1_numbers,k19_sin_cos),k1_numbers))).#trainneg
cnf(i_0_4696, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_909, plain, (r1_fdiff_1(X1,X2)|~v3_rcomp_1(k9_xtuple_0(k3_relat_1(X3,X4)))|~r2_fdiff_1(X1,k9_xtuple_0(k3_relat_1(X3,X4)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k9_xtuple_0(k3_relat_1(X3,X4)),k1_zfmisc_1(k1_numbers))|~r2_hidden(k1_seq_1(X3,X2),k9_xtuple_0(X4))|~r2_hidden(X2,k9_xtuple_0(X3))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v3_valued_0(X3)|~v1_relat_1(X4)|~v1_relat_1(X3))).#trainneg
cnf(i_0_4757, plain, (v3_valued_0(k3_relat_1(k24_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_4791, plain, (v3_valued_0(k3_relat_1(k19_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_4883, plain, (v3_valued_0(k3_relat_1(k29_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_4917, plain, (v3_valued_0(k3_relat_1(k3_taylor_1,k29_sin_cos)))).#trainneg
cnf(i_0_4951, plain, (v3_valued_0(k3_relat_1(k24_sin_cos,k29_sin_cos)))).#trainneg
cnf(i_0_4300, plain, (r1_tarski(X1,X2)|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(X2,k1_partfun1(X2,X3,X4,X5,X6,X7)))))).#trainneg
cnf(i_0_4827, plain, (r1_tarski(k1_relset_1(k1_numbers,k29_sin_cos),k1_numbers))).#trainneg
cnf(i_0_4698, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_912, plain, (r2_hidden(X1,k9_xtuple_0(k1_partfun1(X4,X5,X6,X7,X2,X3)))|~r2_hidden(k1_seq_1(X2,X1),k9_xtuple_0(X3))|~r2_hidden(X1,k9_xtuple_0(X2))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v3_valued_0(X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_4985, plain, (v3_valued_0(k3_relat_1(k19_sin_cos,k29_sin_cos)))).#trainneg
cnf(i_0_5019, plain, (v3_valued_0(k3_relat_1(k29_sin_cos,k29_sin_cos)))).#trainneg
cnf(i_0_5173, plain, (v3_valued_0(k3_relat_1(k3_taylor_1,k19_sin_cos)))).#trainneg
cnf(i_0_5209, plain, (v3_valued_0(k3_relat_1(k24_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_5245, plain, (v3_valued_0(k3_relat_1(k19_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_5281, plain, (v3_valued_0(k3_relat_1(k29_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_5317, plain, (v3_valued_0(k3_relat_1(k3_taylor_1,k24_sin_cos)))).#trainneg
cnf(i_0_5482, plain, (v3_valued_0(k3_relat_1(k24_sin_cos,k24_sin_cos)))).#trainneg
cnf(i_0_5408, plain, (v1_xreal_0(k5_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5995, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_hidden(esk6_2(esk3_0,X1),k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|v1_xboole_0(esk3_0)|r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_4224, plain, (k1_fdiff_1(k3_taylor_1,k1_seq_1(k19_sin_cos,X1))=k7_xcmplx_0(np__1,k1_seq_1(k19_sin_cos,X1))|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k1_seq_1(k19_sin_cos,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_4234, plain, (k1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))=k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1))|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_969, plain, (v1_xboole_0(k9_xtuple_0(k3_relat_1(X1,X2)))|v1_xboole_0(k9_xtuple_0(X1))|m1_subset_1(X3,k9_xtuple_0(X1))|~m1_subset_1(X3,k9_xtuple_0(k3_relat_1(X1,X2)))|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5518, plain, (v3_valued_0(k3_relat_1(k19_sin_cos,k24_sin_cos)))).#trainneg
cnf(i_0_5554, plain, (v3_valued_0(k3_relat_1(k29_sin_cos,k24_sin_cos)))).#trainneg
cnf(i_0_5626, plain, (v1_xboole_0(a_0_0_fdiff_8)|~r1_xxreal_0(np__1,k6_numbers)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(np__1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_5414, negated_conjecture, (r2_fdiff_1(k3_taylor_1,esk3_0)|~v1_xboole_0(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_5653, plain, (v1_xboole_0(a_0_0_fdiff_8)|~r1_xxreal_0(k7_xcmplx_0(np__1,k6_numbers),k6_numbers)|~m2_subset_1(k6_numbers,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_4714, plain, (v4_relat_1(k3_relat_1(k3_taylor_1,k3_taylor_1),k1_numbers))).#trainneg
cnf(i_0_4724, plain, (v5_relat_1(k3_relat_1(k3_taylor_1,k3_taylor_1),k1_numbers))).#trainneg
cnf(i_0_4139, plain, (k1_fdiff_1(k3_taylor_1,X1)=k7_xcmplx_0(np__1,X1)|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(X1,a_0_0_fdiff_8)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4423, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~r1_tarski(X1,esk3_0)|~r2_fdiff_1(X2,esk3_0)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4428, negated_conjecture, (r2_fdiff_1(X1,X2)|m1_subset_1(esk6_2(X2,X1),X2)|~v3_rcomp_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r1_tarski(X2,esk3_0)|~r2_fdiff_1(X1,esk3_0))).#trainneg
cnf(i_0_4419, negated_conjecture, (v1_xreal_0(esk6_2(esk3_0,X1))|r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_981, plain, (m1_subset_1(k1_funct_1(X1,esk6_2(k9_xtuple_0(k3_relat_1(X1,X2)),X3)),k9_xtuple_0(X2))|r2_fdiff_1(X3,k9_xtuple_0(k3_relat_1(X1,X2)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v3_rcomp_1(k9_xtuple_0(k3_relat_1(X1,X2)))|~r1_tarski(k9_xtuple_0(k3_relat_1(X1,X2)),k1_relset_1(k1_numbers,X3))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k9_xtuple_0(k3_relat_1(X1,X2)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5115, plain, (r1_tarski(k1_relset_1(k1_numbers,k24_sin_cos),k1_numbers))).#trainneg
cnf(i_0_4238, plain, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_4748, plain, (v4_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k1_numbers))).#trainneg
cnf(i_0_4758, plain, (v5_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k1_numbers))).#trainneg
cnf(i_0_4782, plain, (v4_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_numbers))).#trainneg
cnf(i_0_4792, plain, (v5_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_numbers))).#trainneg
cnf(i_0_4874, plain, (v4_relat_1(k3_relat_1(k29_sin_cos,k3_taylor_1),k1_numbers))).#trainneg
cnf(i_0_4640, plain, (r2_fdiff_1(k2_fdiff_1(X1,X2),X2)|m1_subset_1(esk6_2(X2,k2_fdiff_1(X1,X2)),X2)|~v3_rcomp_1(X2)|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_983, plain, (m1_subset_1(k1_funct_1(X1,esk1_3(X2,k9_xtuple_0(k3_relat_1(X1,X3)),X4)),k9_xtuple_0(X3))|X4=k2_fdiff_1(X2,k9_xtuple_0(k3_relat_1(X1,X3)))|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X1)|k1_relset_1(k1_numbers,X4)!=k9_xtuple_0(k3_relat_1(X1,X3))|~r2_fdiff_1(X2,k9_xtuple_0(k3_relat_1(X1,X3)))|~v1_funct_1(X4)|~v1_funct_1(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4884, plain, (v5_relat_1(k3_relat_1(k29_sin_cos,k3_taylor_1),k1_numbers))).#trainneg
cnf(i_0_4908, plain, (v4_relat_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k1_numbers))).#trainneg
cnf(i_0_4918, plain, (v5_relat_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k1_numbers))).#trainneg
cnf(i_0_4942, plain, (v4_relat_1(k3_relat_1(k24_sin_cos,k29_sin_cos),k1_numbers))).#trainneg
cnf(i_0_4952, plain, (v5_relat_1(k3_relat_1(k24_sin_cos,k29_sin_cos),k1_numbers))).#trainneg
cnf(i_0_4976, plain, (v4_relat_1(k3_relat_1(k19_sin_cos,k29_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5409, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4296, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k9_xtuple_0(k1_partfun1(X2,X3,X4,X5,X6,X7)))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_4697, negated_conjecture, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_984, plain, (m1_subset_1(k1_funct_1(X1,k1_funct_1(X2,X3)),k9_xtuple_0(X4))|~v1_funct_1(X4)|~v1_funct_1(X1)|~v1_relat_1(X4)|~v1_relat_1(X1)|~r2_hidden(X3,k9_xtuple_0(k3_relat_1(X2,k3_relat_1(X1,X4))))|~v1_funct_1(k3_relat_1(X1,X4))|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_4986, plain, (v5_relat_1(k3_relat_1(k19_sin_cos,k29_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5010, plain, (v4_relat_1(k3_relat_1(k29_sin_cos,k29_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5020, plain, (v5_relat_1(k3_relat_1(k29_sin_cos,k29_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5164, plain, (v4_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5174, plain, (v5_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5200, plain, (v4_relat_1(k3_relat_1(k24_sin_cos,k19_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5210, plain, (v5_relat_1(k3_relat_1(k24_sin_cos,k19_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5236, plain, (v4_relat_1(k3_relat_1(k19_sin_cos,k19_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5899, plain, (~v1_xboole_0(a_0_0_fdiff_8)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4424, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~r1_tarski(X1,X3)|~r1_tarski(X3,esk3_0)|~r2_fdiff_1(X2,esk3_0)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_5894, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(X2,k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_4380, plain, (k1_fdiff_1(k3_taylor_1,k1_seq_1(k29_sin_cos,X1))=k7_xcmplx_0(np__1,k1_seq_1(k29_sin_cos,X1))|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k1_seq_1(k29_sin_cos,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_4328, negated_conjecture, (r1_fdiff_1(X1,esk6_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X2)|~v3_rcomp_1(X2)|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r1_tarski(X2,esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4477, plain, (r1_tarski(X1,k9_xtuple_0(X2))|~v1_funct_1(X7)|~v1_funct_1(X2)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(X3,k1_partfun1(X3,X4,X5,X6,X2,X7)))))).#trainneg
cnf(i_0_6033, plain, (r1_fdiff_1(X1,esk6_2(X2,k2_fdiff_1(X3,X4)))|r2_fdiff_1(k2_fdiff_1(X3,X4),X2)|~v3_rcomp_1(X2)|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r1_tarski(X2,X4)|~r2_fdiff_1(X3,X4)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4699, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_986, plain, (m1_subset_1(k1_seq_1(k1_partfun1(X1,X2,X3,X4,X5,X6),X7),k1_numbers)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~v3_membered(X4)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5246, plain, (v5_relat_1(k3_relat_1(k19_sin_cos,k19_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5272, plain, (v4_relat_1(k3_relat_1(k29_sin_cos,k19_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5282, plain, (v5_relat_1(k3_relat_1(k29_sin_cos,k19_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5308, plain, (v4_relat_1(k3_relat_1(k3_taylor_1,k24_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5318, plain, (v5_relat_1(k3_relat_1(k3_taylor_1,k24_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5473, plain, (v4_relat_1(k3_relat_1(k24_sin_cos,k24_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5483, plain, (v5_relat_1(k3_relat_1(k24_sin_cos,k24_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5509, plain, (v4_relat_1(k3_relat_1(k19_sin_cos,k24_sin_cos),k1_numbers))).#trainneg
cnf(i_0_4299, plain, (r1_tarski(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,X5,X6)),X1)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4259, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4497, negated_conjecture, (r1_fdiff_1(X1,X2)|~v3_rcomp_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_fdiff_1(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_zfmisc_1(k1_numbers))|~r2_hidden(X2,esk3_0))).#trainneg
cnf(i_0_5919, negated_conjecture, (v1_xcmplx_0(X1)|~r2_hidden(X1,esk3_0)|~v3_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_4499, negated_conjecture, (r2_hidden(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,esk3_0)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_4523, plain, (k1_relset_1(X1,k1_partfun1(X2,X3,X4,X5,X6,k19_sin_cos))=k9_xtuple_0(k1_partfun1(X2,X3,X4,X5,X6,k19_sin_cos))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,X7)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_4467, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1002, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(k7_xcmplx_0(np__1,X2),X3))=k7_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_4712, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k3_taylor_1,k3_taylor_1)))).#trainneg
cnf(i_0_4780, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k19_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_4872, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k29_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_5519, plain, (v5_relat_1(k3_relat_1(k19_sin_cos,k24_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5545, plain, (v4_relat_1(k3_relat_1(k29_sin_cos,k24_sin_cos),k1_numbers))).#trainneg
cnf(i_0_5555, plain, (v5_relat_1(k3_relat_1(k29_sin_cos,k24_sin_cos),k1_numbers))).#trainneg
cnf(i_0_4749, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k24_sin_cos)))).#trainneg
cnf(i_0_4750, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k3_taylor_1)))).#trainneg
cnf(i_0_4516, plain, (k1_relset_1(X1,k1_partfun1(X2,X3,X4,X5,X6,k3_taylor_1))=k9_xtuple_0(k1_partfun1(X2,X3,X4,X5,X6,k3_taylor_1))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,X7)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_4736, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5911, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(k19_sin_cos,X2)))|~v1_relat_1(X2))).#trainneg
cnf(i_0_4524, plain, (k1_relset_1(X1,k1_partfun1(X2,X3,X4,X5,X6,k29_sin_cos))=k9_xtuple_0(k1_partfun1(X2,X3,X4,X5,X6,k29_sin_cos))|~v1_funct_1(X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,X7)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_4409, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4478, plain, (r1_tarski(X1,k9_xtuple_0(X2))|~v1_funct_1(X6)|~v1_funct_1(X2)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,X3)))|~v3_rcomp_1(X1)|~r2_fdiff_1(k1_partfun1(k1_numbers,X3,X4,X5,X2,X6),X1)|~v1_funct_1(k1_partfun1(k1_numbers,X3,X4,X5,X2,X6))|~m1_subset_1(k1_partfun1(k1_numbers,X3,X4,X5,X2,X6),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4930, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1012, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(k7_xcmplx_0(np__1,X3),X2))=k7_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_4906, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k3_taylor_1,k29_sin_cos)))).#trainneg
cnf(i_0_4974, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k19_sin_cos,k29_sin_cos)))).#trainneg
cnf(i_0_5008, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k29_sin_cos,k29_sin_cos)))).#trainneg
cnf(i_0_4751, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k19_sin_cos)))).#trainneg
cnf(i_0_4752, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k29_sin_cos)))).#trainneg
cnf(i_0_4783, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k24_sin_cos)))).#trainneg
cnf(i_0_4784, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k3_taylor_1)))).#trainneg
cnf(i_0_4785, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k19_sin_cos)))).#trainneg
cnf(i_0_4416, negated_conjecture, (r2_hidden(esk6_2(esk3_0,X1),k1_numbers)|r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_4683, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4676, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4594, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X7,X8,X9,X10,X5,X6)),k1_zfmisc_1(X1))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_1615, negated_conjecture, (v3_membered(X1)|~v3_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r1_tarski(X1,esk3_0))).#trainneg
cnf(i_0_4427, negated_conjecture, (r2_fdiff_1(X1,X2)|~v3_rcomp_1(X2)|~v1_xboole_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r1_tarski(X2,esk3_0)|~r2_fdiff_1(X1,esk3_0))).#trainneg
cnf(i_0_1016, plain, (k7_xcmplx_0(k3_xcmplx_0(X1,X2),k3_xcmplx_0(X3,X4))=k3_xcmplx_0(X1,k7_xcmplx_0(k7_xcmplx_0(X2,X4),X3))|~v1_xcmplx_0(k3_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4))).#trainneg
cnf(i_0_5162, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k3_taylor_1,k19_sin_cos)))).#trainneg
cnf(i_0_5234, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k19_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_5270, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k29_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_4786, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k29_sin_cos)))).#trainneg
cnf(i_0_4875, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k29_sin_cos,k3_taylor_1),k24_sin_cos)))).#trainneg
cnf(i_0_4876, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k29_sin_cos,k3_taylor_1),k3_taylor_1)))).#trainneg
cnf(i_0_4877, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k29_sin_cos,k3_taylor_1),k19_sin_cos)))).#trainneg
cnf(i_0_4878, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k29_sin_cos,k3_taylor_1),k29_sin_cos)))).#trainneg
cnf(i_0_4465, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4804, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5395, plain, (k7_xcmplx_0(k5_square_1(X1),X2)=k7_xcmplx_0(X1,k7_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4432, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_membered(a_0_0_fdiff_8)|~r2_hidden(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4261, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4431, plain, (v1_xboole_0(a_0_0_fdiff_8)|r2_hidden(k1_funct_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1975, negated_conjecture, (v7_ordinal1(X1)|~v6_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_1021, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(k3_xcmplx_0(X2,X3),X4))=k7_xcmplx_0(k7_xcmplx_0(X1,k7_xcmplx_0(X3,X4)),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k7_xcmplx_0(X3,X4))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_4746, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k24_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_4940, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k24_sin_cos,k29_sin_cos)))).#trainneg
cnf(i_0_4909, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k24_sin_cos)))).#trainneg
cnf(i_0_4910, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k3_taylor_1)))).#trainneg
cnf(i_0_4911, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k19_sin_cos)))).#trainneg
cnf(i_0_4912, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k29_sin_cos)))).#trainneg
cnf(i_0_4943, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k24_sin_cos,k29_sin_cos),k24_sin_cos)))).#trainneg
cnf(i_0_4998, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5032, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4469, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5910, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,X1)),k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_4133, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_4374, plain, (k7_xcmplx_0(k7_xcmplx_0(X1,np__1),X2)=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5682, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k7_xcmplx_0(np__1,X1))|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_1612, negated_conjecture, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|~r1_tarski(X1,esk3_0))).#trainneg
cnf(i_0_4693, plain, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos))))).#trainneg
cnf(i_0_1067, plain, (k7_xcmplx_0(X1,k3_square_1(k5_xcmplx_0(X2)))=k3_xcmplx_0(k7_xcmplx_0(X1,k5_xcmplx_0(X2)),X2)|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k7_xcmplx_0(X1,k5_xcmplx_0(X2)))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_5198, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k24_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_5306, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k3_taylor_1,k24_sin_cos)))).#trainneg
cnf(i_0_5507, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k19_sin_cos,k24_sin_cos)))).#trainneg
cnf(i_0_4944, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k24_sin_cos,k29_sin_cos),k3_taylor_1)))).#trainneg
cnf(i_0_4945, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k24_sin_cos,k29_sin_cos),k19_sin_cos)))).#trainneg
cnf(i_0_4946, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k24_sin_cos,k29_sin_cos),k29_sin_cos)))).#trainneg
cnf(i_0_4977, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k19_sin_cos,k29_sin_cos),k24_sin_cos)))).#trainneg
cnf(i_0_4978, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k19_sin_cos,k29_sin_cos),k3_taylor_1)))).#trainneg
cnf(i_0_5151, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_5413, plain, (m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3222, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,k19_sin_cos)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_3531, plain, (k7_xcmplx_0(k3_square_1(X1),k5_xcmplx_0(X2))=k3_xcmplx_0(X1,k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1943, plain, (r1_tarski(X1,k9_xtuple_0(X2))|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(X2,X7)))|~v1_funct_1(X7)|~v1_funct_1(X2)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_1613, negated_conjecture, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|~r1_tarski(X1,esk3_0))).#trainneg
cnf(i_0_4828, plain, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k29_sin_cos))))).#trainneg
cnf(i_0_1069, plain, (k7_xcmplx_0(X1,k3_square_1(k5_xcmplx_0(X2)))=k7_xcmplx_0(k3_xcmplx_0(X1,X2),k5_xcmplx_0(X2))|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_5543, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k29_sin_cos,k24_sin_cos)))).#trainneg
cnf(i_0_5471, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k24_sin_cos,k24_sin_cos)))).#trainneg
cnf(i_0_4979, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k19_sin_cos,k29_sin_cos),k19_sin_cos)))).#trainneg
cnf(i_0_4980, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k19_sin_cos,k29_sin_cos),k29_sin_cos)))).#trainneg
cnf(i_0_5165, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k24_sin_cos)))).#trainneg
cnf(i_0_5166, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k3_taylor_1)))).#trainneg
cnf(i_0_5167, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k19_sin_cos)))).#trainneg
cnf(i_0_1740, plain, (r2_hidden(esk6_2(X1,k2_fdiff_1(X2,X1)),k1_numbers)|r2_fdiff_1(k2_fdiff_1(X2,X1),X1)|~v3_rcomp_1(X1)|~r2_fdiff_1(X2,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4508, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k29_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4434, plain, (v1_xreal_0(k1_funct_1(X1,X2))|~r2_hidden(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2232, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|v1_xboole_0(esk3_0)|m1_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4896, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4684, plain, (m1_subset_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1614, negated_conjecture, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|~r1_tarski(X1,esk3_0))).#trainneg
cnf(i_0_5116, plain, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k24_sin_cos))))).#trainneg
cnf(i_0_1080, plain, (k7_xcmplx_0(k7_xcmplx_0(k5_xcmplx_0(k3_square_1(X1)),X1),X1)=k3_square_1(k5_xcmplx_0(k3_square_1(X1)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(k3_square_1(X1)))|~v1_xcmplx_0(k3_square_1(X1)))).#trainneg
cnf(i_0_5168, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k29_sin_cos)))).#trainneg
cnf(i_0_5201, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k24_sin_cos,k19_sin_cos),k24_sin_cos)))).#trainneg
cnf(i_0_5202, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k24_sin_cos,k19_sin_cos),k3_taylor_1)))).#trainneg
cnf(i_0_5203, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k24_sin_cos,k19_sin_cos),k19_sin_cos)))).#trainneg
cnf(i_0_5204, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k24_sin_cos,k19_sin_cos),k29_sin_cos)))).#trainneg
cnf(i_0_2045, plain, (k1_relset_1(X1,k2_zfmisc_1(X1,X2))=k9_xtuple_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_4262, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k3_taylor_1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3790, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(np__1,X1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1861, plain, (r1_tarski(X1,k2_zfmisc_1(X2,X3))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X3)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k1_partfun1(X2,X4,X5,X3,X6,X7))))).#trainneg
cnf(i_0_2665, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(X2,k3_rfunct_1(X2,X3,X4,X5)))|~v1_funct_1(X5)|~v1_funct_1(X4)|~v3_membered(X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2211, plain, (m1_subset_1(k7_xcmplx_0(X1,np__1),k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1427, plain, (v1_xboole_0(k1_relset_1(k1_numbers,k19_sin_cos))|m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_5893, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_relset_1(X2,k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1110, plain, (k7_xcmplx_0(k5_xcmplx_0(X1),X2)=k7_xcmplx_0(k7_xcmplx_0(np__1,X2),X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(k7_xcmplx_0(np__1,X2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5273, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k29_sin_cos,k19_sin_cos),k24_sin_cos)))).#trainneg
cnf(i_0_5274, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k29_sin_cos,k19_sin_cos),k3_taylor_1)))).#trainneg
cnf(i_0_5275, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k29_sin_cos,k19_sin_cos),k19_sin_cos)))).#trainneg
cnf(i_0_5276, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k29_sin_cos,k19_sin_cos),k29_sin_cos)))).#trainneg
cnf(i_0_5309, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_taylor_1,k24_sin_cos),k24_sin_cos)))).#trainneg
cnf(i_0_4553, plain, (k1_fdiff_1(k3_taylor_1,k1_seq_1(k24_sin_cos,X1))=k7_xcmplx_0(np__1,k1_seq_1(k24_sin_cos,X1))|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k1_seq_1(k24_sin_cos,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_1385, plain, (r1_tarski(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,X5)),k1_relset_1(k1_numbers,k29_sin_cos))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2692, plain, (m1_subset_1(k1_funct_1(X1,X2),k1_relset_1(k1_numbers,k19_sin_cos))|~r2_hidden(k1_seq_1(X1,X2),k1_relset_1(k1_numbers,k19_sin_cos))|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3404, plain, (v4_relat_1(k3_relat_1(X1,k19_sin_cos),X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_3302, plain, (k7_xcmplx_0(np__1,k3_square_1(k5_xcmplx_0(X1)))=k3_square_1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1464, plain, (v1_xboole_0(k1_relset_1(k1_numbers,k29_sin_cos))|m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_1132, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(k5_xcmplx_0(X1),X2))=k7_xcmplx_0(X1,k5_xcmplx_0(X2))|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5310, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_taylor_1,k24_sin_cos),k3_taylor_1)))).#trainneg
cnf(i_0_5311, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_taylor_1,k24_sin_cos),k19_sin_cos)))).#trainneg
cnf(i_0_5312, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_taylor_1,k24_sin_cos),k29_sin_cos)))).#trainneg
cnf(i_0_5510, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k19_sin_cos,k24_sin_cos),k24_sin_cos)))).#trainneg
cnf(i_0_5511, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k19_sin_cos,k24_sin_cos),k3_taylor_1)))).#trainneg
cnf(i_0_4466, plain, (v1_funct_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3502, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5404, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(np__1,X1))=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2441, plain, (v1_xreal_0(k1_fdiff_1(k29_sin_cos,X1))|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1392, plain, (r1_fdiff_1(X1,k6_numbers)|~v3_rcomp_1(k1_numbers)|~r2_fdiff_1(X1,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_3694, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(X2,X1))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6073, negated_conjecture, (m1_subset_1(X1,a_0_0_fdiff_8)|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_1395, negated_conjecture, (r2_fdiff_1(k3_taylor_1,esk3_0)|m1_subset_1(esk6_2(esk3_0,k3_taylor_1),k1_numbers))).#trainneg
cnf(i_0_6060, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_5512, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k19_sin_cos,k24_sin_cos),k19_sin_cos)))).#trainneg
cnf(i_0_5513, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k19_sin_cos,k24_sin_cos),k29_sin_cos)))).#trainneg
cnf(i_0_5546, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k29_sin_cos,k24_sin_cos),k24_sin_cos)))).#trainneg
cnf(i_0_5547, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k29_sin_cos,k24_sin_cos),k3_taylor_1)))).#trainneg
cnf(i_0_6071, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~v3_rcomp_1(X2)|~r1_fdiff_1(X1,esk6_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1)))|~r1_tarski(X2,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r2_hidden(k1_seq_1(X1,esk6_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1))),esk3_0))).#trainneg
cnf(i_0_6200, negated_conjecture, (~v1_xboole_0(k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_6432, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_1136, plain, (k7_xcmplx_0(k3_xcmplx_0(X1,X2),k5_xcmplx_0(X3))=k7_xcmplx_0(X1,k7_xcmplx_0(k5_xcmplx_0(X2),X3))|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_6076, negated_conjecture, (r2_hidden(X1,k1_numbers)|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_5548, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k29_sin_cos,k24_sin_cos),k19_sin_cos)))).#trainneg
cnf(i_0_5549, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k29_sin_cos,k24_sin_cos),k29_sin_cos)))).#trainneg
cnf(i_0_5690, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_taylor_1,k19_sin_cos)),a_0_0_fdiff_8))).#trainneg
cnf(i_0_6290, plain, (v1_xboole_0(a_0_0_fdiff_8)|r1_fdiff_1(k3_taylor_1,np__1)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(np__1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_8737, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k9_xtuple_0(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_6406, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(esk6_2(X1,X2),X1)|r2_fdiff_1(X2,X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~r2_fdiff_1(X2,esk3_0)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6067, negated_conjecture, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,X2)|~r1_tarski(X2,esk3_0))).#trainneg
cnf(i_0_6433, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_1172, plain, (m1_subset_1(esk6_2(k3_rfunct_1(X1,X2,X3,X4),X5),k2_zfmisc_1(X1,k1_numbers))|r2_fdiff_1(X5,k3_rfunct_1(X1,X2,X3,X4))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v3_membered(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_rcomp_1(k3_rfunct_1(X1,X2,X3,X4))|~r1_tarski(k3_rfunct_1(X1,X2,X3,X4),k1_relset_1(k1_numbers,X5))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_rfunct_1(X1,X2,X3,X4),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5637, plain, (v1_xreal_0(k7_xcmplx_0(np__1,k6_numbers))|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_5635, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,k6_numbers))|v1_xboole_0(a_0_0_fdiff_8)|~v1_membered(k1_numbers)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_5636, plain, (v7_ordinal1(k7_xcmplx_0(np__1,k6_numbers))|v1_xboole_0(a_0_0_fdiff_8)|~v6_membered(k1_numbers)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_4187, plain, (r2_fdiff_1(k29_sin_cos,k1_relset_1(k1_numbers,k29_sin_cos))|~v3_rcomp_1(k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(esk6_2(k1_relset_1(k1_numbers,k29_sin_cos),k29_sin_cos),k1_numbers))).#trainneg
cnf(i_0_6533, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_6587, plain, (r1_tarski(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,X6)),X1)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_8796, plain, (r1_tarski(X1,X2)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v3_membered(X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(X2,k3_rfunct_1(X2,X3,X4,X5)))))).#trainneg
cnf(i_0_6090, negated_conjecture, (m1_subset_1(k1_relset_1(esk3_0,k2_zfmisc_1(esk3_0,X1)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6201, negated_conjecture, (m1_subset_1(k1_seq_1(k3_taylor_1,X1),k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_1202, plain, (m1_subset_1(esk6_2(k1_partfun1(X1,X2,X3,X4,X5,X6),X7),k2_zfmisc_1(X1,X4))|r2_fdiff_1(X7,k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_rcomp_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~r1_tarski(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_relset_1(k1_numbers,X7))|~v1_funct_1(X7)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5640, plain, (r2_hidden(k7_xcmplx_0(np__1,k6_numbers),k1_numbers)|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_5707, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos)),k1_relset_1(k1_numbers,k24_sin_cos)))).#trainneg
cnf(i_0_5713, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k19_sin_cos,k19_sin_cos)),k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_5714, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k29_sin_cos,k19_sin_cos)),k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_6296, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k3_taylor_1)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6579, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(X6,X7)))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_8736, plain, (m1_subset_1(k9_xtuple_0(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_6333, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_8)|k7_xcmplx_0(np__1,k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~m1_subset_1(k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))),a_0_0_fdiff_8)|~v1_xcmplx_0(k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))))).#trainneg
cnf(i_0_6120, plain, (r1_tarski(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(X2,k3_relat_1(X4,X3))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X5))))).#trainneg
cnf(i_0_6493, plain, (v1_xboole_0(X1)|r2_hidden(esk6_2(X1,k2_fdiff_1(X2,X1)),X1)|r2_fdiff_1(k2_fdiff_1(X2,X1),X1)|~v3_rcomp_1(X1)|~r2_fdiff_1(X2,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_6862, plain, (k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6863, plain, (k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k19_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_7106, plain, (k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k29_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k29_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1213, plain, (r1_fdiff_1(X1,k1_funct_1(X2,esk6_2(k9_xtuple_0(k3_relat_1(X2,X3)),X4)))|r2_fdiff_1(X4,k9_xtuple_0(k3_relat_1(X2,X3)))|~v3_rcomp_1(k9_xtuple_0(X3))|~r2_fdiff_1(X1,k9_xtuple_0(X3))|~v1_funct_1(X1)|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k9_xtuple_0(X3),k1_zfmisc_1(k1_numbers))|~v1_relat_1(X3)|~v1_relat_1(X2)|~v3_rcomp_1(k9_xtuple_0(k3_relat_1(X2,X3)))|~r1_tarski(k9_xtuple_0(k3_relat_1(X2,X3)),k1_relset_1(k1_numbers,X4))|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k9_xtuple_0(k3_relat_1(X2,X3)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5460, plain, (v1_xcmplx_0(k7_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6434, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k3_taylor_1)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6439, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6444, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k3_taylor_1)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6517, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6522, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k29_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8523, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_hidden(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_8267, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k3_square_1(X1),X2)|~v1_xcmplx_0(k3_square_1(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6838, plain, (k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_7684, plain, (m1_subset_1(k1_funct_1(X1,X2),k1_numbers)|~r2_hidden(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9175, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(np__1,X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_7105, plain, (k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6697, plain, (r1_tarski(X1,k9_xtuple_0(X2))|~v1_funct_1(X3)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(X4,k3_relat_1(X2,X3))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_7402, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,k29_sin_cos)),k1_zfmisc_1(X6))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_8009, negated_conjecture, (v1_xcmplx_0(X1)|~r1_tarski(X1,esk3_0)|~v3_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))))).#trainneg
cnf(i_0_8828, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k9_xtuple_0(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1270, plain, (k7_xcmplx_0(k7_xcmplx_0(X1,X2),k7_xcmplx_0(X3,np__1))=k7_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(k7_xcmplx_0(X3,np__1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X2,X3)))).#trainneg
cnf(i_0_6527, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k29_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6639, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k29_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6644, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6649, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6654, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8741, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(k3_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8776, plain, (k3_xcmplx_0(np__1,k7_xcmplx_0(X1,np__1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7080, plain, (k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_8308, plain, (k7_xcmplx_0(k5_square_1(X1),k5_xcmplx_0(X2))=k3_xcmplx_0(X1,k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1611, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|r2_hidden(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|~r1_tarski(X1,esk3_0))).#trainneg
cnf(i_0_1941, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(X1,k29_sin_cos))).#trainneg
cnf(i_0_1674, plain, (v1_xcmplx_0(X1)|~r2_hidden(X1,k3_taylor_1)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_6764, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6769, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k24_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6774, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k24_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6779, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k24_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6978, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k24_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6855, plain, (k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k19_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_7019, plain, (k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_8349, plain, (r1_tarski(X1,k9_xtuple_0(X2))|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(X2,k29_sin_cos)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_9307, plain, (k7_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,X1)),np__1)=k7_xcmplx_0(np__1,k1_fdiff_1(k29_sin_cos,X1))|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,X1)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8183, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_6519, plain, (v4_relat_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),k1_numbers)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_8770, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(X1,np__1))=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1616, negated_conjecture, (v1_valued_0(X1)|~v1_valued_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~v1_relat_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r1_tarski(X1,esk3_0))).#trainneg
cnf(i_0_1280, plain, (r1_fdiff_1(X1,esk1_3(X2,X3,X4))|X4=k2_fdiff_1(X2,X3)|~v3_rcomp_1(k1_numbers)|~r2_fdiff_1(X1,k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|k1_relset_1(k1_numbers,X4)!=X3|~r2_fdiff_1(X2,X3)|~v1_funct_1(X4)|~v1_funct_1(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_1791, plain, (v1_xcmplx_0(X1)|~r2_hidden(X1,k19_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_1934, plain, (v1_xcmplx_0(X1)|~r2_hidden(X1,k29_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_6297, plain, (k9_xtuple_0(k3_relat_1(k3_taylor_1,k3_taylor_1))=k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k3_taylor_1)))).#trainneg
cnf(i_0_5654, plain, (m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers)|v1_xboole_0(a_0_0_fdiff_8)|~m2_subset_1(k6_numbers,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_6953, plain, (r2_fdiff_1(k3_taylor_1,k3_relat_1(k3_taylor_1,k3_taylor_1))|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v3_rcomp_1(k3_relat_1(k3_taylor_1,k3_taylor_1))|~r1_tarski(k3_relat_1(k3_taylor_1,k3_taylor_1),a_0_0_fdiff_8)|~m1_subset_1(k3_relat_1(k3_taylor_1,k3_taylor_1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6961, plain, (r2_fdiff_1(k3_taylor_1,k3_relat_1(k19_sin_cos,k3_taylor_1))|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v3_rcomp_1(k3_relat_1(k19_sin_cos,k3_taylor_1))|~r1_tarski(k3_relat_1(k19_sin_cos,k3_taylor_1),a_0_0_fdiff_8)|~m1_subset_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7020, plain, (k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k3_taylor_1))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k3_taylor_1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6116, plain, (v3_valued_0(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6123, plain, (r1_tarski(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X4)|~m1_subset_1(X1,k1_zfmisc_1(k9_xtuple_0(k1_partfun1(X2,X5,X6,X7,X4,X3))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X5))))).#trainneg
cnf(i_0_8793, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k9_xtuple_0(k3_rfunct_1(X2,X3,X4,X5)))|~v1_funct_1(X5)|~v1_funct_1(X4)|~v3_membered(X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_6065, plain, (v3_valued_0(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_7097, plain, (k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k29_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k29_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_8271, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X2,k3_square_1(X1))|~v1_xcmplx_0(k3_square_1(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6995, plain, (k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6620, negated_conjecture, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_6440, plain, (k9_xtuple_0(k3_relat_1(k19_sin_cos,k3_taylor_1))=k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_6445, plain, (k9_xtuple_0(k3_relat_1(k29_sin_cos,k3_taylor_1))=k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_6518, plain, (k9_xtuple_0(k3_relat_1(k3_taylor_1,k29_sin_cos))=k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos)))).#trainneg
cnf(i_0_6969, plain, (r2_fdiff_1(k3_taylor_1,k3_relat_1(k29_sin_cos,k3_taylor_1))|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v3_rcomp_1(k3_relat_1(k29_sin_cos,k3_taylor_1))|~r1_tarski(k3_relat_1(k29_sin_cos,k3_taylor_1),a_0_0_fdiff_8)|~m1_subset_1(k3_relat_1(k29_sin_cos,k3_taylor_1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7243, plain, (r2_fdiff_1(k3_taylor_1,k3_relat_1(k3_taylor_1,k29_sin_cos))|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v3_rcomp_1(k3_relat_1(k3_taylor_1,k29_sin_cos))|~r1_tarski(k3_relat_1(k3_taylor_1,k29_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7251, plain, (r2_fdiff_1(k3_taylor_1,k3_relat_1(k19_sin_cos,k29_sin_cos))|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v3_rcomp_1(k3_relat_1(k19_sin_cos,k29_sin_cos))|~r1_tarski(k3_relat_1(k19_sin_cos,k29_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k3_relat_1(k19_sin_cos,k29_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7259, plain, (r2_fdiff_1(k3_taylor_1,k3_relat_1(k29_sin_cos,k29_sin_cos))|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v3_rcomp_1(k3_relat_1(k29_sin_cos,k29_sin_cos))|~r1_tarski(k3_relat_1(k29_sin_cos,k29_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k3_relat_1(k29_sin_cos,k29_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7485, plain, (r2_fdiff_1(k3_taylor_1,k3_relat_1(k3_taylor_1,k19_sin_cos))|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v3_rcomp_1(k3_relat_1(k3_taylor_1,k19_sin_cos))|~r1_tarski(k3_relat_1(k3_taylor_1,k19_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8317, plain, (r1_tarski(X1,k9_xtuple_0(X2))|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(X2,k3_taylor_1)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_6169, plain, (v3_valued_0(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k29_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_8940, plain, (r1_tarski(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,X5)),k1_relset_1(k1_numbers,k29_sin_cos))|~v1_funct_1(X5)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_6641, plain, (v4_relat_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k29_sin_cos),k1_numbers)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6298, plain, (v4_relat_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1),k1_numbers)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6754, plain, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k5_numbers))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_6528, plain, (k9_xtuple_0(k3_relat_1(k19_sin_cos,k29_sin_cos))=k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k29_sin_cos)))).#trainneg
cnf(i_0_6640, plain, (k9_xtuple_0(k3_relat_1(k29_sin_cos,k29_sin_cos))=k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k29_sin_cos)))).#trainneg
cnf(i_0_6645, plain, (k9_xtuple_0(k3_relat_1(k3_taylor_1,k19_sin_cos))=k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos)))).#trainneg
cnf(i_0_7493, plain, (r2_fdiff_1(k3_taylor_1,k3_relat_1(k19_sin_cos,k19_sin_cos))|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v3_rcomp_1(k3_relat_1(k19_sin_cos,k19_sin_cos))|~r1_tarski(k3_relat_1(k19_sin_cos,k19_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k3_relat_1(k19_sin_cos,k19_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7501, plain, (r2_fdiff_1(k3_taylor_1,k3_relat_1(k29_sin_cos,k19_sin_cos))|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v3_rcomp_1(k3_relat_1(k29_sin_cos,k19_sin_cos))|~r1_tarski(k3_relat_1(k29_sin_cos,k19_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k3_relat_1(k29_sin_cos,k19_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7797, plain, (r2_fdiff_1(k3_taylor_1,k3_relat_1(k24_sin_cos,k3_taylor_1))|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v3_rcomp_1(k3_relat_1(k24_sin_cos,k3_taylor_1))|~r1_tarski(k3_relat_1(k24_sin_cos,k3_taylor_1),a_0_0_fdiff_8)|~m1_subset_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7805, plain, (r2_fdiff_1(k3_taylor_1,k3_relat_1(k24_sin_cos,k29_sin_cos))|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v3_rcomp_1(k3_relat_1(k24_sin_cos,k29_sin_cos))|~r1_tarski(k3_relat_1(k24_sin_cos,k29_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k3_relat_1(k24_sin_cos,k29_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8147, plain, (r2_fdiff_1(k3_taylor_1,k3_relat_1(k24_sin_cos,k19_sin_cos))|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v3_rcomp_1(k3_relat_1(k24_sin_cos,k19_sin_cos))|~r1_tarski(k3_relat_1(k24_sin_cos,k19_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k3_relat_1(k24_sin_cos,k19_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_7373, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,k3_taylor_1)),k1_zfmisc_1(X6))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6660, plain, (r1_tarski(k1_relset_1(X1,k2_zfmisc_1(X1,X2)),X1))).#trainneg
cnf(i_0_8348, plain, (r1_tarski(X1,k9_xtuple_0(X2))|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(X2,k19_sin_cos)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_7401, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,k19_sin_cos)),k1_zfmisc_1(X6))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5690, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos)),a_0_0_fdiff_8))).#trainneg
cnf(i_0_9380, negated_conjecture, (v1_xcmplx_0(esk3_0)|~v1_membered(k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_9381, negated_conjecture, (v7_ordinal1(esk3_0)|~v6_membered(k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_9382, negated_conjecture, (v1_xreal_0(esk3_0)|~v3_membered(k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_9501, plain, (r1_fdiff_1(k3_taylor_1,np__1)|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(np__1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_9618, plain, (r2_fdiff_1(k29_sin_cos,k1_relset_1(k1_numbers,k29_sin_cos))|~v3_rcomp_1(k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_9771, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_9366, negated_conjecture, (m1_subset_1(esk6_2(esk3_0,X1),a_0_0_fdiff_8)|r2_fdiff_1(X1,esk3_0)|~r1_tarski(esk3_0,k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_9364, negated_conjecture, (m1_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_9563, negated_conjecture, (r1_fdiff_1(X1,esk6_2(X2,X3))|r2_fdiff_1(X3,X2)|~v3_rcomp_1(X2)|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r1_tarski(X2,esk3_0)|~r2_fdiff_1(X3,esk3_0)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_9367, negated_conjecture, (m1_subset_1(esk6_2(esk3_0,X1),a_0_0_fdiff_8)|r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_10169, negated_conjecture, (r2_hidden(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8)|v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_9508, plain, (r1_tarski(k9_xtuple_0(k2_zfmisc_1(X1,X2)),X1))).#trainneg
cnf(i_0_9571, negated_conjecture, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_1299, plain, (r1_fdiff_1(X1,k1_partfun1(X2,X3,X4,X5,X6,X7))|v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(X2,X5)))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X2,X5)))|~r2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X5)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X2,X5)),k1_zfmisc_1(k1_numbers))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_9781, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k19_sin_cos,k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_9791, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k29_sin_cos,k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_9379, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(a_0_0_fdiff_8))|r2_hidden(esk3_0,k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_11131, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k3_taylor_1)),a_0_0_fdiff_8))).#trainneg
cnf(i_0_11556, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8))).#trainneg
cnf(i_0_9978, plain, (r1_fdiff_1(X1,esk6_2(X2,k2_fdiff_1(X3,X2)))|v1_xboole_0(X2)|r2_fdiff_1(k2_fdiff_1(X3,X2),X2)|~v3_rcomp_1(X2)|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~r2_fdiff_1(X3,X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_10011, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos)),X1)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10307, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos)),X1)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_9927, negated_conjecture, (m1_subset_1(esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0)),a_0_0_fdiff_8)|r2_fdiff_1(k2_fdiff_1(X1,esk3_0),esk3_0)|~r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_9750, negated_conjecture, (v7_ordinal1(k1_seq_1(k3_taylor_1,X1))|~v6_membered(k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_1300, plain, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|r2_hidden(k3_relat_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X4,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_10046, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k19_sin_cos)),X1)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10757, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos)),X1)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10408, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos)),X1)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10082, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k29_sin_cos)),X1)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10990, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1)),X1)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1306, plain, (r1_fdiff_1(X1,esk6_2(k9_xtuple_0(k3_relat_1(X2,X3)),X4))|r2_fdiff_1(X4,k9_xtuple_0(k3_relat_1(X2,X3)))|~v3_rcomp_1(k9_xtuple_0(X2))|~r2_fdiff_1(X1,k9_xtuple_0(X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k9_xtuple_0(X2),k1_zfmisc_1(k1_numbers))|~v3_rcomp_1(k9_xtuple_0(k3_relat_1(X2,X3)))|~r1_tarski(k9_xtuple_0(k3_relat_1(X2,X3)),k1_relset_1(k1_numbers,X4))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_funct_1(X4)|~m1_subset_1(k9_xtuple_0(k3_relat_1(X2,X3)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_relat_1(X3)|~v1_relat_1(X2))).#trainneg
cnf(i_0_11428, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1)),X1)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10254, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k5_square_1(X1),X2)|~v1_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11822, negated_conjecture, (m1_subset_1(k1_relset_1(esk3_0,k2_zfmisc_1(esk3_0,X1)),k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_11189, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k3_taylor_1)),X1)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10954, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k19_sin_cos)),X1)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_11274, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k29_sin_cos)),X1)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_9385, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_8)|m1_subset_1(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_9926, negated_conjecture, (m1_subset_1(esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0)),k1_numbers)|r2_fdiff_1(k2_fdiff_1(X1,esk3_0),esk3_0)|~r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_10867, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|r2_hidden(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_9761, plain, (m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers)|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_11132, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k3_taylor_1)),k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_11557, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_10305, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10379, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X2,np__1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_10822, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,X2))=k3_xcmplx_0(X2,k5_square_1(X1))|~v1_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11111, negated_conjecture, (v1_valued_0(X1)|~r1_tarski(X1,esk3_0)|~v1_relat_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~v3_valued_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_9509, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k9_xtuple_0(k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_1317, plain, (r1_fdiff_1(X1,esk1_3(X2,k9_xtuple_0(k3_relat_1(X3,X4)),X5))|X5=k2_fdiff_1(X2,k9_xtuple_0(k3_relat_1(X3,X4)))|~v3_rcomp_1(k9_xtuple_0(X3))|~r2_fdiff_1(X1,k9_xtuple_0(X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k9_xtuple_0(X3),k1_zfmisc_1(k1_numbers))|k1_relset_1(k1_numbers,X5)!=k9_xtuple_0(k3_relat_1(X3,X4))|~r2_fdiff_1(X2,k9_xtuple_0(k3_relat_1(X3,X4)))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X5)|~v1_funct_1(X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_relat_1(X4)|~v1_relat_1(X3))).#trainneg
cnf(i_0_11149, plain, (~v1_xboole_0(a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k3_taylor_1))))).#trainneg
cnf(i_0_11487, plain, (~v1_xboole_0(a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_11532, plain, (~v1_xboole_0(a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_11747, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_8155, plain, (r2_fdiff_1(k3_taylor_1,k3_relat_1(k3_taylor_1,k24_sin_cos))|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v3_rcomp_1(k3_relat_1(k3_taylor_1,k24_sin_cos))|~r1_tarski(k3_relat_1(k3_taylor_1,k24_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k3_relat_1(k3_taylor_1,k24_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8163, plain, (r2_fdiff_1(k3_taylor_1,k3_relat_1(k19_sin_cos,k24_sin_cos))|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v3_rcomp_1(k3_relat_1(k19_sin_cos,k24_sin_cos))|~r1_tarski(k3_relat_1(k19_sin_cos,k24_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k3_relat_1(k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8484, plain, (r2_fdiff_1(k3_taylor_1,k3_relat_1(k29_sin_cos,k24_sin_cos))|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v3_rcomp_1(k3_relat_1(k29_sin_cos,k24_sin_cos))|~r1_tarski(k3_relat_1(k29_sin_cos,k24_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k3_relat_1(k29_sin_cos,k24_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_8492, plain, (r2_fdiff_1(k3_taylor_1,k3_relat_1(k24_sin_cos,k24_sin_cos))|~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~v3_rcomp_1(k3_relat_1(k24_sin_cos,k24_sin_cos))|~r1_tarski(k3_relat_1(k24_sin_cos,k24_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k3_relat_1(k24_sin_cos,k24_sin_cos),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_10006, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos)),k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10374, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,np__1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_11615, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k29_sin_cos)),k1_zfmisc_1(k1_numbers))|~v1_relat_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k29_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1620, negated_conjecture, (~v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,X2)|~r1_tarski(X2,esk3_0))).#trainneg
cnf(i_0_10678, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(k7_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1325, plain, (X1=k2_fdiff_1(X2,k9_xtuple_0(k1_partfun1(X5,X6,X7,X8,X3,X4)))|r2_hidden(esk1_3(X2,k9_xtuple_0(k1_partfun1(X5,X6,X7,X8,X3,X4)),X1),k9_xtuple_0(X3))|k1_relset_1(k1_numbers,X1)!=k9_xtuple_0(k1_partfun1(X5,X6,X7,X8,X3,X4))|~r2_fdiff_1(X2,k9_xtuple_0(k1_partfun1(X5,X6,X7,X8,X3,X4)))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_relat_1(X4)|~v1_relat_1(X3)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_9763, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_relset_1(k1_numbers,k24_sin_cos))))).#trainneg
cnf(i_0_9773, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos))))).#trainneg
cnf(i_0_9783, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k29_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_relset_1(k1_numbers,k29_sin_cos))))).#trainneg
cnf(i_0_11469, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1)),k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_10041, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k19_sin_cos)),k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_11061, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos)),k1_zfmisc_1(k1_numbers))|~v1_relat_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10302, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10403, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos)),k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_9697, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(k3_taylor_1,X1,X2,X3,X4,X5)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k3_taylor_1,X1))))).#trainneg
cnf(i_0_10868, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(k29_sin_cos,X1,X2,X3,X4,X5)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k29_sin_cos,X1))))).#trainneg
cnf(i_0_10044, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k19_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1618, negated_conjecture, (m1_subset_1(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,X2)|~r1_tarski(X2,esk3_0))).#trainneg
cnf(i_0_9606, plain, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1326, plain, (r2_hidden(k1_funct_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(X2,k9_xtuple_0(k3_relat_1(X1,k3_taylor_1)))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11514, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k3_taylor_1)),k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_11657, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k29_sin_cos)),k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_11703, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k29_sin_cos)),k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_9486, negated_conjecture, (r2_hidden(esk6_2(esk3_0,X1),k1_numbers)|r2_fdiff_1(X1,esk3_0)|~r1_tarski(esk3_0,k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_9712, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(k19_sin_cos,X1,X2,X3,X4,X5)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k19_sin_cos,X1))))).#trainneg
cnf(i_0_10662, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_11912, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10010, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos)),k1_zfmisc_1(X1))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10077, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k29_sin_cos)),k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10406, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_9719, plain, (r1_tarski(k9_xtuple_0(k1_partfun1(X7,X8,X9,X10,X5,X6)),X1)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X9,X10)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X7,X8))))).#trainneg
cnf(i_0_10755, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_9745, negated_conjecture, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(esk3_0,k2_zfmisc_1(esk3_0,X2))))).#trainneg
cnf(i_0_1327, plain, (r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,k3_taylor_1)))|~r2_hidden(k1_funct_1(X2,X1),a_0_0_fdiff_8)|~r2_hidden(X1,k9_xtuple_0(X2))|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_9905, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k9_xtuple_0(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_4555, negated_conjecture, (k7_xcmplx_0(k7_xcmplx_0(np__1,esk4_0),k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(esk4_0))).#trainneg
cnf(i_0_4556, negated_conjecture, (k7_xcmplx_0(k7_xcmplx_0(np__1,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))),esk4_0)!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(esk4_0)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))).#trainneg
cnf(i_0_5655, plain, (r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,k6_numbers))|v1_xboole_0(a_0_0_fdiff_8)|~m2_subset_1(k6_numbers,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_6294, plain, (v1_xboole_0(a_0_0_fdiff_8)|r2_hidden(k7_xcmplx_0(np__1,k6_numbers),a_0_0_fdiff_8)|~m2_subset_1(k6_numbers,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers))).#trainneg
cnf(i_0_9718, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(X5,X6)),X1)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_11620, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1)),k1_zfmisc_1(k1_numbers))|~v1_relat_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10679, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(np__1,X1))|~m1_subset_1(k3_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_11595, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(X1,k3_taylor_1)),k9_xtuple_0(X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_10179, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,np__1))=k3_square_1(X1)|~v1_xcmplx_0(k3_square_1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9782, plain, (r1_tarski(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos)),k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_11824, negated_conjecture, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,k1_relset_1(esk3_0,k2_zfmisc_1(esk3_0,X2))))).#trainneg
cnf(i_0_9687, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r1_tarski(X1,k9_xtuple_0(k1_partfun1(esk3_0,X2,X3,X4,X5,X6)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(esk3_0,X2))))).#trainneg
cnf(i_0_10629, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k29_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10824, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,k5_square_1(X1)))=k3_square_1(k5_square_1(X1))|~v1_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10985, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1)),k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_9800, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k3_taylor_1))))).#trainneg
cnf(i_0_1328, plain, (X1=k2_fdiff_1(X2,k9_xtuple_0(k3_relat_1(k3_taylor_1,X3)))|r2_hidden(esk1_3(X2,k9_xtuple_0(k3_relat_1(k3_taylor_1,X3)),X1),a_0_0_fdiff_8)|k1_relset_1(k1_numbers,X1)!=k9_xtuple_0(k3_relat_1(k3_taylor_1,X3))|~r2_fdiff_1(X2,k9_xtuple_0(k3_relat_1(k3_taylor_1,X3)))|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_relat_1(X3))).#trainneg
cnf(i_0_6655, plain, (k9_xtuple_0(k3_relat_1(k19_sin_cos,k19_sin_cos))=k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_6765, plain, (k9_xtuple_0(k3_relat_1(k29_sin_cos,k19_sin_cos))=k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_9769, plain, (r2_fdiff_1(k24_sin_cos,k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos)))|~v3_rcomp_1(k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos)))|~v1_xboole_0(k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos)))|~m1_subset_1(k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4216, plain, (r2_fdiff_1(k3_taylor_1,k1_relset_1(k1_numbers,k29_sin_cos))|k1_seq_1(k19_sin_cos,esk6_2(k1_relset_1(k1_numbers,k29_sin_cos),k3_taylor_1))!=k6_numbers|~m1_subset_1(esk6_2(k1_relset_1(k1_numbers,k29_sin_cos),k3_taylor_1),k1_numbers)|~v3_rcomp_1(k1_relset_1(k1_numbers,k29_sin_cos))|~r1_tarski(k1_relset_1(k1_numbers,k29_sin_cos),a_0_0_fdiff_8))).#trainneg
cnf(i_0_4597, negated_conjecture, (k7_xcmplx_0(np__1,k3_xcmplx_0(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers)|~v1_xcmplx_0(esk4_0)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))).#trainneg
cnf(i_0_10306, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos)),k1_zfmisc_1(X1))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10171, negated_conjecture, (r2_hidden(k1_seq_1(k3_taylor_1,esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))),k1_relset_1(k1_numbers,k29_sin_cos))|v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_10988, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_11924, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r2_hidden(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8))).#trainneg
cnf(i_0_14597, negated_conjecture, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_9781, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_12026, negated_conjecture, (r1_tarski(k5_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_13354, plain, (v3_membered(k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_12032, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k24_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_12033, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k19_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_12034, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k29_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_9773, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos))))).#trainneg
cnf(i_0_14602, negated_conjecture, (~v1_xboole_0(a_0_0_fdiff_8)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_12050, negated_conjecture, (r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_12069, negated_conjecture, (k1_zfmisc_1(a_0_0_fdiff_8)=k6_numbers|r2_hidden(esk3_0,k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_12814, plain, (v3_membered(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_13152, plain, (v3_membered(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos))))).#trainneg
cnf(i_0_12022, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_11928, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_11974, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~r1_tarski(X1,k9_xtuple_0(k2_zfmisc_1(esk3_0,X3)))|~r2_fdiff_1(X2,esk3_0)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_14482, negated_conjecture, (r1_tarski(k9_xtuple_0(k1_partfun1(esk3_0,X1,X2,X3,X4,X5)),k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(esk3_0,X1)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_14181, plain, (r2_hidden(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1)|~r2_hidden(k1_funct_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(X2,k9_xtuple_0(X1)))).#trainneg
cnf(i_0_5713, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k19_sin_cos)),k1_relset_1(k1_numbers,k19_sin_cos)))).#trainneg
cnf(i_0_12819, plain, (~v1_xboole_0(a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_13157, plain, (~v1_xboole_0(a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos))))).#trainneg
cnf(i_0_11933, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_12073, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_12081, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_11972, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r1_tarski(X1,k9_xtuple_0(k2_zfmisc_1(esk3_0,X2))))).#trainneg
cnf(i_0_14448, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(X5,X6)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(esk3_0,X2)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_11917, negated_conjecture, (r2_fdiff_1(k3_taylor_1,X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_9791, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_14594, negated_conjecture, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14595, negated_conjecture, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_13398, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_13710, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_13721, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_13732, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_12016, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_relset_1(k1_numbers,k24_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14429, plain, (r1_tarski(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos)),k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_12827, plain, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,X2,X3,X4,k3_taylor_1,k19_sin_cos)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,X2))))).#trainneg
cnf(i_0_13826, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_14035, plain, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,X2,X3,X4,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,X2))))).#trainneg
cnf(i_0_14612, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k19_sin_cos),X1)),k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k19_sin_cos)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_12051, negated_conjecture, (r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(X2,k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_9783, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_relset_1(k1_numbers,k29_sin_cos))))).#trainneg
cnf(i_0_14596, negated_conjecture, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_12632, negated_conjecture, (~v1_xboole_0(a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(esk3_0,k2_zfmisc_1(esk3_0,X2))))).#trainneg
cnf(i_0_5150, plain, (v1_xreal_0(X1)|~v1_xreal_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10618, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10952, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k19_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10081, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k29_sin_cos)),k1_zfmisc_1(X1))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_14259, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(X1,k19_sin_cos)),X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_12085, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10407, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos)),k1_zfmisc_1(X1))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_5714, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k19_sin_cos)),k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_1807, plain, (v1_xcmplx_0(X1)|~r2_hidden(X1,k24_sin_cos)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_11574, plain, (~v1_xboole_0(k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_11675, plain, (~v1_xboole_0(k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k29_sin_cos))))).#trainneg
cnf(i_0_10474, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_9845, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(X3,k19_sin_cos)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X4))))).#trainneg
cnf(i_0_12180, negated_conjecture, (v1_xreal_0(esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0)))|r2_fdiff_1(k2_fdiff_1(X1,esk3_0),esk3_0)|~r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_10045, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k19_sin_cos)),k1_zfmisc_1(X1))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10640, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10882, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_14287, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,np__1),k1_numbers))).#trainneg
cnf(i_0_1329, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_taylor_1,X1)),a_0_0_fdiff_8)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11721, plain, (~v1_xboole_0(k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k29_sin_cos))))).#trainneg
cnf(i_0_11764, plain, (~v1_xboole_0(k1_relset_1(k1_numbers,k19_sin_cos))|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos))))).#trainneg
cnf(i_0_14644, plain, (~v1_xboole_0(k1_relset_1(k1_numbers,k19_sin_cos))|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_14329, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k24_sin_cos,k3_taylor_1)),k1_relset_1(k1_numbers,k24_sin_cos)))).#trainneg
cnf(i_0_13393, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_relset_1(k1_numbers,k19_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_10756, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos)),k1_zfmisc_1(X1))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10301, plain, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_10015, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k19_sin_cos)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_12077, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_13502, plain, (m1_subset_1(k1_relset_1(k3_taylor_1,k1_partfun1(k3_taylor_1,X1,X2,X3,X4,X5)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k3_taylor_1,X1)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_11184, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k3_taylor_1)),k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1330, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_taylor_1,X1)),k1_zfmisc_1(a_0_0_fdiff_8))|~v1_relat_1(X1))).#trainneg
cnf(i_0_14694, plain, (~v1_xboole_0(k1_relset_1(k1_numbers,k19_sin_cos))|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_6064, plain, (v1_xcmplx_0(k1_seq_1(k3_relat_1(k3_taylor_1,k3_taylor_1),X1)))).#trainneg
cnf(i_0_6109, plain, (v1_xcmplx_0(k1_seq_1(k3_relat_1(k19_sin_cos,k3_taylor_1),X1)))).#trainneg
cnf(i_0_13705, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_relset_1(k1_numbers,k29_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_11005, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(X1,k29_sin_cos)),k9_xtuple_0(X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_10154, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10086, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k29_sin_cos)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_11091, plain, (k7_xcmplx_0(k3_xcmplx_0(X1,np__1),np__1)=X1|~v1_xcmplx_0(k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_14296, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(X1,np__1),k1_numbers))).#trainneg
cnf(i_0_1331, plain, (r2_hidden(esk6_2(k9_xtuple_0(k3_relat_1(k3_taylor_1,X1)),X2),a_0_0_fdiff_8)|r2_fdiff_1(X2,k9_xtuple_0(k3_relat_1(k3_taylor_1,X1)))|~v3_rcomp_1(k9_xtuple_0(k3_relat_1(k3_taylor_1,X1)))|~r1_tarski(k9_xtuple_0(k3_relat_1(k3_taylor_1,X1)),k1_relset_1(k1_numbers,X2))|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(k9_xtuple_0(k3_relat_1(k3_taylor_1,X1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_6112, plain, (v1_xcmplx_0(k1_seq_1(k3_relat_1(k29_sin_cos,k3_taylor_1),X1)))).#trainneg
cnf(i_0_6115, plain, (v1_xcmplx_0(k1_seq_1(k3_relat_1(k3_taylor_1,k29_sin_cos),X1)))).#trainneg
cnf(i_0_6165, plain, (v1_xcmplx_0(k1_seq_1(k3_relat_1(k19_sin_cos,k29_sin_cos),X1)))).#trainneg
cnf(i_0_13716, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_relset_1(k1_numbers,k19_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14242, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(X1,k24_sin_cos)),X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_10132, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10475, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k29_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_12333, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k3_taylor_1,X1,X2,X3,k19_sin_cos,k29_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k3_taylor_1,X1))))).#trainneg
cnf(i_0_10989, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1)),k1_zfmisc_1(X1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10751, plain, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_12247, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k3_taylor_1,X1,X2,X3,k29_sin_cos,k19_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k3_taylor_1,X1))))).#trainneg
cnf(i_0_10118, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1332, plain, (r1_tarski(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,X5)),a_0_0_fdiff_8)|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6168, plain, (v1_xcmplx_0(k1_seq_1(k3_relat_1(k29_sin_cos,k29_sin_cos),X1)))).#trainneg
cnf(i_0_6171, plain, (v1_xcmplx_0(k1_seq_1(k3_relat_1(k3_taylor_1,k19_sin_cos),X1)))).#trainneg
cnf(i_0_6177, plain, (v1_xcmplx_0(k1_seq_1(k3_relat_1(k19_sin_cos,k19_sin_cos),X1)))).#trainneg
cnf(i_0_13727, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_relset_1(k1_numbers,k29_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_13119, plain, (r2_fdiff_1(k3_taylor_1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k3_taylor_1)))|~v1_xboole_0(a_0_0_fdiff_8)|~v3_rcomp_1(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k3_taylor_1))))).#trainneg
cnf(i_0_14323, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k19_sin_cos),k3_taylor_1)),k9_xtuple_0(k3_relat_1(k19_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_10710, plain, (k7_xcmplx_0(k7_xcmplx_0(X1,np__1),np__1)=X1|~v1_xcmplx_0(k7_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11427, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1)),k1_zfmisc_1(X1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10303, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_11817, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10129, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1347, plain, (k8_real_1(k1_fdiff_1(X1,k6_numbers),k1_fdiff_1(k19_sin_cos,k1_seq_1(X2,esk6_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k29_sin_cos)))))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,X1),k1_seq_1(X2,esk6_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k29_sin_cos))))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k29_sin_cos),X3)|~r1_fdiff_1(X1,k6_numbers)|~r1_fdiff_1(k19_sin_cos,k1_seq_1(X2,esk6_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k29_sin_cos))))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X2,esk6_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k29_sin_cos))),k1_numbers)|~v3_rcomp_1(X3)|~r1_fdiff_1(X2,esk6_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k29_sin_cos)))|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X2,k29_sin_cos)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6180, plain, (v1_xcmplx_0(k1_seq_1(k3_relat_1(k29_sin_cos,k19_sin_cos),X1)))).#trainneg
cnf(i_0_11969, negated_conjecture, (r1_tarski(k9_xtuple_0(k2_zfmisc_1(esk3_0,X1)),k1_numbers))).#trainneg
cnf(i_0_1937, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r1_tarski(X1,k24_sin_cos))).#trainneg
cnf(i_0_13360, plain, (r2_fdiff_1(k24_sin_cos,k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos)))|~v3_rcomp_1(k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos)))|~v1_xboole_0(k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos)))|~m1_subset_1(k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_13396, plain, (r2_fdiff_1(k19_sin_cos,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1)))|~v3_rcomp_1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1)))|~v1_xboole_0(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_13708, plain, (r2_fdiff_1(k29_sin_cos,k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k3_taylor_1)))|~v3_rcomp_1(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k3_taylor_1)))|~v1_xboole_0(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_13719, plain, (r2_fdiff_1(k19_sin_cos,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k29_sin_cos)))|~v3_rcomp_1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k29_sin_cos)))|~v1_xboole_0(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k29_sin_cos))))).#trainneg
cnf(i_0_11422, plain, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k3_taylor_1)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_11432, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k3_taylor_1)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_11913, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_9803, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_11187, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k3_taylor_1))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k3_taylor_1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10753, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_11943, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),esk3_0)|m1_subset_1(esk6_2(esk3_0,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10140, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_1351, plain, (k1_seq_1(k19_sin_cos,k1_fdiff_1(X1,esk6_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k29_sin_cos))))=k6_numbers|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k29_sin_cos),X3)|~v3_rcomp_1(X3)|~r1_fdiff_1(k2_fdiff_1(X1,X2),esk6_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k29_sin_cos)))|~r1_tarski(X3,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k29_sin_cos)))|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(k1_fdiff_1(X1,esk6_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k29_sin_cos))),k1_numbers)|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers))|~r2_hidden(esk6_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k29_sin_cos)),X2)|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk6_2(X3,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_13359, plain, (~v1_xboole_0(k1_relset_1(k1_numbers,k24_sin_cos))|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_2376, plain, (r2_hidden(X1,k1_numbers)|~v1_xreal_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13730, plain, (r2_fdiff_1(k29_sin_cos,k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k29_sin_cos)))|~v3_rcomp_1(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k29_sin_cos)))|~v1_xboole_0(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k29_sin_cos))))).#trainneg
cnf(i_0_4602, negated_conjecture, (k7_xcmplx_0(k7_xcmplx_0(np__1,esk4_0),k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(esk4_0))).#trainneg
cnf(i_0_4753, plain, (m1_subset_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_11879, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_11188, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k3_taylor_1)),k1_zfmisc_1(X1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_12666, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k19_sin_cos,X1,X2,X3,k29_sin_cos,k3_taylor_1)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k19_sin_cos,X1))))).#trainneg
cnf(i_0_11424, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_11272, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k29_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k29_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_12772, negated_conjecture, (r2_hidden(esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0)),k1_numbers)|r2_fdiff_1(k2_fdiff_1(X1,esk3_0),esk3_0)|~r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_10151, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_1354, plain, (r2_hidden(k1_funct_1(X1,X2),k1_relset_1(X3,k1_partfun1(X3,X4,X5,X6,X7,X8)))|~r2_hidden(X2,k9_xtuple_0(k3_relat_1(X1,k1_partfun1(X3,X4,X5,X6,X7,X8))))|~v1_funct_1(k1_partfun1(X3,X4,X5,X6,X7,X8))|~v1_funct_1(X1)|~v1_relat_1(k1_partfun1(X3,X4,X5,X6,X7,X8))|~v1_relat_1(X1)|~v1_funct_1(X8)|~v1_funct_1(X7)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_8745, plain, (r2_hidden(X1,k1_numbers)|~v1_xreal_0(k3_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10686, plain, (r2_hidden(X1,k1_numbers)|~v1_xreal_0(k7_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11070, plain, (r2_hidden(X1,k1_numbers)|~v1_xreal_0(k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4754, plain, (m1_subset_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4755, plain, (m1_subset_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4756, plain, (m1_subset_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4787, plain, (m1_subset_1(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4788, plain, (m1_subset_1(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_14953, negated_conjecture, (m1_subset_1(np__1,a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14952, negated_conjecture, (m1_subset_1(np__1,a_0_0_fdiff_8)|~m1_subset_1(k5_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14958, negated_conjecture, (m1_subset_1(k6_numbers,a_0_0_fdiff_8)|~m1_subset_1(k5_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14959, negated_conjecture, (m1_subset_1(k6_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14898, negated_conjecture, (m1_subset_1(k5_numbers,k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_15951, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k24_sin_cos,k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_16409, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_taylor_1,k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_15165, negated_conjecture, (r1_tarski(k1_relset_1(esk3_0,k2_zfmisc_1(esk3_0,X1)),k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_14940, negated_conjecture, (m1_subset_1(esk6_2(X1,X2),a_0_0_fdiff_8)|r2_fdiff_1(X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(esk3_0))|~v3_rcomp_1(X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_17024, plain, (r1_tarski(X1,X2)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v3_rcomp_1(X1)|~r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k3_taylor_1),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k3_taylor_1))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_16961, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(k24_sin_cos,X1,X2,X3,X4,X5)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k24_sin_cos,X1))))).#trainneg
cnf(i_0_16182, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_taylor_1,X1)),k1_zfmisc_1(k1_numbers))|~v1_relat_1(X1))).#trainneg
cnf(i_0_15092, negated_conjecture, (r2_hidden(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1)|~r2_hidden(k1_funct_1(X1,X2),esk3_0))).#trainneg
cnf(i_0_1355, plain, (r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,k1_partfun1(X3,X4,X5,X6,X7,X8))))|~r2_hidden(k1_funct_1(X2,X1),k1_relset_1(X3,k1_partfun1(X3,X4,X5,X6,X7,X8)))|~r2_hidden(X1,k9_xtuple_0(X2))|~v1_funct_1(k1_partfun1(X3,X4,X5,X6,X7,X8))|~v1_funct_1(X2)|~v1_relat_1(k1_partfun1(X3,X4,X5,X6,X7,X8))|~v1_relat_1(X2)|~v1_funct_1(X8)|~v1_funct_1(X7)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(X5,X6)))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_16426, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k24_sin_cos,k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_16449, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k19_sin_cos,k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_16450, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k29_sin_cos,k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_17249, plain, (v3_valued_0(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k3_taylor_1)))).#trainneg
cnf(i_0_17655, plain, (v3_valued_0(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k24_sin_cos)))).#trainneg
cnf(i_0_16300, plain, (r1_tarski(X1,X2)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v3_rcomp_1(X1)|~r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k29_sin_cos),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k29_sin_cos))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_14936, negated_conjecture, (m1_subset_1(esk6_2(X1,k3_taylor_1),a_0_0_fdiff_8)|r2_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(X1,k1_zfmisc_1(esk3_0))|~v3_rcomp_1(X1)|~r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_15061, negated_conjecture, (m1_subset_1(k9_xtuple_0(k1_partfun1(esk3_0,X1,X2,X3,X4,X5)),k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(esk3_0,X1)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1357, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6),X7)),k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,X5,X6)))|~v1_relat_1(X7)|~v1_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_15192, negated_conjecture, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk3_0))|~v3_membered(k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_17711, plain, (v3_valued_0(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k19_sin_cos)))).#trainneg
cnf(i_0_17767, plain, (v3_valued_0(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k29_sin_cos)))).#trainneg
cnf(i_0_17823, plain, (v3_valued_0(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k3_taylor_1)))).#trainneg
cnf(i_0_17879, plain, (v3_valued_0(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k24_sin_cos)))).#trainneg
cnf(i_0_14906, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,k24_sin_cos),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_17039, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_17291, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_16086, plain, (r1_tarski(X1,X2)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v3_rcomp_1(X1)|~r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k19_sin_cos),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k19_sin_cos))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_16480, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k29_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_15738, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_15132, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_1358, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6),X7)),k1_zfmisc_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,X5,X6))))|~v1_relat_1(X7)|~v1_relat_1(k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_15346, plain, (~v1_xboole_0(k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_14913, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,k19_sin_cos),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14920, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,k29_sin_cos),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_17234, plain, (v4_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k3_taylor_1),k1_numbers))).#trainneg
cnf(i_0_17250, plain, (v5_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k3_taylor_1),k1_numbers))).#trainneg
cnf(i_0_17640, plain, (v4_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k24_sin_cos),k1_numbers))).#trainneg
cnf(i_0_15906, plain, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(k3_taylor_1,X2)))|~v1_relat_1(X2))).#trainneg
cnf(i_0_17850, plain, (m1_subset_1(k3_relat_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1),k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_16921, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_15076, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(X4,X5)),k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(esk3_0,X1)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_16085, plain, (r1_tarski(X1,X2)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k19_sin_cos)))))).#trainneg
cnf(i_0_14911, negated_conjecture, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,k1_relset_1(k1_numbers,k24_sin_cos))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_1360, plain, (r1_tarski(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k1_partfun1(X5,X6,X7,X8,X9,X10),X11)),k1_relset_1(X5,k1_partfun1(X5,X6,X7,X8,X9,X10)))|~v1_funct_1(X11)|~v1_funct_1(k1_partfun1(X5,X6,X7,X8,X9,X10))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k1_partfun1(X5,X6,X7,X8,X9,X10),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X10)|~v1_funct_1(X9)|~m1_subset_1(X10,k1_zfmisc_1(k2_zfmisc_1(X7,X8)))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X5,X6))))).#trainneg
cnf(i_0_11937, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_17656, plain, (v5_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k24_sin_cos),k1_numbers))).#trainneg
cnf(i_0_17696, plain, (v4_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k19_sin_cos),k1_numbers))).#trainneg
cnf(i_0_17712, plain, (v5_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k19_sin_cos),k1_numbers))).#trainneg
cnf(i_0_17752, plain, (v4_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k29_sin_cos),k1_numbers))).#trainneg
cnf(i_0_17768, plain, (v5_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k29_sin_cos),k1_numbers))).#trainneg
cnf(i_0_16299, plain, (r1_tarski(X1,X2)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k29_sin_cos)))))).#trainneg
cnf(i_0_16964, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k24_sin_cos,X1,X2,X3,k29_sin_cos,k19_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k24_sin_cos,X1))))).#trainneg
cnf(i_0_9855, plain, (m1_subset_1(k9_xtuple_0(k2_zfmisc_1(a_0_0_fdiff_8,X1)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6172, plain, (v3_valued_0(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10143, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k3_taylor_1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_16967, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k24_sin_cos,X1,X2,X3,k19_sin_cos,k29_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k24_sin_cos,X1))))).#trainneg
cnf(i_0_7675, plain, (k7_xcmplx_0(k3_xcmplx_0(X1,X1),X2)=k7_xcmplx_0(X1,k7_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14918, negated_conjecture, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_1363, plain, (k1_relset_1(X1,k3_relat_1(X5,X6))=k9_xtuple_0(k3_relat_1(X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_14944, negated_conjecture, (m1_subset_1(esk3_0,a_0_0_fdiff_8)|k1_zfmisc_1(k1_numbers)=k6_numbers|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_17808, plain, (v4_relat_1(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k3_taylor_1),k1_numbers))).#trainneg
cnf(i_0_17824, plain, (v5_relat_1(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k3_taylor_1),k1_numbers))).#trainneg
cnf(i_0_17864, plain, (v4_relat_1(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k24_sin_cos),k1_numbers))).#trainneg
cnf(i_0_17880, plain, (v5_relat_1(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k24_sin_cos),k1_numbers))).#trainneg
cnf(i_0_9846, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(X3,k29_sin_cos)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X4))))).#trainneg
cnf(i_0_7012, plain, (k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k3_taylor_1))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k3_taylor_1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_7589, plain, (k7_xcmplx_0(k5_square_1(X1),np__1)=k3_xcmplx_0(X1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14870, negated_conjecture, (r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k29_sin_cos),X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k29_sin_cos)))|~r2_hidden(esk6_2(X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k29_sin_cos)),esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_14925, negated_conjecture, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14880, negated_conjecture, (v1_xreal_0(X1)|~r2_hidden(X1,k3_taylor_1)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_15946, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k24_sin_cos,k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_relset_1(k1_numbers,k24_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14895, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_15029, negated_conjecture, (m1_subset_1(k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_6646, plain, (v4_relat_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10441, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_9408, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(a_0_0_fdiff_8,X1,X2,X3,X4,X5)),k1_zfmisc_1(k1_numbers))|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(a_0_0_fdiff_8,X1))))).#trainneg
cnf(i_0_10162, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k29_sin_cos))))).#trainneg
cnf(i_0_1372, plain, (r2_hidden(k1_funct_1(X1,X2),k1_relset_1(k1_numbers,k19_sin_cos))|~r2_hidden(X2,k9_xtuple_0(k3_relat_1(X1,k19_sin_cos)))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_14882, negated_conjecture, (v1_xreal_0(X1)|~r2_hidden(X1,k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14883, negated_conjecture, (v1_xreal_0(X1)|~r2_hidden(X1,k29_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_15350, negated_conjecture, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk3_0))|~m1_subset_1(k1_zfmisc_1(a_0_0_fdiff_8),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_15189, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_15197, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_15204, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_15211, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_15218, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_15166, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(esk3_0,k2_zfmisc_1(esk3_0,X2)))))).#trainneg
cnf(i_0_8522, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~r2_hidden(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7071, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~v1_relat_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,X2)))))).#trainneg
cnf(i_0_11818, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k3_taylor_1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_15765, negated_conjecture, (k1_fdiff_1(k3_taylor_1,esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0)))=k7_xcmplx_0(np__1,esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0)))|v1_xboole_0(a_0_0_fdiff_8)|r2_fdiff_1(k2_fdiff_1(X1,esk3_0),esk3_0)|~m1_subset_1(esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0)),a_0_0_fdiff_8)|~r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_12305, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k19_sin_cos,X1,X2,X3,k3_taylor_1,k29_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k19_sin_cos,X1))))).#trainneg
cnf(i_0_1373, plain, (r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,k19_sin_cos)))|~r2_hidden(k1_funct_1(X2,X1),k1_relset_1(k1_numbers,k19_sin_cos))|~r2_hidden(X1,k9_xtuple_0(X2))|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_9801, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k3_taylor_1))))).#trainneg
cnf(i_0_10130, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_10141, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_14955, negated_conjecture, (m1_subset_1(a_0_0_fdiff_8,a_0_0_fdiff_8)|k1_zfmisc_1(k1_numbers)=k6_numbers|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14961, negated_conjecture, (m1_subset_1(k5_numbers,a_0_0_fdiff_8)|k1_zfmisc_1(k1_numbers)=k6_numbers|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_15025, negated_conjecture, (m1_subset_1(esk3_0,a_0_0_fdiff_8)|k1_zfmisc_1(a_0_0_fdiff_8)=k6_numbers|~m1_subset_1(k1_zfmisc_1(a_0_0_fdiff_8),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_15562, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k3_taylor_1),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_6446, plain, (v4_relat_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k3_taylor_1),k1_numbers)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_11784, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_9812, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(X3,k3_taylor_1)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X4))))).#trainneg
cnf(i_0_12694, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k3_taylor_1,X1,X2,X3,k24_sin_cos,k19_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k3_taylor_1,X1))))).#trainneg
cnf(i_0_16971, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k24_sin_cos,X1,X2,X3,k29_sin_cos,k3_taylor_1)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k24_sin_cos,X1))))).#trainneg
cnf(i_0_10915, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_16968, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k24_sin_cos,X1,X2,X3,k19_sin_cos,k3_taylor_1)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k24_sin_cos,X1))))).#trainneg
cnf(i_0_10615, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k29_sin_cos))))).#trainneg
cnf(i_0_10152, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_10616, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k29_sin_cos))))).#trainneg
cnf(i_0_10626, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k29_sin_cos))))).#trainneg
cnf(i_0_15563, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k24_sin_cos),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_15564, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_15565, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_15566, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_15567, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k24_sin_cos),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_10859, plain, (k7_xcmplx_0(k5_square_1(X1),k7_xcmplx_0(np__1,X2))=k3_xcmplx_0(X1,k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9032, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(k5_xcmplx_0(X1),X1))=k3_square_1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10926, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k24_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_17023, plain, (r1_tarski(X1,X2)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k3_taylor_1)))))).#trainneg
cnf(i_0_10637, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos))))).#trainneg
cnf(i_0_1377, plain, (r2_hidden(esk6_2(k9_xtuple_0(k3_relat_1(k19_sin_cos,X1)),X2),k1_relset_1(k1_numbers,k19_sin_cos))|r2_fdiff_1(X2,k9_xtuple_0(k3_relat_1(k19_sin_cos,X1)))|~v3_rcomp_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,X1)))|~r1_tarski(k9_xtuple_0(k3_relat_1(k19_sin_cos,X1)),k1_relset_1(k1_numbers,X2))|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,X1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_10627, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k29_sin_cos))))).#trainneg
cnf(i_0_10638, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos))))).#trainneg
cnf(i_0_10659, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_15568, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k19_sin_cos),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_15569, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k29_sin_cos),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_15570, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_15571, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k24_sin_cos),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_15572, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k19_sin_cos),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_10651, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_15229, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_14260, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(X1,k29_sin_cos)),X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_12711, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k19_sin_cos)),k1_zfmisc_1(X1))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_7125, plain, (v1_relat_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_12435, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k29_sin_cos,X1,X2,X3,k19_sin_cos,k3_taylor_1)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k29_sin_cos,X1))))).#trainneg
cnf(i_0_10648, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_1379, plain, (r2_hidden(k1_funct_1(X1,X2),k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X2,k9_xtuple_0(k3_relat_1(X1,k29_sin_cos)))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_17925, negated_conjecture, (~v1_xboole_0(a_0_0_fdiff_8)|~r2_hidden(X1,k5_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_17917, negated_conjecture, (v1_xcmplx_0(k5_numbers)|~v1_membered(k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_17918, negated_conjecture, (v7_ordinal1(k5_numbers)|~v6_membered(k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_17919, negated_conjecture, (v1_xreal_0(k5_numbers)|~v3_membered(k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_19240, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),esk3_0)|~r1_tarski(esk3_0,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_20262, negated_conjecture, (r1_tarski(esk3_0,X1)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1),esk3_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_19005, plain, (k1_relset_1(X1,k3_relat_1(X2,k3_taylor_1))=k9_xtuple_0(k3_relat_1(X2,k3_taylor_1))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_1380, plain, (r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,k29_sin_cos)))|~r2_hidden(k1_funct_1(X2,X1),k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X1,k9_xtuple_0(X2))|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_18412, negated_conjecture, (~v1_xboole_0(a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_18424, negated_conjecture, (~v1_xboole_0(a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_18271, negated_conjecture, (~v1_xboole_0(a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k24_sin_cos))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_20458, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k24_sin_cos,k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_17932, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k24_sin_cos,k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_17939, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_taylor_1,k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_18138, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k24_sin_cos,k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_19189, plain, (k7_xcmplx_0(k5_square_1(X1),np__1)=k3_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_18019, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8)|~v3_rcomp_1(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0))|~r1_tarski(X1,esk3_0))).#trainneg
cnf(i_0_18018, negated_conjecture, (r2_fdiff_1(X1,X2)|m1_subset_1(esk6_2(X2,X1),a_0_0_fdiff_8)|~v3_rcomp_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(esk3_0))|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_19197, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(k5_square_1(X2),np__1))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1382, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k29_sin_cos,X1)),k1_relset_1(k1_numbers,k29_sin_cos))|~v1_relat_1(X1))).#trainneg
cnf(i_0_17933, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(k24_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_17940, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(k3_taylor_1,k24_sin_cos))))).#trainneg
cnf(i_0_18146, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(k19_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_18145, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k19_sin_cos,k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_18152, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k29_sin_cos,k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_17928, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k24_sin_cos,k3_taylor_1)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_19060, plain, (k1_relset_1(X1,k3_relat_1(X2,k19_sin_cos))=k9_xtuple_0(k3_relat_1(X2,k19_sin_cos))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_18570, negated_conjecture, (r1_tarski(esk3_0,X1)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_19335, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_14876, negated_conjecture, (v3_membered(X1)|~r1_tarski(X1,esk3_0)|~m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_1383, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k29_sin_cos,X1)),k1_zfmisc_1(k1_relset_1(k1_numbers,k29_sin_cos)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_18153, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(k29_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_10660, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_10879, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_17935, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_taylor_1,k24_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_18134, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k24_sin_cos,k24_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_18141, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k24_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_18148, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k29_sin_cos,k24_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_17916, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(a_0_0_fdiff_8))|r2_hidden(k5_numbers,k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_17942, negated_conjecture, (m1_subset_1(k1_relset_1(esk3_0,k2_zfmisc_1(esk3_0,X1)),k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))))).#trainneg
cnf(i_0_14877, negated_conjecture, (v1_xreal_0(X1)|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_1384, plain, (r2_hidden(esk6_2(k9_xtuple_0(k3_relat_1(k29_sin_cos,X1)),X2),k1_relset_1(k1_numbers,k29_sin_cos))|r2_fdiff_1(X2,k9_xtuple_0(k3_relat_1(k29_sin_cos,X1)))|~v3_rcomp_1(k9_xtuple_0(k3_relat_1(k29_sin_cos,X1)))|~r1_tarski(k9_xtuple_0(k3_relat_1(k29_sin_cos,X1)),k1_relset_1(k1_numbers,X2))|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(k9_xtuple_0(k3_relat_1(k29_sin_cos,X1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_10880, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_18131, negated_conjecture, (r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|~r2_hidden(X1,k9_xtuple_0(X2))|~v1_funct_1(X2)|~v1_relat_1(X2)|~r2_hidden(k1_funct_1(X2,X1),esk3_0))).#trainneg
cnf(i_0_19061, plain, (k1_relset_1(X1,k3_relat_1(X2,k29_sin_cos))=k9_xtuple_0(k3_relat_1(X2,k29_sin_cos))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_20420, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_18703, negated_conjecture, (r1_tarski(esk3_0,X1)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k29_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k29_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_19540, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k3_taylor_1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_18086, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k24_sin_cos,X1,X2,X3,k3_taylor_1,k19_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k24_sin_cos,X1)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_19196, plain, (k7_xcmplx_0(k7_xcmplx_0(k5_square_1(X1),np__1),X2)=k7_xcmplx_0(X1,k7_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14878, negated_conjecture, (v1_xcmplx_0(X1)|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_1388, plain, (r1_fdiff_1(X1,k6_numbers)|~v3_rcomp_1(k5_numbers)|~r2_fdiff_1(X1,k5_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_15573, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k29_sin_cos),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_15574, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k3_taylor_1),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_15575, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k24_sin_cos),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_15576, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k19_sin_cos),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_19716, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_19741, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(X1,k3_taylor_1)),X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_18085, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k24_sin_cos,X1,X2,X3,k3_taylor_1,k29_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k24_sin_cos,X1)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1396, negated_conjecture, (r2_fdiff_1(k2_fdiff_1(X1,X2),esk3_0)|m1_subset_1(esk6_2(esk3_0,k2_fdiff_1(X1,X2)),k1_numbers)|~r1_tarski(esk3_0,X2)|~v1_funct_1(k2_fdiff_1(X1,X2))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_15577, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k29_sin_cos),k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_16239, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k24_sin_cos,k29_sin_cos)),k1_relset_1(k1_numbers,k24_sin_cos)))).#trainneg
cnf(i_0_16410, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k3_taylor_1),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_16411, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k24_sin_cos),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_16412, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_19137, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k3_taylor_1)),X1)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_18923, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(np__1,X1))=k3_xcmplx_0(X1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_19127, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k19_sin_cos,X1,X2,X3,k24_sin_cos,k3_taylor_1)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k19_sin_cos,X1)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_18462, plain, (r1_tarski(X1,a_0_0_fdiff_8)|~v1_relat_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k9_xtuple_0(k3_relat_1(k3_taylor_1,X2)))))).#trainneg
cnf(i_0_19628, plain, (r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,k19_sin_cos)))|~r2_hidden(k1_seq_1(X2,X1),k1_relset_1(k1_numbers,k19_sin_cos))|~r2_hidden(X1,k9_xtuple_0(X2))|~v1_funct_1(X2)|~v1_relat_1(X2)|~v3_valued_0(X2))).#trainneg
cnf(i_0_14888, negated_conjecture, (v1_xreal_0(X1)|~r1_tarski(X1,esk3_0)|~m1_subset_1(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_18102, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k3_taylor_1,X2)))|~v1_relat_1(X2))).#trainneg
cnf(i_0_1403, plain, (r2_hidden(k1_funct_1(X1,X2),k1_relset_1(k1_numbers,k2_fdiff_1(X3,X4)))|~r2_hidden(X2,k9_xtuple_0(k3_relat_1(X1,k2_fdiff_1(X3,X4))))|~v1_funct_1(k2_fdiff_1(X3,X4))|~v1_funct_1(X1)|~v1_relat_1(k2_fdiff_1(X3,X4))|~v1_relat_1(X1)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4097, plain, (k7_xcmplx_0(k3_square_1(X1),np__1)=k3_square_1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6106, plain, (v1_xcmplx_0(k1_seq_1(k3_relat_1(k24_sin_cos,k3_taylor_1),X1)))).#trainneg
cnf(i_0_16413, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_16414, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_16415, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k24_sin_cos),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_16416, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k19_sin_cos),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_16417, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k29_sin_cos),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_19351, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(a_0_0_fdiff_8,X1,X2,X3,k24_sin_cos,k29_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(a_0_0_fdiff_8,X1)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_19225, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(X1,X1))=k5_square_1(X1)|~v1_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6113, plain, (v3_valued_0(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k3_taylor_1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10893, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k24_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_8762, plain, (k7_xcmplx_0(k3_xcmplx_0(np__1,X1),np__1)=X1|~v1_xcmplx_0(k3_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_14613, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k19_sin_cos),X1)),k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k19_sin_cos))))|~v1_relat_1(X1))).#trainneg
cnf(i_0_19191, plain, (m1_subset_1(k7_xcmplx_0(k5_square_1(X1),np__1),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_19128, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k29_sin_cos,X1,X2,X3,k24_sin_cos,k3_taylor_1)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k29_sin_cos,X1)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_18105, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,X1)),k1_zfmisc_1(k1_numbers))|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_12269, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k19_sin_cos)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_14889, negated_conjecture, (v1_xcmplx_0(X1)|~r1_tarski(X1,esk3_0)|~m1_subset_1(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_10890, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k24_sin_cos))))).#trainneg
cnf(i_0_1404, plain, (r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,k2_fdiff_1(X3,X4))))|~r2_hidden(k1_funct_1(X2,X1),k1_relset_1(k1_numbers,k2_fdiff_1(X3,X4)))|~r2_hidden(X1,k9_xtuple_0(X2))|~v1_funct_1(k2_fdiff_1(X3,X4))|~v1_funct_1(X2)|~v1_relat_1(k2_fdiff_1(X3,X4))|~v1_relat_1(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_6118, plain, (v1_xcmplx_0(k1_seq_1(k3_relat_1(k24_sin_cos,k29_sin_cos),X1)))).#trainneg
cnf(i_0_6174, plain, (v1_xcmplx_0(k1_seq_1(k3_relat_1(k24_sin_cos,k19_sin_cos),X1)))).#trainneg
cnf(i_0_6183, plain, (v1_xcmplx_0(k1_seq_1(k3_relat_1(k3_taylor_1,k24_sin_cos),X1)))).#trainneg
cnf(i_0_16418, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_16419, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k24_sin_cos),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_16420, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k19_sin_cos),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_16421, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k29_sin_cos),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_16422, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k3_taylor_1),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_15907, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_taylor_1,X1)),k1_numbers)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9439, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(k5_xcmplx_0(X2),k3_xcmplx_0(X1,X2)))=k5_square_1(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_6166, plain, (v3_valued_0(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_12173, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_6314, plain, (k8_real_1(k7_xcmplx_0(np__1,k1_seq_1(X1,X2)),k1_fdiff_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|v1_xboole_0(a_0_0_fdiff_8)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(k1_seq_1(X1,X2))|~m1_subset_1(k1_seq_1(X1,X2),a_0_0_fdiff_8))).#trainneg
cnf(i_0_10901, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_1405, plain, (X1=k2_fdiff_1(X2,k9_xtuple_0(k3_relat_1(k2_fdiff_1(X3,X4),X5)))|r2_hidden(esk1_3(X2,k9_xtuple_0(k3_relat_1(k2_fdiff_1(X3,X4),X5)),X1),k1_relset_1(k1_numbers,k2_fdiff_1(X3,X4)))|k1_relset_1(k1_numbers,X1)!=k9_xtuple_0(k3_relat_1(k2_fdiff_1(X3,X4),X5))|~r2_fdiff_1(X2,k9_xtuple_0(k3_relat_1(k2_fdiff_1(X3,X4),X5)))|~v1_funct_1(X5)|~v1_funct_1(k2_fdiff_1(X3,X4))|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_relat_1(X5)|~v1_relat_1(k2_fdiff_1(X3,X4))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_6284, plain, (v1_xcmplx_0(k1_seq_1(k3_relat_1(k19_sin_cos,k24_sin_cos),X1)))).#trainneg
cnf(i_0_6287, plain, (v1_xcmplx_0(k1_seq_1(k3_relat_1(k29_sin_cos,k24_sin_cos),X1)))).#trainneg
cnf(i_0_6435, plain, (k9_xtuple_0(k3_relat_1(k24_sin_cos,k3_taylor_1))=k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_16423, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k24_sin_cos),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_16424, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k19_sin_cos),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_16425, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k29_sin_cos),k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_17235, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k3_taylor_1),k24_sin_cos)))).#trainneg
cnf(i_0_17236, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k3_taylor_1),k3_taylor_1)))).#trainneg
cnf(i_0_6529, plain, (v4_relat_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos),k1_numbers)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_7032, plain, (v1_relat_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_9576, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(X1,X2)),k1_numbers)|~m1_subset_1(k9_xtuple_0(X1),k1_zfmisc_1(a_0_0_fdiff_8))|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12312, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_9256, plain, (k7_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,X1)),k1_fdiff_1(k29_sin_cos,X1))=k5_square_1(k5_square_1(k1_seq_1(k19_sin_cos,X1)))|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,X1)))|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1610, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_17933, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_20666, negated_conjecture, (v1_xcmplx_0(k5_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~v3_membered(k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_21356, plain, (v3_membered(k9_xtuple_0(k3_relat_1(k3_taylor_1,k24_sin_cos))))).#trainneg
cnf(i_0_21367, plain, (v3_membered(k9_xtuple_0(k3_relat_1(k24_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_21378, plain, (v3_membered(k9_xtuple_0(k3_relat_1(k19_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_21389, plain, (v3_membered(k9_xtuple_0(k3_relat_1(k29_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_23159, negated_conjecture, (r1_tarski(esk3_0,X1)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_20869, negated_conjecture, (v1_xreal_0(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_23182, plain, (r1_tarski(X1,X2)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v3_rcomp_1(X1)|~r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k29_sin_cos),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_22973, plain, (r1_tarski(X1,X2)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v3_rcomp_1(X1)|~r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_23329, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k19_sin_cos),k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_21359, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k3_taylor_1,k24_sin_cos))))).#trainneg
cnf(i_0_17932, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_20804, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k24_sin_cos,k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_20800, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k24_sin_cos,k29_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_15951, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_22037, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k24_sin_cos,k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_relset_1(k1_numbers,k24_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_22751, plain, (r1_tarski(X1,X2)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v3_rcomp_1(X1)|~r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k19_sin_cos),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k19_sin_cos))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_22972, plain, (r1_tarski(X1,X2)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)))))).#trainneg
cnf(i_0_22750, plain, (r1_tarski(X1,X2)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k19_sin_cos)))))).#trainneg
cnf(i_0_23181, plain, (r1_tarski(X1,X2)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k29_sin_cos)))))).#trainneg
cnf(i_0_23330, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k29_sin_cos),k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_21370, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k24_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_15946, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_relset_1(k1_numbers,k24_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_20815, negated_conjecture, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_taylor_1,k24_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_20823, negated_conjecture, (m1_subset_1(k9_xtuple_0(k3_relat_1(k24_sin_cos,k24_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_21155, negated_conjecture, (m1_subset_1(k9_xtuple_0(k3_relat_1(k19_sin_cos,k24_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_21163, negated_conjecture, (m1_subset_1(k9_xtuple_0(k3_relat_1(k29_sin_cos,k24_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_17237, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k3_taylor_1),k19_sin_cos)))).#trainneg
cnf(i_0_22020, negated_conjecture, (v1_xreal_0(esk6_2(esk3_0,k2_fdiff_1(X1,X2)))|r2_fdiff_1(k2_fdiff_1(X1,X2),esk3_0)|~r1_tarski(esk3_0,X2)|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_22397, plain, (v1_xreal_0(k3_square_1(X1))|~m1_subset_1(k3_square_1(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_23359, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_21381, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k19_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_14329, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k3_taylor_1)),k1_relset_1(k1_numbers,k24_sin_cos)))).#trainneg
cnf(i_0_23088, plain, (~v1_xboole_0(a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_22565, plain, (v1_xreal_0(X1)|~m1_subset_1(k3_xcmplx_0(np__1,X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_20805, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(k24_sin_cos,k29_sin_cos))))).#trainneg
cnf(i_0_17238, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k3_taylor_1),k29_sin_cos)))).#trainneg
cnf(i_0_17641, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k24_sin_cos),k24_sin_cos)))).#trainneg
cnf(i_0_17642, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k24_sin_cos),k3_taylor_1)))).#trainneg
cnf(i_0_17643, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k24_sin_cos),k19_sin_cos)))).#trainneg
cnf(i_0_17644, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k24_sin_cos),k29_sin_cos)))).#trainneg
cnf(i_0_22606, plain, (k7_xcmplx_0(k3_xcmplx_0(X1,np__1),np__1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_20767, plain, (r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,k29_sin_cos)))|~r2_hidden(k1_seq_1(X2,X1),k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X1,k9_xtuple_0(X2))|~v1_funct_1(X2)|~v1_relat_1(X2)|~v3_valued_0(X2))).#trainneg
cnf(i_0_21392, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k29_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_1406, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k2_fdiff_1(X1,X2),X3)),k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2)))|~v1_relat_1(X3)|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_17697, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k19_sin_cos),k24_sin_cos)))).#trainneg
cnf(i_0_17698, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k19_sin_cos),k3_taylor_1)))).#trainneg
cnf(i_0_17699, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k19_sin_cos),k19_sin_cos)))).#trainneg
cnf(i_0_17700, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k19_sin_cos),k29_sin_cos)))).#trainneg
cnf(i_0_17753, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k29_sin_cos),k24_sin_cos)))).#trainneg
cnf(i_0_21384, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10316, plain, (v1_xcmplx_0(k1_funct_1(X1,X2))|~v1_membered(k1_numbers)|~r2_hidden(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11836, plain, (m1_subset_1(k1_relset_1(k3_taylor_1,k2_zfmisc_1(k3_taylor_1,X1)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_10912, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_1407, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k2_fdiff_1(X1,X2),X3)),k1_zfmisc_1(k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2))))|~v1_relat_1(X3)|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_17754, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k29_sin_cos),k3_taylor_1)))).#trainneg
cnf(i_0_17755, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k29_sin_cos),k19_sin_cos)))).#trainneg
cnf(i_0_17756, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k29_sin_cos),k29_sin_cos)))).#trainneg
cnf(i_0_17809, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k3_taylor_1),k24_sin_cos)))).#trainneg
cnf(i_0_17810, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k3_taylor_1),k3_taylor_1)))).#trainneg
cnf(i_0_10165, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k29_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_19377, plain, (m1_subset_1(k1_funct_1(X1,X2),k1_numbers)|~r2_hidden(X2,k9_xtuple_0(k3_relat_1(X1,k19_sin_cos)))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12355, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k29_sin_cos)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_9146, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,k3_xcmplx_0(X1,X2)))=k3_square_1(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12737, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k3_taylor_1,X1,X2,X3,k24_sin_cos,k29_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k3_taylor_1,X1))))).#trainneg
cnf(i_0_12752, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k19_sin_cos,X1,X2,X3,k24_sin_cos,k29_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k19_sin_cos,X1))))).#trainneg
cnf(i_0_10923, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_1408, plain, (r2_hidden(esk6_2(k9_xtuple_0(k3_relat_1(k2_fdiff_1(X1,X2),X3)),X4),k1_relset_1(k1_numbers,k2_fdiff_1(X1,X2)))|r2_fdiff_1(X4,k9_xtuple_0(k3_relat_1(k2_fdiff_1(X1,X2),X3)))|~v3_rcomp_1(k9_xtuple_0(k3_relat_1(k2_fdiff_1(X1,X2),X3)))|~r1_tarski(k9_xtuple_0(k3_relat_1(k2_fdiff_1(X1,X2),X3)),k1_relset_1(k1_numbers,X4))|~v1_funct_1(X3)|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X4)|~m1_subset_1(k9_xtuple_0(k3_relat_1(k2_fdiff_1(X1,X2),X3)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_relat_1(X3)|~v1_relat_1(k2_fdiff_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_6523, plain, (k9_xtuple_0(k3_relat_1(k24_sin_cos,k29_sin_cos))=k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k29_sin_cos)))).#trainneg
cnf(i_0_6650, plain, (k9_xtuple_0(k3_relat_1(k24_sin_cos,k19_sin_cos))=k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_6770, plain, (k9_xtuple_0(k3_relat_1(k3_taylor_1,k24_sin_cos))=k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k24_sin_cos)))).#trainneg
cnf(i_0_17811, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k3_taylor_1),k19_sin_cos)))).#trainneg
cnf(i_0_17812, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k3_taylor_1),k29_sin_cos)))).#trainneg
cnf(i_0_17865, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k24_sin_cos),k24_sin_cos)))).#trainneg
cnf(i_0_17866, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k24_sin_cos),k3_taylor_1)))).#trainneg
cnf(i_0_17867, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k24_sin_cos),k19_sin_cos)))).#trainneg
cnf(i_0_14662, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k19_sin_cos),X1)),k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k19_sin_cos)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_11273, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k29_sin_cos)),k1_zfmisc_1(X1))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_14034, plain, (r2_hidden(X1,a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,X2,X3,X4,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,X2))))).#trainneg
cnf(i_0_10109, plain, (v1_xcmplx_0(k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6181, plain, (v3_valued_0(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_19350, plain, (m1_subset_1(k1_relset_1(a_0_0_fdiff_8,k1_partfun1(a_0_0_fdiff_8,X1,X2,X3,X4,X5)),k1_zfmisc_1(k1_numbers))|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(a_0_0_fdiff_8,X1)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_12826, plain, (r2_hidden(X1,a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,X2,X3,X4,k3_taylor_1,k19_sin_cos)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,X2))))).#trainneg
cnf(i_0_23360, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(X2,k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_18786, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(a_0_0_fdiff_8,k2_zfmisc_1(a_0_0_fdiff_8,X2))))).#trainneg
cnf(i_0_1409, plain, (r1_tarski(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k2_fdiff_1(X5,X6),X7)),k1_relset_1(k1_numbers,k2_fdiff_1(X5,X6)))|~v1_funct_1(X7)|~v1_funct_1(k2_fdiff_1(X5,X6))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k2_fdiff_1(X5,X6),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_6780, plain, (k9_xtuple_0(k3_relat_1(k19_sin_cos,k24_sin_cos))=k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k24_sin_cos)))).#trainneg
cnf(i_0_6979, plain, (k9_xtuple_0(k3_relat_1(k29_sin_cos,k24_sin_cos))=k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k24_sin_cos)))).#trainneg
cnf(i_0_11984, plain, (r1_tarski(k9_xtuple_0(k2_zfmisc_1(a_0_0_fdiff_8,X1)),k1_numbers))).#trainneg
cnf(i_0_17868, plain, (v1_funct_1(k3_relat_1(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k24_sin_cos),k29_sin_cos)))).#trainneg
cnf(i_0_20435, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k3_taylor_1),k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_20436, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k24_sin_cos),k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_20437, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_20438, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_22892, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(k5_xcmplx_0(X1),X1))=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_21078, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(k29_sin_cos,X2)))|~v1_relat_1(X2))).#trainneg
cnf(i_0_8802, plain, (v1_xreal_0(k7_xcmplx_0(X1,np__1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7876, plain, (v1_relat_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k29_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_19360, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(a_0_0_fdiff_8,X1,X2,X3,k24_sin_cos,k3_taylor_1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(a_0_0_fdiff_8,X1)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_12441, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k3_taylor_1)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_6661, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(X2,k2_zfmisc_1(X2,X3)))))).#trainneg
cnf(i_0_1442, plain, (r2_hidden(k1_funct_1(X1,X2),k1_relset_1(k1_numbers,k24_sin_cos))|~r2_hidden(X2,k9_xtuple_0(k3_relat_1(X1,k24_sin_cos)))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11985, plain, (r1_tarski(k9_xtuple_0(k2_zfmisc_1(k5_numbers,X1)),k1_numbers))).#trainneg
cnf(i_0_20445, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_20446, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k24_sin_cos),k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_20447, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k19_sin_cos),k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_20448, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k29_sin_cos),k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_20450, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_6766, plain, (v4_relat_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k19_sin_cos),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_12501, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,np__1))=k5_square_1(X1)|~v1_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12710, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k29_sin_cos,X1,X2,X3,k24_sin_cos,k19_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k29_sin_cos,X1))))).#trainneg
cnf(i_0_8510, plain, (m1_subset_1(esk6_2(X1,k2_fdiff_1(X2,X1)),k1_numbers)|r2_fdiff_1(k2_fdiff_1(X2,X1),X1)|~v3_rcomp_1(X1)|~r2_fdiff_1(X2,X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_11896, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_7332, plain, (v1_relat_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1443, plain, (r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,k24_sin_cos)))|~r2_hidden(k1_funct_1(X2,X1),k1_relset_1(k1_numbers,k24_sin_cos))|~r2_hidden(X1,k9_xtuple_0(X2))|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2546, plain, (v7_ordinal1(X1)|~v6_membered(k1_numbers)|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_17805, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k3_taylor_1)))).#trainneg
cnf(i_0_20451, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k24_sin_cos),k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_20452, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k19_sin_cos),k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_20453, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k29_sin_cos),k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_20454, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k3_taylor_1),k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_20455, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k24_sin_cos),k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_20142, plain, (k7_xcmplx_0(k5_square_1(X1),np__1)=k3_xcmplx_0(X1,k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6178, plain, (v3_valued_0(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_20218, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(k7_xcmplx_0(np__1,X1),X1))=k3_square_1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_17906, plain, (m1_subset_1(k3_relat_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1),k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1445, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k24_sin_cos,X1)),k1_relset_1(k1_numbers,k24_sin_cos))|~v1_relat_1(X1))).#trainneg
cnf(i_0_14903, negated_conjecture, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,k5_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_20456, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k19_sin_cos),k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_20457, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k29_sin_cos),k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_23364, negated_conjecture, (r1_fdiff_1(X1,esk4_0)|~r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_23391, negated_conjecture, (v7_ordinal1(esk4_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_22037, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_relset_1(k1_numbers,k24_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_23367, negated_conjecture, (r2_hidden(esk4_0,k1_numbers))).#trainneg
cnf(i_0_13354, plain, (v3_membered(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_23400, negated_conjecture, (m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v1_xcmplx_0(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_23398, negated_conjecture, (r1_tarski(esk3_0,X1)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_23401, negated_conjecture, (r2_hidden(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v1_xcmplx_0(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_24156, plain, (r1_tarski(X1,X2)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v3_rcomp_1(X1)|~r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k29_sin_cos),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k29_sin_cos))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_23728, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8)|~v3_rcomp_1(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_20805, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k29_sin_cos))))).#trainneg
cnf(i_0_21356, plain, (v3_membered(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k24_sin_cos))))).#trainneg
cnf(i_0_21378, plain, (v3_membered(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_21389, plain, (v3_membered(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_20458, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_23540, negated_conjecture, (r1_tarski(esk3_0,X1)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k19_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k19_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_25207, plain, (r1_tarski(X1,X2)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v3_rcomp_1(X1)|~r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k3_taylor_1),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k3_taylor_1))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_23725, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers)|~v3_rcomp_1(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_24730, negated_conjecture, (r2_hidden(esk6_2(esk3_0,X1),a_0_0_fdiff_8)|v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_25082, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(k7_xcmplx_0(np__1,X1),X1))=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_25825, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(a_0_0_fdiff_8))|~r1_tarski(X1,k5_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_20804, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_24785, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9771, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k19_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_16409, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_24554, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k24_sin_cos)),a_0_0_fdiff_8))).#trainneg
cnf(i_0_23917, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|m1_subset_1(esk6_2(esk3_0,X1),k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_23727, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),X1)|~v3_rcomp_1(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_24130, plain, (m1_subset_1(k1_funct_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(k1_seq_1(X1,X2),k1_relset_1(k1_numbers,k19_sin_cos))|~r2_hidden(X2,k9_xtuple_0(X1))|~v3_valued_0(X1))).#trainneg
cnf(i_0_24155, plain, (r1_tarski(X1,X2)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k29_sin_cos)))))).#trainneg
cnf(i_0_23904, plain, (m1_subset_1(k1_funct_1(X1,X2),k1_relset_1(k1_numbers,k29_sin_cos))|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(k1_seq_1(X1,X2),k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X2,k9_xtuple_0(X1))|~v3_valued_0(X1))).#trainneg
cnf(i_0_25808, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k24_sin_cos,X1)),k1_zfmisc_1(k1_relset_1(k1_numbers,k24_sin_cos)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_24936, plain, (k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_16239, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k29_sin_cos)),k1_relset_1(k1_numbers,k24_sin_cos)))).#trainneg
cnf(i_0_24000, plain, (v1_relat_1(k1_relset_1(k3_taylor_1,k2_zfmisc_1(k3_taylor_1,X1))))).#trainneg
cnf(i_0_23800, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(X1,np__1),k1_numbers))).#trainneg
cnf(i_0_23847, plain, (v1_xreal_0(X1)|~m1_subset_1(k3_xcmplx_0(X1,np__1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_16449, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_16450, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_24985, plain, (k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k24_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k24_sin_cos))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_25206, plain, (r1_tarski(X1,X2)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k3_taylor_1)))))).#trainneg
cnf(i_0_24850, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(esk3_0,k3_rfunct_1(esk3_0,X2,X3,X4))))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v3_membered(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(esk3_0,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk3_0,X2))))).#trainneg
cnf(i_0_24588, plain, (k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k24_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k24_sin_cos))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_23986, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_25461, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_25831, negated_conjecture, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,X2)|~r1_tarski(X2,k5_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_24474, plain, (~v1_xboole_0(k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k29_sin_cos))))).#trainneg
cnf(i_0_24523, plain, (~v1_xboole_0(k1_relset_1(k1_numbers,k19_sin_cos))|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_13359, plain, (~v1_xboole_0(k1_relset_1(k1_numbers,k24_sin_cos))|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_12022, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_18145, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_18152, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_20769, negated_conjecture, (r2_fdiff_1(k3_taylor_1,k1_relset_1(k1_numbers,k19_sin_cos))|~v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~v3_rcomp_1(k1_relset_1(k1_numbers,k19_sin_cos))|~r1_tarski(k1_relset_1(k1_numbers,k19_sin_cos),a_0_0_fdiff_8))).#trainneg
cnf(i_0_25613, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,np__1))=k3_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_25101, plain, (k7_xcmplx_0(np__1,k5_square_1(k5_xcmplx_0(X1)))=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_15029, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k19_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_24574, plain, (~v1_xboole_0(k1_relset_1(k1_numbers,k24_sin_cos))|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k24_sin_cos))))).#trainneg
cnf(i_0_24926, plain, (~v1_xboole_0(k1_relset_1(k1_numbers,k24_sin_cos))|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_24977, plain, (~v1_xboole_0(k1_relset_1(k1_numbers,k24_sin_cos))|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_20780, negated_conjecture, (r2_fdiff_1(k3_taylor_1,k1_relset_1(k1_numbers,k29_sin_cos))|~v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~v3_rcomp_1(k1_relset_1(k1_numbers,k29_sin_cos))|~r1_tarski(k1_relset_1(k1_numbers,k29_sin_cos),a_0_0_fdiff_8))).#trainneg
cnf(i_0_20791, negated_conjecture, (r2_fdiff_1(k3_taylor_1,k1_relset_1(k1_numbers,k24_sin_cos))|~v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~v3_rcomp_1(k1_relset_1(k1_numbers,k24_sin_cos))|~r1_tarski(k1_relset_1(k1_numbers,k24_sin_cos),a_0_0_fdiff_8))).#trainneg
cnf(i_0_21324, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k19_sin_cos),k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_21890, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k3_taylor_1),k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_21891, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k24_sin_cos),k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_24836, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(X1,k1_zfmisc_1(k9_xtuple_0(k1_partfun1(esk3_0,X2,X3,X4,X5,X6))))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(esk3_0,X2))))).#trainneg
cnf(i_0_25628, plain, (k7_xcmplx_0(k5_square_1(X1),np__1)=k3_xcmplx_0(X1,k7_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_25131, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~v1_relat_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k9_xtuple_0(k3_relat_1(k29_sin_cos,X2)))))).#trainneg
cnf(i_0_25612, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,np__1))=k3_xcmplx_0(X1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_25631, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,np__1))=k7_xcmplx_0(k3_xcmplx_0(X1,X1),np__1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_24894, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k24_sin_cos),X1)),k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k24_sin_cos))))|~v1_relat_1(X1))).#trainneg
cnf(i_0_23413, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,k19_sin_cos),k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_13360, plain, (r2_fdiff_1(k24_sin_cos,k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k19_sin_cos)))|~v3_rcomp_1(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k19_sin_cos)))|~v1_xboole_0(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_21892, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_21893, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_21900, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_21901, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k24_sin_cos),k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_21902, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k19_sin_cos),k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_23848, plain, (r2_hidden(X1,k1_numbers)|~m1_subset_1(k3_xcmplx_0(X1,np__1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_24783, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X1))|~v1_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_21067, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(k3_xcmplx_0(X2,X2),np__1))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_8288, plain, (k7_xcmplx_0(k3_square_1(X1),np__1)=k3_xcmplx_0(X1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_17794, plain, (m1_subset_1(k3_relat_1(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k3_taylor_1),k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_12167, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k29_sin_cos,X1,X2,X3,k3_taylor_1,k19_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k29_sin_cos,X1))))).#trainneg
cnf(i_0_12673, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k3_taylor_1)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_25814, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k24_sin_cos))|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(k24_sin_cos,X2)))|~v1_relat_1(X2))).#trainneg
cnf(i_0_10458, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k29_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_23596, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,k29_sin_cos),k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_12016, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_relset_1(k1_numbers,k24_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_21903, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k29_sin_cos),k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_21905, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_21906, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k24_sin_cos),k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_21908, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k29_sin_cos),k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_10855, plain, (k7_xcmplx_0(k5_square_1(X1),k5_xcmplx_0(X1))=k3_xcmplx_0(X1,k5_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10728, plain, (k3_xcmplx_0(np__1,k3_xcmplx_0(X2,X1))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10904, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k24_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_12837, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),X5)|m1_subset_1(esk6_2(X5,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos)),k1_numbers)|~v3_rcomp_1(X5)|~r1_tarski(X5,k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_23057, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),X1)),k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k3_taylor_1))))|~v1_relat_1(X1))).#trainneg
cnf(i_0_8303, plain, (k7_xcmplx_0(k3_square_1(X1),k7_xcmplx_0(np__1,X2))=k3_xcmplx_0(X1,k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_23716, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_11937, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_21909, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k3_taylor_1),k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_21910, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k24_sin_cos),k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_21911, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k19_sin_cos),k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_21912, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k29_sin_cos),k3_taylor_1)),k1_numbers))).#trainneg
cnf(i_0_24555, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k24_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_8299, plain, (k7_xcmplx_0(k3_square_1(X1),k5_xcmplx_0(X1))=k3_xcmplx_0(X1,k5_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8334, plain, (r1_tarski(X1,k9_xtuple_0(X2))|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(X2,k24_sin_cos)))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_9829, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(X3,k24_sin_cos)))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X4))))).#trainneg
cnf(i_0_6110, plain, (v3_valued_0(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_24782, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X1))|~v1_xcmplx_0(k3_square_1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4340, plain, (r1_tarski(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,X5)),k1_relset_1(k1_numbers,k19_sin_cos))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_18139, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(k24_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_5707, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k19_sin_cos)),k1_relset_1(k1_numbers,k24_sin_cos)))).#trainneg
cnf(i_0_15098, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_relset_1(k1_numbers,k19_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_15697, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_relset_1(k1_numbers,k29_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_1633, plain, (v1_xreal_0(k3_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2275, negated_conjecture, (~v1_xboole_0(k1_relset_1(k1_numbers,X1))|~r2_hidden(X2,esk3_0)|~r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_11851, plain, (m1_subset_1(k1_relset_1(k19_sin_cos,k2_zfmisc_1(k19_sin_cos,X1)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_19136, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k3_taylor_1)),k1_zfmisc_1(X1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_17738, plain, (m1_subset_1(k3_relat_1(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k3_taylor_1),k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_17276, plain, (m1_subset_1(k3_relat_1(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k3_taylor_1),k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_23419, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r1_tarski(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_9763, plain, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_relset_1(k1_numbers,k24_sin_cos))))).#trainneg
cnf(i_0_25857, negated_conjecture, (v7_ordinal1(esk4_0)|~v6_membered(a_0_0_fdiff_8))).#trainneg
cnf(i_0_25862, negated_conjecture, (v7_ordinal1(esk4_0)|~v6_membered(esk3_0))).#trainneg
cnf(i_0_25867, negated_conjecture, (v1_int_1(esk4_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_688, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k3_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),esk4_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~m2_subset_1(k3_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),esk4_0),k1_numbers,a_0_0_fdiff_8)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers))).#trainneg
cnf(i_0_1765, negated_conjecture, (k7_xcmplx_0(np__1,k8_real_1(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers))).#trainneg
cnf(i_0_1624, negated_conjecture, (k7_xcmplx_0(np__1,k3_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),esk4_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers))).#trainneg
cnf(i_0_686, negated_conjecture, (k10_real_1(np__1,k3_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),esk4_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers))).#trainneg
cnf(i_0_26212, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_hidden(esk6_2(esk3_0,X1),k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_1552, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k3_xcmplx_0(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~m2_subset_1(k3_xcmplx_0(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))),k1_numbers,a_0_0_fdiff_8)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers))).#trainneg
cnf(i_0_1432, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_8)|k1_fdiff_1(k3_taylor_1,k3_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),esk4_0))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~m1_subset_1(k3_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers))).#trainneg
cnf(i_0_27791, plain, (r1_tarski(X1,X2)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v3_rcomp_1(X1)|~r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k3_taylor_1),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k3_taylor_1))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_17940, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k24_sin_cos))))).#trainneg
cnf(i_0_26322, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),X1)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v3_rcomp_1(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_4597, negated_conjecture, (k7_xcmplx_0(np__1,k3_xcmplx_0(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers))).#trainneg
cnf(i_0_26574, negated_conjecture, (r1_tarski(esk3_0,X1)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1),esk3_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4602, negated_conjecture, (k7_xcmplx_0(k7_xcmplx_0(np__1,esk4_0),k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers))).#trainneg
cnf(i_0_1525, negated_conjecture, (k10_real_1(np__1,k3_xcmplx_0(esk4_0,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(esk4_0)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers))).#trainneg
cnf(i_0_26343, negated_conjecture, (r1_tarski(esk3_0,X1)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_27581, plain, (k7_xcmplx_0(k3_square_1(X1),np__1)=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_26715, negated_conjecture, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,X2)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(X2,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_21155, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k24_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_28355, plain, (~v1_xboole_0(a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k24_sin_cos))))).#trainneg
cnf(i_0_28633, plain, (v1_relat_1(k1_relset_1(k19_sin_cos,k2_zfmisc_1(k19_sin_cos,X1))))).#trainneg
cnf(i_0_27590, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(k3_square_1(X2),np__1))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_26551, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k24_sin_cos)),X1)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_26656, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k24_sin_cos)),X1)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_26423, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos)),X1)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_26159, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_18146, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_25939, negated_conjecture, (r2_hidden(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v3_rcomp_1(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_26539, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k3_taylor_1,X1,X2,X3,k29_sin_cos,k24_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k3_taylor_1,X1)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_26410, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k3_taylor_1,X1,X2,X3,k19_sin_cos,k24_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k3_taylor_1,X1)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_26546, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k24_sin_cos)),k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_27575, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(k3_square_1(X1),np__1))=k3_xcmplx_0(X1,k5_square_1(X1))|~v1_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_26323, negated_conjecture, (v1_xcmplx_0(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v1_membered(X1)|~v3_rcomp_1(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_26421, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_26850, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X1,np__1))=k3_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_18269, negated_conjecture, (m1_subset_1(X1,a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k24_sin_cos))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_21163, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k24_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_5405, plain, (r2_hidden(k5_xcmplx_0(X1),k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_16765, plain, (r2_hidden(X1,k1_numbers)|~m1_subset_1(k7_xcmplx_0(X1,np__1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_15106, plain, (r2_fdiff_1(k3_taylor_1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos)))|~v1_xboole_0(a_0_0_fdiff_8)|~v3_rcomp_1(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_15116, plain, (r2_fdiff_1(k3_taylor_1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos)))|~v1_xboole_0(a_0_0_fdiff_8)|~v3_rcomp_1(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos))))).#trainneg
cnf(i_0_19296, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k19_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_19430, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k19_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_19438, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_26541, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k19_sin_cos,X1,X2,X3,k29_sin_cos,k24_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k19_sin_cos,X1)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_26550, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k24_sin_cos)),k1_zfmisc_1(X1))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_26413, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k29_sin_cos,X1,X2,X3,k19_sin_cos,k24_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k29_sin_cos,X1)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_26549, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k24_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k24_sin_cos))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_26211, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_18410, negated_conjecture, (m1_subset_1(X1,a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_18153, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_22566, plain, (r2_hidden(X1,k1_numbers)|~m1_subset_1(k3_xcmplx_0(np__1,X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6063, plain, (m1_subset_1(k1_seq_1(k3_relat_1(k3_taylor_1,k3_taylor_1),X1),k1_numbers))).#trainneg
cnf(i_0_19446, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k3_taylor_1)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_19454, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k29_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_19462, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k29_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_19694, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k3_taylor_1),k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_19994, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k24_sin_cos),k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_26422, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos)),k1_zfmisc_1(X1))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_26655, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k24_sin_cos)),k1_zfmisc_1(X1))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_26654, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k24_sin_cos))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k24_sin_cos))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_27078, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X1,np__1))=k5_square_1(X1)|~v1_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_27330, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|r2_hidden(k1_relset_1(k1_numbers,k19_sin_cos),k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_18422, negated_conjecture, (m1_subset_1(X1,a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_6108, plain, (m1_subset_1(k1_seq_1(k3_relat_1(k19_sin_cos,k3_taylor_1),X1),k1_numbers))).#trainneg
cnf(i_0_6111, plain, (m1_subset_1(k1_seq_1(k3_relat_1(k29_sin_cos,k3_taylor_1),X1),k1_numbers))).#trainneg
cnf(i_0_6114, plain, (m1_subset_1(k1_seq_1(k3_relat_1(k3_taylor_1,k29_sin_cos),X1),k1_numbers))).#trainneg
cnf(i_0_20002, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_20010, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_20018, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_20026, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k24_sin_cos),k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_20354, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k19_sin_cos),k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_25897, negated_conjecture, (v1_xreal_0(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v3_rcomp_1(X1)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_26901, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(k5_xcmplx_0(X1),X1))=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_28382, plain, (k3_xcmplx_0(k3_square_1(X1),X1)=k3_xcmplx_0(X1,k5_square_1(X1))|~v1_xcmplx_0(k3_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_26646, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k19_sin_cos,X1,X2,X3,k3_taylor_1,k24_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k19_sin_cos,X1)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_26137, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,k3_square_1(X1)))=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_26900, plain, (k7_xcmplx_0(np__1,k5_square_1(k7_xcmplx_0(np__1,X1)))=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_26364, plain, (v1_xboole_0(k1_relset_1(k1_numbers,k29_sin_cos))|r2_hidden(k1_funct_1(X1,X2),k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(k1_seq_1(X1,X2),k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_28424, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(X1,k24_sin_cos)),k9_xtuple_0(X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_26418, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos)),k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_27630, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(np__1,X1))=k3_xcmplx_0(X1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_27834, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|r2_hidden(k1_relset_1(k1_numbers,k29_sin_cos),k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_9746, negated_conjecture, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(esk3_0,k2_zfmisc_1(esk3_0,X2))))).#trainneg
cnf(i_0_1447, plain, (r2_hidden(esk6_2(k9_xtuple_0(k3_relat_1(k24_sin_cos,X1)),X2),k1_relset_1(k1_numbers,k24_sin_cos))|r2_fdiff_1(X2,k9_xtuple_0(k3_relat_1(k24_sin_cos,X1)))|~v3_rcomp_1(k9_xtuple_0(k3_relat_1(k24_sin_cos,X1)))|~r1_tarski(k9_xtuple_0(k3_relat_1(k24_sin_cos,X1)),k1_relset_1(k1_numbers,X2))|~v1_funct_1(X1)|~v1_funct_1(X2)|~m1_subset_1(k9_xtuple_0(k3_relat_1(k24_sin_cos,X1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_6164, plain, (m1_subset_1(k1_seq_1(k3_relat_1(k19_sin_cos,k29_sin_cos),X1),k1_numbers))).#trainneg
cnf(i_0_6167, plain, (m1_subset_1(k1_seq_1(k3_relat_1(k29_sin_cos,k29_sin_cos),X1),k1_numbers))).#trainneg
cnf(i_0_6170, plain, (m1_subset_1(k1_seq_1(k3_relat_1(k3_taylor_1,k19_sin_cos),X1),k1_numbers))).#trainneg
cnf(i_0_20362, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k29_sin_cos),k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_20370, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_20378, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k24_sin_cos),k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_28954, negated_conjecture, (v1_int_1(esk4_0)|~v6_membered(a_0_0_fdiff_8))).#trainneg
cnf(i_0_28955, negated_conjecture, (v1_int_1(esk4_0)|~v6_membered(esk3_0))).#trainneg
cnf(i_0_4555, negated_conjecture, (k7_xcmplx_0(k7_xcmplx_0(np__1,esk4_0),k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))).#trainneg
cnf(i_0_28983, negated_conjecture, (a_0_0_fdiff_8=k6_numbers|r2_hidden(esk4_0,a_0_0_fdiff_8))).#trainneg
cnf(i_0_29000, negated_conjecture, (v1_xcmplx_0(k7_xcmplx_0(np__1,esk4_0))|v1_xboole_0(a_0_0_fdiff_8)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_29001, negated_conjecture, (v7_ordinal1(k7_xcmplx_0(np__1,esk4_0))|v1_xboole_0(a_0_0_fdiff_8)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_28985, negated_conjecture, (r2_hidden(esk4_0,a_0_0_fdiff_8)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_4556, negated_conjecture, (k7_xcmplx_0(k7_xcmplx_0(np__1,k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))),esk4_0)!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))).#trainneg
cnf(i_0_1448, plain, (r1_tarski(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,X5)),k1_relset_1(k1_numbers,k24_sin_cos))|~v1_funct_1(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_29005, negated_conjecture, (r2_hidden(k7_xcmplx_0(np__1,esk4_0),k1_numbers)|v1_xboole_0(a_0_0_fdiff_8))).#trainneg
cnf(i_0_28986, negated_conjecture, (r2_hidden(esk4_0,a_0_0_fdiff_8)|~r2_hidden(X1,k5_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_1474, plain, (v1_xboole_0(k1_relset_1(k1_numbers,k24_sin_cos))|m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k1_relset_1(k1_numbers,k24_sin_cos)))).#trainneg
cnf(i_0_30892, plain, (v1_xreal_0(k1_seq_1(k3_relat_1(k3_taylor_1,k3_taylor_1),X1)))).#trainneg
cnf(i_0_31368, plain, (v1_xreal_0(k1_seq_1(k3_relat_1(k19_sin_cos,k3_taylor_1),X1)))).#trainneg
cnf(i_0_31858, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k24_sin_cos,k24_sin_cos)),k1_relset_1(k1_numbers,k24_sin_cos)))).#trainneg
cnf(i_0_31811, plain, (m1_subset_1(k1_funct_1(X1,X2),k1_numbers)|~r2_hidden(k1_seq_1(X1,X2),k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_29421, negated_conjecture, (r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(X2,k1_zfmisc_1(k5_numbers))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_31373, plain, (v1_xreal_0(k1_seq_1(k3_relat_1(k29_sin_cos,k3_taylor_1),X1)))).#trainneg
cnf(i_0_31378, plain, (v1_xreal_0(k1_seq_1(k3_relat_1(k3_taylor_1,k29_sin_cos),X1)))).#trainneg
cnf(i_0_32088, plain, (v1_xreal_0(k1_seq_1(k3_relat_1(k19_sin_cos,k29_sin_cos),X1)))).#trainneg
cnf(i_0_28975, negated_conjecture, (k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))=k6_numbers|m1_subset_1(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_29159, negated_conjecture, (r1_fdiff_1(X1,esk6_2(X2,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|v1_xboole_0(X2)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X2)|~v3_rcomp_1(X2)|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_32093, plain, (v1_xreal_0(k1_seq_1(k3_relat_1(k29_sin_cos,k29_sin_cos),X1)))).#trainneg
cnf(i_0_32098, plain, (v1_xreal_0(k1_seq_1(k3_relat_1(k3_taylor_1,k19_sin_cos),X1)))).#trainneg
cnf(i_0_28976, negated_conjecture, (m1_subset_1(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,esk3_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_29213, plain, (k7_xcmplx_0(k3_xcmplx_0(X1,X1),np__1)=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_29850, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k29_sin_cos,X1,X2,X3,k3_taylor_1,k24_sin_cos)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k29_sin_cos,X1))))).#trainneg
cnf(i_0_6176, plain, (m1_subset_1(k1_seq_1(k3_relat_1(k19_sin_cos,k19_sin_cos),X1),k1_numbers))).#trainneg
cnf(i_0_29436, plain, (r2_fdiff_1(k3_taylor_1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k24_sin_cos)))|~v1_xboole_0(a_0_0_fdiff_8)|~v3_rcomp_1(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k24_sin_cos))))).#trainneg
cnf(i_0_20386, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k19_sin_cos),k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_21396, negated_conjecture, (k1_zfmisc_1(a_0_0_fdiff_8)=k6_numbers|r2_hidden(k5_numbers,k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_21834, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k29_sin_cos),k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_21842, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k3_taylor_1),k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_27126, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X1,np__1))=k7_xcmplx_0(k3_xcmplx_0(X1,X1),np__1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_27093, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X1,np__1))=k3_xcmplx_0(X1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_27832, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k29_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1484, negated_conjecture, (r1_tarski(esk3_0,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6179, plain, (m1_subset_1(k1_seq_1(k3_relat_1(k29_sin_cos,k19_sin_cos),X1),k1_numbers))).#trainneg
cnf(i_0_14976, negated_conjecture, (m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_21850, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k24_sin_cos),k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_21858, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k19_sin_cos),k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_22029, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k29_sin_cos),k19_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_22050, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k3_taylor_1),k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_22058, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k24_sin_cos),k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_31177, plain, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,X2,X3,X4,k3_taylor_1,k24_sin_cos)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,X2))))).#trainneg
cnf(i_0_27529, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(X2,k7_xcmplx_0(np__1,X2)))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_26899, plain, (k5_square_1(k7_xcmplx_0(np__1,X1))=k7_xcmplx_0(np__1,k5_square_1(X1))|~v1_xcmplx_0(k7_xcmplx_0(np__1,X1))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1488, plain, (k7_xcmplx_0(k3_xcmplx_0(np__1,X1),X2)=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k7_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_14978, negated_conjecture, (m1_subset_1(k1_seq_1(k19_sin_cos,X1),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14979, negated_conjecture, (m1_subset_1(k1_seq_1(k29_sin_cos,X1),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14881, negated_conjecture, (v1_xreal_0(X1)|~r2_hidden(X1,k24_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_22066, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_22430, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_22438, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_22446, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k24_sin_cos),k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_22454, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k19_sin_cos),k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_27098, plain, (k7_xcmplx_0(X1,k3_xcmplx_0(X2,k7_xcmplx_0(X2,np__1)))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_29765, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k24_sin_cos)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_27571, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(k3_square_1(X1),np__1))=k7_xcmplx_0(k5_square_1(X1),k5_xcmplx_0(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_27790, plain, (r1_tarski(X1,X2)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k3_taylor_1)))))).#trainneg
cnf(i_0_29951, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k24_sin_cos)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_10902, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_18245, negated_conjecture, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk3_0))|~m1_subset_1(k1_zfmisc_1(a_0_0_fdiff_8),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_16181, plain, (~v1_xboole_0(a_0_0_fdiff_8)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k3_taylor_1,X2)))|~v1_relat_1(X2))).#trainneg
cnf(i_0_22462, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k29_sin_cos),k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_22784, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_22792, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k24_sin_cos),k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_22800, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k19_sin_cos),k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_22808, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k29_sin_cos),k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_28552, plain, (v1_xreal_0(k7_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_29194, plain, (v1_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(k3_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_29858, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k24_sin_cos)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_31176, plain, (r2_hidden(X1,a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,X2,X3,X4,k3_taylor_1,k24_sin_cos)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,X2))))).#trainneg
cnf(i_0_27588, plain, (k7_xcmplx_0(k7_xcmplx_0(k3_square_1(X1),np__1),X2)=k7_xcmplx_0(X1,k7_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1495, plain, (k7_xcmplx_0(k7_xcmplx_0(X1,X2),np__1)=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4710, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k3_taylor_1,k3_taylor_1)))).#trainneg
cnf(i_0_4711, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k3_relat_1(k3_taylor_1,k3_taylor_1)))).#trainneg
cnf(i_0_4778, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k19_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_22816, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k3_taylor_1),k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_23102, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k24_sin_cos),k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_23110, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k19_sin_cos),k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_23118, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k29_sin_cos),k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_23449, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k3_taylor_1)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_27563, plain, (k7_xcmplx_0(k7_xcmplx_0(k3_square_1(X1),np__1),k5_xcmplx_0(X2))=k3_xcmplx_0(X1,k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_27527, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(k7_xcmplx_0(k5_square_1(X2),np__1),np__1))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_29668, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(k3_xcmplx_0(X2,X2),np__1))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_31494, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(k7_xcmplx_0(np__1,X1),X1))=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_26652, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k24_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_27229, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(np__1,X1))=k3_xcmplx_0(X1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1511, plain, (r2_hidden(esk6_2(a_0_0_fdiff_8,X1),k1_numbers)|r2_fdiff_1(X1,a_0_0_fdiff_8)|~v3_rcomp_1(a_0_0_fdiff_8)|~r1_tarski(a_0_0_fdiff_8,k1_relset_1(k1_numbers,X1))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_4779, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k3_relat_1(k19_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_4870, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k29_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_4871, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k3_relat_1(k29_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_25001, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k3_taylor_1),k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_25010, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k24_sin_cos),k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_25019, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_25028, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_25255, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_7927, plain, (m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers)|~v1_xcmplx_0(k7_xcmplx_0(np__1,X1))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_9029, plain, (k7_xcmplx_0(np__1,k3_square_1(k7_xcmplx_0(np__1,X1)))=k3_square_1(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6441, plain, (v4_relat_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1),k1_numbers)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2683, plain, (v1_xboole_0(k9_xtuple_0(X1))|r2_hidden(k1_funct_1(X2,X3),k9_xtuple_0(X1))|~r2_hidden(k1_seq_1(X2,X3),k9_xtuple_0(X1))|~r2_hidden(X3,k9_xtuple_0(X2))|~v1_funct_1(X1)|~v1_funct_1(X2)|~v3_valued_0(X2)|~v1_relat_1(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_28272, negated_conjecture, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_11139, plain, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k3_taylor_1))))).#trainneg
cnf(i_0_1521, negated_conjecture, (r1_fdiff_1(X1,esk3_0)|k1_zfmisc_1(k1_numbers)=k6_numbers|~v3_rcomp_1(k1_zfmisc_1(k1_numbers))|~r2_fdiff_1(X1,k1_zfmisc_1(k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4904, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k3_taylor_1,k29_sin_cos)))).#trainneg
cnf(i_0_4905, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k3_relat_1(k3_taylor_1,k29_sin_cos)))).#trainneg
cnf(i_0_4972, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k19_sin_cos,k29_sin_cos)))).#trainneg
cnf(i_0_25264, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k24_sin_cos),k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_25273, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k19_sin_cos),k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_25282, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k29_sin_cos),k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_25291, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_25494, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k24_sin_cos),k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_30812, plain, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k24_sin_cos)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_24491, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k19_sin_cos),X1)),k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k19_sin_cos))))|~v1_relat_1(X1))).#trainneg
cnf(i_0_10121, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k3_taylor_1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_28386, plain, (k3_xcmplx_0(X1,k3_square_1(X1))=k3_xcmplx_0(X1,k5_square_1(X1))|~v1_xcmplx_0(k3_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11883, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k3_relat_1(k3_taylor_1,k29_sin_cos),k19_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_28273, negated_conjecture, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_11445, negated_conjecture, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k9_xtuple_0(k2_zfmisc_1(esk3_0,X2)))))).#trainneg
cnf(i_0_1524, plain, (r1_fdiff_1(X1,k5_numbers)|k1_zfmisc_1(k1_numbers)=k6_numbers|~v3_rcomp_1(k1_zfmisc_1(k1_numbers))|~r2_fdiff_1(X1,k1_zfmisc_1(k1_numbers))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4973, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k3_relat_1(k19_sin_cos,k29_sin_cos)))).#trainneg
cnf(i_0_5006, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k29_sin_cos,k29_sin_cos)))).#trainneg
cnf(i_0_5007, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k3_relat_1(k29_sin_cos,k29_sin_cos)))).#trainneg
cnf(i_0_25503, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k19_sin_cos),k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_25512, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k29_sin_cos),k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_25521, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k3_taylor_1),k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_25530, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k24_sin_cos),k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_25838, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k19_sin_cos),k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_29010, negated_conjecture, (k8_real_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k7_xcmplx_0(np__1,esk4_0))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|v1_xboole_0(a_0_0_fdiff_8)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~r1_fdiff_1(k3_taylor_1,esk4_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_12716, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k24_sin_cos,k19_sin_cos)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1530, plain, (m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),a_0_0_fdiff_8)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_32159, negated_conjecture, (v1_int_1(k7_xcmplx_0(np__1,esk4_0))|v1_xboole_0(a_0_0_fdiff_8)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_32261, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))=k6_numbers|k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)!=k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))).#trainneg
cnf(i_0_32344, negated_conjecture, (m1_subset_1(k1_funct_1(k3_taylor_1,X1),k1_numbers)|~r2_hidden(X1,a_0_0_fdiff_8)|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_28274, negated_conjecture, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_1531, plain, (r1_fdiff_1(X1,k1_fdiff_1(k3_taylor_1,X2))|~v3_rcomp_1(a_0_0_fdiff_8)|~r2_fdiff_1(X1,a_0_0_fdiff_8)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X2),k1_numbers))).#trainneg
cnf(i_0_32567, plain, (v1_xreal_0(k1_seq_1(k3_relat_1(k19_sin_cos,k19_sin_cos),X1)))).#trainneg
cnf(i_0_32806, plain, (v1_xreal_0(k1_seq_1(k3_relat_1(k29_sin_cos,k19_sin_cos),X1)))).#trainneg
cnf(i_0_33648, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k24_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_34240, plain, (r1_tarski(k3_relat_1(k3_taylor_1,k3_taylor_1),k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_35148, plain, (r1_tarski(k3_relat_1(k19_sin_cos,k3_taylor_1),k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_35181, plain, (r1_tarski(k3_relat_1(k29_sin_cos,k3_taylor_1),k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_35526, plain, (r1_tarski(k3_relat_1(k3_taylor_1,k29_sin_cos),k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_33573, negated_conjecture, (r1_tarski(esk3_0,X1)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k3_taylor_1),esk3_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k3_taylor_1))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_35462, plain, (r2_hidden(k1_funct_1(X1,X2),k1_relset_1(k1_numbers,k19_sin_cos))|~r2_hidden(k1_seq_1(X1,X2),k1_relset_1(k1_numbers,k19_sin_cos))|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11564, plain, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_1540, plain, (k3_xcmplx_0(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_fdiff_1(X4,k1_fdiff_1(X1,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X2),X4),X3)|~r1_fdiff_1(X4,k1_fdiff_1(X1,X3))|~r1_fdiff_1(k2_fdiff_1(X1,X2),X3)|~v1_funct_1(X4)|~v1_funct_1(k2_fdiff_1(X1,X2))|~v1_xreal_0(k1_fdiff_1(X4,k1_fdiff_1(X1,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k2_fdiff_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k2_fdiff_1(X1,X2),X3),k1_numbers)|~m1_subset_1(X3,k1_numbers)|~r2_hidden(X3,X2)|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_35782, plain, (r1_tarski(k3_relat_1(k19_sin_cos,k29_sin_cos),k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_35815, plain, (r1_tarski(k3_relat_1(k29_sin_cos,k29_sin_cos),k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_5160, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k3_taylor_1,k19_sin_cos)))).#trainneg
cnf(i_0_32331, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k24_sin_cos,k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_relset_1(k1_numbers,k24_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_32294, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_32335, plain, (r2_fdiff_1(k24_sin_cos,k9_xtuple_0(k3_relat_1(k24_sin_cos,k24_sin_cos)))|~v3_rcomp_1(k9_xtuple_0(k3_relat_1(k24_sin_cos,k24_sin_cos)))|~v1_xboole_0(k9_xtuple_0(k3_relat_1(k24_sin_cos,k24_sin_cos))))).#trainneg
cnf(i_0_32151, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_8)|r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,esk4_0))|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers))).#trainneg
cnf(i_0_35122, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X1,np__1))=k3_xcmplx_0(X1,X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11633, plain, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k9_xtuple_0(k2_zfmisc_1(k5_numbers,X2)))))).#trainneg
cnf(i_0_1541, plain, (k3_xcmplx_0(k1_fdiff_1(k19_sin_cos,esk6_2(X1,k29_sin_cos)),k1_fdiff_1(X2,k6_numbers))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,X2),esk6_2(X1,k29_sin_cos))|r2_fdiff_1(k29_sin_cos,X1)|~r1_fdiff_1(X2,k6_numbers)|~r1_fdiff_1(k19_sin_cos,esk6_2(X1,k29_sin_cos))|~v1_funct_1(X2)|~v1_xreal_0(k1_fdiff_1(X2,k6_numbers))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k19_sin_cos,esk6_2(X1,k29_sin_cos)),k1_numbers)|~m1_subset_1(esk6_2(X1,k29_sin_cos),k1_numbers)|~v3_rcomp_1(X1)|~r1_tarski(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5161, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k3_relat_1(k3_taylor_1,k19_sin_cos)))).#trainneg
cnf(i_0_5232, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k19_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_5233, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k3_relat_1(k19_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_25847, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k29_sin_cos),k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_26160, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k29_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_33808, plain, (r1_tarski(X1,X2)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v3_rcomp_1(X1)|~r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k24_sin_cos),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k24_sin_cos))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_33807, plain, (r1_tarski(X1,X2)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k24_sin_cos)))))).#trainneg
cnf(i_0_35916, plain, (r1_tarski(X1,X2)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k24_sin_cos,k19_sin_cos)))))).#trainneg
cnf(i_0_35470, negated_conjecture, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk3_0))|~v3_membered(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))))).#trainneg
cnf(i_0_11754, plain, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos))))).#trainneg
cnf(i_0_1547, negated_conjecture, (v1_xcmplx_0(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v1_membered(k1_numbers)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_5268, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k29_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_5269, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k3_relat_1(k29_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_11138, plain, (r2_hidden(X1,a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k3_taylor_1))))).#trainneg
cnf(i_0_26950, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k19_sin_cos),k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_26959, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k3_taylor_1),k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_26968, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k24_sin_cos),k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_27354, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_27363, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_35113, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,np__1))=k3_xcmplx_0(X1,X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_33626, plain, (r1_tarski(X1,X2)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v3_rcomp_1(X1)|~r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k24_sin_cos),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_33625, plain, (r1_tarski(X1,X2)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k24_sin_cos)))))).#trainneg
cnf(i_0_33489, plain, (r1_tarski(X1,X2)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k24_sin_cos)))))).#trainneg
cnf(i_0_1548, negated_conjecture, (v7_ordinal1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v6_membered(k1_numbers)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_11143, plain, (m1_subset_1(X1,a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k3_taylor_1))))).#trainneg
cnf(i_0_11563, plain, (r2_hidden(X1,a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_11568, plain, (m1_subset_1(X1,a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_27372, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_27381, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k24_sin_cos),k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_27390, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k19_sin_cos),k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_27866, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k29_sin_cos),k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_27875, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_33490, plain, (r1_tarski(X1,X2)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v3_rcomp_1(X1)|~r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k24_sin_cos),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_35917, plain, (r1_tarski(X1,X2)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~v3_rcomp_1(X1)|~r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k24_sin_cos,k19_sin_cos),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k24_sin_cos,k19_sin_cos))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k24_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_35728, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k3_relat_1(k3_taylor_1,k29_sin_cos),k19_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_32436, plain, (v1_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(k3_xcmplx_0(X1,X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1558, negated_conjecture, (k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))=k6_numbers|r2_hidden(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))))).#trainneg
cnf(i_0_11753, plain, (r2_hidden(X1,a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos))))).#trainneg
cnf(i_0_11758, plain, (m1_subset_1(X1,a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos))))).#trainneg
cnf(i_0_27884, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k24_sin_cos),k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_27893, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k29_sin_cos),k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_28315, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k3_taylor_1),k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_28324, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k24_sin_cos),k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_28333, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k19_sin_cos),k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_33225, plain, (k7_xcmplx_0(k3_xcmplx_0(X1,np__1),X2)=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_35241, plain, (v1_xreal_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1559, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos))))|r2_hidden(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos))))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_28342, negated_conjecture, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k29_sin_cos),k3_taylor_1)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_32437, plain, (k3_xcmplx_0(np__1,k5_square_1(X1))=k3_xcmplx_0(X1,X1)|~v1_xcmplx_0(k3_xcmplx_0(X1,X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_35403, plain, (r1_fdiff_1(X1,k1_funct_1(X2,X3))|~v3_rcomp_1(k9_xtuple_0(X4))|~r2_fdiff_1(X1,k9_xtuple_0(X4))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k9_xtuple_0(X4),k1_zfmisc_1(k1_numbers))|~r2_hidden(k1_seq_1(X2,X3),k9_xtuple_0(X4))|~r2_hidden(X3,k9_xtuple_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X2)|~v3_valued_0(X2)|~v1_relat_1(X4)|~v1_relat_1(X2))).#trainneg
cnf(i_0_33735, plain, (m1_subset_1(k7_xcmplx_0(X1,np__1),k1_numbers)|~v1_xcmplx_0(k7_xcmplx_0(X1,np__1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_34474, plain, (k7_xcmplx_0(k7_xcmplx_0(k5_square_1(X1),np__1),k5_xcmplx_0(X2))=k3_xcmplx_0(X1,k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_35713, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,X1))=k3_xcmplx_0(X1,k5_square_1(X1))|~v1_xcmplx_0(k3_xcmplx_0(X1,X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_23602, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r1_tarski(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14980, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,k24_sin_cos),a_0_0_fdiff_8)|v1_xboole_0(k1_zfmisc_1(k1_numbers))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14981, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,k19_sin_cos),a_0_0_fdiff_8)|v1_xboole_0(k1_zfmisc_1(k1_numbers))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14982, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,k29_sin_cos),a_0_0_fdiff_8)|v1_xboole_0(k1_zfmisc_1(k1_numbers))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_15706, plain, (r2_fdiff_1(k3_taylor_1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos)))|~v1_xboole_0(k1_relset_1(k1_numbers,k29_sin_cos))|~v3_rcomp_1(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_15923, plain, (r2_fdiff_1(k3_taylor_1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos)))|~v1_xboole_0(k1_relset_1(k1_numbers,k19_sin_cos))|~v3_rcomp_1(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos))))).#trainneg
cnf(i_0_11734, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),X1)),k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_19029, plain, (k1_relset_1(X1,k3_relat_1(X2,k24_sin_cos))=k9_xtuple_0(k3_relat_1(X2,k24_sin_cos))|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_19223, plain, (k7_xcmplx_0(k3_xcmplx_0(X1,X1),X2)=k7_xcmplx_0(k5_square_1(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_21373, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k24_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_10274, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,X1))=k3_xcmplx_0(k3_square_1(X1),X2)|~v1_xcmplx_0(k3_square_1(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_27337, negated_conjecture, (m1_subset_1(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_13672, plain, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k9_xtuple_0(k2_zfmisc_1(a_0_0_fdiff_8,X2)))))).#trainneg
cnf(i_0_17706, plain, (v4_relat_1(k3_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k19_sin_cos),k19_sin_cos),k1_numbers))).#trainneg
cnf(i_0_13090, plain, (k3_xcmplx_0(X1,k3_square_1(X1))=k3_xcmplx_0(X1,k5_square_1(X1))|~v1_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_20851, plain, (v1_xcmplx_0(k3_square_1(X1))|~v1_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_8273, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,k3_square_1(X1)))=k3_square_1(k3_square_1(X1))|~v1_xcmplx_0(k3_square_1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7832, plain, (v1_relat_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_13788, plain, (m1_subset_1(k1_relset_1(k19_sin_cos,k1_partfun1(k19_sin_cos,X1,X2,X3,X4,X5)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k19_sin_cos,X1)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_5377, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,k3_xcmplx_0(X1,X2)))=k5_square_1(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers))).#trainneg
cnf(i_0_18787, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(a_0_0_fdiff_8,k2_zfmisc_1(a_0_0_fdiff_8,X2))))).#trainneg
cnf(i_0_17762, plain, (v4_relat_1(k3_relat_1(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k29_sin_cos),k19_sin_cos),k1_numbers))).#trainneg
cnf(i_0_17874, plain, (v4_relat_1(k3_relat_1(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k24_sin_cos),k19_sin_cos),k1_numbers))).#trainneg
cnf(i_0_19690, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k3_taylor_1),k19_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_19990, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k24_sin_cos),k19_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_19998, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k19_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_16718, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,X5)),k1_zfmisc_1(a_0_0_fdiff_8))|~v1_funct_1(X5)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_6199, plain, (v1_xreal_0(k7_xcmplx_0(np__1,X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8541, plain, (v1_relat_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k3_taylor_1))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6530, plain, (r2_hidden(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_10854, plain, (k7_xcmplx_0(k5_square_1(X1),k5_xcmplx_0(X1))=k3_xcmplx_0(X1,k3_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_19509, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,X1),X2)),k1_relset_1(k1_numbers,k19_sin_cos))|~v1_relat_1(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_27841, negated_conjecture, (m1_subset_1(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_24563, plain, (r1_tarski(X1,a_0_0_fdiff_8)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k24_sin_cos))))).#trainneg
cnf(i_0_1574, plain, (k1_relset_1(k1_numbers,X1)=X2|X1!=k3_taylor_1|~r2_fdiff_1(X3,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|a_0_0_fdiff_8!=X2|~v1_xboole_0(X2))).#trainneg
cnf(i_0_20006, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k19_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_20014, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k19_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_20022, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k24_sin_cos),k19_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_20350, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k19_sin_cos),k19_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_20358, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k29_sin_cos),k19_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_28952, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_relset_1(k1_numbers,k24_sin_cos)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_8586, plain, (v1_relat_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_20850, plain, (k7_xcmplx_0(k3_square_1(X1),X2)=k7_xcmplx_0(k5_square_1(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_22722, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,X5)),k1_zfmisc_1(k1_numbers))|~v1_funct_1(X5)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_9027, plain, (k3_square_1(k5_xcmplx_0(X1))=k7_xcmplx_0(np__1,k3_square_1(X1))|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_23898, plain, (m1_subset_1(X1,k9_xtuple_0(k3_relat_1(X2,k29_sin_cos)))|~r2_hidden(k1_seq_1(X2,X1),k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X1,k9_xtuple_0(X2))|~v1_funct_1(X2)|~v3_valued_0(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_17231, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k3_taylor_1)))).#trainneg
cnf(i_0_17693, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k19_sin_cos)))).#trainneg
cnf(i_0_17749, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k29_sin_cos)))).#trainneg
cnf(i_0_20366, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k19_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_35927, negated_conjecture, (v7_ordinal1(esk4_0)|~v6_membered(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_35995, negated_conjecture, (r1_tarski(k3_relat_1(k3_taylor_1,k3_taylor_1),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_35958, negated_conjecture, (v1_xreal_0(k1_funct_1(k3_taylor_1,X1))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_35954, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))=k6_numbers|k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)!=k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))).#trainneg
cnf(i_0_37701, negated_conjecture, (v1_xreal_0(k1_relset_1(k1_numbers,k19_sin_cos))|v1_xboole_0(k1_zfmisc_1(k1_numbers))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_35996, plain, (r1_tarski(k3_relat_1(k3_taylor_1,k3_taylor_1),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_35997, plain, (r1_tarski(k3_relat_1(k3_taylor_1,k3_taylor_1),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_35998, negated_conjecture, (r1_tarski(k3_relat_1(k3_taylor_1,k3_taylor_1),a_0_0_fdiff_8)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_36004, negated_conjecture, (r1_tarski(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_36005, plain, (r1_tarski(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_35936, plain, (v1_xreal_0(k1_fdiff_1(k3_taylor_1,X1))|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_35956, negated_conjecture, (v1_xcmplx_0(k1_funct_1(k3_taylor_1,X1))|~v1_membered(k1_numbers)|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_36097, plain, (r1_tarski(X1,a_0_0_fdiff_8)|~v3_rcomp_1(X1)|~r2_fdiff_1(k3_relat_1(k3_taylor_1,k29_sin_cos),X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_38064, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k24_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_36754, negated_conjecture, (r1_tarski(esk3_0,X1)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_36791, negated_conjecture, (r1_tarski(esk3_0,X1)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k24_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k24_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_37706, negated_conjecture, (v1_xreal_0(k1_relset_1(k1_numbers,k29_sin_cos))|v1_xboole_0(k1_zfmisc_1(k1_numbers))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_35955, negated_conjecture, (r2_hidden(k1_funct_1(k3_taylor_1,X1),k1_numbers)|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_36006, plain, (r1_tarski(k3_relat_1(k19_sin_cos,k3_taylor_1),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_36007, negated_conjecture, (r1_tarski(k3_relat_1(k19_sin_cos,k3_taylor_1),a_0_0_fdiff_8)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_36013, negated_conjecture, (r1_tarski(k3_relat_1(k29_sin_cos,k3_taylor_1),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_36014, plain, (r1_tarski(k3_relat_1(k29_sin_cos,k3_taylor_1),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_36015, plain, (r1_tarski(k3_relat_1(k29_sin_cos,k3_taylor_1),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_38877, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,X5)),k1_zfmisc_1(a_0_0_fdiff_8))|~v1_funct_1(X5)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_39092, plain, (k3_xcmplx_0(X1,k3_square_1(X1))=k3_xcmplx_0(X1,k5_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_36028, plain, (r1_tarski(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_25249, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k5_numbers,k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_36382, plain, (r1_tarski(k3_relat_1(k3_taylor_1,k19_sin_cos),k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_36016, negated_conjecture, (r1_tarski(k3_relat_1(k29_sin_cos,k3_taylor_1),a_0_0_fdiff_8)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_36022, negated_conjecture, (r1_tarski(k3_relat_1(k3_taylor_1,k29_sin_cos),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_36023, plain, (r1_tarski(k3_relat_1(k3_taylor_1,k29_sin_cos),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_36024, plain, (r1_tarski(k3_relat_1(k3_taylor_1,k29_sin_cos),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_36025, negated_conjecture, (r1_tarski(k3_relat_1(k3_taylor_1,k29_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_36970, plain, (r2_hidden(X1,a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k19_sin_cos)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_36499, negated_conjecture, (r1_tarski(esk3_0,X1)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k19_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k19_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_36162, plain, (r1_tarski(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos),k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_36171, plain, (r1_tarski(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k29_sin_cos),k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_36001, plain, (r1_tarski(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k3_taylor_1),k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_35929, negated_conjecture, (m1_subset_1(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1595, plain, (r1_fdiff_1(X1,k3_taylor_1)|k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))=k6_numbers|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_36415, plain, (r1_tarski(k3_relat_1(k19_sin_cos,k19_sin_cos),k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_36589, plain, (r1_tarski(k3_relat_1(k29_sin_cos,k19_sin_cos),k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_36156, negated_conjecture, (r1_tarski(k3_relat_1(k19_sin_cos,k29_sin_cos),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_36157, plain, (r1_tarski(k3_relat_1(k19_sin_cos,k29_sin_cos),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_36158, plain, (r1_tarski(k3_relat_1(k19_sin_cos,k29_sin_cos),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_36159, negated_conjecture, (r1_tarski(k3_relat_1(k19_sin_cos,k29_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_36165, negated_conjecture, (r1_tarski(k3_relat_1(k29_sin_cos,k29_sin_cos),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_36010, plain, (r1_tarski(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k3_taylor_1),k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_36830, plain, (r2_hidden(X1,a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_36398, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(k3_relat_1(k19_sin_cos,k19_sin_cos),X1,X2,X3,k19_sin_cos,k19_sin_cos)),k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k3_relat_1(k19_sin_cos,k19_sin_cos),X1))))).#trainneg
cnf(i_0_39109, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,X2),X3)))|~v1_relat_1(X2)|~v1_relat_1(X3))).#trainneg
cnf(i_0_1598, plain, (r1_fdiff_1(X1,k24_sin_cos)|k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))=k6_numbers|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_36166, plain, (r1_tarski(k3_relat_1(k29_sin_cos,k29_sin_cos),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_36167, plain, (r1_tarski(k3_relat_1(k29_sin_cos,k29_sin_cos),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_36168, negated_conjecture, (r1_tarski(k3_relat_1(k29_sin_cos,k29_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_35999, negated_conjecture, (r1_tarski(k3_relat_1(k3_taylor_1,k3_taylor_1),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_36008, negated_conjecture, (r1_tarski(k3_relat_1(k19_sin_cos,k3_taylor_1),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_26547, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k24_sin_cos)),k1_zfmisc_1(k1_relset_1(k1_numbers,k29_sin_cos)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_32382, plain, (k1_fdiff_1(k3_taylor_1,k1_seq_1(k3_relat_1(k19_sin_cos,k29_sin_cos),X1))=k7_xcmplx_0(np__1,k1_seq_1(k3_relat_1(k19_sin_cos,k29_sin_cos),X1))|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k1_seq_1(k3_relat_1(k19_sin_cos,k29_sin_cos),X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_37723, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),X1)),a_0_0_fdiff_8)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3630, plain, (m1_subset_1(k1_funct_1(X1,X2),k9_xtuple_0(X3))|~r2_hidden(X2,k1_relset_1(X4,k1_partfun1(X4,X5,X6,X7,X1,X3)))|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X6,X7)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X4,X5))))).#trainneg
cnf(i_0_28506, plain, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k9_xtuple_0(k3_relat_1(k24_sin_cos,k24_sin_cos)))))).#trainneg
cnf(i_0_1601, plain, (r1_fdiff_1(X1,k19_sin_cos)|k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))=k6_numbers|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_36017, negated_conjecture, (r1_tarski(k3_relat_1(k29_sin_cos,k3_taylor_1),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_36026, negated_conjecture, (r1_tarski(k3_relat_1(k3_taylor_1,k29_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_36160, negated_conjecture, (r1_tarski(k3_relat_1(k19_sin_cos,k29_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_36169, negated_conjecture, (r1_tarski(k3_relat_1(k29_sin_cos,k29_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_38830, plain, (v3_membered(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k3_taylor_1),k19_sin_cos))))).#trainneg
cnf(i_0_15015, plain, (m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos)),k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_19374, plain, (m1_subset_1(k1_funct_1(X1,X2),k1_relset_1(k1_numbers,k19_sin_cos))|~r2_hidden(X2,k9_xtuple_0(k3_relat_1(X1,k19_sin_cos)))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3054, plain, (r1_tarski(k3_rfunct_1(X1,X2,X3,X4),k2_zfmisc_1(X1,k1_numbers))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v3_membered(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_3204, plain, (k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,X1),k5_square_1(k1_seq_1(k19_sin_cos,X1)))=k3_square_1(k1_fdiff_1(k29_sin_cos,X1))|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,X1)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1604, plain, (r1_fdiff_1(X1,k29_sin_cos)|k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))=k6_numbers|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~r2_fdiff_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_38842, plain, (v3_membered(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k24_sin_cos),k19_sin_cos))))).#trainneg
cnf(i_0_38854, plain, (v3_membered(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k19_sin_cos))))).#trainneg
cnf(i_0_39161, plain, (v3_membered(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k19_sin_cos))))).#trainneg
cnf(i_0_39173, plain, (v3_membered(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k19_sin_cos))))).#trainneg
cnf(i_0_39185, plain, (v3_membered(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k24_sin_cos),k19_sin_cos))))).#trainneg
cnf(i_0_6267, plain, (r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)))|v1_xboole_0(a_0_0_fdiff_8)|~m2_subset_1(k1_seq_1(k3_taylor_1,X1),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_14045, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),X5)|m1_subset_1(esk6_2(X5,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos)),k1_numbers)|~v3_rcomp_1(X5)|~r1_tarski(X5,k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6656, plain, (v4_relat_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_6145, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(k5_numbers,X1,X2,X3,X4,X5)),k1_zfmisc_1(k1_numbers))|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))).#trainneg
cnf(i_0_17861, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k24_sin_cos)))).#trainneg
cnf(i_0_39197, plain, (v3_membered(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k19_sin_cos),k19_sin_cos))))).#trainneg
cnf(i_0_39209, plain, (v3_membered(k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k29_sin_cos),k19_sin_cos))))).#trainneg
cnf(i_0_39448, plain, (v3_membered(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k19_sin_cos))))).#trainneg
cnf(i_0_35989, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k24_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_36193, negated_conjecture, (v1_xboole_0(a_0_0_fdiff_8)|r2_hidden(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_19190, plain, (k7_xcmplx_0(k5_square_1(X1),np__1)=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10445, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k3_relat_1(k3_taylor_1,k29_sin_cos),k29_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_11852, plain, (m1_subset_1(k1_relset_1(k29_sin_cos,k2_zfmisc_1(k29_sin_cos,X1)),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_14633, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_28227, plain, (k7_xcmplx_0(k3_square_1(X1),k7_xcmplx_0(np__1,X1))=k3_xcmplx_0(X1,k5_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_28463, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,X5)),k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos)))|~v1_funct_1(X5)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_1621, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k29_sin_cos))))|~r1_tarski(X1,esk3_0)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_2369, plain, (r2_hidden(k3_square_1(X1),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_37696, negated_conjecture, (v1_xreal_0(k1_relset_1(k1_numbers,k24_sin_cos))|v1_xboole_0(k1_zfmisc_1(k1_numbers))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_36183, negated_conjecture, (m1_subset_1(k9_xtuple_0(k3_relat_1(k24_sin_cos,k24_sin_cos)),k1_zfmisc_1(a_0_0_fdiff_8))|~m1_subset_1(k1_relset_1(k1_numbers,k24_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_37694, negated_conjecture, (v1_xcmplx_0(k1_relset_1(k1_numbers,k24_sin_cos))|v1_xboole_0(k1_zfmisc_1(k1_numbers))|~v1_membered(a_0_0_fdiff_8)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14425, plain, (r1_tarski(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos)),k1_numbers)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_11396, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,X1))=k3_xcmplx_0(X2,k3_square_1(X1))|~v1_xcmplx_0(k3_square_1(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_14663, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k19_sin_cos),X1)),k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k19_sin_cos))))|~v1_relat_1(X1))).#trainneg
cnf(i_0_5398, plain, (k7_xcmplx_0(np__1,k3_square_1(k5_xcmplx_0(X1)))=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_31187, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k24_sin_cos),X5)|m1_subset_1(esk6_2(X5,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k24_sin_cos)),k1_numbers)|~v3_rcomp_1(X5)|~r1_tarski(X5,k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k24_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k24_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k24_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_20852, plain, (k3_xcmplx_0(np__1,k3_square_1(X1))=k5_square_1(X1)|~v1_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5933, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_14879, negated_conjecture, (v1_xcmplx_0(X1)|~r2_hidden(X1,k29_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14884, negated_conjecture, (v1_xcmplx_0(X1)|~r2_hidden(X1,k3_taylor_1)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_37695, negated_conjecture, (v7_ordinal1(k1_relset_1(k1_numbers,k24_sin_cos))|v1_xboole_0(k1_zfmisc_1(k1_numbers))|~v6_membered(a_0_0_fdiff_8)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_37699, negated_conjecture, (v1_xcmplx_0(k1_relset_1(k1_numbers,k19_sin_cos))|v1_xboole_0(k1_zfmisc_1(k1_numbers))|~v1_membered(a_0_0_fdiff_8)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_37700, negated_conjecture, (v7_ordinal1(k1_relset_1(k1_numbers,k19_sin_cos))|v1_xboole_0(k1_zfmisc_1(k1_numbers))|~v6_membered(a_0_0_fdiff_8)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_37704, negated_conjecture, (v1_xcmplx_0(k1_relset_1(k1_numbers,k29_sin_cos))|v1_xboole_0(k1_zfmisc_1(k1_numbers))|~v1_membered(a_0_0_fdiff_8)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_37705, negated_conjecture, (v7_ordinal1(k1_relset_1(k1_numbers,k29_sin_cos))|v1_xboole_0(k1_zfmisc_1(k1_numbers))|~v6_membered(a_0_0_fdiff_8)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_3148, plain, (k7_xcmplx_0(np__1,k5_xcmplx_0(X1))=k3_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11801, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k3_taylor_1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_22017, negated_conjecture, (r2_hidden(esk6_2(esk3_0,k2_fdiff_1(X1,X2)),k1_numbers)|r2_fdiff_1(k2_fdiff_1(X1,X2),esk3_0)|~r1_tarski(esk3_0,X2)|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_10820, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k5_square_1(X1),k3_xcmplx_0(np__1,X2))|~v1_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_23765, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k3_taylor_1)),k1_relset_1(k1_numbers,k24_sin_cos))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_11051, plain, (v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(k5_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3710, plain, (k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k7_xcmplx_0(X1,k7_xcmplx_0(X3,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1639, plain, (r1_tarski(esk6_2(k1_zfmisc_1(k9_xtuple_0(k3_relat_1(X1,X2))),k3_taylor_1),k9_xtuple_0(X1))|r2_fdiff_1(k3_taylor_1,k1_zfmisc_1(k9_xtuple_0(k3_relat_1(X1,X2))))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v3_rcomp_1(k1_zfmisc_1(k9_xtuple_0(k3_relat_1(X1,X2))))|~r1_tarski(k1_zfmisc_1(k9_xtuple_0(k3_relat_1(X1,X2))),a_0_0_fdiff_8)|~m1_subset_1(k1_zfmisc_1(k9_xtuple_0(k3_relat_1(X1,X2))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_14885, negated_conjecture, (v1_xcmplx_0(X1)|~r2_hidden(X1,k19_sin_cos)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_20374, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k24_sin_cos),k19_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_20382, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k19_sin_cos),k19_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_21830, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k29_sin_cos),k19_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_32487, plain, (k7_xcmplx_0(k7_xcmplx_0(k3_square_1(X1),np__1),np__1)=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_7685, plain, (r2_hidden(k1_funct_1(X1,X2),k1_numbers)|~r2_hidden(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8795, plain, (r1_tarski(k1_relset_1(X1,k3_rfunct_1(X1,X2,X3,X4)),X1)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v3_membered(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_24945, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k29_sin_cos,k24_sin_cos),X1)),k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k24_sin_cos))))|~v1_relat_1(X1))).#trainneg
cnf(i_0_25572, plain, (k7_xcmplx_0(k3_square_1(X1),np__1)=k7_xcmplx_0(k5_square_1(X1),np__1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_40349, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers)|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,esk3_0))).#trainneg
cnf(i_0_41224, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|v1_xcmplx_0(k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_41172, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_39460, negated_conjecture, (~r1_xxreal_0(k3_xcmplx_0(np__1,esk4_0),k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_39461, negated_conjecture, (~r1_xxreal_0(k3_xcmplx_0(esk4_0,np__1),k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_39462, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(esk4_0,np__1),k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_39481, negated_conjecture, (r2_hidden(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,esk3_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_39480, negated_conjecture, (k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))=k6_numbers|r2_hidden(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_41226, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|v1_xcmplx_0(k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_28999, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers))).#trainneg
cnf(i_0_41191, plain, (v3_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_29005, negated_conjecture, (r2_hidden(k7_xcmplx_0(np__1,esk4_0),k1_numbers))).#trainneg
cnf(i_0_41173, plain, (v4_relat_1(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_40719, negated_conjecture, (v1_xcmplx_0(k7_xcmplx_0(np__1,esk4_0))|~v1_membered(a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_40720, negated_conjecture, (v7_ordinal1(k7_xcmplx_0(np__1,esk4_0))|~v6_membered(a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_40417, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_40718, negated_conjecture, (r2_hidden(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_39479, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))=k6_numbers|k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)!=k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_40318, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|m1_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers))).#trainneg
cnf(i_0_39538, plain, (k1_fdiff_1(k3_taylor_1,k1_fdiff_1(k3_taylor_1,X1))=k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),a_0_0_fdiff_8)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_40149, negated_conjecture, (m1_subset_1(k1_funct_1(k3_taylor_1,X1),k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_40350, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers)|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_41192, plain, (v5_relat_1(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_1640, plain, (v1_relat_1(esk6_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k3_taylor_1))|r2_fdiff_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~r1_tarski(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),a_0_0_fdiff_8)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_41169, plain, (~v1_xboole_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_40502, plain, (v1_relat_1(k1_relset_1(k29_sin_cos,k2_zfmisc_1(k29_sin_cos,X1))))).#trainneg
cnf(i_0_39751, negated_conjecture, (r1_tarski(k3_relat_1(k3_taylor_1,k19_sin_cos),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_39752, plain, (r1_tarski(k3_relat_1(k3_taylor_1,k19_sin_cos),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_39475, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k1_funct_1(k3_taylor_1,X1))=k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,X1))|~m1_subset_1(k1_funct_1(k3_taylor_1,X1),a_0_0_fdiff_8)|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_41599, plain, (m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1641, plain, (v4_relat_1(esk6_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k3_taylor_1),X1)|r2_fdiff_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~r1_tarski(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),a_0_0_fdiff_8)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_41164, plain, (v3_membered(X1)|~v3_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_41165, plain, (v1_valued_0(X1)|~v1_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_39753, plain, (r1_tarski(k3_relat_1(k3_taylor_1,k19_sin_cos),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_39754, negated_conjecture, (r1_tarski(k3_relat_1(k3_taylor_1,k19_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_41740, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(X2,np__1))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_40466, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(X1,np__1))=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_40556, plain, (r1_tarski(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k19_sin_cos),X1)),k1_relset_1(k1_numbers,k19_sin_cos))|~v1_relat_1(X1))).#trainneg
cnf(i_0_1642, plain, (v3_valued_0(esk6_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k3_taylor_1))|r2_fdiff_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_membered(X2)|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~r1_tarski(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),a_0_0_fdiff_8)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_39868, negated_conjecture, (r1_tarski(k3_relat_1(k19_sin_cos,k19_sin_cos),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_39869, plain, (r1_tarski(k3_relat_1(k19_sin_cos,k19_sin_cos),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_39870, plain, (r1_tarski(k3_relat_1(k19_sin_cos,k19_sin_cos),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_39871, negated_conjecture, (r1_tarski(k3_relat_1(k19_sin_cos,k19_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_39483, negated_conjecture, (r2_hidden(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r1_tarski(X1,esk3_0)|~r2_hidden(X2,X1))).#trainneg
cnf(i_0_40465, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,np__1))=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_41600, plain, (m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1643, plain, (v5_relat_1(esk6_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k3_taylor_1),X2)|r2_fdiff_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~r1_tarski(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),a_0_0_fdiff_8)|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_41235, plain, (v1_xreal_0(k3_xcmplx_0(np__1,X1))|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_39877, negated_conjecture, (r1_tarski(k3_relat_1(k29_sin_cos,k19_sin_cos),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_39878, plain, (r1_tarski(k3_relat_1(k29_sin_cos,k19_sin_cos),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_39879, plain, (r1_tarski(k3_relat_1(k29_sin_cos,k19_sin_cos),k1_numbers)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_39880, negated_conjecture, (r1_tarski(k3_relat_1(k29_sin_cos,k19_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_41795, plain, (k7_xcmplx_0(k7_xcmplx_0(k3_square_1(X1),np__1),X2)=k7_xcmplx_0(X1,k7_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_39465, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0)))|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0)),k1_numbers))).#trainneg
cnf(i_0_41197, plain, (m1_subset_1(k3_relat_1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X3,k1_numbers)))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X2,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_39874, plain, (r1_tarski(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos),k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1649, plain, (v1_xboole_0(esk6_2(k1_zfmisc_1(X1),k3_taylor_1))|v1_xboole_0(X1)|m1_subset_1(X2,X1)|r2_fdiff_1(k3_taylor_1,k1_zfmisc_1(X1))|~m1_subset_1(X2,esk6_2(k1_zfmisc_1(X1),k3_taylor_1))|~v3_rcomp_1(k1_zfmisc_1(X1))|~r1_tarski(k1_zfmisc_1(X1),a_0_0_fdiff_8)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_39543, plain, (v1_xreal_0(k7_xcmplx_0(np__1,k6_numbers))|~m2_subset_1(k6_numbers,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_40273, plain, (r1_tarski(k29_sin_cos,k2_zfmisc_1(k1_numbers,k1_numbers))|~v1_funct_1(k16_sin_cos)|~m1_subset_1(k16_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_41299, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,X1))=k3_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_41313, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,k3_taylor_1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_39757, plain, (r1_tarski(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k19_sin_cos),k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1650, plain, (m1_subset_1(X1,X2)|r2_fdiff_1(k3_taylor_1,k1_zfmisc_1(X2))|~r2_hidden(X1,esk6_2(k1_zfmisc_1(X2),k3_taylor_1))|~v3_rcomp_1(k1_zfmisc_1(X2))|~r1_tarski(k1_zfmisc_1(X2),a_0_0_fdiff_8)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_41259, plain, (k3_xcmplx_0(np__1,X1)=k3_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_40099, plain, (k1_fdiff_1(k3_taylor_1,k1_seq_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k29_sin_cos),X1))=k7_xcmplx_0(np__1,k1_seq_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k29_sin_cos),X1))|~m1_subset_1(k1_seq_1(k1_partfun1(X2,X3,X4,X5,k19_sin_cos,k29_sin_cos),X1),a_0_0_fdiff_8)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_41193, plain, (v1_funct_1(k3_relat_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X2,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_39464, negated_conjecture, (r2_hidden(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0)),a_0_0_fdiff_8)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0)),k1_numbers))).#trainneg
cnf(i_0_41181, plain, (m1_subset_1(k3_relat_1(X1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_1651, plain, (r1_tarski(X1,X2)|r2_fdiff_1(k3_taylor_1,k1_zfmisc_1(X2))|~r1_tarski(X1,esk6_2(k1_zfmisc_1(X2),k3_taylor_1))|~v3_rcomp_1(k1_zfmisc_1(X2))|~r1_tarski(k1_zfmisc_1(X2),a_0_0_fdiff_8)|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_6186, plain, (v1_xcmplx_0(k1_seq_1(k3_relat_1(k24_sin_cos,k24_sin_cos),X1)))).#trainneg
cnf(i_0_41161, plain, (v1_xcmplx_0(X1)|~v1_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_39506, negated_conjecture, (r1_tarski(k3_relat_1(k3_taylor_1,k3_taylor_1),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_41901, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k3_relat_1(k19_sin_cos,k19_sin_cos),k19_sin_cos)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_relat_1(k19_sin_cos,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_40577, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~v3_rcomp_1(X1)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k19_sin_cos),X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_39883, plain, (r1_tarski(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k19_sin_cos),k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_41180, plain, (m1_subset_1(k3_relat_1(X1,k19_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_39484, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos)))|r2_hidden(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1652, plain, (r2_fdiff_1(k3_taylor_1,k1_zfmisc_1(X1))|~v1_xboole_0(X1)|~r2_hidden(X2,esk6_2(k1_zfmisc_1(X1),k3_taylor_1))|~v3_rcomp_1(k1_zfmisc_1(X1))|~r1_tarski(k1_zfmisc_1(X1),a_0_0_fdiff_8)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_41162, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_41163, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_5396, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(np__1,X1))=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_32449, plain, (k3_xcmplx_0(X1,k5_square_1(X2))=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,X2))|~v1_xcmplx_0(X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_6809, plain, (r1_tarski(k1_relset_1(X1,k3_relat_1(X5,X6)),X1)|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_41178, plain, (m1_subset_1(k3_relat_1(X1,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_9308, plain, (k7_xcmplx_0(k5_xcmplx_0(X1),np__1)=k7_xcmplx_0(np__1,X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1655, plain, (v7_ordinal1(esk6_2(X1,k3_taylor_1))|r2_fdiff_1(k3_taylor_1,X1)|~v6_membered(X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_20236, plain, (k7_xcmplx_0(np__1,k5_square_1(k5_xcmplx_0(X1)))=k3_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_12759, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k24_sin_cos,k29_sin_cos)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_9343, plain, (k7_xcmplx_0(k5_xcmplx_0(X2),X1)=k7_xcmplx_0(np__1,k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_40567, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k19_sin_cos)),X1,X2,X3,k24_sin_cos,k29_sin_cos)),k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k19_sin_cos)),X1))))).#trainneg
cnf(i_0_36019, plain, (r1_tarski(k1_partfun1(X1,X2,X3,X4,k29_sin_cos,k3_taylor_1),k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_41743, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(X3,k5_xcmplx_0(X2)))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_23136, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(X1,X2)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k9_xtuple_0(X1),k1_zfmisc_1(a_0_0_fdiff_8))|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12621, plain, (k3_xcmplx_0(X1,k3_square_1(X1))=k3_xcmplx_0(k5_square_1(X1),X1)|~v1_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_36551, plain, (r1_tarski(X1,a_0_0_fdiff_8)|~v3_rcomp_1(X1)|~r2_fdiff_1(k3_relat_1(k3_taylor_1,k19_sin_cos),X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_28921, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|~r1_tarski(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_1659, plain, (r1_fdiff_1(X1,X2)|v1_xboole_0(k1_zfmisc_1(X2))|~v3_rcomp_1(k1_zfmisc_1(X2))|~r2_fdiff_1(X1,k1_zfmisc_1(X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_3865, plain, (k7_xcmplx_0(np__1,k5_xcmplx_0(X1))=k3_xcmplx_0(X1,np__1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_39587, negated_conjecture, (r1_tarski(k3_relat_1(k19_sin_cos,k3_taylor_1),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_41985, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r1_tarski(X1,k1_relset_1(esk3_0,k3_rfunct_1(esk3_0,X2,X3,X4)))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v3_membered(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(esk3_0,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(esk3_0,X2))))).#trainneg
cnf(i_0_5811, plain, (k7_xcmplx_0(k3_square_1(k5_xcmplx_0(X1)),X1)=k7_xcmplx_0(k5_xcmplx_0(X1),k5_square_1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_11646, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k29_sin_cos),X1)),k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k29_sin_cos))))|~v1_relat_1(X1))).#trainneg
cnf(i_0_5373, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(X2,k7_xcmplx_0(X1,X2)))=k5_square_1(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~m1_subset_1(k7_xcmplx_0(X1,X2),k1_numbers))).#trainneg
cnf(i_0_13929, plain, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos),X5)|m1_subset_1(esk6_2(X5,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos)),k1_numbers)|~v3_rcomp_1(X5)|~r1_tarski(X5,k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos)))|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_4068, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~m1_subset_1(k3_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_32288, plain, (r1_tarski(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,k24_sin_cos,X5)),k1_relset_1(k1_numbers,k24_sin_cos))|~v1_funct_1(X5)|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_37673, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|~r1_tarski(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_7917, plain, (v3_membered(k9_xtuple_0(k3_relat_1(k19_sin_cos,X1)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_16176, plain, (v3_membered(k9_xtuple_0(k3_relat_1(k3_taylor_1,X1)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_26138, plain, (k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,k5_square_1(X1)))=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_17066, plain, (r1_tarski(X1,k1_relset_1(X2,k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k3_taylor_1)))|~v3_rcomp_1(X1)|~r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k3_taylor_1),X1)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k3_taylor_1))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k29_sin_cos,k3_taylor_1),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_11788, plain, (m1_subset_1(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,k3_relat_1(k3_taylor_1,k29_sin_cos),k3_taylor_1)),k1_zfmisc_1(k1_numbers))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_37710, plain, (m1_subset_1(k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),X1)),k1_zfmisc_1(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos))))|~v1_relat_1(X1))).#trainneg
cnf(i_0_7289, plain, (v1_relat_1(k1_partfun1(X1,X2,X3,X4,k19_sin_cos,k19_sin_cos))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_35493, plain, (r1_tarski(X1,a_0_0_fdiff_8)|~v3_rcomp_1(X1)|~r2_fdiff_1(k3_relat_1(k3_taylor_1,k3_taylor_1),X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1695, negated_conjecture, (m1_subset_1(esk6_2(k5_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k5_numbers)|~v3_rcomp_1(k5_numbers)|~r1_tarski(k5_numbers,esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_21286, plain, (v3_membered(k9_xtuple_0(k3_relat_1(k29_sin_cos,X1)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_12044, negated_conjecture, (r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_8274, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,k3_square_1(X1)))=k5_square_1(k3_square_1(X1))|~v1_xcmplx_0(k3_square_1(X1))|~m1_subset_1(k3_square_1(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_15315, plain, (r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k19_sin_cos)))|~r1_tarski(X1,k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k19_sin_cos),X2)))|~v1_relat_1(X2))).#trainneg
cnf(i_0_19378, plain, (r2_hidden(k1_seq_1(X1,X2),k1_relset_1(k1_numbers,k19_sin_cos))|~r2_hidden(X2,k9_xtuple_0(k3_relat_1(X1,k19_sin_cos)))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v3_valued_0(X1))).#trainneg
cnf(i_0_1696, negated_conjecture, (r1_tarski(esk6_2(k1_zfmisc_1(k9_xtuple_0(k3_relat_1(X1,X2))),k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k9_xtuple_0(X1))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k9_xtuple_0(k3_relat_1(X1,X2))))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v3_rcomp_1(k1_zfmisc_1(k9_xtuple_0(k3_relat_1(X1,X2))))|~r1_tarski(k1_zfmisc_1(k9_xtuple_0(k3_relat_1(X1,X2))),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k9_xtuple_0(k3_relat_1(X1,X2))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_42200, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(np__1,esk4_0),k6_numbers)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_1697, negated_conjecture, (v1_relat_1(esk6_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v3_rcomp_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~r1_tarski(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_42179, negated_conjecture, (v1_xreal_0(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_43462, negated_conjecture, (r1_fdiff_1(k3_taylor_1,esk4_0)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_42213, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,esk4_0))=k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0))|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_1712, negated_conjecture, (v7_ordinal1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v6_membered(X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_42184, negated_conjecture, (v1_xcmplx_0(k7_xcmplx_0(np__1,esk4_0))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_42185, negated_conjecture, (v7_ordinal1(k7_xcmplx_0(np__1,esk4_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_42305, negated_conjecture, (r2_hidden(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0))).#trainneg
cnf(i_0_42348, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,esk4_0))=k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0))|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_42421, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))=k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1))|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_42308, negated_conjecture, (v1_xreal_0(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0))).#trainneg
cnf(i_0_42273, plain, (m1_subset_1(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_1714, negated_conjecture, (r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k29_sin_cos),X1)|m1_subset_1(esk6_2(X1,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k29_sin_cos)),X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~v1_funct_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_42278, negated_conjecture, (v1_xcmplx_0(k7_xcmplx_0(np__1,esk4_0))|~v1_membered(a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_42279, negated_conjecture, (v7_ordinal1(k7_xcmplx_0(np__1,esk4_0))|~v6_membered(a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_42202, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,esk4_0))|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers))).#trainneg
cnf(i_0_42335, plain, (r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1)))|~m2_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_42425, plain, (v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_42299, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))=k6_numbers|k1_seq_1(k2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)!=k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)|~r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_membered(k1_numbers)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_46101, plain, (v3_membered(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_46103, plain, (v3_membered(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_46104, plain, (v3_membered(k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k29_sin_cos))))).#trainneg
cnf(i_0_42188, negated_conjecture, (k8_real_1(k7_xcmplx_0(np__1,esk4_0),k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~r1_fdiff_1(k3_taylor_1,esk4_0)|~v1_funct_1(X1)|~v1_xreal_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers))).#trainneg
cnf(i_0_42195, negated_conjecture, (k8_real_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k7_xcmplx_0(np__1,esk4_0))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~r1_fdiff_1(k3_taylor_1,esk4_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_42190, negated_conjecture, (k3_xcmplx_0(k7_xcmplx_0(np__1,esk4_0),k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~r1_fdiff_1(k3_taylor_1,esk4_0)|~v1_funct_1(X1)|~v1_xreal_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers))).#trainneg
cnf(i_0_42409, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,X1)))|~m2_subset_1(k1_funct_1(k3_taylor_1,X1),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,X1)),k1_numbers)|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_funct_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_42187, negated_conjecture, (k8_real_1(k7_xcmplx_0(np__1,esk4_0),k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~r1_fdiff_1(k3_taylor_1,esk4_0)|~v1_funct_1(X1)|~v1_xreal_0(k7_xcmplx_0(np__1,esk4_0))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_42268, plain, (v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_46106, plain, (v3_membered(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k3_taylor_1))))).#trainneg
cnf(i_0_46253, plain, (v3_membered(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k3_taylor_1))))).#trainneg
cnf(i_0_46255, plain, (v3_membered(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k19_sin_cos))))).#trainneg
cnf(i_0_46256, plain, (v3_membered(k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k29_sin_cos))))).#trainneg
cnf(i_0_42176, negated_conjecture, (r2_hidden(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_42337, plain, (r2_hidden(k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1)),a_0_0_fdiff_8)|~m2_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_42223, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k3_xcmplx_0(np__1,esk4_0))|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k3_xcmplx_0(np__1,esk4_0),k1_numbers))).#trainneg
cnf(i_0_42342, plain, (m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1)),a_0_0_fdiff_8)|~m2_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_42267, plain, (m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_1720, plain, (k7_xcmplx_0(np__1,k3_square_1(k5_square_1(k1_seq_1(k19_sin_cos,X1))))=k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,X1),k5_square_1(k1_seq_1(k19_sin_cos,X1)))|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,X1)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_44956, plain, (k3_square_1(X1)=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_42557, plain, (v1_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1))|~v3_valued_0(k2_zfmisc_1(k1_numbers,k1_numbers)))).#trainneg
cnf(i_0_42191, negated_conjecture, (k8_real_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,esk4_0)),k7_xcmplx_0(np__1,esk4_0))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),esk4_0)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,esk4_0))|~r1_fdiff_1(k3_taylor_1,esk4_0)|~m1_subset_1(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_42360, negated_conjecture, (v1_xcmplx_0(k1_funct_1(k3_taylor_1,X1))|~v1_membered(k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_42419, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,X1)))=k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,X1)))|~m2_subset_1(k1_funct_1(k3_taylor_1,X1),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,X1)),k1_numbers)|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_funct_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_1721, plain, (k7_xcmplx_0(k3_xcmplx_0(X1,np__1),k5_square_1(k1_seq_1(k19_sin_cos,X2)))=k3_xcmplx_0(X1,k1_fdiff_1(k29_sin_cos,X2))|k1_seq_1(k19_sin_cos,X2)=k6_numbers|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,X2)))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_42274, plain, (k9_xtuple_0(X1)=k1_relset_1(k1_numbers,X1)|~v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_39624, negated_conjecture, (r1_tarski(k3_relat_1(k29_sin_cos,k3_taylor_1),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_39755, negated_conjecture, (r1_tarski(k3_relat_1(k3_taylor_1,k19_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_42358, plain, (k1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,X1))=k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,X1))|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_42398, negated_conjecture, (k8_real_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,k1_funct_1(k3_taylor_1,X1))),k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,X1)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),k1_funct_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,k1_funct_1(k3_taylor_1,X1)))|~r1_fdiff_1(k3_taylor_1,k1_funct_1(k3_taylor_1,X1))|~m1_subset_1(k1_seq_1(k3_taylor_1,k1_funct_1(k3_taylor_1,X1)),a_0_0_fdiff_8)|~m1_subset_1(k1_funct_1(k3_taylor_1,X1),k1_numbers)|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_funct_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_46185, plain, (k7_xcmplx_0(np__1,k1_fdiff_1(k29_sin_cos,X1))=k5_square_1(k1_seq_1(k19_sin_cos,X1))|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_43875, plain, (m1_subset_1(k3_xcmplx_0(X1,np__1),k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_42333, plain, (~r1_xxreal_0(k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1)),k6_numbers)|~m2_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_42331, plain, (k8_real_1(k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1)),k1_fdiff_1(k2_fdiff_1(k3_taylor_1,X2),X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k3_taylor_1,X2),k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_fdiff_1(k3_taylor_1,X1))|~r1_fdiff_1(k2_fdiff_1(k3_taylor_1,X2),X1)|~r2_hidden(X1,X2)|~r2_fdiff_1(k3_taylor_1,X2)|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_42319, plain, (k8_real_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,k1_fdiff_1(k3_taylor_1,X1))),k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),k1_fdiff_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,k1_fdiff_1(k3_taylor_1,X1)))|~r1_fdiff_1(k3_taylor_1,k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k1_seq_1(k3_taylor_1,k1_fdiff_1(k3_taylor_1,X1)),a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_39148, plain, (r1_tarski(X1,a_0_0_fdiff_8)|~v3_rcomp_1(X1)|~r2_fdiff_1(k3_relat_1(k3_taylor_1,k24_sin_cos),X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_1722, plain, (k5_xcmplx_0(k1_fdiff_1(k29_sin_cos,X1))=k5_square_1(k1_seq_1(k19_sin_cos,X1))|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,X1)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_43792, plain, (m1_subset_1(k3_xcmplx_0(np__1,X1),k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_42368, negated_conjecture, (v1_xreal_0(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_45091, plain, (k7_xcmplx_0(np__1,k3_xcmplx_0(X2,X1))=k7_xcmplx_0(k7_xcmplx_0(np__1,X2),X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_42222, negated_conjecture, (r2_hidden(k3_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k3_xcmplx_0(np__1,esk4_0),k1_numbers))).#trainneg
cnf(i_0_1723, plain, (v1_xcmplx_0(k1_fdiff_1(k29_sin_cos,X1))|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,X1)))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_45445, plain, (v1_xreal_0(k3_xcmplx_0(X1,np__1))|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_42231, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k3_xcmplx_0(esk4_0,np__1))|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k3_xcmplx_0(esk4_0,np__1),k1_numbers))).#trainneg
cnf(i_0_42323, plain, (k8_real_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,k1_fdiff_1(k3_taylor_1,X2))),k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),k1_fdiff_1(k3_taylor_1,X2))|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,k1_fdiff_1(k3_taylor_1,X2)))|~r1_fdiff_1(k3_taylor_1,k1_fdiff_1(k3_taylor_1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X2),k1_numbers)|~m2_subset_1(X2,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_42402, negated_conjecture, (k8_real_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,k1_funct_1(k3_taylor_1,X2))),k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),k1_funct_1(k3_taylor_1,X2))|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,k1_funct_1(k3_taylor_1,X2)))|~r1_fdiff_1(k3_taylor_1,k1_funct_1(k3_taylor_1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_funct_1(k3_taylor_1,X2),k1_numbers)|~r2_hidden(X2,esk3_0)|~m1_subset_1(k1_funct_1(k3_taylor_1,X2),a_0_0_fdiff_8))).#trainneg
cnf(i_0_42311, plain, (m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~v1_xreal_0(k1_fdiff_1(k3_taylor_1,X1))|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_42411, negated_conjecture, (r2_hidden(k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,X1)),a_0_0_fdiff_8)|~m2_subset_1(k1_funct_1(k3_taylor_1,X1),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,X1)),k1_numbers)|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_funct_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_44079, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0)),a_0_0_fdiff_8)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0)),k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_45090, plain, (k7_xcmplx_0(k5_xcmplx_0(X2),X1)=k7_xcmplx_0(np__1,k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_18607, negated_conjecture, (r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k24_sin_cos))))).#trainneg
cnf(i_0_1724, plain, (k7_xcmplx_0(X1,k5_square_1(k1_seq_1(k19_sin_cos,X2)))=k3_xcmplx_0(X1,k1_fdiff_1(k29_sin_cos,X2))|k1_seq_1(k19_sin_cos,X2)=k6_numbers|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,X2)))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_46422, negated_conjecture, (~v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_48026, plain, (~r1_xxreal_0(np__1,k6_numbers)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(np__1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_46591, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0)),k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_46651, plain, (v3_membered(k1_relset_1(k1_numbers,X1))|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_1747, plain, (r1_tarski(esk6_2(k1_zfmisc_1(k9_xtuple_0(k3_relat_1(X1,X2))),X3),k9_xtuple_0(X1))|r2_fdiff_1(X3,k1_zfmisc_1(k9_xtuple_0(k3_relat_1(X1,X2))))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v3_rcomp_1(k1_zfmisc_1(k9_xtuple_0(k3_relat_1(X1,X2))))|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_zfmisc_1(k9_xtuple_0(k3_relat_1(X1,X2))),k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(k1_zfmisc_1(k9_xtuple_0(k3_relat_1(X1,X2))),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_46567, negated_conjecture, (v1_int_1(k7_xcmplx_0(np__1,esk4_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_48028, plain, (~r1_xxreal_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,k6_numbers)),k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(k6_numbers,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_46526, negated_conjecture, (k1_fdiff_1(k3_taylor_1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))=k7_xcmplx_0(np__1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_47307, plain, (k5_square_1(X1)=k3_xcmplx_0(X1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_47800, plain, (r1_tarski(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,X5,X6)),k1_relset_1(k1_numbers,X5))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_46642, negated_conjecture, (k1_fdiff_1(k3_taylor_1,esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))=k7_xcmplx_0(np__1,esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~m1_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8))).#trainneg
cnf(i_0_47733, plain, (k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,X5,X6))=k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,X6))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_46625, negated_conjecture, (k8_real_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)),k1_fdiff_1(k3_taylor_1,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_47164, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k3_xcmplx_0(np__1,esk4_0))|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k3_xcmplx_0(np__1,esk4_0),k1_numbers))).#trainneg
cnf(i_0_46429, negated_conjecture, (m1_subset_1(k1_funct_1(k3_taylor_1,esk4_0),k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_48684, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))=k6_numbers|k7_xcmplx_0(k3_xcmplx_0(np__1,k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),esk4_0)!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))).#trainneg
cnf(i_0_46654, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_1778, plain, (k1_fdiff_1(k3_taylor_1,esk6_2(X1,k3_taylor_1))=k7_xcmplx_0(np__1,esk6_2(X1,k3_taylor_1))|v1_xboole_0(a_0_0_fdiff_8)|r2_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(esk6_2(X1,k3_taylor_1),a_0_0_fdiff_8)|~v3_rcomp_1(X1)|~r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_47357, negated_conjecture, (v1_xcmplx_0(k1_funct_1(k3_taylor_1,X1))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_47857, plain, (k1_fdiff_1(k3_taylor_1,np__1)=np__1|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(np__1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_46425, negated_conjecture, (r1_fdiff_1(X1,esk4_0)|~v3_rcomp_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_fdiff_1(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_47339, plain, (k7_xcmplx_0(X1,k5_square_1(X2))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k5_xcmplx_0(X2),k1_numbers))).#trainneg
cnf(i_0_48719, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))=k6_numbers|k3_xcmplx_0(k7_xcmplx_0(np__1,esk4_0),k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(k7_xcmplx_0(np__1,esk4_0)))).#trainneg
cnf(i_0_48199, negated_conjecture, (k1_fdiff_1(k3_taylor_1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))=k7_xcmplx_0(np__1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8)|~v3_rcomp_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_47891, plain, (k1_fdiff_1(k3_taylor_1,k1_fdiff_1(k29_sin_cos,X1))=k7_xcmplx_0(np__1,k1_fdiff_1(k29_sin_cos,X1))|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~m2_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,X1)),k1_numbers,a_0_0_fdiff_8)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,X1)))|~m1_subset_1(k1_fdiff_1(k29_sin_cos,X1),k1_numbers)|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,X1)),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_47892, plain, (k1_fdiff_1(k3_taylor_1,X1)=k7_xcmplx_0(np__1,X1)|~m2_subset_1(k5_xcmplx_0(X1),k1_numbers,a_0_0_fdiff_8)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k5_xcmplx_0(X1),a_0_0_fdiff_8)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_46647, plain, (v1_xboole_0(k1_zfmisc_1(k1_numbers))|r2_hidden(k1_relset_1(k1_numbers,X1),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_46655, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_relset_1(k1_numbers,X2))|~m1_subset_1(X2,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_1779, plain, (k1_fdiff_1(k3_taylor_1,esk6_2(X1,X2))=k7_xcmplx_0(np__1,esk6_2(X1,X2))|v1_xboole_0(a_0_0_fdiff_8)|r2_fdiff_1(X2,X1)|~m1_subset_1(esk6_2(X1,X2),a_0_0_fdiff_8)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_42556, negated_conjecture, (v3_membered(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1))|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_46577, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0)))|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_48031, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0))),k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0)),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0)),k1_numbers)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_46626, negated_conjecture, (k3_xcmplx_0(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)),k1_fdiff_1(k3_taylor_1,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~v1_xreal_0(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_48381, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k3_xcmplx_0(esk4_0,np__1))|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k3_xcmplx_0(esk4_0,np__1),k1_numbers))).#trainneg
cnf(i_0_47190, plain, (v1_xreal_0(k1_seq_1(X1,X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_18989, negated_conjecture, (r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k19_sin_cos))))).#trainneg
cnf(i_0_1780, plain, (k1_fdiff_1(k3_taylor_1,esk1_3(X1,X2,X3))=k7_xcmplx_0(np__1,esk1_3(X1,X2,X3))|v1_xboole_0(a_0_0_fdiff_8)|X3=k2_fdiff_1(X1,X2)|~m1_subset_1(esk1_3(X1,X2,X3),a_0_0_fdiff_8)|k1_relset_1(k1_numbers,X3)!=X2|~r2_fdiff_1(X1,X2)|~v1_funct_1(X3)|~v1_funct_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_47341, plain, (v1_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1))|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_44179, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1))|~v3_membered(k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_39791, negated_conjecture, (r1_tarski(k3_relat_1(k3_taylor_1,k29_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_39872, negated_conjecture, (r1_tarski(k3_relat_1(k19_sin_cos,k19_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_46245, plain, (r1_tarski(X1,a_0_0_fdiff_8)|~v3_rcomp_1(X1)|~r2_fdiff_1(k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k3_taylor_1),X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_42353, plain, (k1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)))=k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)))|~m2_subset_1(k1_seq_1(k3_taylor_1,X1),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_47801, plain, (m1_subset_1(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,X5,X6)),k1_zfmisc_1(k1_relset_1(k1_numbers,X5)))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_43456, plain, (k7_xcmplx_0(k3_xcmplx_0(X1,X1),X2)=k7_xcmplx_0(X1,k7_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_47797, plain, (r1_tarski(k9_xtuple_0(k1_partfun1(X1,X2,X3,X4,X5,X6)),k1_relset_1(k1_numbers,X5))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X5,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_44283, negated_conjecture, (k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos))=k6_numbers|r2_hidden(esk4_0,k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_47189, plain, (v7_ordinal1(k1_seq_1(X1,X2))|~v6_membered(k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_19268, negated_conjecture, (r1_tarski(X1,a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,k29_sin_cos))))).#trainneg
cnf(i_0_1783, plain, (k1_fdiff_1(k3_taylor_1,k5_square_1(X1))=k7_xcmplx_0(np__1,k5_square_1(X1))|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k5_square_1(X1),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_39881, negated_conjecture, (r1_tarski(k3_relat_1(k29_sin_cos,k19_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_42407, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,X1)),k6_numbers)|~m2_subset_1(k1_funct_1(k3_taylor_1,X1),k1_numbers,a_0_0_fdiff_8)|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_funct_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_42197, negated_conjecture, (k8_real_1(k1_fdiff_1(X1,k7_xcmplx_0(np__1,esk4_0)),k1_fdiff_1(k2_fdiff_1(k3_taylor_1,X2),esk4_0))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k3_taylor_1,X2),X1),esk4_0)|~r1_fdiff_1(X1,k7_xcmplx_0(np__1,esk4_0))|~r1_fdiff_1(k2_fdiff_1(k3_taylor_1,X2),esk4_0)|~r2_hidden(esk4_0,X2)|~r2_fdiff_1(k3_taylor_1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_42359, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k29_sin_cos))|r2_hidden(k1_funct_1(k3_taylor_1,X1),k1_relset_1(k1_numbers,k29_sin_cos))|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_46624, negated_conjecture, (k3_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)))|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_43877, plain, (v1_xreal_0(k3_xcmplx_0(X1,np__1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_47688, plain, (r1_tarski(k1_relset_1(k1_numbers,k3_relat_1(X1,X2)),k9_xtuple_0(X1))|~v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(k3_relat_1(X1,X2),k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_48956, plain, (k5_square_1(X1)=k3_xcmplx_0(X1,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_48781, negated_conjecture, (~v1_xboole_0(k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_1784, plain, (k1_fdiff_1(k3_taylor_1,k10_real_1(X1,X2))=k7_xcmplx_0(np__1,k10_real_1(X1,X2))|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k10_real_1(X1,X2),a_0_0_fdiff_8)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_49632, negated_conjecture, (v1_xreal_0(k1_funct_1(k3_taylor_1,esk4_0)))).#trainneg
cnf(i_0_48777, negated_conjecture, (~r1_xxreal_0(esk4_0,k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_48774, negated_conjecture, (~r1_xxreal_0(k3_xcmplx_0(np__1,esk4_0),k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_48775, negated_conjecture, (~r1_xxreal_0(k3_xcmplx_0(esk4_0,np__1),k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_51065, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_funct_1(k3_taylor_1,X1))!=k6_numbers|~m1_subset_1(k1_funct_1(k3_taylor_1,X1),k1_numbers)|~r2_hidden(X1,esk3_0))).#trainneg
cnf(i_0_48930, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~m2_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))),k1_numbers)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_48782, negated_conjecture, (m1_subset_1(k1_seq_1(k3_taylor_1,esk4_0),k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_49091, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~m2_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))),k1_numbers)|~m1_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8))).#trainneg
cnf(i_0_50261, plain, (r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,esk6_2(X1,X2)))|r2_fdiff_1(X2,X1)|~m2_subset_1(esk6_2(X1,X2),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk6_2(X1,X2)),k1_numbers)|~v3_rcomp_1(X1)|~v1_funct_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2)))|~m1_subset_1(esk6_2(X1,X2),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_48955, plain, (k5_square_1(X1)=k3_xcmplx_0(X1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_48903, negated_conjecture, (r1_fdiff_1(k3_taylor_1,esk4_0)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_1785, plain, (k1_fdiff_1(k3_taylor_1,k7_xcmplx_0(X1,X2))=k7_xcmplx_0(np__1,k7_xcmplx_0(X1,X2))|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(X1,X2),a_0_0_fdiff_8)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_49735, plain, (k1_fdiff_1(k3_taylor_1,np__1)=np__1|~m1_subset_1(np__1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_48789, negated_conjecture, (m1_subset_1(k1_funct_1(k3_taylor_1,esk4_0),k1_numbers))).#trainneg
cnf(i_0_48776, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(esk4_0,np__1),k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_50106, plain, (k1_fdiff_1(k3_taylor_1,X1)=k7_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k5_xcmplx_0(X1),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_50231, plain, (r1_tarski(k1_relset_1(k1_numbers,X1),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_50232, plain, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k3_taylor_1))|~m1_subset_1(X1,k1_zfmisc_1(k1_relset_1(k1_numbers,X2))))).#trainneg
cnf(i_0_1786, plain, (k1_fdiff_1(k3_taylor_1,k1_fdiff_1(k3_taylor_1,X1))=k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1))|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),a_0_0_fdiff_8)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_49630, negated_conjecture, (v1_xcmplx_0(k1_funct_1(k3_taylor_1,esk4_0))|~v1_membered(k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_48907, plain, (~r1_xxreal_0(k6_numbers,k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(k6_numbers,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_48772, plain, (r2_hidden(np__1,a_0_0_fdiff_8)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(np__1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_48773, plain, (r1_fdiff_1(k3_taylor_1,np__1)|~m2_subset_1(np__1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(np__1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_48904, plain, (~r1_xxreal_0(k3_xcmplx_0(np__1,k6_numbers),k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(k6_numbers,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_50078, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~m2_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))),k1_numbers)|~v3_rcomp_1(X1)|~m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_49628, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(esk4_0,np__1))|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(esk4_0,np__1),k1_numbers))).#trainneg
cnf(i_0_48940, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))=k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~m2_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))),k1_numbers)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_15741, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k19_sin_cos)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1787, plain, (k1_fdiff_1(k3_taylor_1,k1_fdiff_1(X1,X2))=k7_xcmplx_0(np__1,k1_fdiff_1(X1,X2))|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(X1,X2),a_0_0_fdiff_8)|~v1_funct_1(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_49631, negated_conjecture, (v7_ordinal1(k1_funct_1(k3_taylor_1,esk4_0))|~v6_membered(k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_50312, negated_conjecture, (v1_xreal_0(k3_xcmplx_0(np__1,esk4_0))|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_48905, plain, (~r1_xxreal_0(k3_xcmplx_0(k6_numbers,np__1),k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(k6_numbers,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_48906, plain, (~r1_xxreal_0(k7_xcmplx_0(k6_numbers,np__1),k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(k6_numbers,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_50313, negated_conjecture, (v1_xreal_0(k3_xcmplx_0(esk4_0,np__1))|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_50314, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(esk4_0,np__1))|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_50101, plain, (k1_fdiff_1(k3_taylor_1,k1_fdiff_1(k29_sin_cos,X1))=k7_xcmplx_0(np__1,k1_fdiff_1(k29_sin_cos,X1))|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,X1)))|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,X1)),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_50989, plain, (k8_real_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,k5_square_1(X2))),k7_xcmplx_0(np__1,k5_square_1(X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),k5_square_1(X2))|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,k5_square_1(X2)))|~r1_fdiff_1(k3_taylor_1,k5_square_1(X2))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k5_square_1(X2),k1_numbers)|~m1_subset_1(k5_square_1(X2),a_0_0_fdiff_8)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_49089, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~r1_xxreal_0(k7_xcmplx_0(np__1,esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))),k6_numbers)|~m2_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8))).#trainneg
cnf(i_0_50994, plain, (~r1_xxreal_0(k7_xcmplx_0(np__1,k5_square_1(X1)),k6_numbers)|~m2_subset_1(k5_square_1(X1),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k5_square_1(X1),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_51116, negated_conjecture, (k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),k1_seq_1(k3_taylor_1,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~r2_hidden(X1,esk3_0)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)))|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_49629, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k29_sin_cos))|r2_hidden(k1_funct_1(k3_taylor_1,esk4_0),k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_48928, negated_conjecture, (r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~r1_xxreal_0(k7_xcmplx_0(np__1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))),k6_numbers)|~m2_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers,a_0_0_fdiff_8)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_49072, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))),k1_numbers)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8))).#trainneg
cnf(i_0_15742, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k19_sin_cos)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1788, plain, (k1_fdiff_1(k3_taylor_1,k8_real_1(X1,X2))=k7_xcmplx_0(np__1,k8_real_1(X1,X2))|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k8_real_1(X1,X2),a_0_0_fdiff_8)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_50317, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(np__1,k3_xcmplx_0(np__1,esk4_0)),k6_numbers)|~m2_subset_1(k3_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k3_xcmplx_0(np__1,esk4_0),k1_numbers))).#trainneg
cnf(i_0_50318, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(np__1,k3_xcmplx_0(esk4_0,np__1)),k6_numbers)|~m2_subset_1(k3_xcmplx_0(esk4_0,np__1),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k3_xcmplx_0(esk4_0,np__1),k1_numbers))).#trainneg
cnf(i_0_50319, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(np__1,k7_xcmplx_0(esk4_0,np__1)),k6_numbers)|~m2_subset_1(k7_xcmplx_0(esk4_0,np__1),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(esk4_0,np__1),k1_numbers))).#trainneg
cnf(i_0_50253, plain, (k8_real_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk6_2(X2,X3))),k7_xcmplx_0(np__1,esk6_2(X2,X3)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk6_2(X2,X3))|r2_fdiff_1(X3,X2)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk6_2(X2,X3)))|~r1_fdiff_1(k3_taylor_1,esk6_2(X2,X3))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk6_2(X2,X3),k1_numbers)|~v3_rcomp_1(X2)|~v1_funct_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_zfmisc_1(k1_relset_1(k1_numbers,X3)))|~m1_subset_1(esk6_2(X2,X3),a_0_0_fdiff_8)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_51115, negated_conjecture, (k7_xcmplx_0(k3_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),np__1),k1_seq_1(k3_taylor_1,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~r2_hidden(X1,esk3_0)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)))|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_49635, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))=k6_numbers|k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)!=k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))).#trainneg
cnf(i_0_50379, negated_conjecture, (k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),k1_seq_1(k3_taylor_1,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~r2_hidden(X1,esk3_0)|~v1_xreal_0(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_49075, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(np__1,esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8))).#trainneg
cnf(i_0_15743, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k19_sin_cos)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1789, plain, (k1_fdiff_1(k3_taylor_1,k8_real_1(X1,X2))=k7_xcmplx_0(np__1,k8_real_1(X1,X2))|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k8_real_1(X1,X2),a_0_0_fdiff_8)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_48783, negated_conjecture, (m1_subset_1(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k1_relset_1(k1_numbers,k29_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_46532, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0)),k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_48929, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))),k1_numbers)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~m2_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers,a_0_0_fdiff_8)|~v1_xreal_0(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_47730, plain, (r1_tarski(k1_relset_1(k1_numbers,k1_partfun1(X1,X2,X3,X4,X5,X6)),k9_xtuple_0(X5))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_42416, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,X1)),a_0_0_fdiff_8)|~m2_subset_1(k1_funct_1(k3_taylor_1,X1),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,X1)),k1_numbers)|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_funct_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_42390, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,X1)),k1_numbers)|~v1_xreal_0(k1_funct_1(k3_taylor_1,X1))|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_funct_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_42345, plain, (k1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1)))=k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1)))|~m2_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_47187, plain, (r2_hidden(k1_seq_1(X1,X2),k1_numbers)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_19078, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1790, plain, (k1_fdiff_1(k3_taylor_1,k3_xcmplx_0(X1,X2))=k7_xcmplx_0(np__1,k3_xcmplx_0(X1,X2))|v1_xboole_0(a_0_0_fdiff_8)|~m1_subset_1(k3_xcmplx_0(X1,X2),a_0_0_fdiff_8)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_4744, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k24_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_4745, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k3_relat_1(k24_sin_cos,k3_taylor_1)))).#trainneg
cnf(i_0_4938, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k24_sin_cos,k29_sin_cos)))).#trainneg
cnf(i_0_49098, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))),a_0_0_fdiff_8)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~m2_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))),k1_numbers)|~m1_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8))).#trainneg
cnf(i_0_49604, negated_conjecture, (k8_real_1(k1_fdiff_1(k3_taylor_1,X1),k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)))|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_43144, plain, (v1_xreal_0(k3_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_48911, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))),k1_numbers)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v1_xreal_0(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_49626, negated_conjecture, (k8_real_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)),k7_xcmplx_0(np__1,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_42343, plain, (r2_hidden(k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~m2_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_52262, negated_conjecture, (k1_seq_1(k19_sin_cos,k1_funct_1(k3_taylor_1,esk4_0))!=k6_numbers)).#trainneg
cnf(i_0_51984, negated_conjecture, (v1_xcmplx_0(k1_funct_1(k3_taylor_1,esk4_0)))).#trainneg
cnf(i_0_51301, plain, (k3_square_1(X1)=k5_square_1(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_52983, plain, (k1_fdiff_1(k3_taylor_1,X1)=k7_xcmplx_0(np__1,X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_51688, plain, (k1_fdiff_1(k3_taylor_1,X1)=k7_xcmplx_0(np__1,X1)|~m1_subset_1(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_51421, negated_conjecture, (v7_ordinal1(k1_seq_1(k3_taylor_1,esk4_0))|~v6_membered(k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_1810, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k29_sin_cos)))|~r1_tarski(X1,X6)|~r1_tarski(X6,esk3_0)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_51986, plain, (r2_hidden(np__1,a_0_0_fdiff_8)|~m1_subset_1(np__1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_51398, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k1_funct_1(k3_taylor_1,esk4_0))=k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,esk4_0))|~m1_subset_1(k1_funct_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_51400, negated_conjecture, (r2_hidden(k1_funct_1(k3_taylor_1,esk4_0),k1_numbers)|~v1_xcmplx_0(k1_funct_1(k3_taylor_1,esk4_0)))).#trainneg
cnf(i_0_51314, plain, (k7_xcmplx_0(X1,k5_square_1(X2))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1811, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),X5)|m1_subset_1(esk6_2(X5,k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos)),k1_numbers)|~v3_rcomp_1(X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~r1_tarski(X5,esk3_0)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_51987, plain, (r1_fdiff_1(k3_taylor_1,np__1)|~m1_subset_1(np__1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_51279, plain, (k7_xcmplx_0(k5_square_1(X1),X2)=k7_xcmplx_0(X1,k7_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1812, negated_conjecture, (r2_fdiff_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),X5)|~v3_rcomp_1(X5)|~v1_xboole_0(X5)|~v1_funct_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(X1,X2,X3,X4,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X5,k1_zfmisc_1(k1_numbers))|~r1_tarski(X5,esk3_0)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_51770, negated_conjecture, (v7_ordinal1(k1_funct_1(k3_taylor_1,esk4_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_52153, negated_conjecture, (v1_int_1(k1_funct_1(k3_taylor_1,esk4_0))|~v6_membered(k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_51985, plain, (~r1_xxreal_0(k6_numbers,k6_numbers)|~m2_subset_1(k6_numbers,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),a_0_0_fdiff_8))).#trainneg
cnf(i_0_51780, plain, (k1_fdiff_1(k3_taylor_1,X1)=k7_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_51308, plain, (k7_xcmplx_0(k3_square_1(X1),np__1)=k5_square_1(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_52094, plain, (k8_real_1(k7_xcmplx_0(np__1,k1_fdiff_1(X1,X2)),k1_fdiff_1(k2_fdiff_1(X1,X3),X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(X1,X3),k3_taylor_1),X2)|~r1_fdiff_1(k3_taylor_1,k1_fdiff_1(X1,X2))|~r1_fdiff_1(k2_fdiff_1(X1,X3),X2)|~r2_hidden(X2,X3)|~r2_fdiff_1(X1,X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(k1_fdiff_1(X1,X2),a_0_0_fdiff_8))).#trainneg
cnf(i_0_51321, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(k5_square_1(X2),np__1))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_19079, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k29_sin_cos)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_52252, negated_conjecture, (k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),k1_seq_1(k3_taylor_1,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~r2_hidden(X1,esk3_0)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_51588, plain, (k1_fdiff_1(k3_taylor_1,X1)=k7_xcmplx_0(np__1,X1)|~v1_xreal_0(k5_xcmplx_0(X1))|~m1_subset_1(X1,a_0_0_fdiff_8)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_51488, plain, (k3_square_1(X1)=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_51280, plain, (v1_xreal_0(k5_square_1(X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k5_square_1(X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_51501, plain, (k7_xcmplx_0(X1,k5_square_1(X2))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(np__1,X2),k1_numbers))).#trainneg
cnf(i_0_19080, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k29_sin_cos)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_1814, negated_conjecture, (r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos)))|~r1_tarski(X1,esk3_0)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_51283, plain, (k3_xcmplx_0(X1,k3_square_1(X1))=k3_xcmplx_0(X1,k5_square_1(X1))|~v1_xcmplx_0(k3_square_1(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_51927, plain, (k8_real_1(k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1)),k1_fdiff_1(k2_fdiff_1(k3_taylor_1,X2),X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k2_fdiff_1(k3_taylor_1,X2),k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_fdiff_1(k3_taylor_1,X1))|~r1_fdiff_1(k2_fdiff_1(k3_taylor_1,X2),X1)|~r2_hidden(X1,X2)|~r2_fdiff_1(k3_taylor_1,X2)|~m1_subset_1(X1,k1_numbers)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~v1_xreal_0(X1)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_51395, plain, (k1_fdiff_1(k3_taylor_1,k1_fdiff_1(k29_sin_cos,X1))=k7_xcmplx_0(np__1,k1_fdiff_1(k29_sin_cos,X1))|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,X1)))|~m1_subset_1(k1_fdiff_1(k29_sin_cos,X1),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_51272, plain, (k7_xcmplx_0(k5_square_1(X1),k5_xcmplx_0(X2))=k3_xcmplx_0(X1,k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_19719, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k3_taylor_1)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_52266, negated_conjecture, (r2_hidden(esk4_0,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_52692, negated_conjecture, (m1_subset_1(k3_xcmplx_0(esk4_0,np__1),k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_53118, negated_conjecture, (r2_hidden(k7_xcmplx_0(np__1,esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))),a_0_0_fdiff_8)|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~m2_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8))).#trainneg
cnf(i_0_52169, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k3_xcmplx_0(esk4_0,np__1))=k7_xcmplx_0(np__1,k3_xcmplx_0(esk4_0,np__1))|~m1_subset_1(k3_xcmplx_0(esk4_0,np__1),a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_52693, negated_conjecture, (m1_subset_1(k7_xcmplx_0(esk4_0,np__1),k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_52507, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)!=k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))).#trainneg
cnf(i_0_51289, plain, (k3_xcmplx_0(np__1,k5_square_1(X1))=k3_xcmplx_0(X1,X1)|~v1_xcmplx_0(k3_xcmplx_0(X1,X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_19720, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k3_taylor_1)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_52264, negated_conjecture, (m1_subset_1(k1_funct_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_52691, negated_conjecture, (m1_subset_1(k3_xcmplx_0(np__1,esk4_0),k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_52384, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(np__1,esk4_0),np__1),k6_numbers)|~m2_subset_1(k3_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k3_xcmplx_0(np__1,esk4_0),k1_numbers))).#trainneg
cnf(i_0_51273, plain, (k7_xcmplx_0(k5_square_1(X1),np__1)=k3_xcmplx_0(X1,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_51931, plain, (r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1)))|~m2_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1)),k1_numbers)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~v1_xreal_0(X1)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_19721, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k3_taylor_1)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_4939, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k3_relat_1(k24_sin_cos,k29_sin_cos)))).#trainneg
cnf(i_0_5196, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k24_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_5197, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k3_relat_1(k24_sin_cos,k19_sin_cos)))).#trainneg
cnf(i_0_52395, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(np__1,esk4_0),np__1),k6_numbers)|~m2_subset_1(k3_xcmplx_0(esk4_0,np__1),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(k7_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k3_xcmplx_0(esk4_0,np__1),k1_numbers))).#trainneg
cnf(i_0_52259, negated_conjecture, (m1_subset_1(k1_funct_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k1_relset_1(k1_numbers,k29_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_43882, plain, (m1_subset_1(k3_xcmplx_0(X1,np__1),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_49593, negated_conjecture, (k8_real_1(k1_fdiff_1(k3_taylor_1,X1),k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~v1_xreal_0(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)),k1_numbers)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_55819, negated_conjecture, (~v1_xboole_0(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_53607, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,esk4_0)),k6_numbers)|~m2_subset_1(k1_funct_1(k3_taylor_1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_funct_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_54308, negated_conjecture, (v1_int_1(k1_funct_1(k3_taylor_1,esk4_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_55871, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(np__1,k3_xcmplx_0(esk4_0,np__1)),k6_numbers)|~m2_subset_1(k3_xcmplx_0(esk4_0,np__1),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k3_xcmplx_0(esk4_0,np__1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_53638, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,esk4_0))=k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,esk4_0))|~m1_subset_1(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_53609, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,esk4_0)))|~m2_subset_1(k1_funct_1(k3_taylor_1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,esk4_0)),k1_numbers)|~m1_subset_1(k1_funct_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_54310, negated_conjecture, (v1_int_1(k1_seq_1(k3_taylor_1,esk4_0))|~v6_membered(k1_relset_1(k1_numbers,k29_sin_cos)))).#trainneg
cnf(i_0_23096, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k3_taylor_1)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_42195, negated_conjecture, (k8_real_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k7_xcmplx_0(np__1,esk4_0))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_55833, negated_conjecture, (v1_xreal_0(k3_xcmplx_0(esk4_0,np__1))|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_55914, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(esk4_0,np__1))|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_55998, negated_conjecture, (v1_xreal_0(k3_xcmplx_0(np__1,esk4_0))|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_53414, negated_conjecture, (k1_fdiff_1(k3_taylor_1,esk6_2(esk3_0,X1))=k7_xcmplx_0(np__1,esk6_2(esk3_0,X1))|r2_fdiff_1(X1,esk3_0)|~v1_xcmplx_0(esk6_2(esk3_0,X1))|~m1_subset_1(esk6_2(esk3_0,X1),a_0_0_fdiff_8)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_53505, plain, (k8_real_1(k7_xcmplx_0(np__1,k1_seq_1(X1,X2)),k1_fdiff_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(k1_seq_1(X1,X2))|~m1_subset_1(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers))).#trainneg
cnf(i_0_53589, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,esk4_0)),k1_numbers)|~m1_subset_1(k1_funct_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_24481, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k29_sin_cos)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_42191, negated_conjecture, (k8_real_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,esk4_0)),k7_xcmplx_0(np__1,esk4_0))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),esk4_0)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,esk4_0))|~m1_subset_1(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_55991, negated_conjecture, (v7_ordinal1(k1_funct_1(k3_taylor_1,esk4_0))|~v6_membered(a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_53616, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,esk4_0)),a_0_0_fdiff_8)|~m2_subset_1(k1_funct_1(k3_taylor_1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,esk4_0)),k1_numbers)|~m1_subset_1(k1_funct_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_53416, negated_conjecture, (k1_fdiff_1(k3_taylor_1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))=k7_xcmplx_0(np__1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v1_xcmplx_0(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_53459, plain, (k1_fdiff_1(k3_taylor_1,k1_seq_1(k19_sin_cos,X1))=k7_xcmplx_0(np__1,k1_seq_1(k19_sin_cos,X1))|~m1_subset_1(k1_seq_1(k19_sin_cos,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_24532, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k24_sin_cos,k19_sin_cos)))|~m1_subset_1(k24_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_42190, negated_conjecture, (k3_xcmplx_0(k7_xcmplx_0(np__1,esk4_0),k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(X1)|~v1_xreal_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_53426, plain, (k1_fdiff_1(k3_taylor_1,k6_numbers)=k7_xcmplx_0(np__1,k6_numbers)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_55820, negated_conjecture, (m1_subset_1(esk4_0,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_55831, negated_conjecture, (v1_xcmplx_0(k3_xcmplx_0(esk4_0,np__1))|~v1_membered(k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_55832, negated_conjecture, (v7_ordinal1(k3_xcmplx_0(esk4_0,np__1))|~v6_membered(k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_55912, negated_conjecture, (v1_xcmplx_0(k7_xcmplx_0(esk4_0,np__1))|~v1_membered(k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_53611, negated_conjecture, (r2_hidden(k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,esk4_0)),a_0_0_fdiff_8)|~m2_subset_1(k1_funct_1(k3_taylor_1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,esk4_0)),k1_numbers)|~m1_subset_1(k1_funct_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_42187, negated_conjecture, (k8_real_1(k7_xcmplx_0(np__1,esk4_0),k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_53595, negated_conjecture, (r2_hidden(k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,esk4_0)),k1_numbers)|~m1_subset_1(k1_funct_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_53461, plain, (k1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))=k7_xcmplx_0(np__1,k1_seq_1(X1,X2))|~v1_xcmplx_0(k1_seq_1(X1,X2))|~m1_subset_1(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_56075, plain, (k7_xcmplx_0(k5_square_1(X1),np__1)=k3_square_1(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_55558, plain, (k1_fdiff_1(k3_taylor_1,k1_fdiff_1(k29_sin_cos,X1))=k7_xcmplx_0(np__1,k1_fdiff_1(k29_sin_cos,X1))|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~m1_subset_1(k1_fdiff_1(k29_sin_cos,X1),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_53506, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,k1_seq_1(X1,X2)),k1_fdiff_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(k1_fdiff_1(X1,X2))|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(X1,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(k1_seq_1(X1,X2))|~m1_subset_1(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers))).#trainneg
cnf(i_0_53425, plain, (k1_fdiff_1(k3_taylor_1,k1_funct_1(X1,X2))=k7_xcmplx_0(np__1,k1_funct_1(X1,X2))|~v1_xcmplx_0(k1_funct_1(X1,X2))|~m1_subset_1(k1_funct_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_53504, plain, (k3_xcmplx_0(k1_fdiff_1(X1,X2),k7_xcmplx_0(np__1,k1_seq_1(X1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_seq_1(X1,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(k1_seq_1(X1,X2))|~m1_subset_1(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers))).#trainneg
cnf(i_0_55919, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)!=k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_28427, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k24_sin_cos)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_42188, negated_conjecture, (k8_real_1(k7_xcmplx_0(np__1,esk4_0),k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(X1)|~v1_xreal_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_55913, negated_conjecture, (v7_ordinal1(k7_xcmplx_0(esk4_0,np__1))|~v6_membered(k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_55996, negated_conjecture, (v1_xcmplx_0(k3_xcmplx_0(np__1,esk4_0))|~v1_membered(k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_55997, negated_conjecture, (v7_ordinal1(k3_xcmplx_0(np__1,esk4_0))|~v6_membered(k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_56300, negated_conjecture, (v7_ordinal1(k1_funct_1(k3_taylor_1,esk4_0))|~v6_membered(a_0_0_fdiff_8)|~m1_subset_1(k1_relset_1(k1_numbers,k29_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_55989, negated_conjecture, (r2_hidden(k1_funct_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_54930, negated_conjecture, (k7_xcmplx_0(k7_xcmplx_0(np__1,esk4_0),k1_seq_1(k3_taylor_1,esk4_0))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),esk4_0)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_xcmplx_0(k7_xcmplx_0(np__1,esk4_0))|~m1_subset_1(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_53460, plain, (k1_fdiff_1(k3_taylor_1,k1_seq_1(k29_sin_cos,X1))=k7_xcmplx_0(np__1,k1_seq_1(k29_sin_cos,X1))|~m1_subset_1(k1_seq_1(k29_sin_cos,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_55680, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,X2))=k3_xcmplx_0(k5_square_1(X1),X2)|~v1_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_53503, plain, (k8_real_1(k1_fdiff_1(X1,X2),k7_xcmplx_0(np__1,k1_seq_1(X1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_seq_1(X1,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(k1_seq_1(X1,X2))|~m1_subset_1(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers))).#trainneg
cnf(i_0_53532, plain, (r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,X1))|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_28428, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k19_sin_cos,k24_sin_cos)))|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_55830, negated_conjecture, (r2_hidden(k3_xcmplx_0(esk4_0,np__1),k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_55911, negated_conjecture, (r2_hidden(k7_xcmplx_0(esk4_0,np__1),k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_53592, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,esk4_0)))|~m1_subset_1(k1_funct_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_53502, plain, (k8_real_1(k1_fdiff_1(X1,X2),k7_xcmplx_0(np__1,k1_seq_1(X1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(k1_fdiff_1(X1,X2))|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(X1,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(k1_seq_1(X1,X2))|~m1_subset_1(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers))).#trainneg
cnf(i_0_53448, plain, (k1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))=k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1))|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_53412, negated_conjecture, (k1_fdiff_1(k3_taylor_1,esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0)))=k7_xcmplx_0(np__1,esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0)))|r2_fdiff_1(k2_fdiff_1(X1,esk3_0),esk3_0)|~v1_xcmplx_0(esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0)))|~m1_subset_1(esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0)),a_0_0_fdiff_8)|~r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_28429, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k3_relat_1(k29_sin_cos,k24_sin_cos)))|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_55995, negated_conjecture, (r2_hidden(k3_xcmplx_0(np__1,esk4_0),k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_5304, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k3_taylor_1,k24_sin_cos)))).#trainneg
cnf(i_0_56298, negated_conjecture, (r2_hidden(k1_funct_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k1_relset_1(k1_numbers,k29_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_54311, plain, (r2_hidden(k6_numbers,a_0_0_fdiff_8)|~m2_subset_1(k6_numbers,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),a_0_0_fdiff_8)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_54312, plain, (r1_fdiff_1(k3_taylor_1,k6_numbers)|~m2_subset_1(k6_numbers,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers)|~m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),a_0_0_fdiff_8)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_54965, plain, (k1_fdiff_1(k3_taylor_1,X1)=k7_xcmplx_0(np__1,X1)|~v1_xreal_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_56448, negated_conjecture, (v1_int_1(k1_seq_1(k3_taylor_1,esk4_0))|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_47775, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_relset_1(k1_numbers,k2_zfmisc_1(X2,X3)))|~m1_subset_1(k2_zfmisc_1(X2,X3),k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_1827, plain, (k3_xcmplx_0(k7_xcmplx_0(X2,X3),X1)=k7_xcmplx_0(k3_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k7_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k7_xcmplx_0(X2,X3),k1_numbers))).#trainneg
cnf(i_0_56993, plain, (~r1_xxreal_0(k7_xcmplx_0(np__1,k6_numbers),k6_numbers)|~m2_subset_1(k6_numbers,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_57115, plain, (k8_real_1(k7_xcmplx_0(np__1,k1_seq_1(X1,X2)),k1_fdiff_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~r2_hidden(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(X2,k9_xtuple_0(X1))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_56620, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,esk6_2(esk3_0,X1)))|r2_fdiff_1(X1,esk3_0)|~m2_subset_1(esk6_2(esk3_0,X1),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk6_2(esk3_0,X1)),k1_numbers)|~v1_funct_1(X1)|~v1_xcmplx_0(esk6_2(esk3_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_58313, plain, (k8_real_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)),k1_fdiff_1(k3_taylor_1,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_56796, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~m2_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))),k1_numbers)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~v1_xcmplx_0(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_57030, negated_conjecture, (v7_ordinal1(esk4_0)|~v6_membered(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos))))).#trainneg
cnf(i_0_57672, plain, (r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,k1_seq_1(k29_sin_cos,X1)))|~m2_subset_1(k1_seq_1(k29_sin_cos,X1),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(k29_sin_cos,X1)),k1_numbers)|~m1_subset_1(k1_seq_1(k29_sin_cos,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_56960, negated_conjecture, (k7_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~v1_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)))|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(X1)|~v1_xreal_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_56562, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0))=k7_xcmplx_0(np__1,k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0))|~m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_41167, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_56445, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,esk4_0)),k6_numbers)|~m2_subset_1(k1_seq_1(k3_taylor_1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_56449, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(np__1,k3_xcmplx_0(np__1,esk4_0)),k6_numbers)|~m2_subset_1(k3_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k3_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_56858, plain, (r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,k1_seq_1(k19_sin_cos,X1)))|~m2_subset_1(k1_seq_1(k19_sin_cos,X1),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(k19_sin_cos,X1)),k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_57336, plain, (k7_xcmplx_0(k1_fdiff_1(X1,X2),k1_seq_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~v1_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_xcmplx_0(k1_seq_1(X1,X2))|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(k1_fdiff_1(X1,X2))|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(X1,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_57488, plain, (k7_xcmplx_0(k1_fdiff_1(X1,X2),k1_seq_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_xcmplx_0(k1_fdiff_1(X1,X2))|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_seq_1(X1,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_56773, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,esk4_0)),a_0_0_fdiff_8)|~m2_subset_1(k1_seq_1(k3_taylor_1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_41168, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_56970, plain, (v1_xreal_0(k7_xcmplx_0(np__1,k6_numbers))|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_56453, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(np__1,k7_xcmplx_0(esk4_0,np__1)),k6_numbers)|~m2_subset_1(k7_xcmplx_0(esk4_0,np__1),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(esk4_0,np__1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_56456, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(np__1,esk4_0),np__1),k6_numbers)|~m2_subset_1(k3_xcmplx_0(esk4_0,np__1),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k3_xcmplx_0(esk4_0,np__1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_58314, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)),k1_fdiff_1(k3_taylor_1,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~v1_xreal_0(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_58513, plain, (k1_fdiff_1(k3_taylor_1,k5_xcmplx_0(X1))=X1|~v1_xreal_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~m1_subset_1(k5_xcmplx_0(X1),a_0_0_fdiff_8)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_57830, plain, (k3_xcmplx_0(X1,k5_square_1(X2))=k3_xcmplx_0(X2,k3_xcmplx_0(X2,X1))|~v1_xcmplx_0(k5_square_1(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_57668, plain, (k8_real_1(k7_xcmplx_0(np__1,k1_seq_1(k29_sin_cos,X1)),k1_fdiff_1(k29_sin_cos,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k29_sin_cos,k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k29_sin_cos,X1))|~r1_fdiff_1(k29_sin_cos,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_seq_1(k29_sin_cos,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_56720, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_56563, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0))=k7_xcmplx_0(np__1,k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0))|~v1_xreal_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)))|~m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0),a_0_0_fdiff_8)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_57608, negated_conjecture, (r2_hidden(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_56854, plain, (k8_real_1(k7_xcmplx_0(np__1,k1_seq_1(k19_sin_cos,X1)),k1_fdiff_1(k19_sin_cos,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k19_sin_cos,X1))|~r1_fdiff_1(k19_sin_cos,X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_57708, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X1,np__1))=k5_square_1(X1)|~v1_xcmplx_0(k5_square_1(X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_56388, negated_conjecture, (k3_xcmplx_0(k7_xcmplx_0(np__1,esk4_0),k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~v1_xcmplx_0(k7_xcmplx_0(np__1,esk4_0))|~v1_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)))|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_57624, negated_conjecture, (k7_xcmplx_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),esk4_0),np__1)=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),esk4_0)|~v1_xcmplx_0(k7_xcmplx_0(np__1,esk4_0))|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,esk4_0))|~m1_subset_1(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_56968, plain, (v1_xcmplx_0(k7_xcmplx_0(np__1,k6_numbers))|~v1_membered(k1_numbers)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_56969, plain, (v7_ordinal1(k7_xcmplx_0(np__1,k6_numbers))|~v6_membered(k1_numbers)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_58535, plain, (k1_fdiff_1(k3_taylor_1,X1)=k7_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8)|~v1_xcmplx_0(k7_xcmplx_0(np__1,X1))|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers))).#trainneg
cnf(i_0_58312, plain, (k3_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)))|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_57107, negated_conjecture, (k3_xcmplx_0(k7_xcmplx_0(np__1,esk4_0),k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,esk4_0)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),esk4_0)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_xreal_0(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,esk4_0)))|~r2_hidden(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_56962, negated_conjecture, (k3_xcmplx_0(k7_xcmplx_0(np__1,esk4_0),k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,esk4_0)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),esk4_0)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_xreal_0(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,esk4_0)))|~m1_subset_1(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_38833, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k3_taylor_1),k19_sin_cos))))).#trainneg
cnf(i_0_56967, plain, (m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_56973, plain, (r2_hidden(k7_xcmplx_0(np__1,k6_numbers),k1_numbers)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_56766, negated_conjecture, (v7_ordinal1(k1_seq_1(k3_taylor_1,esk4_0))|~v6_membered(a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_58222, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,esk4_0)))|~m1_subset_1(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_57105, negated_conjecture, (k3_xcmplx_0(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,esk4_0)),k7_xcmplx_0(np__1,esk4_0))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),esk4_0)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,esk4_0))|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers)|~r2_hidden(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_57669, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,k1_seq_1(k29_sin_cos,X1)),k1_fdiff_1(k29_sin_cos,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k29_sin_cos,k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k29_sin_cos,X1))|~r1_fdiff_1(k29_sin_cos,X1)|~v1_xreal_0(k1_fdiff_1(k29_sin_cos,X1))|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(k29_sin_cos,X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_seq_1(k29_sin_cos,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_56855, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,k1_seq_1(k19_sin_cos,X1)),k1_fdiff_1(k19_sin_cos,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k19_sin_cos,k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k19_sin_cos,X1))|~r1_fdiff_1(k19_sin_cos,X1)|~v1_xreal_0(k1_fdiff_1(k19_sin_cos,X1))|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(k19_sin_cos,X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_seq_1(k19_sin_cos,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_57116, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,k1_seq_1(X1,X2)),k1_fdiff_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(k1_fdiff_1(X1,X2))|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(X1,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~r2_hidden(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(X2,k9_xtuple_0(X1))|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_58150, plain, (r1_fdiff_1(k3_taylor_1,X1)|~m2_subset_1(k5_xcmplx_0(X1),k1_numbers,a_0_0_fdiff_8)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k5_xcmplx_0(X1),a_0_0_fdiff_8)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_56709, plain, (k8_real_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)),k7_xcmplx_0(np__1,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_38845, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k24_sin_cos),k19_sin_cos))))).#trainneg
cnf(i_0_1857, plain, (r1_tarski(X1,k2_zfmisc_1(X2,X3))|~r1_tarski(X1,k3_relat_1(X6,X7))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X5,X3)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,X4))))).#trainneg
cnf(i_0_56440, negated_conjecture, (k3_xcmplx_0(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,esk4_0)),k7_xcmplx_0(np__1,esk4_0))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),esk4_0)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,esk4_0))|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_57610, negated_conjecture, (m1_subset_1(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),esk4_0),k1_numbers)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_xcmplx_0(k7_xcmplx_0(np__1,esk4_0))|~m1_subset_1(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_54379, plain, (m1_subset_1(k5_square_1(X1),k1_numbers)|~m1_subset_1(k3_square_1(X1),k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_38857, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k19_sin_cos),k19_sin_cos))))).#trainneg
cnf(i_0_1864, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(np__1,X2))=k7_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))|~v1_xcmplx_0(k7_xcmplx_0(np__1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_57180, plain, (k7_xcmplx_0(k3_square_1(X1),X2)=k7_xcmplx_0(k5_square_1(X1),X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k5_square_1(X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_53520, plain, (k8_real_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2)),k7_xcmplx_0(np__1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),X2)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))|~r1_fdiff_1(k3_taylor_1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(X2)|~m1_subset_1(X2,a_0_0_fdiff_8))).#trainneg
cnf(i_0_58512, plain, (k1_fdiff_1(k3_taylor_1,k5_square_1(k1_seq_1(k19_sin_cos,X1)))=k1_fdiff_1(k29_sin_cos,X1)|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~v1_xreal_0(k1_fdiff_1(k29_sin_cos,X1))|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,X1)))|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,X1)),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_39164, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k3_relat_1(k3_taylor_1,k29_sin_cos),k19_sin_cos))))).#trainneg
cnf(i_0_1865, plain, (k3_xcmplx_0(X1,k7_xcmplx_0(np__1,X2))=k7_xcmplx_0(X1,k3_xcmplx_0(np__1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k7_xcmplx_0(np__1,X2))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_61180, plain, (r1_fdiff_1(k3_taylor_1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k5_xcmplx_0(X1),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_59077, negated_conjecture, (k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk4_0)|~r1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))|~v1_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))).#trainneg
cnf(i_0_59390, plain, (k7_xcmplx_0(k1_fdiff_1(X1,X2),k1_seq_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_xcmplx_0(k1_fdiff_1(X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(k1_seq_1(X1,X2)))).#trainneg
cnf(i_0_1878, plain, (k3_xcmplx_0(X1,k5_xcmplx_0(X2))=k7_xcmplx_0(X1,k3_xcmplx_0(np__1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_59392, plain, (k7_xcmplx_0(k1_fdiff_1(X1,X2),k1_seq_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_xcmplx_0(k1_fdiff_1(X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers))).#trainneg
cnf(i_0_59955, plain, (k5_xcmplx_0(X1)=k1_fdiff_1(k3_taylor_1,X1)|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(k5_xcmplx_0(X1))|~m1_subset_1(X1,a_0_0_fdiff_8)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_51782, negated_conjecture, (r1_tarski(k1_relset_1(k1_numbers,X1),k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_39176, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k3_taylor_1),k19_sin_cos))))).#trainneg
cnf(i_0_1881, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(k7_xcmplx_0(np__1,X2),X3))=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_59289, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(k7_xcmplx_0(np__1,esk4_0),np__1),k6_numbers)|~m2_subset_1(k3_xcmplx_0(np__1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k3_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_62106, plain, (k7_xcmplx_0(X1,k3_xcmplx_0(np__1,X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_59078, negated_conjecture, (k7_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),esk4_0)=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(X1)|~v1_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_59384, plain, (k7_xcmplx_0(k1_fdiff_1(X1,X2),k1_seq_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_xcmplx_0(k1_fdiff_1(X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(k7_xcmplx_0(np__1,k1_seq_1(X1,X2)))|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(X1,X2)),k1_numbers))).#trainneg
cnf(i_0_59957, plain, (k1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,X1))=X1|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_39188, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k24_sin_cos),k19_sin_cos))))).#trainneg
cnf(i_0_1882, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(k7_xcmplx_0(np__1,X3),X2))=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_57601, negated_conjecture, (v7_ordinal1(k1_seq_1(k3_taylor_1,esk4_0))|~v6_membered(a_0_0_fdiff_8)|~m1_subset_1(k1_relset_1(k1_numbers,k29_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_38323, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_hidden(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,k1_relset_1(k1_numbers,k19_sin_cos))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_39200, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k19_sin_cos),k19_sin_cos))))).#trainneg
cnf(i_0_1883, plain, (k3_xcmplx_0(X1,k5_square_1(k1_seq_1(k19_sin_cos,X2)))=k7_xcmplx_0(X1,k1_fdiff_1(k29_sin_cos,X2))|k1_seq_1(k19_sin_cos,X2)=k6_numbers|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,X2)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_59175, negated_conjecture, (m1_subset_1(esk4_0,k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(esk3_0,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_59176, negated_conjecture, (m1_subset_1(esk4_0,k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(a_0_0_fdiff_8,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_59180, negated_conjecture, (m1_subset_1(esk4_0,k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_59193, plain, (m1_subset_1(np__1,k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(k5_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_59194, plain, (m1_subset_1(np__1,k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_62017, negated_conjecture, (k7_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k3_xcmplx_0(np__1,esk4_0))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(X1)|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~v1_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))))).#trainneg
cnf(i_0_60539, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(k7_xcmplx_0(np__1,esk4_0))|~r1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))|~v1_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))|~m1_subset_1(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers))).#trainneg
cnf(i_0_61125, plain, (k7_xcmplx_0(k1_fdiff_1(X1,X2),k1_seq_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~v1_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_xcmplx_0(k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~r2_hidden(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_xreal_0(k1_fdiff_1(X1,X2))|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(X1,X2)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_58912, plain, (k3_xcmplx_0(k1_fdiff_1(X1,X2),k7_xcmplx_0(np__1,k1_seq_1(X1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_seq_1(X1,X2)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~r1_fdiff_1(X1,X2)|~r2_hidden(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_60716, plain, (k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),k1_seq_1(k3_taylor_1,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)))|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_61809, plain, (k3_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2)),k7_xcmplx_0(np__1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),X2)|~v1_xreal_0(k7_xcmplx_0(np__1,X2))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2)),k1_numbers)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))|~r1_fdiff_1(k3_taylor_1,X2)|~v1_funct_1(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X2,a_0_0_fdiff_8))).#trainneg
cnf(i_0_61812, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,X2),k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),X2)|~v1_xreal_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2)))|~m1_subset_1(k7_xcmplx_0(np__1,X2),k1_numbers)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))|~r1_fdiff_1(k3_taylor_1,X2)|~v1_funct_1(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X2,a_0_0_fdiff_8))).#trainneg
cnf(i_0_39118, negated_conjecture, (v1_xboole_0(k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_hidden(X1,k1_relset_1(k1_numbers,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,k1_relset_1(k1_numbers,k29_sin_cos))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_39212, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k3_relat_1(k24_sin_cos,k29_sin_cos),k19_sin_cos))))).#trainneg
cnf(i_0_1886, plain, (k7_xcmplx_0(X1,k5_square_1(k5_xcmplx_0(X2)))=k3_xcmplx_0(k7_xcmplx_0(X1,k5_xcmplx_0(X2)),X2)|~v1_xcmplx_0(k7_xcmplx_0(X1,k5_xcmplx_0(X2)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(k5_xcmplx_0(X2),k1_numbers))).#trainneg
cnf(i_0_59407, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_59404, negated_conjecture, (r1_tarski(esk3_0,k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(a_0_0_fdiff_8,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_59214, plain, (m1_subset_1(k6_numbers,k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(k5_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_59215, plain, (m1_subset_1(k6_numbers,k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_59570, plain, (r1_tarski(a_0_0_fdiff_8,k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_59572, plain, (r1_tarski(k5_numbers,k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_60629, plain, (k1_fdiff_1(k3_taylor_1,k5_xcmplx_0(X1))=X1|~v1_xcmplx_0(k5_xcmplx_0(X1))|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k5_xcmplx_0(X1),a_0_0_fdiff_8)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_59843, plain, (k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),k1_seq_1(k3_taylor_1,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~v1_xreal_0(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_59924, plain, (m1_subset_1(X1,a_0_0_fdiff_8)|~m2_subset_1(k5_xcmplx_0(X1),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k5_xcmplx_0(X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_61451, negated_conjecture, (v1_xreal_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),esk4_0))|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_xcmplx_0(k7_xcmplx_0(np__1,esk4_0))|~m1_subset_1(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_39451, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k3_taylor_1),k19_sin_cos))))).#trainneg
cnf(i_0_1919, plain, (k7_xcmplx_0(np__1,k5_square_1(k5_square_1(k1_seq_1(k19_sin_cos,X1))))=k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,X1),k5_square_1(k1_seq_1(k19_sin_cos,X1)))|k1_seq_1(k19_sin_cos,X1)=k6_numbers|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,X1)))|~m1_subset_1(k5_square_1(k1_seq_1(k19_sin_cos,X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_59768, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_58471, negated_conjecture, (r2_hidden(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k1_relset_1(k1_numbers,k29_sin_cos),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_60243, plain, (k3_xcmplx_0(X1,k3_square_1(X1))=k3_xcmplx_0(X1,k5_square_1(X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(k5_square_1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_61011, plain, (k7_xcmplx_0(k1_fdiff_1(k29_sin_cos,X1),k1_seq_1(k29_sin_cos,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k29_sin_cos,k3_taylor_1),X1)|~v1_xcmplx_0(k1_fdiff_1(k29_sin_cos,X1))|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k29_sin_cos,X1))|~r1_fdiff_1(k29_sin_cos,X1)|~v1_xreal_0(k1_fdiff_1(k29_sin_cos,X1))|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(k29_sin_cos,X1)),k1_numbers)|~m1_subset_1(k1_seq_1(k29_sin_cos,X1),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_56961, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(k7_xcmplx_0(np__1,esk4_0))|~r1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))|~v1_xreal_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))).#trainneg
cnf(i_0_58311, plain, (k8_real_1(k1_fdiff_1(k3_taylor_1,X1),k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)))|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_41885, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k24_sin_cos),k19_sin_cos))))).#trainneg
cnf(i_0_1920, plain, (k7_xcmplx_0(np__1,k5_square_1(k5_xcmplx_0(X1)))=k7_xcmplx_0(X1,k5_xcmplx_0(X1))|~v1_xcmplx_0(k5_xcmplx_0(X1))|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5305, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k3_relat_1(k3_taylor_1,k24_sin_cos)))).#trainneg
cnf(i_0_5505, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k19_sin_cos,k24_sin_cos)))).#trainneg
cnf(i_0_5506, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k3_relat_1(k19_sin_cos,k24_sin_cos)))).#trainneg
cnf(i_0_60887, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_52156, negated_conjecture, (m1_subset_1(k3_xcmplx_0(np__1,esk4_0),k1_numbers)|~v1_xcmplx_0(k3_xcmplx_0(np__1,esk4_0))|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_52157, negated_conjecture, (r2_hidden(k3_xcmplx_0(np__1,esk4_0),k1_numbers)|~v1_xcmplx_0(k3_xcmplx_0(np__1,esk4_0))|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_57201, plain, (k7_xcmplx_0(X1,k3_square_1(X2))=k7_xcmplx_0(X1,k5_square_1(X2))|~v1_xcmplx_0(k5_square_1(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_60805, negated_conjecture, (k7_xcmplx_0(k7_xcmplx_0(np__1,esk4_0),k1_seq_1(k3_taylor_1,esk4_0))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),esk4_0)|~v1_xcmplx_0(k7_xcmplx_0(np__1,esk4_0))|~r2_hidden(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,esk4_0))))).#trainneg
cnf(i_0_57152, plain, (k1_fdiff_1(k3_taylor_1,k1_fdiff_1(X1,X3))=k7_xcmplx_0(np__1,k1_fdiff_1(X1,X3))|~r2_hidden(k1_fdiff_1(X1,X3),a_0_0_fdiff_8)|~r2_hidden(X3,k9_xtuple_0(k2_fdiff_1(X1,X2)))|~v1_funct_1(k2_fdiff_1(X1,X2))|~v3_valued_0(k2_fdiff_1(X1,X2))|~v1_relat_1(k2_fdiff_1(X1,X2))|~r2_hidden(X3,X2)|~r2_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_61449, negated_conjecture, (v1_xcmplx_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),esk4_0))|~v1_membered(k1_numbers)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_xcmplx_0(k7_xcmplx_0(np__1,esk4_0))|~m1_subset_1(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_58385, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0))))|r2_fdiff_1(k2_fdiff_1(X1,esk3_0),esk3_0)|~m2_subset_1(esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0)),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0))),k1_numbers)|~r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~v1_xcmplx_0(esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_58310, plain, (k8_real_1(k1_fdiff_1(k3_taylor_1,X1),k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~v1_xreal_0(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_41897, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k19_sin_cos),k19_sin_cos))))).#trainneg
cnf(i_0_1926, plain, (k7_xcmplx_0(X1,k5_square_1(k5_xcmplx_0(X2)))=k7_xcmplx_0(k3_xcmplx_0(X1,X2),k5_xcmplx_0(X2))|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(k5_xcmplx_0(X2),k1_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_5541, plain, (m1_subset_1(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r2_hidden(X1,k3_relat_1(k29_sin_cos,k24_sin_cos)))).#trainneg
cnf(i_0_63775, negated_conjecture, (v7_ordinal1(esk4_0)|~v6_membered(k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(esk3_0,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_62502, plain, (k5_xcmplx_0(X1)=k1_fdiff_1(k3_taylor_1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_62501, plain, (k5_xcmplx_0(X1)=k1_fdiff_1(k3_taylor_1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_63038, plain, (k7_xcmplx_0(X1,k3_xcmplx_0(X2,np__1))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_62174, plain, (r1_fdiff_1(k3_taylor_1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_64838, plain, (m1_subset_1(X1,a_0_0_fdiff_8)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k5_xcmplx_0(X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_65478, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(np__1,k6_numbers))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_65476, negated_conjecture, (v1_xcmplx_0(k7_xcmplx_0(np__1,k6_numbers))|~v1_membered(a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_64041, plain, (k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),k1_seq_1(k3_taylor_1,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_65409, plain, (k7_xcmplx_0(X1,k5_xcmplx_0(X1))=k5_square_1(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_63981, plain, (k7_xcmplx_0(k1_fdiff_1(X1,X2),k1_seq_1(X1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2)|~v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_xcmplx_0(k1_fdiff_1(X1,X2))|~r1_fdiff_1(X1,X2)|~r2_hidden(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_xreal_0(k7_xcmplx_0(np__1,k1_seq_1(X1,X2)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1936, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,X2)|~r1_tarski(X2,k3_taylor_1))).#trainneg
cnf(i_0_65477, negated_conjecture, (v7_ordinal1(k7_xcmplx_0(np__1,k6_numbers))|~v6_membered(a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_65475, negated_conjecture, (r2_hidden(k7_xcmplx_0(np__1,k6_numbers),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_65479, negated_conjecture, (m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_65486, negated_conjecture, (m1_subset_1(k3_xcmplx_0(np__1,esk4_0),k1_numbers)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_62495, plain, (k5_xcmplx_0(X1)=k1_fdiff_1(k3_taylor_1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~m1_subset_1(k5_xcmplx_0(X1),k1_numbers))).#trainneg
cnf(i_0_62435, plain, (k7_xcmplx_0(k5_square_1(k1_seq_1(X1,X2)),k1_fdiff_1(X1,X2))=k7_xcmplx_0(k1_seq_1(X1,X2),k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2))|~v1_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_xcmplx_0(k1_seq_1(X1,X2))|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_65145, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk4_0)|~r1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(k7_xcmplx_0(np__1,esk4_0)))).#trainneg
cnf(i_0_41909, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(k3_relat_1(k19_sin_cos,k29_sin_cos),k19_sin_cos))))).#trainneg
cnf(i_0_1938, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,X2)|~r1_tarski(X2,k24_sin_cos))).#trainneg
cnf(i_0_65480, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k6_numbers)=k7_xcmplx_0(np__1,k6_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_65491, negated_conjecture, (r2_hidden(k3_xcmplx_0(np__1,esk4_0),k1_numbers)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_65481, negated_conjecture, (r2_hidden(k6_numbers,a_0_0_fdiff_8)|~m2_subset_1(k6_numbers,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_65482, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k6_numbers)|~m2_subset_1(k6_numbers,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers)|~m1_subset_1(k6_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_59177, negated_conjecture, (m1_subset_1(esk4_0,k2_zfmisc_1(k1_numbers,k1_numbers))|~m1_subset_1(k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k29_sin_cos)),k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_58201, negated_conjecture, (m1_subset_1(k3_xcmplx_0(esk4_0,np__1),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_58212, negated_conjecture, (m1_subset_1(k7_xcmplx_0(esk4_0,np__1),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_62287, plain, (k7_xcmplx_0(k7_xcmplx_0(np__1,X1),k1_seq_1(k3_taylor_1,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~v1_xcmplx_0(k7_xcmplx_0(np__1,X1))|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_62236, plain, (k7_xcmplx_0(k5_square_1(k1_seq_1(X1,X2)),k1_fdiff_1(X1,X2))=k7_xcmplx_0(k1_seq_1(X1,X2),k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2))|~v1_xreal_0(k1_seq_1(X1,X2))|~v1_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_xcmplx_0(k1_seq_1(X1,X2))|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_62434, plain, (k7_xcmplx_0(k3_square_1(k1_seq_1(X1,X2)),k1_fdiff_1(X1,X2))=k7_xcmplx_0(k1_seq_1(X1,X2),k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2))|~v1_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_xcmplx_0(k1_seq_1(X1,X2))|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_63303, plain, (k1_fdiff_1(k3_taylor_1,k7_xcmplx_0(X2,X1))=k7_xcmplx_0(X1,X2)|~v1_xreal_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~m1_subset_1(k7_xcmplx_0(X2,X1),a_0_0_fdiff_8)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_52072, negated_conjecture, (r1_tarski(X1,X2)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r1_tarski(X1,esk3_0)|~r2_fdiff_1(k3_relat_1(k3_taylor_1,k19_sin_cos),esk3_0))).#trainneg
cnf(i_0_1940, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,X2)|~r1_tarski(X2,k19_sin_cos))).#trainneg
cnf(i_0_5542, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,k3_relat_1(k29_sin_cos,k24_sin_cos)))).#trainneg
cnf(i_0_12630, negated_conjecture, (m1_subset_1(X1,a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(esk3_0,k2_zfmisc_1(esk3_0,X2))))).#trainneg
cnf(i_0_24562, plain, (r2_hidden(X1,a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k24_sin_cos))))).#trainneg
cnf(i_0_58450, negated_conjecture, (m1_subset_1(k3_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_64626, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_64635, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(a_0_0_fdiff_8,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_64652, plain, (m1_subset_1(a_0_0_fdiff_8,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_64661, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_63829, negated_conjecture, (k7_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k7_xcmplx_0(esk4_0,np__1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),esk4_0)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0))|~v1_funct_1(X1)|~v1_xcmplx_0(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)))|~m1_subset_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_63840, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk4_0)|~r1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))|~v1_xcmplx_0(k7_xcmplx_0(np__1,esk4_0))|~m1_subset_1(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)),k1_numbers))).#trainneg
cnf(i_0_64745, plain, (k1_fdiff_1(k3_taylor_1,k7_xcmplx_0(np__1,X1))=X1|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_62232, plain, (k7_xcmplx_0(k3_square_1(k1_seq_1(X1,X2)),k1_fdiff_1(X1,X2))=k7_xcmplx_0(k1_seq_1(X1,X2),k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2))|~v1_xcmplx_0(k1_fdiff_1(X1,X2))|~v1_xcmplx_0(k1_seq_1(X1,X2))|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~v1_xreal_0(k1_seq_1(X1,X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_63043, plain, (k7_xcmplx_0(X1,k3_square_1(X2))=k7_xcmplx_0(X1,k5_square_1(X2))|~v1_xcmplx_0(k3_square_1(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_square_1(X2))|~v1_xcmplx_0(X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_62408, plain, (k7_xcmplx_0(k1_fdiff_1(X1,X2),k5_square_1(k1_seq_1(X1,X2)))=k7_xcmplx_0(k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,X1,k3_taylor_1),X2),k1_seq_1(X1,X2))|~v1_xcmplx_0(k1_seq_1(X1,X2))|~v1_xcmplx_0(k1_fdiff_1(X1,X2))|~m1_subset_1(k1_seq_1(X1,X2),k1_numbers)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(X1,X2))|~r1_fdiff_1(X1,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k1_fdiff_1(X1,X2),k1_numbers)|~m1_subset_1(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_52308, negated_conjecture, (r1_tarski(X1,X2)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r1_tarski(X1,esk3_0)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k19_sin_cos),esk3_0))).#trainneg
cnf(i_0_1942, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,X2)|~r1_tarski(X2,k29_sin_cos))).#trainneg
cnf(i_0_24567, plain, (m1_subset_1(X1,a_0_0_fdiff_8)|~r2_hidden(X1,k1_relset_1(k1_numbers,k3_relat_1(k3_taylor_1,k24_sin_cos))))).#trainneg
cnf(i_0_39907, negated_conjecture, (r1_tarski(k3_relat_1(k19_sin_cos,k29_sin_cos),a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(k2_zfmisc_1(k1_numbers,k1_numbers),k1_zfmisc_1(a_0_0_fdiff_8)))).#trainneg
cnf(i_0_63018, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(k3_xcmplx_0(np__1,X2),X3))=k7_xcmplx_0(X1,k7_xcmplx_0(X2,X3))|~v1_xcmplx_0(k7_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_63658, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k7_xcmplx_0(esk4_0,k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~m1_subset_1(k7_xcmplx_0(esk4_0,k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),a_0_0_fdiff_8)|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))).#trainneg
cnf(i_0_63238, plain, (k8_real_1(k1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2)),k7_xcmplx_0(np__1,X2))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X1),X2)|~r1_fdiff_1(X1,k1_seq_1(k3_taylor_1,X2))|~m2_subset_1(k7_xcmplx_0(np__1,X2),k1_numbers,a_0_0_fdiff_8)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2)|~v1_xcmplx_0(X2)|~m1_subset_1(k7_xcmplx_0(np__1,X2),a_0_0_fdiff_8))).#trainneg
cnf(i_0_63659, negated_conjecture, (k7_xcmplx_0(np__1,k7_xcmplx_0(esk4_0,k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))).#trainneg
cnf(i_0_64832, plain, (k7_xcmplx_0(k7_xcmplx_0(np__1,X1),k1_seq_1(k3_taylor_1,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~v1_xreal_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(k7_xcmplx_0(np__1,X1))|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)),k1_numbers)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_64223, plain, (k7_xcmplx_0(k1_fdiff_1(X2,k1_seq_1(k3_taylor_1,X1)),X1)=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,X2),X1)|~v1_xcmplx_0(k1_fdiff_1(X2,k1_seq_1(k3_taylor_1,X1)))|~v1_xcmplx_0(X1)|~r1_fdiff_1(X2,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~v1_funct_1(X2)|~v1_xreal_0(k1_fdiff_1(X2,k1_seq_1(k3_taylor_1,X1)))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(k7_xcmplx_0(np__1,X1),k1_numbers)|~m1_subset_1(X1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_52609, negated_conjecture, (r1_tarski(X1,X2)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r1_tarski(X1,esk3_0)|~r2_fdiff_1(k3_relat_1(k29_sin_cos,k19_sin_cos),esk3_0))).#trainneg
cnf(i_0_1951, plain, (v1_xboole_0(k1_zfmisc_1(X1))|r2_hidden(k1_relset_1(X1,k1_partfun1(X1,X2,X3,X4,X5,X6)),k1_zfmisc_1(X1))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_63657, negated_conjecture, (k1_fdiff_1(k3_taylor_1,k7_xcmplx_0(esk4_0,k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~m2_subset_1(k7_xcmplx_0(esk4_0,k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))),k1_numbers,a_0_0_fdiff_8)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))).#trainneg
cnf(i_0_63656, negated_conjecture, (k10_real_1(np__1,k7_xcmplx_0(esk4_0,k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))!=k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))))|~v1_xcmplx_0(k5_square_1(k1_seq_1(k19_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))))).#trainneg
cnf(i_0_64556, plain, (k3_xcmplx_0(k7_xcmplx_0(X1,X2),k5_xcmplx_0(X2))=k7_xcmplx_0(X1,k5_square_1(X2))|~v1_xcmplx_0(k7_xcmplx_0(X1,X2))|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_63241, negated_conjecture, (k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,esk4_0))=esk4_0|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_64168, plain, (k3_xcmplx_0(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)),k7_xcmplx_0(np__1,X1))=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k3_taylor_1),X1)|~r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~r1_fdiff_1(k3_taylor_1,X1)|~v1_xreal_0(k7_xcmplx_0(np__1,X1))|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(np__1,k1_seq_1(k3_taylor_1,X1)),k1_numbers)|~m1_subset_1(X1,a_0_0_fdiff_8)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_62399, plain, (k7_xcmplx_0(X1,k3_xcmplx_0(np__1,k5_xcmplx_0(X2)))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X2))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_66533, negated_conjecture, (r1_fdiff_1(k3_taylor_1,esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xcmplx_0(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers)|~m2_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8))).#trainneg
cnf(i_0_67700, negated_conjecture, (v1_relat_1(esk3_0)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_1956, plain, (v1_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),k1_partfun1(k2_zfmisc_1(X1,X2),X3,X4,X5,X6,X7)))|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_67719, negated_conjecture, (v3_valued_0(esk3_0)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_67028, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(np__1,k6_numbers),k6_numbers)|~m2_subset_1(k6_numbers,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_67767, negated_conjecture, (v1_relat_1(esk3_0)|~m1_subset_1(a_0_0_fdiff_8,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_67786, negated_conjecture, (v3_valued_0(esk3_0)|~m1_subset_1(a_0_0_fdiff_8,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_67834, plain, (v1_relat_1(a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_67853, plain, (v3_valued_0(a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_69400, negated_conjecture, (k3_xcmplx_0(esk4_0,np__1)=esk4_0|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_66541, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,esk4_0))|~m2_subset_1(k1_seq_1(k3_taylor_1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_66944, negated_conjecture, (k1_seq_1(k2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0),esk4_0)!=k1_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk4_0)|~r1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0))|~r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xreal_0(k1_fdiff_1(k29_sin_cos,k1_seq_1(k3_taylor_1,esk4_0)))|~v1_xcmplx_0(k7_xcmplx_0(np__1,esk4_0))|~v1_membered(k1_numbers))).#trainneg
cnf(i_0_67701, negated_conjecture, (v4_relat_1(esk3_0,k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_1957, plain, (v4_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),k1_partfun1(k2_zfmisc_1(X1,X2),X3,X4,X5,X6,X7)),X1)|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_67901, plain, (v1_relat_1(k5_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_67920, plain, (v3_valued_0(k5_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_66536, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k1_funct_1(k3_taylor_1,X1))|~v1_xcmplx_0(k1_funct_1(k3_taylor_1,X1))|~m1_subset_1(k1_funct_1(k3_taylor_1,X1),k1_numbers)|~m2_subset_1(k1_funct_1(k3_taylor_1,X1),k1_numbers,a_0_0_fdiff_8)|~r2_hidden(X1,esk3_0)|~m1_subset_1(k7_xcmplx_0(np__1,k1_funct_1(k3_taylor_1,X1)),k1_numbers)|~m1_subset_1(k1_funct_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_66289, plain, (r1_fdiff_1(k3_taylor_1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_69399, negated_conjecture, (k3_xcmplx_0(np__1,esk4_0)=esk4_0|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_67720, negated_conjecture, (v5_relat_1(esk3_0,k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_1959, plain, (v5_relat_1(k1_relset_1(k2_zfmisc_1(X1,X2),k1_partfun1(k2_zfmisc_1(X1,X2),X3,X4,X5,X6,X7)),X2)|~v1_funct_1(X7)|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(k2_zfmisc_1(X1,X2),X3))))).#trainneg
cnf(i_0_66538, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k6_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_67768, negated_conjecture, (v4_relat_1(esk3_0,k1_numbers)|~m1_subset_1(a_0_0_fdiff_8,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_67787, negated_conjecture, (v5_relat_1(esk3_0,k1_numbers)|~m1_subset_1(a_0_0_fdiff_8,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_67835, plain, (v4_relat_1(a_0_0_fdiff_8,k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_66215, plain, (k3_xcmplx_0(X1,k1_fdiff_1(k3_taylor_1,X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(X2,a_0_0_fdiff_8))).#trainneg
cnf(i_0_66129, plain, (k3_xcmplx_0(X1,k1_fdiff_1(k3_taylor_1,X2))=k7_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X2,a_0_0_fdiff_8)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_66130, plain, (k7_xcmplx_0(X1,k1_fdiff_1(k3_taylor_1,X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(X2,a_0_0_fdiff_8)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_66537, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k1_funct_1(k3_taylor_1,esk4_0))|~m2_subset_1(k1_funct_1(k3_taylor_1,esk4_0),k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_funct_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_66216, plain, (k7_xcmplx_0(X1,k1_fdiff_1(k3_taylor_1,X2))=k3_xcmplx_0(X1,X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(X2)|~m1_subset_1(X2,a_0_0_fdiff_8))).#trainneg
cnf(i_0_66988, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k24_sin_cos)))).#trainneg
cnf(i_0_1969, plain, (m1_subset_1(k1_relset_1(X1,k3_relat_1(X5,X6)),k1_zfmisc_1(X1))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_67854, plain, (v5_relat_1(a_0_0_fdiff_8,k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_67902, plain, (v4_relat_1(k5_numbers,k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_67921, plain, (v5_relat_1(k5_numbers,k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_66293, plain, (k7_xcmplx_0(X1,k7_xcmplx_0(X2,k1_fdiff_1(k3_taylor_1,X3)))=k7_xcmplx_0(k7_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xreal_0(X3)|~m1_subset_1(X3,a_0_0_fdiff_8))).#trainneg
cnf(i_0_66241, plain, (k3_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2)=k7_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_66253, plain, (m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_66545, plain, (m1_subset_1(X1,a_0_0_fdiff_8)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(k7_xcmplx_0(np__1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_66239, plain, (k5_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))=X1|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_69793, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,esk4_0))|~m1_subset_1(k1_seq_1(k3_taylor_1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_69766, negated_conjecture, (r1_fdiff_1(k3_taylor_1,esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),esk3_0)|~v1_xcmplx_0(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(esk6_2(esk3_0,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8))).#trainneg
cnf(i_0_67636, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k19_sin_cos)))).#trainneg
cnf(i_0_1977, negated_conjecture, (m1_subset_1(X1,k1_relset_1(k1_numbers,k1_partfun1(X2,X3,X4,X5,k3_taylor_1,k29_sin_cos)))|~r2_hidden(X1,esk3_0)|~m1_subset_1(k29_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_71933, negated_conjecture, (k5_xcmplx_0(k7_xcmplx_0(np__1,esk4_0))=esk4_0)).#trainneg
cnf(i_0_67077, negated_conjecture, (r2_hidden(k6_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(k6_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_69822, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k1_seq_1(k3_taylor_1,X1))|~m2_subset_1(k1_seq_1(k3_taylor_1,X1),k1_numbers,a_0_0_fdiff_8)|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_seq_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_69831, negated_conjecture, (r1_fdiff_1(k3_taylor_1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v1_xreal_0(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~v1_xcmplx_0(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(k7_xcmplx_0(np__1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))),a_0_0_fdiff_8)|~m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~v1_funct_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))|~m1_subset_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_69823, negated_conjecture, (r1_fdiff_1(k3_taylor_1,k1_funct_1(k3_taylor_1,X1))|~r2_hidden(X1,esk3_0)|~m1_subset_1(k1_funct_1(k3_taylor_1,X1),a_0_0_fdiff_8))).#trainneg
cnf(i_0_69827, negated_conjecture, (r1_fdiff_1(k3_taylor_1,esk6_2(esk3_0,X1))|r2_fdiff_1(X1,esk3_0)|~v1_xreal_0(esk6_2(esk3_0,X1))|~v1_xcmplx_0(esk6_2(esk3_0,X1))|~m1_subset_1(k7_xcmplx_0(np__1,esk6_2(esk3_0,X1)),a_0_0_fdiff_8)|~m1_subset_1(esk6_2(esk3_0,X1),a_0_0_fdiff_8)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k1_relset_1(k1_numbers,X1))))).#trainneg
cnf(i_0_69829, negated_conjecture, (r1_fdiff_1(k3_taylor_1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|r2_fdiff_1(k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos),X1)|~v1_xreal_0(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~v1_xcmplx_0(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)))|~m1_subset_1(k7_xcmplx_0(np__1,esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos))),a_0_0_fdiff_8)|~m1_subset_1(esk6_2(X1,k1_partfun1(k1_numbers,k1_numbers,k1_numbers,k1_numbers,k3_taylor_1,k29_sin_cos)),a_0_0_fdiff_8)|~v3_rcomp_1(X1)|~r1_tarski(X1,esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_69784, negated_conjecture, (v1_xcmplx_0(k1_seq_1(esk3_0,X1))|~v1_funct_1(esk3_0)|~v1_relat_1(esk3_0)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_68531, plain, (r1_tarski(X1,k2_zfmisc_1(k1_numbers,k1_numbers))|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k29_sin_cos)))).#trainneg
cnf(i_0_67083, negated_conjecture, (v1_xreal_0(k3_xcmplx_0(esk4_0,np__1))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_67095, negated_conjecture, (v1_xreal_0(k7_xcmplx_0(esk4_0,np__1))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_67684, negated_conjecture, (v1_xreal_0(k3_xcmplx_0(np__1,esk4_0))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_69857, plain, (r1_fdiff_1(k3_taylor_1,k1_fdiff_1(k3_taylor_1,X1))|~v1_xreal_0(k1_fdiff_1(k3_taylor_1,X1))|~v1_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))|~m1_subset_1(k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1)),a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),a_0_0_fdiff_8)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_71806, plain, (m1_subset_1(k5_xcmplx_0(X1),a_0_0_fdiff_8)|~v1_xreal_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(k5_xcmplx_0(X1))|~m1_subset_1(X1,a_0_0_fdiff_8)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_69876, plain, (r1_fdiff_1(k3_taylor_1,X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8)|~m2_subset_1(X1,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k1_fdiff_1(k3_taylor_1,X1),k1_numbers))).#trainneg
cnf(i_0_69800, negated_conjecture, (m1_subset_1(k1_relset_1(k1_numbers,esk3_0),k1_zfmisc_1(k1_numbers))|~v1_relat_1(esk3_0)|~m1_subset_1(k1_numbers,k1_zfmisc_1(k3_taylor_1)))).#trainneg
cnf(i_0_52926, negated_conjecture, (r1_tarski(X1,X2)|~m1_subset_1(k3_taylor_1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r1_tarski(X1,esk3_0)|~r2_fdiff_1(k3_relat_1(k3_taylor_1,k29_sin_cos),esk3_0))).#trainneg
cnf(i_0_66833, negated_conjecture, (v1_xcmplx_0(k7_xcmplx_0(np__1,k6_numbers))|~v1_membered(k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_66834, negated_conjecture, (v7_ordinal1(k7_xcmplx_0(np__1,k6_numbers))|~v6_membered(k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_66840, negated_conjecture, (v1_xcmplx_0(k3_xcmplx_0(np__1,esk4_0))|~v1_membered(k1_numbers)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_66841, negated_conjecture, (v7_ordinal1(k3_xcmplx_0(np__1,esk4_0))|~v6_membered(k1_numbers)|~m2_subset_1(esk4_0,k1_numbers,a_0_0_fdiff_8))).#trainneg
cnf(i_0_66153, plain, (k5_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1))=X1|~v1_xcmplx_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_71696, plain, (v1_xreal_0(k1_fdiff_1(k3_taylor_1,X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_69401, negated_conjecture, (k5_xcmplx_0(esk4_0)=k7_xcmplx_0(np__1,esk4_0)|~v1_xcmplx_0(k7_xcmplx_0(np__1,esk4_0))|~m1_subset_1(k7_xcmplx_0(np__1,esk4_0),a_0_0_fdiff_8))).#trainneg
cnf(i_0_66261, plain, (k7_xcmplx_0(np__1,k1_fdiff_1(k3_taylor_1,X1))=k3_xcmplx_0(np__1,X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8))).#trainneg
cnf(i_0_66155, plain, (k3_xcmplx_0(k1_fdiff_1(k3_taylor_1,X1),X2)=k7_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_69825, negated_conjecture, (r1_fdiff_1(k3_taylor_1,esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0)))|r2_fdiff_1(k2_fdiff_1(X1,esk3_0),esk3_0)|~v1_xreal_0(esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0)))|~v1_xcmplx_0(esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0)))|~m1_subset_1(k7_xcmplx_0(np__1,esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0))),a_0_0_fdiff_8)|~m1_subset_1(esk6_2(esk3_0,k2_fdiff_1(X1,esk3_0)),a_0_0_fdiff_8)|~r2_fdiff_1(X1,esk3_0)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k1_numbers,k1_numbers))))).#trainneg
cnf(i_0_69838, plain, (r1_fdiff_1(k3_taylor_1,k1_funct_1(X1,X2))|~v1_xreal_0(k1_funct_1(X1,X2))|~v1_xcmplx_0(k1_funct_1(X1,X2))|~m1_subset_1(k7_xcmplx_0(np__1,k1_funct_1(X1,X2)),a_0_0_fdiff_8)|~m1_subset_1(k1_funct_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(k1_seq_1(X1,X2),a_0_0_fdiff_8)|~r2_hidden(X2,k9_xtuple_0(X1))|~v1_funct_1(X1)|~v3_valued_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_54887, negated_conjecture, (r1_tarski(X1,X2)|~m1_subset_1(k19_sin_cos,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))|~r1_tarski(X1,esk3_0)|~r2_fdiff_1(k3_relat_1(k19_sin_cos,k29_sin_cos),esk3_0))).#trainneg
cnf(i_0_1994, plain, (k7_xcmplx_0(np__1,k3_xcmplx_0(np__1,X1))=k5_xcmplx_0(X1)|~v1_xcmplx_0(k5_xcmplx_0(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_69792, negated_conjecture, (k3_xcmplx_0(esk4_0,np__1)=esk4_0|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_69878, negated_conjecture, (k3_xcmplx_0(np__1,esk4_0)=esk4_0|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_67045, negated_conjecture, (~r1_xxreal_0(k7_xcmplx_0(np__1,k7_xcmplx_0(np__1,k6_numbers)),k6_numbers)|~m2_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers,a_0_0_fdiff_8)|~m2_subset_1(k6_numbers,k1_numbers,a_0_0_fdiff_8)|~m1_subset_1(k7_xcmplx_0(np__1,k6_numbers),k1_numbers)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_67081, negated_conjecture, (v1_xcmplx_0(k3_xcmplx_0(esk4_0,np__1))|~v1_membered(a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_67082, negated_conjecture, (v7_ordinal1(k3_xcmplx_0(esk4_0,np__1))|~v6_membered(a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_67093, negated_conjecture, (v1_xcmplx_0(k7_xcmplx_0(esk4_0,np__1))|~v1_membered(a_0_0_fdiff_8)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_66131, plain, (k7_xcmplx_0(k5_xcmplx_0(X1),X2)=k7_xcmplx_0(k1_fdiff_1(k3_taylor_1,X2),X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(X2,a_0_0_fdiff_8)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_66250, plain, (k7_xcmplx_0(X1,k1_fdiff_1(k3_taylor_1,X1))=k3_square_1(X1)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,a_0_0_fdiff_8))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 85
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 135
# Removed in clause preprocessing      : 2
# Initial clauses in saturation        : 133
# Processed clauses                    : 5927
# ...of these trivial                  : 138
# ...subsumed                          : 2653
# ...remaining for further processing  : 3136
# Other redundant clauses eliminated   : 5
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 135
# Backward-rewritten                   : 110
# Generated clauses                    : 73937
# ...of the previous two non-trivial   : 26625
# Contextual simplify-reflections      : 395
# Paramodulations                      : 73925
# Factorizations                       : 0
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 2885
#    Positive orientable unit clauses  : 551
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 10
#    Non-unit-clauses                  : 2323
# Current number of unprocessed clauses: 20534
# ...number of literals in the above   : 109062
# Current number of archived formulas  : 0
# Current number of archived clauses   : 3385
# Proof object given clauses           : 218
# Proof search given clauses           : 3136
# Clause-clause subsumption calls (NU) : 1167256
# Rec. Clause-clause subsumption calls : 252524
# Non-unit clause-clause subsumptions  : 3046
# Unit Clause-clause subsumption calls : 52493
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 3884
# BW rewrite match successes           : 30
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2002524

# -------------------------------------------------
# User time                : 31.024 s
# System time              : 0.783 s
# Total time               : 31.807 s
# Maximum resident set size: 222856 pages
31.02user 0.81system 1:06.30elapsed 48%CPU (0avgtext+0avgdata 223076maxresident)k
32inputs+1568outputs (0major+217767minor)pagefaults 0swaps
