2020-08-31 10:45:46.516151: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_79
2020-08-31 10:45:46.596702: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-08-31 10:45:46.707770: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-08-31 10:45:47.035020: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-08-31 10:45:48.620232: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 2104082 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_79' loaded (query=768; context=1536; weigths=real; len_mult=0.000000).
# Preprocessing time       : 2.538 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t140_zf_lang1, axiom, ![X1]:((v1_zf_lang(X1)&m2_finseq_1(X1,k5_numbers))=>k3_zf_lang1(k6_zf_lang(X1))=k3_zf_lang1(X1)), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', t140_zf_lang1)).
fof(t59_zf_lang1, axiom, ![X1]:(m2_subset_1(X1,k5_numbers,k1_zf_lang)=>![X2]:(m2_subset_1(X2,k5_numbers,k1_zf_lang)=>k2_zf_model(k5_zf_lang(X1,X2))=k2_tarski(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', t59_zf_lang1)).
fof(t139_zf_lang1, axiom, ![X1]:(m2_subset_1(X1,k5_numbers,k1_zf_lang)=>![X2]:(m2_subset_1(X2,k5_numbers,k1_zf_lang)=>k3_zf_lang1(k5_zf_lang(X1,X2))=k2_tarski(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', t139_zf_lang1)).
fof(t60_zf_lang1, axiom, ![X1]:((v1_zf_lang(X1)&m2_finseq_1(X1,k5_numbers))=>k2_zf_model(k6_zf_lang(X1))=k2_zf_model(X1)), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', t60_zf_lang1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', commutativity_k2_tarski)).
fof(reflexivity_r1_tarski, axiom, ![X1, X2]:r1_tarski(X1,X1), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', reflexivity_r1_tarski)).
fof(t58_zf_lang1, axiom, ![X1]:(m2_subset_1(X1,k5_numbers,k1_zf_lang)=>![X2]:(m2_subset_1(X2,k5_numbers,k1_zf_lang)=>k2_zf_model(k4_zf_lang(X1,X2))=k2_tarski(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', t58_zf_lang1)).
fof(t138_zf_lang1, axiom, ![X1]:(m2_subset_1(X1,k5_numbers,k1_zf_lang)=>![X2]:(m2_subset_1(X2,k5_numbers,k1_zf_lang)=>k3_zf_lang1(k4_zf_lang(X1,X2))=k2_tarski(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', t138_zf_lang1)).
fof(t13_xboole_1, axiom, ![X1, X2, X3, X4]:((r1_tarski(X1,X2)&r1_tarski(X3,X4))=>r1_tarski(k2_xboole_0(X1,X3),k2_xboole_0(X2,X4))), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', t13_xboole_1)).
fof(t141_zf_lang1, axiom, ![X1]:((v1_zf_lang(X1)&m2_finseq_1(X1,k5_numbers))=>![X2]:((v1_zf_lang(X2)&m2_finseq_1(X2,k5_numbers))=>k3_zf_lang1(k7_zf_lang(X1,X2))=k2_xboole_0(k3_zf_lang1(X1),k3_zf_lang1(X2)))), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', t141_zf_lang1)).
fof(redefinition_k4_subset_1, axiom, ![X1, X2, X3]:((m1_subset_1(X2,k1_zfmisc_1(X1))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>k4_subset_1(X1,X2,X3)=k2_xboole_0(X2,X3)), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', redefinition_k4_subset_1)).
fof(t61_zf_lang1, axiom, ![X1]:((v1_zf_lang(X1)&m2_finseq_1(X1,k5_numbers))=>![X2]:((v1_zf_lang(X2)&m2_finseq_1(X2,k5_numbers))=>k2_zf_model(k7_zf_lang(X1,X2))=k4_subset_1(k1_zf_lang,k2_zf_model(X1),k2_zf_model(X2)))), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', t61_zf_lang1)).
fof(dt_k2_zf_model, axiom, ![X1]:((v1_zf_lang(X1)&m1_finseq_1(X1,k5_numbers))=>m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', dt_k2_zf_model)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', t1_xboole_1)).
fof(idempotence_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X1)=X1, file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', idempotence_k2_xboole_0)).
fof(t7_xboole_1, axiom, ![X1, X2]:r1_tarski(X1,k2_xboole_0(X1,X2)), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', t7_xboole_1)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', commutativity_k2_xboole_0)).
fof(redefinition_m2_finseq_1, axiom, ![X1, X2]:(m2_finseq_1(X2,X1)<=>m1_finseq_1(X2,X1)), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', redefinition_m2_finseq_1)).
fof(t36_xboole_1, axiom, ![X1, X2]:r1_tarski(k4_xboole_0(X1,X2),X1), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', t36_xboole_1)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', t1_boole)).
fof(t142_zf_lang1, axiom, ![X1]:((v1_zf_lang(X1)&m2_finseq_1(X1,k5_numbers))=>![X2]:(m2_subset_1(X2,k5_numbers,k1_zf_lang)=>k3_zf_lang1(k8_zf_lang(X2,X1))=k2_xboole_0(k3_zf_lang1(X1),k1_tarski(X2)))), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', t142_zf_lang1)).
fof(s1_zf_lang1__e5_143__zf_lang1, axiom, ((((![X1]:(m2_subset_1(X1,k5_numbers,k1_zf_lang)=>![X2]:(m2_subset_1(X2,k5_numbers,k1_zf_lang)=>(r1_tarski(k2_zf_model(k4_zf_lang(X1,X2)),k3_zf_lang1(k4_zf_lang(X1,X2)))&r1_tarski(k2_zf_model(k5_zf_lang(X1,X2)),k3_zf_lang1(k5_zf_lang(X1,X2))))))&![X1]:((v1_zf_lang(X1)&m2_finseq_1(X1,k5_numbers))=>(r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))=>r1_tarski(k2_zf_model(k6_zf_lang(X1)),k3_zf_lang1(k6_zf_lang(X1))))))&![X1]:((v1_zf_lang(X1)&m2_finseq_1(X1,k5_numbers))=>![X2]:((v1_zf_lang(X2)&m2_finseq_1(X2,k5_numbers))=>((r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))&r1_tarski(k2_zf_model(X2),k3_zf_lang1(X2)))=>r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k3_zf_lang1(k7_zf_lang(X1,X2)))))))&![X1]:((v1_zf_lang(X1)&m2_finseq_1(X1,k5_numbers))=>![X2]:(m2_subset_1(X2,k5_numbers,k1_zf_lang)=>(r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))=>r1_tarski(k2_zf_model(k8_zf_lang(X2,X1)),k3_zf_lang1(k8_zf_lang(X2,X1)))))))=>![X1]:((v1_zf_lang(X1)&m2_finseq_1(X1,k5_numbers))=>r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1)))), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', s1_zf_lang1__e5_143__zf_lang1)).
fof(t151_zf_lang1, conjecture, ![X1]:((v1_zf_lang(X1)&m2_finseq_1(X1,k5_numbers))=>r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', t151_zf_lang1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', t3_subset)).
fof(redefinition_k7_subset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>k7_subset_1(X1,X2,X3)=k4_xboole_0(X2,X3)), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', redefinition_k7_subset_1)).
fof(t62_zf_lang1, axiom, ![X1]:((v1_zf_lang(X1)&m2_finseq_1(X1,k5_numbers))=>![X2]:(m2_subset_1(X2,k5_numbers,k1_zf_lang)=>k2_zf_model(k8_zf_lang(X2,X1))=k7_subset_1(k1_zf_lang,k2_zf_model(X1),k1_tarski(X2)))), file('/scratch/mptp/hardpred/preds__0.005//t151_zf_lang1', t62_zf_lang1)).
fof(c_0_26, plain, (epred1_0<=>((![X1]:(m2_subset_1(X1,k5_numbers,k1_zf_lang)=>![X2]:(m2_subset_1(X2,k5_numbers,k1_zf_lang)=>(r1_tarski(k2_zf_model(k4_zf_lang(X1,X2)),k3_zf_lang1(k4_zf_lang(X1,X2)))&r1_tarski(k2_zf_model(k5_zf_lang(X1,X2)),k3_zf_lang1(k5_zf_lang(X1,X2))))))&![X1]:((v1_zf_lang(X1)&m2_finseq_1(X1,k5_numbers))=>(r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))=>r1_tarski(k2_zf_model(k6_zf_lang(X1)),k3_zf_lang1(k6_zf_lang(X1))))))&![X1]:((v1_zf_lang(X1)&m2_finseq_1(X1,k5_numbers))=>![X2]:((v1_zf_lang(X2)&m2_finseq_1(X2,k5_numbers))=>((r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))&r1_tarski(k2_zf_model(X2),k3_zf_lang1(X2)))=>r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k3_zf_lang1(k7_zf_lang(X1,X2)))))))), introduced(definition)).
fof(c_0_27, plain, (((![X1]:(m2_subset_1(X1,k5_numbers,k1_zf_lang)=>![X2]:(m2_subset_1(X2,k5_numbers,k1_zf_lang)=>(r1_tarski(k2_zf_model(k4_zf_lang(X1,X2)),k3_zf_lang1(k4_zf_lang(X1,X2)))&r1_tarski(k2_zf_model(k5_zf_lang(X1,X2)),k3_zf_lang1(k5_zf_lang(X1,X2))))))&![X1]:((v1_zf_lang(X1)&m2_finseq_1(X1,k5_numbers))=>(r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))=>r1_tarski(k2_zf_model(k6_zf_lang(X1)),k3_zf_lang1(k6_zf_lang(X1))))))&![X1]:((v1_zf_lang(X1)&m2_finseq_1(X1,k5_numbers))=>![X2]:((v1_zf_lang(X2)&m2_finseq_1(X2,k5_numbers))=>((r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))&r1_tarski(k2_zf_model(X2),k3_zf_lang1(X2)))=>r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k3_zf_lang1(k7_zf_lang(X1,X2)))))))=>epred1_0), inference(split_equiv,[status(thm)],[c_0_26])).
fof(c_0_28, plain, ((((((v1_zf_lang(esk16_0)|(v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0)&(m2_finseq_1(esk16_0,k5_numbers)|(v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0))&(((v1_zf_lang(esk17_0)|(v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0)&(m2_finseq_1(esk17_0,k5_numbers)|(v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0))&(((r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|(v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0)&(r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|(v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0))&(~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|(v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0))))&(((v1_zf_lang(esk16_0)|(m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0)&(m2_finseq_1(esk16_0,k5_numbers)|(m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0))&(((v1_zf_lang(esk17_0)|(m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0)&(m2_finseq_1(esk17_0,k5_numbers)|(m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0))&(((r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|(m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0)&(r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|(m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0))&(~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|(m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0)))))&((((v1_zf_lang(esk16_0)|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0)&(m2_finseq_1(esk16_0,k5_numbers)|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0))&(((v1_zf_lang(esk17_0)|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0)&(m2_finseq_1(esk17_0,k5_numbers)|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0))&(((r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0)&(r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0))&(~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0))))&(((v1_zf_lang(esk16_0)|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0)&(m2_finseq_1(esk16_0,k5_numbers)|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0))&(((v1_zf_lang(esk17_0)|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0)&(m2_finseq_1(esk17_0,k5_numbers)|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0))&(((r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0)&(r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0))&(~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))|epred1_0))))))&((((((v1_zf_lang(esk16_0)|(v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0)&(m2_finseq_1(esk16_0,k5_numbers)|(v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0))&(((v1_zf_lang(esk17_0)|(v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0)&(m2_finseq_1(esk17_0,k5_numbers)|(v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0))&(((r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|(v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0)&(r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|(v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0))&(~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|(v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0))))&(((v1_zf_lang(esk16_0)|(m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0)&(m2_finseq_1(esk16_0,k5_numbers)|(m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0))&(((v1_zf_lang(esk17_0)|(m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0)&(m2_finseq_1(esk17_0,k5_numbers)|(m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0))&(((r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|(m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0)&(r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|(m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0))&(~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|(m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0)))))&((((v1_zf_lang(esk16_0)|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0)&(m2_finseq_1(esk16_0,k5_numbers)|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0))&(((v1_zf_lang(esk17_0)|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0)&(m2_finseq_1(esk17_0,k5_numbers)|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0))&(((r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0)&(r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0))&(~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0))))&(((v1_zf_lang(esk16_0)|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0)&(m2_finseq_1(esk16_0,k5_numbers)|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0))&(((v1_zf_lang(esk17_0)|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0)&(m2_finseq_1(esk17_0,k5_numbers)|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0))&(((r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0)&(r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0))&(~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))|epred1_0))))))&(((((v1_zf_lang(esk16_0)|(v1_zf_lang(esk15_0)|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0)&(m2_finseq_1(esk16_0,k5_numbers)|(v1_zf_lang(esk15_0)|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0))&(((v1_zf_lang(esk17_0)|(v1_zf_lang(esk15_0)|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0)&(m2_finseq_1(esk17_0,k5_numbers)|(v1_zf_lang(esk15_0)|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0))&(((r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|(v1_zf_lang(esk15_0)|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0)&(r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|(v1_zf_lang(esk15_0)|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0))&(~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|(v1_zf_lang(esk15_0)|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0))))&(((v1_zf_lang(esk16_0)|(m2_finseq_1(esk15_0,k5_numbers)|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0)&(m2_finseq_1(esk16_0,k5_numbers)|(m2_finseq_1(esk15_0,k5_numbers)|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0))&(((v1_zf_lang(esk17_0)|(m2_finseq_1(esk15_0,k5_numbers)|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0)&(m2_finseq_1(esk17_0,k5_numbers)|(m2_finseq_1(esk15_0,k5_numbers)|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0))&(((r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|(m2_finseq_1(esk15_0,k5_numbers)|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0)&(r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|(m2_finseq_1(esk15_0,k5_numbers)|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0))&(~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|(m2_finseq_1(esk15_0,k5_numbers)|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0)))))&((((v1_zf_lang(esk16_0)|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0)&(m2_finseq_1(esk16_0,k5_numbers)|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0))&(((v1_zf_lang(esk17_0)|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0)&(m2_finseq_1(esk17_0,k5_numbers)|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0))&(((r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0)&(r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0))&(~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|(r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0))))&(((v1_zf_lang(esk16_0)|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0)&(m2_finseq_1(esk16_0,k5_numbers)|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0))&(((v1_zf_lang(esk17_0)|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0)&(m2_finseq_1(esk17_0,k5_numbers)|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0))&(((r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0)&(r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0))&(~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|(~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|(~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))))|epred1_0)))))))), inference(distribute,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_27])])])])).
fof(c_0_29, plain, ![X23]:(~v1_zf_lang(X23)|~m2_finseq_1(X23,k5_numbers)|k3_zf_lang1(k6_zf_lang(X23))=k3_zf_lang1(X23)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t140_zf_lang1])])).
cnf(c_0_30, plain, (r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_31, plain, (k3_zf_lang1(k6_zf_lang(X1))=k3_zf_lang1(X1)|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_29])).
fof(c_0_32, plain, ![X38, X39]:(~m2_subset_1(X38,k5_numbers,k1_zf_lang)|(~m2_subset_1(X39,k5_numbers,k1_zf_lang)|k2_zf_model(k5_zf_lang(X38,X39))=k2_tarski(X38,X39))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t59_zf_lang1])])])).
fof(c_0_33, plain, ![X29, X30]:(~m2_subset_1(X29,k5_numbers,k1_zf_lang)|(~m2_subset_1(X30,k5_numbers,k1_zf_lang)|k3_zf_lang1(k5_zf_lang(X29,X30))=k2_tarski(X29,X30))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t139_zf_lang1])])])).
cnf(c_0_34, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(evalgc,[status(thm)],[c_0_30])).
cnf(c_0_35, plain, (k3_zf_lang1(k6_zf_lang(X1))=k3_zf_lang1(X1)|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_31])).
cnf(c_0_36, plain, (r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_37, plain, (r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_38, plain, ![X24]:(~v1_zf_lang(X24)|~m2_finseq_1(X24,k5_numbers)|k2_zf_model(k6_zf_lang(X24))=k2_zf_model(X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t60_zf_lang1])])).
fof(c_0_39, plain, ![X45, X46]:k2_tarski(X45,X46)=k2_tarski(X46,X45), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_40, plain, (r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_41, plain, (v1_zf_lang(esk16_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_42, plain, (k2_zf_model(k5_zf_lang(X1,X2))=k2_tarski(X1,X2)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_43, plain, (k3_zf_lang1(k5_zf_lang(X1,X2))=k2_tarski(X1,X2)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_44, plain, ![X106]:r1_tarski(X106,X106), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r1_tarski])])).
cnf(c_0_45, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_35])])).
cnf(c_0_46, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))), inference(evalgc,[status(thm)],[c_0_36])).
cnf(c_0_47, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))), inference(evalgc,[status(thm)],[c_0_37])).
cnf(c_0_48, plain, (k2_zf_model(k6_zf_lang(X1))=k2_zf_model(X1)|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_49, plain, ![X40, X41]:(~m2_subset_1(X40,k5_numbers,k1_zf_lang)|(~m2_subset_1(X41,k5_numbers,k1_zf_lang)|k2_zf_model(k4_zf_lang(X40,X41))=k2_tarski(X40,X41))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t58_zf_lang1])])])).
cnf(c_0_50, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_51, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(evalgc,[status(thm)],[c_0_40])).
cnf(c_0_52, plain, (r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_53, plain, (r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_54, plain, (v1_zf_lang(esk16_0)|epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_55, plain, (epred1_0|v1_zf_lang(esk16_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(evalgc,[status(thm)],[c_0_41])).
cnf(c_0_56, plain, (v1_zf_lang(esk16_0)|v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_57, plain, (v1_zf_lang(esk16_0)|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_58, plain, (v1_zf_lang(esk17_0)|epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
fof(c_0_59, plain, ![X31, X32]:(~m2_subset_1(X31,k5_numbers,k1_zf_lang)|(~m2_subset_1(X32,k5_numbers,k1_zf_lang)|k3_zf_lang1(k4_zf_lang(X31,X32))=k2_tarski(X31,X32))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t138_zf_lang1])])])).
cnf(c_0_60, plain, (r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_61, plain, (r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_62, plain, (k2_zf_model(k5_zf_lang(X1,X2))=k2_tarski(X1,X2)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_42])).
cnf(c_0_63, plain, (k3_zf_lang1(k5_zf_lang(X1,X2))=k2_tarski(X1,X2)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_43])).
cnf(c_0_64, plain, (r1_tarski(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_65, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_45, c_0_46]), c_0_47])).
cnf(c_0_66, plain, (k2_zf_model(k6_zf_lang(X1))=k2_zf_model(X1)|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_48])).
cnf(c_0_67, plain, (r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_68, plain, (r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_69, plain, (k2_zf_model(k4_zf_lang(X1,X2))=k2_tarski(X1,X2)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_70, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(evalgc,[status(thm)],[c_0_50])).
cnf(c_0_71, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_35])])).
cnf(c_0_72, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))), inference(evalgc,[status(thm)],[c_0_52])).
cnf(c_0_73, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))), inference(evalgc,[status(thm)],[c_0_53])).
cnf(c_0_74, plain, (epred1_0|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_54])).
cnf(c_0_75, plain, (v1_zf_lang(esk16_0)|v1_zf_lang(esk15_0)|epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_76, plain, (v1_zf_lang(esk16_0)|m2_finseq_1(esk15_0,k5_numbers)|epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_77, plain, (v1_zf_lang(esk16_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_78, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_55, c_0_35])])).
cnf(c_0_79, plain, (epred1_0|v1_zf_lang(esk15_0)|v1_zf_lang(esk16_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_56])).
cnf(c_0_80, plain, (epred1_0|v1_zf_lang(esk16_0)|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_57])).
cnf(c_0_81, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_58])).
cnf(c_0_82, plain, (v1_zf_lang(esk17_0)|v1_zf_lang(esk15_0)|epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_83, plain, (v1_zf_lang(esk17_0)|m2_finseq_1(esk15_0,k5_numbers)|epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_84, plain, (k3_zf_lang1(k4_zf_lang(X1,X2))=k2_tarski(X1,X2)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_85, plain, (v1_zf_lang(esk17_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_86, plain, (v1_zf_lang(esk17_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_87, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_60])).
cnf(c_0_88, plain, (r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|v1_zf_lang(esk15_0)|epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_89, plain, (r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|m2_finseq_1(esk15_0,k5_numbers)|epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_90, plain, (r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_91, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_61])).
cnf(c_0_92, plain, (k3_zf_lang1(k5_zf_lang(X1,X2))=k2_zf_model(k5_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63])])).
cnf(c_0_93, plain, (r1_tarski(X1,X1)), inference(evalgc,[status(thm)],[c_0_64])).
cnf(c_0_94, plain, (r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_95, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66])])).
cnf(c_0_96, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))), inference(evalgc,[status(thm)],[c_0_67])).
cnf(c_0_97, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_68])).
cnf(c_0_98, plain, (r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|v1_zf_lang(esk15_0)|epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_99, plain, (r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|m2_finseq_1(esk15_0,k5_numbers)|epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_100, plain, (r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_101, plain, (k2_zf_model(k4_zf_lang(X1,X2))=k2_tarski(X1,X2)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_69])).
cnf(c_0_102, plain, (k2_tarski(X2,X1)=k2_zf_model(k5_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_62])])).
cnf(c_0_103, plain, (r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_104, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_71, c_0_72]), c_0_73])).
cnf(c_0_105, plain, (r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_106, plain, (epred1_0|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_66])])).
cnf(c_0_107, plain, (epred1_0|v1_zf_lang(esk15_0)|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_75])).
cnf(c_0_108, plain, (epred1_0|v1_zf_lang(esk16_0)|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_109, plain, (epred1_0|v1_zf_lang(esk16_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(evalgc,[status(thm)],[c_0_77])).
cnf(c_0_110, plain, (v1_zf_lang(esk16_0)|v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_111, plain, (v1_zf_lang(esk16_0)|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_112, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_78, c_0_79]), c_0_80])).
cnf(c_0_113, plain, (v1_zf_lang(esk16_0)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_114, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_66])])).
cnf(c_0_115, plain, (epred1_0|v1_zf_lang(esk15_0)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_82])).
cnf(c_0_116, plain, (epred1_0|v1_zf_lang(esk17_0)|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_83])).
cnf(c_0_117, plain, (k3_zf_lang1(k4_zf_lang(X1,X2))=k2_tarski(X1,X2)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_84])).
cnf(c_0_118, plain, (epred1_0|v1_zf_lang(esk17_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_119, plain, (v1_zf_lang(esk17_0)|v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_120, plain, (v1_zf_lang(esk17_0)|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_121, plain, (epred1_0|v1_zf_lang(esk17_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_122, plain, (v1_zf_lang(esk17_0)|v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_123, plain, (v1_zf_lang(esk17_0)|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_124, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_66])])).
cnf(c_0_125, plain, (epred1_0|v1_zf_lang(esk15_0)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_88])).
cnf(c_0_126, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_127, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_128, plain, (r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_129, plain, (r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_130, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_93])])])).
cnf(c_0_131, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))), inference(evalgc,[status(thm)],[c_0_94])).
cnf(c_0_132, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_95, c_0_46]), c_0_47]), c_0_96])).
cnf(c_0_133, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_66])])).
cnf(c_0_134, plain, (epred1_0|v1_zf_lang(esk15_0)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_135, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_99])).
cnf(c_0_136, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(evalgc,[status(thm)],[c_0_100])).
cnf(c_0_137, plain, (r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_138, plain, (r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_139, plain, (k2_zf_model(k5_zf_lang(X2,X1))=k2_zf_model(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102])])).
cnf(c_0_140, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_103])).
cnf(c_0_141, plain, (r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_142, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_66])])).
cnf(c_0_143, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_144, plain, (epred1_0|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_106, c_0_107]), c_0_108])).
cnf(c_0_145, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_66])])).
cnf(c_0_146, plain, (epred1_0|v1_zf_lang(esk15_0)|v1_zf_lang(esk16_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_110])).
cnf(c_0_147, plain, (epred1_0|v1_zf_lang(esk16_0)|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_111])).
cnf(c_0_148, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_66])])).
cnf(c_0_149, plain, (epred1_0|v1_zf_lang(esk16_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))), inference(evalgc,[status(thm)],[c_0_113])).
cnf(c_0_150, plain, (v1_zf_lang(esk16_0)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_151, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_114, c_0_115]), c_0_116])).
cnf(c_0_152, plain, (k3_zf_lang1(k5_zf_lang(X1,X2))=k3_zf_lang1(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_63])])).
cnf(c_0_153, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_66])])).
cnf(c_0_154, plain, (epred1_0|v1_zf_lang(esk15_0)|v1_zf_lang(esk17_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_119])).
cnf(c_0_155, plain, (epred1_0|v1_zf_lang(esk17_0)|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_120])).
cnf(c_0_156, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_66])])).
cnf(c_0_157, plain, (epred1_0|v1_zf_lang(esk15_0)|v1_zf_lang(esk17_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_122])).
cnf(c_0_158, plain, (epred1_0|v1_zf_lang(esk17_0)|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_123])).
cnf(c_0_159, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_35])])).
cnf(c_0_160, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_124, c_0_125]), c_0_126])).
cnf(c_0_161, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_66])])).
cnf(c_0_162, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))), inference(evalgc,[status(thm)],[c_0_128])).
cnf(c_0_163, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))), inference(evalgc,[status(thm)],[c_0_129])).
cnf(c_0_164, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_130, c_0_131]), c_0_132])).
cnf(c_0_165, plain, (k3_zf_lang1(k4_zf_lang(X1,X2))=k2_zf_model(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_117])])).
cnf(c_0_166, plain, (m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_167, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_133, c_0_134]), c_0_135])).
cnf(c_0_168, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_66])])).
cnf(c_0_169, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_170, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))), inference(evalgc,[status(thm)],[c_0_138])).
cnf(c_0_171, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_139])])).
cnf(c_0_172, plain, (k2_tarski(X2,X1)=k3_zf_lang1(k5_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_70, c_0_63])])).
cnf(c_0_173, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_92]), c_0_93])])])).
cnf(c_0_174, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))), inference(evalgc,[status(thm)],[c_0_141])).
cnf(c_0_175, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_142, c_0_72]), c_0_73]), c_0_143])).
cnf(c_0_176, plain, (epred1_0|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_92]), c_0_93])])])).
cnf(c_0_177, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_145, c_0_146]), c_0_147])).
cnf(c_0_178, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk16_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_148, c_0_79]), c_0_80]), c_0_149])).
cnf(c_0_179, plain, (v1_zf_lang(esk16_0)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_180, plain, (epred1_0|v1_zf_lang(esk16_0)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_150])).
cnf(c_0_181, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152])])).
cnf(c_0_182, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_153, c_0_154]), c_0_155])).
cnf(c_0_183, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_156, c_0_157]), c_0_158])).
cnf(c_0_184, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_159, c_0_157]), c_0_158])).
cnf(c_0_185, plain, (v1_zf_lang(esk17_0)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_186, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_92]), c_0_93])])])).
cnf(c_0_187, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_161, c_0_162]), c_0_163])).
cnf(c_0_188, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_165]), c_0_93])])])).
cnf(c_0_189, plain, (m2_finseq_1(esk16_0,k5_numbers)|epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_190, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(evalgc,[status(thm)],[c_0_166])).
cnf(c_0_191, plain, (m2_finseq_1(esk16_0,k5_numbers)|v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_192, plain, (m2_finseq_1(esk16_0,k5_numbers)|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_193, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_92]), c_0_93])])])).
cnf(c_0_194, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_168, c_0_169]), c_0_170])).
cnf(c_0_195, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_171, c_0_72]), c_0_169])).
cnf(c_0_196, plain, (k3_zf_lang1(k5_zf_lang(X2,X1))=k2_zf_model(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_172])])).
cnf(c_0_197, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_35])])).
cnf(c_0_198, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_173, c_0_174]), c_0_175])).
cnf(c_0_199, plain, (epred1_0|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_176, c_0_177]), c_0_178])).
cnf(c_0_200, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_35])])).
cnf(c_0_201, plain, (epred1_0|v1_zf_lang(esk16_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))), inference(evalgc,[status(thm)],[c_0_179])).
cnf(c_0_202, plain, (epred1_0|v1_zf_lang(esk16_0)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_92]), c_0_93])])])).
cnf(c_0_203, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_92]), c_0_93])])])).
cnf(c_0_204, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_181, c_0_182]), c_0_183])).
cnf(c_0_205, plain, (k2_zf_model(k5_zf_lang(X1,X2))=k2_zf_model(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_62])])).
cnf(c_0_206, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_66])])).
cnf(c_0_207, plain, (epred1_0|v1_zf_lang(esk17_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))), inference(evalgc,[status(thm)],[c_0_185])).
cnf(c_0_208, plain, (v1_zf_lang(esk17_0)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_209, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_186, c_0_187]), c_0_132])).
cnf(c_0_210, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_35])])).
cnf(c_0_211, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_188, c_0_131]), c_0_132])).
cnf(c_0_212, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_92]), c_0_93])])])).
cnf(c_0_213, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_189])).
cnf(c_0_214, plain, (m2_finseq_1(esk16_0,k5_numbers)|v1_zf_lang(esk15_0)|epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_215, plain, (m2_finseq_1(esk16_0,k5_numbers)|m2_finseq_1(esk15_0,k5_numbers)|epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_216, plain, (m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_217, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_35])])).
cnf(c_0_218, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_191])).
cnf(c_0_219, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_192])).
cnf(c_0_220, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_193, c_0_194]), c_0_175])).
cnf(c_0_221, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_196]), c_0_93])])])).
cnf(c_0_222, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_197, c_0_169]), c_0_170])).
cnf(c_0_223, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_165]), c_0_93])])])).
cnf(c_0_224, plain, (epred1_0|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_165]), c_0_93])])])).
cnf(c_0_225, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk16_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_200, c_0_146]), c_0_147]), c_0_201])).
cnf(c_0_226, plain, (epred1_0|v1_zf_lang(esk15_0)|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_202, c_0_146]), c_0_178])).
cnf(c_0_227, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_203, c_0_201]), c_0_178])).
cnf(c_0_228, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_205])])).
cnf(c_0_229, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk17_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_206, c_0_157]), c_0_158]), c_0_207])).
cnf(c_0_230, plain, (epred1_0|v1_zf_lang(esk17_0)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_208])).
cnf(c_0_231, plain, (v1_zf_lang(esk17_0)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_232, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_165]), c_0_93])])])).
cnf(c_0_233, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_210, c_0_162]), c_0_163]), c_0_211])).
cnf(c_0_234, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_212, c_0_162]), c_0_132])).
cnf(c_0_235, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_66])])).
cnf(c_0_236, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_214])).
cnf(c_0_237, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_215])).
cnf(c_0_238, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(evalgc,[status(thm)],[c_0_216])).
cnf(c_0_239, plain, (m2_finseq_1(esk16_0,k5_numbers)|v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_240, plain, (m2_finseq_1(esk16_0,k5_numbers)|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_241, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_217, c_0_218]), c_0_219])).
cnf(c_0_242, plain, (m2_finseq_1(esk16_0,k5_numbers)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_243, plain, (m2_finseq_1(esk17_0,k5_numbers)|m2_finseq_1(esk15_0,k5_numbers)|epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_244, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_165]), c_0_93])])])).
cnf(c_0_245, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_221, c_0_175]), c_0_169])).
cnf(c_0_246, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_66])])).
cnf(c_0_247, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_223, c_0_174]), c_0_175])).
fof(c_0_248, plain, ![X12, X13, X14, X15]:(~r1_tarski(X12,X13)|~r1_tarski(X14,X15)|r1_tarski(k2_xboole_0(X12,X14),k2_xboole_0(X13,X15))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t13_xboole_1])])).
fof(c_0_249, plain, ![X27, X28]:(~v1_zf_lang(X27)|~m2_finseq_1(X27,k5_numbers)|(~v1_zf_lang(X28)|~m2_finseq_1(X28,k5_numbers)|k3_zf_lang1(k7_zf_lang(X27,X28))=k2_xboole_0(k3_zf_lang1(X27),k3_zf_lang1(X28)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t141_zf_lang1])])])).
fof(c_0_250, plain, ![X66, X67, X68]:(~m1_subset_1(X67,k1_zfmisc_1(X66))|~m1_subset_1(X68,k1_zfmisc_1(X66))|k4_subset_1(X66,X67,X68)=k2_xboole_0(X67,X68)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_subset_1])])).
fof(c_0_251, plain, ![X42, X43]:(~v1_zf_lang(X42)|~m2_finseq_1(X42,k5_numbers)|(~v1_zf_lang(X43)|~m2_finseq_1(X43,k5_numbers)|k2_zf_model(k7_zf_lang(X42,X43))=k4_subset_1(k1_zf_lang,k2_zf_model(X42),k2_zf_model(X43)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t61_zf_lang1])])])).
cnf(c_0_252, plain, (epred1_0|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_224, c_0_225]), c_0_178])).
cnf(c_0_253, plain, (epred1_0|v1_zf_lang(esk15_0)|v1_zf_lang(esk16_0)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_165]), c_0_93])])])).
cnf(c_0_254, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|v1_zf_lang(esk16_0)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_165]), c_0_93])])])).
cnf(c_0_255, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_228, c_0_182]), c_0_229])).
cnf(c_0_256, plain, (epred1_0|v1_zf_lang(esk17_0)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_92]), c_0_93])])])).
cnf(c_0_257, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_92]), c_0_93])])])).
cnf(c_0_258, plain, (epred1_0|v1_zf_lang(esk17_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))), inference(evalgc,[status(thm)],[c_0_231])).
cnf(c_0_259, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_232, c_0_233]), c_0_132])).
cnf(c_0_260, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|v1_zf_lang(esk15_0)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_165]), c_0_93])])])).
cnf(c_0_261, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_235, c_0_236]), c_0_237])).
cnf(c_0_262, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_66])])).
cnf(c_0_263, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_239])).
cnf(c_0_264, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_240])).
cnf(c_0_265, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_66])])).
cnf(c_0_266, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))), inference(evalgc,[status(thm)],[c_0_242])).
cnf(c_0_267, plain, (m2_finseq_1(esk16_0,k5_numbers)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_268, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_243])).
cnf(c_0_269, plain, (m2_finseq_1(esk17_0,k5_numbers)|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_270, plain, (m2_finseq_1(esk17_0,k5_numbers)|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_271, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_244, c_0_194]), c_0_175])).
cnf(c_0_272, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|v1_zf_lang(esk15_0)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_165]), c_0_93])])])).
cnf(c_0_273, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_246, c_0_169]), c_0_170]), c_0_247])).
cnf(c_0_274, plain, (r1_tarski(k2_xboole_0(X1,X3),k2_xboole_0(X2,X4))|~r1_tarski(X1,X2)|~r1_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_248])).
cnf(c_0_275, plain, (k3_zf_lang1(k7_zf_lang(X1,X2))=k2_xboole_0(k3_zf_lang1(X1),k3_zf_lang1(X2))|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~m2_finseq_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_249])).
cnf(c_0_276, plain, (k4_subset_1(X2,X1,X3)=k2_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_250])).
cnf(c_0_277, plain, (k2_zf_model(k7_zf_lang(X1,X2))=k4_subset_1(k1_zf_lang,k2_zf_model(X1),k2_zf_model(X2))|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~m2_finseq_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_251])).
cnf(c_0_278, plain, (epred1_0|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_35])])).
cnf(c_0_279, plain, (epred1_0|v1_zf_lang(esk16_0)|v1_zf_lang(esk15_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_253, c_0_225]), c_0_178])).
cnf(c_0_280, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|v1_zf_lang(esk16_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_254, c_0_225]), c_0_178])).
cnf(c_0_281, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_165]), c_0_93])])])).
cnf(c_0_282, plain, (epred1_0|v1_zf_lang(esk15_0)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_256, c_0_154]), c_0_229])).
cnf(c_0_283, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_257, c_0_258]), c_0_207])).
cnf(c_0_284, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_35])])).
cnf(c_0_285, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|v1_zf_lang(esk15_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_260, c_0_233]), c_0_132])).
cnf(c_0_286, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_92]), c_0_93])])])).
cnf(c_0_287, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_262, c_0_263]), c_0_264])).
cnf(c_0_288, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk16_0,k5_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_265, c_0_218]), c_0_219]), c_0_266])).
cnf(c_0_289, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_267])).
cnf(c_0_290, plain, (m2_finseq_1(esk16_0,k5_numbers)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_291, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_139])])).
cnf(c_0_292, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_finseq_1(esk17_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_269])).
cnf(c_0_293, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_finseq_1(esk17_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_270])).
cnf(c_0_294, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_35])])).
cnf(c_0_295, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|v1_zf_lang(esk15_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_272, c_0_273]), c_0_175])).
cnf(c_0_296, plain, (r1_tarski(k2_xboole_0(X1,X3),k2_xboole_0(X2,X4))|~r1_tarski(X3,X4)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_274])).
cnf(c_0_297, plain, (k3_zf_lang1(k7_zf_lang(X1,X2))=k2_xboole_0(k3_zf_lang1(X1),k3_zf_lang1(X2))|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_275])).
cnf(c_0_298, plain, (k4_subset_1(X2,X1,X3)=k2_xboole_0(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_276])).
cnf(c_0_299, plain, (k2_zf_model(k7_zf_lang(X1,X2))=k4_subset_1(k1_zf_lang,k2_zf_model(X1),k2_zf_model(X2))|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_277])).
cnf(c_0_300, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_92]), c_0_93])])])).
cnf(c_0_301, plain, (epred1_0|v1_zf_lang(esk16_0)|~m2_finseq_1(esk15_0,k5_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_278, c_0_279]), c_0_280])).
cnf(c_0_302, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_139])])).
cnf(c_0_303, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_281, c_0_182]), c_0_229])).
cnf(c_0_304, plain, (epred1_0|v1_zf_lang(esk15_0)|v1_zf_lang(esk17_0)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_282, c_0_165]), c_0_93])])])).
cnf(c_0_305, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|v1_zf_lang(esk17_0)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_165]), c_0_93])])])).
cnf(c_0_306, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_139])])).
cnf(c_0_307, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk15_0,k5_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_284, c_0_285]), c_0_211])).
cnf(c_0_308, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_286, c_0_287]), c_0_288])).
cnf(c_0_309, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_92]), c_0_93])])])).
cnf(c_0_310, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_92]), c_0_93])])])).
cnf(c_0_311, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))), inference(evalgc,[status(thm)],[c_0_290])).
cnf(c_0_312, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_291, c_0_292]), c_0_293])).
cnf(c_0_313, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_139])])).
cnf(c_0_314, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~m2_finseq_1(esk15_0,k5_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_294, c_0_295]), c_0_247])).
cnf(c_0_315, plain, (v1_zf_lang(esk15_0)|epred1_0|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_316, plain, (r1_tarski(k2_xboole_0(X1,X2),k3_zf_lang1(k7_zf_lang(X3,X4)))|~r1_tarski(X2,k3_zf_lang1(X4))|~r1_tarski(X1,k3_zf_lang1(X3))|~m2_finseq_1(X4,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_296, c_0_297])])).
cnf(c_0_317, plain, (k2_xboole_0(k2_zf_model(X1),k2_zf_model(X2))=k2_zf_model(k7_zf_lang(X1,X2))|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_299])])).
cnf(c_0_318, plain, (epred1_0|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_300, c_0_225]), c_0_178]), c_0_301])).
cnf(c_0_319, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_302, c_0_229]), c_0_155])).
cnf(c_0_320, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_303, c_0_35])])).
cnf(c_0_321, plain, (epred1_0|v1_zf_lang(esk17_0)|v1_zf_lang(esk15_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_304, c_0_154]), c_0_229])).
cnf(c_0_322, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|v1_zf_lang(esk17_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_305, c_0_258]), c_0_229])).
cnf(c_0_323, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_306, c_0_132]), c_0_233]), c_0_307])).
cnf(c_0_324, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_308, c_0_165]), c_0_93])])])).
cnf(c_0_325, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_309, c_0_264]), c_0_288])).
cnf(c_0_326, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_310, c_0_311]), c_0_288])).
cnf(c_0_327, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_196]), c_0_93])])])).
cnf(c_0_328, plain, (epred1_0|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_329, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_313, c_0_175]), c_0_273]), c_0_314])).
cnf(c_0_330, plain, (epred1_0|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))), inference(evalgc,[status(thm)],[c_0_315])).
cnf(c_0_331, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k3_zf_lang1(k7_zf_lang(X3,X4)))|~r1_tarski(k2_zf_model(X2),k3_zf_lang1(X4))|~r1_tarski(k2_zf_model(X1),k3_zf_lang1(X3))|~m2_finseq_1(X4,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X3)|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_316, c_0_317])])).
cnf(c_0_332, plain, (epred1_0|v1_zf_lang(esk16_0)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_165]), c_0_93])])])).
cnf(c_0_333, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_319, c_0_196]), c_0_93])])])).
cnf(c_0_334, plain, (epred1_0|v1_zf_lang(esk17_0)|~m2_finseq_1(esk15_0,k5_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_320, c_0_321]), c_0_322])).
cnf(c_0_335, plain, (m2_finseq_1(esk17_0,k5_numbers)|v1_zf_lang(esk15_0)|epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_336, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_196]), c_0_93])])])).
cnf(c_0_337, plain, (v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_338, plain, (v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_339, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_324, c_0_287]), c_0_288])).
cnf(c_0_340, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_finseq_1(esk16_0,k5_numbers)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_165]), c_0_93])])])).
cnf(c_0_341, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_finseq_1(esk16_0,k5_numbers)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_326, c_0_165]), c_0_93])])])).
cnf(c_0_342, plain, (m2_finseq_1(esk15_0,k5_numbers)|epred1_0|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_343, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_327, c_0_292]), c_0_293])).
cnf(c_0_344, plain, (epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))), inference(evalgc,[status(thm)],[c_0_328])).
cnf(c_0_345, plain, (m2_finseq_1(esk17_0,k5_numbers)|epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_346, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_196]), c_0_93])])])).
cnf(c_0_347, plain, (epred1_0|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_331])])).
cnf(c_0_348, plain, (epred1_0|v1_zf_lang(esk16_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_332, c_0_225]), c_0_178])).
cnf(c_0_349, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_333, c_0_229]), c_0_155]), c_0_334])).
cnf(c_0_350, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_335])).
cnf(c_0_351, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_336, c_0_132]), c_0_233])).
cnf(c_0_352, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))), inference(evalgc,[status(thm)],[c_0_337])).
cnf(c_0_353, plain, (m2_finseq_1(esk17_0,k5_numbers)|v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_354, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))), inference(evalgc,[status(thm)],[c_0_338])).
cnf(c_0_355, plain, (m2_finseq_1(esk17_0,k5_numbers)|v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_356, plain, (m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_357, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_339, c_0_35])])).
cnf(c_0_358, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|m2_finseq_1(esk15_0,k5_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_340, c_0_264]), c_0_288])).
cnf(c_0_359, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_finseq_1(esk16_0,k5_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_341, c_0_311]), c_0_288])).
cnf(c_0_360, plain, (m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_361, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))), inference(evalgc,[status(thm)],[c_0_342])).
cnf(c_0_362, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|m2_finseq_1(esk15_0,k5_numbers)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_343, c_0_165]), c_0_93])])])).
cnf(c_0_363, plain, (m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_364, plain, (m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_365, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_331])])).
cnf(c_0_366, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_345])).
cnf(c_0_367, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_346, c_0_175]), c_0_273])).
cnf(c_0_368, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_347, c_0_348]), c_0_349]), c_0_236]), c_0_350]), c_0_351]), c_0_295])).
cnf(c_0_369, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_352, c_0_331])])).
cnf(c_0_370, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_finseq_1(esk17_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_353])).
cnf(c_0_371, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_354, c_0_331])])).
cnf(c_0_372, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_finseq_1(esk17_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_355])).
cnf(c_0_373, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))), inference(evalgc,[status(thm)],[c_0_356])).
cnf(c_0_374, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_357, c_0_358]), c_0_359])).
cnf(c_0_375, plain, (m2_finseq_1(esk17_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_376, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))), inference(evalgc,[status(thm)],[c_0_360])).
cnf(c_0_377, plain, (m2_finseq_1(esk17_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_378, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_331])])).
cnf(c_0_379, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_finseq_1(esk17_0,k5_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_362, c_0_293]), c_0_292])).
cnf(c_0_380, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))), inference(evalgc,[status(thm)],[c_0_363])).
cnf(c_0_381, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))), inference(evalgc,[status(thm)],[c_0_364])).
cnf(c_0_382, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_365, c_0_348]), c_0_349]), c_0_213]), c_0_366]), c_0_351]), c_0_367])).
cnf(c_0_383, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_368, c_0_92]), c_0_93])])])).
cnf(c_0_384, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_369, c_0_348]), c_0_321]), c_0_263]), c_0_370]), c_0_233]), c_0_273])).
cnf(c_0_385, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_371, c_0_348]), c_0_229]), c_0_288]), c_0_372]), c_0_132]), c_0_175])).
cnf(c_0_386, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_331])])).
cnf(c_0_387, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_263])])).
cnf(c_0_388, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(evalgc,[status(thm)],[c_0_375])).
cnf(c_0_389, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_376, c_0_331])])).
cnf(c_0_390, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(evalgc,[status(thm)],[c_0_377])).
cnf(c_0_391, plain, (r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|epred1_0|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_392, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_378, c_0_348]), c_0_349]), c_0_358]), c_0_379]), c_0_351]), c_0_367])).
cnf(c_0_393, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_380, c_0_331])])).
cnf(c_0_394, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_381, c_0_331])])).
cnf(c_0_395, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_66])])).
cnf(c_0_396, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_383, c_0_384]), c_0_385])).
cnf(c_0_397, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_386, c_0_348]), c_0_118]), c_0_387]), c_0_388]), c_0_233]), c_0_273])).
cnf(c_0_398, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_389, c_0_348]), c_0_229]), c_0_288]), c_0_390]), c_0_132]), c_0_175])).
cnf(c_0_399, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))), inference(evalgc,[status(thm)],[c_0_391])).
cnf(c_0_400, plain, (m2_finseq_1(esk17_0,k5_numbers)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_401, plain, (r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_402, plain, (r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_403, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_92]), c_0_93])])])).
cnf(c_0_404, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_393, c_0_348]), c_0_155]), c_0_387]), c_0_379]), c_0_233]), c_0_273])).
cnf(c_0_405, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_394, c_0_348]), c_0_229]), c_0_288]), c_0_379]), c_0_132]), c_0_175])).
cnf(c_0_406, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_395, c_0_396]), c_0_392])).
cnf(c_0_407, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_397, c_0_66])])).
cnf(c_0_408, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_398, c_0_66])])).
cnf(c_0_409, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_399, c_0_331])])).
cnf(c_0_410, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(evalgc,[status(thm)],[c_0_400])).
cnf(c_0_411, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))), inference(evalgc,[status(thm)],[c_0_401])).
cnf(c_0_412, plain, (m2_finseq_1(esk17_0,k5_numbers)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_413, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))), inference(evalgc,[status(thm)],[c_0_402])).
cnf(c_0_414, plain, (m2_finseq_1(esk17_0,k5_numbers)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|epred1_0), inference(split_conjunct,[status(thm)],[c_0_28])).
cnf(c_0_415, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_403, c_0_404]), c_0_405])).
fof(c_0_416, plain, ![X65]:(~v1_zf_lang(X65)|~m1_finseq_1(X65,k5_numbers)|m1_subset_1(k2_zf_model(X65),k1_zfmisc_1(k1_zf_lang))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_zf_model])])).
cnf(c_0_417, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_35])])).
cnf(c_0_418, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_92]), c_0_93])])])).
cnf(c_0_419, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_407, c_0_384]), c_0_404])).
cnf(c_0_420, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_408, c_0_385]), c_0_405])).
cnf(c_0_421, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_409, c_0_348]), c_0_349]), c_0_359]), c_0_410]), c_0_351]), c_0_367])).
cnf(c_0_422, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_411, c_0_331])])).
cnf(c_0_423, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))), inference(evalgc,[status(thm)],[c_0_412])).
cnf(c_0_424, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_413, c_0_331])])).
cnf(c_0_425, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))), inference(evalgc,[status(thm)],[c_0_414])).
cnf(c_0_426, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_415, c_0_165]), c_0_93])])])).
cnf(c_0_427, plain, (m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~v1_zf_lang(X1)|~m1_finseq_1(X1,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_416])).
cnf(c_0_428, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_417, c_0_350]), c_0_268])).
cnf(c_0_429, plain, (k3_zf_lang1(k5_zf_lang(X1,X2))=k2_zf_model(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_63])])).
cnf(c_0_430, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_388, c_0_35])])).
cnf(c_0_431, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_390, c_0_35])])).
fof(c_0_432, plain, ![X16, X17, X18]:(~r1_tarski(X16,X17)|~r1_tarski(X17,X18)|r1_tarski(X16,X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_433, plain, ![X73]:k2_xboole_0(X73,X73)=X73, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k2_xboole_0])])).
cnf(c_0_434, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_418, c_0_419]), c_0_420])).
cnf(c_0_435, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_165]), c_0_93])])])).
cnf(c_0_436, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_421, c_0_92]), c_0_93])])])).
cnf(c_0_437, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_422, c_0_348]), c_0_322]), c_0_387]), c_0_423]), c_0_233]), c_0_273])).
cnf(c_0_438, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_424, c_0_348]), c_0_229]), c_0_288]), c_0_425]), c_0_132]), c_0_175])).
cnf(c_0_439, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_426, c_0_404]), c_0_405])).
cnf(c_0_440, plain, (m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~v1_zf_lang(X1)|~m1_finseq_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_427])).
cnf(c_0_441, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_428, c_0_429])])).
cnf(c_0_442, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_430, c_0_370]), c_0_293])).
cnf(c_0_443, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_431, c_0_372]), c_0_292])).
cnf(c_0_444, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_432])).
cnf(c_0_445, plain, (k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_433])).
fof(c_0_446, plain, ![X10, X11]:r1_tarski(X10,k2_xboole_0(X10,X11)), inference(variable_rename,[status(thm)],[t7_xboole_1])).
fof(c_0_447, plain, ![X51, X52]:k2_xboole_0(X51,X52)=k2_xboole_0(X52,X51), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_448, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_434, c_0_165]), c_0_93])])])).
cnf(c_0_449, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_435, c_0_384]), c_0_385])).
cnf(c_0_450, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_350, c_0_139])])).
cnf(c_0_451, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_436, c_0_437]), c_0_438])).
cnf(c_0_452, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_439, c_0_440])])).
fof(c_0_453, plain, ![X47, X48]:((~m2_finseq_1(X48,X47)|m1_finseq_1(X48,X47))&(~m1_finseq_1(X48,X47)|m2_finseq_1(X48,X47))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_m2_finseq_1])])).
cnf(c_0_454, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_139])])).
cnf(c_0_455, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_441, c_0_442]), c_0_443])).
cnf(c_0_456, plain, (r1_tarski(X1,X3)|~r1_tarski(X2,X3)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_444])).
cnf(c_0_457, plain, (k2_xboole_0(X1,X1)=X1), inference(evalgc,[status(thm)],[c_0_445])).
cnf(c_0_458, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_446])).
cnf(c_0_459, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_447])).
cnf(c_0_460, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_448, c_0_419]), c_0_420])).
cnf(c_0_461, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_449, c_0_440])])).
cnf(c_0_462, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_450, c_0_372]), c_0_370])).
cnf(c_0_463, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_451, c_0_165]), c_0_93])])])).
cnf(c_0_464, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0)|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_452, c_0_440])])).
cnf(c_0_465, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_453])).
cnf(c_0_466, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_454, c_0_288]), c_0_387]), c_0_374])).
cnf(c_0_467, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_455, c_0_165]), c_0_93])])])).
cnf(c_0_468, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,k2_xboole_0(X4,X5))|~r1_tarski(X5,X3)|~r1_tarski(X4,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_456, c_0_296])])).
cnf(c_0_469, plain, (r1_tarski(k2_xboole_0(X1,X2),X3)|~r1_tarski(X2,X3)|~r1_tarski(X1,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_296, c_0_457])])).
cnf(c_0_470, plain, (r1_tarski(X1,k2_xboole_0(X1,X2))), inference(evalgc,[status(thm)],[c_0_458])).
cnf(c_0_471, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(evalgc,[status(thm)],[c_0_459])).
cnf(c_0_472, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_460, c_0_440])])).
cnf(c_0_473, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_finseq_1(esk17_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))), inference(csr,[status(thm)],[c_0_461, c_0_321])).
cnf(c_0_474, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_462, c_0_196]), c_0_93])])])).
cnf(c_0_475, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_463, c_0_437]), c_0_438])).
cnf(c_0_476, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m1_finseq_1(esk17_0,k5_numbers)|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)), inference(csr,[status(thm)],[c_0_464, c_0_348])).
cnf(c_0_477, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_465])).
cnf(c_0_478, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_466, c_0_196]), c_0_93])])])).
cnf(c_0_479, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_467, c_0_442]), c_0_443])).
cnf(c_0_480, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X6,X4)|~r1_tarski(X5,X3)|~r1_tarski(X2,k2_xboole_0(X5,X6))|~r1_tarski(X1,k2_xboole_0(X5,X6))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_468, c_0_469])])).
cnf(c_0_481, plain, (r1_tarski(X1,k2_xboole_0(X2,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_470, c_0_471])])).
fof(c_0_482, plain, ![X19, X20]:r1_tarski(k4_xboole_0(X19,X20),X19), inference(variable_rename,[status(thm)],[t36_xboole_1])).
fof(c_0_483, plain, ![X102]:k2_xboole_0(X102,k1_xboole_0)=X102, inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_484, plain, (epred1_0|~m1_finseq_1(esk17_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[c_0_472, c_0_303])).
cnf(c_0_485, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_finseq_1(esk17_0,k5_numbers)|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_440])])).
cnf(c_0_486, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_474, c_0_372]), c_0_370])).
cnf(c_0_487, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_475, c_0_440])])).
cnf(c_0_488, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~m2_finseq_1(esk17_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_476, c_0_477])])).
cnf(c_0_489, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_478, c_0_288]), c_0_387])).
cnf(c_0_490, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_479, c_0_205]), c_0_93])])])).
cnf(c_0_491, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X1,k2_xboole_0(X5,X2))|~r1_tarski(X2,X4)|~r1_tarski(X5,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_480, c_0_481])])).
cnf(c_0_492, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_456, c_0_481])])).
cnf(c_0_493, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(split_conjunct,[status(thm)],[c_0_482])).
cnf(c_0_494, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_483])).
cnf(c_0_495, plain, (epred1_0|~m1_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_484, c_0_440])])).
cnf(c_0_496, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_finseq_1(esk17_0,k5_numbers)|~m1_finseq_1(esk16_0,k5_numbers)), inference(csr,[status(thm)],[c_0_485, c_0_348])).
cnf(c_0_497, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|v1_zf_lang(esk15_0)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_486, c_0_165]), c_0_93])])])).
cnf(c_0_498, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_finseq_1(esk17_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))), inference(csr,[status(thm)],[c_0_487, c_0_322])).
cnf(c_0_499, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)), inference(csr,[status(thm)],[c_0_488, c_0_379])).
cnf(c_0_500, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_489, c_0_165]), c_0_93])])])).
cnf(c_0_501, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_490, c_0_442]), c_0_443])).
cnf(c_0_502, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X2,X4)|~r1_tarski(X5,X3)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_491, c_0_492])])).
cnf(c_0_503, plain, (r1_tarski(k4_xboole_0(X1,X2),X1)), inference(evalgc,[status(thm)],[c_0_493])).
cnf(c_0_504, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(evalgc,[status(thm)],[c_0_494])).
cnf(c_0_505, plain, (epred1_0|~m1_finseq_1(esk17_0,k5_numbers)|~m1_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))), inference(csr,[status(thm)],[c_0_495, c_0_348])).
cnf(c_0_506, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_finseq_1(esk16_0,k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_496, c_0_477])])).
cnf(c_0_507, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|v1_zf_lang(esk15_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_497, c_0_370]), c_0_372])).
cnf(c_0_508, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_finseq_1(esk17_0,k5_numbers)|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_498, c_0_440])])).
cnf(c_0_509, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk17_0)|~m2_finseq_1(esk16_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_499, c_0_477])])).
cnf(c_0_510, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_500, c_0_387]), c_0_288])).
cnf(c_0_511, plain, (epred1_0|v1_zf_lang(esk17_0)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_349, c_0_165]), c_0_93])])])).
cnf(c_0_512, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_501, c_0_66])])).
cnf(c_0_513, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X2,X4)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_502, c_0_503])])).
cnf(c_0_514, plain, (k2_xboole_0(k1_xboole_0,X1)=X1), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_504, c_0_471])])).
cnf(c_0_515, plain, (epred1_0|~m1_finseq_1(esk17_0,k5_numbers)|~m1_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_505, c_0_35])])).
cnf(c_0_516, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_finseq_1(esk16_0,k5_numbers)), inference(csr,[status(thm)],[c_0_506, c_0_507])).
cnf(c_0_517, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_finseq_1(esk17_0,k5_numbers)|~m1_finseq_1(esk16_0,k5_numbers)), inference(csr,[status(thm)],[c_0_508, c_0_348])).
cnf(c_0_518, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk17_0)), inference(csr,[status(thm)],[c_0_509, c_0_510])).
cnf(c_0_519, plain, (epred1_0|v1_zf_lang(esk17_0)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(csr,[status(thm)],[c_0_511, c_0_229])).
cnf(c_0_520, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_410, c_0_139])])).
cnf(c_0_521, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_512, c_0_507]), c_0_379])).
cnf(c_0_522, plain, (r1_tarski(k2_xboole_0(X1,X2),X3)|~r1_tarski(X2,X3)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_513, c_0_514])])).
fof(c_0_523, plain, ![X21, X22]:(~v1_zf_lang(X21)|~m2_finseq_1(X21,k5_numbers)|(~m2_subset_1(X22,k5_numbers,k1_zf_lang)|k3_zf_lang1(k8_zf_lang(X22,X21))=k2_xboole_0(k3_zf_lang1(X21),k1_tarski(X22)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t142_zf_lang1])])])).
cnf(c_0_524, plain, (epred1_0|~m1_finseq_1(esk17_0,k5_numbers)|~m1_finseq_1(esk16_0,k5_numbers)|~m2_finseq_1(esk15_0,k5_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_515, c_0_516]), c_0_517])).
cnf(c_0_525, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_518, c_0_155])])).
cnf(c_0_526, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_518, c_0_519])])).
cnf(c_0_527, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_finseq_1(esk16_0,k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_517, c_0_477])])).
cnf(c_0_528, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_520, c_0_425]), c_0_423]), c_0_521])).
cnf(c_0_529, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k2_xboole_0(X3,X4))|~r1_tarski(X4,X2)|~r1_tarski(X3,X4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_456, c_0_522])])).
cnf(c_0_530, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_456, c_0_470])])).
fof(c_0_531, axiom, ((epred1_0&![X1]:((v1_zf_lang(X1)&m2_finseq_1(X1,k5_numbers))=>![X2]:(m2_subset_1(X2,k5_numbers,k1_zf_lang)=>(r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))=>r1_tarski(k2_zf_model(k8_zf_lang(X2,X1)),k3_zf_lang1(k8_zf_lang(X2,X1)))))))=>![X1]:((v1_zf_lang(X1)&m2_finseq_1(X1,k5_numbers))=>r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1)))), inference(apply_def,[status(thm)],[s1_zf_lang1__e5_143__zf_lang1, c_0_26])).
cnf(c_0_532, plain, (k3_zf_lang1(k8_zf_lang(X2,X1))=k2_xboole_0(k3_zf_lang1(X1),k1_tarski(X2))|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)), inference(split_conjunct,[status(thm)],[c_0_523])).
fof(c_0_533, negated_conjecture, ~(![X1]:((v1_zf_lang(X1)&m2_finseq_1(X1,k5_numbers))=>r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1)))), inference(assume_negation,[status(cth)],[t151_zf_lang1])).
cnf(c_0_534, plain, (epred1_0|~m1_finseq_1(esk16_0,k5_numbers)|~m2_finseq_1(esk15_0,k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_524, c_0_477])])).
cnf(c_0_535, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)), inference(csr,[status(thm)],[c_0_525, c_0_526])).
cnf(c_0_536, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_527, c_0_477])])).
cnf(c_0_537, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_528, c_0_196]), c_0_93])])])).
cnf(c_0_538, plain, (r1_tarski(X1,X2)|~r1_tarski(X4,X2)|~r1_tarski(X3,X4)|~r1_tarski(X1,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_529, c_0_530])])).
fof(c_0_539, plain, ![X110, X111]:((~m1_subset_1(X110,k1_zfmisc_1(X111))|r1_tarski(X110,X111))&(~r1_tarski(X110,X111)|m1_subset_1(X110,k1_zfmisc_1(X111)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_540, plain, ![X9]:(((v1_zf_lang(esk2_0)|~epred1_0|(~v1_zf_lang(X9)|~m2_finseq_1(X9,k5_numbers)|r1_tarski(k2_zf_model(X9),k3_zf_lang1(X9))))&(m2_finseq_1(esk2_0,k5_numbers)|~epred1_0|(~v1_zf_lang(X9)|~m2_finseq_1(X9,k5_numbers)|r1_tarski(k2_zf_model(X9),k3_zf_lang1(X9)))))&((m2_subset_1(esk3_0,k5_numbers,k1_zf_lang)|~epred1_0|(~v1_zf_lang(X9)|~m2_finseq_1(X9,k5_numbers)|r1_tarski(k2_zf_model(X9),k3_zf_lang1(X9))))&((r1_tarski(k2_zf_model(esk2_0),k3_zf_lang1(esk2_0))|~epred1_0|(~v1_zf_lang(X9)|~m2_finseq_1(X9,k5_numbers)|r1_tarski(k2_zf_model(X9),k3_zf_lang1(X9))))&(~r1_tarski(k2_zf_model(k8_zf_lang(esk3_0,esk2_0)),k3_zf_lang1(k8_zf_lang(esk3_0,esk2_0)))|~epred1_0|(~v1_zf_lang(X9)|~m2_finseq_1(X9,k5_numbers)|r1_tarski(k2_zf_model(X9),k3_zf_lang1(X9))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_531])])])])])).
cnf(c_0_541, plain, (k3_zf_lang1(k8_zf_lang(X2,X1))=k2_xboole_0(k3_zf_lang1(X1),k1_tarski(X2))|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_532])).
fof(c_0_542, negated_conjecture, ((v1_zf_lang(esk1_0)&m2_finseq_1(esk1_0,k5_numbers))&~r1_tarski(k2_zf_model(esk1_0),k3_zf_lang1(esk1_0))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_533])])])).
cnf(c_0_543, plain, (epred1_0|~m1_finseq_1(esk16_0,k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)), inference(csr,[status(thm)],[c_0_534, c_0_535])).
cnf(c_0_544, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk17_0,k5_numbers)), inference(csr,[status(thm)],[c_0_536, c_0_510])).
cnf(c_0_545, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_537, c_0_165]), c_0_93])])])).
cnf(c_0_546, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X4,X2)|~r1_tarski(X1,X4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_538, c_0_470])])).
cnf(c_0_547, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_539])).
cnf(c_0_548, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~r1_tarski(k2_zf_model(k8_zf_lang(esk3_0,esk2_0)),k3_zf_lang1(k8_zf_lang(esk3_0,esk2_0)))|~epred1_0|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_540])).
cnf(c_0_549, plain, (k2_xboole_0(k1_tarski(X2),k3_zf_lang1(X1))=k3_zf_lang1(k8_zf_lang(X2,X1))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_471, c_0_541])])).
cnf(c_0_550, plain, (v1_zf_lang(esk2_0)|r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~epred1_0|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_540])).
cnf(c_0_551, negated_conjecture, (v1_zf_lang(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_542])).
cnf(c_0_552, negated_conjecture, (m2_finseq_1(esk1_0,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_542])).
cnf(c_0_553, negated_conjecture, (~r1_tarski(k2_zf_model(esk1_0),k3_zf_lang1(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_542])).
cnf(c_0_554, plain, (m2_finseq_1(esk2_0,k5_numbers)|r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~epred1_0|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_540])).
cnf(c_0_555, plain, (m2_subset_1(esk3_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~epred1_0|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_540])).
cnf(c_0_556, plain, (epred1_0|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_543, c_0_477])])).
cnf(c_0_557, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_544, c_0_545])])).
cnf(c_0_558, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_544, c_0_425])])).
cnf(c_0_559, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,k4_xboole_0(X2,X4))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_546, c_0_503])])).
fof(c_0_560, plain, ![X95, X96, X97]:(~m1_subset_1(X96,k1_zfmisc_1(X95))|k7_subset_1(X95,X96,X97)=k4_xboole_0(X96,X97)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_subset_1])])).
fof(c_0_561, plain, ![X25, X26]:(~v1_zf_lang(X25)|~m2_finseq_1(X25,k5_numbers)|(~m2_subset_1(X26,k5_numbers,k1_zf_lang)|k2_zf_model(k8_zf_lang(X26,X25))=k7_subset_1(k1_zf_lang,k2_zf_model(X25),k1_tarski(X26)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t62_zf_lang1])])])).
cnf(c_0_562, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_547])).
cnf(c_0_563, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~epred1_0|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)|~r1_tarski(k2_zf_model(k8_zf_lang(esk3_0,esk2_0)),k3_zf_lang1(k8_zf_lang(esk3_0,esk2_0)))), inference(evalgc,[status(thm)],[c_0_548])).
cnf(c_0_564, plain, (r1_tarski(X1,k3_zf_lang1(k8_zf_lang(X2,X3)))|~r1_tarski(X1,k3_zf_lang1(X3))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_492, c_0_549])])).
cnf(c_0_565, plain, (v1_zf_lang(esk2_0)|r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~epred1_0|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_550])).
cnf(c_0_566, negated_conjecture, (v1_zf_lang(esk1_0)), inference(evalgc,[status(thm)],[c_0_551])).
cnf(c_0_567, negated_conjecture, (m2_finseq_1(esk1_0,k5_numbers)), inference(evalgc,[status(thm)],[c_0_552])).
cnf(c_0_568, negated_conjecture, (~r1_tarski(k2_zf_model(esk1_0),k3_zf_lang1(esk1_0))), inference(evalgc,[status(thm)],[c_0_553])).
cnf(c_0_569, plain, (m2_finseq_1(esk2_0,k5_numbers)|r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~epred1_0|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_554])).
cnf(c_0_570, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|m2_subset_1(esk3_0,k5_numbers,k1_zf_lang)|~epred1_0|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_555])).
cnf(c_0_571, plain, (epred1_0|~m2_finseq_1(esk17_0,k5_numbers)), inference(csr,[status(thm)],[c_0_556, c_0_510])).
cnf(c_0_572, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_557, c_0_558])])).
cnf(c_0_573, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_544, c_0_423])])).
cnf(c_0_574, plain, (r1_tarski(k4_xboole_0(X1,X2),k2_xboole_0(X1,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_559, c_0_93])])).
cnf(c_0_575, plain, (k7_subset_1(X2,X1,X3)=k4_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_560])).
cnf(c_0_576, plain, (k2_zf_model(k8_zf_lang(X2,X1))=k7_subset_1(k1_zf_lang,k2_zf_model(X1),k1_tarski(X2))|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)), inference(split_conjunct,[status(thm)],[c_0_561])).
cnf(c_0_577, plain, (r1_tarski(k2_zf_model(esk2_0),k3_zf_lang1(esk2_0))|r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~epred1_0|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_540])).
cnf(c_0_578, plain, (r1_tarski(k2_zf_model(X1),k1_zf_lang)|~m1_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_562, c_0_440])])).
cnf(c_0_579, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~epred1_0|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)|~m2_subset_1(esk3_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k8_zf_lang(esk3_0,esk2_0)),k3_zf_lang1(esk2_0))|~m2_finseq_1(esk2_0,k5_numbers)|~v1_zf_lang(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_563, c_0_564])])).
cnf(c_0_580, negated_conjecture, (v1_zf_lang(esk2_0)|~epred1_0), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_565, c_0_566]), c_0_567])]), c_0_568])])).
cnf(c_0_581, negated_conjecture, (m2_finseq_1(esk2_0,k5_numbers)|~epred1_0), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_569, c_0_566]), c_0_567])]), c_0_568])])).
cnf(c_0_582, negated_conjecture, (m2_subset_1(esk3_0,k5_numbers,k1_zf_lang)|~epred1_0), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_570, c_0_566]), c_0_567])]), c_0_568])])).
cnf(c_0_583, plain, (epred1_0|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_571, c_0_521])])).
cnf(c_0_584, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))), inference(csr,[status(thm)],[c_0_572, c_0_573])).
cnf(c_0_585, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X4,X3)|~r1_tarski(X1,X4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_529, c_0_574])])).
cnf(c_0_586, plain, (k7_subset_1(X2,X1,X3)=k4_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_575])).
cnf(c_0_587, plain, (k2_zf_model(k8_zf_lang(X2,X1))=k7_subset_1(k1_zf_lang,k2_zf_model(X1),k1_tarski(X2))|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[c_0_576])).
cnf(c_0_588, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|r1_tarski(k2_zf_model(esk2_0),k3_zf_lang1(esk2_0))|~epred1_0|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_577])).
cnf(c_0_589, plain, (r1_tarski(k2_zf_model(X1),k1_zf_lang)|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_578, c_0_477])])).
cnf(c_0_590, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~epred1_0|~r1_tarski(k2_zf_model(k8_zf_lang(esk3_0,esk2_0)),k3_zf_lang1(esk2_0))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_579, c_0_580]), c_0_581]), c_0_582])).
cnf(c_0_591, plain, (epred1_0), inference(csr,[status(thm)],[c_0_583, c_0_584])).
cnf(c_0_592, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_585, c_0_93])])).
cnf(c_0_593, plain, (k4_xboole_0(k2_zf_model(X1),k1_tarski(X2))=k2_zf_model(k8_zf_lang(X2,X1))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_586, c_0_587])])).
cnf(c_0_594, negated_conjecture, (r1_tarski(k2_zf_model(esk2_0),k3_zf_lang1(esk2_0))|~epred1_0), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_588, c_0_566]), c_0_567])]), c_0_568])])).
cnf(c_0_595, negated_conjecture, (r1_tarski(k2_zf_model(esk2_0),k1_zf_lang)|~m2_finseq_1(esk2_0,k5_numbers)|~epred1_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_589, c_0_580])])).
cnf(c_0_596, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~r1_tarski(k2_zf_model(k8_zf_lang(esk3_0,esk2_0)),k3_zf_lang1(esk2_0))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_590, c_0_591])])])).
cnf(c_0_597, plain, (r1_tarski(k2_zf_model(k8_zf_lang(X2,X1)),X3)|~r1_tarski(k2_zf_model(X1),X3)|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_592, c_0_593])])).
cnf(c_0_598, negated_conjecture, (m2_subset_1(esk3_0,k5_numbers,k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_582, c_0_591])])])).
cnf(c_0_599, negated_conjecture, (r1_tarski(k2_zf_model(esk2_0),k3_zf_lang1(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_594, c_0_591])])])).
cnf(c_0_600, negated_conjecture, (m2_finseq_1(esk2_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_581, c_0_591])])])).
cnf(c_0_601, negated_conjecture, (v1_zf_lang(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_580, c_0_591])])])).
cnf(c_0_602, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_539])).
cnf(c_0_603, negated_conjecture, (r1_tarski(k2_zf_model(esk2_0),k1_zf_lang)|~epred1_0), inference(csr,[status(thm)],[c_0_595, c_0_581])).
cnf(c_0_604, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)|~m1_subset_1(k2_zf_model(esk2_0),k1_zfmisc_1(k1_zf_lang))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_596, c_0_597]), c_0_598]), c_0_599]), c_0_600]), c_0_601])])])).
cnf(c_0_605, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_602])).
cnf(c_0_606, negated_conjecture, (r1_tarski(k2_zf_model(esk2_0),k1_zf_lang)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_603, c_0_591])])])).
cnf(c_0_607, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_604, c_0_605]), c_0_606])])])).
cnf(c_0_608, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_607, c_0_566]), c_0_567])]), c_0_568]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 609
# Proof object clause steps            : 552
# Proof object formula steps           : 57
# Proof object conjectures             : 21
# Proof object clause conjectures      : 18
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 117
# Proof object initial formulas used   : 26
# Proof object generating inferences   : 182
# Proof object simplifying inferences  : 445
# Training examples: 304 positive, 1024 negative
# Training: Positive examples begin
cnf(c_0_566, negated_conjecture, (v1_zf_lang(esk1_0))).# trainpos
cnf(c_0_567, negated_conjecture, (m2_finseq_1(esk1_0,k5_numbers))).# trainpos
cnf(c_0_504, plain, (k2_xboole_0(X1,k1_xboole_0)=X1)).# trainpos
cnf(c_0_93, plain, (r1_tarski(X1,X1))).# trainpos
cnf(c_0_457, plain, (k2_xboole_0(X1,X1)=X1)).# trainpos
cnf(c_0_471, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1))).# trainpos
cnf(c_0_70, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1))).# trainpos
cnf(c_0_568, negated_conjecture, (~r1_tarski(k2_zf_model(esk1_0),k3_zf_lang1(esk1_0)))).# trainpos
cnf(c_0_470, plain, (r1_tarski(X1,k2_xboole_0(X1,X2)))).# trainpos
cnf(c_0_503, plain, (r1_tarski(k4_xboole_0(X1,X2),X1))).# trainpos
cnf(c_0_477, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_605, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_562, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_66, plain, (k2_zf_model(k6_zf_lang(X1))=k2_zf_model(X1)|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers))).# trainpos
cnf(c_0_35, plain, (k3_zf_lang1(k6_zf_lang(X1))=k3_zf_lang1(X1)|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers))).# trainpos
cnf(c_0_456, plain, (r1_tarski(X1,X3)|~r1_tarski(X2,X3)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_586, plain, (k7_subset_1(X2,X1,X3)=k4_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_440, plain, (m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~v1_zf_lang(X1)|~m1_finseq_1(X1,k5_numbers))).# trainpos
cnf(c_0_146, plain, (epred1_0|v1_zf_lang(esk15_0)|v1_zf_lang(esk16_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_79, plain, (epred1_0|v1_zf_lang(esk15_0)|v1_zf_lang(esk16_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_154, plain, (epred1_0|v1_zf_lang(esk15_0)|v1_zf_lang(esk17_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_157, plain, (epred1_0|v1_zf_lang(esk15_0)|v1_zf_lang(esk17_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_263, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_218, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_370, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_finseq_1(esk17_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_372, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_finseq_1(esk17_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_147, plain, (epred1_0|v1_zf_lang(esk16_0)|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_80, plain, (epred1_0|v1_zf_lang(esk16_0)|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_155, plain, (epred1_0|v1_zf_lang(esk17_0)|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_158, plain, (epred1_0|v1_zf_lang(esk17_0)|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_264, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_219, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_293, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_finseq_1(esk17_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_292, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_finseq_1(esk17_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_162, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0)))).# trainpos
cnf(c_0_169, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0)))).# trainpos
cnf(c_0_46, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0)))).# trainpos
cnf(c_0_72, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0)))).# trainpos
cnf(c_0_201, plain, (epred1_0|v1_zf_lang(esk16_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0)))).# trainpos
cnf(c_0_149, plain, (epred1_0|v1_zf_lang(esk16_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0)))).# trainpos
cnf(c_0_258, plain, (epred1_0|v1_zf_lang(esk17_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0)))).# trainpos
cnf(c_0_207, plain, (epred1_0|v1_zf_lang(esk17_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0)))).# trainpos
cnf(c_0_296, plain, (r1_tarski(k2_xboole_0(X1,X3),k2_xboole_0(X2,X4))|~r1_tarski(X3,X4)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_163, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0)))).# trainpos
cnf(c_0_170, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0)))).# trainpos
cnf(c_0_47, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0)))).# trainpos
cnf(c_0_73, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0)))).# trainpos
cnf(c_0_311, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0)))).# trainpos
cnf(c_0_266, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0)))).# trainpos
cnf(c_0_423, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0)))).# trainpos
cnf(c_0_425, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0)))).# trainpos
cnf(c_0_298, plain, (k4_subset_1(X2,X1,X3)=k2_xboole_0(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_101, plain, (k2_zf_model(k4_zf_lang(X1,X2))=k2_tarski(X1,X2)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_62, plain, (k2_zf_model(k5_zf_lang(X1,X2))=k2_tarski(X1,X2)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_117, plain, (k3_zf_lang1(k4_zf_lang(X1,X2))=k2_tarski(X1,X2)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_63, plain, (k3_zf_lang1(k5_zf_lang(X1,X2))=k2_tarski(X1,X2)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_565, plain, (v1_zf_lang(esk2_0)|r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~epred1_0|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers))).# trainpos
cnf(c_0_109, plain, (epred1_0|v1_zf_lang(esk16_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0))))).# trainpos
cnf(c_0_55, plain, (epred1_0|v1_zf_lang(esk16_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0))))).# trainpos
cnf(c_0_118, plain, (epred1_0|v1_zf_lang(esk17_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0))))).# trainpos
cnf(c_0_121, plain, (epred1_0|v1_zf_lang(esk17_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0))))).# trainpos
cnf(c_0_569, plain, (m2_finseq_1(esk2_0,k5_numbers)|r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~epred1_0|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers))).# trainpos
cnf(c_0_131, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0)))).# trainpos
cnf(c_0_174, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0)))).# trainpos
cnf(c_0_96, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0)))).# trainpos
cnf(c_0_143, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0)))).# trainpos
cnf(c_0_238, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0))))).# trainpos
cnf(c_0_190, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0))))).# trainpos
cnf(c_0_388, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0))))).# trainpos
cnf(c_0_390, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0))))).# trainpos
cnf(c_0_570, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|m2_subset_1(esk3_0,k5_numbers,k1_zf_lang)|~epred1_0|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers))).# trainpos
cnf(c_0_352, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0))))).# trainpos
cnf(c_0_354, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0))))).# trainpos
cnf(c_0_588, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|r1_tarski(k2_zf_model(esk2_0),k3_zf_lang1(esk2_0))|~epred1_0|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers))).# trainpos
cnf(c_0_127, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0))))).# trainpos
cnf(c_0_136, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0))))).# trainpos
cnf(c_0_34, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0))))).# trainpos
cnf(c_0_51, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0))))).# trainpos
cnf(c_0_541, plain, (k3_zf_lang1(k8_zf_lang(X2,X1))=k2_xboole_0(k3_zf_lang1(X1),k1_tarski(X2))|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_380, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0))))).# trainpos
cnf(c_0_381, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0))))).# trainpos
cnf(c_0_587, plain, (k2_zf_model(k8_zf_lang(X2,X1))=k7_subset_1(k1_zf_lang,k2_zf_model(X1),k1_tarski(X2))|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_297, plain, (k3_zf_lang1(k7_zf_lang(X1,X2))=k2_xboole_0(k3_zf_lang1(X1),k3_zf_lang1(X2))|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers))).# trainpos
cnf(c_0_411, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0))))).# trainpos
cnf(c_0_413, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0))))).# trainpos
cnf(c_0_299, plain, (k2_zf_model(k7_zf_lang(X1,X2))=k4_subset_1(k1_zf_lang,k2_zf_model(X1),k2_zf_model(X2))|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers))).# trainpos
cnf(c_0_563, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~epred1_0|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)|~r1_tarski(k2_zf_model(k8_zf_lang(esk3_0,esk2_0)),k3_zf_lang1(k8_zf_lang(esk3_0,esk2_0))))).# trainpos
cnf(c_0_373, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0))))).# trainpos
cnf(c_0_376, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0))))).# trainpos
cnf(c_0_107, plain, (epred1_0|v1_zf_lang(esk15_0)|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_115, plain, (epred1_0|v1_zf_lang(esk15_0)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_236, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_350, plain, (epred1_0|v1_zf_lang(esk15_0)|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_108, plain, (epred1_0|v1_zf_lang(esk16_0)|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_116, plain, (epred1_0|v1_zf_lang(esk17_0)|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_237, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_268, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_125, plain, (epred1_0|v1_zf_lang(esk15_0)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_134, plain, (epred1_0|v1_zf_lang(esk15_0)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_180, plain, (epred1_0|v1_zf_lang(esk16_0)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_230, plain, (epred1_0|v1_zf_lang(esk17_0)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_126, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_135, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_289, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_410, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_74, plain, (epred1_0|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_81, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_91, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_140, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_213, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_366, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_330, plain, (epred1_0|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0))))).# trainpos
cnf(c_0_87, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_97, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0))))).# trainpos
cnf(c_0_361, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0))))).# trainpos
cnf(c_0_399, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0))))).# trainpos
cnf(c_0_344, plain, (epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0))))).# trainpos
cnf(c_0_317, plain, (k2_xboole_0(k2_zf_model(X1),k2_zf_model(X2))=k2_zf_model(k7_zf_lang(X1,X2))|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).# trainpos
cnf(c_0_549, plain, (k2_xboole_0(k1_tarski(X2),k3_zf_lang1(X1))=k3_zf_lang1(k8_zf_lang(X2,X1))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).# trainpos
cnf(c_0_152, plain, (k3_zf_lang1(k5_zf_lang(X1,X2))=k3_zf_lang1(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_165, plain, (k3_zf_lang1(k4_zf_lang(X1,X2))=k2_zf_model(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_92, plain, (k3_zf_lang1(k5_zf_lang(X1,X2))=k2_zf_model(k5_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_429, plain, (k3_zf_lang1(k5_zf_lang(X1,X2))=k2_zf_model(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_205, plain, (k2_zf_model(k5_zf_lang(X1,X2))=k2_zf_model(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_593, plain, (k4_xboole_0(k2_zf_model(X1),k1_tarski(X2))=k2_zf_model(k8_zf_lang(X2,X1))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).# trainpos
cnf(c_0_316, plain, (r1_tarski(k2_xboole_0(X1,X2),k3_zf_lang1(k7_zf_lang(X3,X4)))|~r1_tarski(X2,k3_zf_lang1(X4))|~r1_tarski(X1,k3_zf_lang1(X3))|~m2_finseq_1(X4,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X3))).# trainpos
cnf(c_0_582, negated_conjecture, (m2_subset_1(esk3_0,k5_numbers,k1_zf_lang)|~epred1_0)).# trainpos
cnf(c_0_469, plain, (r1_tarski(k2_xboole_0(X1,X2),X3)|~r1_tarski(X2,X3)|~r1_tarski(X1,X3))).# trainpos
cnf(c_0_172, plain, (k2_tarski(X2,X1)=k3_zf_lang1(k5_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_468, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,k2_xboole_0(X4,X5))|~r1_tarski(X5,X3)|~r1_tarski(X4,X2))).# trainpos
cnf(c_0_481, plain, (r1_tarski(X1,k2_xboole_0(X2,X1)))).# trainpos
cnf(c_0_578, plain, (r1_tarski(k2_zf_model(X1),k1_zf_lang)|~m1_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).# trainpos
cnf(c_0_530, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_102, plain, (k2_tarski(X2,X1)=k2_zf_model(k5_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_145, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_161, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_124, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_580, negated_conjecture, (v1_zf_lang(esk2_0)|~epred1_0)).# trainpos
cnf(c_0_106, plain, (epred1_0|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_581, negated_conjecture, (m2_finseq_1(esk2_0,k5_numbers)|~epred1_0)).# trainpos
cnf(c_0_133, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_168, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_197, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_45, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_78, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_114, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_153, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_594, negated_conjecture, (r1_tarski(k2_zf_model(esk2_0),k3_zf_lang1(esk2_0))|~epred1_0)).# trainpos
cnf(c_0_71, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_417, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_235, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_430, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_159, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_156, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_262, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_514, plain, (k2_xboole_0(k1_xboole_0,X1)=X1)).# trainpos
cnf(c_0_431, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_217, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_492, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,X3))).# trainpos
cnf(c_0_148, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_196, plain, (k3_zf_lang1(k5_zf_lang(X2,X1))=k2_zf_model(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_203, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_202, plain, (epred1_0|v1_zf_lang(esk16_0)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_176, plain, (epred1_0|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_265, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_200, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_309, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_95, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_257, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_130, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_181, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_206, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_256, plain, (epred1_0|v1_zf_lang(esk17_0)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_139, plain, (k2_zf_model(k5_zf_lang(X2,X1))=k2_zf_model(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_253, plain, (epred1_0|v1_zf_lang(esk15_0)|v1_zf_lang(esk16_0)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_304, plain, (epred1_0|v1_zf_lang(esk15_0)|v1_zf_lang(esk17_0)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_305, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|v1_zf_lang(esk17_0)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_254, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|v1_zf_lang(esk16_0)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_188, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_171, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_228, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_224, plain, (epred1_0|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_340, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_finseq_1(esk16_0,k5_numbers)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_589, plain, (r1_tarski(k2_zf_model(X1),k1_zf_lang)|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers))).# trainpos
cnf(c_0_186, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_212, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_310, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_142, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_173, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_291, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_210, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_286, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_480, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X6,X4)|~r1_tarski(X5,X3)|~r1_tarski(X2,k2_xboole_0(X5,X6))|~r1_tarski(X1,k2_xboole_0(X5,X6)))).# trainpos
cnf(c_0_193, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_278, plain, (epred1_0|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_281, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_324, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_221, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_341, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_finseq_1(esk16_0,k5_numbers)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_223, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_260, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|v1_zf_lang(esk15_0)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_491, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X1,k2_xboole_0(X5,X2))|~r1_tarski(X2,X4)|~r1_tarski(X5,X3))).# trainpos
cnf(c_0_244, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_232, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_327, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_246, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_441, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_302, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_320, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_272, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|v1_zf_lang(esk15_0)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_294, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_284, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_357, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_362, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|m2_finseq_1(esk15_0,k5_numbers)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_333, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_467, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_502, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X2,X4)|~r1_tarski(X5,X3)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_450, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_300, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_313, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_331, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k3_zf_lang1(k7_zf_lang(X3,X4)))|~r1_tarski(k2_zf_model(X2),k3_zf_lang1(X4))|~r1_tarski(k2_zf_model(X1),k3_zf_lang1(X3))|~m2_finseq_1(X4,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X3)|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).# trainpos
cnf(c_0_306, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_564, plain, (r1_tarski(X1,k3_zf_lang1(k8_zf_lang(X2,X3)))|~r1_tarski(X1,k3_zf_lang1(X3))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X3))).# trainpos
cnf(c_0_332, plain, (epred1_0|v1_zf_lang(esk16_0)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_511, plain, (epred1_0|v1_zf_lang(esk17_0)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_371, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).# trainpos
cnf(c_0_369, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).# trainpos
cnf(c_0_336, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_347, plain, (epred1_0|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).# trainpos
cnf(c_0_387, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_346, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_474, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_579, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~epred1_0|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)|~m2_subset_1(esk3_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k8_zf_lang(esk3_0,esk2_0)),k3_zf_lang1(esk2_0))|~m2_finseq_1(esk2_0,k5_numbers)|~v1_zf_lang(esk2_0))).# trainpos
cnf(c_0_378, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).# trainpos
cnf(c_0_394, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).# trainpos
cnf(c_0_393, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).# trainpos
cnf(c_0_490, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_409, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).# trainpos
cnf(c_0_424, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).# trainpos
cnf(c_0_365, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).# trainpos
cnf(c_0_422, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).# trainpos
cnf(c_0_389, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).# trainpos
cnf(c_0_386, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).# trainpos
cnf(c_0_513, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X2,X4)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_383, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_497, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|v1_zf_lang(esk15_0)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_512, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_395, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_436, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_408, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_407, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_403, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_520, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_435, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_418, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_426, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_463, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_537, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_545, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_461, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0))).# trainpos
cnf(c_0_448, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_452, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0))).# trainpos
cnf(c_0_487, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0))).# trainpos
cnf(c_0_522, plain, (r1_tarski(k2_xboole_0(X1,X2),X3)|~r1_tarski(X2,X3)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_485, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_finseq_1(esk17_0,k5_numbers)|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0))).# trainpos
cnf(c_0_464, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0)|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0))).# trainpos
cnf(c_0_508, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_finseq_1(esk17_0,k5_numbers)|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0))).# trainpos
cnf(c_0_472, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0))).# trainpos
cnf(c_0_529, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k2_xboole_0(X3,X4))|~r1_tarski(X4,X2)|~r1_tarski(X3,X4))).# trainpos
cnf(c_0_454, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_595, negated_conjecture, (r1_tarski(k2_zf_model(esk2_0),k1_zf_lang)|~m2_finseq_1(esk2_0,k5_numbers)|~epred1_0)).# trainpos
cnf(c_0_506, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_finseq_1(esk16_0,k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers))).# trainpos
cnf(c_0_478, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_527, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_finseq_1(esk16_0,k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers))).# trainpos
cnf(c_0_488, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~m2_finseq_1(esk17_0,k5_numbers))).# trainpos
cnf(c_0_495, plain, (epred1_0|~m1_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0))).# trainpos
cnf(c_0_538, plain, (r1_tarski(X1,X2)|~r1_tarski(X4,X2)|~r1_tarski(X3,X4)|~r1_tarski(X1,X3))).# trainpos
cnf(c_0_500, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_509, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk17_0)|~m2_finseq_1(esk16_0,k5_numbers))).# trainpos
cnf(c_0_536, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers))).# trainpos
cnf(c_0_515, plain, (epred1_0|~m1_finseq_1(esk17_0,k5_numbers)|~m1_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).# trainpos
cnf(c_0_526, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_525, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_557, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_534, plain, (epred1_0|~m1_finseq_1(esk16_0,k5_numbers)|~m2_finseq_1(esk15_0,k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers))).# trainpos
cnf(c_0_558, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_573, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_546, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X4,X2)|~r1_tarski(X1,X4))).# trainpos
cnf(c_0_556, plain, (epred1_0|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers))).# trainpos
cnf(c_0_572, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_583, plain, (epred1_0|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0)))).# trainpos
cnf(c_0_601, negated_conjecture, (v1_zf_lang(esk2_0))).# trainpos
cnf(c_0_600, negated_conjecture, (m2_finseq_1(esk2_0,k5_numbers))).# trainpos
cnf(c_0_598, negated_conjecture, (m2_subset_1(esk3_0,k5_numbers,k1_zf_lang))).# trainpos
cnf(c_0_596, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~r1_tarski(k2_zf_model(k8_zf_lang(esk3_0,esk2_0)),k3_zf_lang1(esk2_0))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).# trainpos
cnf(c_0_599, negated_conjecture, (r1_tarski(k2_zf_model(esk2_0),k3_zf_lang1(esk2_0)))).# trainpos
cnf(c_0_606, negated_conjecture, (r1_tarski(k2_zf_model(esk2_0),k1_zf_lang))).# trainpos
cnf(c_0_559, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,k4_xboole_0(X2,X4)))).# trainpos
cnf(c_0_574, plain, (r1_tarski(k4_xboole_0(X1,X2),k2_xboole_0(X1,X3)))).# trainpos
cnf(c_0_585, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X4,X3)|~r1_tarski(X1,X4))).# trainpos
cnf(c_0_592, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,X3))).# trainpos
cnf(c_0_597, plain, (r1_tarski(k2_zf_model(k8_zf_lang(X2,X1)),X3)|~r1_tarski(k2_zf_model(X1),X3)|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).# trainpos
cnf(c_0_604, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)|~m1_subset_1(k2_zf_model(esk2_0),k1_zfmisc_1(k1_zf_lang)))).# trainpos
cnf(c_0_607, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_16, plain, (k4_ordinal1=k5_numbers)).#trainneg
cnf(i_0_52, plain, (a_0_0_zf_lang=k1_zf_lang)).#trainneg
cnf(i_0_55, plain, (k1_xboole_0=k6_numbers)).#trainneg
cnf(i_0_139, plain, (o_0_0_xboole_0=k1_xboole_0)).#trainneg
cnf(i_0_125, plain, (v1_zf_lang(esk4_0))).#trainneg
cnf(i_0_124, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_184, plain, (v1_xboole_0(esk6_0))).#trainneg
cnf(i_0_33, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_188, plain, (v3_ordinal1(esk7_0))).#trainneg
cnf(i_0_43, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_88, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_105, plain, (v2_xxreal_0(np__3))).#trainneg
cnf(i_0_82, plain, (v2_xxreal_0(np__4))).#trainneg
cnf(i_0_138, plain, (v2_xxreal_0(np__5))).#trainneg
cnf(i_0_128, plain, (v1_finset_1(esk4_0))).#trainneg
cnf(i_0_127, plain, (v1_finseq_1(esk4_0))).#trainneg
cnf(i_0_132, plain, (v1_relat_1(esk4_0))).#trainneg
cnf(i_0_129, plain, (v1_funct_1(esk4_0))).#trainneg
cnf(i_0_230, plain, (v7_ordinal1(esk12_0))).#trainneg
cnf(i_0_126, plain, (v2_finseq_1(esk4_0))).#trainneg
cnf(i_0_183, plain, (v1_xxreal_0(esk6_0))).#trainneg
cnf(i_0_212, plain, (v1_xxreal_0(esk9_0))).#trainneg
cnf(i_0_211, plain, (v3_xxreal_0(esk9_0))).#trainneg
cnf(i_0_94, plain, (~v1_xboole_0(k1_zf_lang))).#trainneg
cnf(i_0_34, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_172, plain, (~v1_xboole_0(np__2))).#trainneg
cnf(i_0_156, plain, (~v1_xboole_0(np__3))).#trainneg
cnf(i_0_185, plain, (~v1_xboole_0(np__4))).#trainneg
cnf(i_0_177, plain, (~v1_xboole_0(np__5))).#trainneg
cnf(i_0_189, plain, (~v1_xboole_0(esk7_0))).#trainneg
cnf(i_0_231, plain, (~v1_xboole_0(esk12_0))).#trainneg
cnf(i_0_214, plain, (v1_xboole_0(esk10_1(X1)))).#trainneg
cnf(i_0_186, plain, (v1_finset_1(k1_tarski(X1)))).#trainneg
cnf(i_0_101, plain, (v1_finseq_1(k5_finseq_1(X1)))).#trainneg
cnf(i_0_153, plain, (v1_relat_1(k5_finseq_1(X1)))).#trainneg
cnf(i_0_152, plain, (v1_funct_1(k5_finseq_1(X1)))).#trainneg
cnf(i_0_165, plain, (k2_xcmplx_0(np__2,np__1)=np__3)).#trainneg
cnf(i_0_112, plain, (k2_xcmplx_0(np__1,np__1)=np__2)).#trainneg
cnf(i_0_171, plain, (k2_xcmplx_0(np__1,np__3)=np__4)).#trainneg
cnf(i_0_196, plain, (k2_xcmplx_0(np__1,np__4)=np__5)).#trainneg
cnf(i_0_151, plain, (k2_xcmplx_0(np__3,np__1)=np__4)).#trainneg
cnf(i_0_169, plain, (k2_xcmplx_0(np__4,np__1)=np__5)).#trainneg
cnf(i_0_113, plain, (k3_xcmplx_0(np__2,np__2)=np__4)).#trainneg
cnf(i_0_187, plain, (k3_xcmplx_0(np__2,np__1)=np__2)).#trainneg
cnf(i_0_229, plain, (k3_xcmplx_0(np__1,np__1)=np__1)).#trainneg
cnf(i_0_41, plain, (m1_subset_1(np__2,k5_numbers))).#trainneg
cnf(i_0_40, plain, (m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_86, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_85, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_103, plain, (m1_subset_1(np__3,k5_numbers))).#trainneg
cnf(i_0_102, plain, (m1_subset_1(np__3,k1_numbers))).#trainneg
cnf(i_0_80, plain, (m1_subset_1(np__4,k5_numbers))).#trainneg
cnf(i_0_79, plain, (m1_subset_1(np__4,k1_numbers))).#trainneg
cnf(i_0_155, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_136, plain, (m1_subset_1(np__5,k5_numbers))).#trainneg
cnf(i_0_135, plain, (m1_subset_1(np__5,k1_numbers))).#trainneg
cnf(i_0_133, plain, (m1_finseq_1(esk4_0,k5_numbers))).#trainneg
cnf(i_0_131, plain, (v4_relat_1(esk4_0,k5_numbers))).#trainneg
cnf(i_0_130, plain, (v5_relat_1(esk4_0,k5_numbers))).#trainneg
cnf(i_0_208, plain, (r1_xxreal_0(np__2,np__2))).#trainneg
cnf(i_0_228, plain, (r1_xxreal_0(np__2,np__3))).#trainneg
cnf(i_0_210, plain, (r1_xxreal_0(np__1,np__2))).#trainneg
cnf(i_0_143, plain, (r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_198, plain, (r1_xxreal_0(np__3,np__3))).#trainneg
cnf(i_0_58, plain, (k4_xboole_0(k1_xboole_0,X1)=k1_xboole_0)).#trainneg
cnf(i_0_93, plain, (k4_xboole_0(X1,k1_xboole_0)=X1)).#trainneg
cnf(i_0_159, plain, (~v1_xboole_0(k1_tarski(X1)))).#trainneg
cnf(i_0_157, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_175, plain, (~r1_xxreal_0(np__2,np__1))).#trainneg
cnf(i_0_163, plain, (~r1_xxreal_0(np__3,np__2))).#trainneg
cnf(i_0_173, plain, (~r1_xxreal_0(np__3,np__1))).#trainneg
cnf(i_0_150, plain, (~r1_xxreal_0(np__4,np__2))).#trainneg
cnf(i_0_147, plain, (~r1_xxreal_0(np__4,np__3))).#trainneg
cnf(i_0_207, plain, (~r1_xxreal_0(np__5,np__2))).#trainneg
cnf(i_0_218, plain, (~r1_xxreal_0(np__5,np__3))).#trainneg
cnf(i_0_209, plain, (~r1_xxreal_0(np__5,np__4))).#trainneg
cnf(i_0_182, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_141, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_46, plain, (m1_subset_1(k1_zf_lang,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_190, plain, (m1_subset_1(esk7_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_154, plain, (m1_subset_1(esk5_1(X1),X1))).#trainneg
cnf(i_0_203, plain, (v1_xboole_0(esk8_2(X1,X2)))).#trainneg
cnf(i_0_170, plain, (v1_finset_1(k2_tarski(X1,X2)))).#trainneg
cnf(i_0_202, plain, (v1_relat_1(esk8_2(X1,X2)))).#trainneg
cnf(i_0_148, plain, (~v1_xboole_0(k2_tarski(X1,X2)))).#trainneg
cnf(i_0_42, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_87, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_104, plain, (m2_subset_1(np__3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_81, plain, (m2_subset_1(np__4,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_137, plain, (m2_subset_1(np__5,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_161, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_215, plain, (m1_subset_1(esk10_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_201, plain, (v4_relat_1(esk8_2(X1,X2),X1))).#trainneg
cnf(i_0_200, plain, (v5_relat_1(esk8_2(X1,X2),X2))).#trainneg
cnf(i_0_116, plain, (k3_xcmplx_0(X1,k6_numbers)=k6_numbers|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_164, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_233, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_168, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_225, plain, (X1=esk11_1(X1)|~r2_hidden(X1,a_0_0_zf_lang))).#trainneg
cnf(i_0_121, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_109, plain, (v1_finseq_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_111, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_122, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_110, plain, (v1_funct_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_146, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_36, plain, (m1_subset_1(k4_xboole_0(X1,X2),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_232, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_100, plain, (k5_finseq_1(X1)=k3_zf_lang(X1)|~m1_subset_1(X1,k1_zf_lang))).#trainneg
cnf(i_0_28, plain, (m2_finseq_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_115, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_227, plain, (v5_relat_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_118, plain, (v1_finset_1(k3_enumset1(X1,X2,X3,X4,X5)))).#trainneg
cnf(i_0_134, plain, (~v1_xboole_0(k3_enumset1(X1,X2,X3,X4,X5)))).#trainneg
cnf(i_0_224, plain, (r1_xxreal_0(np__5,esk11_1(X1))|~r2_hidden(X1,a_0_0_zf_lang))).#trainneg
cnf(i_0_26, plain, (m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m1_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_83, plain, (m2_finseq_1(k3_zf_lang(X1),k5_numbers)|~m1_subset_1(X1,k1_zf_lang))).#trainneg
cnf(i_0_226, plain, (m1_subset_1(esk11_1(X1),k5_numbers)|~r2_hidden(X1,a_0_0_zf_lang))).#trainneg
cnf(i_0_204, plain, (m1_subset_1(esk8_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_213, plain, (v1_finset_1(k10_xtuple_0(X1))|~v1_finset_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_140, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_197, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_119, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_77, plain, (k1_zf_model(X1)=k2_zf_model(X1)|~v1_zf_lang(X1)|~m1_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_32, plain, (v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~m1_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_91, plain, (v1_finset_1(k1_zf_model(X1))|~v1_zf_lang(X1)|~m1_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_108, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_78, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_123, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_191, plain, (v7_ordinal1(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_178, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_120, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_167, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_166, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_192, plain, (v1_finset_1(X1)|~v1_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_158, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_195, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_59, plain, (k12_finseq_1(X1,X2)=k5_finseq_1(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_176, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_106, plain, (k6_zf_lang(X1)=k8_finseq_1(k5_numbers,k12_finseq_1(k5_numbers,np__2),X1)|~m2_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_221, plain, (v4_relat_1(X1,k5_numbers)|~v1_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_205, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_160, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_21, plain, (v1_xboole_0(X1)|m2_finseq_1(k12_finseq_1(X1,X2),X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_149, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_117, plain, (v1_zf_lang(k4_zf_lang(X1,X2))|~m1_subset_1(X2,k1_zf_lang)|~m1_subset_1(X1,k1_zf_lang))).#trainneg
cnf(i_0_114, plain, (v1_zf_lang(k5_zf_lang(X1,X2))|~m1_subset_1(X2,k1_zf_lang)|~m1_subset_1(X1,k1_zf_lang))).#trainneg
cnf(i_0_216, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_217, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_47, plain, (m1_subset_1(k7_subset_1(X2,X1,X3),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_142, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_223, plain, (r2_hidden(X2,a_0_0_zf_lang)|X1!=X2|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(np__5,X1))).#trainneg
cnf(i_0_84, plain, (m2_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m1_subset_1(X2,k1_zf_lang)|~m1_subset_1(X1,k1_zf_lang))).#trainneg
cnf(i_0_95, plain, (m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m1_subset_1(X2,k1_zf_lang)|~m1_subset_1(X1,k1_zf_lang))).#trainneg
cnf(i_0_39, plain, (m2_finseq_1(k7_zf_lang(X1,X2),k5_numbers)|~m1_finseq_1(X2,k5_numbers)|~m1_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_54, plain, (m2_finseq_1(k8_zf_lang(X1,X2),k5_numbers)|~m1_subset_1(X1,k1_zf_lang)|~m1_finseq_1(X2,k5_numbers))).#trainneg
cnf(i_0_144, plain, (m2_subset_1(k1_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_107, plain, (k8_finseq_1(X2,X1,X3)=k7_finseq_1(X1,X3)|~m1_finseq_1(X3,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_90, plain, (k4_subset_1(X2,X1,X1)=X1|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_22, plain, (m2_finseq_1(k8_finseq_1(X2,X1,X3),X2)|~m1_finseq_1(X3,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_53, plain, (v1_zf_lang(k8_zf_lang(X1,X2))|~v1_zf_lang(X2)|~m1_subset_1(X1,k1_zf_lang)|~m1_finseq_1(X2,k5_numbers))).#trainneg
cnf(i_0_51, plain, (k4_subset_1(X2,X1,X3)=k4_subset_1(X2,X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_37, plain, (m1_subset_1(k4_subset_1(X2,X1,X3),k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_48, plain, (k3_zf_lang1(X1)=k4_xboole_0(k10_xtuple_0(X1),k3_enumset1(k6_numbers,np__1,np__2,np__3,np__4))|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_38, plain, (v1_zf_lang(k7_zf_lang(X1,X2))|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~m1_finseq_1(X2,k5_numbers)|~m1_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_206, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_56, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_57, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_97, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_174, plain, (k7_zf_lang(X1,X2)=k8_finseq_1(k5_numbers,k8_finseq_1(k5_numbers,k12_finseq_1(k5_numbers,np__3),X1),X2)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_199, plain, (k3_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(k3_xcmplx_0(X1,X3),k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_162, plain, (k8_zf_lang(X1,X2)=k8_finseq_1(k5_numbers,k8_finseq_1(k5_numbers,k12_finseq_1(k5_numbers,np__4),k3_zf_lang(X1)),X2)|~m2_finseq_1(X2,k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_179, plain, (v1_finseq_1(k7_finseq_1(X1,X2))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_181, plain, (v1_relat_1(k7_finseq_1(X1,X2))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_180, plain, (v1_funct_1(k7_finseq_1(X1,X2))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_35, plain, (k4_zf_lang(X1,X2)=k8_finseq_1(k5_numbers,k8_finseq_1(k5_numbers,k12_finseq_1(k5_numbers,k6_numbers),k3_zf_lang(X1)),k3_zf_lang(X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_50, plain, (k5_zf_lang(X1,X2)=k8_finseq_1(k5_numbers,k8_finseq_1(k5_numbers,k12_finseq_1(k5_numbers,np__1),k3_zf_lang(X1)),k3_zf_lang(X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_808, plain, (r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(k7_zf_lang(X1,X2)))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_815, plain, (k4_subset_1(k1_zf_lang,k2_zf_model(X1),k2_zf_model(X2))=k2_zf_model(k7_zf_lang(X2,X1))|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X1)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_810, plain, (k2_xboole_0(k3_zf_lang1(X2),k3_zf_lang1(X1))=k3_zf_lang1(k7_zf_lang(X1,X2))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_820, plain, (k4_subset_1(k1_zf_lang,k2_zf_model(X1),k2_zf_model(X2))=k2_zf_model(k7_zf_lang(k6_zf_lang(X1),X2))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(k6_zf_lang(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_582, plain, (k3_zf_lang1(k4_zf_lang(X1,X2))=k2_zf_model(k5_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_641, plain, (r1_tarski(k3_zf_lang1(X1),k10_xtuple_0(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_794, plain, (r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(k8_zf_lang(X2,X1)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_814, plain, (k2_xboole_0(k3_zf_lang1(X1),k3_zf_lang1(X2))=k3_zf_lang1(k7_zf_lang(k6_zf_lang(X1),X2))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(k6_zf_lang(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_550, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,X3)|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_807, plain, (k3_zf_lang1(k7_zf_lang(X1,X1))=k3_zf_lang1(X1)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_809, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),k2_xboole_0(X3,X4))|~r1_tarski(k3_zf_lang1(X2),X4)|~r1_tarski(k3_zf_lang1(X1),X3)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_578, plain, (k2_tarski(X2,X1)=k3_zf_lang1(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_450, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k4_xboole_0(X2,X3)))).#trainneg
cnf(i_0_802, plain, (k7_subset_1(k1_zf_lang,k2_zf_model(X1),k1_tarski(X2))=k2_zf_model(k8_zf_lang(X2,k6_zf_lang(X1)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_718, plain, (m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m1_subset_1(X1,k1_zf_lang))).#trainneg
cnf(i_0_569, plain, (k2_tarski(X2,X1)=k2_zf_model(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_799, plain, (k2_xboole_0(k3_zf_lang1(X1),k1_tarski(X2))=k3_zf_lang1(k8_zf_lang(X2,k6_zf_lang(X1)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_786, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_617, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_831, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_742, plain, (m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_813, plain, (k2_xboole_0(k3_zf_lang1(X1),k3_zf_lang1(X2))=k3_zf_lang1(k7_zf_lang(X1,k6_zf_lang(X2)))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X1)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_819, plain, (k4_subset_1(k1_zf_lang,k2_zf_model(X1),k2_zf_model(X2))=k2_zf_model(k7_zf_lang(X1,k6_zf_lang(X2)))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X1)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_823, plain, (epred1_0|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_798, plain, (r1_tarski(k2_xboole_0(X1,X2),k3_zf_lang1(k8_zf_lang(X4,X3)))|~r1_tarski(X2,k1_tarski(X4))|~r1_tarski(X1,k3_zf_lang1(X3))|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_547, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X4,X3))|~r1_tarski(X2,X4)|~r1_tarski(X1,X3))).#trainneg
cnf(i_0_833, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_722, plain, (v1_xboole_0(X1)|m2_subset_1(X2,X3,X1)|~m1_subset_1(X2,X1)|~r1_tarski(X1,X3))).#trainneg
cnf(i_0_818, plain, (m1_subset_1(k2_zf_model(k7_zf_lang(X1,X2)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_466, plain, (m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m1_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_791, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_795, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X2,X1)),k2_xboole_0(X3,X4))|~r1_tarski(k1_tarski(X2),X4)|~r1_tarski(k3_zf_lang1(X1),X3)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_620, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_365, plain, (r1_tarski(k5_numbers,k1_numbers))).#trainneg
cnf(i_0_772, plain, (r1_tarski(k2_zf_model(esk2_0),k3_zf_lang1(esk2_0))|r1_tarski(k2_zf_model(k6_zf_lang(X1)),k3_zf_lang1(k6_zf_lang(X1)))|~epred1_0|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m1_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_636, plain, (r1_tarski(k4_subset_1(X1,X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_825, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_732, plain, (m1_subset_1(X1,k5_numbers)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_364, plain, (r1_tarski(k1_zf_lang,k5_numbers))).#trainneg
cnf(i_0_621, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_830, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_793, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_539, plain, (k4_subset_1(X1,X2,X2)=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))|~r1_tarski(X3,X1))).#trainneg
cnf(i_0_720, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_730, plain, (m1_subset_1(X1,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_717, plain, (v1_xboole_0(k4_xboole_0(X1,X2))|m2_subset_1(X3,X1,k4_xboole_0(X1,X2))|~m1_subset_1(X3,k4_xboole_0(X1,X2)))).#trainneg
cnf(i_0_681, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_406, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_827, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_529, plain, (k4_subset_1(X1,X2,X2)=X2|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_349, plain, (m2_finseq_1(esk4_0,k5_numbers))).#trainneg
cnf(i_0_481, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X3)|~r1_tarski(X3,X2))).#trainneg
cnf(i_0_719, plain, (v1_xboole_0(k2_zf_model(X1))|m2_subset_1(X2,k1_zf_lang,k2_zf_model(X1))|~m1_subset_1(X2,k2_zf_model(X1))|~m1_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_449, plain, (k12_finseq_1(X1,X2)=k12_finseq_1(X3,X2)|v1_xboole_0(X1)|v1_xboole_0(X3)|~m1_subset_1(X2,X1)|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_676, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_801, plain, (m1_subset_1(k2_zf_model(k8_zf_lang(X2,X1)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_401, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_368, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_503, plain, (r1_tarski(k7_subset_1(X1,X2,X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_642, plain, (m1_subset_1(k3_zf_lang1(X1),k1_zfmisc_1(k10_xtuple_0(X1)))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_625, plain, (k4_subset_1(X1,X3,X2)=k2_xboole_0(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_359, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k1_zf_lang))).#trainneg
cnf(i_0_223, plain, (r2_hidden(X1,k1_zf_lang)|~r1_xxreal_0(np__5,X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_519, plain, (m2_subset_1(k1_nat_1(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_683, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_721, plain, (v1_xboole_0(X1)|m2_subset_1(X2,k2_zfmisc_1(k5_numbers,X3),X1)|~m1_subset_1(X2,X1)|~m2_finseq_1(X1,X3))).#trainneg
cnf(i_0_716, plain, (v1_xboole_0(k7_subset_1(X1,X2,X3))|m2_subset_1(X4,X1,k7_subset_1(X1,X2,X3))|~m1_subset_1(X4,k7_subset_1(X1,X2,X3))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_571, plain, (~v1_xboole_0(k2_zf_model(k4_zf_lang(X1,X2)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_747, plain, (m2_finseq_1(k7_zf_lang(X1,X2),k5_numbers)|~m1_finseq_1(X2,k5_numbers)|~m1_finseq_1(k8_finseq_1(k5_numbers,k12_finseq_1(k5_numbers,np__3),X1),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_575, plain, (~v1_xboole_0(k2_zf_model(k5_zf_lang(X1,X2)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_679, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_639, plain, (m1_subset_1(k2_xboole_0(X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_528, plain, (k8_finseq_1(X1,X2,X3)=k8_finseq_1(X4,X2,X3)|~m1_finseq_1(X3,X1)|~m1_finseq_1(X2,X1)|~m1_finseq_1(X3,X4)|~m1_finseq_1(X2,X4))).#trainneg
cnf(i_0_1069, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k2_xboole_0(X3,X4))|~r1_tarski(X4,X2)|~r1_tarski(X3,X2))).#trainneg
cnf(i_0_1134, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,X3)|~r1_tarski(X4,X2))).#trainneg
cnf(i_0_1464, plain, (v1_xboole_0(X1)|m2_subset_1(X2,X1,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_852, plain, (r1_tarski(k1_tarski(X1),k3_zf_lang1(k8_zf_lang(X1,X2)))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_838, plain, (k2_xboole_0(k2_zf_model(X2),k2_zf_model(X1))=k2_zf_model(k7_zf_lang(X1,X2))|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_1072, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),X3)|~r1_tarski(k3_zf_lang1(X2),X3)|~r1_tarski(k3_zf_lang1(X1),X3)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_1105, plain, (k3_zf_lang1(k4_zf_lang(X2,X1))=k3_zf_lang1(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_886, plain, (r1_tarski(k2_xboole_0(X1,X2),k3_zf_lang1(k7_zf_lang(X4,X3)))|~r1_tarski(X2,k3_zf_lang1(X4))|~r1_tarski(X1,k3_zf_lang1(X3))|~m2_finseq_1(X3,k5_numbers)|~m2_finseq_1(X4,k5_numbers)|~v1_zf_lang(X3)|~v1_zf_lang(X4))).#trainneg
cnf(i_0_1562, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k7_subset_1(X2,X3,X4))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1097, plain, (k3_zf_lang1(k5_zf_lang(X2,X1))=k3_zf_lang1(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1104, plain, (k3_zf_lang1(k4_zf_lang(X2,X1))=k2_zf_model(k5_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1103, plain, (k3_zf_lang1(k4_zf_lang(X2,X1))=k2_zf_model(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1062, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X2,X1)),k3_zf_lang1(k7_zf_lang(X3,X4)))|~r1_tarski(k3_zf_lang1(X2),k3_zf_lang1(X4))|~r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(X3))|~m2_finseq_1(X4,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X3)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X1)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_836, plain, (r1_tarski(k2_zf_model(X1),k2_zf_model(k7_zf_lang(X1,X2)))|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_882, plain, (r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(k7_zf_lang(X2,X1)))|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X1)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_1098, plain, (k3_zf_lang1(k5_zf_lang(X2,X1))=k3_zf_lang1(k5_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1140, plain, (r1_tarski(k2_zf_model(X1),k2_zf_model(k7_zf_lang(X2,X1)))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X1)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_1070, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X1,X2)),X3)|~r1_tarski(k3_zf_lang1(X2),X3)|~r1_tarski(k1_tarski(X1),X3)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_1059, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X1,X2)),k3_zf_lang1(k7_zf_lang(X3,X4)))|~r1_tarski(k3_zf_lang1(X2),k3_zf_lang1(X4))|~r1_tarski(k1_tarski(X1),k3_zf_lang1(X3))|~m2_finseq_1(X4,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X3)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_1075, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),X3)|~r1_tarski(k2_zf_model(X2),X3)|~r1_tarski(k2_zf_model(X1),X3)|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_1096, plain, (k3_zf_lang1(k5_zf_lang(X2,X1))=k2_zf_model(k5_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1083, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),k3_zf_lang1(k7_zf_lang(X3,X4)))|~r1_tarski(k3_zf_lang1(X2),k3_zf_lang1(X4))|~r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(X3))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~m2_finseq_1(X4,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_1397, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k4_subset_1(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1272, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),k3_zf_lang1(k8_zf_lang(X3,X4)))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X2),k1_tarski(X3))|~r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(X4))|~m2_finseq_1(X4,k5_numbers)|~v1_zf_lang(X4)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_1060, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X2,X1)),k3_zf_lang1(k7_zf_lang(X3,X4)))|~r1_tarski(k1_tarski(X2),k3_zf_lang1(X4))|~r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(X3))|~m2_finseq_1(X4,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X3)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_1066, plain, (r1_tarski(X1,k3_zf_lang1(k7_zf_lang(X2,X3)))|~r1_tarski(X1,k3_zf_lang1(X3))|~r1_tarski(X1,k3_zf_lang1(X2))|~m2_finseq_1(X3,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X3)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_1353, plain, (m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers))).#trainneg
cnf(i_0_1265, plain, (epred1_0|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1359, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X1,X2)),k3_zf_lang1(k8_zf_lang(X3,X4)))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~r1_tarski(k1_tarski(X1),k3_zf_lang1(X4))|~r1_tarski(k3_zf_lang1(X2),k1_tarski(X3))|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_finseq_1(X4,k5_numbers)|~v1_zf_lang(X4))).#trainneg
cnf(i_0_1010, plain, (k3_zf_lang1(k7_zf_lang(k6_zf_lang(X1),X1))=k3_zf_lang1(X1)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1))).#trainneg
cnf(i_0_1178, plain, (k2_zf_model(k4_zf_lang(X2,X1))=k2_zf_model(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1041, plain, (r1_tarski(k2_zf_model(k8_zf_lang(X2,X1)),k2_zf_model(X1))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_1081, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),k3_zf_lang1(k8_zf_lang(X3,X4)))|~r1_tarski(k3_zf_lang1(X2),k3_zf_lang1(X4))|~r1_tarski(k3_zf_lang1(X1),k1_tarski(X3))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_finseq_1(X4,k5_numbers)|~v1_zf_lang(X4))).#trainneg
cnf(i_0_874, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1124, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,k2_xboole_0(X5,X4))|~r1_tarski(X5,X3)|~r1_tarski(X4,X2))).#trainneg
cnf(i_0_1153, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X5,X3)|~r1_tarski(X4,X2)|~r1_tarski(X1,X4))).#trainneg
cnf(i_0_997, plain, (r1_tarski(k3_zf_lang1(X1),k10_xtuple_0(k6_zf_lang(X1)))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_856, plain, (r1_tarski(k2_xboole_0(X1,X2),k3_zf_lang1(k8_zf_lang(X3,X4)))|~r1_tarski(X2,k3_zf_lang1(X4))|~r1_tarski(X1,k1_tarski(X3))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_finseq_1(X4,k5_numbers)|~v1_zf_lang(X4))).#trainneg
cnf(i_0_1486, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1503, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1501, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1488, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1901, plain, (k2_zf_model(k5_zf_lang(X2,X1))=k2_zf_model(k5_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2228, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),k3_zf_lang1(X1))|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(X1)|~v1_zf_lang(k6_zf_lang(X1)))).#trainneg
cnf(i_0_1742, plain, (m2_subset_1(X1,k5_numbers,k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1740, plain, (m2_subset_1(X1,k1_zf_lang,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2148, plain, (r1_tarski(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~r1_tarski(X1,X3))).#trainneg
cnf(i_0_1736, plain, (m2_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang),k1_zfmisc_1(k1_zf_lang))|~m1_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_1948, plain, (r1_tarski(k2_zf_model(X1),k2_zf_model(k7_zf_lang(X1,X2)))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~r1_tarski(k2_zf_model(X2),k1_zf_lang))).#trainneg
cnf(i_0_1860, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk14_0,esk13_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2464, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1861, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk14_0,esk13_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2198, plain, (r1_tarski(k2_zf_model(X1),k2_zf_model(k7_zf_lang(X1,X2)))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~v1_zf_lang(k6_zf_lang(X2)))).#trainneg
cnf(i_0_2158, plain, (r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(k8_zf_lang(X2,X3)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X1),k1_tarski(X2))|~r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(X3))|~m2_finseq_1(X3,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X3)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_1869, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk14_0,esk13_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2470, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2043, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk14_0,esk13_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1733, plain, (m2_subset_1(k3_zf_lang1(X1),k1_zfmisc_1(k10_xtuple_0(X1)),k1_zfmisc_1(k10_xtuple_0(X1)))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_2044, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk14_0,esk13_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1879, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk14_0,esk13_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2052, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk14_0,esk13_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2478, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1881, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk14_0,esk13_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1951, plain, (r1_tarski(k2_zf_model(X1),k2_zf_model(k7_zf_lang(X1,X2)))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~m1_finseq_1(X2,k5_numbers))).#trainneg
cnf(i_0_2480, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2062, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk14_0,esk13_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1872, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk14_0,esk13_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2064, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk14_0,esk13_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2108, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k5_zf_lang(esk14_0,esk13_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2471, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2109, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k5_zf_lang(esk14_0,esk13_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2202, plain, (epred1_0|v1_zf_lang(esk16_0)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1878, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk14_0,esk13_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2126, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k5_zf_lang(esk14_0,esk13_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2116, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k5_zf_lang(esk14_0,esk13_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2287, plain, (r1_tarski(X1,k2_zf_model(X2))|~r1_tarski(X1,k2_zf_model(k8_zf_lang(X3,X2)))|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_2055, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk14_0,esk13_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2477, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2403, plain, (r1_tarski(X1,k3_zf_lang1(k8_zf_lang(X2,X3)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k3_zf_lang1(X3))|~r1_tarski(X1,k1_tarski(X2))|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_2174, plain, (epred1_0|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_2189, plain, (m2_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1))).#trainneg
cnf(i_0_2128, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k5_zf_lang(esk14_0,esk13_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1880, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk14_0,esk13_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2080, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),X3)|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(X2),X3)|~r1_tarski(k2_zf_model(X1),X3)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~r1_tarski(k2_zf_model(X2),k1_zf_lang))).#trainneg
cnf(i_0_1221, plain, (k3_zf_lang1(k7_zf_lang(X1,k6_zf_lang(X1)))=k3_zf_lang1(X1)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1))).#trainneg
cnf(i_0_1410, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2061, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk14_0,esk13_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1209, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2479, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_877, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1728, plain, (m2_subset_1(k4_subset_1(X1,X2,X3),k1_zfmisc_1(X1),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2221, plain, (r1_tarski(k3_zf_lang1(X1),k10_xtuple_0(k7_zf_lang(k6_zf_lang(X1),X1)))|~m2_finseq_1(k7_zf_lang(k6_zf_lang(X1),X1),k5_numbers)|~v1_zf_lang(k7_zf_lang(k6_zf_lang(X1),X1))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1))).#trainneg
cnf(i_0_1906, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X5,X3)|~r1_tarski(X4,X2)|~r1_tarski(X1,X5))).#trainneg
cnf(i_0_1207, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1210, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_871, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2009, plain, (r1_tarski(k2_zf_model(X1),k2_zf_model(k7_zf_lang(X2,X1)))|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X1)|~v1_zf_lang(X2)|~r1_tarski(k2_zf_model(X1),k1_zf_lang))).#trainneg
cnf(i_0_1487, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_962, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2167, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_1270, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X2,X1)),k3_zf_lang1(k8_zf_lang(X3,X4)))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~r1_tarski(k1_tarski(X2),k1_tarski(X3))|~r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(X4))|~m2_finseq_1(X4,k5_numbers)|~v1_zf_lang(X4)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_2063, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk14_0,esk13_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1217, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2166, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_880, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1215, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_965, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1218, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_995, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1216, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1859, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk14_0,esk13_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2127, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k5_zf_lang(esk14_0,esk13_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1117, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,k3_zf_lang1(k8_zf_lang(X5,X4)))|~r1_tarski(k1_tarski(X5),X3)|~r1_tarski(k3_zf_lang1(X4),X2)|~m2_subset_1(X5,k5_numbers,k1_zf_lang)|~m2_finseq_1(X4,k5_numbers)|~v1_zf_lang(X4))).#trainneg
cnf(i_0_1311, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_843, plain, (r1_tarski(X1,k3_zf_lang1(k7_zf_lang(X2,X3)))|~r1_tarski(X1,k3_zf_lang1(X2))|~m2_finseq_1(X3,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X3)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_2083, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),X3)|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(X2),X3)|~r1_tarski(k2_zf_model(X1),X3)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~m1_finseq_1(X2,k5_numbers))).#trainneg
cnf(i_0_1309, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1312, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_883, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X2,X1)),k2_xboole_0(X3,X4))|~r1_tarski(k3_zf_lang1(X2),X4)|~r1_tarski(k3_zf_lang1(X1),X3)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X1)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_1315, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1016, plain, (k3_zf_lang1(k7_zf_lang(k6_zf_lang(X1),X2))=k3_zf_lang1(k7_zf_lang(X1,X2))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(k6_zf_lang(X1)))).#trainneg
cnf(i_0_2541, negated_conjecture, (m2_subset_1(esk3_0,k5_numbers,k5_numbers)|~epred1_0)).#trainneg
cnf(i_0_3179, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3180, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3182, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3181, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3187, plain, (epred1_0|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3010, negated_conjecture, (r1_tarski(k2_zf_model(esk1_0),k1_zf_lang))).#trainneg
cnf(i_0_3184, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3183, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3029, plain, (r1_tarski(k2_zf_model(X1),k2_zf_model(k7_zf_lang(X2,X1)))|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X1)|~v1_zf_lang(X2)|~r1_tarski(k2_zf_model(X2),k1_zf_lang))).#trainneg
cnf(i_0_2886, plain, (epred1_0|v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_3188, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3186, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3185, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_2848, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),X3)|~r1_tarski(k2_zf_model(X2),k1_zf_lang)|~r1_tarski(k2_zf_model(X2),X3)|~r1_tarski(k2_zf_model(X1),X3)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~r1_tarski(k2_zf_model(X1),k1_zf_lang))).#trainneg
cnf(i_0_3189, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_2576, plain, (r1_tarski(k2_zf_model(X1),k2_zf_model(k7_zf_lang(X1,X2)))|~r1_tarski(k2_zf_model(X2),k1_zf_lang)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~r1_tarski(k2_zf_model(X1),k1_zf_lang))).#trainneg
cnf(i_0_3190, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3208, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),X3)|~m1_finseq_1(X2,k5_numbers)|~r1_tarski(k2_zf_model(X2),X3)|~r1_tarski(k2_zf_model(X1),X3)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~m1_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_2731, plain, (r1_tarski(k2_zf_model(X1),k2_zf_model(k7_zf_lang(X1,X2)))|~m1_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~m1_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_2536, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),k3_zf_lang1(X1))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)|~m1_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_3205, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),X3)|~m1_finseq_1(X2,k5_numbers)|~r1_tarski(k2_zf_model(X2),X3)|~r1_tarski(k2_zf_model(X1),X3)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~r1_tarski(k2_zf_model(X1),k1_zf_lang))).#trainneg
cnf(i_0_2728, plain, (r1_tarski(k2_zf_model(X1),k2_zf_model(k7_zf_lang(X1,X2)))|~m1_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~r1_tarski(k2_zf_model(X1),k1_zf_lang))).#trainneg
cnf(i_0_3038, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2954, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X6,X2)|~r1_tarski(X1,k2_xboole_0(X4,X5))|~r1_tarski(X5,X3)|~r1_tarski(X4,X3))).#trainneg
cnf(i_0_3281, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),X3)|~r1_tarski(k3_zf_lang1(X2),X3)|~r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),X3)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(k6_zf_lang(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_3275, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),k3_zf_lang1(k7_zf_lang(X1,X2)))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(k6_zf_lang(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_2836, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_3283, plain, (k3_zf_lang1(k7_zf_lang(X1,k6_zf_lang(X1)))=k3_zf_lang1(k6_zf_lang(X1))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_3095, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2812, plain, (r1_tarski(X1,k3_zf_lang1(k8_zf_lang(X2,X3)))|~r1_tarski(X1,X4)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~r1_tarski(X4,k3_zf_lang1(X3))|~r1_tarski(X4,k1_tarski(X2))|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_3143, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X1,k2_xboole_0(X5,X6))|~r1_tarski(X6,X4)|~r1_tarski(X5,X3)|~r1_tarski(X2,X5))).#trainneg
cnf(i_0_3274, plain, (k3_zf_lang1(k7_zf_lang(X1,k6_zf_lang(k6_zf_lang(X1))))=k3_zf_lang1(k6_zf_lang(X1))|~m2_finseq_1(k6_zf_lang(k6_zf_lang(X1)),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(k6_zf_lang(k6_zf_lang(X1)))|~v1_zf_lang(k6_zf_lang(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_2810, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~epred1_0|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)|~m2_subset_1(esk3_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k8_zf_lang(esk3_0,esk2_0)),k3_zf_lang1(esk2_0))|~r1_tarski(k2_zf_model(k8_zf_lang(esk3_0,esk2_0)),k1_tarski(esk3_0))|~m2_finseq_1(esk2_0,k5_numbers)|~v1_zf_lang(esk2_0))).#trainneg
cnf(i_0_3144, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X1,k2_xboole_0(X5,X6))|~r1_tarski(X6,X4)|~r1_tarski(X5,X3)|~r1_tarski(X2,X6))).#trainneg
cnf(i_0_1768, plain, (k2_xboole_0(k2_zf_model(k4_zf_lang(X1,X2)),k2_zf_model(X3))=k2_zf_model(k7_zf_lang(X3,k5_zf_lang(X1,X2)))|~m1_subset_1(k2_zf_model(k4_zf_lang(X1,X2)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X3),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~v1_zf_lang(X3)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_3002, plain, (r1_tarski(k2_zf_model(esk15_0),k1_zf_lang)|epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk16_0,k5_numbers)|~m2_finseq_1(esk15_0,k5_numbers))).#trainneg
cnf(i_0_2554, plain, (r1_tarski(X1,k10_xtuple_0(X2))|~r1_tarski(X1,k3_zf_lang1(X2))|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_1744, plain, (r1_tarski(X1,k3_zf_lang1(k8_zf_lang(X2,X3)))|~r1_tarski(X1,k1_tarski(X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_842, plain, (k2_zf_model(k7_zf_lang(k6_zf_lang(X1),X2))=k2_xboole_0(k2_zf_model(X1),k2_zf_model(X2))|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(k6_zf_lang(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_3068, plain, (m2_subset_1(esk14_0,k5_numbers,k5_numbers)|epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0)))).#trainneg
cnf(i_0_2257, plain, (k2_xboole_0(k2_zf_model(X1),k2_zf_model(k4_zf_lang(X3,X2)))=k2_zf_model(k7_zf_lang(X1,k4_zf_lang(X2,X3)))|~m1_subset_1(k2_zf_model(k4_zf_lang(X3,X2)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k4_zf_lang(X2,X3),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k4_zf_lang(X2,X3))|~v1_zf_lang(X1)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2501, plain, (k2_xboole_0(k2_zf_model(X1),k2_zf_model(k5_zf_lang(X3,X2)))=k2_zf_model(k7_zf_lang(X1,k5_zf_lang(X2,X3)))|~m1_subset_1(k2_zf_model(k5_zf_lang(X3,X2)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k5_zf_lang(X2,X3),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k5_zf_lang(X2,X3))|~v1_zf_lang(X1)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2452, plain, (k2_xboole_0(k2_zf_model(X1),k2_zf_model(k4_zf_lang(X3,X2)))=k2_zf_model(k7_zf_lang(X1,k5_zf_lang(X2,X3)))|~m1_subset_1(k2_zf_model(k4_zf_lang(X3,X2)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k5_zf_lang(X2,X3),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k5_zf_lang(X2,X3))|~v1_zf_lang(X1)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1276, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k3_zf_lang1(k8_zf_lang(X3,X4)))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(X2),k1_tarski(X3))|~r1_tarski(k2_zf_model(X1),k3_zf_lang1(X4))|~m2_finseq_1(X4,k5_numbers)|~v1_zf_lang(X4)|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_2294, plain, (r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(k8_zf_lang(X2,X3)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(X3))|~r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),k1_tarski(X2))|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X1)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X3))).#trainneg
cnf(i_0_2239, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),k3_zf_lang1(X3))|~r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(X3))|~r1_tarski(k3_zf_lang1(X2),k3_zf_lang1(k6_zf_lang(X3)))|~m2_finseq_1(X3,k5_numbers)|~m2_finseq_1(k6_zf_lang(X3),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X3)|~v1_zf_lang(k6_zf_lang(X3))|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_2882, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X1,X2)),k3_zf_lang1(X3))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~r1_tarski(k1_tarski(X1),k3_zf_lang1(k6_zf_lang(X3)))|~r1_tarski(k3_zf_lang1(X2),k3_zf_lang1(X3))|~m2_finseq_1(k6_zf_lang(X3),k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X3))|~v1_zf_lang(X3)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_1114, plain, (r1_tarski(k2_xboole_0(X1,X2),X3)|~r1_tarski(X2,k4_xboole_0(X3,X4))|~r1_tarski(X1,k4_xboole_0(X3,X4)))).#trainneg
cnf(i_0_3421, plain, (epred1_0|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m1_finseq_1(esk17_0,k5_numbers)|~m1_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3523, plain, (epred1_0|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~m1_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3358, plain, (epred1_0|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3295, plain, (epred1_0|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_3351, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3345, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3344, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3612, plain, (r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(k6_zf_lang(X1)))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1))).#trainneg
cnf(i_0_3350, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3587, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_3389, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_3364, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_3388, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_3320, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_3359, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3360, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3328, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_3361, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3347, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3353, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3327, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_3572, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X5,X3)|~r1_tarski(X2,X4)|~r1_tarski(X1,X4))).#trainneg
cnf(i_0_3391, plain, (r1_tarski(X1,k3_zf_lang1(esk17_0))|epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(X1,k2_zf_model(esk17_0)))).#trainneg
cnf(i_0_3330, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_3349, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3346, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3438, plain, (r1_tarski(X1,k3_zf_lang1(X2))|~r1_tarski(X1,k3_zf_lang1(k6_zf_lang(X2)))|~m1_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_3355, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3352, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3348, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3329, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_3354, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3785, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X5,X4)|~r1_tarski(X1,X3)|~r1_tarski(X2,X5))).#trainneg
cnf(i_0_3919, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk14_0,esk13_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_3357, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_3918, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k2_zf_model(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_3446, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(k6_zf_lang(X1))),k3_zf_lang1(X1))|~m1_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_3834, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~epred1_0|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)|~m2_subset_1(esk3_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k8_zf_lang(esk3_0,esk2_0)),k1_tarski(esk3_0))|~m2_finseq_1(esk2_0,k5_numbers)|~v1_zf_lang(esk2_0))).#trainneg
cnf(i_0_3304, negated_conjecture, (r1_tarski(X1,k1_zf_lang)|~r1_tarski(X1,k2_zf_model(esk1_0)))).#trainneg
cnf(i_0_3577, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),X3)|~r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),X3)|~r1_tarski(k3_zf_lang1(X2),X3)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~m1_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_3333, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),X3)|~r1_tarski(k2_zf_model(X2),k4_xboole_0(X3,X4))|~r1_tarski(k2_zf_model(X1),k4_xboole_0(X3,X4))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_3794, plain, (k2_xboole_0(k2_zf_model(X3),k2_zf_model(k4_zf_lang(X1,X2)))=k2_zf_model(k7_zf_lang(X3,k5_zf_lang(X1,X2)))|~m1_subset_1(k2_zf_model(k4_zf_lang(X1,X2)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X3),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~v1_zf_lang(X3))).#trainneg
cnf(i_0_3456, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(k4_zf_lang(X1,X2))),k2_zf_model(k5_zf_lang(X1,X2)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_3497, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),X3)|~m1_finseq_1(X2,k5_numbers)|~r1_tarski(k2_zf_model(X2),k4_xboole_0(X3,X4))|~r1_tarski(k2_zf_model(X1),k4_xboole_0(X3,X4))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_3583, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(k6_zf_lang(X1))),k3_zf_lang1(X1))|~m2_finseq_1(k6_zf_lang(k6_zf_lang(X1)),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(k6_zf_lang(k6_zf_lang(X1)))|~v1_zf_lang(X1)|~v1_zf_lang(k6_zf_lang(X1)))).#trainneg
cnf(i_0_3916, plain, (m2_subset_1(esk13_0,k5_numbers,k5_numbers)|epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0)))).#trainneg
cnf(i_0_1584, plain, (k2_zf_model(k7_zf_lang(X1,k6_zf_lang(X2)))=k2_xboole_0(k2_zf_model(X2),k2_zf_model(X1))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X1)|~v1_zf_lang(X2)|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang)))).#trainneg
cnf(i_0_1420, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k2_zf_model(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2538, plain, (m2_subset_1(esk14_0,k5_numbers,k5_numbers)|epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0)))).#trainneg
cnf(i_0_1855, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk14_0,esk13_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1865, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk14_0,esk13_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_3332, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),X3)|~m1_subset_1(X4,k1_zfmisc_1(X3))|~r1_tarski(k2_zf_model(X2),k7_subset_1(X3,X4,X5))|~r1_tarski(k2_zf_model(X1),k7_subset_1(X3,X4,X5))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_1427, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1422, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_2240, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),k3_zf_lang1(X3))|~r1_tarski(k3_zf_lang1(X2),k3_zf_lang1(X3))|~r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(k6_zf_lang(X3)))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~m2_finseq_1(k6_zf_lang(X3),k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~v1_zf_lang(X3)|~v1_zf_lang(k6_zf_lang(X3)))).#trainneg
cnf(i_0_3004, plain, (r1_tarski(k2_zf_model(esk15_0),k1_zf_lang)|epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk15_0,k5_numbers))).#trainneg
cnf(i_0_2673, plain, (m2_subset_1(k2_zf_model(k5_zf_lang(X1,X2)),k1_zfmisc_1(k10_xtuple_0(k5_zf_lang(X1,X2))),k1_zfmisc_1(k10_xtuple_0(k5_zf_lang(X1,X2))))|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_3628, plain, (r1_tarski(X1,k3_zf_lang1(k6_zf_lang(X2)))|~r1_tarski(X1,k3_zf_lang1(X2))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X2))).#trainneg
cnf(i_0_1428, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_3003, plain, (r1_tarski(k2_zf_model(esk15_0),k1_zf_lang)|epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk15_0,k5_numbers))).#trainneg
cnf(i_0_2496, plain, (k2_xboole_0(k2_zf_model(k5_zf_lang(X2,X1)),k2_zf_model(X3))=k2_zf_model(k7_zf_lang(X3,k5_zf_lang(X1,X2)))|~m1_subset_1(k2_zf_model(k5_zf_lang(X2,X1)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X3),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~v1_zf_lang(X3)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1426, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_1864, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk14_0,esk13_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_3447, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(k5_zf_lang(X1,X2))),k2_zf_model(k4_zf_lang(X1,X2)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_3111, plain, (m2_subset_1(esk13_0,k5_numbers,k5_numbers)|epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0)))).#trainneg
cnf(i_0_857, plain, (k2_xboole_0(k1_tarski(X1),k3_zf_lang1(X2))=k3_zf_lang1(k8_zf_lang(X1,k6_zf_lang(X2)))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_4087, plain, (epred1_0|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4094, plain, (epred1_0|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4103, plain, (r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(k6_zf_lang(X1)))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)|~m1_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_4144, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4135, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4183, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(k5_numbers,X3))|~r1_tarski(X2,X3)|~r1_tarski(X1,X3))).#trainneg
cnf(i_0_4654, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4152, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4163, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_4111, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0)))).#trainneg
cnf(i_0_4237, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_zf_model(k8_zf_lang(X1,esk17_0)),k3_zf_lang1(esk17_0))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_4118, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk14_0,esk13_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4120, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4110, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0)))).#trainneg
cnf(i_0_4115, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4114, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4132, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk14_0,esk13_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4459, plain, (r1_tarski(k3_zf_lang1(X1),X2)|~m1_finseq_1(X1,k5_numbers)|~r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),X2)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_4121, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4134, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk14_0,esk13_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4129, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4326, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_4133, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4128, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4175, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X2,X4)|~r1_tarski(X1,X4))).#trainneg
cnf(i_0_4162, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_4159, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_4239, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_4530, plain, (r1_tarski(k3_zf_lang1(k4_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X1,X2)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2))|~m2_finseq_1(k4_zf_lang(X1,X2),k5_numbers))).#trainneg
cnf(i_0_4325, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_4270, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X1,X3)|~r1_tarski(X2,k4_xboole_0(X4,X5)))).#trainneg
cnf(i_0_4822, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X1,X2)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X1,X2))|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers))).#trainneg
cnf(i_0_4169, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_4263, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_4850, plain, (k3_zf_lang1(k8_zf_lang(X1,k6_zf_lang(X2)))=k3_zf_lang1(k8_zf_lang(X1,X2))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~v1_zf_lang(k6_zf_lang(X2)))).#trainneg
cnf(i_0_4691, plain, (r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(X2))|~r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(k6_zf_lang(X2)))|~r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),k3_zf_lang1(X2))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_4520, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~r1_tarski(X1,k3_zf_lang1(k6_zf_lang(k4_zf_lang(X2,X3))))|~m1_finseq_1(k4_zf_lang(X2,X3),k5_numbers)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X2,X3)))).#trainneg
cnf(i_0_4230, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k4_xboole_0(k2_zf_model(esk17_0),X1),k3_zf_lang1(esk17_0)))).#trainneg
cnf(i_0_4490, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X1,X2)),k2_zf_model(k7_zf_lang(X3,k5_zf_lang(X1,X2))))|~m1_subset_1(k2_zf_model(k4_zf_lang(X1,X2)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X3),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~v1_zf_lang(X3))).#trainneg
cnf(i_0_4242, plain, (m2_subset_1(esk14_0,k5_numbers,k5_numbers)|epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0)))).#trainneg
cnf(i_0_4532, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k10_xtuple_0(X3))|~m1_finseq_1(X2,k5_numbers)|~r1_tarski(k2_zf_model(X2),k3_zf_lang1(X3))|~r1_tarski(k2_zf_model(X1),k3_zf_lang1(X3))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_4484, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,k5_zf_lang(X2,X3))),X4)|~r1_tarski(k2_zf_model(k4_zf_lang(X2,X3)),X4)|~r1_tarski(k2_zf_model(X1),X4)|~m1_subset_1(k2_zf_model(k4_zf_lang(X2,X3)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X2,X3),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k5_zf_lang(X2,X3))|~v1_zf_lang(X1))).#trainneg
cnf(i_0_4262, plain, (m2_subset_1(esk13_0,k5_numbers,k5_numbers)|epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0)))).#trainneg
cnf(i_0_4231, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|r1_tarski(k2_xboole_0(X1,X2),k3_zf_lang1(esk17_0))|~r1_tarski(X2,k2_zf_model(esk17_0))|~r1_tarski(X1,k2_zf_model(esk17_0)))).#trainneg
cnf(i_0_1403, plain, (r1_tarski(k2_xboole_0(X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1917, plain, (r1_tarski(X1,k2_zf_model(k7_zf_lang(X3,X2)))|~r1_tarski(X1,k2_zf_model(X3))|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X3),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_3141, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X1,k2_xboole_0(X2,X5))|~r1_tarski(X5,X4)|~r1_tarski(X2,X3))).#trainneg
cnf(i_0_3492, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X2,X4)|~r1_tarski(X5,X3)|~r1_tarski(X1,X5))).#trainneg
cnf(i_0_1116, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,k3_zf_lang1(k8_zf_lang(X4,X5)))|~r1_tarski(k3_zf_lang1(X5),X3)|~r1_tarski(k1_tarski(X4),X2)|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~m2_finseq_1(X5,k5_numbers)|~v1_zf_lang(X5))).#trainneg
cnf(i_0_4879, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_4882, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_4880, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_4883, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_4887, plain, (epred1_0|r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4932, plain, (epred1_0|r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk17_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4884, plain, (epred1_0|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_4885, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_4878, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang))).#trainneg
cnf(i_0_4881, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang))).#trainneg
cnf(i_0_4886, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_finseq_1(esk17_0,k5_numbers)|~m1_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_4874, plain, (epred1_0|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_5155, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X4,X3))|~r1_tarski(X2,X4)|~r1_tarski(X1,X4))).#trainneg
cnf(i_0_5207, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_5086, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_5227, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_5005, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_5018, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_4995, plain, (epred1_0|m2_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_5223, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_5186, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_4952, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_5187, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_5087, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_5007, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4940, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_4999, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_5004, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_5017, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_5082, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),k3_zf_lang1(X1))|~m1_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_5052, plain, (epred1_0|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_5246, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_5179, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_5292, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_5110, plain, (epred1_0|~r1_tarski(k2_zf_model(k7_zf_lang(esk16_0,esk17_0)),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_5198, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X1,X2)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_5206, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_5226, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_5274, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),k3_zf_lang1(k8_zf_lang(X2,X1)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_5003, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang))).#trainneg
cnf(i_0_5205, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang))).#trainneg
cnf(i_0_5016, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang))).#trainneg
cnf(i_0_5225, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang))).#trainneg
cnf(i_0_5237, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X1,X2)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_5575, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X2,X4)|~r1_tarski(X1,k2_xboole_0(X5,X6))|~r1_tarski(X6,X3)|~r1_tarski(X5,X3))).#trainneg
cnf(i_0_5195, plain, (r1_tarski(k3_zf_lang1(k4_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X1,X2)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_4872, plain, (r1_tarski(k2_zf_model(esk16_0),k1_zf_lang)|epred1_0|~m2_finseq_1(esk16_0,k5_numbers))).#trainneg
cnf(i_0_5577, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(k5_numbers,X3))|~r1_tarski(X2,X3)|~r1_tarski(X1,k1_zf_lang))).#trainneg
cnf(i_0_5123, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,k2_xboole_0(X4,X5))|~r1_tarski(X5,X3)|~r1_tarski(X4,X3))).#trainneg
cnf(i_0_5201, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X2,X1)),k2_zf_model(k5_zf_lang(X1,X2)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_5612, negated_conjecture, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(k3_zf_lang1(esk2_0),X3))|~r1_tarski(X2,X3)|~r1_tarski(X1,k2_zf_model(esk2_0))|~epred1_0)).#trainneg
cnf(i_0_5291, plain, (r1_tarski(X1,k3_zf_lang1(k8_zf_lang(X2,X3)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k3_zf_lang1(k6_zf_lang(X3)))|~m2_finseq_1(k6_zf_lang(X3),k5_numbers)|~v1_zf_lang(k6_zf_lang(X3))|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_5328, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~m1_finseq_1(k4_zf_lang(X2,X3),k5_numbers)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k3_zf_lang1(k4_zf_lang(X2,X3)))|~v1_zf_lang(k4_zf_lang(X2,X3))|~m2_finseq_1(k4_zf_lang(X2,X3),k5_numbers))).#trainneg
cnf(i_0_5240, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X2,X1)),k2_zf_model(k4_zf_lang(X1,X2)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_1738, plain, (m2_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1)))|~m2_finseq_1(X2,X1))).#trainneg
cnf(i_0_2145, plain, (r1_tarski(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~r1_tarski(X1,k2_xboole_0(X4,X3)))).#trainneg
cnf(i_0_4255, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(k6_zf_lang(X1))),k3_zf_lang1(X1))|~m1_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)|~m1_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(k6_zf_lang(X1)))).#trainneg
cnf(i_0_3443, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(k7_zf_lang(k6_zf_lang(X1),X1))),k3_zf_lang1(X1))|~m1_finseq_1(k7_zf_lang(k6_zf_lang(X1),X1),k5_numbers)|~v1_zf_lang(k7_zf_lang(k6_zf_lang(X1),X1))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1))).#trainneg
cnf(i_0_4902, plain, (r1_tarski(X1,k3_zf_lang1(k6_zf_lang(X2)))|~r1_tarski(X1,k3_zf_lang1(X2))|~m1_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_1765, plain, (k2_zf_model(k7_zf_lang(k6_zf_lang(X2),X1))=k2_xboole_0(k2_zf_model(X1),k2_zf_model(X2))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~v1_zf_lang(X1)|~v1_zf_lang(k6_zf_lang(X2))|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_890, plain, (k2_xboole_0(k3_zf_lang1(X1),k3_zf_lang1(X2))=k3_zf_lang1(k7_zf_lang(X2,k6_zf_lang(X1)))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X2)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_5137, plain, (r1_tarski(k2_xboole_0(X1,X2),k3_zf_lang1(k8_zf_lang(X3,X4)))|~r1_tarski(X2,k3_zf_lang1(X4))|~r1_tarski(X1,k3_zf_lang1(X4))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_finseq_1(X4,k5_numbers)|~v1_zf_lang(X4))).#trainneg
cnf(i_0_5705, plain, (epred1_0|m2_finseq_1(esk16_0,k5_numbers)|m2_finseq_1(esk17_0,k5_numbers))).#trainneg
cnf(i_0_5819, plain, (epred1_0|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_5961, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X1,X2)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_5671, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_5860, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_5797, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_5858, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X2,X1)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_6295, plain, (r1_tarski(X1,k3_zf_lang1(k6_zf_lang(X2)))|~r1_tarski(X1,k3_zf_lang1(X2))|~v1_zf_lang(X2)|~m2_finseq_1(X2,k5_numbers))).#trainneg
cnf(i_0_5842, plain, (epred1_0|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_5836, plain, (epred1_0|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_6293, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(k6_zf_lang(X1),X1)),k3_zf_lang1(X1))|~m1_finseq_1(k7_zf_lang(k6_zf_lang(X1),X1),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k7_zf_lang(k6_zf_lang(X1),X1))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~m2_finseq_1(k7_zf_lang(k6_zf_lang(X1),X1),k5_numbers))).#trainneg
cnf(i_0_5901, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_6065, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X1,X2)))|~m1_finseq_1(k4_zf_lang(X2,X1),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X2,X1)))).#trainneg
cnf(i_0_5674, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_5958, plain, (r1_tarski(k3_zf_lang1(k4_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_5810, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_5862, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X2,X1)),k2_zf_model(k5_zf_lang(X1,X2)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_6389, plain, (k3_zf_lang1(k7_zf_lang(X2,k6_zf_lang(X1)))=k3_zf_lang1(k7_zf_lang(X1,k6_zf_lang(X2)))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X1)|~v1_zf_lang(X2)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(k6_zf_lang(X1)))).#trainneg
cnf(i_0_6446, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X1,X2)),k3_zf_lang1(k8_zf_lang(X3,X4)))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X2),k3_zf_lang1(X4))|~r1_tarski(k1_tarski(X1),k3_zf_lang1(X4))|~m2_finseq_1(X4,k5_numbers)|~v1_zf_lang(X4)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_6131, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~m1_finseq_1(k4_zf_lang(X2,X3),k5_numbers)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~v1_zf_lang(k4_zf_lang(X2,X3)))).#trainneg
cnf(i_0_5905, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X2,X1)),k2_zf_model(k4_zf_lang(X1,X2)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_6388, plain, (k3_zf_lang1(k7_zf_lang(X2,k6_zf_lang(X1)))=k3_zf_lang1(k7_zf_lang(X2,X1))|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X1)|~v1_zf_lang(X2)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(k6_zf_lang(X1)))).#trainneg
cnf(i_0_5963, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X2,X1)),k2_zf_model(k4_zf_lang(X1,X2)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_5669, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang))).#trainneg
cnf(i_0_5962, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X2,X1)),k2_zf_model(k4_zf_lang(X1,X2)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_6398, plain, (k3_zf_lang1(k7_zf_lang(k6_zf_lang(X2),k6_zf_lang(X1)))=k2_xboole_0(k3_zf_lang1(X1),k3_zf_lang1(X2))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X1)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_6159, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X1,X2)))|~m1_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X2,X1)))).#trainneg
cnf(i_0_5702, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers))).#trainneg
cnf(i_0_6118, plain, (r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(k8_zf_lang(X2,X3)))|~m1_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),k3_zf_lang1(k6_zf_lang(X3)))|~m2_finseq_1(k6_zf_lang(X3),k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(k6_zf_lang(X3))|~v1_zf_lang(X3))).#trainneg
cnf(i_0_5717, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,k2_xboole_0(X4,X5))|~r1_tarski(X5,X2)|~r1_tarski(X4,X2))).#trainneg
cnf(i_0_5896, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~m1_finseq_1(k5_zf_lang(X2,X3),k5_numbers)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X2,X3)))).#trainneg
cnf(i_0_6331, plain, (k2_xboole_0(k2_zf_model(k6_zf_lang(X2)),k2_zf_model(X1))=k2_xboole_0(k2_zf_model(X2),k2_zf_model(k6_zf_lang(X1)))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X2)|~m1_subset_1(k2_zf_model(k6_zf_lang(X2)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_5964, plain, (r1_tarski(k3_zf_lang1(k4_zf_lang(X2,X1)),k2_zf_model(k4_zf_lang(X1,X2)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_6387, plain, (k3_zf_lang1(k7_zf_lang(X2,k6_zf_lang(X1)))=k3_zf_lang1(k7_zf_lang(X1,X2))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(k6_zf_lang(X1)))).#trainneg
cnf(i_0_5807, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_5953, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~r1_tarski(X1,k3_zf_lang1(k4_zf_lang(X2,X3)))|~m1_finseq_1(k4_zf_lang(X2,X3),k5_numbers)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X2,X3)))).#trainneg
cnf(i_0_6056, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X5,X3)|~r1_tarski(X4,X3)|~r1_tarski(X1,X5))).#trainneg
cnf(i_0_6467, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang))).#trainneg
cnf(i_0_6469, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_6479, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_6500, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_7154, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_7152, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang))).#trainneg
cnf(i_0_6584, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_6511, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_6481, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang))).#trainneg
cnf(i_0_6821, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang)|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_6477, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang))).#trainneg
cnf(i_0_6483, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_6905, plain, (epred1_0|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk16_0,k5_numbers))).#trainneg
cnf(i_0_6557, plain, (epred1_0|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_6556, plain, (epred1_0|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_7153, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_6807, plain, (k3_zf_lang1(k7_zf_lang(X1,X1))=k3_zf_lang1(k6_zf_lang(X1))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1))).#trainneg
cnf(i_0_6738, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~v1_zf_lang(k4_zf_lang(X2,X3))|~m2_finseq_1(k4_zf_lang(X2,X3),k5_numbers))).#trainneg
cnf(i_0_6794, plain, (k3_zf_lang1(k7_zf_lang(k6_zf_lang(X1),X1))=k3_zf_lang1(k6_zf_lang(X1))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_7132, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),k3_zf_lang1(k7_zf_lang(X2,X1)))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X1)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_6795, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),k3_zf_lang1(k7_zf_lang(X2,X1)))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X2)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_6490, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_6792, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),X3)|~r1_tarski(k3_zf_lang1(k6_zf_lang(X2)),X3)|~r1_tarski(k3_zf_lang1(X1),X3)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X1)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_6569, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_6597, plain, (r1_tarski(k3_zf_lang1(k4_zf_lang(X2,X1)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_6559, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(k6_zf_lang(X1),X1)),k3_zf_lang1(X1))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k7_zf_lang(k6_zf_lang(X1),X1))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~m2_finseq_1(k7_zf_lang(k6_zf_lang(X1),X1),k5_numbers))).#trainneg
cnf(i_0_7063, plain, (k2_xboole_0(k2_zf_model(X1),k2_zf_model(k6_zf_lang(X2)))=k2_zf_model(k7_zf_lang(X2,k6_zf_lang(X1)))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X2)|~m1_subset_1(k2_zf_model(k6_zf_lang(X2)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X1))).#trainneg
cnf(i_0_7131, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X2,X1)),X3)|~m1_finseq_1(X1,k5_numbers)|~r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),X3)|~r1_tarski(k3_zf_lang1(k6_zf_lang(X2)),X3)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_6525, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X2,X1)),k2_zf_model(k5_zf_lang(X2,X1)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_6586, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~r1_tarski(X1,k3_zf_lang1(k4_zf_lang(X3,X2)))|~m1_finseq_1(k4_zf_lang(X3,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X3,X2)))).#trainneg
cnf(i_0_6502, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~r1_tarski(X1,k2_zf_model(k4_zf_lang(X3,X2)))|~m1_finseq_1(k4_zf_lang(X3,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X3,X2)))).#trainneg
cnf(i_0_7000, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_finseq_1(k5_zf_lang(X3,X4),k5_numbers)|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X3,X4))|~m1_subset_1(X2,k1_zfmisc_1(k2_zf_model(k5_zf_lang(X3,X4))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zf_model(k5_zf_lang(X3,X4)))))).#trainneg
cnf(i_0_6516, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~r1_tarski(X1,k2_zf_model(k4_zf_lang(X3,X2)))|~m1_finseq_1(k4_zf_lang(X3,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X3,X2)))).#trainneg
cnf(i_0_7139, plain, (r1_tarski(k2_xboole_0(X1,X2),k3_zf_lang1(k7_zf_lang(X4,X3)))|~r1_tarski(X2,k3_zf_lang1(X3))|~r1_tarski(X1,k3_zf_lang1(k6_zf_lang(X4)))|~m2_finseq_1(k6_zf_lang(X4),k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(k6_zf_lang(X4))|~v1_zf_lang(X3)|~m2_finseq_1(X4,k5_numbers)|~v1_zf_lang(X4))).#trainneg
cnf(i_0_3287, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X1,X2)),k3_zf_lang1(k7_zf_lang(X3,X4)))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X2),k3_zf_lang1(X4))|~r1_tarski(k1_tarski(X1),k3_zf_lang1(k6_zf_lang(X3)))|~m2_finseq_1(X4,k5_numbers)|~m2_finseq_1(k6_zf_lang(X3),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(k6_zf_lang(X3))|~v1_zf_lang(X2)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_2868, plain, (k3_zf_lang1(k7_zf_lang(k7_zf_lang(X2,k6_zf_lang(X2)),X1))=k2_xboole_0(k3_zf_lang1(X1),k3_zf_lang1(X2))|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(k7_zf_lang(X2,k6_zf_lang(X2)),k5_numbers)|~v1_zf_lang(X1)|~v1_zf_lang(k7_zf_lang(X2,k6_zf_lang(X2)))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X2))).#trainneg
cnf(i_0_7138, plain, (r1_tarski(k2_xboole_0(X1,X2),k3_zf_lang1(k7_zf_lang(X4,X3)))|~r1_tarski(X2,k3_zf_lang1(k6_zf_lang(X4)))|~r1_tarski(X1,k3_zf_lang1(X3))|~m2_finseq_1(k6_zf_lang(X4),k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(k6_zf_lang(X4))|~v1_zf_lang(X3)|~m2_finseq_1(X4,k5_numbers)|~v1_zf_lang(X4))).#trainneg
cnf(i_0_2881, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X1,X2)),k3_zf_lang1(X3))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X2),k3_zf_lang1(k6_zf_lang(X3)))|~r1_tarski(k1_tarski(X1),k3_zf_lang1(X3))|~m2_finseq_1(k6_zf_lang(X3),k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X3))|~v1_zf_lang(X3)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_5849, plain, (epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_5020, plain, (r1_tarski(X1,k3_zf_lang1(esk17_0))|epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(X1,k2_zf_model(k8_zf_lang(X2,esk17_0)))|~m1_subset_1(k2_zf_model(esk17_0),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(esk17_0,k5_numbers))).#trainneg
cnf(i_0_4462, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k10_xtuple_0(X3))|~r1_tarski(k2_zf_model(X2),k3_zf_lang1(X3))|~r1_tarski(k2_zf_model(X1),k3_zf_lang1(X3))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_5280, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X1,X2)),X3)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~r1_tarski(k3_zf_lang1(k6_zf_lang(X2)),X3)|~r1_tarski(k1_tarski(X1),X3)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_3791, plain, (k2_zf_model(k7_zf_lang(k4_zf_lang(X1,X2),k5_zf_lang(X1,X2)))=k2_zf_model(k4_zf_lang(X1,X2))|~m1_subset_1(k2_zf_model(k4_zf_lang(X1,X2)),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_6561, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~r1_tarski(X1,k2_zf_model(k5_zf_lang(X3,X2)))|~m1_finseq_1(k5_zf_lang(X3,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X3,X2)))).#trainneg
cnf(i_0_1074, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(k6_zf_lang(X1),X2)),X3)|~r1_tarski(k3_zf_lang1(X2),X3)|~r1_tarski(k3_zf_lang1(X1),X3)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_7290, plain, (epred1_0|v1_zf_lang(esk15_0)|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang))).#trainneg
cnf(i_0_7291, plain, (epred1_0|v1_zf_lang(esk15_0)|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_7305, plain, (epred1_0|m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|v1_zf_lang(esk15_0)|~m1_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_7292, plain, (epred1_0|v1_zf_lang(esk15_0)|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_7286, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_7288, plain, (epred1_0|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_7287, plain, (epred1_0|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang))).#trainneg
cnf(i_0_7289, plain, (epred1_0|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_7284, plain, (epred1_0|v1_zf_lang(esk15_0)|~m1_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang))).#trainneg
cnf(i_0_7298, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~v1_zf_lang(esk17_0)|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_7297, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~v1_zf_lang(esk17_0)|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_7303, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0)|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_7314, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_7296, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang))).#trainneg
cnf(i_0_7324, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m1_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_7315, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_7300, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_7283, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_7304, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang)|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_7313, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang))).#trainneg
cnf(i_0_7311, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_7318, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_7316, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang))).#trainneg
cnf(i_0_7317, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_7301, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0)|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang))).#trainneg
cnf(i_0_7389, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_7310, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_7293, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang))).#trainneg
cnf(i_0_7295, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_7309, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang))).#trainneg
cnf(i_0_7322, plain, (epred1_0|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk16_0,k5_numbers))).#trainneg
cnf(i_0_7308, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_7387, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m1_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang))).#trainneg
cnf(i_0_7306, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang))).#trainneg
cnf(i_0_7307, plain, (epred1_0|~m1_subset_1(k2_zf_model(esk16_0),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_8112, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),X2)|~r1_tarski(k3_zf_lang1(X1),X2)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1))).#trainneg
cnf(i_0_7480, plain, (r1_tarski(k2_zf_model(X1),k2_zf_model(k7_zf_lang(X2,k6_zf_lang(X1))))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(k6_zf_lang(X2)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_7326, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~m2_finseq_1(k4_zf_lang(X2,X3),k5_numbers)|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X2,k1_zf_lang))).#trainneg
cnf(i_0_7642, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k2_zf_model(k4_zf_lang(X3,X2)))|~v1_zf_lang(k4_zf_lang(X3,X2))|~m2_finseq_1(k4_zf_lang(X3,X2),k5_numbers))).#trainneg
cnf(i_0_8035, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X1,X2)))|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~v1_zf_lang(k4_zf_lang(X1,X2))|~m1_subset_1(k2_zf_model(k4_zf_lang(X1,X2)),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_7586, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_7869, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X6,X4)|~r1_tarski(X5,X3)|~r1_tarski(X2,k2_xboole_0(X5,X6))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_7482, plain, (k2_xboole_0(k2_zf_model(k6_zf_lang(X2)),k2_zf_model(X1))=k2_zf_model(k7_zf_lang(X2,k6_zf_lang(X1)))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(k6_zf_lang(X2)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_7896, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X6,X2)|~r1_tarski(X1,k2_xboole_0(X4,X5))|~r1_tarski(X5,X3)|~r1_tarski(X4,X5))).#trainneg
cnf(i_0_7491, negated_conjecture, (r1_tarski(k2_zf_model(k7_zf_lang(X2,k6_zf_lang(X1))),k2_xboole_0(k3_zf_lang1(esk2_0),X3))|~epred1_0|~r1_tarski(k2_zf_model(X1),k2_zf_model(esk2_0))|~r1_tarski(k2_zf_model(k6_zf_lang(X2)),X3)|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(k6_zf_lang(X2)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_7476, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(k6_zf_lang(X1),X1)),k3_zf_lang1(X1))|~m2_finseq_1(k7_zf_lang(k6_zf_lang(X1),X1),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~m1_finseq_1(X1,k5_numbers)|~m1_finseq_1(k6_zf_lang(X1),k5_numbers))).#trainneg
cnf(i_0_7554, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~m1_finseq_1(k4_zf_lang(X3,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~v1_zf_lang(k4_zf_lang(X3,X2)))).#trainneg
cnf(i_0_7581, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_finseq_1(k4_zf_lang(X4,X3),k5_numbers)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X4,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zf_model(k4_zf_lang(X4,X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zf_model(k4_zf_lang(X4,X3)))))).#trainneg
cnf(i_0_7918, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X2,X1)),X3)|~r1_tarski(k2_zf_model(X2),X3)|~r1_tarski(k2_zf_model(X1),k2_zf_model(X2))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X1)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_7437, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~r1_tarski(X1,k2_zf_model(k5_zf_lang(X3,X2)))|~m1_finseq_1(k4_zf_lang(X3,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X3,X2)))).#trainneg
cnf(i_0_7174, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(k4_zf_lang(X1,X2))),k2_zf_model(k4_zf_lang(X1,X2)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_6014, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,k3_zf_lang1(k8_zf_lang(X5,X4)))|~r1_tarski(k1_tarski(X5),X3)|~r1_tarski(k3_zf_lang1(X4),X3)|~m2_subset_1(X5,k5_numbers,k1_zf_lang)|~m2_finseq_1(X4,k5_numbers)|~v1_zf_lang(X4))).#trainneg
cnf(i_0_3451, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(k5_zf_lang(X1,X2))),k2_zf_model(k5_zf_lang(X2,X1)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_6058, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~r1_tarski(X1,k2_zf_model(k5_zf_lang(X3,X2)))|~m1_finseq_1(k4_zf_lang(X2,X3),k5_numbers)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X2,X3)))).#trainneg
cnf(i_0_3231, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),X3)|~r1_tarski(X5,X3)|~r1_tarski(X4,X3)|~r1_tarski(k3_zf_lang1(X1),X5)|~r1_tarski(k3_zf_lang1(X2),X4)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_2796, plain, (r1_tarski(k2_zf_model(k8_zf_lang(X1,k8_zf_lang(X2,X3))),k2_zf_model(X3))|~m1_subset_1(k2_zf_model(X3),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X3)|~m1_subset_1(k2_zf_model(k8_zf_lang(X2,X3)),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k8_zf_lang(X2,X3),k5_numbers)|~v1_zf_lang(k8_zf_lang(X2,X3)))).#trainneg
cnf(i_0_3458, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(k4_zf_lang(X1,X2))),k2_zf_model(k5_zf_lang(X2,X1)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k4_zf_lang(X1,X2))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_6135, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~m1_finseq_1(k4_zf_lang(X2,X3),k5_numbers)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k3_zf_lang1(k4_zf_lang(X3,X2)))|~v1_zf_lang(k4_zf_lang(X2,X3)))).#trainneg
cnf(i_0_8125, plain, (epred1_0|v1_zf_lang(esk15_0)|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang)|~m2_finseq_1(esk17_0,k5_numbers))).#trainneg
cnf(i_0_8152, plain, (epred1_0|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_8151, plain, (epred1_0|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_8168, plain, (epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_8166, plain, (epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang))).#trainneg
cnf(i_0_8150, plain, (epred1_0|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang))).#trainneg
cnf(i_0_8149, plain, (epred1_0|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_8167, plain, (epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~v1_zf_lang(k6_zf_lang(esk17_0))|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_8175, plain, (epred1_0|~m1_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_8172, plain, (epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_8131, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m1_finseq_1(esk16_0,k5_numbers)|~m1_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_8136, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m1_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang)|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_8169, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang)|~m2_finseq_1(esk17_0,k5_numbers))).#trainneg
cnf(i_0_8231, plain, (epred1_0|~m1_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m2_finseq_1(k6_zf_lang(esk16_0),k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(k6_zf_lang(esk16_0))|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_8146, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~m2_finseq_1(esk16_0,k5_numbers))).#trainneg
cnf(i_0_8144, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang)|~v1_zf_lang(esk17_0)|~m2_finseq_1(esk17_0,k5_numbers))).#trainneg
cnf(i_0_8383, plain, (k2_zf_model(k7_zf_lang(X1,k6_zf_lang(X2)))=k2_zf_model(k7_zf_lang(X2,k6_zf_lang(X1)))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X2)|~m1_subset_1(k2_zf_model(k6_zf_lang(X2)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X1))).#trainneg
cnf(i_0_8249, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X3,X2)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X2,k1_zf_lang)|~r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_8894, plain, (r1_tarski(k3_zf_lang1(k4_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X2,X1)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2))|~m2_finseq_1(k4_zf_lang(X1,X2),k5_numbers))).#trainneg
cnf(i_0_8198, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k2_xboole_0(X4,X3))|~r1_tarski(X4,X2)|~r1_tarski(X3,X4))).#trainneg
cnf(i_0_8254, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X1,X2)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~v1_zf_lang(k4_zf_lang(X1,X2))|~r1_tarski(k2_zf_model(k4_zf_lang(X1,X2)),k1_zf_lang))).#trainneg
cnf(i_0_8148, plain, (epred1_0|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_8247, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X3,X2)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X2,k1_zf_lang)|~r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_8714, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X2,X1)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X1,X2))|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers))).#trainneg
cnf(i_0_8382, plain, (k2_zf_model(k7_zf_lang(k6_zf_lang(X1),X2))=k2_zf_model(k7_zf_lang(X1,k6_zf_lang(X2)))|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(k6_zf_lang(X1))|~m1_subset_1(k2_zf_model(k6_zf_lang(X2)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X1))).#trainneg
cnf(i_0_8232, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),X2)|~r1_tarski(k3_zf_lang1(X1),X2)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)|~m1_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_8234, plain, (r1_tarski(k2_zf_model(X1),k2_zf_model(k7_zf_lang(X2,k6_zf_lang(X1))))|~m1_subset_1(k2_zf_model(k6_zf_lang(X2)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~r1_tarski(k2_zf_model(k6_zf_lang(X1)),k1_zf_lang))).#trainneg
cnf(i_0_8129, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~m2_finseq_1(esk16_0,k5_numbers))).#trainneg
cnf(i_0_8171, plain, (epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_8331, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),k2_zf_model(k7_zf_lang(X1,k6_zf_lang(X2))))|~m1_subset_1(k2_zf_model(k6_zf_lang(X2)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_8230, plain, (epred1_0|~m1_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~m1_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_8613, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),X3)|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(X2),k2_zf_model(X1))|~r1_tarski(k2_zf_model(X1),X3)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~r1_tarski(k2_zf_model(X2),k1_zf_lang))).#trainneg
cnf(i_0_8724, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_zf_model(k5_zf_lang(X3,X4)))|~m1_finseq_1(k4_zf_lang(X3,X4),k5_numbers)|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X3,X4))|~m1_subset_1(X2,k1_zfmisc_1(k2_zf_model(k5_zf_lang(X4,X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zf_model(k5_zf_lang(X4,X3)))))).#trainneg
cnf(i_0_8599, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~m1_finseq_1(k4_zf_lang(X3,X2),k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zf_model(k4_zf_lang(X3,X2))))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X3,X2)))).#trainneg
cnf(i_0_8923, plain, (r1_tarski(k3_zf_lang1(k4_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X2,X1)))|~m1_finseq_1(k4_zf_lang(X2,X1),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X2,X1)))).#trainneg
cnf(i_0_8147, plain, (epred1_0|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang))).#trainneg
cnf(i_0_8384, plain, (k2_xboole_0(k2_zf_model(X1),k2_zf_model(k6_zf_lang(X2)))=k2_zf_model(k7_zf_lang(X1,k6_zf_lang(X2)))|~m1_subset_1(k2_zf_model(k6_zf_lang(X2)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X1)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_8170, plain, (epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang))).#trainneg
cnf(i_0_8627, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_finseq_1(k4_zf_lang(X4,X3),k5_numbers)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X4,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zf_model(k5_zf_lang(X4,X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zf_model(k5_zf_lang(X4,X3)))))).#trainneg
cnf(i_0_8716, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~m1_finseq_1(k4_zf_lang(X2,X3),k5_numbers)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k2_zf_model(k4_zf_lang(X3,X2)))|~v1_zf_lang(k4_zf_lang(X2,X3)))).#trainneg
cnf(i_0_8548, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_finseq_1(k4_zf_lang(X4,X3),k5_numbers)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X4,X3))|~m1_subset_1(X2,k1_zfmisc_1(k2_zf_model(k5_zf_lang(X3,X4))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zf_model(k5_zf_lang(X3,X4)))))).#trainneg
cnf(i_0_8273, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X2,X1)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2))|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_8210, plain, (r1_tarski(k2_xboole_0(X1,X2),X3)|~r1_tarski(X5,X3)|~r1_tarski(X4,X5)|~r1_tarski(X2,X5)|~r1_tarski(X1,X5))).#trainneg
cnf(i_0_8924, plain, (epred1_0|v1_zf_lang(esk15_0)|~m2_finseq_1(esk16_0,k5_numbers))).#trainneg
cnf(i_0_8949, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang)|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_8943, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m1_finseq_1(esk16_0,k5_numbers)|~m2_finseq_1(esk15_0,k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers))).#trainneg
cnf(i_0_8946, plain, (epred1_0|r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang))).#trainneg
cnf(i_0_9108, plain, (epred1_0|~m1_finseq_1(esk16_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_9267, plain, (epred1_0|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang)|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_9012, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_subset_1(X2,k1_zf_lang)|~m1_subset_1(X1,k1_zf_lang))).#trainneg
cnf(i_0_8944, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang)|~m2_finseq_1(esk15_0,k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers))).#trainneg
cnf(i_0_9158, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),X2)|~r1_tarski(k3_zf_lang1(X1),X2)|~v1_zf_lang(X1)|~m2_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_8996, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~m1_finseq_1(X2,k5_numbers)|~r1_tarski(k2_zf_model(X2),k2_zf_model(k4_zf_lang(X4,X3)))|~r1_tarski(k2_zf_model(X1),k2_zf_model(k4_zf_lang(X4,X3)))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_8995, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~r1_tarski(k2_zf_model(X2),k2_zf_model(k4_zf_lang(X4,X3)))|~r1_tarski(k2_zf_model(X1),k2_zf_model(k4_zf_lang(X4,X3)))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_8997, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~m1_finseq_1(X2,k5_numbers)|~m1_finseq_1(X1,k5_numbers)|~r1_tarski(k2_zf_model(X2),k2_zf_model(k4_zf_lang(X4,X3)))|~r1_tarski(k2_zf_model(X1),k2_zf_model(k4_zf_lang(X4,X3)))|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_9101, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_subset_1(X2,k1_zf_lang)|~m1_subset_1(X1,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_9192, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),X3)|~r1_tarski(k2_zf_model(X2),k2_zf_model(X1))|~r1_tarski(k2_zf_model(X1),X3)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~m1_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_9102, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_9106, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X2,X1)),k2_zf_model(k4_zf_lang(X1,X2)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_9274, plain, (k2_xboole_0(k2_zf_model(k6_zf_lang(X2)),k2_zf_model(X1))=k2_zf_model(k7_zf_lang(X1,k6_zf_lang(X2)))|~m1_subset_1(k2_zf_model(k6_zf_lang(X2)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X1)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_9104, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X2,X1)),k2_zf_model(k4_zf_lang(X1,X2)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_8930, plain, (epred1_0|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang))).#trainneg
cnf(i_0_9420, plain, (r1_tarski(X1,X2)|~r1_tarski(X4,k4_xboole_0(X2,X3))|~r1_tarski(X1,X4))).#trainneg
cnf(i_0_9175, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),k2_xboole_0(k2_zf_model(X2),k2_zf_model(X1)))|~m1_subset_1(k2_zf_model(k6_zf_lang(X2)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~v1_zf_lang(X2)|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang)))).#trainneg
cnf(i_0_8985, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X2),k2_zf_model(k4_zf_lang(X4,X3)))|~r1_tarski(k3_zf_lang1(X1),k2_zf_model(k4_zf_lang(X4,X3)))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_9663, negated_conjecture, (r1_tarski(k2_xboole_0(X1,X2),k3_zf_lang1(esk2_0))|~r1_tarski(X3,k2_zf_model(esk2_0))|~r1_tarski(X2,k2_zf_model(esk2_0))|~r1_tarski(X1,k2_zf_model(esk2_0))|~epred1_0)).#trainneg
cnf(i_0_9193, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),X3)|~r1_tarski(k2_zf_model(X2),k2_zf_model(X1))|~r1_tarski(k2_zf_model(X1),X3)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~v1_zf_lang(k6_zf_lang(X1)))).#trainneg
cnf(i_0_8979, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~m1_subset_1(X2,k1_zfmisc_1(k2_zf_model(k4_zf_lang(X4,X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zf_model(k4_zf_lang(X4,X3)))))).#trainneg
cnf(i_0_9189, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),X3)|~r1_tarski(k2_zf_model(X2),k2_zf_model(X1))|~r1_tarski(k2_zf_model(X1),X3)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~r1_tarski(k2_zf_model(X1),k1_zf_lang))).#trainneg
cnf(i_0_9284, negated_conjecture, (r1_tarski(k2_zf_model(k7_zf_lang(X1,k6_zf_lang(X2))),k2_xboole_0(k3_zf_lang1(esk2_0),X3))|~epred1_0|~r1_tarski(k2_zf_model(X1),k2_zf_model(esk2_0))|~r1_tarski(k2_zf_model(k6_zf_lang(X2)),X3)|~m1_subset_1(k2_zf_model(k6_zf_lang(X2)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X1)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_9160, plain, (r1_tarski(k2_zf_model(X1),k2_zf_model(k7_zf_lang(X2,k6_zf_lang(X1))))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~r1_tarski(k2_zf_model(k6_zf_lang(X2)),k1_zf_lang))).#trainneg
cnf(i_0_9094, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X2,X3),k5_numbers)|~m2_finseq_1(k4_zf_lang(X2,X3),k5_numbers)|~v1_zf_lang(k5_zf_lang(X2,X3))|~v1_zf_lang(k4_zf_lang(X2,X3)))).#trainneg
cnf(i_0_9378, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k2_zf_model(k4_zf_lang(X3,X2)))|~v1_zf_lang(k4_zf_lang(X2,X3))|~m2_finseq_1(k4_zf_lang(X2,X3),k5_numbers))).#trainneg
cnf(i_0_8929, plain, (r1_tarski(k2_zf_model(esk15_0),k1_zf_lang)|epred1_0|~m2_finseq_1(esk15_0,k5_numbers)|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang))).#trainneg
cnf(i_0_9146, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X2,X1)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_8983, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X2),k2_zf_model(k4_zf_lang(X4,X3)))|~r1_tarski(k1_tarski(X1),k2_zf_model(k4_zf_lang(X4,X3)))|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_8936, plain, (epred1_0|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang)|~m2_finseq_1(k6_zf_lang(esk17_0),k5_numbers)|~m1_finseq_1(esk17_0,k5_numbers)|~v1_zf_lang(esk17_0))).#trainneg
cnf(i_0_9387, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~m1_finseq_1(k4_zf_lang(X3,X4),k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zf_model(k5_zf_lang(X4,X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zf_model(k5_zf_lang(X4,X3))))|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X3,X4)))).#trainneg
cnf(i_0_9531, plain, (r1_tarski(X1,k2_zf_model(k7_zf_lang(X2,X3)))|~r1_tarski(X4,k2_zf_model(X2))|~r1_tarski(X1,X4)|~m2_finseq_1(X3,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X3)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_9718, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~m2_finseq_1(esk15_0,k5_numbers)|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers))).#trainneg
cnf(i_0_9875, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X1,X2)))|~m1_subset_1(X2,k1_zf_lang)|~m1_subset_1(X1,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_9724, plain, (epred1_0|m2_finseq_1(esk15_0,k5_numbers)|v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_10148, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),k2_zf_model(X1))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1))).#trainneg
cnf(i_0_9727, plain, (epred1_0|~r1_tarski(k2_zf_model(esk17_0),k1_zf_lang)|~m2_finseq_1(esk16_0,k5_numbers))).#trainneg
cnf(i_0_10470, plain, (epred1_0|~m1_finseq_1(esk17_0,k5_numbers)|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk16_0),k1_zf_lang)|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_10309, plain, (epred1_0|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(esk17_0),k2_zf_model(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_9969, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_10163, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),k2_zf_model(k7_zf_lang(X1,X2)))|~m1_subset_1(k2_zf_model(k6_zf_lang(X2)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~v1_zf_lang(X2)|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang)))).#trainneg
cnf(i_0_10308, plain, (epred1_0|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k2_zf_model(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(esk16_0))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_9716, plain, (epred1_0|~r1_tarski(k2_zf_model(esk17_0),k3_zf_lang1(esk16_0)))).#trainneg
cnf(i_0_10088, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k2_xboole_0(X3,X4))|~r1_tarski(X4,k4_xboole_0(X2,X5))|~r1_tarski(X3,k4_xboole_0(X2,X5)))).#trainneg
cnf(i_0_10162, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),k2_zf_model(k7_zf_lang(X2,X1)))|~m1_subset_1(k2_zf_model(k6_zf_lang(X2)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~v1_zf_lang(X2)|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang)))).#trainneg
cnf(i_0_10323, plain, (epred1_0|~r1_tarski(k2_zf_model(k4_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k4_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k5_zf_lang(esk13_0,esk14_0)),k3_zf_lang1(k5_zf_lang(esk13_0,esk14_0)))|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(esk17_0),k2_zf_model(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_10452, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X2,X1)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_10410, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~m2_finseq_1(k4_zf_lang(X2,X1),k5_numbers)|~v1_zf_lang(k5_zf_lang(X2,X1))|~v1_zf_lang(k4_zf_lang(X2,X1))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_9865, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~r1_tarski(X1,k2_zf_model(k5_zf_lang(X3,X2)))|~m1_subset_1(X2,k1_zf_lang)|~m1_subset_1(X3,k1_zf_lang))).#trainneg
cnf(i_0_9968, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X2,X1)),k2_zf_model(k4_zf_lang(X2,X1)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_10149, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),k2_xboole_0(k2_zf_model(X1),k2_zf_model(X2)))|~m1_subset_1(k2_zf_model(k6_zf_lang(X2)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_9753, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),X2)|~r1_tarski(k3_zf_lang1(X1),k4_xboole_0(X2,X3))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_9956, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~r1_tarski(X1,k2_zf_model(k5_zf_lang(X3,X2)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X3,X2),k5_numbers)|~m2_finseq_1(k4_zf_lang(X3,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X3,X2))|~v1_zf_lang(k4_zf_lang(X3,X2)))).#trainneg
cnf(i_0_10125, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),X2)|~r1_tarski(k2_zf_model(X1),X2)|~r1_tarski(k2_zf_model(X3),X2)|~m1_subset_1(k2_zf_model(k6_zf_lang(X3)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X3),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(k6_zf_lang(X3))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_10433, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~r1_tarski(X1,k2_zf_model(k5_zf_lang(X3,X2)))|~m1_finseq_1(k5_zf_lang(X3,X2),k5_numbers)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X3,X2)))).#trainneg
cnf(i_0_10453, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X2,X1)),k2_zf_model(k5_zf_lang(X2,X1)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_10225, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zf_model(k4_zf_lang(X3,X2))))|~m1_subset_1(k1_xboole_0,k1_zfmisc_1(k2_zf_model(k4_zf_lang(X3,X2))))|~m1_subset_1(X2,k1_zf_lang)|~m1_subset_1(X3,k1_zf_lang))).#trainneg
cnf(i_0_10091, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k3_zf_lang1(k8_zf_lang(X3,X4)))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X4),k4_xboole_0(X2,X5))|~r1_tarski(k1_tarski(X3),k4_xboole_0(X2,X5))|~m2_finseq_1(X4,k5_numbers)|~v1_zf_lang(X4))).#trainneg
cnf(i_0_10319, plain, (epred1_0|m2_subset_1(esk14_0,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(k6_zf_lang(esk15_0)),k3_zf_lang1(k6_zf_lang(esk15_0)))|~r1_tarski(k2_zf_model(esk17_0),k2_zf_model(esk16_0))|~r1_tarski(k2_zf_model(esk16_0),k3_zf_lang1(k7_zf_lang(esk16_0,esk17_0)))|~m2_finseq_1(esk17_0,k5_numbers)|~m2_finseq_1(esk16_0,k5_numbers)|~v1_zf_lang(esk17_0)|~v1_zf_lang(esk16_0))).#trainneg
cnf(i_0_8989, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_subset_1(X2,k1_zf_lang)|~m1_subset_1(X1,k1_zf_lang)|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_10483, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_zf_model(k4_zf_lang(X4,X3)))|~m1_finseq_1(k4_zf_lang(X3,X4),k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zf_model(k5_zf_lang(X4,X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zf_model(k5_zf_lang(X4,X3))))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X4,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_9561, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_subset_1(X2,k1_zf_lang)|~m1_subset_1(X1,k1_zf_lang)|~m1_finseq_1(k4_zf_lang(X2,X1),k5_numbers)|~m1_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X2,X1))|~v1_zf_lang(k5_zf_lang(X2,X1)))).#trainneg
cnf(i_0_9505, negated_conjecture, (r1_tarski(X1,k3_zf_lang1(esk2_0))|~r1_tarski(X2,k2_zf_model(esk2_0))|~r1_tarski(X1,X2)|~epred1_0)).#trainneg
cnf(i_0_8978, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~r1_tarski(X2,k2_zf_model(k4_zf_lang(X4,X3)))|~r1_tarski(X1,k2_zf_model(k4_zf_lang(X4,X3))))).#trainneg
cnf(i_0_9692, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_zf_model(k7_zf_lang(X3,X4)))|~r1_tarski(X5,k2_zf_model(X4))|~r1_tarski(X2,k2_zf_model(X4))|~r1_tarski(X1,k2_zf_model(X4))|~m2_finseq_1(X4,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_10639, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),k2_zf_model(X1))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~m1_finseq_1(k6_zf_lang(X1),k5_numbers))).#trainneg
cnf(i_0_11214, negated_conjecture, (r1_tarski(X1,k3_zf_lang1(esk2_0))|~epred1_0|~r1_tarski(X1,k2_zf_model(esk2_0)))).#trainneg
cnf(i_0_10638, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),k2_zf_model(X1))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~r1_tarski(k2_zf_model(k6_zf_lang(X1)),k1_zf_lang))).#trainneg
cnf(i_0_10756, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X1,X2)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_10645, plain, (epred1_0|~r1_tarski(k2_zf_model(esk15_0),k3_zf_lang1(esk15_0))|~r1_tarski(k2_zf_model(esk17_0),k2_zf_model(esk16_0))|~m2_finseq_1(esk15_0,k5_numbers)|~v1_zf_lang(esk15_0))).#trainneg
cnf(i_0_10667, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),k2_zf_model(k7_zf_lang(X1,X2)))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~v1_zf_lang(X2)|~r1_tarski(k2_zf_model(k6_zf_lang(X2)),k1_zf_lang))).#trainneg
cnf(i_0_11121, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X1,X2)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m1_subset_1(X2,k1_zf_lang)|~m1_subset_1(X1,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_10734, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),k2_zf_model(k7_zf_lang(X2,X1)))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~v1_zf_lang(X2)|~r1_tarski(k2_zf_model(k6_zf_lang(X2)),k1_zf_lang))).#trainneg
cnf(i_0_11101, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X1,X2)),X3)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X2),k4_xboole_0(X3,X4))|~r1_tarski(k1_tarski(X1),k4_xboole_0(X3,X4))|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_10813, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~m1_finseq_1(X2,k5_numbers)|~m1_finseq_1(X1,k5_numbers)|~r1_tarski(k2_zf_model(X2),k2_zf_model(k5_zf_lang(X4,X3)))|~r1_tarski(k2_zf_model(X1),k2_zf_model(k5_zf_lang(X4,X3)))|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_10812, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~m1_finseq_1(X2,k5_numbers)|~r1_tarski(k2_zf_model(X2),k2_zf_model(k5_zf_lang(X4,X3)))|~r1_tarski(k2_zf_model(X1),k2_zf_model(k5_zf_lang(X4,X3)))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_10811, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~r1_tarski(k2_zf_model(X2),k2_zf_model(k5_zf_lang(X4,X3)))|~r1_tarski(k2_zf_model(X1),k2_zf_model(k5_zf_lang(X4,X3)))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_10797, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X2),k2_zf_model(k5_zf_lang(X4,X3)))|~r1_tarski(k1_tarski(X1),k2_zf_model(k5_zf_lang(X4,X3)))|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_10990, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),X2)|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~r1_tarski(k2_zf_model(X1),X2)|~r1_tarski(k2_zf_model(X3),X2)|~m2_finseq_1(k6_zf_lang(X3),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(k6_zf_lang(X3))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~v1_zf_lang(X3)|~r1_tarski(k2_zf_model(k6_zf_lang(X3)),k1_zf_lang))).#trainneg
cnf(i_0_11053, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X2,X1)))|~m1_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X2,X1)))).#trainneg
cnf(i_0_10779, plain, (r1_tarski(k3_zf_lang1(k4_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~m2_finseq_1(k4_zf_lang(X2,X1),k5_numbers)|~v1_zf_lang(k5_zf_lang(X2,X1))|~v1_zf_lang(k4_zf_lang(X2,X1))|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_10785, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~m1_subset_1(X2,k1_zf_lang)|~m1_subset_1(X3,k1_zf_lang)|~r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_10794, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~m1_subset_1(X2,k1_zfmisc_1(k2_zf_model(k5_zf_lang(X4,X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zf_model(k5_zf_lang(X4,X3)))))).#trainneg
cnf(i_0_10727, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),X2)|~r1_tarski(k2_zf_model(X1),k4_xboole_0(X2,X4))|~r1_tarski(k2_zf_model(X3),k4_xboole_0(X2,X4))|~m1_subset_1(k2_zf_model(k6_zf_lang(X3)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X3),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(k6_zf_lang(X3))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_10651, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~r1_tarski(X1,k2_zf_model(k4_zf_lang(X3,X2)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X3,X2),k5_numbers)|~m2_finseq_1(k4_zf_lang(X3,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X3,X2))|~v1_zf_lang(k4_zf_lang(X3,X2)))).#trainneg
cnf(i_0_11390, plain, (r1_tarski(X1,k2_xboole_0(k5_numbers,X2))|~r1_tarski(X1,k1_zf_lang))).#trainneg
cnf(i_0_11314, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_zf_model(k7_zf_lang(X3,X4)))|~r1_tarski(X2,k2_zf_model(X4))|~r1_tarski(X1,k2_zf_model(X4))|~m2_finseq_1(X4,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X3)|~m1_subset_1(k2_zf_model(X4),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X5,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_10799, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X2),k2_zf_model(k5_zf_lang(X4,X3)))|~r1_tarski(k3_zf_lang1(X1),k2_zf_model(k5_zf_lang(X4,X3)))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_10867, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),X2)|~r1_tarski(k2_zf_model(X1),X2)|~r1_tarski(k2_zf_model(X3),k2_zf_model(X1))|~m1_subset_1(k2_zf_model(k6_zf_lang(X3)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X3),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(k6_zf_lang(X3))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_11021, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X3,X4)))|~m1_finseq_1(k5_zf_lang(X4,X3),k5_numbers)|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X4,X3))|~m1_finseq_1(X2,k5_numbers)|~m1_finseq_1(X1,k5_numbers)|~r1_tarski(k2_zf_model(X2),k2_zf_model(k5_zf_lang(X4,X3)))|~r1_tarski(k2_zf_model(X1),k2_zf_model(k5_zf_lang(X4,X3)))|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_10783, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X2,X1)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~m2_finseq_1(k4_zf_lang(X2,X1),k5_numbers)|~v1_zf_lang(k5_zf_lang(X2,X1))|~v1_zf_lang(k4_zf_lang(X2,X1))|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_9016, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~r1_tarski(X1,k3_zf_lang1(k4_zf_lang(X3,X2)))|~m1_finseq_1(k4_zf_lang(X3,X2),k5_numbers)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X3,X2)))).#trainneg
cnf(i_0_11267, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_zf_model(k7_zf_lang(X3,X4)))|~r1_tarski(X2,k2_zf_model(X4))|~r1_tarski(X1,k2_zf_model(X4))|~m2_finseq_1(X4,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X3)|~m1_subset_1(X6,k1_zfmisc_1(k2_zf_model(X4)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zf_model(X4))))).#trainneg
cnf(i_0_11192, plain, (r1_tarski(k3_zf_lang1(k4_zf_lang(X2,X1)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_finseq_1(k4_zf_lang(X2,X1),k5_numbers)|~m1_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~m1_subset_1(X2,k1_zf_lang)|~m1_subset_1(X1,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_9141, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X1,X2)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_10782, plain, (r1_tarski(k3_zf_lang1(k4_zf_lang(X2,X1)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~m2_finseq_1(k4_zf_lang(X2,X1),k5_numbers)|~v1_zf_lang(k5_zf_lang(X2,X1))|~v1_zf_lang(k4_zf_lang(X2,X1))|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_10866, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),X2)|~r1_tarski(k2_zf_model(X3),X2)|~r1_tarski(k2_zf_model(X1),k2_zf_model(X3))|~m1_subset_1(k2_zf_model(k6_zf_lang(X3)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X3),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(k6_zf_lang(X3))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_8969, plain, (k2_zf_model(k7_zf_lang(X1,k6_zf_lang(X2)))=k2_zf_model(k7_zf_lang(X2,k6_zf_lang(X1)))|~m1_subset_1(k2_zf_model(k6_zf_lang(X2)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X2)|~v1_zf_lang(X1)|~r1_tarski(k2_zf_model(k6_zf_lang(X1)),k1_zf_lang))).#trainneg
cnf(i_0_9532, plain, (r1_tarski(X1,k2_zf_model(k7_zf_lang(X2,X3)))|~r1_tarski(X4,k2_zf_model(X2))|~r1_tarski(X1,X4)|~m1_finseq_1(X3,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X3)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_9350, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_zf_model(k4_zf_lang(X4,X3)))|~m1_finseq_1(k4_zf_lang(X4,X3),k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k2_zf_model(k5_zf_lang(X4,X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zf_model(k5_zf_lang(X4,X3))))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X4,X3)))).#trainneg
cnf(i_0_11525, plain, (epred1_0|~m2_subset_1(esk13_0,k5_numbers,k1_zf_lang)|~m2_subset_1(esk14_0,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_11552, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),k2_zf_model(X1))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)|~m1_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_11541, negated_conjecture, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~epred1_0|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1)|~r1_tarski(k2_zf_model(k8_zf_lang(esk3_0,esk2_0)),k2_zf_model(esk2_0)))).#trainneg
cnf(i_0_11584, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),k2_zf_model(k7_zf_lang(X1,X2)))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~v1_zf_lang(X2)|~m1_finseq_1(k6_zf_lang(X1),k5_numbers))).#trainneg
cnf(i_0_11612, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),k2_zf_model(k7_zf_lang(X2,X1)))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~v1_zf_lang(X2)|~m1_finseq_1(k6_zf_lang(X1),k5_numbers))).#trainneg
cnf(i_0_12051, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_subset_1(X2,k1_zf_lang)|~m1_subset_1(X1,k1_zf_lang)|~m1_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~v1_zf_lang(k4_zf_lang(X1,X2))|~v1_zf_lang(k5_zf_lang(X2,X1)))).#trainneg
cnf(i_0_12183, plain, (r1_tarski(k3_zf_lang1(k4_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_subset_1(X2,k1_zf_lang)|~m1_subset_1(X1,k1_zf_lang)|~m1_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~v1_zf_lang(k4_zf_lang(X1,X2))|~v1_zf_lang(k5_zf_lang(X2,X1)))).#trainneg
cnf(i_0_12152, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_12165, plain, (r1_tarski(k3_zf_lang1(k4_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X1,X2)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_11769, plain, (r1_tarski(k3_zf_lang1(k4_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X1,X2)))|~m1_subset_1(X1,k1_zf_lang)|~m1_subset_1(X2,k1_zf_lang)|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~m2_finseq_1(k4_zf_lang(X2,X1),k5_numbers)|~v1_zf_lang(k5_zf_lang(X2,X1))|~v1_zf_lang(k4_zf_lang(X2,X1))|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_11729, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),X2)|~r1_tarski(k2_zf_model(X1),X2)|~r1_tarski(k2_zf_model(X3),X2)|~m2_finseq_1(k6_zf_lang(X3),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(k6_zf_lang(X3))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~v1_zf_lang(X3)|~m1_finseq_1(k6_zf_lang(X1),k5_numbers))).#trainneg
cnf(i_0_11583, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),k2_zf_model(k7_zf_lang(X1,X2)))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~v1_zf_lang(X2)|~r1_tarski(k2_zf_model(k6_zf_lang(X1)),k1_zf_lang))).#trainneg
cnf(i_0_11611, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),k2_zf_model(k7_zf_lang(X2,X1)))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~v1_zf_lang(X2)|~r1_tarski(k2_zf_model(k6_zf_lang(X1)),k1_zf_lang))).#trainneg
cnf(i_0_9129, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~r1_tarski(X1,k3_zf_lang1(k5_zf_lang(X3,X2)))|~m1_finseq_1(k5_zf_lang(X3,X2),k5_numbers)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X3,X2)))).#trainneg
cnf(i_0_9418, plain, (r1_tarski(X1,k3_zf_lang1(k8_zf_lang(X2,X3)))|~r1_tarski(X4,k1_tarski(X2))|~r1_tarski(X1,X4)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_11125, plain, (r1_tarski(k3_zf_lang1(k4_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m1_subset_1(X2,k1_zf_lang)|~m1_subset_1(X1,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_9973, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~r1_tarski(X1,k2_zf_model(k4_zf_lang(X3,X2)))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X2,X3),k5_numbers)|~m2_finseq_1(k4_zf_lang(X2,X3),k5_numbers)|~v1_zf_lang(k5_zf_lang(X2,X3))|~v1_zf_lang(k4_zf_lang(X2,X3)))).#trainneg
cnf(i_0_9375, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~m1_finseq_1(k4_zf_lang(X3,X2),k5_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k2_zf_model(k5_zf_lang(X3,X2))))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X3,X2)))).#trainneg
cnf(i_0_10098, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k2_zf_model(k7_zf_lang(X3,X4)))|~m1_finseq_1(X4,k5_numbers)|~r1_tarski(k2_zf_model(X4),k4_xboole_0(X2,X5))|~r1_tarski(k2_zf_model(X3),k4_xboole_0(X2,X5))|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_11724, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X4,X3)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X2),k2_zf_model(k5_zf_lang(X4,X3)))|~r1_tarski(k1_tarski(X1),k2_zf_model(k5_zf_lang(X4,X3)))|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X4,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_11005, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X3,X4)))|~m1_finseq_1(k5_zf_lang(X4,X3),k5_numbers)|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X4,X3))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X2),k2_zf_model(k5_zf_lang(X4,X3)))|~r1_tarski(k1_tarski(X1),k2_zf_model(k5_zf_lang(X4,X3)))|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_11964, plain, (r1_tarski(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(k5_numbers,k1_zfmisc_1(X2))|~r1_tarski(X1,k1_zf_lang))).#trainneg
cnf(i_0_9033, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X2,X1)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_11861, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~m1_subset_1(k1_xboole_0,k1_zfmisc_1(k2_zf_model(k5_zf_lang(X3,X2))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zf_model(k5_zf_lang(X3,X2))))|~m1_subset_1(X2,k1_zf_lang)|~m1_subset_1(X3,k1_zf_lang))).#trainneg
cnf(i_0_10390, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k2_zf_model(k5_zf_lang(X3,X2)))|~m2_finseq_1(k5_zf_lang(X2,X3),k5_numbers)|~m2_finseq_1(k4_zf_lang(X2,X3),k5_numbers)|~v1_zf_lang(k5_zf_lang(X2,X3))|~v1_zf_lang(k4_zf_lang(X2,X3)))).#trainneg
cnf(i_0_10827, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X3,X2),k5_numbers)|~m2_finseq_1(k4_zf_lang(X3,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X3,X2))|~v1_zf_lang(k4_zf_lang(X3,X2)))).#trainneg
cnf(i_0_12168, plain, (r1_tarski(k3_zf_lang1(k4_zf_lang(X2,X1)),k2_zf_model(k5_zf_lang(X1,X2)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_9560, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2))|~m1_finseq_1(k4_zf_lang(X2,X1),k5_numbers)|~m1_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~v1_zf_lang(k4_zf_lang(X2,X1))|~v1_zf_lang(k5_zf_lang(X2,X1)))).#trainneg
cnf(i_0_11268, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_zf_model(k7_zf_lang(X3,X4)))|~r1_tarski(X2,k2_zf_model(X4))|~r1_tarski(X1,k2_zf_model(X4))|~m2_finseq_1(X4,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zf_model(X4))))).#trainneg
cnf(i_0_9153, plain, (k2_zf_model(k7_zf_lang(k6_zf_lang(X1),X2))=k2_zf_model(k7_zf_lang(X1,k6_zf_lang(X2)))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(k6_zf_lang(X2)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_11020, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X3,X4)))|~m1_finseq_1(k5_zf_lang(X4,X3),k5_numbers)|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X4,X3))|~m1_finseq_1(X2,k5_numbers)|~r1_tarski(k2_zf_model(X2),k2_zf_model(k5_zf_lang(X4,X3)))|~r1_tarski(k2_zf_model(X1),k2_zf_model(k5_zf_lang(X4,X3)))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_10976, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X3,X4)))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X4,X3),k5_numbers)|~m2_finseq_1(k4_zf_lang(X4,X3),k5_numbers)|~v1_zf_lang(k5_zf_lang(X4,X3))|~v1_zf_lang(k4_zf_lang(X4,X3))|~m1_finseq_1(X2,k5_numbers)|~m1_finseq_1(X1,k5_numbers)|~r1_tarski(k2_zf_model(X2),k2_zf_model(k5_zf_lang(X4,X3)))|~r1_tarski(k2_zf_model(X1),k2_zf_model(k5_zf_lang(X4,X3)))|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_12052, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~v1_zf_lang(k4_zf_lang(X1,X2))|~m1_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~v1_zf_lang(k5_zf_lang(X2,X1)))).#trainneg
cnf(i_0_4, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~r1_tarski(k2_zf_model(k8_zf_lang(esk3_0,esk2_0)),k3_zf_lang1(k8_zf_lang(esk3_0,esk2_0)))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_2541, negated_conjecture, (m2_subset_1(esk3_0,k5_numbers,k5_numbers))).#trainneg
cnf(i_0_3834, plain, (r1_tarski(k2_zf_model(X1),k3_zf_lang1(X1))|~r1_tarski(k2_zf_model(k8_zf_lang(esk3_0,esk2_0)),k1_tarski(esk3_0))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_12793, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),X3)|~m1_finseq_1(X4,k5_numbers)|~r1_tarski(k2_zf_model(X4),k4_xboole_0(X3,X6))|~r1_tarski(k2_zf_model(X5),k4_xboole_0(X3,X6))|~m2_finseq_1(X5,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X5)|~m1_finseq_1(X2,k5_numbers)|~r1_tarski(k2_zf_model(X2),k2_zf_model(k7_zf_lang(X5,X4)))|~r1_tarski(k2_zf_model(X1),k2_zf_model(k7_zf_lang(X5,X4)))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_12785, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X1,X2)),X3)|~m1_finseq_1(X4,k5_numbers)|~r1_tarski(k2_zf_model(X4),k4_xboole_0(X3,X6))|~r1_tarski(k2_zf_model(X5),k4_xboole_0(X3,X6))|~m2_finseq_1(X5,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X5)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X2),k2_zf_model(k7_zf_lang(X5,X4)))|~r1_tarski(k1_tarski(X1),k2_zf_model(k7_zf_lang(X5,X4)))|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_12794, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),X3)|~m1_finseq_1(X4,k5_numbers)|~r1_tarski(k2_zf_model(X4),k4_xboole_0(X3,X6))|~r1_tarski(k2_zf_model(X5),k4_xboole_0(X3,X6))|~m2_finseq_1(X5,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X5)|~m1_finseq_1(X2,k5_numbers)|~m1_finseq_1(X1,k5_numbers)|~r1_tarski(k2_zf_model(X2),k2_zf_model(k7_zf_lang(X5,X4)))|~r1_tarski(k2_zf_model(X1),k2_zf_model(k7_zf_lang(X5,X4)))|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_12792, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),X3)|~m1_finseq_1(X4,k5_numbers)|~r1_tarski(k2_zf_model(X4),k4_xboole_0(X3,X6))|~r1_tarski(k2_zf_model(X5),k4_xboole_0(X3,X6))|~m2_finseq_1(X5,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X5)|~r1_tarski(k2_zf_model(X2),k2_zf_model(k7_zf_lang(X5,X4)))|~r1_tarski(k2_zf_model(X1),k2_zf_model(k7_zf_lang(X5,X4)))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_12787, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),X3)|~m1_finseq_1(X4,k5_numbers)|~r1_tarski(k2_zf_model(X4),k4_xboole_0(X3,X6))|~r1_tarski(k2_zf_model(X5),k4_xboole_0(X3,X6))|~m2_finseq_1(X5,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X5)|~r1_tarski(k3_zf_lang1(X2),k2_zf_model(k7_zf_lang(X5,X4)))|~r1_tarski(k3_zf_lang1(X1),k2_zf_model(k7_zf_lang(X5,X4)))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_12363, plain, (r1_tarski(X1,k2_zf_model(X2))|~r1_tarski(X1,k2_zf_model(k6_zf_lang(X2)))|~m1_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_12544, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X2,X1)),k2_zf_model(k5_zf_lang(X1,X2)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_12790, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),X3)|~m1_finseq_1(X4,k5_numbers)|~r1_tarski(k2_zf_model(X4),k4_xboole_0(X3,X6))|~r1_tarski(k2_zf_model(X5),k4_xboole_0(X3,X6))|~m2_finseq_1(X5,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X5)|~r1_tarski(k2_zf_model(X2),k2_zf_model(X1))|~r1_tarski(k2_zf_model(X1),k2_zf_model(k7_zf_lang(X5,X4)))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_12789, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),X2)|~m1_finseq_1(X3,k5_numbers)|~r1_tarski(k2_zf_model(X3),k4_xboole_0(X2,X5))|~r1_tarski(k2_zf_model(X4),k4_xboole_0(X2,X5))|~m2_finseq_1(X4,k5_numbers)|~v1_zf_lang(X3)|~v1_zf_lang(X4)|~r1_tarski(k3_zf_lang1(X1),k2_zf_model(k7_zf_lang(X4,X3)))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_11109, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~r1_tarski(X1,k3_zf_lang1(k5_zf_lang(X3,X2)))|~m1_finseq_1(k5_zf_lang(X3,X2),k5_numbers)|~m1_subset_1(X2,k1_zf_lang)|~m1_subset_1(X3,k1_zf_lang))).#trainneg
cnf(i_0_12138, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~r1_tarski(X1,k3_zf_lang1(k5_zf_lang(X2,X3)))|~m1_finseq_1(k5_zf_lang(X2,X3),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X2,X3)))).#trainneg
cnf(i_0_10793, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~r1_tarski(X2,k2_zf_model(k5_zf_lang(X4,X3)))|~r1_tarski(X1,k2_zf_model(k5_zf_lang(X4,X3))))).#trainneg
cnf(i_0_9088, plain, (r1_tarski(k2_xboole_0(X1,X2),X3)|~r1_tarski(X1,X3)|~r1_tarski(X2,X1))).#trainneg
cnf(i_0_11004, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X3,X4)))|~m1_finseq_1(k5_zf_lang(X4,X3),k5_numbers)|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X4,X3))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~r1_tarski(k3_zf_lang1(k6_zf_lang(X2)),k2_zf_model(k5_zf_lang(X4,X3)))|~r1_tarski(k1_tarski(X1),k2_zf_model(k5_zf_lang(X4,X3)))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X2))).#trainneg
cnf(i_0_11099, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),X3)|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X5),k4_xboole_0(X3,X6))|~r1_tarski(k1_tarski(X4),k4_xboole_0(X3,X6))|~m2_finseq_1(X5,k5_numbers)|~v1_zf_lang(X5)|~m1_finseq_1(X2,k5_numbers)|~m1_finseq_1(X1,k5_numbers)|~r1_tarski(k2_zf_model(X2),k3_zf_lang1(k8_zf_lang(X4,X5)))|~r1_tarski(k2_zf_model(X1),k3_zf_lang1(k8_zf_lang(X4,X5)))|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_10798, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~m1_finseq_1(X1,k5_numbers)|~r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),k2_zf_model(k5_zf_lang(X4,X3)))|~r1_tarski(k3_zf_lang1(X2),k2_zf_model(k5_zf_lang(X4,X3)))|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_12791, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),X3)|~m1_finseq_1(X4,k5_numbers)|~r1_tarski(k2_zf_model(X4),k4_xboole_0(X3,X6))|~r1_tarski(k2_zf_model(X5),k4_xboole_0(X3,X6))|~m2_finseq_1(X5,k5_numbers)|~v1_zf_lang(X4)|~v1_zf_lang(X5)|~m1_finseq_1(X1,k5_numbers)|~r1_tarski(k2_zf_model(X2),k2_zf_model(X1))|~r1_tarski(k2_zf_model(X1),k2_zf_model(k7_zf_lang(X5,X4)))|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_12642, plain, (r1_tarski(X1,k3_zf_lang1(k8_zf_lang(X2,X3)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k4_xboole_0(k1_tarski(X2),X4))|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_12186, plain, (r1_tarski(k3_zf_lang1(k4_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~v1_zf_lang(k4_zf_lang(X1,X2))|~v1_zf_lang(k5_zf_lang(X2,X1)))).#trainneg
cnf(i_0_11793, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X2),k2_zf_model(k5_zf_lang(X3,X4)))|~r1_tarski(k3_zf_lang1(X1),k2_zf_model(k5_zf_lang(X3,X4)))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_12169, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X2,X1)),k2_zf_model(k5_zf_lang(X1,X2)))|~m1_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_13868, plain, (r1_tarski(k4_xboole_0(k1_tarski(X1),X2),k3_zf_lang1(k8_zf_lang(X1,X3)))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_13676, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,k2_xboole_0(X4,X5))|~r1_tarski(X4,X2)|~r1_tarski(X5,X4))).#trainneg
cnf(i_0_13679, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X6,X2)|~r1_tarski(X1,k2_xboole_0(X4,X5))|~r1_tarski(X4,X3)|~r1_tarski(X5,X4))).#trainneg
cnf(i_0_13684, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X2,X1)),X3)|~r1_tarski(k3_zf_lang1(X1),X3)|~r1_tarski(k1_tarski(X2),k3_zf_lang1(X1))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_13686, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),X3)|~r1_tarski(k3_zf_lang1(X1),X3)|~r1_tarski(k3_zf_lang1(X2),k3_zf_lang1(X1))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_13680, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X1,X3)|~r1_tarski(X2,k2_xboole_0(X5,X6))|~r1_tarski(X5,X4)|~r1_tarski(X6,X5))).#trainneg
cnf(i_0_13682, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X1,X2)),X3)|~r1_tarski(k1_tarski(X1),X3)|~r1_tarski(k3_zf_lang1(X2),k1_tarski(X1))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_13687, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X2,X1)),X3)|~r1_tarski(k3_zf_lang1(X1),X3)|~r1_tarski(k3_zf_lang1(X2),k3_zf_lang1(X1))|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X1)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_13513, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_finseq_1(k5_zf_lang(X4,X3),k5_numbers)|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X2),k3_zf_lang1(k5_zf_lang(X4,X3)))|~r1_tarski(k1_tarski(X1),k3_zf_lang1(k5_zf_lang(X4,X3)))|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_13534, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~m1_finseq_1(k5_zf_lang(X2,X3),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k3_zf_lang1(k4_zf_lang(X2,X3)))|~v1_zf_lang(k5_zf_lang(X2,X3)))).#trainneg
cnf(i_0_13674, plain, (r1_tarski(X1,X2)|~r1_tarski(X5,k2_xboole_0(X3,X4))|~r1_tarski(X1,X5)|~r1_tarski(X3,X2)|~r1_tarski(X4,X3))).#trainneg
cnf(i_0_13532, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~m1_finseq_1(k5_zf_lang(X2,X3),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~v1_zf_lang(k5_zf_lang(X2,X3)))).#trainneg
cnf(i_0_13393, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~r1_tarski(X1,k2_zf_model(k5_zf_lang(X3,X2)))|~m1_finseq_1(k5_zf_lang(X2,X3),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X2,X3)))).#trainneg
cnf(i_0_13538, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~m1_finseq_1(k5_zf_lang(X2,X3),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k3_zf_lang1(k5_zf_lang(X3,X2)))|~v1_zf_lang(k5_zf_lang(X2,X3)))).#trainneg
cnf(i_0_13677, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X6,X3)|~r1_tarski(X1,k2_xboole_0(X4,X5))|~r1_tarski(X4,X3)|~r1_tarski(X5,X4))).#trainneg
cnf(i_0_14578, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k2_xboole_0(X3,X4))|~r1_tarski(X5,X2)|~r1_tarski(X6,X5)|~r1_tarski(X4,X5)|~r1_tarski(X3,X5))).#trainneg
cnf(i_0_14305, plain, (r1_tarski(k4_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X5,X4)|~r1_tarski(X1,X3))).#trainneg
cnf(i_0_14303, plain, (r1_tarski(k4_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X5,X3)|~r1_tarski(X1,X3))).#trainneg
cnf(i_0_14344, plain, (r1_tarski(k4_xboole_0(k2_zf_model(X1),X2),k2_zf_model(k7_zf_lang(X3,X1)))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X3),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X1)|~v1_zf_lang(X3))).#trainneg
cnf(i_0_14476, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X2,X1)),X3)|~r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(k6_zf_lang(X2)))|~r1_tarski(k3_zf_lang1(k6_zf_lang(X2)),X3)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X1)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_13940, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2))|~m1_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~v1_zf_lang(k5_zf_lang(X2,X1)))).#trainneg
cnf(i_0_13378, plain, (r1_tarski(k3_zf_lang1(k8_zf_lang(X1,X2)),k2_zf_model(X3))|~m1_finseq_1(X3,k5_numbers)|~v1_zf_lang(X3)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X2),k2_zf_model(k6_zf_lang(X3)))|~r1_tarski(k1_tarski(X1),k2_zf_model(k6_zf_lang(X3)))|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_13854, plain, (r1_tarski(k2_xboole_0(X1,X2),k3_zf_lang1(k8_zf_lang(X3,X4)))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_finseq_1(X4,k5_numbers)|~v1_zf_lang(X4)|~r1_tarski(X1,k4_xboole_0(k1_tarski(X3),X5))|~r1_tarski(X2,X1))).#trainneg
cnf(i_0_13582, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_zf_model(k4_zf_lang(X4,X3)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~r1_tarski(X2,k2_zf_model(k5_zf_lang(X4,X3)))|~r1_tarski(X1,k2_zf_model(k5_zf_lang(X4,X3)))|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X4,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_14835, plain, (r1_tarski(k4_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X1,X3))).#trainneg
cnf(i_0_14343, plain, (r1_tarski(k4_xboole_0(k2_zf_model(X1),X2),k2_zf_model(k7_zf_lang(X1,X3)))|~m1_subset_1(k2_zf_model(X3),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(X3,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X3)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_13537, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~m1_finseq_1(k5_zf_lang(X2,X3),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k3_zf_lang1(k4_zf_lang(X3,X2)))|~v1_zf_lang(k5_zf_lang(X2,X3)))).#trainneg
cnf(i_0_15239, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~r1_tarski(X1,k4_xboole_0(X3,X4)))).#trainneg
cnf(i_0_13535, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~m1_finseq_1(k5_zf_lang(X2,X3),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k2_zf_model(k4_zf_lang(X3,X2)))|~v1_zf_lang(k5_zf_lang(X2,X3)))).#trainneg
cnf(i_0_13551, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),k2_zf_model(k5_zf_lang(X2,X3)))|~m1_finseq_1(k5_zf_lang(X2,X3),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X2,X3))|~r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(k5_zf_lang(X2,X3)))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_13388, plain, (r1_tarski(k2_zf_model(k8_zf_lang(X1,k6_zf_lang(X2))),k2_zf_model(X2))|~m1_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2)|~m1_subset_1(k2_zf_model(k6_zf_lang(X2)),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~v1_zf_lang(k6_zf_lang(X2)))).#trainneg
cnf(i_0_14222, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X2,X1)),X3)|~r1_tarski(k3_zf_lang1(k6_zf_lang(X2)),k3_zf_lang1(X1))|~r1_tarski(k3_zf_lang1(X1),X3)|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X1)|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_14604, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~v1_zf_lang(k5_zf_lang(X2,X3))|~m2_finseq_1(k5_zf_lang(X2,X3),k5_numbers))).#trainneg
cnf(i_0_13549, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X3,X4)))|~m1_finseq_1(k5_zf_lang(X3,X4),k5_numbers)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X3,X4))|~r1_tarski(k3_zf_lang1(X2),k3_zf_lang1(k5_zf_lang(X3,X4)))|~r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(k5_zf_lang(X3,X4)))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_13892, plain, (r1_tarski(k3_zf_lang1(k4_zf_lang(X2,X1)),k2_zf_model(k4_zf_lang(X2,X1)))|~m1_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~m2_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~v1_zf_lang(k4_zf_lang(X1,X2))|~v1_zf_lang(k5_zf_lang(X2,X1)))).#trainneg
cnf(i_0_14350, plain, (r1_tarski(k3_zf_lang1(X1),k2_xboole_0(k10_xtuple_0(X1),X2))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_13562, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X3,X4)))|~m1_finseq_1(k5_zf_lang(X3,X4),k5_numbers)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X3,X4))|~r1_tarski(k2_zf_model(X2),k3_zf_lang1(k5_zf_lang(X3,X4)))|~r1_tarski(k2_zf_model(X1),k3_zf_lang1(k5_zf_lang(X3,X4)))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_13852, plain, (r1_tarski(k4_xboole_0(k4_xboole_0(k1_tarski(X1),X2),X3),k3_zf_lang1(k8_zf_lang(X1,X4)))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(X4,k5_numbers)|~v1_zf_lang(X4))).#trainneg
cnf(i_0_14575, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k2_xboole_0(X3,X4))|~r1_tarski(X5,X2)|~r1_tarski(X6,X5)|~r1_tarski(X4,X6)|~r1_tarski(X3,X5))).#trainneg
cnf(i_0_14327, plain, (r1_tarski(k4_xboole_0(k3_zf_lang1(X1),X2),k3_zf_lang1(k7_zf_lang(X1,X3)))|~m2_finseq_1(X3,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X3)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_13382, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),k2_zf_model(X2))|~m1_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2)|~r1_tarski(k3_zf_lang1(X1),k2_zf_model(k6_zf_lang(X2)))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_13944, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X1,X2)))|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2))|~m1_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~v1_zf_lang(k5_zf_lang(X2,X1)))).#trainneg
cnf(i_0_14569, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k4_xboole_0(X3,X4))|~r1_tarski(X3,X2)|~r1_tarski(X5,X3))).#trainneg
cnf(i_0_13947, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X1,X2)))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~m2_finseq_1(k4_zf_lang(X2,X1),k5_numbers)|~v1_zf_lang(k5_zf_lang(X2,X1))|~v1_zf_lang(k4_zf_lang(X2,X1))|~m1_finseq_1(k5_zf_lang(X2,X1),k5_numbers))).#trainneg
cnf(i_0_15709, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k4_xboole_0(X3,X4))|~r1_tarski(X3,X2))).#trainneg
cnf(i_0_16178, plain, (r1_tarski(k4_xboole_0(k3_zf_lang1(X1),X2),k3_zf_lang1(k6_zf_lang(X1)))|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1))).#trainneg
cnf(i_0_16157, plain, (r1_tarski(X1,X2)|~r1_tarski(X5,X2)|~r1_tarski(X6,X5)|~r1_tarski(X4,X6)|~r1_tarski(X3,X5)|~r1_tarski(X1,X4))).#trainneg
cnf(i_0_15702, plain, (r1_tarski(k4_xboole_0(X1,X2),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~r1_tarski(X1,k2_zf_model(k4_zf_lang(X4,X3))))).#trainneg
cnf(i_0_16026, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X1,X2)),k2_xboole_0(k10_xtuple_0(k5_zf_lang(X1,X2)),X3))|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_15711, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X2,X4)|~r1_tarski(X1,k4_xboole_0(X5,X6))|~r1_tarski(X5,X3))).#trainneg
cnf(i_0_15719, plain, (r1_tarski(k3_zf_lang1(X1),X2)|~r1_tarski(k10_xtuple_0(X1),X2)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_15710, plain, (r1_tarski(X1,X2)|~r1_tarski(X5,k4_xboole_0(X3,X4))|~r1_tarski(X1,X5)|~r1_tarski(X3,X2))).#trainneg
cnf(i_0_15698, plain, (r1_tarski(k4_xboole_0(X1,X2),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~r1_tarski(X1,k2_zf_model(k5_zf_lang(X4,X3))))).#trainneg
cnf(i_0_15572, plain, (r1_tarski(X1,k2_xboole_0(X2,X3))|~r1_tarski(X1,k4_xboole_0(X4,X5))|~r1_tarski(X4,X2))).#trainneg
cnf(i_0_16008, plain, (r1_tarski(k3_zf_lang1(X1),k2_xboole_0(X2,k10_xtuple_0(X1)))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_16025, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X1,X2)),k2_xboole_0(k10_xtuple_0(k5_zf_lang(X1,X2)),X3))|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_15815, plain, (r1_tarski(X1,k2_zf_model(k5_zf_lang(X2,X3)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~r1_tarski(X1,k2_zf_model(k4_zf_lang(X3,X2)))|~v1_zf_lang(k5_zf_lang(X2,X3))|~m2_finseq_1(k5_zf_lang(X2,X3),k5_numbers))).#trainneg
cnf(i_0_16275, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X1,X2)))|~m1_subset_1(X1,k1_zf_lang)|~m1_subset_1(X2,k1_zf_lang)|~m1_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~m1_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k4_zf_lang(X1,X2))|~v1_zf_lang(k5_zf_lang(X2,X1)))).#trainneg
cnf(i_0_16033, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X1,X2)),k2_xboole_0(k10_xtuple_0(k4_zf_lang(X1,X2)),X3))|~m2_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k4_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_14306, plain, (r1_tarski(k4_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X1,X4)|~r1_tarski(X5,X3))).#trainneg
cnf(i_0_17105, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X1,X2)),k10_xtuple_0(k4_zf_lang(X1,X2)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k4_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k4_zf_lang(X1,X2)))).#trainneg
cnf(i_0_16802, plain, (r1_tarski(k2_zf_model(k5_zf_lang(X1,X2)),k10_xtuple_0(k5_zf_lang(X1,X2)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_17034, plain, (r1_tarski(k2_zf_model(k4_zf_lang(X1,X2)),k10_xtuple_0(k5_zf_lang(X1,X2)))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2)))).#trainneg
cnf(i_0_16699, plain, (r1_tarski(k2_zf_model(k8_zf_lang(X1,X2)),k2_zf_model(X3))|~m1_finseq_1(X3,k5_numbers)|~v1_zf_lang(X3)|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(X2),k2_zf_model(k6_zf_lang(X3)))|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_16650, plain, (r1_tarski(k2_zf_model(k8_zf_lang(X1,X2)),X3)|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(X2),k4_xboole_0(X3,X4))|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_16697, plain, (r1_tarski(k2_zf_model(k8_zf_lang(X1,X2)),k2_zf_model(X3))|~m1_subset_1(k2_zf_model(X3),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~m2_finseq_1(X3,k5_numbers)|~v1_zf_lang(X3)|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(X2),k2_zf_model(k8_zf_lang(X4,X3)))|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_16912, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X5,X3)|~r1_tarski(X2,k4_xboole_0(X5,X6))|~r1_tarski(X1,k4_xboole_0(X5,X6)))).#trainneg
cnf(i_0_16910, plain, (r1_tarski(k2_xboole_0(X1,X2),k2_xboole_0(X3,X4))|~r1_tarski(X5,X3)|~r1_tarski(X2,k4_xboole_0(X5,X6))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_13870, plain, (r1_tarski(X1,k2_zf_model(k4_zf_lang(X2,X3)))|~r1_tarski(X1,k3_zf_lang1(k4_zf_lang(X3,X2)))|~m1_finseq_1(k5_zf_lang(X2,X3),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X3,X2),k5_numbers)|~m2_finseq_1(k4_zf_lang(X3,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X3,X2))|~v1_zf_lang(k4_zf_lang(X3,X2))|~v1_zf_lang(k5_zf_lang(X2,X3)))).#trainneg
cnf(i_0_13515, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_finseq_1(k5_zf_lang(X4,X3),k5_numbers)|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~r1_tarski(k3_zf_lang1(X2),k3_zf_lang1(k5_zf_lang(X4,X3)))|~r1_tarski(k3_zf_lang1(X1),k3_zf_lang1(k5_zf_lang(X4,X3)))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_13563, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X3,X4)))|~m1_finseq_1(k5_zf_lang(X3,X4),k5_numbers)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X3,X4))|~m1_finseq_1(X2,k5_numbers)|~r1_tarski(k2_zf_model(X2),k3_zf_lang1(k5_zf_lang(X3,X4)))|~r1_tarski(k2_zf_model(X1),k3_zf_lang1(k5_zf_lang(X3,X4)))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_13754, plain, (r1_tarski(k3_zf_lang1(k6_zf_lang(X1)),k2_xboole_0(X2,X3))|~r1_tarski(k3_zf_lang1(X1),k4_xboole_0(X2,X4))|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_16352, plain, (r1_tarski(k3_zf_lang1(k4_zf_lang(X2,X1)),k2_zf_model(k4_zf_lang(X1,X2)))|~m1_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(k4_zf_lang(X2,X1),k5_numbers)|~v1_zf_lang(k5_zf_lang(X2,X1))|~v1_zf_lang(k4_zf_lang(X2,X1)))).#trainneg
cnf(i_0_13941, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X2,X1)))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~m2_finseq_1(k4_zf_lang(X2,X1),k5_numbers)|~v1_zf_lang(k5_zf_lang(X2,X1))|~v1_zf_lang(k4_zf_lang(X2,X1))|~m1_finseq_1(k5_zf_lang(X2,X1),k5_numbers))).#trainneg
cnf(i_0_16017, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),k2_xboole_0(k10_xtuple_0(k7_zf_lang(X1,k6_zf_lang(X2))),X3))|~m2_finseq_1(k7_zf_lang(X1,k6_zf_lang(X2)),k5_numbers)|~v1_zf_lang(k7_zf_lang(X1,k6_zf_lang(X2)))|~m2_finseq_1(k6_zf_lang(X2),k5_numbers)|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(k6_zf_lang(X2))|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_16027, plain, (r1_tarski(k3_zf_lang1(k4_zf_lang(X1,X2)),k2_xboole_0(k10_xtuple_0(k5_zf_lang(X1,X2)),X3))|~m2_finseq_1(k5_zf_lang(X1,X2),k5_numbers)|~v1_zf_lang(k5_zf_lang(X1,X2))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang))).#trainneg
cnf(i_0_14315, plain, (r1_tarski(k4_xboole_0(X1,X2),X3)|~m1_subset_1(X4,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_13752, plain, (r1_tarski(k3_zf_lang1(k7_zf_lang(X1,X2)),k2_xboole_0(X3,X4))|~r1_tarski(k3_zf_lang1(X2),k4_xboole_0(X3,X5))|~r1_tarski(k3_zf_lang1(X1),k4_xboole_0(X3,X5))|~m2_finseq_1(X2,k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_13564, plain, (r1_tarski(k2_zf_model(k7_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X3,X4)))|~m1_finseq_1(k5_zf_lang(X3,X4),k5_numbers)|~m2_subset_1(X3,k5_numbers,k1_zf_lang)|~m2_subset_1(X4,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X3,X4))|~m1_finseq_1(X2,k5_numbers)|~m1_finseq_1(X1,k5_numbers)|~r1_tarski(k2_zf_model(X2),k3_zf_lang1(k5_zf_lang(X3,X4)))|~r1_tarski(k2_zf_model(X1),k3_zf_lang1(k5_zf_lang(X3,X4)))|~v1_zf_lang(X2)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_14351, plain, (r1_tarski(k2_zf_model(k8_zf_lang(X2,X1)),k2_xboole_0(k2_zf_model(X1),X3))|~m1_subset_1(k2_zf_model(X1),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_finseq_1(X1,k5_numbers)|~v1_zf_lang(X1))).#trainneg
cnf(i_0_16714, plain, (r1_tarski(k2_zf_model(k8_zf_lang(X1,X2)),k2_zf_model(k4_zf_lang(X3,X4)))|~m1_subset_1(X3,k1_zf_lang)|~m1_subset_1(X4,k1_zf_lang)|~m1_subset_1(k2_zf_model(X2),k1_zfmisc_1(k1_zf_lang))|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~r1_tarski(k2_zf_model(X2),k2_zf_model(k4_zf_lang(X4,X3)))|~m2_finseq_1(X2,k5_numbers)|~v1_zf_lang(X2))).#trainneg
cnf(i_0_14032, plain, (r1_tarski(k2_zf_model(k6_zf_lang(X1)),k2_xboole_0(X2,X3))|~r1_tarski(k2_zf_model(X1),X2)|~r1_tarski(k2_zf_model(X4),k2_zf_model(X1))|~m1_subset_1(k2_zf_model(k6_zf_lang(X4)),k1_zfmisc_1(k1_zf_lang))|~m1_subset_1(k2_zf_model(k6_zf_lang(X1)),k1_zfmisc_1(k1_zf_lang))|~m2_finseq_1(k6_zf_lang(X4),k5_numbers)|~m2_finseq_1(k6_zf_lang(X1),k5_numbers)|~m2_finseq_1(X1,k5_numbers)|~m2_finseq_1(X4,k5_numbers)|~v1_zf_lang(k6_zf_lang(X4))|~v1_zf_lang(k6_zf_lang(X1))|~v1_zf_lang(X1)|~v1_zf_lang(X4))).#trainneg
cnf(i_0_14347, plain, (r1_tarski(k4_xboole_0(X1,X2),k2_xboole_0(X3,X1)))).#trainneg
cnf(i_0_13946, plain, (r1_tarski(k3_zf_lang1(k5_zf_lang(X1,X2)),k2_zf_model(k5_zf_lang(X1,X2)))|~m1_finseq_1(k5_zf_lang(X2,X1),k5_numbers)|~m2_subset_1(X2,k5_numbers,k1_zf_lang)|~m2_subset_1(X1,k5_numbers,k1_zf_lang)|~v1_zf_lang(k5_zf_lang(X2,X1)))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 174
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 317
# Removed in clause preprocessing      : 23
# Initial clauses in saturation        : 294
# Processed clauses                    : 2596
# ...of these trivial                  : 7
# ...subsumed                          : 1261
# ...remaining for further processing  : 1328
# Other redundant clauses eliminated   : 1
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 511
# Backward-rewritten                   : 103
# Generated clauses                    : 17824
# ...of the previous two non-trivial   : 17302
# Contextual simplify-reflections      : 1056
# Paramodulations                      : 17823
# Factorizations                       : 0
# Equation resolutions                 : 1
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 713
#    Positive orientable unit clauses  : 98
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 22
#    Non-unit-clauses                  : 591
# Current number of unprocessed clauses: 14045
# ...number of literals in the above   : 127788
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1943
# Proof object given clauses           : 304
# Proof search given clauses           : 1328
# Clause-clause subsumption calls (NU) : 497922
# Rec. Clause-clause subsumption calls : 45421
# Non-unit clause-clause subsumptions  : 2815
# Unit Clause-clause subsumption calls : 2432
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 16
# BW rewrite match successes           : 7
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 516433

# -------------------------------------------------
# User time                : 57.532 s
# System time              : 1.859 s
# Total time               : 59.392 s
# Maximum resident set size: 547404 pages
57.53user 1.98system 1:02.52elapsed 95%CPU (0avgtext+0avgdata 547932maxresident)k
0inputs+1024outputs (0major+571423minor)pagefaults 0swaps
