2020-05-31 11:50:07.536805: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_55
2020-05-31 11:50:07.642750: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-05-31 11:50:07.914556: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-05-31 11:50:08.684465: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-05-31 11:50:10.967882: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 3431038 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_55' loaded (query=512; context=768; weigths=real; len_mult=0.000000).
# Preprocessing time       : 4.690 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(fc6_ordinal1, axiom, (~(v1_xboole_0(k4_ordinal1))&v3_ordinal1(k4_ordinal1)), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', fc6_ordinal1)).
fof(fc6_card_1, axiom, ![X1]:(v1_finset_1(X1)=>(v1_finset_1(k1_card_1(X1))&v1_card_1(k1_card_1(X1)))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', fc6_card_1)).
fof(redefinition_k3_finseq_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>k3_finseq_1(X1)=k1_card_1(X1)), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', redefinition_k3_finseq_1)).
fof(fc7_finseq_1, axiom, ![X1]:v1_finseq_1(k5_finseq_1(X1)), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', fc7_finseq_1)).
fof(redefinition_k9_finseq_1, axiom, ![X1]:k9_finseq_1(X1)=k5_finseq_1(X1), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', redefinition_k9_finseq_1)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', t6_boole)).
fof(redefinition_k2_matrlin, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_finseq_1(X2,X1))=>k2_matrlin(X1,X2)=k5_finseq_1(X2)), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', redefinition_k2_matrlin)).
fof(cc5_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(m1_subset_1(X2,X1)=>v3_ordinal1(X2))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', cc5_ordinal1)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', redefinition_k5_numbers)).
fof(cc2_finseq_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>((v1_relat_1(X1)&v1_funct_1(X1))&v1_finset_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', cc2_finseq_1)).
fof(t40_finseq_1, axiom, ![X1, X2]:(((v1_relat_1(X2)&v1_funct_1(X2))&v1_finseq_1(X2))=>(X2=k9_finseq_1(X1)<=>(k3_finseq_1(X2)=np__1&k1_funct_1(X2,np__1)=X1))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', t40_finseq_1)).
fof(dt_k9_finseq_1, axiom, ![X1]:(v1_relat_1(k9_finseq_1(X1))&v1_funct_1(k9_finseq_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', dt_k9_finseq_1)).
fof(t20_matrix_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_matrix_1(X2)&m2_finseq_1(X2,k3_finseq_2(X1)))=>(~(r1_xxreal_0(k3_finseq_1(X2),k1_xboole_0))=>![X3]:(v7_ordinal1(X3)=>(m1_matrix_1(X2,X1,k3_finseq_1(X2),X3)<=>X3=k1_matrix_1(X2)))))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', t20_matrix_1)).
fof(spc0_boole, axiom, v1_xboole_0(np__0), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', spc0_boole)).
fof(dt_k2_matrlin, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_finseq_1(X2,X1))=>m1_matrix_1(k2_matrlin(X1,X2),X1,np__1,k3_finseq_1(X2))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', dt_k2_matrlin)).
fof(cc6_card_1, axiom, ![X1]:((v3_ordinal1(X1)&v1_finset_1(X1))=>v7_ordinal1(X1)), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', cc6_card_1)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', spc1_numerals)).
fof(d3_finseq_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>![X2]:(m2_subset_1(X2,k1_numbers,k5_numbers)=>(X2=k3_finseq_1(X1)<=>k2_finseq_1(X2)=k9_xtuple_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', d3_finseq_1)).
fof(rqLessOrEqual__r1_xxreal_0__r1_r0, axiom, ~(r1_xxreal_0(np__1,np__0)), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', rqLessOrEqual__r1_xxreal_0__r1_r0)).
fof(dt_m1_matrix_1, axiom, ![X1, X2, X3]:(((~(v1_xboole_0(X1))&v7_ordinal1(X2))&v7_ordinal1(X3))=>![X4]:(m1_matrix_1(X4,X1,X2,X3)=>(v1_matrix_1(X4)&m2_finseq_1(X4,k3_finseq_2(X1))))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', dt_m1_matrix_1)).
fof(dt_k3_finseq_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', dt_k3_finseq_1)).
fof(dt_k4_matrix_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&v1_matrix_1(X2))&m1_finseq_1(X2,k3_finseq_2(X1)))=>(v1_matrix_1(k4_matrix_1(X1,X2))&m2_finseq_1(k4_matrix_1(X1,X2),k3_finseq_2(X1)))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', dt_k4_matrix_1)).
fof(redefinition_m2_finseq_1, axiom, ![X1, X2]:(m2_finseq_1(X2,X1)<=>m1_finseq_1(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', redefinition_m2_finseq_1)).
fof(t87_zfmisc_1, axiom, ![X1, X2, X3, X4]:(r2_hidden(k4_tarski(X1,X2),k2_zfmisc_1(X3,X4))<=>(r2_hidden(X1,X3)&r2_hidden(X2,X4))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', t87_zfmisc_1)).
fof(d5_tarski, axiom, ![X1, X2]:k4_tarski(X1,X2)=k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', d5_tarski)).
fof(d4_matrix_1, axiom, ![X1]:((((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))&v1_matrix_1(X1))=>k2_matrix_1(X1)=k2_zfmisc_1(k4_finseq_1(X1),k2_finseq_1(k1_matrix_1(X1)))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', d4_matrix_1)).
fof(redefinition_k4_finseq_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>k4_finseq_1(X1)=k9_xtuple_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', redefinition_k4_finseq_1)).
fof(d6_matrix_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_matrix_1(X2)&m2_finseq_1(X2,k3_finseq_2(X1)))=>![X3]:((v1_matrix_1(X3)&m2_finseq_1(X3,k3_finseq_2(X1)))=>(X3=k4_matrix_1(X1,X2)<=>((k3_finseq_1(X3)=k1_matrix_1(X2)&![X4]:(v7_ordinal1(X4)=>![X5]:(v7_ordinal1(X5)=>(r2_hidden(k4_tarski(X4,X5),k2_matrix_1(X3))<=>r2_hidden(k4_tarski(X5,X4),k2_matrix_1(X2))))))&![X4]:(v7_ordinal1(X4)=>![X5]:(v7_ordinal1(X5)=>(r2_hidden(k4_tarski(X5,X4),k2_matrix_1(X2))=>k3_matrix_1(X1,X3,X4,X5)=k3_matrix_1(X1,X2,X5,X4))))))))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', d6_matrix_1)).
fof(commutativity_k2_tarski, axiom, ![X1, X2]:k2_tarski(X1,X2)=k2_tarski(X2,X1), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', commutativity_k2_tarski)).
fof(t2_finseq_1, axiom, (k2_finseq_1(np__1)=k1_tarski(np__1)&k2_finseq_1(np__2)=k2_tarski(np__1,np__2)), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', t2_finseq_1)).
fof(t21_matrix_1, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_matrix_1(X2)&m2_finseq_1(X2,k3_finseq_2(X1)))=>![X3]:((v1_matrix_1(X3)&m2_finseq_1(X3,k3_finseq_2(X1)))=>(((k3_finseq_1(X2)=k3_finseq_1(X3)&k1_matrix_1(X2)=k1_matrix_1(X3))&![X4]:(v7_ordinal1(X4)=>![X5]:(v7_ordinal1(X5)=>(r2_hidden(k4_tarski(X4,X5),k2_matrix_1(X2))=>k3_matrix_1(X1,X2,X4,X5)=k3_matrix_1(X1,X3,X4,X5)))))=>X2=X3)))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', t21_matrix_1)).
fof(dt_m2_finseq_2, axiom, ![X1, X2]:(m1_finseq_2(X2,X1)=>![X3]:(m2_finseq_2(X3,X1,X2)=>m2_finseq_1(X3,X1))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', dt_m2_finseq_2)).
fof(dt_k3_pre_poly, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>m2_finseq_2(k3_pre_poly(X1,X2),X1,k3_finseq_2(X1))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', dt_k3_pre_poly)).
fof(dt_k3_finseq_2, axiom, ![X1]:m1_finseq_2(k3_finseq_2(X1),X1), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', dt_k3_finseq_2)).
fof(redefinition_k3_pre_poly, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k3_pre_poly(X1,X2)=k5_finseq_1(X2)), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', redefinition_k3_pre_poly)).
fof(dt_m1_finseq_1, axiom, ![X1, X2]:(m1_finseq_1(X2,X1)=>((v1_relat_1(X2)&v1_funct_1(X2))&v1_finseq_1(X2))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', dt_m1_finseq_1)).
fof(dt_m2_finseq_1, axiom, ![X1, X2]:(m2_finseq_1(X2,X1)=>((v1_funct_1(X2)&v1_finseq_1(X2))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', dt_m2_finseq_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', t3_subset)).
fof(reflexivity_r1_tarski, axiom, ![X1, X2]:r1_tarski(X1,X1), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', reflexivity_r1_tarski)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', fc1_subset_1)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', d1_tarski)).
fof(t10_matrix_2, axiom, ![X1]:(~(v1_xboole_0(X1))=>![X2]:((v1_matrix_1(X2)&m2_finseq_1(X2,k3_finseq_2(X1)))=>(~(r1_xxreal_0(k1_matrix_1(X2),k6_numbers))=>(k3_finseq_1(k4_matrix_1(X1,X2))=k1_matrix_1(X2)&k1_matrix_1(k4_matrix_1(X1,X2))=k3_finseq_1(X2))))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', t10_matrix_2)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', redefinition_k6_numbers)).
fof(fc12_finseq_1, axiom, ![X1]:~(v1_xboole_0(k13_finseq_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', fc12_finseq_1)).
fof(redefinition_k3_finseq_2, axiom, ![X1]:k3_finseq_2(X1)=k13_finseq_1(X1), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', redefinition_k3_finseq_2)).
fof(redefinition_m2_finseq_2, axiom, ![X1, X2]:(m1_finseq_2(X2,X1)=>![X3]:(m2_finseq_2(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', redefinition_m2_finseq_2)).
fof(t15_matrlin, conjecture, ![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,X1)=>k2_matrlin(X1,k3_pre_poly(X1,X2))=k4_matrix_1(X1,k2_matrlin(X1,k3_pre_poly(X1,X2))))), file('/scratch/mptp/allpred/preds__0.005/t15_matrlin', t15_matrlin)).
fof(c_0_47, plain, (~v1_xboole_0(k4_ordinal1)&v3_ordinal1(k4_ordinal1)), inference(fof_simplification,[status(thm)],[fc6_ordinal1])).
fof(c_0_48, plain, ![X85]:((v1_finset_1(k1_card_1(X85))|~v1_finset_1(X85))&(v1_card_1(k1_card_1(X85))|~v1_finset_1(X85))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc6_card_1])])])).
fof(c_0_49, plain, ![X68]:(~v1_relat_1(X68)|~v1_funct_1(X68)|~v1_finseq_1(X68)|k3_finseq_1(X68)=k1_card_1(X68)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k3_finseq_1])])).
fof(c_0_50, plain, ![X71]:v1_finseq_1(k5_finseq_1(X71)), inference(variable_rename,[status(thm)],[fc7_finseq_1])).
fof(c_0_51, plain, ![X39]:k9_finseq_1(X39)=k5_finseq_1(X39), inference(variable_rename,[status(thm)],[redefinition_k9_finseq_1])).
fof(c_0_52, plain, ![X88]:(~v1_xboole_0(X88)|X88=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
fof(c_0_53, plain, ![X42, X43]:(v1_xboole_0(X42)|~m1_finseq_1(X43,X42)|k2_matrlin(X42,X43)=k5_finseq_1(X43)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k2_matrlin])])])).
fof(c_0_54, plain, ![X141, X142]:(~v3_ordinal1(X141)|(~m1_subset_1(X142,X141)|v3_ordinal1(X142))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_ordinal1])])])).
cnf(c_0_55, plain, (v3_ordinal1(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_56, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
cnf(c_0_57, plain, (v1_finset_1(k1_card_1(X1))|~v1_finset_1(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_58, plain, (k3_finseq_1(X1)=k1_card_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_49])).
fof(c_0_59, plain, ![X196]:(((v1_relat_1(X196)|(~v1_relat_1(X196)|~v1_funct_1(X196)|~v1_finseq_1(X196)))&(v1_funct_1(X196)|(~v1_relat_1(X196)|~v1_funct_1(X196)|~v1_finseq_1(X196))))&(v1_finset_1(X196)|(~v1_relat_1(X196)|~v1_funct_1(X196)|~v1_finseq_1(X196)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_finseq_1])])])).
fof(c_0_60, plain, ![X11, X12]:(((k3_finseq_1(X12)=np__1|X12!=k9_finseq_1(X11)|(~v1_relat_1(X12)|~v1_funct_1(X12)|~v1_finseq_1(X12)))&(k1_funct_1(X12,np__1)=X11|X12!=k9_finseq_1(X11)|(~v1_relat_1(X12)|~v1_funct_1(X12)|~v1_finseq_1(X12))))&(k3_finseq_1(X12)!=np__1|k1_funct_1(X12,np__1)!=X11|X12=k9_finseq_1(X11)|(~v1_relat_1(X12)|~v1_funct_1(X12)|~v1_finseq_1(X12)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t40_finseq_1])])])).
cnf(c_0_61, plain, (v1_finseq_1(k5_finseq_1(X1))), inference(split_conjunct,[status(thm)],[c_0_50])).
cnf(c_0_62, plain, (k9_finseq_1(X1)=k5_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_51])).
fof(c_0_63, plain, ![X66]:(v1_relat_1(k9_finseq_1(X66))&v1_funct_1(k9_finseq_1(X66))), inference(variable_rename,[status(thm)],[dt_k9_finseq_1])).
fof(c_0_64, plain, ![X8, X9, X10]:((~m1_matrix_1(X9,X8,k3_finseq_1(X9),X10)|X10=k1_matrix_1(X9)|~v7_ordinal1(X10)|r1_xxreal_0(k3_finseq_1(X9),k1_xboole_0)|(~v1_matrix_1(X9)|~m2_finseq_1(X9,k3_finseq_2(X8)))|v1_xboole_0(X8))&(X10!=k1_matrix_1(X9)|m1_matrix_1(X9,X8,k3_finseq_1(X9),X10)|~v7_ordinal1(X10)|r1_xxreal_0(k3_finseq_1(X9),k1_xboole_0)|(~v1_matrix_1(X9)|~m2_finseq_1(X9,k3_finseq_2(X8)))|v1_xboole_0(X8))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t20_matrix_1])])])])])).
cnf(c_0_65, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_66, plain, (v1_xboole_0(np__0)), inference(split_conjunct,[status(thm)],[spc0_boole])).
fof(c_0_67, plain, ![X46, X47]:(v1_xboole_0(X46)|~m1_finseq_1(X47,X46)|m1_matrix_1(k2_matrlin(X46,X47),X46,np__1,k3_finseq_1(X47))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k2_matrlin])])])).
cnf(c_0_68, plain, (v1_xboole_0(X1)|k2_matrlin(X1,X2)=k5_finseq_1(X2)|~m1_finseq_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
fof(c_0_69, plain, ![X146]:(~v3_ordinal1(X146)|~v1_finset_1(X146)|v7_ordinal1(X146)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc6_card_1])])).
cnf(c_0_70, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_71, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_72, plain, (v3_ordinal1(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_55])).
cnf(c_0_73, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_56])).
cnf(c_0_74, plain, (v1_finset_1(k1_card_1(X1))|~v1_finset_1(X1)), inference(evalgc,[status(thm)],[c_0_57])).
cnf(c_0_75, plain, (k1_card_1(X1)=k3_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_58])).
cnf(c_0_76, plain, (v1_finset_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_77, plain, (k3_finseq_1(X1)=np__1|X1!=k9_finseq_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_78, plain, (v1_finseq_1(k9_finseq_1(X1))), inference(rw,[status(thm)],[c_0_61, c_0_62])).
cnf(c_0_79, plain, (v1_funct_1(k9_finseq_1(X1))), inference(split_conjunct,[status(thm)],[c_0_63])).
cnf(c_0_80, plain, (v1_relat_1(k9_finseq_1(X1))), inference(split_conjunct,[status(thm)],[c_0_63])).
fof(c_0_81, plain, ![X15, X16]:((X16!=k3_finseq_1(X15)|k2_finseq_1(X16)=k9_xtuple_0(X15)|~m2_subset_1(X16,k1_numbers,k5_numbers)|(~v1_relat_1(X15)|~v1_funct_1(X15)|~v1_finseq_1(X15)))&(k2_finseq_1(X16)!=k9_xtuple_0(X15)|X16=k3_finseq_1(X15)|~m2_subset_1(X16,k1_numbers,k5_numbers)|(~v1_relat_1(X15)|~v1_funct_1(X15)|~v1_finseq_1(X15)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_finseq_1])])])])).
cnf(c_0_82, plain, (X3=k1_matrix_1(X1)|r1_xxreal_0(k3_finseq_1(X1),k1_xboole_0)|v1_xboole_0(X2)|~m1_matrix_1(X1,X2,k3_finseq_1(X1),X3)|~v7_ordinal1(X3)|~v1_matrix_1(X1)|~m2_finseq_1(X1,k3_finseq_2(X2))), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_83, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_84, plain, (v1_xboole_0(np__0)), inference(evalgc,[status(thm)],[c_0_66])).
fof(c_0_85, plain, ~r1_xxreal_0(np__1,np__0), inference(fof_simplification,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r1_r0])).
fof(c_0_86, plain, ![X72, X73, X74, X75]:((v1_matrix_1(X75)|~m1_matrix_1(X75,X72,X73,X74)|(v1_xboole_0(X72)|~v7_ordinal1(X73)|~v7_ordinal1(X74)))&(m2_finseq_1(X75,k3_finseq_2(X72))|~m1_matrix_1(X75,X72,X73,X74)|(v1_xboole_0(X72)|~v7_ordinal1(X73)|~v7_ordinal1(X74)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_matrix_1])])])])])).
cnf(c_0_87, plain, (v1_xboole_0(X1)|m1_matrix_1(k2_matrlin(X1,X2),X1,np__1,k3_finseq_1(X2))|~m1_finseq_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_88, plain, (k2_matrlin(X1,X2)=k9_finseq_1(X2)|v1_xboole_0(X1)|~m1_finseq_1(X2,X1)), inference(rw,[status(thm)],[c_0_68, c_0_62])).
cnf(c_0_89, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_90, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_70])).
cnf(c_0_91, plain, (m1_subset_1(np__1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_71])).
cnf(c_0_92, plain, (v3_ordinal1(k5_numbers)), inference(rw,[status(thm)],[c_0_72, c_0_73])).
cnf(c_0_93, plain, (v1_finset_1(k3_finseq_1(X1))|~v1_finset_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75])])).
cnf(c_0_94, plain, (v1_finset_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_95, plain, (k3_finseq_1(X1)=np__1|X1!=k9_finseq_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_77])).
cnf(c_0_96, plain, (v1_finseq_1(k9_finseq_1(X1))), inference(evalgc,[status(thm)],[c_0_78])).
cnf(c_0_97, plain, (v1_funct_1(k9_finseq_1(X1))), inference(evalgc,[status(thm)],[c_0_79])).
cnf(c_0_98, plain, (v1_relat_1(k9_finseq_1(X1))), inference(evalgc,[status(thm)],[c_0_80])).
cnf(c_0_99, plain, (k2_finseq_1(X1)=k9_xtuple_0(X2)|X1!=k3_finseq_1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
fof(c_0_100, plain, ![X76]:(~v1_relat_1(X76)|~v1_funct_1(X76)|~v1_finseq_1(X76)|m2_subset_1(k3_finseq_1(X76),k1_numbers,k5_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_finseq_1])])).
fof(c_0_101, plain, ![X55, X56]:((v1_matrix_1(k4_matrix_1(X55,X56))|(v1_xboole_0(X55)|~v1_matrix_1(X56)|~m1_finseq_1(X56,k3_finseq_2(X55))))&(m2_finseq_1(k4_matrix_1(X55,X56),k3_finseq_2(X55))|(v1_xboole_0(X55)|~v1_matrix_1(X56)|~m1_finseq_1(X56,k3_finseq_2(X55))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_matrix_1])])])])).
fof(c_0_102, plain, ![X48, X49]:((~m2_finseq_1(X49,X48)|m1_finseq_1(X49,X48))&(~m1_finseq_1(X49,X48)|m2_finseq_1(X49,X48))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_m2_finseq_1])])).
cnf(c_0_103, plain, (X3=k1_matrix_1(X1)|v1_xboole_0(X2)|r1_xxreal_0(k3_finseq_1(X1),k1_xboole_0)|~v1_matrix_1(X1)|~v7_ordinal1(X3)|~m2_finseq_1(X1,k3_finseq_2(X2))|~m1_matrix_1(X1,X2,k3_finseq_1(X1),X3)), inference(evalgc,[status(thm)],[c_0_82])).
cnf(c_0_104, plain, (k1_xboole_0=np__0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_84])])).
cnf(c_0_105, plain, (~r1_xxreal_0(np__1,np__0)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_106, plain, (m2_finseq_1(X1,k3_finseq_2(X2))|v1_xboole_0(X2)|~m1_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X4)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_107, plain, (v1_xboole_0(X1)|m1_matrix_1(k2_matrlin(X1,X2),X1,np__1,k3_finseq_1(X2))|~m1_finseq_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_87])).
cnf(c_0_108, plain, (k2_matrlin(X1,X2)=k9_finseq_1(X2)|v1_xboole_0(X1)|~m1_finseq_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_88])).
cnf(c_0_109, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1)), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_110, plain, (v3_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_91]), c_0_92])])])).
cnf(c_0_111, plain, (v1_finset_1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_112, plain, (k3_finseq_1(k9_finseq_1(X1))=np__1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_95]), c_0_96]), c_0_97]), c_0_98])])])).
fof(c_0_113, plain, ![X17, X18, X19, X20]:(((r2_hidden(X17,X19)|~r2_hidden(k4_tarski(X17,X18),k2_zfmisc_1(X19,X20)))&(r2_hidden(X18,X20)|~r2_hidden(k4_tarski(X17,X18),k2_zfmisc_1(X19,X20))))&(~r2_hidden(X17,X19)|~r2_hidden(X18,X20)|r2_hidden(k4_tarski(X17,X18),k2_zfmisc_1(X19,X20)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t87_zfmisc_1])])])).
fof(c_0_114, plain, ![X62, X63]:k4_tarski(X62,X63)=k2_tarski(k2_tarski(X62,X63),k1_tarski(X62)), inference(variable_rename,[status(thm)],[d5_tarski])).
fof(c_0_115, plain, ![X109]:(~v1_relat_1(X109)|~v1_funct_1(X109)|~v1_finseq_1(X109)|~v1_matrix_1(X109)|k2_matrix_1(X109)=k2_zfmisc_1(k4_finseq_1(X109),k2_finseq_1(k1_matrix_1(X109)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d4_matrix_1])])).
fof(c_0_116, plain, ![X67]:(~v1_relat_1(X67)|~v1_funct_1(X67)|~v1_finseq_1(X67)|k4_finseq_1(X67)=k9_xtuple_0(X67)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_finseq_1])])).
cnf(c_0_117, plain, (k9_xtuple_0(X2)=k2_finseq_1(X1)|X1!=k3_finseq_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_99])).
cnf(c_0_118, plain, (m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_100])).
fof(c_0_119, plain, ![X28, X29, X30, X31, X32, X33, X34]:((((k3_finseq_1(X30)=k1_matrix_1(X29)|X30!=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28))&((~r2_hidden(k4_tarski(X31,X32),k2_matrix_1(X30))|r2_hidden(k4_tarski(X32,X31),k2_matrix_1(X29))|~v7_ordinal1(X32)|~v7_ordinal1(X31)|X30!=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28))&(~r2_hidden(k4_tarski(X32,X31),k2_matrix_1(X29))|r2_hidden(k4_tarski(X31,X32),k2_matrix_1(X30))|~v7_ordinal1(X32)|~v7_ordinal1(X31)|X30!=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28))))&(~v7_ordinal1(X33)|(~v7_ordinal1(X34)|(~r2_hidden(k4_tarski(X34,X33),k2_matrix_1(X29))|k3_matrix_1(X28,X30,X33,X34)=k3_matrix_1(X28,X29,X34,X33)))|X30!=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28)))&(((v7_ordinal1(esk6_3(X28,X29,X30))|(v7_ordinal1(esk4_3(X28,X29,X30))|k3_finseq_1(X30)!=k1_matrix_1(X29))|X30=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28))&((v7_ordinal1(esk7_3(X28,X29,X30))|(v7_ordinal1(esk4_3(X28,X29,X30))|k3_finseq_1(X30)!=k1_matrix_1(X29))|X30=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28))&((r2_hidden(k4_tarski(esk7_3(X28,X29,X30),esk6_3(X28,X29,X30)),k2_matrix_1(X29))|(v7_ordinal1(esk4_3(X28,X29,X30))|k3_finseq_1(X30)!=k1_matrix_1(X29))|X30=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28))&(k3_matrix_1(X28,X30,esk6_3(X28,X29,X30),esk7_3(X28,X29,X30))!=k3_matrix_1(X28,X29,esk7_3(X28,X29,X30),esk6_3(X28,X29,X30))|(v7_ordinal1(esk4_3(X28,X29,X30))|k3_finseq_1(X30)!=k1_matrix_1(X29))|X30=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28)))))&(((v7_ordinal1(esk6_3(X28,X29,X30))|(v7_ordinal1(esk5_3(X28,X29,X30))|k3_finseq_1(X30)!=k1_matrix_1(X29))|X30=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28))&((v7_ordinal1(esk7_3(X28,X29,X30))|(v7_ordinal1(esk5_3(X28,X29,X30))|k3_finseq_1(X30)!=k1_matrix_1(X29))|X30=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28))&((r2_hidden(k4_tarski(esk7_3(X28,X29,X30),esk6_3(X28,X29,X30)),k2_matrix_1(X29))|(v7_ordinal1(esk5_3(X28,X29,X30))|k3_finseq_1(X30)!=k1_matrix_1(X29))|X30=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28))&(k3_matrix_1(X28,X30,esk6_3(X28,X29,X30),esk7_3(X28,X29,X30))!=k3_matrix_1(X28,X29,esk7_3(X28,X29,X30),esk6_3(X28,X29,X30))|(v7_ordinal1(esk5_3(X28,X29,X30))|k3_finseq_1(X30)!=k1_matrix_1(X29))|X30=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28)))))&(((v7_ordinal1(esk6_3(X28,X29,X30))|(~r2_hidden(k4_tarski(esk4_3(X28,X29,X30),esk5_3(X28,X29,X30)),k2_matrix_1(X30))|~r2_hidden(k4_tarski(esk5_3(X28,X29,X30),esk4_3(X28,X29,X30)),k2_matrix_1(X29))|k3_finseq_1(X30)!=k1_matrix_1(X29))|X30=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28))&((v7_ordinal1(esk7_3(X28,X29,X30))|(~r2_hidden(k4_tarski(esk4_3(X28,X29,X30),esk5_3(X28,X29,X30)),k2_matrix_1(X30))|~r2_hidden(k4_tarski(esk5_3(X28,X29,X30),esk4_3(X28,X29,X30)),k2_matrix_1(X29))|k3_finseq_1(X30)!=k1_matrix_1(X29))|X30=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28))&((r2_hidden(k4_tarski(esk7_3(X28,X29,X30),esk6_3(X28,X29,X30)),k2_matrix_1(X29))|(~r2_hidden(k4_tarski(esk4_3(X28,X29,X30),esk5_3(X28,X29,X30)),k2_matrix_1(X30))|~r2_hidden(k4_tarski(esk5_3(X28,X29,X30),esk4_3(X28,X29,X30)),k2_matrix_1(X29))|k3_finseq_1(X30)!=k1_matrix_1(X29))|X30=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28))&(k3_matrix_1(X28,X30,esk6_3(X28,X29,X30),esk7_3(X28,X29,X30))!=k3_matrix_1(X28,X29,esk7_3(X28,X29,X30),esk6_3(X28,X29,X30))|(~r2_hidden(k4_tarski(esk4_3(X28,X29,X30),esk5_3(X28,X29,X30)),k2_matrix_1(X30))|~r2_hidden(k4_tarski(esk5_3(X28,X29,X30),esk4_3(X28,X29,X30)),k2_matrix_1(X29))|k3_finseq_1(X30)!=k1_matrix_1(X29))|X30=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28)))))&((v7_ordinal1(esk6_3(X28,X29,X30))|(r2_hidden(k4_tarski(esk4_3(X28,X29,X30),esk5_3(X28,X29,X30)),k2_matrix_1(X30))|r2_hidden(k4_tarski(esk5_3(X28,X29,X30),esk4_3(X28,X29,X30)),k2_matrix_1(X29))|k3_finseq_1(X30)!=k1_matrix_1(X29))|X30=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28))&((v7_ordinal1(esk7_3(X28,X29,X30))|(r2_hidden(k4_tarski(esk4_3(X28,X29,X30),esk5_3(X28,X29,X30)),k2_matrix_1(X30))|r2_hidden(k4_tarski(esk5_3(X28,X29,X30),esk4_3(X28,X29,X30)),k2_matrix_1(X29))|k3_finseq_1(X30)!=k1_matrix_1(X29))|X30=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28))&((r2_hidden(k4_tarski(esk7_3(X28,X29,X30),esk6_3(X28,X29,X30)),k2_matrix_1(X29))|(r2_hidden(k4_tarski(esk4_3(X28,X29,X30),esk5_3(X28,X29,X30)),k2_matrix_1(X30))|r2_hidden(k4_tarski(esk5_3(X28,X29,X30),esk4_3(X28,X29,X30)),k2_matrix_1(X29))|k3_finseq_1(X30)!=k1_matrix_1(X29))|X30=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28))&(k3_matrix_1(X28,X30,esk6_3(X28,X29,X30),esk7_3(X28,X29,X30))!=k3_matrix_1(X28,X29,esk7_3(X28,X29,X30),esk6_3(X28,X29,X30))|(r2_hidden(k4_tarski(esk4_3(X28,X29,X30),esk5_3(X28,X29,X30)),k2_matrix_1(X30))|r2_hidden(k4_tarski(esk5_3(X28,X29,X30),esk4_3(X28,X29,X30)),k2_matrix_1(X29))|k3_finseq_1(X30)!=k1_matrix_1(X29))|X30=k4_matrix_1(X28,X29)|(~v1_matrix_1(X30)|~m2_finseq_1(X30,k3_finseq_2(X28)))|(~v1_matrix_1(X29)|~m2_finseq_1(X29,k3_finseq_2(X28)))|v1_xboole_0(X28))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_matrix_1])])])])])])).
cnf(c_0_120, plain, (m2_finseq_1(k4_matrix_1(X1,X2),k3_finseq_2(X1))|v1_xboole_0(X1)|~v1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(X1))), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_121, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_122, plain, (X1=k1_matrix_1(X2)|r1_xxreal_0(k3_finseq_1(X2),np__0)|v1_xboole_0(X3)|~m1_matrix_1(X2,X3,k3_finseq_1(X2),X1)|~v7_ordinal1(X1)|~m2_finseq_1(X2,k3_finseq_2(X3))|~v1_matrix_1(X2)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_104])])).
cnf(c_0_123, plain, (~r1_xxreal_0(np__1,np__0)), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_124, plain, (v1_xboole_0(X2)|m2_finseq_1(X1,k3_finseq_2(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4)), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_125, plain, (m1_matrix_1(k9_finseq_1(X2),X1,np__1,k3_finseq_1(X2))|v1_xboole_0(X1)|~m1_finseq_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_108])])).
cnf(c_0_126, plain, (v7_ordinal1(np__1)|~v1_finset_1(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_110])])).
cnf(c_0_127, plain, (v1_finset_1(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_112]), c_0_96]), c_0_97]), c_0_98])])])).
cnf(c_0_128, plain, (v1_matrix_1(X1)|v1_xboole_0(X2)|~m1_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X4)), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_129, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X3,X1),k2_zfmisc_1(X4,X2))), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_130, plain, (k4_tarski(X1,X2)=k2_tarski(k2_tarski(X1,X2),k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_114])).
fof(c_0_131, plain, ![X57, X58]:k2_tarski(X57,X58)=k2_tarski(X58,X57), inference(variable_rename,[status(thm)],[commutativity_k2_tarski])).
cnf(c_0_132, plain, (k2_matrix_1(X1)=k2_zfmisc_1(k4_finseq_1(X1),k2_finseq_1(k1_matrix_1(X1)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)|~v1_matrix_1(X1)), inference(split_conjunct,[status(thm)],[c_0_115])).
cnf(c_0_133, plain, (k4_finseq_1(X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_116])).
cnf(c_0_134, plain, (k2_finseq_1(k3_finseq_1(X1))=k9_xtuple_0(X1)|~m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_117])])).
cnf(c_0_135, plain, (m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_118])).
cnf(c_0_136, plain, (k2_finseq_1(np__1)=k1_tarski(np__1)), inference(split_conjunct,[status(thm)],[t2_finseq_1])).
fof(c_0_137, plain, ![X80, X81, X82]:((v7_ordinal1(esk8_3(X80,X81,X82))|(k3_finseq_1(X81)!=k3_finseq_1(X82)|k1_matrix_1(X81)!=k1_matrix_1(X82))|X81=X82|(~v1_matrix_1(X82)|~m2_finseq_1(X82,k3_finseq_2(X80)))|(~v1_matrix_1(X81)|~m2_finseq_1(X81,k3_finseq_2(X80)))|v1_xboole_0(X80))&((v7_ordinal1(esk9_3(X80,X81,X82))|(k3_finseq_1(X81)!=k3_finseq_1(X82)|k1_matrix_1(X81)!=k1_matrix_1(X82))|X81=X82|(~v1_matrix_1(X82)|~m2_finseq_1(X82,k3_finseq_2(X80)))|(~v1_matrix_1(X81)|~m2_finseq_1(X81,k3_finseq_2(X80)))|v1_xboole_0(X80))&((r2_hidden(k4_tarski(esk8_3(X80,X81,X82),esk9_3(X80,X81,X82)),k2_matrix_1(X81))|(k3_finseq_1(X81)!=k3_finseq_1(X82)|k1_matrix_1(X81)!=k1_matrix_1(X82))|X81=X82|(~v1_matrix_1(X82)|~m2_finseq_1(X82,k3_finseq_2(X80)))|(~v1_matrix_1(X81)|~m2_finseq_1(X81,k3_finseq_2(X80)))|v1_xboole_0(X80))&(k3_matrix_1(X80,X81,esk8_3(X80,X81,X82),esk9_3(X80,X81,X82))!=k3_matrix_1(X80,X82,esk8_3(X80,X81,X82),esk9_3(X80,X81,X82))|(k3_finseq_1(X81)!=k3_finseq_1(X82)|k1_matrix_1(X81)!=k1_matrix_1(X82))|X81=X82|(~v1_matrix_1(X82)|~m2_finseq_1(X82,k3_finseq_2(X80)))|(~v1_matrix_1(X81)|~m2_finseq_1(X81,k3_finseq_2(X80)))|v1_xboole_0(X80))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t21_matrix_1])])])])])])).
cnf(c_0_138, plain, (k3_finseq_1(X1)=k1_matrix_1(X2)|v1_xboole_0(X3)|X1!=k4_matrix_1(X3,X2)|~v1_matrix_1(X1)|~m2_finseq_1(X1,k3_finseq_2(X3))|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X3))), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_139, plain, (v1_xboole_0(X1)|m2_finseq_1(k4_matrix_1(X1,X2),k3_finseq_2(X1))|~v1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(X1))), inference(evalgc,[status(thm)],[c_0_120])).
cnf(c_0_140, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_121])).
cnf(c_0_141, plain, (v1_matrix_1(k4_matrix_1(X1,X2))|v1_xboole_0(X1)|~v1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(X1))), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_142, plain, (X1=k1_matrix_1(k9_finseq_1(X2))|v1_xboole_0(X3)|~m1_matrix_1(k9_finseq_1(X2),X3,np__1,X1)|~v7_ordinal1(X1)|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X3))|~v1_matrix_1(k9_finseq_1(X2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_112]), c_0_123])])).
cnf(c_0_143, plain, (m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|v1_xboole_0(X2)|~v7_ordinal1(k3_finseq_1(X1))|~v7_ordinal1(np__1)|~m1_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125])])).
cnf(c_0_144, plain, (v7_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_126, c_0_127])])])).
cnf(c_0_145, plain, (v1_xboole_0(X2)|v1_matrix_1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4)), inference(evalgc,[status(thm)],[c_0_128])).
fof(c_0_146, plain, ![X50, X51, X52]:(~m1_finseq_2(X51,X50)|(~m2_finseq_2(X52,X50,X51)|m2_finseq_1(X52,X50))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m2_finseq_2])])])).
fof(c_0_147, plain, ![X69, X70]:(v1_xboole_0(X69)|~m1_subset_1(X70,X69)|m2_finseq_2(k3_pre_poly(X69,X70),X69,k3_finseq_2(X69))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_pre_poly])])])).
fof(c_0_148, plain, ![X126]:m1_finseq_2(k3_finseq_2(X126),X126), inference(variable_rename,[status(thm)],[dt_k3_finseq_2])).
fof(c_0_149, plain, ![X40, X41]:(v1_xboole_0(X40)|~m1_subset_1(X41,X40)|k3_pre_poly(X40,X41)=k5_finseq_1(X41)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_pre_poly])])])).
cnf(c_0_150, plain, (m1_matrix_1(k2_matrlin(X1,k9_finseq_1(X2)),X1,np__1,np__1)|v1_xboole_0(X1)|~m1_finseq_1(k9_finseq_1(X2),X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_112])])).
cnf(c_0_151, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k2_tarski(X3,X1),k1_tarski(X3)),k2_zfmisc_1(X4,X2))), inference(rw,[status(thm)],[c_0_129, c_0_130])).
cnf(c_0_152, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_131])).
cnf(c_0_153, plain, (k2_zfmisc_1(k4_finseq_1(X1),k2_finseq_1(k1_matrix_1(X1)))=k2_matrix_1(X1)|~v1_matrix_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_132])).
cnf(c_0_154, plain, (k4_finseq_1(X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_133])).
cnf(c_0_155, plain, (k2_finseq_1(k3_finseq_1(X1))=k9_xtuple_0(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_134, c_0_135])).
cnf(c_0_156, plain, (k2_finseq_1(np__1)=k1_tarski(np__1)), inference(evalgc,[status(thm)],[c_0_136])).
cnf(c_0_157, plain, (r2_hidden(k4_tarski(esk8_3(X1,X2,X3),esk9_3(X1,X2,X3)),k2_matrix_1(X2))|X2=X3|v1_xboole_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))), inference(split_conjunct,[status(thm)],[c_0_137])).
cnf(c_0_158, plain, (k1_matrix_1(X2)=k3_finseq_1(X1)|v1_xboole_0(X3)|X1!=k4_matrix_1(X3,X2)|~v1_matrix_1(X2)|~v1_matrix_1(X1)|~m2_finseq_1(X2,k3_finseq_2(X3))|~m2_finseq_1(X1,k3_finseq_2(X3))), inference(evalgc,[status(thm)],[c_0_138])).
cnf(c_0_159, plain, (m2_finseq_1(k4_matrix_1(X1,X2),k3_finseq_2(X1))|v1_xboole_0(X1)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140])])).
cnf(c_0_160, plain, (v1_xboole_0(X1)|v1_matrix_1(k4_matrix_1(X1,X2))|~v1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(X1))), inference(evalgc,[status(thm)],[c_0_141])).
fof(c_0_161, plain, ![X64, X65]:(((v1_relat_1(X65)|~m1_finseq_1(X65,X64))&(v1_funct_1(X65)|~m1_finseq_1(X65,X64)))&(v1_finseq_1(X65)|~m1_finseq_1(X65,X64))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_finseq_1])])])).
cnf(c_0_162, plain, (k1_matrix_1(k9_finseq_1(X1))=k3_finseq_1(X1)|v1_xboole_0(X2)|~v7_ordinal1(k3_finseq_1(X1))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~v1_matrix_1(k9_finseq_1(X1))|~m1_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_125])])).
cnf(c_0_163, plain, (m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|v1_xboole_0(X2)|~m1_finseq_1(X1,X2)|~v7_ordinal1(k3_finseq_1(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_143, c_0_144])])])).
cnf(c_0_164, plain, (v1_matrix_1(k9_finseq_1(X1))|v1_xboole_0(X2)|~v7_ordinal1(k3_finseq_1(X1))|~v7_ordinal1(np__1)|~m1_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_125])])).
cnf(c_0_165, plain, (m2_finseq_1(X3,X2)|~m1_finseq_2(X1,X2)|~m2_finseq_2(X3,X2,X1)), inference(split_conjunct,[status(thm)],[c_0_146])).
cnf(c_0_166, plain, (v1_xboole_0(X1)|m2_finseq_2(k3_pre_poly(X1,X2),X1,k3_finseq_2(X1))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_147])).
cnf(c_0_167, plain, (m1_finseq_2(k3_finseq_2(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_148])).
cnf(c_0_168, plain, (v1_xboole_0(X1)|k3_pre_poly(X1,X2)=k5_finseq_1(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_149])).
cnf(c_0_169, plain, (v1_matrix_1(k2_matrlin(X1,k9_finseq_1(X2)))|v1_xboole_0(X1)|~v7_ordinal1(np__1)|~m1_finseq_1(k9_finseq_1(X2),X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_150])])).
cnf(c_0_170, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k2_tarski(X3,X1),k1_tarski(X3)),k2_zfmisc_1(X4,X2))), inference(evalgc,[status(thm)],[c_0_151])).
cnf(c_0_171, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1)), inference(evalgc,[status(thm)],[c_0_152])).
cnf(c_0_172, plain, (k2_zfmisc_1(k9_xtuple_0(X1),k2_finseq_1(k1_matrix_1(X1)))=k2_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_154])])).
cnf(c_0_173, plain, (k9_xtuple_0(k9_finseq_1(X1))=k1_tarski(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_112]), c_0_156]), c_0_96]), c_0_97]), c_0_98])])])).
cnf(c_0_174, plain, (X2=X3|v1_xboole_0(X1)|r2_hidden(k2_tarski(k2_tarski(esk8_3(X1,X2,X3),esk9_3(X1,X2,X3)),k1_tarski(esk8_3(X1,X2,X3))),k2_matrix_1(X2))|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))), inference(rw,[status(thm)],[c_0_157, c_0_130])).
cnf(c_0_175, plain, (k1_matrix_1(X1)=k3_finseq_1(k4_matrix_1(X2,X3))|v1_xboole_0(X2)|k4_matrix_1(X2,X3)!=k4_matrix_1(X2,X1)|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(X1)|~v1_matrix_1(k4_matrix_1(X2,X3))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_159])])).
cnf(c_0_176, plain, (v1_matrix_1(k4_matrix_1(X1,X2))|v1_xboole_0(X1)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_140])])).
cnf(c_0_177, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_161])).
fof(c_0_178, plain, ![X44, X45]:(((v1_funct_1(X45)|~m2_finseq_1(X45,X44))&(v1_finseq_1(X45)|~m2_finseq_1(X45,X44)))&(m1_subset_1(X45,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X44)))|~m2_finseq_1(X45,X44))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m2_finseq_1])])])).
cnf(c_0_179, plain, (k9_xtuple_0(k9_finseq_1(X1))=k2_finseq_1(X2)|X2!=np__1|~m2_subset_1(X2,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_112]), c_0_96]), c_0_97]), c_0_98])])])).
cnf(c_0_180, plain, (k1_matrix_1(k9_finseq_1(X1))=k3_finseq_1(X1)|v1_xboole_0(X2)|~m1_finseq_1(X1,X2)|~v7_ordinal1(k3_finseq_1(X1))|~v1_matrix_1(k9_finseq_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_163])])).
cnf(c_0_181, plain, (v1_matrix_1(k9_finseq_1(X1))|v1_xboole_0(X2)|~m1_finseq_1(X1,X2)|~v7_ordinal1(k3_finseq_1(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_164, c_0_144])])])).
cnf(c_0_182, plain, (m2_finseq_1(X3,X2)|~m1_finseq_2(X1,X2)|~m2_finseq_2(X3,X2,X1)), inference(evalgc,[status(thm)],[c_0_165])).
cnf(c_0_183, plain, (v1_xboole_0(X1)|m2_finseq_2(k3_pre_poly(X1,X2),X1,k3_finseq_2(X1))|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_166])).
cnf(c_0_184, plain, (m1_finseq_2(k3_finseq_2(X1),X1)), inference(evalgc,[status(thm)],[c_0_167])).
cnf(c_0_185, plain, (k3_pre_poly(X1,X2)=k9_finseq_1(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_168, c_0_62])).
fof(c_0_186, plain, ![X77, X78]:((~m1_subset_1(X77,k1_zfmisc_1(X78))|r1_tarski(X77,X78))&(~r1_tarski(X77,X78)|m1_subset_1(X77,k1_zfmisc_1(X78)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
fof(c_0_187, plain, ![X110]:r1_tarski(X110,X110), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r1_tarski])])).
cnf(c_0_188, plain, (v1_matrix_1(k2_matrlin(X1,k9_finseq_1(X2)))|v1_xboole_0(X1)|~m1_finseq_1(k9_finseq_1(X2),X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_169, c_0_144])])).
cnf(c_0_189, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X3,X1)),k2_zfmisc_1(X4,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_171])])).
cnf(c_0_190, plain, (k2_zfmisc_1(k1_tarski(np__1),k2_finseq_1(k1_matrix_1(k9_finseq_1(X1))))=k2_matrix_1(k9_finseq_1(X1))|~v1_matrix_1(k9_finseq_1(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_173]), c_0_96]), c_0_97]), c_0_98])])])).
cnf(c_0_191, plain, (X2=X3|v1_xboole_0(X1)|r2_hidden(k2_tarski(k2_tarski(esk8_3(X1,X2,X3),esk9_3(X1,X2,X3)),k1_tarski(esk8_3(X1,X2,X3))),k2_matrix_1(X2))|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))), inference(evalgc,[status(thm)],[c_0_174])).
cnf(c_0_192, plain, (k1_matrix_1(X1)=k3_finseq_1(k4_matrix_1(X2,X3))|v1_xboole_0(X2)|k4_matrix_1(X2,X3)!=k4_matrix_1(X2,X1)|~m2_finseq_1(X1,k3_finseq_2(X2))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X1)|~v1_matrix_1(X3)), inference(csr,[status(thm)],[c_0_175, c_0_176])).
cnf(c_0_193, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_177])).
cnf(c_0_194, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_178])).
cnf(c_0_195, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_178])).
cnf(c_0_196, plain, (k2_finseq_1(X1)=k1_tarski(np__1)|X1!=np__1|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(rw,[status(thm)],[c_0_179, c_0_173])).
cnf(c_0_197, plain, (k1_matrix_1(k9_finseq_1(X1))=k3_finseq_1(X1)|v1_xboole_0(X2)|~m1_finseq_1(X1,X2)|~v7_ordinal1(k3_finseq_1(X1))), inference(csr,[status(thm)],[c_0_180, c_0_181])).
cnf(c_0_198, plain, (m2_finseq_1(k3_pre_poly(X1,X2),X1)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_183]), c_0_184])])])).
cnf(c_0_199, plain, (k3_pre_poly(X1,X2)=k9_finseq_1(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_185])).
cnf(c_0_200, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_186])).
cnf(c_0_201, plain, (r1_tarski(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_187])).
fof(c_0_202, plain, ![X121]:~v1_xboole_0(k1_zfmisc_1(X121)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
cnf(c_0_203, plain, (v1_matrix_1(k9_finseq_1(k9_finseq_1(X2)))|v1_xboole_0(X1)|~m1_finseq_1(k9_finseq_1(X2),X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_108])])).
cnf(c_0_204, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(k9_finseq_1(X2))))|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X3,X1)),k2_matrix_1(k9_finseq_1(X2)))|~v1_matrix_1(k9_finseq_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_190])])).
cnf(c_0_205, plain, (X1=X2|r2_hidden(k2_tarski(k1_tarski(esk8_3(X3,X1,X2)),k2_tarski(esk8_3(X3,X1,X2),esk9_3(X3,X1,X2))),k2_matrix_1(X1))|v1_xboole_0(X3)|k3_finseq_1(X1)!=k3_finseq_1(X2)|k1_matrix_1(X1)!=k1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X3))|~m2_finseq_1(X1,k3_finseq_2(X3))|~v1_matrix_1(X2)|~v1_matrix_1(X1)), inference(rw,[status(thm)],[c_0_191, c_0_171])).
cnf(c_0_206, plain, (k3_finseq_1(k4_matrix_1(X2,X1))=k1_matrix_1(X1)|v1_xboole_0(X2)|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_192])])).
cnf(c_0_207, plain, (v1_relat_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_140])])).
cnf(c_0_208, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_194])).
cnf(c_0_209, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_195])).
cnf(c_0_210, plain, (k3_matrix_1(X4,X5,X1,X2)=k3_matrix_1(X4,X3,X2,X1)|v1_xboole_0(X4)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r2_hidden(k4_tarski(X2,X1),k2_matrix_1(X3))|X5!=k4_matrix_1(X4,X3)|~v1_matrix_1(X5)|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X4))), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_211, plain, (k9_xtuple_0(X1)=k1_tarski(np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__1|~m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_196])])).
cnf(c_0_212, plain, (k1_matrix_1(k9_finseq_1(X1))=k3_finseq_1(X1)|v1_xboole_0(X2)|~v7_ordinal1(k3_finseq_1(X1))|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_140])])).
cnf(c_0_213, plain, (m2_finseq_1(k9_finseq_1(X2),X1)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_199])])).
cnf(c_0_214, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_200])).
cnf(c_0_215, plain, (r1_tarski(X1,X1)), inference(evalgc,[status(thm)],[c_0_201])).
cnf(c_0_216, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_202])).
cnf(c_0_217, plain, (v1_matrix_1(k9_finseq_1(k9_finseq_1(X1)))|v1_xboole_0(X2)|~m2_finseq_1(k9_finseq_1(X1),X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_140])])).
cnf(c_0_218, plain, (X1=k3_finseq_1(X2)|k2_finseq_1(X1)!=k9_xtuple_0(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
fof(c_0_219, plain, ![X21, X22, X23, X24, X25, X26]:(((~r2_hidden(X23,X22)|X23=X21|X22!=k1_tarski(X21))&(X24!=X21|r2_hidden(X24,X22)|X22!=k1_tarski(X21)))&((~r2_hidden(esk3_2(X25,X26),X26)|esk3_2(X25,X26)!=X25|X26=k1_tarski(X25))&(r2_hidden(esk3_2(X25,X26),X26)|esk3_2(X25,X26)=X25|X26=k1_tarski(X25)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_220, plain, (r2_hidden(esk9_3(X1,k9_finseq_1(X2),X3),k2_finseq_1(k1_matrix_1(k9_finseq_1(X2))))|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|~v1_matrix_1(k9_finseq_1(X2))|k3_finseq_1(X3)!=np__1|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~v1_matrix_1(X3)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_205]), c_0_112])])).
cnf(c_0_221, plain, (m2_subset_1(k1_matrix_1(X2),k1_numbers,k5_numbers)|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X1,X2))|~v1_funct_1(k4_matrix_1(X1,X2))|~v1_relat_1(k4_matrix_1(X1,X2))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_206])])).
cnf(c_0_222, plain, (v1_relat_1(k4_matrix_1(X1,X2))|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_159])])).
cnf(c_0_223, plain, (v1_funct_1(k4_matrix_1(X1,X2))|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_159])])).
cnf(c_0_224, plain, (v1_finseq_1(k4_matrix_1(X1,X2))|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_159])])).
cnf(c_0_225, plain, (k3_matrix_1(X4,X3,X2,X1)=k3_matrix_1(X4,X5,X1,X2)|v1_xboole_0(X4)|X5!=k4_matrix_1(X4,X3)|~v1_matrix_1(X5)|~v1_matrix_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_finseq_1(X5,k3_finseq_2(X4))|~m2_finseq_1(X3,k3_finseq_2(X4))|~r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X2)),k2_matrix_1(X3))), inference(rw,[status(thm)],[c_0_210, c_0_130])).
cnf(c_0_226, plain, (k9_xtuple_0(X1)=k1_tarski(np__1)|k3_finseq_1(X1)!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_211, c_0_135])).
cnf(c_0_227, plain, (k1_matrix_1(k9_finseq_1(k9_finseq_1(X1)))=np__1|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_213]), c_0_112]), c_0_112]), c_0_144])])])).
cnf(c_0_228, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_215])])).
cnf(c_0_229, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_216])).
cnf(c_0_230, plain, (v1_matrix_1(k9_finseq_1(k9_finseq_1(X1)))|v1_xboole_0(X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_213])])).
cnf(c_0_231, plain, (r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))|~r2_hidden(X1,X2)|~r2_hidden(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_232, plain, (X1=k3_finseq_1(X2)|k9_xtuple_0(X2)!=k2_finseq_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_218])).
cnf(c_0_233, plain, (X1=X3|~r2_hidden(X1,X2)|X2!=k1_tarski(X3)), inference(split_conjunct,[status(thm)],[c_0_219])).
cnf(c_0_234, plain, (k9_finseq_1(X1)=X2|r2_hidden(esk9_3(X3,k9_finseq_1(X1),X2),k1_tarski(np__1))|v1_xboole_0(X3)|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(X2)|k3_finseq_1(X2)!=np__1|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X3))|~m2_finseq_1(X2,k3_finseq_2(X3))|~v1_matrix_1(k9_finseq_1(X1))|~v1_matrix_1(X2)|k1_matrix_1(k9_finseq_1(X1))!=np__1|~m2_subset_1(k1_matrix_1(k9_finseq_1(X1)),k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_196])])).
cnf(c_0_235, plain, (m2_subset_1(k1_matrix_1(X1),k1_numbers,k5_numbers)|v1_xboole_0(X2)|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_221, c_0_222]), c_0_223]), c_0_224])).
cnf(c_0_236, plain, (k3_matrix_1(X4,X3,X2,X1)=k3_matrix_1(X4,X5,X1,X2)|v1_xboole_0(X4)|X5!=k4_matrix_1(X4,X3)|~v1_matrix_1(X5)|~v1_matrix_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_finseq_1(X5,k3_finseq_2(X4))|~m2_finseq_1(X3,k3_finseq_2(X4))|~r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X2)),k2_matrix_1(X3))), inference(evalgc,[status(thm)],[c_0_225])).
cnf(c_0_237, plain, (k2_zfmisc_1(k1_tarski(np__1),k2_finseq_1(k1_matrix_1(X1)))=k2_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1)|k3_finseq_1(X1)!=np__1), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_226])])).
cnf(c_0_238, plain, (k1_matrix_1(k9_finseq_1(k9_finseq_1(X1)))=np__1), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_228]), c_0_229])])).
cnf(c_0_239, plain, (v1_matrix_1(k9_finseq_1(k9_finseq_1(X1)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_228]), c_0_229])])).
cnf(c_0_240, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_zfmisc_1(X2,X4))|~r2_hidden(X3,X4)|~r2_hidden(X1,X2)), inference(rw,[status(thm)],[c_0_231, c_0_130])).
cnf(c_0_241, plain, (r2_hidden(X1,X2)|~r2_hidden(k4_tarski(X1,X3),k2_zfmisc_1(X2,X4))), inference(split_conjunct,[status(thm)],[c_0_113])).
cnf(c_0_242, plain, (k3_finseq_1(X1)=k3_finseq_1(X2)|k9_xtuple_0(X2)!=k9_xtuple_0(X1)|~m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_155])])).
cnf(c_0_243, plain, (X1=X3|X2!=k1_tarski(X3)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_233])).
cnf(c_0_244, plain, (k9_finseq_1(X1)=X2|r2_hidden(esk9_3(X3,k9_finseq_1(X1),X2),k1_tarski(np__1))|v1_xboole_0(X3)|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(X2)|k1_matrix_1(k9_finseq_1(X1))!=np__1|k3_finseq_1(X2)!=np__1|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X3))|~m2_finseq_1(X2,k3_finseq_2(X3))|~v1_matrix_1(k9_finseq_1(X1))|~v1_matrix_1(X2)), inference(csr,[status(thm)],[c_0_234, c_0_235])).
cnf(c_0_245, plain, (k3_matrix_1(X1,X2,X3,X4)=k3_matrix_1(X1,X5,X4,X3)|v1_xboole_0(X1)|X5!=k4_matrix_1(X1,X2)|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X3,X4)),k2_matrix_1(X2))|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m2_finseq_1(X5,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X5)|~v1_matrix_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_171])])).
cnf(c_0_246, plain, (k2_matrix_1(k9_finseq_1(k9_finseq_1(X1)))=k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_238]), c_0_156]), c_0_112]), c_0_96]), c_0_97]), c_0_98]), c_0_239])])])).
cnf(c_0_247, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_zfmisc_1(X2,X4))|~r2_hidden(X3,X4)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_240])).
cnf(c_0_248, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_219])).
cnf(c_0_249, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_zfmisc_1(X2,X4))), inference(rw,[status(thm)],[c_0_241, c_0_130])).
cnf(c_0_250, plain, (k3_finseq_1(X1)=k3_finseq_1(X2)|k9_xtuple_0(X2)!=k9_xtuple_0(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)), inference(csr,[status(thm)],[c_0_242, c_0_135])).
cnf(c_0_251, plain, (k9_xtuple_0(k4_matrix_1(X1,X2))=k2_finseq_1(k1_matrix_1(X2))|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X1,X2))|~v1_funct_1(k4_matrix_1(X1,X2))|~v1_relat_1(k4_matrix_1(X1,X2))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_206])])).
cnf(c_0_252, plain, (X2=X3|v1_xboole_0(X1)|k3_matrix_1(X1,X2,esk8_3(X1,X2,X3),esk9_3(X1,X2,X3))!=k3_matrix_1(X1,X3,esk8_3(X1,X2,X3),esk9_3(X1,X2,X3))|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))), inference(split_conjunct,[status(thm)],[c_0_137])).
cnf(c_0_253, plain, (esk9_3(X1,k9_finseq_1(X2),X3)=X4|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|k1_tarski(np__1)!=k1_tarski(X4)|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k1_matrix_1(k9_finseq_1(X2))!=np__1|k3_finseq_1(X3)!=np__1|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_244])])).
cnf(c_0_254, plain, (k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),X3,X4)=k3_matrix_1(X1,X5,X4,X3)|v1_xboole_0(X1)|X5!=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X3,X4)),k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1)))|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m2_finseq_1(X5,k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~v1_matrix_1(X5)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_246]), c_0_239])])])).
cnf(c_0_255, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_171])])).
cnf(c_0_256, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(evalgc,[status(thm)],[c_0_248])).
cnf(c_0_257, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_zfmisc_1(X2,X4))), inference(evalgc,[status(thm)],[c_0_249])).
cnf(c_0_258, plain, (k3_finseq_1(X2)=np__1|k9_xtuple_0(X2)!=k1_tarski(np__1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_173]), c_0_112]), c_0_96]), c_0_97]), c_0_98])])])).
cnf(c_0_259, plain, (k9_xtuple_0(k4_matrix_1(X1,X2))=k2_finseq_1(k1_matrix_1(X2))|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_251, c_0_222]), c_0_223]), c_0_224])).
cnf(c_0_260, plain, (X2=X3|v1_xboole_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|k3_matrix_1(X1,X2,esk8_3(X1,X2,X3),esk9_3(X1,X2,X3))!=k3_matrix_1(X1,X3,esk8_3(X1,X2,X3),esk9_3(X1,X2,X3))|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))), inference(evalgc,[status(thm)],[c_0_252])).
cnf(c_0_261, plain, (esk9_3(X1,k9_finseq_1(X2),X3)=np__1|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k1_matrix_1(k9_finseq_1(X2))!=np__1|k3_finseq_1(X3)!=np__1|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_253])])).
cnf(c_0_262, plain, (k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),X3,X4)=k3_matrix_1(X1,X5,X4,X3)|v1_xboole_0(X1)|X5!=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X5,k3_finseq_2(X1))|~v1_matrix_1(X5)|~r2_hidden(X4,k1_tarski(np__1))|~r2_hidden(X3,k1_tarski(np__1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_255])])).
cnf(c_0_263, plain, (r2_hidden(X1,X2)|X2!=k1_tarski(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_256])])).
cnf(c_0_264, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X3)),k2_zfmisc_1(X2,X4))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_171])])).
cnf(c_0_265, plain, (k3_finseq_1(k4_matrix_1(X1,X2))=np__1|v1_xboole_0(X1)|k2_finseq_1(k1_matrix_1(X2))!=k1_tarski(np__1)|~v1_finseq_1(k4_matrix_1(X1,X2))|~v1_funct_1(k4_matrix_1(X1,X2))|~v1_relat_1(k4_matrix_1(X1,X2))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_259])])).
fof(c_0_266, plain, ![X13, X14]:((k3_finseq_1(k4_matrix_1(X13,X14))=k1_matrix_1(X14)|r1_xxreal_0(k1_matrix_1(X14),k6_numbers)|(~v1_matrix_1(X14)|~m2_finseq_1(X14,k3_finseq_2(X13)))|v1_xboole_0(X13))&(k1_matrix_1(k4_matrix_1(X13,X14))=k3_finseq_1(X14)|r1_xxreal_0(k1_matrix_1(X14),k6_numbers)|(~v1_matrix_1(X14)|~m2_finseq_1(X14,k3_finseq_2(X13)))|v1_xboole_0(X13))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t10_matrix_2])])])])])).
cnf(c_0_267, plain, (k9_finseq_1(X1)=X2|v1_xboole_0(X3)|k3_matrix_1(X3,k9_finseq_1(X1),esk8_3(X3,k9_finseq_1(X1),X2),np__1)!=k3_matrix_1(X3,X2,esk8_3(X3,k9_finseq_1(X1),X2),np__1)|k3_finseq_1(X2)!=np__1|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X3))|~v1_matrix_1(X2)|~v1_matrix_1(k9_finseq_1(X1))|k1_matrix_1(k9_finseq_1(X1))!=np__1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_261]), c_0_112])])])).
cnf(c_0_268, plain, (k3_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X4,X3)=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),X3,X4)|v1_xboole_0(X1)|~r2_hidden(X4,k1_tarski(np__1))|~r2_hidden(X3,k1_tarski(np__1))|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_262])])).
cnf(c_0_269, plain, (r2_hidden(X1,k1_tarski(X1))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_263])])).
cnf(c_0_270, plain, (r2_hidden(X1,k1_tarski(np__1))|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_matrix_1(k9_finseq_1(X3)))|~v1_matrix_1(k9_finseq_1(X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_190])])).
cnf(c_0_271, plain, (k3_finseq_1(k4_matrix_1(X1,X2))=np__1|v1_xboole_0(X1)|k2_finseq_1(k1_matrix_1(X2))!=k1_tarski(np__1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_265, c_0_222]), c_0_223]), c_0_224])).
cnf(c_0_272, plain, (X1=k9_finseq_1(X2)|k3_finseq_1(X1)!=np__1|k1_funct_1(X1,np__1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_273, plain, (k1_matrix_1(k4_matrix_1(X1,X2))=k3_finseq_1(X2)|r1_xxreal_0(k1_matrix_1(X2),k6_numbers)|v1_xboole_0(X1)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))), inference(split_conjunct,[status(thm)],[c_0_266])).
cnf(c_0_274, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
fof(c_0_275, plain, ![X153]:~v1_xboole_0(k13_finseq_1(X153)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc12_finseq_1])])).
fof(c_0_276, plain, ![X149]:k3_finseq_2(X149)=k13_finseq_1(X149), inference(variable_rename,[status(thm)],[redefinition_k3_finseq_2])).
fof(c_0_277, plain, ![X59, X60, X61]:((~m2_finseq_2(X61,X59,X60)|m1_subset_1(X61,X60)|~m1_finseq_2(X60,X59))&(~m1_subset_1(X61,X60)|m2_finseq_2(X61,X59,X60)|~m1_finseq_2(X60,X59))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_m2_finseq_2])])])])).
cnf(c_0_278, plain, (k9_finseq_1(X1)=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,k9_finseq_1(X1),esk8_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,esk8_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k9_finseq_1(X1))!=np__1|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=np__1|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~v1_matrix_1(k9_finseq_1(X1))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~r2_hidden(esk8_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~v7_ordinal1(esk8_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_268]), c_0_269]), c_0_144])])])).
cnf(c_0_279, plain, (r2_hidden(esk8_3(X1,k9_finseq_1(X2),X3),k1_tarski(np__1))|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|~v1_matrix_1(k9_finseq_1(X2))|k3_finseq_1(X3)!=np__1|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~v1_matrix_1(X3)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_270, c_0_205]), c_0_112])])).
cnf(c_0_280, plain, (k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))=np__1|v1_xboole_0(X1)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_238]), c_0_156]), c_0_239])])])).
cnf(c_0_281, plain, (X1=k9_finseq_1(X2)|k3_finseq_1(X1)!=np__1|k1_funct_1(X1,np__1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_272])).
cnf(c_0_282, plain, (k1_matrix_1(k4_matrix_1(X1,X2))=k3_finseq_1(X2)|v1_xboole_0(X1)|r1_xxreal_0(k1_matrix_1(X2),k6_numbers)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))), inference(evalgc,[status(thm)],[c_0_273])).
cnf(c_0_283, plain, (k6_numbers=k1_xboole_0), inference(evalgc,[status(thm)],[c_0_274])).
cnf(c_0_284, plain, (~v1_xboole_0(k13_finseq_1(X1))), inference(split_conjunct,[status(thm)],[c_0_275])).
cnf(c_0_285, plain, (k3_finseq_2(X1)=k13_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_276])).
cnf(c_0_286, plain, (m1_subset_1(X1,X3)|~m2_finseq_2(X1,X2,X3)|~m1_finseq_2(X3,X2)), inference(split_conjunct,[status(thm)],[c_0_277])).
cnf(c_0_287, plain, (k9_finseq_1(X1)=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,k9_finseq_1(X1),esk8_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,esk8_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k9_finseq_1(X1))!=np__1|~v7_ordinal1(esk8_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k9_finseq_1(X1))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_278, c_0_279]), c_0_280])).
cnf(c_0_288, plain, (k9_finseq_1(k1_funct_1(X1,np__1))=X1|k3_finseq_1(X1)!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_281])])).
cnf(c_0_289, plain, (k1_matrix_1(k4_matrix_1(X1,X2))=k3_finseq_1(X2)|r1_xxreal_0(k1_matrix_1(X2),k1_xboole_0)|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)), inference(rw,[status(thm)],[c_0_282, c_0_283])).
cnf(c_0_290, plain, (~v1_xboole_0(k3_finseq_2(X1))), inference(rw,[status(thm)],[c_0_284, c_0_285])).
cnf(c_0_291, plain, (m1_subset_1(X1,X3)|~m1_finseq_2(X3,X2)|~m2_finseq_2(X1,X2,X3)), inference(evalgc,[status(thm)],[c_0_286])).
cnf(c_0_292, plain, (X1=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,X1,esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|k1_matrix_1(X1)!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(X1)!=np__1|~v7_ordinal1(esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X1)|k3_finseq_1(X1)!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_287, c_0_288])])).
cnf(c_0_293, plain, (esk8_3(X1,k9_finseq_1(X2),X3)=X4|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|k1_tarski(np__1)!=k1_tarski(X4)|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k3_finseq_1(X3)!=np__1|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_279])])).
cnf(c_0_294, plain, (k1_matrix_1(k4_matrix_1(X1,X2))=k3_finseq_1(X2)|r1_xxreal_0(k1_matrix_1(X2),np__0)|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_289, c_0_104])])).
cnf(c_0_295, plain, (~v1_xboole_0(k3_finseq_2(X1))), inference(evalgc,[status(thm)],[c_0_290])).
cnf(c_0_296, plain, (m1_subset_1(k3_pre_poly(X1,X2),k3_finseq_2(X1))|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_183]), c_0_184])])])).
cnf(c_0_297, plain, (X1=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,X1,esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|k1_matrix_1(X1)!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(X1)!=np__1|k3_finseq_1(X1)!=np__1|~v7_ordinal1(esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_292, c_0_207]), c_0_208]), c_0_209])).
cnf(c_0_298, plain, (esk8_3(X1,k9_finseq_1(X2),X3)=np__1|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k3_finseq_1(X3)!=np__1|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_293])])).
cnf(c_0_299, plain, (k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))=np__1|r1_xxreal_0(k1_matrix_1(k9_finseq_1(X2)),np__0)|v1_xboole_0(X1)|~v1_matrix_1(k9_finseq_1(X2))|~m1_subset_1(X2,k3_finseq_2(X1))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_294, c_0_213]), c_0_112]), c_0_295])])).
cnf(c_0_300, plain, (m1_subset_1(k9_finseq_1(X2),k3_finseq_2(X1))|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_296, c_0_199])])).
fof(c_0_301, negated_conjecture, ~(![X1]:(~(v1_xboole_0(X1))=>![X2]:(m1_subset_1(X2,X1)=>k2_matrlin(X1,k3_pre_poly(X1,X2))=k4_matrix_1(X1,k2_matrlin(X1,k3_pre_poly(X1,X2)))))), inference(assume_negation,[status(cth)],[t15_matrlin])).
cnf(c_0_302, plain, (k9_finseq_1(X1)=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,k9_finseq_1(X1),np__1,np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k9_finseq_1(X1))!=np__1|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k9_finseq_1(X1))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=np__1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_297, c_0_298]), c_0_112]), c_0_144])])])).
cnf(c_0_303, plain, (k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))=np__1|r1_xxreal_0(k1_matrix_1(k9_finseq_1(k9_finseq_1(X2))),np__0)|v1_xboole_0(X1)|~v1_matrix_1(k9_finseq_1(k9_finseq_1(X2)))|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_299, c_0_300])])).
fof(c_0_304, negated_conjecture, (~v1_xboole_0(esk1_0)&(m1_subset_1(esk2_0,esk1_0)&k2_matrlin(esk1_0,k3_pre_poly(esk1_0,esk2_0))!=k4_matrix_1(esk1_0,k2_matrlin(esk1_0,k3_pre_poly(esk1_0,esk2_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_301])])])])).
cnf(c_0_305, plain, (k9_finseq_1(X1)=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,k9_finseq_1(X1),np__1,np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k9_finseq_1(X1))!=np__1|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k9_finseq_1(X1))), inference(csr,[status(thm)],[c_0_302, c_0_280])).
cnf(c_0_306, plain, (k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))=np__1|r1_xxreal_0(k1_matrix_1(k9_finseq_1(k9_finseq_1(X2))),np__0)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_303, c_0_239])])])).
cnf(c_0_307, negated_conjecture, (k2_matrlin(esk1_0,k3_pre_poly(esk1_0,esk2_0))!=k4_matrix_1(esk1_0,k2_matrlin(esk1_0,k3_pre_poly(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_304])).
cnf(c_0_308, negated_conjecture, (m1_subset_1(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_304])).
cnf(c_0_309, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_304])).
cnf(c_0_310, plain, (k9_finseq_1(X1)=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,k9_finseq_1(X1),np__1,np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k9_finseq_1(X1))!=np__1|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k9_finseq_1(X1))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_159]), c_0_239])])])).
cnf(c_0_311, plain, (k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))=np__1|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_306, c_0_238]), c_0_123])).
cnf(c_0_312, negated_conjecture, (k4_matrix_1(esk1_0,k2_matrlin(esk1_0,k3_pre_poly(esk1_0,esk2_0)))!=k2_matrlin(esk1_0,k3_pre_poly(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[c_0_307])).
cnf(c_0_313, negated_conjecture, (m1_subset_1(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_308])).
cnf(c_0_314, negated_conjecture, (~v1_xboole_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_309])).
cnf(c_0_315, plain, (k9_finseq_1(X1)=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,k9_finseq_1(X1),np__1,np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k1_matrix_1(k9_finseq_1(X1))!=np__1|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k9_finseq_1(X1))|~m1_subset_1(X3,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_311])])).
cnf(c_0_316, negated_conjecture, (k4_matrix_1(esk1_0,k2_matrlin(esk1_0,k9_finseq_1(esk2_0)))!=k2_matrlin(esk1_0,k9_finseq_1(esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_199]), c_0_313])]), c_0_314])])).
cnf(c_0_317, plain, (k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X1)))=k9_finseq_1(k9_finseq_1(X1))|v1_xboole_0(X2)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X1)),k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X1))))|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_315]), c_0_238]), c_0_239])])])).
cnf(c_0_318, negated_conjecture, (k4_matrix_1(esk1_0,k9_finseq_1(k9_finseq_1(esk2_0)))!=k9_finseq_1(k9_finseq_1(esk2_0))|~m1_finseq_1(k9_finseq_1(esk2_0),esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_316, c_0_108]), c_0_314])])).
cnf(c_0_319, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k9_finseq_1(k9_finseq_1(X2))|v1_xboole_0(X1)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_317, c_0_176]), c_0_239])])])).
cnf(c_0_320, negated_conjecture, (~m1_finseq_1(k9_finseq_1(esk2_0),esk1_0)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(esk2_0)),k3_finseq_2(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_319]), c_0_313])]), c_0_314])])).
cnf(c_0_321, negated_conjecture, (~m1_finseq_1(k9_finseq_1(esk2_0),esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_320, c_0_163]), c_0_112]), c_0_144])]), c_0_314])])).
cnf(c_0_322, negated_conjecture, (~m2_finseq_1(k9_finseq_1(esk2_0),esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_321, c_0_140])])).
cnf(c_0_323, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_322, c_0_213]), c_0_313])]), c_0_314]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 324
# Proof object clause steps            : 234
# Proof object formula steps           : 90
# Proof object conjectures             : 15
# Proof object clause conjectures      : 12
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 60
# Proof object initial formulas used   : 47
# Proof object generating inferences   : 82
# Proof object simplifying inferences  : 134
# Training examples: 145 positive, 1362 negative
# Training: Positive examples begin
cnf(c_0_283, plain, (k6_numbers=k1_xboole_0)).# trainpos
cnf(c_0_73, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_84, plain, (v1_xboole_0(np__0))).# trainpos
cnf(c_0_72, plain, (v3_ordinal1(k4_ordinal1))).# trainpos
cnf(c_0_314, negated_conjecture, (~v1_xboole_0(esk1_0))).# trainpos
cnf(c_0_313, negated_conjecture, (m1_subset_1(esk2_0,esk1_0))).# trainpos
cnf(c_0_312, negated_conjecture, (k4_matrix_1(esk1_0,k2_matrlin(esk1_0,k3_pre_poly(esk1_0,esk2_0)))!=k2_matrlin(esk1_0,k3_pre_poly(esk1_0,esk2_0)))).# trainpos
cnf(c_0_156, plain, (k2_finseq_1(np__1)=k1_tarski(np__1))).# trainpos
cnf(c_0_98, plain, (v1_relat_1(k9_finseq_1(X1)))).# trainpos
cnf(c_0_295, plain, (~v1_xboole_0(k3_finseq_2(X1)))).# trainpos
cnf(c_0_232, plain, (X1=k3_finseq_1(X2)|k9_xtuple_0(X2)!=k2_finseq_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_83, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_123, plain, (~r1_xxreal_0(np__1,np__0))).# trainpos
cnf(c_0_97, plain, (v1_funct_1(k9_finseq_1(X1)))).# trainpos
cnf(c_0_96, plain, (v1_finseq_1(k9_finseq_1(X1)))).# trainpos
cnf(c_0_91, plain, (m1_subset_1(np__1,k5_numbers))).# trainpos
cnf(c_0_229, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_215, plain, (r1_tarski(X1,X1))).# trainpos
cnf(c_0_184, plain, (m1_finseq_2(k3_finseq_2(X1),X1))).# trainpos
cnf(c_0_171, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1))).# trainpos
cnf(c_0_74, plain, (v1_finset_1(k1_card_1(X1))|~v1_finset_1(X1))).# trainpos
cnf(c_0_140, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_193, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2))).# trainpos
cnf(c_0_208, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_209, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_109, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1))).# trainpos
cnf(c_0_94, plain, (v1_finset_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_256, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2))).# trainpos
cnf(c_0_243, plain, (X1=X3|X2!=k1_tarski(X3)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_214, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_90, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_154, plain, (k4_finseq_1(X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_75, plain, (k1_card_1(X1)=k3_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_199, plain, (k3_pre_poly(X1,X2)=k9_finseq_1(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_108, plain, (k2_matrlin(X1,X2)=k9_finseq_1(X2)|v1_xboole_0(X1)|~m1_finseq_1(X2,X1))).# trainpos
cnf(c_0_291, plain, (m1_subset_1(X1,X3)|~m1_finseq_2(X3,X2)|~m2_finseq_2(X1,X2,X3))).# trainpos
cnf(c_0_182, plain, (m2_finseq_1(X3,X2)|~m1_finseq_2(X1,X2)|~m2_finseq_2(X3,X2,X1))).# trainpos
cnf(c_0_95, plain, (k3_finseq_1(X1)=np__1|X1!=k9_finseq_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_135, plain, (m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_183, plain, (v1_xboole_0(X1)|m2_finseq_2(k3_pre_poly(X1,X2),X1,k3_finseq_2(X1))|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_170, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k2_tarski(X3,X1),k1_tarski(X3)),k2_zfmisc_1(X4,X2)))).# trainpos
cnf(c_0_257, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_zfmisc_1(X2,X4)))).# trainpos
cnf(c_0_160, plain, (v1_xboole_0(X1)|v1_matrix_1(k4_matrix_1(X1,X2))|~v1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(X1)))).# trainpos
cnf(c_0_107, plain, (v1_xboole_0(X1)|m1_matrix_1(k2_matrlin(X1,X2),X1,np__1,k3_finseq_1(X2))|~m1_finseq_1(X2,X1))).# trainpos
cnf(c_0_139, plain, (v1_xboole_0(X1)|m2_finseq_1(k4_matrix_1(X1,X2),k3_finseq_2(X1))|~v1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(X1)))).# trainpos
cnf(c_0_145, plain, (v1_xboole_0(X2)|v1_matrix_1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4))).# trainpos
cnf(c_0_281, plain, (X1=k9_finseq_1(X2)|k3_finseq_1(X1)!=np__1|k1_funct_1(X1,np__1)!=X2|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_247, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_zfmisc_1(X2,X4))|~r2_hidden(X3,X4)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_117, plain, (k9_xtuple_0(X2)=k2_finseq_1(X1)|X1!=k3_finseq_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_153, plain, (k2_zfmisc_1(k4_finseq_1(X1),k2_finseq_1(k1_matrix_1(X1)))=k2_matrix_1(X1)|~v1_matrix_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_260, plain, (X2=X3|v1_xboole_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|k3_matrix_1(X1,X2,esk8_3(X1,X2,X3),esk9_3(X1,X2,X3))!=k3_matrix_1(X1,X3,esk8_3(X1,X2,X3),esk9_3(X1,X2,X3))|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).# trainpos
cnf(c_0_124, plain, (v1_xboole_0(X2)|m2_finseq_1(X1,k3_finseq_2(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4))).# trainpos
cnf(c_0_158, plain, (k1_matrix_1(X2)=k3_finseq_1(X1)|v1_xboole_0(X3)|X1!=k4_matrix_1(X3,X2)|~v1_matrix_1(X2)|~v1_matrix_1(X1)|~m2_finseq_1(X2,k3_finseq_2(X3))|~m2_finseq_1(X1,k3_finseq_2(X3)))).# trainpos
cnf(c_0_282, plain, (k1_matrix_1(k4_matrix_1(X1,X2))=k3_finseq_1(X2)|v1_xboole_0(X1)|r1_xxreal_0(k1_matrix_1(X2),k6_numbers)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1)))).# trainpos
cnf(c_0_103, plain, (X3=k1_matrix_1(X1)|v1_xboole_0(X2)|r1_xxreal_0(k3_finseq_1(X1),k1_xboole_0)|~v1_matrix_1(X1)|~v7_ordinal1(X3)|~m2_finseq_1(X1,k3_finseq_2(X2))|~m1_matrix_1(X1,X2,k3_finseq_1(X1),X3))).# trainpos
cnf(c_0_191, plain, (X2=X3|v1_xboole_0(X1)|r2_hidden(k2_tarski(k2_tarski(esk8_3(X1,X2,X3),esk9_3(X1,X2,X3)),k1_tarski(esk8_3(X1,X2,X3))),k2_matrix_1(X2))|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).# trainpos
cnf(c_0_236, plain, (k3_matrix_1(X4,X3,X2,X1)=k3_matrix_1(X4,X5,X1,X2)|v1_xboole_0(X4)|X5!=k4_matrix_1(X4,X3)|~v1_matrix_1(X5)|~v1_matrix_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_finseq_1(X5,k3_finseq_2(X4))|~m2_finseq_1(X3,k3_finseq_2(X4))|~r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X2)),k2_matrix_1(X3)))).# trainpos
cnf(c_0_316, negated_conjecture, (k4_matrix_1(esk1_0,k2_matrlin(esk1_0,k9_finseq_1(esk2_0)))!=k2_matrlin(esk1_0,k9_finseq_1(esk2_0)))).# trainpos
cnf(c_0_112, plain, (k3_finseq_1(k9_finseq_1(X1))=np__1)).# trainpos
cnf(c_0_207, plain, (v1_relat_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_263, plain, (r2_hidden(X1,X2)|X2!=k1_tarski(X1))).# trainpos
cnf(c_0_104, plain, (k1_xboole_0=np__0)).# trainpos
cnf(c_0_159, plain, (m2_finseq_1(k4_matrix_1(X1,X2),k3_finseq_2(X1))|v1_xboole_0(X1)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1)))).# trainpos
cnf(c_0_198, plain, (m2_finseq_1(k3_pre_poly(X1,X2),X1)|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_255, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3))).# trainpos
cnf(c_0_288, plain, (k9_finseq_1(k1_funct_1(X1,np__1))=X1|k3_finseq_1(X1)!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_296, plain, (m1_subset_1(k3_pre_poly(X1,X2),k3_finseq_2(X1))|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_189, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X3,X1)),k2_zfmisc_1(X4,X2)))).# trainpos
cnf(c_0_134, plain, (k2_finseq_1(k3_finseq_1(X1))=k9_xtuple_0(X1)|~m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_110, plain, (v3_ordinal1(np__1))).# trainpos
cnf(c_0_176, plain, (v1_matrix_1(k4_matrix_1(X1,X2))|v1_xboole_0(X1)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1)))).# trainpos
cnf(c_0_264, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X3)),k2_zfmisc_1(X2,X4)))).# trainpos
cnf(c_0_172, plain, (k2_zfmisc_1(k9_xtuple_0(X1),k2_finseq_1(k1_matrix_1(X1)))=k2_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).# trainpos
cnf(c_0_228, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_93, plain, (v1_finset_1(k3_finseq_1(X1))|~v1_finset_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_125, plain, (m1_matrix_1(k9_finseq_1(X2),X1,np__1,k3_finseq_1(X2))|v1_xboole_0(X1)|~m1_finseq_1(X2,X1))).# trainpos
cnf(c_0_245, plain, (k3_matrix_1(X1,X2,X3,X4)=k3_matrix_1(X1,X5,X4,X3)|v1_xboole_0(X1)|X5!=k4_matrix_1(X1,X2)|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X3,X4)),k2_matrix_1(X2))|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m2_finseq_1(X5,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X5)|~v1_matrix_1(X2))).# trainpos
cnf(c_0_269, plain, (r2_hidden(X1,k1_tarski(X1)))).# trainpos
cnf(c_0_175, plain, (k1_matrix_1(X1)=k3_finseq_1(k4_matrix_1(X2,X3))|v1_xboole_0(X2)|k4_matrix_1(X2,X3)!=k4_matrix_1(X2,X1)|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(X1)|~v1_matrix_1(k4_matrix_1(X2,X3))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3))).# trainpos
cnf(c_0_150, plain, (m1_matrix_1(k2_matrlin(X1,k9_finseq_1(X2)),X1,np__1,np__1)|v1_xboole_0(X1)|~m1_finseq_1(k9_finseq_1(X2),X1))).# trainpos
cnf(c_0_126, plain, (v7_ordinal1(np__1)|~v1_finset_1(np__1))).# trainpos
cnf(c_0_294, plain, (k1_matrix_1(k4_matrix_1(X1,X2))=k3_finseq_1(X2)|r1_xxreal_0(k1_matrix_1(X2),np__0)|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).# trainpos
cnf(c_0_127, plain, (v1_finset_1(np__1))).# trainpos
cnf(c_0_300, plain, (m1_subset_1(k9_finseq_1(X2),k3_finseq_2(X1))|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_213, plain, (m2_finseq_1(k9_finseq_1(X2),X1)|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_173, plain, (k9_xtuple_0(k9_finseq_1(X1))=k1_tarski(np__1))).# trainpos
cnf(c_0_318, negated_conjecture, (k4_matrix_1(esk1_0,k9_finseq_1(k9_finseq_1(esk2_0)))!=k9_finseq_1(k9_finseq_1(esk2_0))|~m1_finseq_1(k9_finseq_1(esk2_0),esk1_0))).# trainpos
cnf(c_0_122, plain, (X1=k1_matrix_1(X2)|r1_xxreal_0(k3_finseq_1(X2),np__0)|v1_xboole_0(X3)|~m1_matrix_1(X2,X3,k3_finseq_1(X2),X1)|~v7_ordinal1(X1)|~m2_finseq_1(X2,k3_finseq_2(X3))|~v1_matrix_1(X2))).# trainpos
cnf(c_0_242, plain, (k3_finseq_1(X1)=k3_finseq_1(X2)|k9_xtuple_0(X2)!=k9_xtuple_0(X1)|~m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_179, plain, (k9_xtuple_0(k9_finseq_1(X1))=k2_finseq_1(X2)|X2!=np__1|~m2_subset_1(X2,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_224, plain, (v1_finseq_1(k4_matrix_1(X1,X2))|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).# trainpos
cnf(c_0_143, plain, (m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|v1_xboole_0(X2)|~v7_ordinal1(k3_finseq_1(X1))|~v7_ordinal1(np__1)|~m1_finseq_1(X1,X2))).# trainpos
cnf(c_0_223, plain, (v1_funct_1(k4_matrix_1(X1,X2))|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).# trainpos
cnf(c_0_164, plain, (v1_matrix_1(k9_finseq_1(X1))|v1_xboole_0(X2)|~v7_ordinal1(k3_finseq_1(X1))|~v7_ordinal1(np__1)|~m1_finseq_1(X1,X2))).# trainpos
cnf(c_0_222, plain, (v1_relat_1(k4_matrix_1(X1,X2))|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).# trainpos
cnf(c_0_206, plain, (k3_finseq_1(k4_matrix_1(X2,X1))=k1_matrix_1(X1)|v1_xboole_0(X2)|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(X1))).# trainpos
cnf(c_0_142, plain, (X1=k1_matrix_1(k9_finseq_1(X2))|v1_xboole_0(X3)|~m1_matrix_1(k9_finseq_1(X2),X3,np__1,X1)|~v7_ordinal1(X1)|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X3))|~v1_matrix_1(k9_finseq_1(X2)))).# trainpos
cnf(c_0_144, plain, (v7_ordinal1(np__1))).# trainpos
cnf(c_0_299, plain, (k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))=np__1|r1_xxreal_0(k1_matrix_1(k9_finseq_1(X2)),np__0)|v1_xboole_0(X1)|~v1_matrix_1(k9_finseq_1(X2))|~m1_subset_1(X2,k3_finseq_2(X1)))).# trainpos
cnf(c_0_169, plain, (v1_matrix_1(k2_matrlin(X1,k9_finseq_1(X2)))|v1_xboole_0(X1)|~v7_ordinal1(np__1)|~m1_finseq_1(k9_finseq_1(X2),X1))).# trainpos
cnf(c_0_190, plain, (k2_zfmisc_1(k1_tarski(np__1),k2_finseq_1(k1_matrix_1(k9_finseq_1(X1))))=k2_matrix_1(k9_finseq_1(X1))|~v1_matrix_1(k9_finseq_1(X1)))).# trainpos
cnf(c_0_211, plain, (k9_xtuple_0(X1)=k1_tarski(np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__1|~m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers))).# trainpos
cnf(c_0_258, plain, (k3_finseq_1(X2)=np__1|k9_xtuple_0(X2)!=k1_tarski(np__1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).# trainpos
cnf(c_0_162, plain, (k1_matrix_1(k9_finseq_1(X1))=k3_finseq_1(X1)|v1_xboole_0(X2)|~v7_ordinal1(k3_finseq_1(X1))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~v1_matrix_1(k9_finseq_1(X1))|~m1_finseq_1(X1,X2))).# trainpos
cnf(c_0_181, plain, (v1_matrix_1(k9_finseq_1(X1))|v1_xboole_0(X2)|~m1_finseq_1(X1,X2)|~v7_ordinal1(k3_finseq_1(X1)))).# trainpos
cnf(c_0_270, plain, (r2_hidden(X1,k1_tarski(np__1))|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_matrix_1(k9_finseq_1(X3)))|~v1_matrix_1(k9_finseq_1(X3)))).# trainpos
cnf(c_0_203, plain, (v1_matrix_1(k9_finseq_1(k9_finseq_1(X2)))|v1_xboole_0(X1)|~m1_finseq_1(k9_finseq_1(X2),X1))).# trainpos
cnf(c_0_303, plain, (k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))=np__1|r1_xxreal_0(k1_matrix_1(k9_finseq_1(k9_finseq_1(X2))),np__0)|v1_xboole_0(X1)|~v1_matrix_1(k9_finseq_1(k9_finseq_1(X2)))|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_204, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(k9_finseq_1(X2))))|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X3,X1)),k2_matrix_1(k9_finseq_1(X2)))|~v1_matrix_1(k9_finseq_1(X2)))).# trainpos
cnf(c_0_163, plain, (m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|v1_xboole_0(X2)|~m1_finseq_1(X1,X2)|~v7_ordinal1(k3_finseq_1(X1)))).# trainpos
cnf(c_0_251, plain, (k9_xtuple_0(k4_matrix_1(X1,X2))=k2_finseq_1(k1_matrix_1(X2))|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X1,X2))|~v1_funct_1(k4_matrix_1(X1,X2))|~v1_relat_1(k4_matrix_1(X1,X2))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).# trainpos
cnf(c_0_237, plain, (k2_zfmisc_1(k1_tarski(np__1),k2_finseq_1(k1_matrix_1(X1)))=k2_matrix_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1)|k3_finseq_1(X1)!=np__1)).# trainpos
cnf(c_0_221, plain, (m2_subset_1(k1_matrix_1(X2),k1_numbers,k5_numbers)|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X1,X2))|~v1_funct_1(k4_matrix_1(X1,X2))|~v1_relat_1(k4_matrix_1(X1,X2))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).# trainpos
cnf(c_0_217, plain, (v1_matrix_1(k9_finseq_1(k9_finseq_1(X1)))|v1_xboole_0(X2)|~m2_finseq_1(k9_finseq_1(X1),X2))).# trainpos
cnf(c_0_279, plain, (r2_hidden(esk8_3(X1,k9_finseq_1(X2),X3),k1_tarski(np__1))|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|~v1_matrix_1(k9_finseq_1(X2))|k3_finseq_1(X3)!=np__1|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~v1_matrix_1(X3))).# trainpos
cnf(c_0_180, plain, (k1_matrix_1(k9_finseq_1(X1))=k3_finseq_1(X1)|v1_xboole_0(X2)|~m1_finseq_1(X1,X2)|~v7_ordinal1(k3_finseq_1(X1))|~v1_matrix_1(k9_finseq_1(X1)))).# trainpos
cnf(c_0_220, plain, (r2_hidden(esk9_3(X1,k9_finseq_1(X2),X3),k2_finseq_1(k1_matrix_1(k9_finseq_1(X2))))|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|~v1_matrix_1(k9_finseq_1(X2))|k3_finseq_1(X3)!=np__1|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~v1_matrix_1(X3))).# trainpos
cnf(c_0_265, plain, (k3_finseq_1(k4_matrix_1(X1,X2))=np__1|v1_xboole_0(X1)|k2_finseq_1(k1_matrix_1(X2))!=k1_tarski(np__1)|~v1_finseq_1(k4_matrix_1(X1,X2))|~v1_funct_1(k4_matrix_1(X1,X2))|~v1_relat_1(k4_matrix_1(X1,X2))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).# trainpos
cnf(c_0_212, plain, (k1_matrix_1(k9_finseq_1(X1))=k3_finseq_1(X1)|v1_xboole_0(X2)|~v7_ordinal1(k3_finseq_1(X1))|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_293, plain, (esk8_3(X1,k9_finseq_1(X2),X3)=X4|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|k1_tarski(np__1)!=k1_tarski(X4)|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k3_finseq_1(X3)!=np__1|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3))).# trainpos
cnf(c_0_230, plain, (v1_matrix_1(k9_finseq_1(k9_finseq_1(X1)))|v1_xboole_0(X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_234, plain, (k9_finseq_1(X1)=X2|r2_hidden(esk9_3(X3,k9_finseq_1(X1),X2),k1_tarski(np__1))|v1_xboole_0(X3)|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(X2)|k3_finseq_1(X2)!=np__1|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X3))|~m2_finseq_1(X2,k3_finseq_2(X3))|~v1_matrix_1(k9_finseq_1(X1))|~v1_matrix_1(X2)|k1_matrix_1(k9_finseq_1(X1))!=np__1|~m2_subset_1(k1_matrix_1(k9_finseq_1(X1)),k1_numbers,k5_numbers))).# trainpos
cnf(c_0_239, plain, (v1_matrix_1(k9_finseq_1(k9_finseq_1(X1))))).# trainpos
cnf(c_0_298, plain, (esk8_3(X1,k9_finseq_1(X2),X3)=np__1|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k3_finseq_1(X3)!=np__1|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3))).# trainpos
cnf(c_0_253, plain, (esk9_3(X1,k9_finseq_1(X2),X3)=X4|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|k1_tarski(np__1)!=k1_tarski(X4)|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k1_matrix_1(k9_finseq_1(X2))!=np__1|k3_finseq_1(X3)!=np__1|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3))).# trainpos
cnf(c_0_227, plain, (k1_matrix_1(k9_finseq_1(k9_finseq_1(X1)))=np__1|v1_xboole_0(X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_238, plain, (k1_matrix_1(k9_finseq_1(k9_finseq_1(X1)))=np__1)).# trainpos
cnf(c_0_261, plain, (esk9_3(X1,k9_finseq_1(X2),X3)=np__1|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k1_matrix_1(k9_finseq_1(X2))!=np__1|k3_finseq_1(X3)!=np__1|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3))).# trainpos
cnf(c_0_306, plain, (k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))=np__1|r1_xxreal_0(k1_matrix_1(k9_finseq_1(k9_finseq_1(X2))),np__0)|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_246, plain, (k2_matrix_1(k9_finseq_1(k9_finseq_1(X1)))=k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1)))).# trainpos
cnf(c_0_254, plain, (k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),X3,X4)=k3_matrix_1(X1,X5,X4,X3)|v1_xboole_0(X1)|X5!=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X3,X4)),k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1)))|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m2_finseq_1(X5,k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~v1_matrix_1(X5))).# trainpos
cnf(c_0_280, plain, (k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))=np__1|v1_xboole_0(X1)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1)))).# trainpos
cnf(c_0_262, plain, (k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),X3,X4)=k3_matrix_1(X1,X5,X4,X3)|v1_xboole_0(X1)|X5!=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X5,k3_finseq_2(X1))|~v1_matrix_1(X5)|~r2_hidden(X4,k1_tarski(np__1))|~r2_hidden(X3,k1_tarski(np__1)))).# trainpos
cnf(c_0_267, plain, (k9_finseq_1(X1)=X2|v1_xboole_0(X3)|k3_matrix_1(X3,k9_finseq_1(X1),esk8_3(X3,k9_finseq_1(X1),X2),np__1)!=k3_matrix_1(X3,X2,esk8_3(X3,k9_finseq_1(X1),X2),np__1)|k3_finseq_1(X2)!=np__1|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X3))|~v1_matrix_1(X2)|~v1_matrix_1(k9_finseq_1(X1))|k1_matrix_1(k9_finseq_1(X1))!=np__1)).# trainpos
cnf(c_0_268, plain, (k3_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X4,X3)=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),X3,X4)|v1_xboole_0(X1)|~r2_hidden(X4,k1_tarski(np__1))|~r2_hidden(X3,k1_tarski(np__1))|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))))).# trainpos
cnf(c_0_278, plain, (k9_finseq_1(X1)=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,k9_finseq_1(X1),esk8_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,esk8_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k9_finseq_1(X1))!=np__1|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=np__1|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~v1_matrix_1(k9_finseq_1(X1))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~r2_hidden(esk8_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~v7_ordinal1(esk8_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2)))).# trainpos
cnf(c_0_292, plain, (X1=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,X1,esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|k1_matrix_1(X1)!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(X1)!=np__1|~v7_ordinal1(esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X1)|k3_finseq_1(X1)!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_302, plain, (k9_finseq_1(X1)=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,k9_finseq_1(X1),np__1,np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k9_finseq_1(X1))!=np__1|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k9_finseq_1(X1))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=np__1)).# trainpos
cnf(c_0_310, plain, (k9_finseq_1(X1)=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,k9_finseq_1(X1),np__1,np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k9_finseq_1(X1))!=np__1|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k9_finseq_1(X1)))).# trainpos
cnf(c_0_315, plain, (k9_finseq_1(X1)=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,k9_finseq_1(X1),np__1,np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k1_matrix_1(k9_finseq_1(X1))!=np__1|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k9_finseq_1(X1))|~m1_subset_1(X3,X2))).# trainpos
cnf(c_0_317, plain, (k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X1)))=k9_finseq_1(k9_finseq_1(X1))|v1_xboole_0(X2)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X1)),k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X1))))|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_319, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k9_finseq_1(k9_finseq_1(X2))|v1_xboole_0(X1)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_320, negated_conjecture, (~m1_finseq_1(k9_finseq_1(esk2_0),esk1_0)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(esk2_0)),k3_finseq_2(esk1_0)))).# trainpos
cnf(c_0_321, negated_conjecture, (~m1_finseq_1(k9_finseq_1(esk2_0),esk1_0))).# trainpos
cnf(c_0_322, negated_conjecture, (~m2_finseq_1(k9_finseq_1(esk2_0),esk1_0))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_105, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_182, plain, (v1_xboole_0(esk15_0))).#trainneg
cnf(i_0_215, plain, (v1_xboole_0(esk18_0))).#trainneg
cnf(i_0_249, plain, (v1_xboole_0(esk23_0))).#trainneg
cnf(i_0_247, plain, (v7_ordinal1(esk22_0))).#trainneg
cnf(i_0_275, plain, (v3_ordinal1(esk28_0))).#trainneg
cnf(i_0_180, plain, (v1_xxreal_0(esk15_0))).#trainneg
cnf(i_0_190, plain, (v1_xxreal_0(esk17_0))).#trainneg
cnf(i_0_214, plain, (v1_xxreal_0(esk18_0))).#trainneg
cnf(i_0_258, plain, (v1_xxreal_0(esk25_0))).#trainneg
cnf(i_0_125, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_74, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_248, plain, (~v1_xboole_0(esk22_0))).#trainneg
cnf(i_0_172, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_165, plain, (v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_179, plain, (v1_xreal_0(esk15_0))).#trainneg
cnf(i_0_256, plain, (v1_xreal_0(esk25_0))).#trainneg
cnf(i_0_126, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_156, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_181, plain, (v1_xcmplx_0(esk15_0))).#trainneg
cnf(i_0_259, plain, (v1_xcmplx_0(esk25_0))).#trainneg
cnf(i_0_189, plain, (v3_xxreal_0(esk17_0))).#trainneg
cnf(i_0_257, plain, (v3_xxreal_0(esk25_0))).#trainneg
cnf(i_0_276, plain, (~v1_xboole_0(esk28_0))).#trainneg
cnf(i_0_216, plain, (~v1_finset_1(k4_ordinal1))).#trainneg
cnf(i_0_97, plain, (v1_xboole_0(esk10_1(X1)))).#trainneg
cnf(i_0_138, plain, (v1_xboole_0(esk12_1(X1)))).#trainneg
cnf(i_0_166, plain, (~v1_xboole_0(k9_finseq_1(X1)))).#trainneg
cnf(i_0_129, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_155, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_168, plain, (~r1_xxreal_0(np__2,np__1))).#trainneg
cnf(i_0_147, plain, (v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_184, plain, (v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_142, plain, (v1_relat_1(esk12_1(X1)))).#trainneg
cnf(i_0_139, plain, (v1_funct_1(esk12_1(X1)))).#trainneg
cnf(i_0_136, plain, (v1_finseq_1(esk12_1(X1)))).#trainneg
cnf(i_0_268, plain, (v1_finset_1(k1_tarski(X1)))).#trainneg
cnf(i_0_137, plain, (v1_finset_1(esk12_1(X1)))).#trainneg
cnf(i_0_94, plain, (v1_card_1(k1_card_1(X1)))).#trainneg
cnf(i_0_274, plain, (v4_funct_1(k3_finseq_2(X1)))).#trainneg
cnf(i_0_123, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_153, plain, (m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_87, plain, (~v1_xboole_0(k1_tarski(X1)))).#trainneg
cnf(i_0_154, plain, (m1_subset_1(np__2,k5_numbers))).#trainneg
cnf(i_0_213, plain, (~v1_xboole_0(k2_tarski(X1,X2)))).#trainneg
cnf(i_0_110, plain, (m1_subset_1(np__0,k1_numbers))).#trainneg
cnf(i_0_112, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_133, plain, (~r1_xxreal_0(np__2,np__0))).#trainneg
cnf(i_0_111, plain, (m1_subset_1(np__0,k5_numbers))).#trainneg
cnf(i_0_194, plain, (v3_ordinal1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_167, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_120, plain, (r1_xxreal_0(np__1,np__1))).#trainneg
cnf(i_0_192, plain, (r1_xxreal_0(np__1,np__2))).#trainneg
cnf(i_0_255, plain, (r1_xxreal_0(np__2,np__2))).#trainneg
cnf(i_0_197, plain, (r1_xxreal_0(np__0,np__1))).#trainneg
cnf(i_0_253, plain, (r1_xxreal_0(np__0,np__2))).#trainneg
cnf(i_0_151, plain, (r1_xxreal_0(np__0,np__0))).#trainneg
cnf(i_0_41, plain, (k2_finseq_1(np__2)=k2_tarski(np__1,np__2))).#trainneg
cnf(i_0_114, plain, (k1_card_1(k1_card_1(X1))=k1_card_1(X1))).#trainneg
cnf(i_0_185, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_277, plain, (m1_subset_1(esk28_0,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_141, plain, (v4_relat_1(esk12_1(X1),k5_numbers))).#trainneg
cnf(i_0_108, plain, (m1_subset_1(esk11_1(X1),X1))).#trainneg
cnf(i_0_143, plain, (m1_finseq_1(esk12_1(X1),X1))).#trainneg
cnf(i_0_100, plain, (v1_xboole_0(X1)|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_169, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk13_1(X1)))).#trainneg
cnf(i_0_242, plain, (v1_xboole_0(X1)|v1_relat_1(esk21_1(X1)))).#trainneg
cnf(i_0_231, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk20_1(X1)))).#trainneg
cnf(i_0_200, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_196, plain, (~v1_xboole_0(X1)|~v3_card_1(X1,np__1))).#trainneg
cnf(i_0_187, plain, (~v1_subset_1(esk16_1(X1),X1))).#trainneg
cnf(i_0_116, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_191, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_186, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_239, plain, (v1_xboole_0(X1)|v1_funct_1(esk21_1(X1)))).#trainneg
cnf(i_0_236, plain, (v1_xboole_0(X1)|v1_finseq_1(esk21_1(X1)))).#trainneg
cnf(i_0_140, plain, (v5_relat_1(esk12_1(X1),X1))).#trainneg
cnf(i_0_176, plain, (v1_xboole_0(esk14_2(X1,X2)))).#trainneg
cnf(i_0_118, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_175, plain, (v1_relat_1(esk14_2(X1,X2)))).#trainneg
cnf(i_0_267, plain, (v1_finset_1(k2_tarski(X1,X2)))).#trainneg
cnf(i_0_98, plain, (m1_subset_1(esk10_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_188, plain, (m1_subset_1(esk16_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_173, plain, (v5_relat_1(esk14_2(X1,X2),X2))).#trainneg
cnf(i_0_174, plain, (v4_relat_1(esk14_2(X1,X2),X1))).#trainneg
cnf(i_0_237, plain, (v1_xboole_0(X1)|v1_finset_1(esk21_1(X1)))).#trainneg
cnf(i_0_238, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk21_1(X1)))).#trainneg
cnf(i_0_262, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk26_1(X1)))).#trainneg
cnf(i_0_261, plain, (v1_xboole_0(X1)|v1_finset_1(esk26_1(X1)))).#trainneg
cnf(i_0_157, plain, (v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(k9_xtuple_0(X1)))).#trainneg
cnf(i_0_217, plain, (v1_xboole_0(X1)|~v3_card_1(X1,k1_xboole_0))).#trainneg
cnf(i_0_227, plain, (~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_225, plain, (~v1_xxreal_0(X1)|~v2_xxreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_269, plain, (v1_finseq_1(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_230, plain, (v1_xboole_0(X1)|v1_zfmisc_1(esk20_1(X1)))).#trainneg
cnf(i_0_235, plain, (v1_xboole_0(X1)|v2_finseq_1(esk21_1(X1)))).#trainneg
cnf(i_0_77, plain, (k1_finseq_1(X1)=k2_finseq_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_132, plain, (v1_xboole_0(k9_xtuple_0(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_93, plain, (v1_xboole_0(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_224, plain, (v1_relat_1(esk19_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_266, plain, (v1_relat_1(esk27_1(X1))|~v1_card_1(X1))).#trainneg
cnf(i_0_222, plain, (v1_funct_1(esk19_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_265, plain, (v1_funct_1(esk27_1(X1))|~v1_card_1(X1))).#trainneg
cnf(i_0_219, plain, (v1_finseq_1(esk19_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_252, plain, (v1_finset_1(k1_finseq_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_221, plain, (v1_finset_1(esk19_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_91, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_183, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_241, plain, (v1_xboole_0(X1)|v4_relat_1(esk21_1(X1),k5_numbers))).#trainneg
cnf(i_0_95, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_128, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_145, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_159, plain, (v1_finset_1(X1)|~v1_finset_1(k1_card_1(X1)))).#trainneg
cnf(i_0_84, plain, (v5_relat_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_218, plain, (v2_finseq_1(esk19_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_171, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_115, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_150, plain, (v1_card_3(X1)|~m1_finseq_1(X1,k5_numbers))).#trainneg
cnf(i_0_195, plain, (v1_zfmisc_1(X1)|~v3_card_1(X1,np__1))).#trainneg
cnf(i_0_61, plain, (v1_funct_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_60, plain, (v1_finseq_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_250, plain, (v1_xboole_0(X1)|v3_card_1(esk24_1(X1),np__1))).#trainneg
cnf(i_0_243, plain, (v1_xboole_0(X1)|m1_finseq_1(esk21_1(X1),X1))).#trainneg
cnf(i_0_240, plain, (v1_xboole_0(X1)|v5_relat_1(esk21_1(X1),X1))).#trainneg
cnf(i_0_164, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).#trainneg
cnf(i_0_229, plain, (v1_relat_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_223, plain, (v4_relat_1(esk19_1(X1),k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_121, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_203, plain, (v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_193, plain, (m1_subset_1(k2_finseq_1(X1),k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_212, plain, (v1_xboole_0(k1_finseq_1(X1))|~v1_xboole_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_53, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_152, plain, (v3_card_1(k1_finseq_1(X1),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_220, plain, (v3_card_1(esk19_1(X1),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_264, plain, (v3_card_1(esk27_1(X1),X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_49, plain, (m2_finseq_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_170, plain, (v1_xboole_0(X1)|m1_subset_1(esk13_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_232, plain, (v1_xboole_0(X1)|m1_subset_1(esk20_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_251, plain, (v1_xboole_0(X1)|m1_subset_1(esk24_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_263, plain, (v1_xboole_0(X1)|m1_subset_1(esk26_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_76, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_177, plain, (m1_subset_1(esk14_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_178, plain, (v1_relat_1(k1_tarski(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)))))).#trainneg
cnf(i_0_210, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_228, plain, (v1_funct_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_254, plain, (v3_ordinal1(k9_xtuple_0(X1))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v5_ordinal1(X1))).#trainneg
cnf(i_0_122, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_96, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_130, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_88, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_57, plain, (m2_finseq_2(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_finseq_2(X2,X3))).#trainneg
cnf(i_0_198, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_134, plain, (v1_funct_1(k1_tarski(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)))))).#trainneg
cnf(i_0_244, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_45, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_103, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_104, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_16, plain, (X2=k1_tarski(X1)|esk3_2(X1,X2)=X1|r2_hidden(esk3_2(X1,X2),X2))).#trainneg
cnf(i_0_127, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_17, plain, (X2=k1_tarski(X1)|esk3_2(X1,X2)!=X1|~r2_hidden(esk3_2(X1,X2),X2))).#trainneg
cnf(i_0_260, plain, (v1_funct_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_109, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_160, plain, (m1_subset_1(k4_finseq_1(X1),k1_zfmisc_1(k5_numbers))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_211, plain, (k3_finseq_1(k3_finseq_1(X1))=k3_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_89, plain, (m1_subset_1(k1_matrix_1(X1),k5_numbers)|~v1_matrix_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_233, plain, (v2_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_246, plain, (v1_xboole_0(X2)|v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_7, plain, (k1_funct_1(X1,np__1)=X2|X1!=k9_finseq_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_273, plain, (v5_relat_1(X3,X2)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_278, plain, (v3_card_1(k9_xtuple_0(X2),X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_card_1(X1)|~v3_card_1(X2,X1))).#trainneg
cnf(i_0_245, plain, (v1_xboole_0(X1)|v2_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_107, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_10, plain, (k3_finseq_1(k4_matrix_1(X1,X2))=k1_matrix_1(X2)|v1_xboole_0(X1)|r1_xxreal_0(k1_matrix_1(X2),k6_numbers)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_161, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_106, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_86, plain, (m1_subset_1(k3_matrix_1(X2,X1,X3,X4),X2)|~v1_matrix_1(X1)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_finseq_1(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_81, plain, (X2=X3|v1_xboole_0(X1)|v7_ordinal1(esk8_3(X1,X2,X3))|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_4, plain, (v1_xboole_0(X3)|r1_xxreal_0(k3_finseq_1(X2),k1_xboole_0)|m1_matrix_1(X2,X3,k3_finseq_1(X2),X1)|X1!=k1_matrix_1(X2)|~v1_matrix_1(X2)|~v7_ordinal1(X1)|~m2_finseq_1(X2,k3_finseq_2(X3)))).#trainneg
cnf(i_0_80, plain, (X2=X3|v1_xboole_0(X1)|v7_ordinal1(esk9_3(X1,X2,X3))|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_35, plain, (X3=k4_matrix_1(X1,X2)|v1_xboole_0(X1)|v7_ordinal1(esk4_3(X1,X2,X3))|v7_ordinal1(esk6_3(X1,X2,X3))|k1_matrix_1(X2)!=k3_finseq_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_34, plain, (X3=k4_matrix_1(X1,X2)|v1_xboole_0(X1)|v7_ordinal1(esk4_3(X1,X2,X3))|v7_ordinal1(esk7_3(X1,X2,X3))|k1_matrix_1(X2)!=k3_finseq_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_31, plain, (X3=k4_matrix_1(X1,X2)|v1_xboole_0(X1)|v7_ordinal1(esk5_3(X1,X2,X3))|v7_ordinal1(esk6_3(X1,X2,X3))|k1_matrix_1(X2)!=k3_finseq_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_30, plain, (X3=k4_matrix_1(X1,X2)|v1_xboole_0(X1)|v7_ordinal1(esk5_3(X1,X2,X3))|v7_ordinal1(esk7_3(X1,X2,X3))|k1_matrix_1(X2)!=k3_finseq_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_37, plain, (v1_xboole_0(X5)|r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X2)),k2_matrix_1(X4))|X4!=k4_matrix_1(X5,X3)|~v1_matrix_1(X4)|~v1_matrix_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_finseq_1(X4,k3_finseq_2(X5))|~m2_finseq_1(X3,k3_finseq_2(X5))|~r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3)))).#trainneg
cnf(i_0_38, plain, (v1_xboole_0(X5)|r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X2)),k2_matrix_1(X4))|X3!=k4_matrix_1(X5,X4)|~v1_matrix_1(X4)|~v1_matrix_1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_finseq_1(X4,k3_finseq_2(X5))|~m2_finseq_1(X3,k3_finseq_2(X5))|~r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3)))).#trainneg
cnf(i_0_33, plain, (X3=k4_matrix_1(X1,X2)|v1_xboole_0(X1)|v7_ordinal1(esk4_3(X1,X2,X3))|r2_hidden(k2_tarski(k2_tarski(esk7_3(X1,X2,X3),esk6_3(X1,X2,X3)),k1_tarski(esk7_3(X1,X2,X3))),k2_matrix_1(X2))|k1_matrix_1(X2)!=k3_finseq_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_32, plain, (X2=k4_matrix_1(X1,X3)|v1_xboole_0(X1)|v7_ordinal1(esk4_3(X1,X3,X2))|k1_matrix_1(X3)!=k3_finseq_1(X2)|k3_matrix_1(X1,X2,esk6_3(X1,X3,X2),esk7_3(X1,X3,X2))!=k3_matrix_1(X1,X3,esk7_3(X1,X3,X2),esk6_3(X1,X3,X2))|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_27, plain, (X3=k4_matrix_1(X1,X2)|v1_xboole_0(X1)|v7_ordinal1(esk6_3(X1,X2,X3))|k1_matrix_1(X2)!=k3_finseq_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~r2_hidden(k2_tarski(k2_tarski(esk4_3(X1,X2,X3),esk5_3(X1,X2,X3)),k1_tarski(esk4_3(X1,X2,X3))),k2_matrix_1(X3))|~r2_hidden(k2_tarski(k2_tarski(esk5_3(X1,X2,X3),esk4_3(X1,X2,X3)),k1_tarski(esk5_3(X1,X2,X3))),k2_matrix_1(X2)))).#trainneg
cnf(i_0_29, plain, (X3=k4_matrix_1(X1,X2)|v1_xboole_0(X1)|v7_ordinal1(esk5_3(X1,X2,X3))|r2_hidden(k2_tarski(k2_tarski(esk7_3(X1,X2,X3),esk6_3(X1,X2,X3)),k1_tarski(esk7_3(X1,X2,X3))),k2_matrix_1(X2))|k1_matrix_1(X2)!=k3_finseq_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_28, plain, (X2=k4_matrix_1(X1,X3)|v1_xboole_0(X1)|v7_ordinal1(esk5_3(X1,X3,X2))|k1_matrix_1(X3)!=k3_finseq_1(X2)|k3_matrix_1(X1,X2,esk6_3(X1,X3,X2),esk7_3(X1,X3,X2))!=k3_matrix_1(X1,X3,esk7_3(X1,X3,X2),esk6_3(X1,X3,X2))|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_23, plain, (X3=k4_matrix_1(X1,X2)|v1_xboole_0(X1)|v7_ordinal1(esk6_3(X1,X2,X3))|r2_hidden(k2_tarski(k2_tarski(esk4_3(X1,X2,X3),esk5_3(X1,X2,X3)),k1_tarski(esk4_3(X1,X2,X3))),k2_matrix_1(X3))|r2_hidden(k2_tarski(k2_tarski(esk5_3(X1,X2,X3),esk4_3(X1,X2,X3)),k1_tarski(esk5_3(X1,X2,X3))),k2_matrix_1(X2))|k1_matrix_1(X2)!=k3_finseq_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_26, plain, (X3=k4_matrix_1(X1,X2)|v1_xboole_0(X1)|v7_ordinal1(esk7_3(X1,X2,X3))|k1_matrix_1(X2)!=k3_finseq_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~r2_hidden(k2_tarski(k2_tarski(esk4_3(X1,X2,X3),esk5_3(X1,X2,X3)),k1_tarski(esk4_3(X1,X2,X3))),k2_matrix_1(X3))|~r2_hidden(k2_tarski(k2_tarski(esk5_3(X1,X2,X3),esk4_3(X1,X2,X3)),k1_tarski(esk5_3(X1,X2,X3))),k2_matrix_1(X2)))).#trainneg
cnf(i_0_24, plain, (X2=k4_matrix_1(X1,X3)|v1_xboole_0(X1)|k1_matrix_1(X3)!=k3_finseq_1(X2)|k3_matrix_1(X1,X2,esk6_3(X1,X3,X2),esk7_3(X1,X3,X2))!=k3_matrix_1(X1,X3,esk7_3(X1,X3,X2),esk6_3(X1,X3,X2))|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~r2_hidden(k2_tarski(k2_tarski(esk4_3(X1,X3,X2),esk5_3(X1,X3,X2)),k1_tarski(esk4_3(X1,X3,X2))),k2_matrix_1(X2))|~r2_hidden(k2_tarski(k2_tarski(esk5_3(X1,X3,X2),esk4_3(X1,X3,X2)),k1_tarski(esk5_3(X1,X3,X2))),k2_matrix_1(X3)))).#trainneg
cnf(i_0_25, plain, (X3=k4_matrix_1(X1,X2)|v1_xboole_0(X1)|r2_hidden(k2_tarski(k2_tarski(esk7_3(X1,X2,X3),esk6_3(X1,X2,X3)),k1_tarski(esk7_3(X1,X2,X3))),k2_matrix_1(X2))|k1_matrix_1(X2)!=k3_finseq_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~r2_hidden(k2_tarski(k2_tarski(esk4_3(X1,X2,X3),esk5_3(X1,X2,X3)),k1_tarski(esk4_3(X1,X2,X3))),k2_matrix_1(X3))|~r2_hidden(k2_tarski(k2_tarski(esk5_3(X1,X2,X3),esk4_3(X1,X2,X3)),k1_tarski(esk5_3(X1,X2,X3))),k2_matrix_1(X2)))).#trainneg
cnf(i_0_22, plain, (X3=k4_matrix_1(X1,X2)|v1_xboole_0(X1)|v7_ordinal1(esk7_3(X1,X2,X3))|r2_hidden(k2_tarski(k2_tarski(esk4_3(X1,X2,X3),esk5_3(X1,X2,X3)),k1_tarski(esk4_3(X1,X2,X3))),k2_matrix_1(X3))|r2_hidden(k2_tarski(k2_tarski(esk5_3(X1,X2,X3),esk4_3(X1,X2,X3)),k1_tarski(esk5_3(X1,X2,X3))),k2_matrix_1(X2))|k1_matrix_1(X2)!=k3_finseq_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_21, plain, (X3=k4_matrix_1(X1,X2)|v1_xboole_0(X1)|r2_hidden(k2_tarski(k2_tarski(esk4_3(X1,X2,X3),esk5_3(X1,X2,X3)),k1_tarski(esk4_3(X1,X2,X3))),k2_matrix_1(X3))|r2_hidden(k2_tarski(k2_tarski(esk5_3(X1,X2,X3),esk4_3(X1,X2,X3)),k1_tarski(esk5_3(X1,X2,X3))),k2_matrix_1(X2))|r2_hidden(k2_tarski(k2_tarski(esk7_3(X1,X2,X3),esk6_3(X1,X2,X3)),k1_tarski(esk7_3(X1,X2,X3))),k2_matrix_1(X2))|k1_matrix_1(X2)!=k3_finseq_1(X3)|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_20, plain, (X2=k4_matrix_1(X1,X3)|v1_xboole_0(X1)|r2_hidden(k2_tarski(k2_tarski(esk4_3(X1,X3,X2),esk5_3(X1,X3,X2)),k1_tarski(esk4_3(X1,X3,X2))),k2_matrix_1(X2))|r2_hidden(k2_tarski(k2_tarski(esk5_3(X1,X3,X2),esk4_3(X1,X3,X2)),k1_tarski(esk5_3(X1,X3,X2))),k2_matrix_1(X3))|k1_matrix_1(X3)!=k3_finseq_1(X2)|k3_matrix_1(X1,X2,esk6_3(X1,X3,X2),esk7_3(X1,X3,X2))!=k3_matrix_1(X1,X3,esk7_3(X1,X3,X2),esk6_3(X1,X3,X2))|~v1_matrix_1(X3)|~v1_matrix_1(X2)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_279, plain, (esk23_0=k1_xboole_0)).#trainneg
cnf(i_0_598, plain, (k1_funct_1(k9_finseq_1(X1),np__1)=X1)).#trainneg
cnf(i_0_315, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_541, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk16_1(X2)))).#trainneg
cnf(i_0_545, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk10_1(X2)))).#trainneg
cnf(i_0_280, plain, (esk18_0=k1_xboole_0)).#trainneg
cnf(i_0_281, plain, (esk15_0=k1_xboole_0)).#trainneg
cnf(i_0_296, plain, (~v2_xxreal_0(esk25_0))).#trainneg
cnf(i_0_297, plain, (~v2_xxreal_0(esk17_0))).#trainneg
cnf(i_0_461, negated_conjecture, (v3_ordinal1(esk2_0)|~v3_ordinal1(esk1_0))).#trainneg
cnf(i_0_348, negated_conjecture, (v1_relat_1(esk2_0)|~v4_funct_1(esk1_0))).#trainneg
cnf(i_0_304, plain, (v1_finset_1(k2_finseq_1(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_308, plain, (X1=esk14_2(X2,X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_533, negated_conjecture, (k4_matrix_1(esk1_0,k9_finseq_1(k3_pre_poly(esk1_0,esk2_0)))!=k9_finseq_1(k3_pre_poly(esk1_0,esk2_0))|~m1_finseq_1(k3_pre_poly(esk1_0,esk2_0),esk1_0))).#trainneg
cnf(i_0_495, plain, (v1_card_1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_512, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_544, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk11_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_462, plain, (v3_ordinal1(np__0))).#trainneg
cnf(i_0_657, plain, (k9_xtuple_0(k3_finseq_1(X1))=k2_finseq_1(X2)|X2!=k3_finseq_1(X1)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_finseq_1(k3_finseq_1(X1))|~v1_funct_1(k3_finseq_1(X1))|~v1_relat_1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_309, plain, (X1=esk12_1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_464, plain, (v3_ordinal1(np__2))).#trainneg
cnf(i_0_329, plain, (~m1_subset_1(k5_numbers,k5_numbers))).#trainneg
cnf(i_0_294, plain, (~v1_xxreal_0(np__2)|~v1_xboole_0(np__2))).#trainneg
cnf(i_0_427, negated_conjecture, (v1_funct_1(esk2_0)|~v4_funct_1(esk1_0))).#trainneg
cnf(i_0_446, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_310, plain, (X1=esk10_1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_311, plain, (X1=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_295, plain, (~v1_xxreal_0(np__1)|~v1_xboole_0(np__1))).#trainneg
cnf(i_0_655, plain, (m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_668, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3))|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X1,k4_finseq_1(X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_552, plain, (m2_finseq_1(X1,X2)|~m1_finseq_2(X3,X2)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_601, plain, (m2_finseq_2(k9_finseq_1(X2),X1,k3_finseq_2(X1))|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_402, plain, (X1=X2|v1_xboole_0(X3)|X3!=k1_tarski(X2)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_665, plain, (r2_hidden(X1,k4_finseq_1(X2))|~r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_matrix_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_383, plain, (m2_finseq_1(esk21_1(X1),X1)|v1_xboole_0(X1))).#trainneg
cnf(i_0_509, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_376, plain, (~v1_xboole_0(k1_finseq_1(np__1))|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_312, plain, (X1=k1_card_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_349, plain, (v1_relat_1(np__0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_379, plain, (~v1_xboole_0(esk19_1(np__1))|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_381, plain, (~v1_xboole_0(esk27_1(np__1))|~v1_card_1(np__1))).#trainneg
cnf(i_0_522, plain, (v5_relat_1(X1,X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_362, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_313, plain, (X1=k9_xtuple_0(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_316, plain, (v1_xboole_0(esk16_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_648, plain, (r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X1)),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_499, plain, (v1_xboole_0(X1)|~v1_xboole_0(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_606, plain, (r2_hidden(X1,X2)|v1_xboole_0(k2_zfmisc_1(X4,X2))|~m1_subset_1(k2_tarski(k2_tarski(X3,X1),k1_tarski(X3)),k2_zfmisc_1(X4,X2)))).#trainneg
cnf(i_0_343, plain, (v1_xboole_0(X1)|~v1_relat_1(X1)|~v1_xboole_0(X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_471, plain, (v1_relat_1(esk28_0)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_630, plain, (~v1_xboole_0(k9_xtuple_0(X1))|~v3_card_1(X1,np__1)|~v1_card_1(np__1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_317, plain, (v1_xboole_0(esk11_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_350, plain, (v1_relat_1(np__0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_374, plain, (~v1_finseq_1(k5_numbers)|~v1_funct_1(k5_numbers)|~v1_relat_1(k5_numbers))).#trainneg
cnf(i_0_351, plain, (v1_relat_1(np__2)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_569, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_570, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_326, plain, (v1_finset_1(X1)|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_328, plain, (v1_finset_1(X1)|~m1_subset_1(k1_card_1(X1),k5_numbers))).#trainneg
cnf(i_0_384, plain, (m2_finseq_1(esk12_1(X1),X1))).#trainneg
cnf(i_0_705, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3))|v1_xboole_0(X4)|X5!=k4_matrix_1(X4,X3)|~r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k2_matrix_1(X5))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(X5))).#trainneg
cnf(i_0_674, plain, (m1_subset_1(k3_finseq_1(X1),k1_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_666, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|~r2_hidden(k2_tarski(k2_tarski(X3,X1),k1_tarski(X3)),k2_matrix_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_602, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X3)),k2_zfmisc_1(X4,X2)))).#trainneg
cnf(i_0_607, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k2_tarski(X3,X1),k1_tarski(X1)),k2_zfmisc_1(X2,X4)))).#trainneg
cnf(i_0_700, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3))|v1_xboole_0(X4)|X3!=k4_matrix_1(X4,X5)|~r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k2_matrix_1(X5))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(X5))).#trainneg
cnf(i_0_537, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,esk28_0))).#trainneg
cnf(i_0_352, plain, (v1_relat_1(np__2)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_508, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,X3))|~m2_finseq_1(X2,X3))).#trainneg
cnf(i_0_354, plain, (v1_relat_1(np__1)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_355, plain, (v1_relat_1(np__1)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_360, plain, (v7_ordinal1(esk28_0)|~v1_finset_1(esk28_0))).#trainneg
cnf(i_0_492, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_368, plain, (X1=k1_finseq_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_344, plain, (v1_relat_1(esk28_0)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_644, plain, (m1_subset_1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_578, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k4_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_675, plain, (m2_finseq_1(k2_matrlin(X1,X2),k3_finseq_2(X1))|v1_xboole_0(X1)|~v7_ordinal1(k3_finseq_1(X2))|~v7_ordinal1(np__1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_645, plain, (k2_tarski(k2_tarski(X1,X2),k1_tarski(X1))=X3|k2_zfmisc_1(X4,X5)!=k1_tarski(X3)|~r2_hidden(X2,X5)|~r2_hidden(X1,X4))).#trainneg
cnf(i_0_497, plain, (v1_xboole_0(k3_finseq_1(X1))|~v1_xboole_0(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_611, plain, (r2_hidden(X1,X2)|v1_xboole_0(k2_zfmisc_1(X2,X4))|~m1_subset_1(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_zfmisc_1(X2,X4)))).#trainneg
cnf(i_0_702, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3))|v1_xboole_0(X4)|X3!=k4_matrix_1(X4,X5)|~r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X2,X1)),k2_matrix_1(X5))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(X5))).#trainneg
cnf(i_0_547, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_441, plain, (v1_xboole_0(esk16_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_345, plain, (v1_relat_1(esk16_1(X1))|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_428, plain, (v1_funct_1(np__0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_429, plain, (v1_funct_1(np__0)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_430, plain, (v1_funct_1(np__2)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_531, plain, (v4_relat_1(X1,k5_numbers)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_565, plain, (r1_xxreal_0(X1,esk25_0)|r1_xxreal_0(esk25_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_346, plain, (v1_relat_1(esk11_1(X1))|~v4_funct_1(X1))).#trainneg
cnf(i_0_347, plain, (v1_relat_1(esk10_1(X1))|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_643, plain, (~v1_xboole_0(k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_677, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_707, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3))|v1_xboole_0(X4)|X5!=k4_matrix_1(X4,X3)|~r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X2,X1)),k2_matrix_1(X5))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(X5))).#trainneg
cnf(i_0_585, plain, (v3_ordinal1(k1_matrix_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_378, plain, (v3_card_1(k2_finseq_1(X1),X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_710, plain, (k3_matrix_1(X1,X2,X3,X4)=k3_matrix_1(X1,X5,X4,X3)|v1_xboole_0(X1)|X5!=k4_matrix_1(X1,X2)|~r2_hidden(k2_tarski(k2_tarski(X4,X3),k1_tarski(X3)),k2_matrix_1(X2))|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m2_finseq_1(X5,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X5)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_628, plain, (m1_matrix_1(k2_matrlin(X1,k3_finseq_1(X2)),X1,np__1,k3_finseq_1(X2))|v1_xboole_0(X1)|~m1_finseq_1(k3_finseq_1(X2),X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_579, plain, (m1_subset_1(k9_xtuple_0(X1),k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_538, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk26_1(X2)))).#trainneg
cnf(i_0_370, plain, (v1_xboole_0(k2_finseq_1(X1))|~v1_xboole_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_431, plain, (v1_funct_1(np__2)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_536, plain, (esk3_2(X1,X2)=X3|esk3_2(X1,X2)=X1|X2=k1_tarski(X1)|X2!=k1_tarski(X3))).#trainneg
cnf(i_0_353, plain, (v1_relat_1(k5_numbers)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_434, plain, (v1_funct_1(np__1)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_435, plain, (v1_funct_1(np__1)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_567, plain, (r1_xxreal_0(X1,esk17_0)|r1_xxreal_0(esk17_0,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_548, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~r2_hidden(X1,X3)|~m2_finseq_1(X3,X2))).#trainneg
cnf(i_0_357, plain, (v1_xboole_0(esk16_1(k2_zfmisc_1(X1,X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_358, plain, (v1_xboole_0(esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_496, plain, (k1_card_1(k3_finseq_1(X1))=k3_finseq_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_494, plain, (k3_finseq_1(k1_card_1(X1))=k1_card_1(X1)|~v1_finseq_1(k1_card_1(X1))|~v1_funct_1(k1_card_1(X1))|~v1_relat_1(k1_card_1(X1)))).#trainneg
cnf(i_0_546, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_367, plain, (m1_subset_1(k1_tarski(np__1),k1_zfmisc_1(k5_numbers))|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_217, plain, (v1_xboole_0(X1)|~v3_card_1(X1,np__0))).#trainneg
cnf(i_0_281, plain, (esk15_0=np__0)).#trainneg
cnf(i_0_363, plain, (v1_relat_1(k2_finseq_1(X1))|~v4_funct_1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_280, plain, (esk18_0=np__0)).#trainneg
cnf(i_0_934, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_1216, plain, (~v1_xboole_0(np__2))).#trainneg
cnf(i_0_817, plain, (r1_xxreal_0(X1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_366, plain, (m1_subset_1(k2_tarski(np__1,np__2),k1_zfmisc_1(k5_numbers))|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_373, plain, (v1_finset_1(X1)|~v1_finseq_1(k1_card_1(X1))|~v1_funct_1(k1_card_1(X1))|~v1_relat_1(k1_card_1(X1)))).#trainneg
cnf(i_0_795, plain, (v1_card_1(np__1))).#trainneg
cnf(i_0_214, plain, (v1_xxreal_0(k1_xboole_0))).#trainneg
cnf(i_0_279, plain, (esk23_0=np__0)).#trainneg
cnf(i_0_386, plain, (v1_relat_1(esk13_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_51, plain, (k6_numbers=np__0)).#trainneg
cnf(i_0_818, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_739, negated_conjecture, (v7_ordinal1(esk2_0)|~v1_finset_1(esk2_0)|~v3_ordinal1(esk1_0))).#trainneg
cnf(i_0_387, plain, (v1_xboole_0(esk13_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_390, plain, (v1_relat_1(esk20_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_944, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X3)),k2_matrix_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_948, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3))|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X1,k9_xtuple_0(X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_1007, plain, (r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X1)),k2_matrix_1(X3))|v1_xboole_0(X4)|X5!=k4_matrix_1(X4,X3)|~r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X2)),k2_matrix_1(X5))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(X5))).#trainneg
cnf(i_0_728, plain, (k3_finseq_1(np__1)=np__1)).#trainneg
cnf(i_0_899, plain, (v7_ordinal1(k3_finseq_1(X1))|~v1_finset_1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_179, plain, (v1_xreal_0(k1_xboole_0))).#trainneg
cnf(i_0_181, plain, (v1_xcmplx_0(k1_xboole_0))).#trainneg
cnf(i_0_10, plain, (k3_finseq_1(k4_matrix_1(X1,X2))=k1_matrix_1(X2)|r1_xxreal_0(k1_matrix_1(X2),np__0)|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_815, plain, (r1_xxreal_0(esk25_0,esk25_0))).#trainneg
cnf(i_0_763, plain, (v1_funct_1(k3_pre_poly(X1,X2))|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_764, plain, (v1_finseq_1(k3_pre_poly(X1,X2))|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_394, plain, (v1_relat_1(esk24_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1040, plain, (r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X1)),k2_matrix_1(X3))|v1_xboole_0(X4)|X3!=k4_matrix_1(X4,X5)|~r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X2)),k2_matrix_1(X5))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(X5))).#trainneg
cnf(i_0_943, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_matrix_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_1261, plain, (k1_card_1(np__1)=np__1)).#trainneg
cnf(i_0_1009, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3))|v1_xboole_0(X4)|X5!=k4_matrix_1(X4,X3)|~r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X1,X2)),k2_matrix_1(X5))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(X5))).#trainneg
cnf(i_0_398, plain, (v1_relat_1(esk26_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_816, plain, (r1_xxreal_0(esk17_0,esk17_0))).#trainneg
cnf(i_0_765, plain, (v1_relat_1(k3_pre_poly(X1,X2))|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_409, plain, (v1_relat_1(k1_tarski(k2_tarski(k2_tarski(X2,X1),k1_tarski(X1)))))).#trainneg
cnf(i_0_950, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_matrix_1(X3))|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X1,k9_xtuple_0(X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_865, plain, (k3_matrix_1(X1,X2,X3,X4)=X5|v1_xboole_0(X1)|X1!=k1_tarski(X5)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_875, plain, (r2_hidden(X1,k4_finseq_1(X2))|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X3)),k2_matrix_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_1093, plain, (k1_matrix_1(X1)=k3_finseq_1(k2_matrlin(X2,X3))|v1_xboole_0(X2)|k2_matrlin(X2,X3)!=k4_matrix_1(X2,X1)|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(X1)|~v1_matrix_1(k2_matrlin(X2,X3))|~m1_finseq_1(X3,X2)|~v7_ordinal1(k3_finseq_1(X3))|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_539, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk24_1(X2)))).#trainneg
cnf(i_0_4, plain, (m1_matrix_1(X1,X2,k3_finseq_1(X1),X3)|r1_xxreal_0(k3_finseq_1(X1),np__0)|v1_xboole_0(X2)|X3!=k1_matrix_1(X1)|~v7_ordinal1(X3)|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(X1))).#trainneg
cnf(i_0_1150, plain, (r1_xxreal_0(esk25_0,esk17_0)|r1_xxreal_0(esk17_0,esk25_0))).#trainneg
cnf(i_0_986, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_411, plain, (v1_relat_1(k1_tarski(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)))))).#trainneg
cnf(i_0_1042, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3))|v1_xboole_0(X4)|X3!=k4_matrix_1(X4,X5)|~r2_hidden(k2_tarski(k1_tarski(X2),k2_tarski(X1,X2)),k2_matrix_1(X5))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(X5))).#trainneg
cnf(i_0_788, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X3,X1)),k2_matrix_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_540, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk20_1(X2)))).#trainneg
cnf(i_0_967, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_859, negated_conjecture, (esk2_0=X1|k1_tarski(X1)!=esk1_0)).#trainneg
cnf(i_0_418, plain, (v1_funct_1(esk28_0)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_481, plain, (v1_relat_1(k5_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_987, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_419, plain, (v1_funct_1(esk26_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_420, plain, (v1_funct_1(esk24_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_771, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X2,X1)),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_898, plain, (v7_ordinal1(k1_card_1(X1))|~v1_finset_1(k1_card_1(X1)))).#trainneg
cnf(i_0_1078, plain, (k3_matrix_1(X1,X2,X3,X4)=k3_matrix_1(X1,X5,X4,X3)|v1_xboole_0(X1)|X5!=k4_matrix_1(X1,X2)|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X4,X3)),k2_matrix_1(X2))|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m2_finseq_1(X5,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X5)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_770, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_matrix_1(X3))|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X1,k4_finseq_1(X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_543, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk13_1(X2)))).#trainneg
cnf(i_0_801, plain, (v7_ordinal1(np__0)|~v1_finset_1(np__0))).#trainneg
cnf(i_0_860, plain, (np__0=X1|k1_tarski(X1)!=k5_numbers)).#trainneg
cnf(i_0_421, plain, (v1_funct_1(esk20_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_812, plain, (v7_ordinal1(np__2)|~v1_finset_1(np__2))).#trainneg
cnf(i_0_433, plain, (v1_funct_1(k5_numbers)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_553, plain, (v1_funct_1(esk28_0)|~v1_funct_1(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1051, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v7_ordinal1(X2))).#trainneg
cnf(i_0_907, plain, (X1=esk16_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_422, plain, (v1_funct_1(esk16_1(X1))|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_423, plain, (v1_funct_1(esk14_2(X1,X2))|~v4_funct_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_789, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X1,X3)),k2_zfmisc_1(X4,X2)))).#trainneg
cnf(i_0_846, plain, (m2_finseq_1(X1,X2)|~m1_subset_1(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_949, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k2_matrix_1(X3))|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X2,k9_xtuple_0(X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_1075, plain, (m1_matrix_1(k9_finseq_1(k9_finseq_1(X1)),X2,np__1,np__1)|v1_xboole_0(X2)|~m1_finseq_1(k9_finseq_1(X1),X2))).#trainneg
cnf(i_0_1178, plain, (v7_ordinal1(k1_matrix_1(X1))|~v1_finset_1(k1_matrix_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_1268, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k1_tarski(np__1))|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_563, plain, (v1_funct_1(k5_numbers)|~v1_funct_1(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_861, plain, (np__0=X1|k1_tarski(X1)!=k1_numbers)).#trainneg
cnf(i_0_424, plain, (v1_funct_1(esk13_1(X1))|v1_xboole_0(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1049, plain, (v7_ordinal1(esk28_0)|~v1_finseq_1(esk28_0)|~v1_funct_1(esk28_0)|~v1_relat_1(esk28_0))).#trainneg
cnf(i_0_1278, plain, (v1_relat_1(k1_tarski(np__1))|~v1_relat_1(k5_numbers)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1272, plain, (v1_relat_1(k1_tarski(np__1))|~v4_funct_1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1107, plain, (X1=k3_finseq_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1208, plain, (X1=k2_finseq_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_425, plain, (v1_funct_1(esk11_1(X1))|~v4_funct_1(X1))).#trainneg
cnf(i_0_426, plain, (v1_funct_1(esk10_1(X1))|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1190, plain, (m1_matrix_1(k2_matrlin(X1,np__1),X1,np__1,np__1)|v1_xboole_0(X1)|~m1_finseq_1(np__1,X1))).#trainneg
cnf(i_0_757, plain, (X1=k4_matrix_1(X2,k4_matrix_1(X2,X3))|r2_hidden(k2_tarski(k1_tarski(esk5_3(X2,k4_matrix_1(X2,X3),X1)),k2_tarski(esk4_3(X2,k4_matrix_1(X2,X3),X1),esk5_3(X2,k4_matrix_1(X2,X3),X1))),k2_matrix_1(k4_matrix_1(X2,X3)))|r2_hidden(k2_tarski(k1_tarski(esk4_3(X2,k4_matrix_1(X2,X3),X1)),k2_tarski(esk4_3(X2,k4_matrix_1(X2,X3),X1),esk5_3(X2,k4_matrix_1(X2,X3),X1))),k2_matrix_1(X1))|v1_xboole_0(X2)|k3_matrix_1(X2,X1,esk6_3(X2,k4_matrix_1(X2,X3),X1),esk7_3(X2,k4_matrix_1(X2,X3),X1))!=k3_matrix_1(X2,k4_matrix_1(X2,X3),esk7_3(X2,k4_matrix_1(X2,X3),X1),esk6_3(X2,k4_matrix_1(X2,X3),X1))|k1_matrix_1(k4_matrix_1(X2,X3))!=k3_finseq_1(X1)|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,X3))|~v1_matrix_1(X1)|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_803, plain, (k9_xtuple_0(np__1)=k2_finseq_1(X2)|X2!=np__1|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_finseq_1(np__1)|~v1_funct_1(np__1)|~v1_relat_1(np__1))).#trainneg
cnf(i_0_758, plain, (k1_matrix_1(X1)=k3_finseq_1(k3_pre_poly(k3_finseq_2(X2),X3))|v1_xboole_0(X2)|k3_pre_poly(k3_finseq_2(X2),X3)!=k4_matrix_1(X2,X1)|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(X1)|~v1_matrix_1(k3_pre_poly(k3_finseq_2(X2),X3))|~m1_subset_1(X3,k3_finseq_2(X2)))).#trainneg
cnf(i_0_1032, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(k2_tarski(k2_tarski(X3,X1),k1_tarski(X1)),k2_matrix_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_1181, plain, (v3_card_1(k1_tarski(np__1),np__1)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_863, plain, (np__2=X1|k1_tarski(X1)!=k5_numbers)).#trainneg
cnf(i_0_432, plain, (v1_funct_1(k2_finseq_1(X1))|~v4_funct_1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_626, plain, (v5_relat_1(k5_numbers,X1)|~v5_relat_1(k1_numbers,X1)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1273, plain, (v1_funct_1(k1_tarski(np__1))|~v4_funct_1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_1274, plain, (v3_ordinal1(k1_tarski(np__1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_976, plain, (X1=esk11_1(k1_zfmisc_1(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1134, plain, (X1=esk16_1(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_437, plain, (v7_ordinal1(k9_xtuple_0(X1))|~v1_finset_1(k9_xtuple_0(X1))|~v5_ordinal1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_926, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k2_matrix_1(X3))|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X2,k4_finseq_1(X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_1189, plain, (m2_finseq_1(k2_matrlin(X1,k3_finseq_1(X2)),k3_finseq_2(X1))|v1_xboole_0(X1)|~v7_ordinal1(k3_finseq_1(X2))|~v7_ordinal1(np__1)|~m1_finseq_1(k3_finseq_1(X2),X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_871, plain, (k3_pre_poly(X1,X2)=X3|v1_xboole_0(X1)|k3_finseq_2(X1)!=k1_tarski(X3)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1106, plain, (k2_tarski(k2_tarski(X1,X2),k1_tarski(X1))=X3|v1_xboole_0(X5)|k2_zfmisc_1(X4,X5)!=k1_tarski(X3)|~r2_hidden(X1,X4)|~m1_subset_1(X2,X5))).#trainneg
cnf(i_0_1017, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|~r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X3)),k2_matrix_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_806, plain, (esk12_1(X1)=esk12_1(X2))).#trainneg
cnf(i_0_809, plain, (esk12_1(X1)=np__0)).#trainneg
cnf(i_0_873, plain, (r2_hidden(X1,k4_finseq_1(X2))|~r2_hidden(k2_tarski(k2_tarski(X3,X1),k1_tarski(X1)),k2_matrix_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_960, plain, (v1_relat_1(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_821, plain, (esk10_1(X1)=esk10_1(X2))).#trainneg
cnf(i_0_1279, plain, (v1_funct_1(k1_tarski(np__1))|~v1_funct_1(k5_numbers)|~v1_relat_1(k5_numbers)|~v7_ordinal1(np__1))).#trainneg
cnf(i_0_864, plain, (np__2=X1|k1_tarski(X1)!=k1_numbers)).#trainneg
cnf(i_0_615, plain, (v5_relat_1(esk28_0,X1)|~v5_relat_1(k1_numbers,X1)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1180, plain, (v3_card_1(k2_tarski(np__1,np__2),np__2)|~v7_ordinal1(np__2))).#trainneg
cnf(i_0_321, plain, (~r2_hidden(X1,esk16_1(X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1232, plain, (X1=esk16_1(k2_zfmisc_1(X2,X3))|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1839, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk10_1(X3)))).#trainneg
cnf(i_0_1578, plain, (v7_ordinal1(np__0))).#trainneg
cnf(i_0_1582, plain, (v7_ordinal1(np__2))).#trainneg
cnf(i_0_381, plain, (~v1_xboole_0(esk27_1(np__1)))).#trainneg
cnf(i_0_444, plain, (v1_xboole_0(esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_137, plain, (v1_finset_1(np__0))).#trainneg
cnf(i_0_1294, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X1)|v1_xboole_0(X1)|~v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_453, plain, (v3_ordinal1(esk26_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_454, plain, (v3_ordinal1(esk24_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_136, plain, (v1_finseq_1(np__0))).#trainneg
cnf(i_0_139, plain, (v1_funct_1(np__0))).#trainneg
cnf(i_0_630, plain, (~v3_card_1(X1,np__1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(k9_xtuple_0(X1)))).#trainneg
cnf(i_0_455, plain, (v3_ordinal1(esk20_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_141, plain, (v4_relat_1(np__0,k5_numbers))).#trainneg
cnf(i_0_142, plain, (v1_relat_1(np__0))).#trainneg
cnf(i_0_1333, negated_conjecture, (v7_ordinal1(esk2_0)|~v1_finset_1(esk2_0)|~v1_card_1(esk1_0))).#trainneg
cnf(i_0_1494, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_456, plain, (v3_ordinal1(esk16_1(X1))|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_457, plain, (v3_ordinal1(esk14_2(X1,X2))|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_1421, plain, (X1=np__1|k2_finseq_1(X1)!=k1_tarski(np__1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_143, plain, (m1_finseq_1(np__0,X1))).#trainneg
cnf(i_0_1554, plain, (v7_ordinal1(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_1284, plain, (m1_subset_1(X1,k1_tarski(X1)))).#trainneg
cnf(i_0_384, plain, (m2_finseq_1(np__0,X1))).#trainneg
cnf(i_0_140, plain, (v5_relat_1(np__0,X1))).#trainneg
cnf(i_0_458, plain, (v3_ordinal1(esk13_1(X1))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1321, plain, (r1_xxreal_0(esk25_0,np__0)|r1_xxreal_0(np__0,esk25_0))).#trainneg
cnf(i_0_1289, plain, (v1_xboole_0(esk27_1(np__0))|~v1_card_1(np__0))).#trainneg
cnf(i_0_1290, plain, (v1_xboole_0(esk19_1(np__0))|~v7_ordinal1(np__0))).#trainneg
cnf(i_0_1591, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_459, plain, (v3_ordinal1(esk11_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_460, plain, (v3_ordinal1(esk10_1(X1))|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1573, plain, (~v1_xboole_0(X1)|X1!=np__1|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1313, plain, (m2_finseq_1(k2_matrlin(X1,k9_finseq_1(X2)),k3_finseq_2(X1))|v1_xboole_0(X1)|~v7_ordinal1(np__1)|~m1_finseq_1(k9_finseq_1(X2),X1))).#trainneg
cnf(i_0_1327, plain, (k1_matrix_1(k4_matrix_1(X1,k4_matrix_1(X1,X2)))=k3_finseq_1(k4_matrix_1(X1,X2))|r1_xxreal_0(k1_matrix_1(k4_matrix_1(X1,X2)),np__0)|v1_xboole_0(X1)|~v1_matrix_1(k4_matrix_1(X1,X2))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_1725, plain, (k1_matrix_1(X1)=np__1|v1_xboole_0(X2)|k9_finseq_1(X3)!=k4_matrix_1(X2,X1)|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(k9_finseq_1(X3))|~v1_matrix_1(X1)|~m1_subset_1(X3,k3_finseq_2(X2)))).#trainneg
cnf(i_0_1291, plain, (v1_xboole_0(k1_finseq_1(np__0))|~v7_ordinal1(np__0))).#trainneg
cnf(i_0_1562, plain, (k1_tarski(np__1)=k2_tarski(np__1,np__2)|np__2!=np__1)).#trainneg
cnf(i_0_868, plain, (np__1=X1|k1_tarski(X1)!=k5_numbers)).#trainneg
cnf(i_0_466, plain, (v3_ordinal1(k2_finseq_1(X1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1322, plain, (r1_xxreal_0(esk17_0,np__0)|r1_xxreal_0(np__0,esk17_0))).#trainneg
cnf(i_0_1293, plain, (v1_xboole_0(k2_finseq_1(np__0))|~v7_ordinal1(np__0))).#trainneg
cnf(i_0_1334, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1515, plain, (X1=k2_zfmisc_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_472, plain, (v1_relat_1(esk26_1(X1))|v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_473, plain, (v1_relat_1(esk24_1(X1))|v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1542, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X2,X1)),k2_matrix_1(X3))|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X1,k9_xtuple_0(X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_1360, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k2_matrix_1(X3))|v1_xboole_0(X4)|X5!=k4_matrix_1(X4,X3)|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_matrix_1(X5))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(X5))).#trainneg
cnf(i_0_961, plain, (v1_funct_1(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_869, plain, (np__1=X1|k1_tarski(X1)!=k1_numbers)).#trainneg
cnf(i_0_474, plain, (v1_relat_1(esk20_1(X1))|v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1732, plain, (v5_relat_1(k5_numbers,X1)|~v1_relat_1(k1_numbers)|~m1_finseq_1(k1_numbers,X1))).#trainneg
cnf(i_0_1426, plain, (m1_subset_1(k1_tarski(np__1),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1388, plain, (v3_xxreal_0(X1)|v2_xxreal_0(np__0)|r1_xxreal_0(np__0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1332, plain, (r1_xxreal_0(X1,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_475, plain, (v1_relat_1(esk16_1(X1))|~v1_relat_1(X1))).#trainneg
cnf(i_0_477, plain, (v1_relat_1(esk13_1(X1))|v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1324, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1687, plain, (v1_matrix_1(k2_matrlin(X1,np__1))|v1_xboole_0(X1)|~v7_ordinal1(np__1)|~m1_finseq_1(np__1,X1))).#trainneg
cnf(i_0_1664, plain, (k3_finseq_1(X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_962, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1378, plain, (k9_xtuple_0(np__1)=k1_tarski(np__1)|~v1_finseq_1(np__1)|~v1_funct_1(np__1)|~v1_relat_1(np__1))).#trainneg
cnf(i_0_850, plain, (esk28_0=X1|k1_zfmisc_1(k1_numbers)!=k1_tarski(X1))).#trainneg
cnf(i_0_478, plain, (v1_relat_1(esk11_1(k1_zfmisc_1(X1)))|~v1_relat_1(X1))).#trainneg
cnf(i_0_1733, plain, (v5_relat_1(k5_numbers,X1)|~v1_relat_1(k1_numbers)|~m2_finseq_1(k1_numbers,X1))).#trainneg
cnf(i_0_1424, plain, (v3_ordinal1(k1_tarski(np__1))|~v5_ordinal1(k9_finseq_1(X1)))).#trainneg
cnf(i_0_984, plain, (r1_xxreal_0(X1,esk25_0)|r1_xxreal_0(esk25_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_985, plain, (r1_xxreal_0(X1,esk17_0)|r1_xxreal_0(esk17_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_479, plain, (v1_relat_1(esk10_1(X1))|~v1_relat_1(X1))).#trainneg
cnf(i_0_480, plain, (v1_relat_1(k2_finseq_1(X1))|~v1_relat_1(k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1401, plain, (v7_ordinal1(k3_finseq_1(k3_pre_poly(X1,X2)))|v1_xboole_0(X1)|~v1_funct_1(k3_pre_poly(X1,X2))|~v1_relat_1(k3_pre_poly(X1,X2))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1418, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k2_matrix_1(X3))|v1_xboole_0(X4)|X3!=k4_matrix_1(X4,X5)|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_matrix_1(X5))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(X5))).#trainneg
cnf(i_0_1688, plain, (m2_finseq_1(k2_matrlin(X1,np__1),k3_finseq_2(X1))|v1_xboole_0(X1)|~v7_ordinal1(np__1)|~m1_finseq_1(np__1,X1))).#trainneg
cnf(i_0_1680, plain, (k2_finseq_1(X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1285, plain, (X1=X2|k1_tarski(X1)!=k1_tarski(X2))).#trainneg
cnf(i_0_823, plain, (esk10_1(X1)=np__0)).#trainneg
cnf(i_0_376, plain, (~v1_xboole_0(k1_finseq_1(np__1)))).#trainneg
cnf(i_0_482, plain, (v1_relat_1(esk26_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1934, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2016, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_483, plain, (v1_relat_1(esk24_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_484, plain, (v1_relat_1(esk20_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2200, plain, (r2_hidden(X1,k1_tarski(np__1))|~r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(k9_finseq_1(X3)))|~v1_matrix_1(k9_finseq_1(X3)))).#trainneg
cnf(i_0_1867, plain, (X1=k1_matrix_1(k4_matrix_1(X2,X3))|r1_xxreal_0(k1_matrix_1(X3),np__0)|v1_xboole_0(X4)|v1_xboole_0(X2)|~m1_matrix_1(k4_matrix_1(X2,X3),X4,k1_matrix_1(X3),X1)|~v7_ordinal1(X1)|~m2_finseq_1(k4_matrix_1(X2,X3),k3_finseq_2(X4))|~v1_matrix_1(k4_matrix_1(X2,X3))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_2203, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(k9_finseq_1(X2))))|~r2_hidden(k2_tarski(k2_tarski(X3,X1),k1_tarski(X3)),k2_matrix_1(k9_finseq_1(X2)))|~v1_matrix_1(k9_finseq_1(X2)))).#trainneg
cnf(i_0_1881, plain, (k1_matrix_1(k9_finseq_1(k9_finseq_1(X1)))=np__1|v1_xboole_0(X2)|~v7_ordinal1(np__1)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X1)),k3_finseq_2(X2))|~v1_matrix_1(k9_finseq_1(k9_finseq_1(X1)))|~m1_finseq_1(k9_finseq_1(X1),X2))).#trainneg
cnf(i_0_1880, plain, (X1=k1_matrix_1(X2)|v1_xboole_0(X3)|~m1_matrix_1(X2,X3,np__1,X1)|~v7_ordinal1(X1)|~m2_finseq_1(X2,k3_finseq_2(X3))|~v1_matrix_1(X2)|k3_finseq_1(X2)!=np__1|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1839, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,np__0))).#trainneg
cnf(i_0_2022, plain, (np__0!=np__1)).#trainneg
cnf(i_0_379, plain, (~v1_xboole_0(esk19_1(np__1)))).#trainneg
cnf(i_0_485, plain, (v1_relat_1(esk16_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_2264, plain, (r1_xxreal_0(np__0,esk22_0)|r1_xxreal_0(esk22_0,np__0))).#trainneg
cnf(i_0_2288, plain, (v1_card_1(np__0)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1992, plain, (X1=esk19_1(np__0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_487, plain, (v1_relat_1(esk13_1(k2_zfmisc_1(X1,X2)))|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_488, plain, (v1_relat_1(esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_2201, plain, (r2_hidden(X1,k1_tarski(np__1))|~r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X1)),k2_matrix_1(k9_finseq_1(X3)))|~v1_matrix_1(k9_finseq_1(X3)))).#trainneg
cnf(i_0_2032, plain, (k1_matrix_1(X1)=k3_finseq_1(k2_matrlin(X2,k9_finseq_1(X3)))|v1_xboole_0(X2)|k2_matrlin(X2,k9_finseq_1(X3))!=k4_matrix_1(X2,X1)|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(X1)|~v1_matrix_1(k2_matrlin(X2,k9_finseq_1(X3)))|~m1_finseq_1(k9_finseq_1(X3),X2))).#trainneg
cnf(i_0_675, plain, (m2_finseq_1(k2_matrlin(X1,X2),k3_finseq_2(X1))|v1_xboole_0(X1)|~m1_finseq_1(X2,X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_1883, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1268, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k1_tarski(np__1)))).#trainneg
cnf(i_0_98, plain, (m1_subset_1(np__0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1995, plain, (esk19_1(np__0)=np__0)).#trainneg
cnf(i_0_1093, plain, (k1_matrix_1(X1)=k3_finseq_1(k2_matrlin(X2,X3))|v1_xboole_0(X2)|k2_matrlin(X2,X3)!=k4_matrix_1(X2,X1)|~m1_finseq_1(X3,X2)|~v7_ordinal1(k3_finseq_1(X3))|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(k2_matrlin(X2,X3))|~v1_matrix_1(X1))).#trainneg
cnf(i_0_493, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_2312, plain, (r1_xxreal_0(esk25_0,esk22_0)|r1_xxreal_0(esk22_0,esk25_0))).#trainneg
cnf(i_0_2051, plain, (X1=k1_finseq_1(np__0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2128, plain, (X1=k2_finseq_1(np__0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_501, plain, (v1_funct_1(k1_tarski(k2_tarski(k2_tarski(X2,X1),k1_tarski(X1)))))).#trainneg
cnf(i_0_2131, plain, (k2_finseq_1(np__0)=np__0)).#trainneg
cnf(i_0_2049, plain, (k1_matrix_1(X1)=np__1|v1_xboole_0(X2)|k9_finseq_1(k9_finseq_1(X3))!=k4_matrix_1(X2,X1)|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(k9_finseq_1(k9_finseq_1(X3)))|~v1_matrix_1(X1)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_2040, plain, (m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|v1_xboole_0(X1)|~m1_finseq_1(k9_finseq_1(X2),X1))).#trainneg
cnf(i_0_2028, plain, (v1_finseq_1(k2_matrlin(X1,k9_finseq_1(X2)))|v1_xboole_0(X1)|~m1_finseq_1(k9_finseq_1(X2),X1))).#trainneg
cnf(i_0_2054, plain, (k1_finseq_1(np__0)=np__0)).#trainneg
cnf(i_0_1865, plain, (m1_matrix_1(k2_matrlin(X1,k4_matrix_1(X2,X3)),X1,np__1,k1_matrix_1(X3))|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_finseq_1(k4_matrix_1(X2,X3),X1)|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_1181, plain, (v3_card_1(k1_tarski(np__1),np__1))).#trainneg
cnf(i_0_1982, plain, (esk27_1(np__0)=np__0|~v1_card_1(np__0))).#trainneg
cnf(i_0_1274, plain, (v3_ordinal1(k1_tarski(np__1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_503, plain, (v1_funct_1(k1_tarski(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)))))).#trainneg
cnf(i_0_2005, plain, (v3_xxreal_0(X1)|v2_xxreal_0(esk22_0)|r1_xxreal_0(esk22_0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2012, plain, (v3_xxreal_0(X1)|v2_xxreal_0(np__0)|r1_xxreal_0(np__0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2027, plain, (v1_funct_1(k2_matrlin(X1,k9_finseq_1(X2)))|v1_xboole_0(X1)|~m1_finseq_1(k9_finseq_1(X2),X1))).#trainneg
cnf(i_0_2026, plain, (X1!=np__1|~v1_xboole_0(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2154, plain, (k2_zfmisc_1(X1,X2)=np__0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2210, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(k9_finseq_1(X3)))|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(k9_finseq_1(X3))))|~r2_hidden(X1,k1_tarski(np__1))|~v1_matrix_1(k9_finseq_1(X3)))).#trainneg
cnf(i_0_1959, plain, (v1_funct_1(X1)|~v4_funct_1(k1_tarski(X1)))).#trainneg
cnf(i_0_2204, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(k9_finseq_1(X2))))|~r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X3)),k2_matrix_1(k9_finseq_1(X2)))|~v1_matrix_1(k9_finseq_1(X2)))).#trainneg
cnf(i_0_1452, plain, (X1=k1_matrix_1(k3_finseq_1(X2))|r1_xxreal_0(k3_finseq_1(X2),np__0)|v1_xboole_0(X3)|~m1_matrix_1(k3_finseq_1(X2),X3,k3_finseq_1(X2),X1)|~v7_ordinal1(X1)|~m2_finseq_1(k3_finseq_1(X2),k3_finseq_2(X3))|~v1_matrix_1(k3_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1273, plain, (v1_funct_1(k1_tarski(np__1))|~v4_funct_1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2321, plain, (r1_xxreal_0(esk25_0,np__2)|r1_xxreal_0(np__2,esk25_0))).#trainneg
cnf(i_0_1278, plain, (v1_relat_1(k1_tarski(np__1))|~v1_relat_1(k5_numbers))).#trainneg
cnf(i_0_1924, plain, (r1_xxreal_0(X1,esk22_0)|r1_xxreal_0(esk22_0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1931, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_507, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_514, plain, (v5_relat_1(esk26_1(k2_zfmisc_1(X1,X2)),X2)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_1958, plain, (v1_relat_1(X1)|~v4_funct_1(k1_tarski(X1)))).#trainneg
cnf(i_0_2207, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(k9_finseq_1(X2))))|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X1,X3)),k2_matrix_1(k9_finseq_1(X2)))|~v1_matrix_1(k9_finseq_1(X2)))).#trainneg
cnf(i_0_2039, plain, (m2_finseq_1(k2_matrlin(X1,X2),k3_finseq_2(X1))|v1_xboole_0(X1)|~m1_finseq_1(X2,X1)|k3_finseq_1(X2)!=np__1|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1694, plain, (k4_matrix_1(X1,X2)=k4_matrix_1(X1,k4_matrix_1(X1,X3))|r2_hidden(k2_tarski(k1_tarski(esk5_3(X1,k4_matrix_1(X1,X3),k4_matrix_1(X1,X2))),k2_tarski(esk4_3(X1,k4_matrix_1(X1,X3),k4_matrix_1(X1,X2)),esk5_3(X1,k4_matrix_1(X1,X3),k4_matrix_1(X1,X2)))),k2_matrix_1(k4_matrix_1(X1,X3)))|r2_hidden(k2_tarski(k1_tarski(esk4_3(X1,k4_matrix_1(X1,X3),k4_matrix_1(X1,X2))),k2_tarski(esk4_3(X1,k4_matrix_1(X1,X3),k4_matrix_1(X1,X2)),esk5_3(X1,k4_matrix_1(X1,X3),k4_matrix_1(X1,X2)))),k2_matrix_1(k4_matrix_1(X1,X2)))|v1_xboole_0(X1)|k3_matrix_1(X1,k4_matrix_1(X1,X2),esk6_3(X1,k4_matrix_1(X1,X3),k4_matrix_1(X1,X2)),esk7_3(X1,k4_matrix_1(X1,X3),k4_matrix_1(X1,X2)))!=k3_matrix_1(X1,k4_matrix_1(X1,X3),esk7_3(X1,k4_matrix_1(X1,X3),k4_matrix_1(X1,X2)),esk6_3(X1,k4_matrix_1(X1,X3),k4_matrix_1(X1,X2)))|k1_matrix_1(k4_matrix_1(X1,X3))!=k3_finseq_1(k4_matrix_1(X1,X2))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,X2))|~v1_matrix_1(X3)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_2029, plain, (v1_relat_1(k2_matrlin(X1,k9_finseq_1(X2)))|v1_xboole_0(X1)|~m1_finseq_1(k9_finseq_1(X2),X1))).#trainneg
cnf(i_0_1571, plain, (X1=k3_finseq_1(X2)|k9_xtuple_0(X2)!=k1_tarski(np__1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|X1!=np__1)).#trainneg
cnf(i_0_2212, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k2_matrix_1(k9_finseq_1(X3)))|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(k9_finseq_1(X3))))|~r2_hidden(X2,k1_tarski(np__1))|~v1_matrix_1(k9_finseq_1(X3)))).#trainneg
cnf(i_0_1541, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X2,X1)),k2_matrix_1(X3))|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X1,k4_finseq_1(X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_1409, plain, (X1=k4_matrix_1(X2,k9_finseq_1(X3))|r2_hidden(k2_tarski(k1_tarski(esk5_3(X2,k9_finseq_1(X3),X1)),k2_tarski(esk4_3(X2,k9_finseq_1(X3),X1),esk5_3(X2,k9_finseq_1(X3),X1))),k2_matrix_1(k9_finseq_1(X3)))|r2_hidden(k2_tarski(k1_tarski(esk4_3(X2,k9_finseq_1(X3),X1)),k2_tarski(esk4_3(X2,k9_finseq_1(X3),X1),esk5_3(X2,k9_finseq_1(X3),X1))),k2_matrix_1(X1))|v1_xboole_0(X2)|k3_matrix_1(X2,X1,esk6_3(X2,k9_finseq_1(X3),X1),esk7_3(X2,k9_finseq_1(X3),X1))!=k3_matrix_1(X2,k9_finseq_1(X3),esk7_3(X2,k9_finseq_1(X3),X1),esk6_3(X2,k9_finseq_1(X3),X1))|k1_matrix_1(k9_finseq_1(X3))!=k3_finseq_1(X1)|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(k9_finseq_1(X3))|~v1_matrix_1(X1)|~m1_subset_1(X3,k3_finseq_2(X2)))).#trainneg
cnf(i_0_2782, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(np__0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_2494, plain, (v1_card_1(np__0))).#trainneg
cnf(i_0_515, plain, (v5_relat_1(esk24_1(k2_zfmisc_1(X1,X2)),X2)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2608, plain, (v3_card_1(np__0,np__0))).#trainneg
cnf(i_0_516, plain, (v5_relat_1(esk20_1(k2_zfmisc_1(X1,X2)),X2)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_517, plain, (v5_relat_1(esk16_1(k2_zfmisc_1(X1,X2)),X2))).#trainneg
cnf(i_0_2702, plain, (k2_zfmisc_1(k2_finseq_1(k1_matrix_1(X2)),k2_finseq_1(k1_matrix_1(k4_matrix_1(X1,X2))))=k2_matrix_1(k4_matrix_1(X1,X2))|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X1,X2))|~v1_funct_1(k4_matrix_1(X1,X2))|~v1_relat_1(k4_matrix_1(X1,X2))|~v1_matrix_1(k4_matrix_1(X1,X2))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_2472, plain, (v1_matrix_1(k9_finseq_1(X1))|v1_xboole_0(X2)|~v7_ordinal1(k3_finseq_1(X1))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_2609, plain, (k1_matrix_1(X1)=np__1|v1_xboole_0(X2)|k9_finseq_1(X3)!=k4_matrix_1(X2,X1)|~m1_finseq_1(X3,X2)|~v7_ordinal1(k3_finseq_1(X3))|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(k9_finseq_1(X3))|~v1_matrix_1(X1))).#trainneg
cnf(i_0_2606, plain, (v2_finseq_1(np__0))).#trainneg
cnf(i_0_1272, plain, (v1_relat_1(k1_tarski(np__1))|~v4_funct_1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_519, plain, (v5_relat_1(esk13_1(k2_zfmisc_1(X1,X2)),X2)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2767, plain, (v1_xboole_0(esk16_1(np__0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2777, plain, (v1_relat_1(esk16_1(np__0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_520, plain, (v5_relat_1(esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_523, plain, (v4_relat_1(esk26_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2732, plain, (v1_matrix_1(k2_matrlin(X1,k4_matrix_1(X2,X3)))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v7_ordinal1(k1_matrix_1(X3))|~m1_finseq_1(k4_matrix_1(X2,X3),X1)|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_2580, plain, (r1_xxreal_0(np__1,X1)|r1_xxreal_0(X1,np__1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2686, plain, (k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))=np__1|r1_xxreal_0(k1_matrix_1(k9_finseq_1(X2)),np__0)|v1_xboole_0(X1)|~v1_matrix_1(k9_finseq_1(X2))|~m1_finseq_1(X2,X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_2594, plain, (v4_relat_1(np__0,X1))).#trainneg
cnf(i_0_366, plain, (m1_subset_1(k2_tarski(np__1,np__2),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_524, plain, (v4_relat_1(esk24_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2322, plain, (r1_xxreal_0(esk25_0,np__1)|r1_xxreal_0(np__1,esk25_0))).#trainneg
cnf(i_0_2663, plain, (k1_tarski(np__1)!=np__0)).#trainneg
cnf(i_0_1932, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_525, plain, (v4_relat_1(esk20_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_526, plain, (v4_relat_1(esk16_1(k2_zfmisc_1(X1,X2)),X1))).#trainneg
cnf(i_0_2733, plain, (m2_finseq_1(k2_matrlin(X1,k4_matrix_1(X2,X3)),k3_finseq_2(X1))|v1_xboole_0(X1)|v1_xboole_0(X2)|~v7_ordinal1(k1_matrix_1(X3))|~m1_finseq_1(k4_matrix_1(X2,X3),X1)|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_2924, plain, (k1_matrix_1(X1)=k3_finseq_1(k2_matrlin(X2,X3))|v1_xboole_0(X2)|k2_matrlin(X2,X3)!=k4_matrix_1(X2,X1)|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(X1)|~v1_matrix_1(k2_matrlin(X2,X3))|k3_finseq_1(X3)!=np__1|~m1_finseq_1(X3,X2))).#trainneg
cnf(i_0_2693, plain, (k3_finseq_1(X1)=k3_finseq_1(k4_matrix_1(X2,X3))|v1_xboole_0(X2)|k2_finseq_1(k1_matrix_1(X3))!=k9_xtuple_0(X1)|~v1_finseq_1(k4_matrix_1(X2,X3))|~v1_finseq_1(X1)|~v1_funct_1(k4_matrix_1(X2,X3))|~v1_funct_1(X1)|~v1_relat_1(k4_matrix_1(X2,X3))|~v1_relat_1(X1)|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_1279, plain, (v1_funct_1(k1_tarski(np__1))|~v1_funct_1(k5_numbers)|~v1_relat_1(k5_numbers))).#trainneg
cnf(i_0_528, plain, (v4_relat_1(esk13_1(k2_zfmisc_1(X1,X2)),X1)|v1_xboole_0(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2325, plain, (r1_xxreal_0(esk17_0,esk22_0)|r1_xxreal_0(esk22_0,esk17_0))).#trainneg
cnf(i_0_2752, plain, (esk11_1(k5_numbers)!=np__1|~v1_xboole_0(esk11_1(k5_numbers)))).#trainneg
cnf(i_0_1933, plain, (r1_xxreal_0(X1,np__1)|r1_xxreal_0(np__1,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_529, plain, (v4_relat_1(esk11_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_535, plain, (m1_subset_1(esk3_2(X1,X2),X2)|esk3_2(X1,X2)=X1|X2=k1_tarski(X1))).#trainneg
cnf(i_0_2552, plain, (v1_finseq_1(k2_matrlin(X1,X2))|v1_xboole_0(X1)|~m1_finseq_1(X2,X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_2478, plain, (k1_matrix_1(k2_matrlin(X1,k9_finseq_1(X2)))=np__1|v1_xboole_0(X1)|k3_finseq_1(k2_matrlin(X1,k9_finseq_1(X2)))!=np__1|~m2_finseq_1(k2_matrlin(X1,k9_finseq_1(X2)),k3_finseq_2(X1))|~v1_matrix_1(k2_matrlin(X1,k9_finseq_1(X2)))|~m1_finseq_1(k9_finseq_1(X2),X1))).#trainneg
cnf(i_0_2961, plain, (X1=k3_finseq_1(X2)|k9_xtuple_0(X2)!=k1_tarski(np__1)|X1!=np__1|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1802, plain, (k2_tarski(k2_tarski(X1,X2),k1_tarski(X1))=X3|v1_xboole_0(X4)|v1_xboole_0(X5)|k2_zfmisc_1(X5,X4)!=k1_tarski(X3)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X5))).#trainneg
cnf(i_0_2174, plain, (k2_zfmisc_1(X1,X2)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2561, plain, (k1_matrix_1(k4_matrix_1(X1,k2_matrlin(X1,X2)))=k3_finseq_1(k2_matrlin(X1,X2))|r1_xxreal_0(k1_matrix_1(k2_matrlin(X1,X2)),np__0)|v1_xboole_0(X1)|~v1_matrix_1(k2_matrlin(X1,X2))|~m1_finseq_1(X2,X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_2182, plain, (v1_matrix_1(k2_matrlin(X1,X2))|v1_xboole_0(X1)|~m1_finseq_1(X2,X1)|k3_finseq_1(X2)!=np__1|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2551, plain, (v1_funct_1(k2_matrlin(X1,X2))|v1_xboole_0(X1)|~m1_finseq_1(X2,X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_1463, plain, (k3_matrix_1(k1_tarski(X1),X2,X3,X4)=X1|~m1_finseq_1(X2,k3_finseq_2(k1_tarski(X1)))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_1567, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_matrix_1(X3))|~r2_hidden(X2,k1_tarski(np__1))|~r2_hidden(X1,k9_xtuple_0(X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v1_matrix_1(X3)|k1_matrix_1(X3)!=np__1|~m2_subset_1(k1_matrix_1(X3),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1960, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_tarski(X1)))).#trainneg
cnf(i_0_2616, plain, (v3_xxreal_0(esk22_0)|r1_xxreal_0(esk25_0,esk22_0)|~v1_xreal_0(esk22_0))).#trainneg
cnf(i_0_1180, plain, (v3_card_1(k2_tarski(np__1,np__2),np__2))).#trainneg
cnf(i_0_542, plain, (m1_subset_1(X1,k2_zfmisc_1(X2,X3))|~r2_hidden(X1,esk14_2(X2,X3)))).#trainneg
cnf(i_0_2323, plain, (r1_xxreal_0(esk25_0,X1)|r1_xxreal_0(X1,esk25_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2336, plain, (r1_xxreal_0(esk17_0,X1)|r1_xxreal_0(X1,esk17_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_549, plain, (X1=k1_tarski(X2)|v1_xboole_0(X1)|esk3_2(X2,X1)!=X2|~m1_subset_1(esk3_2(X2,X1),X1))).#trainneg
cnf(i_0_554, plain, (v1_funct_1(esk26_1(X1))|v1_xboole_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1871, plain, (k9_xtuple_0(k4_matrix_1(X1,X2))=k2_finseq_1(X3)|v1_xboole_0(X1)|X3!=k1_matrix_1(X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~v1_finseq_1(k4_matrix_1(X1,X2))|~v1_funct_1(k4_matrix_1(X1,X2))|~v1_relat_1(k4_matrix_1(X1,X2))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_2553, plain, (v1_relat_1(k2_matrlin(X1,X2))|v1_xboole_0(X1)|~m1_finseq_1(X2,X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_1566, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3))|~r2_hidden(X2,k1_tarski(np__1))|~r2_hidden(X1,k9_xtuple_0(X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v1_matrix_1(X3)|k1_matrix_1(X3)!=np__1|~m2_subset_1(k1_matrix_1(X3),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1386, plain, (r1_xxreal_0(np__0,X1)|r1_xxreal_0(X1,np__0)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1400, plain, (m1_subset_1(k3_finseq_1(k3_pre_poly(X1,X2)),k5_numbers)|v1_xboole_0(X1)|~v1_funct_1(k3_pre_poly(X1,X2))|~v1_relat_1(k3_pre_poly(X1,X2))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1627, plain, (v7_ordinal1(k3_finseq_1(k4_matrix_1(X1,X2)))|v1_xboole_0(X1)|~v1_funct_1(k4_matrix_1(X1,X2))|~v1_relat_1(k4_matrix_1(X1,X2))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_2841, plain, (k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1))=k2_matrix_1(X1)|k3_finseq_1(X1)!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1)|k1_matrix_1(X1)!=np__1|~m2_subset_1(k1_matrix_1(X1),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2858, plain, (X1=k1_matrix_1(np__1)|v1_xboole_0(X3)|~m1_matrix_1(np__1,X3,np__1,X1)|~v7_ordinal1(X1)|~m2_finseq_1(np__1,k3_finseq_2(X3))|~v1_matrix_1(np__1))).#trainneg
cnf(i_0_2296, plain, (v3_ordinal1(X1)|~v1_card_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_964, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_2334, plain, (r1_xxreal_0(esk17_0,np__2)|r1_xxreal_0(np__2,esk17_0))).#trainneg
cnf(i_0_1189, plain, (m2_finseq_1(k2_matrlin(X1,k3_finseq_1(X2)),k3_finseq_2(X1))|v1_xboole_0(X1)|~m1_finseq_1(k3_finseq_1(X2),X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_555, plain, (v1_funct_1(esk24_1(X1))|v1_xboole_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2335, plain, (r1_xxreal_0(esk17_0,np__1)|r1_xxreal_0(np__1,esk17_0))).#trainneg
cnf(i_0_3201, plain, (k2_tarski(np__1,np__2)!=np__0|np__2!=np__1)).#trainneg
cnf(i_0_556, plain, (v1_funct_1(esk20_1(X1))|v1_xboole_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_557, plain, (v1_funct_1(esk16_1(X1))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3044, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|v1_xboole_0(X4)|~r2_hidden(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_matrix_1(k4_matrix_1(X4,X2)))|~m2_finseq_1(X2,k3_finseq_2(X4))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_3072, plain, (v1_matrix_1(k9_finseq_1(k2_matrlin(X1,k9_finseq_1(X2))))|v1_xboole_0(X1)|~v7_ordinal1(k3_finseq_1(k2_matrlin(X1,k9_finseq_1(X2))))|~m1_finseq_1(k9_finseq_1(X2),X1))).#trainneg
cnf(i_0_3236, plain, (v1_matrix_1(k9_finseq_1(k2_matrlin(X1,k4_matrix_1(X2,X3))))|v1_xboole_0(X2)|v1_xboole_0(X1)|~v7_ordinal1(k3_finseq_1(k2_matrlin(X1,k4_matrix_1(X2,X3))))|~m1_finseq_1(k4_matrix_1(X2,X3),X1)|~v7_ordinal1(k1_matrix_1(X3))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_3521, plain, (v7_ordinal1(k1_matrix_1(X2))|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_3070, plain, (v1_matrix_1(k9_finseq_1(k4_matrix_1(X1,X2)))|v1_xboole_0(X1)|~v7_ordinal1(k3_finseq_1(k4_matrix_1(X1,X2)))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_1289, plain, (v1_xboole_0(esk27_1(np__0)))).#trainneg
cnf(i_0_1982, plain, (esk27_1(np__0)=np__0)).#trainneg
cnf(i_0_2103, negated_conjecture, (esk2_0=np__1|k2_tarski(np__1,np__2)!=esk1_0|np__2!=np__1)).#trainneg
cnf(i_0_558, plain, (v1_funct_1(esk14_2(X1,X2))|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2567, plain, (r1_xxreal_0(np__1,esk22_0)|r1_xxreal_0(esk22_0,np__1))).#trainneg
cnf(i_0_3101, plain, (v1_xboole_0(esk16_1(np__0)))).#trainneg
cnf(i_0_3116, plain, (v1_relat_1(esk16_1(np__0)))).#trainneg
cnf(i_0_559, plain, (v1_funct_1(esk13_1(X1))|v1_xboole_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_560, plain, (v1_funct_1(esk11_1(k1_zfmisc_1(X1)))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3075, plain, (v1_matrix_1(k9_finseq_1(k3_pre_poly(X1,X2)))|v1_xboole_0(X1)|~v7_ordinal1(k3_finseq_1(k3_pre_poly(X1,X2)))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_3046, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|v1_xboole_0(X4)|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X3)),k2_matrix_1(k4_matrix_1(X4,X2)))|~m2_finseq_1(X2,k3_finseq_2(X4))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_3445, plain, (X1=X2|k3_matrix_1(k1_tarski(X3),X1,esk8_3(k1_tarski(X3),X1,X2),esk9_3(k1_tarski(X3),X1,X2))!=X3|k3_finseq_1(X1)!=k3_finseq_1(X2)|k1_matrix_1(X1)!=k1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(X1,k3_finseq_2(k1_tarski(X3)))|~v1_matrix_1(X2)|~v1_matrix_1(X1)|~m1_finseq_1(X2,k3_finseq_2(k1_tarski(X3)))|~v7_ordinal1(esk9_3(k1_tarski(X3),X1,X2))|~v7_ordinal1(esk8_3(k1_tarski(X3),X1,X2)))).#trainneg
cnf(i_0_3064, plain, (k2_zfmisc_1(k1_tarski(np__1),k2_finseq_1(k1_matrix_1(k4_matrix_1(X2,X1))))=k2_matrix_1(k4_matrix_1(X2,X1))|v1_xboole_0(X2)|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(X1)|k1_matrix_1(X1)!=np__1|~m2_subset_1(k1_matrix_1(X1),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2986, plain, (v1_matrix_1(k9_finseq_1(k9_finseq_1(X1)))|v1_xboole_0(X2)|~m1_finseq_1(X1,X2)|~v7_ordinal1(k3_finseq_1(X1)))).#trainneg
cnf(i_0_3045, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|v1_xboole_0(X4)|~r2_hidden(k2_tarski(k2_tarski(X3,X1),k1_tarski(X1)),k2_matrix_1(k4_matrix_1(X4,X2)))|~m2_finseq_1(X2,k3_finseq_2(X4))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_3237, plain, (k1_matrix_1(X1)=k3_finseq_1(k2_matrlin(X2,k4_matrix_1(X3,X4)))|v1_xboole_0(X2)|v1_xboole_0(X3)|k2_matrlin(X2,k4_matrix_1(X3,X4))!=k4_matrix_1(X2,X1)|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(X1)|~v1_matrix_1(k2_matrlin(X2,k4_matrix_1(X3,X4)))|~m1_finseq_1(k4_matrix_1(X3,X4),X2)|~v7_ordinal1(k1_matrix_1(X4))|~m2_finseq_1(X4,k3_finseq_2(X3))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_3151, plain, (esk9_3(X1,k9_finseq_1(X2),X3)=X4|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|k2_finseq_1(k1_matrix_1(k9_finseq_1(X2)))!=k1_tarski(X4)|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k3_finseq_1(X3)!=np__1|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_3413, plain, (v4_relat_1(esk16_1(np__0),X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1940, plain, (np__2=np__1|k1_tarski(np__1)!=k2_tarski(np__1,np__2))).#trainneg
cnf(i_0_562, plain, (v1_funct_1(k2_finseq_1(X1))|~v1_funct_1(k5_numbers)|~v1_relat_1(k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_3230, plain, (v4_relat_1(esk16_1(np__0),X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_3408, plain, (v5_relat_1(esk16_1(np__0),X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_571, plain, (v1_relat_1(k4_finseq_1(X1))|~v4_funct_1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3073, plain, (v1_matrix_1(k9_finseq_1(k2_matrlin(X1,X2)))|v1_xboole_0(X1)|~v7_ordinal1(k3_finseq_1(k2_matrlin(X1,X2)))|~m1_finseq_1(X2,X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_3047, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(k4_matrix_1(X2,X3))))|v1_xboole_0(X2)|~r2_hidden(k2_tarski(k2_tarski(X4,X1),k1_tarski(X4)),k2_matrix_1(k4_matrix_1(X2,X3)))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_3054, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(k4_matrix_1(X4,X3)))|v1_xboole_0(X4)|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(k4_matrix_1(X4,X3))))|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(X3)))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_3232, plain, (v1_finseq_1(k2_matrlin(X1,k4_matrix_1(X2,X3)))|v1_xboole_0(X2)|v1_xboole_0(X1)|~m1_finseq_1(k4_matrix_1(X2,X3),X1)|~v7_ordinal1(k1_matrix_1(X3))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_3321, plain, (k2_zfmisc_1(k4_finseq_1(k2_matrlin(X1,k9_finseq_1(X2))),k1_tarski(np__1))=k2_matrix_1(k2_matrlin(X1,k9_finseq_1(X2)))|v1_xboole_0(X1)|~v1_finseq_1(k2_matrlin(X1,k9_finseq_1(X2)))|~v1_funct_1(k2_matrlin(X1,k9_finseq_1(X2)))|~v1_relat_1(k2_matrlin(X1,k9_finseq_1(X2)))|~v1_matrix_1(k2_matrlin(X1,k9_finseq_1(X2)))|k3_finseq_1(k2_matrlin(X1,k9_finseq_1(X2)))!=np__1|~m1_finseq_1(k9_finseq_1(X2),X1))).#trainneg
cnf(i_0_3325, plain, (k2_zfmisc_1(k9_xtuple_0(k2_matrlin(X1,k9_finseq_1(X2))),k1_tarski(np__1))=k2_matrix_1(k2_matrlin(X1,k9_finseq_1(X2)))|v1_xboole_0(X1)|~v1_finseq_1(k2_matrlin(X1,k9_finseq_1(X2)))|~v1_funct_1(k2_matrlin(X1,k9_finseq_1(X2)))|~v1_relat_1(k2_matrlin(X1,k9_finseq_1(X2)))|~v1_matrix_1(k2_matrlin(X1,k9_finseq_1(X2)))|k3_finseq_1(k2_matrlin(X1,k9_finseq_1(X2)))!=np__1|~m1_finseq_1(k9_finseq_1(X2),X1))).#trainneg
cnf(i_0_3536, plain, (v1_finseq_1(k2_matrlin(X1,k3_finseq_1(X2)))|v1_xboole_0(X1)|~m1_finseq_1(k3_finseq_1(X2),X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_3163, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,esk9_3(X2,k9_finseq_1(X3),X4))),k2_matrix_1(k9_finseq_1(X3)))|k9_finseq_1(X3)=X4|v1_xboole_0(X2)|~r2_hidden(X1,k4_finseq_1(k9_finseq_1(X3)))|~v1_matrix_1(k9_finseq_1(X3))|k1_matrix_1(k9_finseq_1(X3))!=k1_matrix_1(X4)|k3_finseq_1(X4)!=np__1|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X2))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_3231, plain, (v1_funct_1(k2_matrlin(X1,k4_matrix_1(X2,X3)))|v1_xboole_0(X2)|v1_xboole_0(X1)|~m1_finseq_1(k4_matrix_1(X2,X3),X1)|~v7_ordinal1(k1_matrix_1(X3))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_3157, plain, (r2_hidden(k2_tarski(k2_tarski(esk9_3(X1,k9_finseq_1(X2),X3),X4),k1_tarski(X4)),k2_matrix_1(k9_finseq_1(X2)))|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|~r2_hidden(X4,k1_tarski(np__1))|~v1_matrix_1(k9_finseq_1(X2))|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k3_finseq_1(X3)!=np__1|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_3188, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k2_tarski(np__1,np__2)))).#trainneg
cnf(i_0_2104, plain, (np__0=np__1|k2_tarski(np__1,np__2)!=k5_numbers|np__2!=np__1)).#trainneg
cnf(i_0_572, plain, (v1_funct_1(k4_finseq_1(X1))|~v4_funct_1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2742, plain, (v3_ordinal1(k1_tarski(np__1))|~v1_card_1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_2849, plain, (v1_funct_1(np__1)|~v4_funct_1(k2_tarski(np__1,np__2))|np__2!=np__1)).#trainneg
cnf(i_0_573, plain, (v3_ordinal1(k4_finseq_1(X1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_576, plain, (v1_relat_1(k4_finseq_1(X1))|~v1_relat_1(k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2984, plain, (v1_matrix_1(k9_finseq_1(X1))|v1_xboole_0(X2)|~m2_finseq_1(X1,X2)|k3_finseq_1(X1)!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3156, plain, (r2_hidden(k2_tarski(k2_tarski(X1,esk9_3(X2,k9_finseq_1(X3),X4)),k1_tarski(X1)),k2_matrix_1(k9_finseq_1(X3)))|k9_finseq_1(X3)=X4|v1_xboole_0(X2)|~r2_hidden(X1,k1_tarski(np__1))|~v1_matrix_1(k9_finseq_1(X3))|k1_matrix_1(k9_finseq_1(X3))!=k1_matrix_1(X4)|k3_finseq_1(X4)!=np__1|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X2))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_3087, plain, (k1_matrix_1(X1)=np__1|v1_xboole_0(X2)|k9_finseq_1(X3)!=k4_matrix_1(X2,X1)|~v7_ordinal1(k3_finseq_1(X3))|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(X1)|~m2_finseq_1(X3,X2))).#trainneg
cnf(i_0_3233, plain, (v1_relat_1(k2_matrlin(X1,k4_matrix_1(X2,X3)))|v1_xboole_0(X2)|v1_xboole_0(X1)|~m1_finseq_1(k4_matrix_1(X2,X3),X1)|~v7_ordinal1(k1_matrix_1(X3))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_3318, plain, (v7_ordinal1(k3_finseq_1(k2_matrlin(X1,X2)))|v1_xboole_0(X1)|~v1_funct_1(k2_matrlin(X1,X2))|~v1_relat_1(k2_matrlin(X1,X2))|~m1_finseq_1(X2,X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_2932, plain, (m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|v1_xboole_0(X1)|k3_finseq_1(X2)!=np__1|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_3691, plain, (r2_hidden(esk8_3(X1,k4_matrix_1(X2,X3),X4),k2_finseq_1(k1_matrix_1(X3)))|v1_xboole_0(X2)|k4_matrix_1(X2,X3)=X4|v1_xboole_0(X1)|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3)|k3_finseq_1(k4_matrix_1(X2,X3))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,X3))!=k1_matrix_1(X4)|~m2_finseq_1(X4,k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X2,X3),k3_finseq_2(X1))|~v1_matrix_1(X4)|~v1_matrix_1(k4_matrix_1(X2,X3)))).#trainneg
cnf(i_0_3455, plain, (v3_ordinal1(X1)|~v1_card_1(k1_tarski(X1)))).#trainneg
cnf(i_0_3777, plain, (v4_relat_1(esk16_1(np__0),X1))).#trainneg
cnf(i_0_3670, plain, (esk16_1(np__0)=np__0)).#trainneg
cnf(i_0_2105, plain, (np__0=np__1|k2_tarski(np__1,np__2)!=k1_numbers|np__2!=np__1)).#trainneg
cnf(i_0_577, plain, (v1_funct_1(k4_finseq_1(X1))|~v1_funct_1(k5_numbers)|~v1_relat_1(k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2911, plain, (v1_relat_1(np__1)|~v4_funct_1(k2_tarski(np__1,np__2))|np__2!=np__1)).#trainneg
cnf(i_0_4001, plain, (v7_ordinal1(k1_tarski(np__1))|~v1_card_1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3946, plain, (v1_finseq_1(k2_matrlin(X1,np__0))|v1_xboole_0(X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_583, plain, (v1_relat_1(k1_matrix_1(X1))|~v4_funct_1(k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_584, plain, (v1_funct_1(k1_matrix_1(X1))|~v4_funct_1(k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_3988, plain, (r2_hidden(esk9_3(X1,k9_finseq_1(X2),X3),k4_finseq_1(k9_finseq_1(X2)))|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|~v1_matrix_1(k9_finseq_1(X2))|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k3_finseq_1(X3)!=np__1|~r2_hidden(esk9_3(X1,k9_finseq_1(X2),X3),k1_tarski(np__1))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_3697, plain, (X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|k1_matrix_1(X1)!=k1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(k1_tarski(X3)))|~v7_ordinal1(esk9_3(k1_tarski(X3),X1,X2))|~v7_ordinal1(esk8_3(k1_tarski(X3),X1,X2))|~m2_finseq_1(X1,k3_finseq_2(k1_tarski(X3)))|~v1_matrix_1(X2)|~v1_matrix_1(X1)|~m1_finseq_1(X1,k3_finseq_2(k1_tarski(X3))))).#trainneg
cnf(i_0_735, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk16_1(X2))|~m1_subset_1(X1,esk16_1(X2)))).#trainneg
cnf(i_0_2487, plain, (v3_xxreal_0(esk22_0)|v2_xxreal_0(np__0)|r1_xxreal_0(np__0,esk22_0)|~v1_xreal_0(esk22_0))).#trainneg
cnf(i_0_3341, plain, (esk11_1(k5_numbers)=k3_finseq_1(X1)|k9_xtuple_0(X1)!=k1_tarski(np__1)|esk11_1(k5_numbers)!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_616, plain, (v5_relat_1(esk26_1(X1),X2)|v1_xboole_0(X1)|~v5_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1918, plain, (v7_ordinal1(k3_finseq_1(np__0))|~v1_funct_1(np__0)|~v1_relat_1(np__0))).#trainneg
cnf(i_0_2440, plain, (v3_xxreal_0(np__0)|v2_xxreal_0(X1)|r1_xxreal_0(X1,np__0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_617, plain, (v5_relat_1(esk24_1(X1),X2)|v1_xboole_0(X1)|~v5_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_618, plain, (v5_relat_1(esk20_1(X1),X2)|v1_xboole_0(X1)|~v5_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3990, plain, (r2_hidden(k2_tarski(k2_tarski(esk9_3(X1,k9_finseq_1(X2),X3),X4),k1_tarski(esk9_3(X1,k9_finseq_1(X2),X3))),k2_matrix_1(X5))|v1_xboole_0(X6)|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|k9_finseq_1(X2)!=k4_matrix_1(X6,X5)|~v7_ordinal1(esk9_3(X1,k9_finseq_1(X2),X3))|~v7_ordinal1(X4)|~m2_finseq_1(X5,k3_finseq_2(X6))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X6))|~v1_matrix_1(X5)|~v1_matrix_1(k9_finseq_1(X2))|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k3_finseq_1(X3)!=np__1|~r2_hidden(X4,k1_tarski(np__1))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_3568, plain, (k1_matrix_1(k9_finseq_1(k2_matrlin(X1,k9_finseq_1(X2))))=k3_finseq_1(k2_matrlin(X1,k9_finseq_1(X2)))|v1_xboole_0(X1)|~v7_ordinal1(k3_finseq_1(k2_matrlin(X1,k9_finseq_1(X2))))|~m1_finseq_1(k9_finseq_1(X2),X1))).#trainneg
cnf(i_0_968, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_619, plain, (v5_relat_1(esk16_1(X1),X2)|~v5_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2617, plain, (r1_xxreal_0(esk25_0,esk22_0)|~v2_xxreal_0(esk22_0)|~v1_xreal_0(esk22_0))).#trainneg
cnf(i_0_1152, plain, (r1_xxreal_0(esk25_0,X1)|r1_xxreal_0(X1,esk25_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_620, plain, (v5_relat_1(esk14_2(X1,X2),X3)|~v5_relat_1(k2_zfmisc_1(X1,X2),X3))).#trainneg
cnf(i_0_621, plain, (v5_relat_1(esk13_1(X1),X2)|v1_xboole_0(X1)|~v5_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3989, plain, (r2_hidden(k2_tarski(k2_tarski(esk9_3(X1,k9_finseq_1(X2),X3),X4),k1_tarski(esk9_3(X1,k9_finseq_1(X2),X3))),k2_matrix_1(X5))|v1_xboole_0(X6)|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|X5!=k4_matrix_1(X6,k9_finseq_1(X2))|~v7_ordinal1(esk9_3(X1,k9_finseq_1(X2),X3))|~v7_ordinal1(X4)|~m2_finseq_1(X5,k3_finseq_2(X6))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X6))|~v1_matrix_1(X5)|~v1_matrix_1(k9_finseq_1(X2))|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k3_finseq_1(X3)!=np__1|~r2_hidden(X4,k1_tarski(np__1))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_3955, plain, (r2_hidden(k2_tarski(k2_tarski(esk9_3(X1,k9_finseq_1(X2),X3),X4),k1_tarski(esk9_3(X1,k9_finseq_1(X2),X3))),k2_matrix_1(X5))|v1_xboole_0(X6)|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|k9_finseq_1(X2)!=k4_matrix_1(X6,X5)|~v7_ordinal1(esk9_3(X1,k9_finseq_1(X2),X3))|~v7_ordinal1(X4)|~m2_finseq_1(X5,k3_finseq_2(X6))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X6))|~v1_matrix_1(X5)|~v1_matrix_1(k9_finseq_1(X2))|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k3_finseq_1(X3)!=np__1|~r2_hidden(X4,k4_finseq_1(k9_finseq_1(X2)))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_3623, plain, (v7_ordinal1(k1_matrix_1(k4_matrix_1(X1,X2)))|v1_xboole_0(X1)|~v1_matrix_1(k4_matrix_1(X1,X2))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_3954, plain, (r2_hidden(k2_tarski(k2_tarski(esk9_3(X1,k9_finseq_1(X2),X3),X4),k1_tarski(esk9_3(X1,k9_finseq_1(X2),X3))),k2_matrix_1(X5))|v1_xboole_0(X6)|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|X5!=k4_matrix_1(X6,k9_finseq_1(X2))|~v7_ordinal1(esk9_3(X1,k9_finseq_1(X2),X3))|~v7_ordinal1(X4)|~m2_finseq_1(X5,k3_finseq_2(X6))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X6))|~v1_matrix_1(X5)|~v1_matrix_1(k9_finseq_1(X2))|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k3_finseq_1(X3)!=np__1|~r2_hidden(X4,k4_finseq_1(k9_finseq_1(X2)))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_3937, plain, (k2_matrix_1(k2_matrlin(X1,k9_finseq_1(X2)))=k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1))|v1_xboole_0(X1)|k3_finseq_1(k2_matrlin(X1,k9_finseq_1(X2)))!=np__1|~m1_finseq_1(k9_finseq_1(X2),X1)|~v1_finseq_1(k2_matrlin(X1,k9_finseq_1(X2)))|~v1_funct_1(k2_matrlin(X1,k9_finseq_1(X2)))|~v1_relat_1(k2_matrlin(X1,k9_finseq_1(X2))))).#trainneg
cnf(i_0_4014, plain, (v1_matrix_1(k9_finseq_1(k2_matrlin(X1,k9_finseq_1(X2))))|v1_xboole_0(X1)|k3_finseq_1(k2_matrlin(X1,k9_finseq_1(X2)))!=np__1|~m1_finseq_1(k9_finseq_1(X2),X1))).#trainneg
cnf(i_0_2781, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(np__0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_1620, plain, (m2_finseq_1(esk11_1(k3_finseq_2(X1)),X1))).#trainneg
cnf(i_0_2403, plain, (X1=np__1|k1_tarski(X1)!=k2_tarski(np__1,np__2)|np__2!=np__1)).#trainneg
cnf(i_0_622, plain, (v5_relat_1(esk11_1(k1_zfmisc_1(X1)),X2)|~v5_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2488, plain, (v2_xxreal_0(np__0)|r1_xxreal_0(np__0,esk22_0)|~v2_xxreal_0(esk22_0)|~v1_xreal_0(esk22_0))).#trainneg
cnf(i_0_2340, plain, (v1_relat_1(k2_tarski(np__1,np__2))|~v1_relat_1(k5_numbers))).#trainneg
cnf(i_0_1223, plain, (r1_xxreal_0(esk17_0,X1)|r1_xxreal_0(X1,esk17_0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_624, plain, (v5_relat_1(k4_finseq_1(X1),X2)|~v5_relat_1(k5_numbers,X2)|~v1_relat_1(k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_625, plain, (v5_relat_1(k2_finseq_1(X1),X2)|~v5_relat_1(k5_numbers,X2)|~v1_relat_1(k5_numbers)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4049, plain, (k1_matrix_1(X1)=np__1|v1_xboole_0(X2)|k9_finseq_1(k9_finseq_1(X3))!=k4_matrix_1(X2,X1)|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(X1)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_3735, plain, (r2_hidden(X1,k1_tarski(np__1))|v1_xboole_0(X3)|~r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(k4_matrix_1(X3,X4)))|k1_matrix_1(X4)!=np__1|~m2_finseq_1(X4,k3_finseq_2(X3))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_3566, plain, (k1_matrix_1(k9_finseq_1(k4_matrix_1(X1,X2)))=k3_finseq_1(k4_matrix_1(X1,X2))|v1_xboole_0(X1)|~v7_ordinal1(k3_finseq_1(k4_matrix_1(X1,X2)))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_3571, plain, (k1_matrix_1(k9_finseq_1(k2_matrlin(X1,X2)))=k3_finseq_1(k2_matrlin(X1,X2))|v1_xboole_0(X1)|~v7_ordinal1(k3_finseq_1(k2_matrlin(X1,X2)))|~m1_finseq_1(X2,X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_3981, plain, (r2_hidden(k2_tarski(k2_tarski(esk9_3(X1,k9_finseq_1(X2),X3),esk9_3(X1,k9_finseq_1(X2),X3)),k1_tarski(esk9_3(X1,k9_finseq_1(X2),X3))),k2_matrix_1(X4))|v1_xboole_0(X5)|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|k9_finseq_1(X2)!=k4_matrix_1(X5,X4)|~v7_ordinal1(esk9_3(X1,k9_finseq_1(X2),X3))|~m2_finseq_1(X4,k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X5))|~v1_matrix_1(X4)|~v1_matrix_1(k9_finseq_1(X2))|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k3_finseq_1(X3)!=np__1|~r2_hidden(esk9_3(X1,k9_finseq_1(X2),X3),k1_tarski(np__1))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_4088, plain, (k1_matrix_1(X1)=np__1|v1_xboole_0(X3)|k9_finseq_1(X2)!=k4_matrix_1(X3,X1)|~m2_finseq_1(X1,k3_finseq_2(X3))|~v1_matrix_1(X1)|~v1_matrix_1(k9_finseq_1(X2))|k3_finseq_1(X2)!=np__1|~m1_finseq_1(X2,X3))).#trainneg
cnf(i_0_3963, plain, (r2_hidden(k2_tarski(k2_tarski(esk9_3(X1,k9_finseq_1(X2),X3),esk9_3(X1,k9_finseq_1(X2),X3)),k1_tarski(esk9_3(X1,k9_finseq_1(X2),X3))),k2_matrix_1(X4))|v1_xboole_0(X5)|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|X4!=k4_matrix_1(X5,k9_finseq_1(X2))|~v7_ordinal1(esk9_3(X1,k9_finseq_1(X2),X3))|~m2_finseq_1(X4,k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X5))|~v1_matrix_1(X4)|~v1_matrix_1(k9_finseq_1(X2))|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k3_finseq_1(X3)!=np__1|~r2_hidden(esk9_3(X1,k9_finseq_1(X2),X3),k4_finseq_1(k9_finseq_1(X2)))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_3685, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|v1_xboole_0(X3)|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X4,X1)),k2_matrix_1(k4_matrix_1(X3,X2)))|~m2_finseq_1(X2,k3_finseq_2(X3))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_745, plain, (k1_finseq_1(X1)=esk14_2(X2,X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3143, plain, (r1_xxreal_0(esk11_1(k5_numbers),np__1)|r1_xxreal_0(np__1,esk11_1(k5_numbers)))).#trainneg
cnf(i_0_627, plain, (v5_relat_1(X1,X2)|~v5_relat_1(k2_zfmisc_1(k5_numbers,X3),X2)|~m2_finseq_1(X1,X3))).#trainneg
cnf(i_0_2120, plain, (v3_ordinal1(k2_tarski(np__1,np__2))|~v3_ordinal1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1487, plain, (r1_xxreal_0(X1,esk22_0)|r1_xxreal_0(esk22_0,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_635, plain, (v1_matrix_1(k2_matrlin(X1,X2))|v1_xboole_0(X1)|~v7_ordinal1(k3_finseq_1(X2))|~v7_ordinal1(np__1)|~m1_finseq_1(X2,X1))).#trainneg
cnf(i_0_658, plain, (v1_relat_1(k2_matrix_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_3980, plain, (r2_hidden(k2_tarski(k2_tarski(esk9_3(X1,k9_finseq_1(X2),X3),esk9_3(X1,k9_finseq_1(X2),X3)),k1_tarski(esk9_3(X1,k9_finseq_1(X2),X3))),k2_matrix_1(X4))|v1_xboole_0(X5)|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|X4!=k4_matrix_1(X5,k9_finseq_1(X2))|~v7_ordinal1(esk9_3(X1,k9_finseq_1(X2),X3))|~m2_finseq_1(X4,k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X5))|~v1_matrix_1(X4)|~v1_matrix_1(k9_finseq_1(X2))|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k3_finseq_1(X3)!=np__1|~r2_hidden(esk9_3(X1,k9_finseq_1(X2),X3),k1_tarski(np__1))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_3964, plain, (r2_hidden(k2_tarski(k2_tarski(esk9_3(X1,k9_finseq_1(X2),X3),esk9_3(X1,k9_finseq_1(X2),X3)),k1_tarski(esk9_3(X1,k9_finseq_1(X2),X3))),k2_matrix_1(X4))|v1_xboole_0(X5)|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|k9_finseq_1(X2)!=k4_matrix_1(X5,X4)|~v7_ordinal1(esk9_3(X1,k9_finseq_1(X2),X3))|~m2_finseq_1(X4,k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X5))|~v1_matrix_1(X4)|~v1_matrix_1(k9_finseq_1(X2))|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k3_finseq_1(X3)!=np__1|~r2_hidden(esk9_3(X1,k9_finseq_1(X2),X3),k4_finseq_1(k9_finseq_1(X2)))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_4050, plain, (k1_matrix_1(X1)=np__1|v1_xboole_0(X2)|k9_finseq_1(k4_matrix_1(X2,X3))!=k4_matrix_1(k3_finseq_2(X2),X1)|~v7_ordinal1(k3_finseq_1(k4_matrix_1(X2,X3)))|~m2_finseq_1(X1,k3_finseq_2(k3_finseq_2(X2)))|~v1_matrix_1(X1)|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_3573, plain, (k1_matrix_1(k9_finseq_1(k3_pre_poly(X1,X2)))=k3_finseq_1(k3_pre_poly(X1,X2))|v1_xboole_0(X1)|~v7_ordinal1(k3_finseq_1(k3_pre_poly(X1,X2)))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_747, plain, (k1_card_1(X1)=esk14_2(X2,X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_659, plain, (v1_xboole_0(k4_finseq_1(X1))|v1_xboole_0(k2_finseq_1(k1_matrix_1(X1)))|~v1_xboole_0(k2_matrix_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_4506, plain, (r1_xxreal_0(esk22_0,esk22_0))).#trainneg
cnf(i_0_4129, plain, (~v1_subset_1(np__0,np__0))).#trainneg
cnf(i_0_660, plain, (m1_subset_1(esk14_2(k4_finseq_1(X1),k2_finseq_1(k1_matrix_1(X1))),k1_zfmisc_1(k2_matrix_1(X1)))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_661, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_matrix_1(X2)))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_4536, negated_conjecture, (k1_matrix_1(k9_finseq_1(k9_finseq_1(esk2_0)))=np__1)).#trainneg
cnf(i_0_4372, negated_conjecture, (k1_matrix_1(X1)=np__1|k9_finseq_1(k9_finseq_1(esk2_0))!=k4_matrix_1(esk1_0,X1)|~m2_finseq_1(X1,k3_finseq_2(esk1_0))|~v1_matrix_1(X1))).#trainneg
cnf(i_0_4106, plain, (esk8_3(X1,k4_matrix_1(X2,X3),X4)=X5|k4_matrix_1(X2,X3)=X4|v1_xboole_0(X1)|v1_xboole_0(X2)|k2_finseq_1(k1_matrix_1(X3))!=k1_tarski(X5)|k3_finseq_1(k4_matrix_1(X2,X3))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,X3))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X2,X3),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X3)|~v1_matrix_1(X4))).#trainneg
cnf(i_0_4115, plain, (r2_hidden(k2_tarski(k2_tarski(X1,esk8_3(X2,k4_matrix_1(X3,X4),X5)),k1_tarski(X1)),k2_matrix_1(X4))|k4_matrix_1(X3,X4)=X5|v1_xboole_0(X2)|v1_xboole_0(X3)|~r2_hidden(X1,k4_finseq_1(X4))|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4)|~v1_matrix_1(X4)|k3_finseq_1(k4_matrix_1(X3,X4))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X3,X4))!=k1_matrix_1(X5)|~m2_finseq_1(k4_matrix_1(X3,X4),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X5,k3_finseq_2(X2))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_4116, plain, (r2_hidden(k2_tarski(k2_tarski(X1,esk8_3(X2,k4_matrix_1(X3,X4),X5)),k1_tarski(X1)),k2_matrix_1(X4))|k4_matrix_1(X3,X4)=X5|v1_xboole_0(X2)|v1_xboole_0(X3)|~r2_hidden(X1,k9_xtuple_0(X4))|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4)|~v1_matrix_1(X4)|k3_finseq_1(k4_matrix_1(X3,X4))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X3,X4))!=k1_matrix_1(X5)|~m2_finseq_1(k4_matrix_1(X3,X4),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X5,k3_finseq_2(X2))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_4103, plain, (v1_matrix_1(k9_finseq_1(X1))|k3_finseq_1(X1)!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4111, plain, (r2_hidden(k2_tarski(k2_tarski(X1,esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5)),k1_tarski(X1)),k2_matrix_1(k9_finseq_1(X4)))|k4_matrix_1(X3,k9_finseq_1(X4))=X5|v1_xboole_0(X2)|v1_xboole_0(X3)|~r2_hidden(X1,k1_tarski(np__1))|~v1_matrix_1(k9_finseq_1(X4))|k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(X4)))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(X4)))!=k1_matrix_1(X5)|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(X4)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X4),k3_finseq_2(X3))|~m2_finseq_1(X5,k3_finseq_2(X2))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_4174, plain, (X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|k1_matrix_1(X1)!=k1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(k1_tarski(X3)))|~m1_finseq_1(X1,k3_finseq_2(k1_tarski(X3)))|~v7_ordinal1(esk8_3(k1_tarski(X3),X1,X2))|~v1_matrix_1(X2)|~v1_matrix_1(X1)|~m1_subset_1(esk9_3(k1_tarski(X3),X1,X2),k5_numbers))).#trainneg
cnf(i_0_4264, plain, (k9_finseq_1(X1)=X2|r2_hidden(k2_tarski(k2_tarski(esk9_3(X3,k9_finseq_1(X1),X2),X4),k1_tarski(esk9_3(X3,k9_finseq_1(X1),X2))),k2_matrix_1(k4_matrix_1(X5,k9_finseq_1(X1))))|v1_xboole_0(X3)|v1_xboole_0(X5)|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(X2)|k3_finseq_1(X2)!=np__1|~r2_hidden(X4,k1_tarski(np__1))|~v7_ordinal1(esk9_3(X3,k9_finseq_1(X1),X2))|~v7_ordinal1(X4)|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X5,k9_finseq_1(X1)),k3_finseq_2(X5))|~m2_finseq_1(X2,k3_finseq_2(X3))|~v1_matrix_1(k9_finseq_1(X1))|~v1_matrix_1(k4_matrix_1(X5,k9_finseq_1(X1)))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_662, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_matrix_1(X2)))|~v1_xboole_0(k2_finseq_1(k1_matrix_1(X2)))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_3180, plain, (v1_relat_1(k2_tarski(np__1,np__2))|~v4_funct_1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3181, plain, (v1_funct_1(k2_tarski(np__1,np__2))|~v4_funct_1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_4500, plain, (v5_relat_1(X1,X2)|~m2_finseq_1(X1,X3)|~v1_xboole_0(X3))).#trainneg
cnf(i_0_4608, plain, (v1_card_1(esk14_2(X2,X3))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_663, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_matrix_1(X2)))|~v1_xboole_0(k4_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_664, plain, (v4_relat_1(X1,k4_finseq_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_matrix_1(X2)))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_4279, plain, (k9_finseq_1(X1)=X2|r2_hidden(k2_tarski(k2_tarski(esk9_3(X3,k9_finseq_1(X1),X2),X4),k1_tarski(esk9_3(X3,k9_finseq_1(X1),X2))),k2_matrix_1(k4_matrix_1(X5,k9_finseq_1(X1))))|v1_xboole_0(X3)|v1_xboole_0(X5)|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(X2)|k3_finseq_1(X2)!=np__1|~r2_hidden(X4,k4_finseq_1(k9_finseq_1(X1)))|~v7_ordinal1(esk9_3(X3,k9_finseq_1(X1),X2))|~v7_ordinal1(X4)|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X5,k9_finseq_1(X1)),k3_finseq_2(X5))|~m2_finseq_1(X2,k3_finseq_2(X3))|~v1_matrix_1(k9_finseq_1(X1))|~v1_matrix_1(k4_matrix_1(X5,k9_finseq_1(X1)))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_4112, plain, (r2_hidden(k2_tarski(k2_tarski(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),X5),k1_tarski(X5)),k2_matrix_1(k9_finseq_1(X3)))|k4_matrix_1(X2,k9_finseq_1(X3))=X4|v1_xboole_0(X1)|v1_xboole_0(X2)|~r2_hidden(X5,k1_tarski(np__1))|~v1_matrix_1(k9_finseq_1(X3))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_4175, plain, (X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|k1_matrix_1(X1)!=k1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(k1_tarski(X3)))|~m1_finseq_1(X1,k3_finseq_2(k1_tarski(X3)))|~v7_ordinal1(esk8_3(k1_tarski(X3),X1,X2))|~v1_matrix_1(X2)|~v1_matrix_1(X1)|~m2_finseq_1(X2,k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(X1,k3_finseq_2(k1_tarski(X3))))).#trainneg
cnf(i_0_4120, plain, (r2_hidden(k2_tarski(k2_tarski(esk8_3(X1,k4_matrix_1(X2,X3),X4),X5),k1_tarski(X5)),k2_matrix_1(X3))|k4_matrix_1(X2,X3)=X4|v1_xboole_0(X1)|v1_xboole_0(X2)|~r2_hidden(X5,k4_finseq_1(X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v1_matrix_1(X3)|k3_finseq_1(k4_matrix_1(X2,X3))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,X3))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X2,X3),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_4236, plain, (k4_matrix_1(X1,k9_finseq_1(k2_matrlin(X2,k9_finseq_1(X3))))=X4|r2_hidden(esk8_3(X5,k4_matrix_1(X1,k9_finseq_1(k2_matrlin(X2,k9_finseq_1(X3)))),X4),k2_finseq_1(k3_finseq_1(k2_matrlin(X2,k9_finseq_1(X3)))))|v1_xboole_0(X5)|v1_xboole_0(X1)|v1_xboole_0(X2)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k2_matrlin(X2,k9_finseq_1(X3)))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k2_matrlin(X2,k9_finseq_1(X3)))))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k2_matrlin(X2,k9_finseq_1(X3)))),k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(k2_matrlin(X2,k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(X4,k3_finseq_2(X5))|~v1_matrix_1(k9_finseq_1(k2_matrlin(X2,k9_finseq_1(X3))))|~v1_matrix_1(X4)|~m1_finseq_1(k9_finseq_1(X3),X2)|~v7_ordinal1(k3_finseq_1(k2_matrlin(X2,k9_finseq_1(X3)))))).#trainneg
cnf(i_0_4125, plain, (k4_matrix_1(X1,X2)=X3|r2_hidden(esk8_3(X4,k4_matrix_1(X1,X2),X3),k1_tarski(np__1))|v1_xboole_0(X4)|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,X2))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,X2))!=k1_matrix_1(X3)|~m2_finseq_1(k4_matrix_1(X1,X2),k3_finseq_2(X4))|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(X2)|~v1_matrix_1(X3)|k1_matrix_1(X2)!=np__1|~m2_subset_1(k1_matrix_1(X2),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_4119, plain, (r2_hidden(k2_tarski(k2_tarski(X1,esk8_3(X2,k4_matrix_1(X3,k4_matrix_1(X4,X5)),X6)),k1_tarski(X1)),k2_matrix_1(k4_matrix_1(X4,X5)))|v1_xboole_0(X4)|k4_matrix_1(X3,k4_matrix_1(X4,X5))=X6|v1_xboole_0(X2)|v1_xboole_0(X3)|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(X5)))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X5)|k3_finseq_1(k4_matrix_1(X3,k4_matrix_1(X4,X5)))!=k3_finseq_1(X6)|k1_matrix_1(k4_matrix_1(X3,k4_matrix_1(X4,X5)))!=k1_matrix_1(X6)|~m2_finseq_1(k4_matrix_1(X3,k4_matrix_1(X4,X5)),k3_finseq_2(X2))|~m2_finseq_1(k4_matrix_1(X4,X5),k3_finseq_2(X3))|~m2_finseq_1(X6,k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X4,X5))|~v1_matrix_1(X6))).#trainneg
cnf(i_0_4342, plain, (v1_xboole_0(esk11_1(k3_finseq_2(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_667, plain, (v5_relat_1(X1,k2_finseq_1(k1_matrix_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(k2_matrix_1(X2)))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_4505, plain, (v3_ordinal1(k2_tarski(np__1,np__2))|~v1_card_1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_678, plain, (v5_relat_1(k3_matrix_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4,X5),X2)|~m1_finseq_1(X3,k3_finseq_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_679, plain, (v1_relat_1(k3_matrix_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4,X5))|~m1_finseq_1(X3,k3_finseq_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_4121, plain, (r2_hidden(k2_tarski(k2_tarski(esk8_3(X1,k4_matrix_1(X2,X3),X4),X5),k1_tarski(X5)),k2_matrix_1(X3))|k4_matrix_1(X2,X3)=X4|v1_xboole_0(X1)|v1_xboole_0(X2)|~r2_hidden(X5,k9_xtuple_0(X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v1_matrix_1(X3)|k3_finseq_1(k4_matrix_1(X2,X3))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,X3))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X2,X3),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_4437, plain, (k2_zfmisc_1(k2_finseq_1(k3_finseq_1(k2_matrlin(X1,X2))),k2_finseq_1(k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k2_matrlin(X1,X2))))))=k2_matrix_1(k4_matrix_1(X3,k9_finseq_1(k2_matrlin(X1,X2))))|v1_xboole_0(X3)|v1_xboole_0(X1)|~m2_finseq_1(k9_finseq_1(k2_matrlin(X1,X2)),k3_finseq_2(X3))|~v1_matrix_1(k9_finseq_1(k2_matrlin(X1,X2)))|~m1_finseq_1(X2,X1)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_5031, plain, (r2_hidden(esk8_3(X1,k4_matrix_1(X2,X3),X4),k9_xtuple_0(X3))|k4_matrix_1(X2,X3)=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v1_matrix_1(X3)|k3_finseq_1(k4_matrix_1(X2,X3))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,X3))!=k1_matrix_1(X4)|~r2_hidden(esk8_3(X1,k4_matrix_1(X2,X3),X4),k4_finseq_1(X3))|~m2_finseq_1(k4_matrix_1(X2,X3),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_4758, plain, (esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4)=X5|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|k1_tarski(np__1)!=k1_tarski(X5)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_4997, plain, (r2_hidden(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),k4_finseq_1(k9_finseq_1(X3)))|k4_matrix_1(X2,k9_finseq_1(X3))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_matrix_1(k9_finseq_1(X3))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k1_matrix_1(X4)|~r2_hidden(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_4727, plain, (k2_zfmisc_1(k4_finseq_1(k9_finseq_1(k9_finseq_1(X1))),k1_tarski(np__1))=k2_matrix_1(k9_finseq_1(k9_finseq_1(X1))))).#trainneg
cnf(i_0_4750, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|r2_hidden(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3),k1_tarski(np__1))|v1_xboole_0(X4)|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_5195, plain, (r2_hidden(esk8_3(X1,k4_matrix_1(X2,X3),X4),k4_finseq_1(X3))|k4_matrix_1(X2,X3)=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v1_matrix_1(X3)|k3_finseq_1(k4_matrix_1(X2,X3))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,X3))!=k1_matrix_1(X4)|~r2_hidden(esk8_3(X1,k4_matrix_1(X2,X3),X4),k9_xtuple_0(X3))|~m2_finseq_1(k4_matrix_1(X2,X3),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_4339, plain, (v1_funct_1(esk11_1(k3_finseq_2(X1))))).#trainneg
cnf(i_0_4654, plain, (v3_xxreal_0(esk22_0)|v2_xxreal_0(esk22_0)|~v1_xreal_0(esk22_0))).#trainneg
cnf(i_0_2958, plain, (k3_finseq_1(X1)=np__2|k9_xtuple_0(X1)!=k1_tarski(np__1)|np__2!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_680, plain, (v1_xboole_0(k3_matrix_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4,X5))|~v1_xboole_0(X2)|~m1_finseq_1(X3,k3_finseq_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_4515, plain, (r1_xxreal_0(esk22_0,np__2)|r1_xxreal_0(np__2,esk22_0))).#trainneg
cnf(i_0_1915, plain, (m1_subset_1(k3_finseq_1(np__0),k1_numbers)|~v1_funct_1(np__0)|~v1_relat_1(np__0))).#trainneg
cnf(i_0_681, plain, (v1_xboole_0(k3_matrix_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4,X5))|~v1_xboole_0(X1)|~m1_finseq_1(X3,k3_finseq_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_682, plain, (v4_relat_1(k3_matrix_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)),X3,X4,X5),X1)|~m1_finseq_1(X3,k3_finseq_2(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_5019, plain, (r2_hidden(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),k1_tarski(np__1))|k4_matrix_1(X2,k9_finseq_1(X3))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_matrix_1(k9_finseq_1(X3))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k1_matrix_1(X4)|~r2_hidden(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),k4_finseq_1(k9_finseq_1(X3)))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_5126, plain, (esk8_3(X1,k4_matrix_1(X2,k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),X5)=X6|k4_matrix_1(X2,k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))=X5|v1_xboole_0(X2)|v1_xboole_0(X1)|v1_xboole_0(X3)|k3_finseq_1(k4_matrix_1(X2,k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X2,k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))!=k1_matrix_1(X5)|k1_tarski(np__1)!=k1_tarski(X6)|~m2_finseq_1(k4_matrix_1(X2,k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X2))|~m2_finseq_1(X5,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(X5)|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_4804, plain, (r2_hidden(k2_tarski(k2_tarski(esk8_3(X1,k4_matrix_1(X2,X3),X4),X5),k1_tarski(esk8_3(X1,k4_matrix_1(X2,X3),X4))),k2_matrix_1(X6))|v1_xboole_0(X7)|k4_matrix_1(X2,X3)=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|X6!=k4_matrix_1(X7,X3)|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,X3),X4))|~v7_ordinal1(X5)|~m2_finseq_1(X6,k3_finseq_2(X7))|~m2_finseq_1(X3,k3_finseq_2(X7))|~v1_matrix_1(X6)|~v1_matrix_1(X3)|k3_finseq_1(k4_matrix_1(X2,X3))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,X3))!=k1_matrix_1(X4)|~r2_hidden(X5,k9_xtuple_0(X3))|~m2_finseq_1(k4_matrix_1(X2,X3),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_2196, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_matrix_1(k9_finseq_1(X2))))|~v1_xboole_0(k2_finseq_1(k1_matrix_1(k9_finseq_1(X2))))|~v1_matrix_1(k9_finseq_1(X2)))).#trainneg
cnf(i_0_4340, plain, (v1_finseq_1(esk11_1(k3_finseq_2(X1))))).#trainneg
cnf(i_0_684, plain, (v1_xboole_0(k3_matrix_1(k1_zfmisc_1(X1),X2,X3,X4))|~v1_xboole_0(X1)|~m1_finseq_1(X2,k3_finseq_2(k1_zfmisc_1(X1)))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_2304, plain, (v7_ordinal1(k1_tarski(np__1))|~v5_ordinal1(np__1)|~v1_funct_1(np__1)|~v1_relat_1(np__1)|~v1_finseq_1(np__1))).#trainneg
cnf(i_0_1917, plain, (m1_subset_1(k3_finseq_1(np__0),k5_numbers)|~v1_funct_1(np__0)|~v1_relat_1(np__0))).#trainneg
cnf(i_0_2301, plain, (v3_ordinal1(k1_tarski(np__1))|~v5_ordinal1(np__1)|~v1_funct_1(np__1)|~v1_relat_1(np__1)|~v1_finseq_1(np__1))).#trainneg
cnf(i_0_685, plain, (v1_relat_1(k3_matrix_1(k1_zfmisc_1(X1),X2,X3,X4))|~v1_relat_1(X1)|~m1_finseq_1(X2,k3_finseq_2(k1_zfmisc_1(X1)))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_686, plain, (v1_funct_1(k3_matrix_1(k1_zfmisc_1(X1),X2,X3,X4))|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_finseq_1(X2,k3_finseq_2(k1_zfmisc_1(X1)))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_5057, plain, (esk8_3(X1,k4_matrix_1(X2,X3),X4)=X5|k4_matrix_1(X2,X3)=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|k1_tarski(np__1)!=k1_tarski(X5)|k3_finseq_1(k4_matrix_1(X2,X3))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,X3))!=k1_matrix_1(X4)|k1_matrix_1(X3)!=np__1|~m2_finseq_1(k4_matrix_1(X2,X3),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X3)|~v1_matrix_1(X4))).#trainneg
cnf(i_0_4742, plain, (X1=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|v1_xboole_0(X2)|v1_xboole_0(X4)|~m1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4,np__1,X1)|~v7_ordinal1(X1)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2)))).#trainneg
cnf(i_0_5052, plain, (k4_matrix_1(X1,k9_finseq_1(k2_matrlin(X2,k9_finseq_1(X3))))=X4|r2_hidden(esk8_3(X5,k4_matrix_1(X1,k9_finseq_1(k2_matrlin(X2,k9_finseq_1(X3)))),X4),k9_xtuple_0(k2_matrlin(X2,k9_finseq_1(X3))))|v1_xboole_0(X2)|v1_xboole_0(X1)|v1_xboole_0(X5)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k2_matrlin(X2,k9_finseq_1(X3)))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k2_matrlin(X2,k9_finseq_1(X3)))))!=k1_matrix_1(X4)|~m1_finseq_1(k9_finseq_1(X3),X2)|~v7_ordinal1(k3_finseq_1(k2_matrlin(X2,k9_finseq_1(X3))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k2_matrlin(X2,k9_finseq_1(X3)))),k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(k2_matrlin(X2,k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(X4,k3_finseq_2(X5))|~v1_matrix_1(X4)|~v1_finseq_1(k2_matrlin(X2,k9_finseq_1(X3)))|~v1_funct_1(k2_matrlin(X2,k9_finseq_1(X3)))|~v1_relat_1(k2_matrlin(X2,k9_finseq_1(X3))))).#trainneg
cnf(i_0_5135, plain, (k4_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))=X4|r2_hidden(esk8_3(X5,k4_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),X4),k1_tarski(np__1))|v1_xboole_0(X5)|v1_xboole_0(X1)|v1_xboole_0(X2)|k3_finseq_1(k4_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k3_finseq_2(X5))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(X4,k3_finseq_2(X5))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_5111, plain, (k2_zfmisc_1(k4_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))),k1_tarski(np__1))=k2_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_funct_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_relat_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_4827, plain, (r2_hidden(k2_tarski(k2_tarski(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),X5),k1_tarski(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4))),k2_matrix_1(X6))|v1_xboole_0(X7)|k4_matrix_1(X2,k9_finseq_1(X3))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|X6!=k4_matrix_1(X7,k9_finseq_1(X3))|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4))|~v7_ordinal1(X5)|~m2_finseq_1(X6,k3_finseq_2(X7))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X7))|~v1_matrix_1(X6)|~v1_matrix_1(k9_finseq_1(X3))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k1_matrix_1(X4)|~r2_hidden(X5,k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_4778, plain, (r2_hidden(k2_tarski(k2_tarski(esk8_3(X1,k4_matrix_1(X2,X3),X4),X5),k1_tarski(esk8_3(X1,k4_matrix_1(X2,X3),X4))),k2_matrix_1(X6))|v1_xboole_0(X7)|k4_matrix_1(X2,X3)=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|X6!=k4_matrix_1(X7,X3)|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,X3),X4))|~v7_ordinal1(X5)|~m2_finseq_1(X6,k3_finseq_2(X7))|~m2_finseq_1(X3,k3_finseq_2(X7))|~v1_matrix_1(X6)|~v1_matrix_1(X3)|k3_finseq_1(k4_matrix_1(X2,X3))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,X3))!=k1_matrix_1(X4)|~r2_hidden(X5,k4_finseq_1(X3))|~m2_finseq_1(k4_matrix_1(X2,X3),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_2993, plain, (v1_xboole_0(esk11_1(k1_zfmisc_1(np__0)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4341, plain, (v1_relat_1(esk11_1(k3_finseq_2(X1))))).#trainneg
cnf(i_0_4185, plain, (r1_xxreal_0(np__1,k3_finseq_1(np__0))|r1_xxreal_0(k3_finseq_1(np__0),np__1))).#trainneg
cnf(i_0_687, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_matrix_1(k1_zfmisc_1(X2),X3,X4,X5))|~m1_finseq_1(X3,k3_finseq_2(k1_zfmisc_1(X2)))|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_2090, plain, (r2_hidden(np__1,X1)|X1!=k2_tarski(np__1,np__2)|np__2!=np__1)).#trainneg
cnf(i_0_3187, plain, (v1_funct_1(k2_tarski(np__1,np__2))|~v1_funct_1(k5_numbers)|~v1_relat_1(k5_numbers))).#trainneg
cnf(i_0_688, plain, (v5_relat_1(k3_matrix_1(k1_zfmisc_1(X1),X2,X3,X4),X5)|~v5_relat_1(X1,X5)|~v1_relat_1(X1)|~m1_finseq_1(X2,k3_finseq_2(k1_zfmisc_1(X1)))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_689, plain, (v1_relat_1(k3_matrix_1(X1,X2,X3,X4))|~v4_funct_1(X1)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_5117, plain, (k2_zfmisc_1(k9_xtuple_0(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))),k1_tarski(np__1))=k2_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_funct_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_relat_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_5678, plain, (k4_matrix_1(X1,k9_finseq_1(X2))=X3|r2_hidden(k2_tarski(k2_tarski(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3),X5),k1_tarski(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3))),k2_matrix_1(k4_matrix_1(X6,k9_finseq_1(X2))))|v1_xboole_0(X4)|v1_xboole_0(X1)|v1_xboole_0(X6)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k1_matrix_1(X3)|~r2_hidden(X5,k1_tarski(np__1))|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3))|~v7_ordinal1(X5)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X6))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X6,k9_finseq_1(X2)),k3_finseq_2(X6))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(k4_matrix_1(X6,k9_finseq_1(X2)))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_5269, plain, (esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4)=np__1|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X1)|v1_xboole_0(X2)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_5361, plain, (k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),X3,X4)=k3_matrix_1(X1,X5,X4,X3)|v1_xboole_0(X1)|X5!=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))|~r2_hidden(k2_tarski(k2_tarski(X3,X4),k1_tarski(X3)),k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1)))|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m2_finseq_1(X5,k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_4727, plain, (k2_zfmisc_1(k4_finseq_1(k9_finseq_1(k9_finseq_1(X1))),k1_tarski(np__1))=k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1)))).#trainneg
cnf(i_0_5412, plain, (k2_matrix_1(k9_finseq_1(X1))=k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1))|k3_finseq_1(X1)!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4754, plain, (esk9_3(X1,k9_finseq_1(k9_finseq_1(X2)),X3)=X4|k9_finseq_1(k9_finseq_1(X2))=X3|v1_xboole_0(X1)|k1_tarski(np__1)!=k1_tarski(X4)|k1_matrix_1(X3)!=np__1|k3_finseq_1(X3)!=np__1|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_5817, plain, (k2_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))=k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1))|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_funct_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_relat_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m1_subset_1(X2,X1)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1)))).#trainneg
cnf(i_0_4777, plain, (r2_hidden(esk8_3(X1,k4_matrix_1(X2,k4_matrix_1(X3,X4)),X5),k2_finseq_1(k1_matrix_1(X4)))|v1_xboole_0(X3)|k4_matrix_1(X2,k4_matrix_1(X3,X4))=X5|v1_xboole_0(X2)|v1_xboole_0(X1)|~m2_finseq_1(X4,k3_finseq_2(X3))|~v1_matrix_1(X4)|k3_finseq_1(k4_matrix_1(X2,k4_matrix_1(X3,X4)))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X2,k4_matrix_1(X3,X4)))!=k1_matrix_1(X5)|~r2_hidden(esk8_3(X1,k4_matrix_1(X2,k4_matrix_1(X3,X4)),X5),k4_finseq_1(k4_matrix_1(X3,X4)))|~m2_finseq_1(k4_matrix_1(X2,k4_matrix_1(X3,X4)),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X3,X4),k3_finseq_2(X2))|~m2_finseq_1(X5,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X3,X4))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_5698, plain, (v1_xboole_0(esk11_1(k1_zfmisc_1(np__0))))).#trainneg
cnf(i_0_4923, plain, (v1_card_1(esk14_2(X1,X2)))).#trainneg
cnf(i_0_5582, plain, (v3_ordinal1(k3_finseq_1(np__0)))).#trainneg
cnf(i_0_5725, plain, (~v1_xboole_0(X1)|esk16_1(X1)!=k2_tarski(np__1,np__2)|np__2!=np__1)).#trainneg
cnf(i_0_690, plain, (v1_funct_1(k3_matrix_1(X1,X2,X3,X4))|~v4_funct_1(X1)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_5592, plain, (v3_ordinal1(k2_tarski(np__1,np__2))|~v5_ordinal1(np__1)|~v1_finseq_1(np__1)|~v1_funct_1(np__1)|~v1_relat_1(np__1)|np__2!=np__1)).#trainneg
cnf(i_0_5586, plain, (k3_finseq_1(np__0)!=np__1|~v1_xboole_0(k3_finseq_1(np__0)))).#trainneg
cnf(i_0_5458, plain, (v1_relat_1(k3_finseq_1(np__0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_691, plain, (v3_ordinal1(k3_matrix_1(X1,X2,X3,X4))|~v3_ordinal1(X1)|~m1_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_698, plain, (m1_subset_1(k2_tarski(k1_tarski(esk8_3(X1,X2,X3)),k2_tarski(esk8_3(X1,X2,X3),esk9_3(X1,X2,X3))),k2_matrix_1(X2))|X2=X3|v1_xboole_0(X1)|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X3)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_5127, plain, (X1=k1_matrix_1(k4_matrix_1(X2,k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|v1_xboole_0(X2)|v1_xboole_0(X5)|v1_xboole_0(X3)|~m1_matrix_1(k4_matrix_1(X2,k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),X5,np__1,X1)|~v7_ordinal1(X1)|~m2_finseq_1(k4_matrix_1(X2,k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),k3_finseq_2(X5))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_4803, plain, (r2_hidden(esk8_3(X1,k4_matrix_1(X2,k4_matrix_1(X3,X4)),X5),k2_finseq_1(k1_matrix_1(X4)))|v1_xboole_0(X3)|k4_matrix_1(X2,k4_matrix_1(X3,X4))=X5|v1_xboole_0(X2)|v1_xboole_0(X1)|~m2_finseq_1(X4,k3_finseq_2(X3))|~v1_matrix_1(X4)|k3_finseq_1(k4_matrix_1(X2,k4_matrix_1(X3,X4)))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X2,k4_matrix_1(X3,X4)))!=k1_matrix_1(X5)|~r2_hidden(esk8_3(X1,k4_matrix_1(X2,k4_matrix_1(X3,X4)),X5),k9_xtuple_0(k4_matrix_1(X3,X4)))|~m2_finseq_1(k4_matrix_1(X2,k4_matrix_1(X3,X4)),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X3,X4),k3_finseq_2(X2))|~m2_finseq_1(X5,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X3,X4))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_4756, plain, (k9_finseq_1(k9_finseq_1(X1))=X2|r2_hidden(esk9_3(X3,k9_finseq_1(k9_finseq_1(X1)),X2),k1_tarski(np__1))|v1_xboole_0(X3)|k1_matrix_1(X2)!=np__1|k3_finseq_1(X2)!=np__1|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X1)),k3_finseq_2(X3))|~m2_finseq_1(X2,k3_finseq_2(X3))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_4805, plain, (r2_hidden(k2_tarski(k2_tarski(esk8_3(X1,k4_matrix_1(X2,X3),X4),X5),k1_tarski(esk8_3(X1,k4_matrix_1(X2,X3),X4))),k2_matrix_1(X6))|v1_xboole_0(X7)|k4_matrix_1(X2,X3)=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|X3!=k4_matrix_1(X7,X6)|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,X3),X4))|~v7_ordinal1(X5)|~m2_finseq_1(X6,k3_finseq_2(X7))|~m2_finseq_1(X3,k3_finseq_2(X7))|~v1_matrix_1(X6)|~v1_matrix_1(X3)|k3_finseq_1(k4_matrix_1(X2,X3))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,X3))!=k1_matrix_1(X4)|~r2_hidden(X5,k9_xtuple_0(X3))|~m2_finseq_1(k4_matrix_1(X2,X3),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_4141, plain, (v1_finseq_1(k2_matrlin(X1,np__0))|v1_xboole_0(X1))).#trainneg
cnf(i_0_5459, plain, (v1_funct_1(k3_finseq_1(np__0))|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_5748, plain, (~v1_xboole_0(X1)|X1!=k2_tarski(np__1,np__2)|np__2!=np__1)).#trainneg
cnf(i_0_699, plain, (k2_tarski(k1_tarski(esk8_3(X1,X2,X3)),k2_tarski(esk8_3(X1,X2,X3),esk9_3(X1,X2,X3)))=X4|X2=X3|v1_xboole_0(X1)|k2_matrix_1(X2)!=k1_tarski(X4)|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X3)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_5580, plain, (v1_relat_1(k3_finseq_1(np__0))|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_4493, plain, (v1_finset_1(esk14_2(X2,X3))|~v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_704, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3))|v1_xboole_0(X4)|v1_xboole_0(k2_matrix_1(X5))|X3!=k4_matrix_1(X4,X5)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(X5)|~m1_subset_1(k2_tarski(k2_tarski(X2,X1),k1_tarski(X2)),k2_matrix_1(X5)))).#trainneg
cnf(i_0_709, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3))|v1_xboole_0(X4)|v1_xboole_0(k2_matrix_1(X5))|X5!=k4_matrix_1(X4,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(X5)|~m1_subset_1(k2_tarski(k2_tarski(X2,X1),k1_tarski(X2)),k2_matrix_1(X5)))).#trainneg
cnf(i_0_5274, plain, (esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4)=X5|k4_matrix_1(X2,k9_finseq_1(X3))=X4|v1_xboole_0(X1)|v1_xboole_0(X2)|k4_finseq_1(k9_finseq_1(X3))!=k1_tarski(X5)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k1_matrix_1(X4)|~r2_hidden(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X3))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_5689, plain, (k4_matrix_1(X1,X2)=X3|r2_hidden(k2_tarski(k2_tarski(esk8_3(X4,k4_matrix_1(X1,X2),X3),X5),k1_tarski(esk8_3(X4,k4_matrix_1(X1,X2),X3))),k2_matrix_1(k2_matrlin(X6,k9_finseq_1(X7))))|v1_xboole_0(X4)|v1_xboole_0(X1)|v1_xboole_0(X6)|k3_finseq_1(k4_matrix_1(X1,X2))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,X2))!=k1_matrix_1(X3)|k2_matrlin(X6,k9_finseq_1(X7))!=k4_matrix_1(X6,X2)|~r2_hidden(X5,k4_finseq_1(X2))|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X1,X2),X3))|~v7_ordinal1(X5)|~m2_finseq_1(k4_matrix_1(X1,X2),k3_finseq_2(X4))|~m2_finseq_1(X2,k3_finseq_2(X6))|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(k2_matrlin(X6,k9_finseq_1(X7)))|~v1_matrix_1(X2)|~v1_matrix_1(X3)|~m1_finseq_1(k9_finseq_1(X7),X6))).#trainneg
cnf(i_0_5113, plain, (k2_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))=k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1))|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=np__1|~v1_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_funct_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_relat_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_5532, plain, (k4_matrix_1(X1,X2)=X3|r2_hidden(k2_tarski(k2_tarski(esk8_3(X4,k4_matrix_1(X1,X2),X3),X5),k1_tarski(esk8_3(X4,k4_matrix_1(X1,X2),X3))),k2_matrix_1(k2_matrlin(X6,k9_finseq_1(X7))))|v1_xboole_0(X4)|v1_xboole_0(X1)|v1_xboole_0(X6)|k3_finseq_1(k4_matrix_1(X1,X2))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,X2))!=k1_matrix_1(X3)|k2_matrlin(X6,k9_finseq_1(X7))!=k4_matrix_1(X6,X2)|~r2_hidden(X5,k9_xtuple_0(X2))|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X1,X2),X3))|~v7_ordinal1(X5)|~m2_finseq_1(k4_matrix_1(X1,X2),k3_finseq_2(X4))|~m2_finseq_1(X2,k3_finseq_2(X6))|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(k2_matrlin(X6,k9_finseq_1(X7)))|~v1_matrix_1(X2)|~v1_matrix_1(X3)|~m1_finseq_1(k9_finseq_1(X7),X6))).#trainneg
cnf(i_0_5125, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|v1_xboole_0(X3)|~r2_hidden(X2,k1_tarski(np__1))|~r2_hidden(X1,k4_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|~v1_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_funct_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_relat_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_4814, plain, (r2_hidden(k2_tarski(k2_tarski(esk8_3(X1,k4_matrix_1(X2,X3),X4),esk8_3(X1,k4_matrix_1(X2,X3),X4)),k1_tarski(esk8_3(X1,k4_matrix_1(X2,X3),X4))),k2_matrix_1(X5))|v1_xboole_0(X6)|k4_matrix_1(X2,X3)=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|X3!=k4_matrix_1(X6,X5)|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,X3),X4))|~m2_finseq_1(X5,k3_finseq_2(X6))|~m2_finseq_1(X3,k3_finseq_2(X6))|~v1_matrix_1(X5)|~v1_matrix_1(X3)|k3_finseq_1(k4_matrix_1(X2,X3))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,X3))!=k1_matrix_1(X4)|~r2_hidden(esk8_3(X1,k4_matrix_1(X2,X3),X4),k9_xtuple_0(X3))|~m2_finseq_1(k4_matrix_1(X2,X3),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_4836, plain, (r2_hidden(k2_tarski(k2_tarski(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4)),k1_tarski(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4))),k2_matrix_1(X5))|v1_xboole_0(X6)|k4_matrix_1(X2,k9_finseq_1(X3))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|X5!=k4_matrix_1(X6,k9_finseq_1(X3))|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4))|~m2_finseq_1(X5,k3_finseq_2(X6))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X6))|~v1_matrix_1(X5)|~v1_matrix_1(k9_finseq_1(X3))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k1_matrix_1(X4)|~r2_hidden(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_4482, plain, (v1_relat_1(k1_finseq_1(X3))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_5581, plain, (v1_funct_1(k3_finseq_1(np__0))|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_714, plain, (k3_matrix_1(X1,X2,X3,X4)=k3_matrix_1(X1,X5,X4,X3)|v1_xboole_0(X1)|v1_xboole_0(k2_matrix_1(X2))|X5!=k4_matrix_1(X1,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m2_finseq_1(X5,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X5)|~v1_matrix_1(X2)|~m1_subset_1(k2_tarski(k2_tarski(X3,X4),k1_tarski(X3)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_2109, plain, (m1_subset_1(np__1,k2_tarski(np__1,np__2))|np__2!=np__1)).#trainneg
cnf(i_0_4616, plain, (v7_ordinal1(esk14_2(X2,X3))|~v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_716, plain, (m1_subset_1(k2_tarski(k1_tarski(esk7_3(X1,X2,X3)),k2_tarski(esk7_3(X1,X2,X3),esk6_3(X1,X2,X3))),k2_matrix_1(X2))|X3=k4_matrix_1(X1,X2)|v7_ordinal1(esk4_3(X1,X2,X3))|v1_xboole_0(X1)|k1_matrix_1(X2)!=k3_finseq_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X3)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_717, plain, (k2_tarski(k1_tarski(esk7_3(X1,X2,X3)),k2_tarski(esk7_3(X1,X2,X3),esk6_3(X1,X2,X3)))=X4|X3=k4_matrix_1(X1,X2)|v7_ordinal1(esk4_3(X1,X2,X3))|v1_xboole_0(X1)|k2_matrix_1(X2)!=k1_tarski(X4)|k1_matrix_1(X2)!=k3_finseq_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X3)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_4828, plain, (r2_hidden(k2_tarski(k2_tarski(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),X5),k1_tarski(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4))),k2_matrix_1(X6))|v1_xboole_0(X7)|k4_matrix_1(X2,k9_finseq_1(X3))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|k9_finseq_1(X3)!=k4_matrix_1(X7,X6)|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4))|~v7_ordinal1(X5)|~m2_finseq_1(X6,k3_finseq_2(X7))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X7))|~v1_matrix_1(X6)|~v1_matrix_1(k9_finseq_1(X3))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k1_matrix_1(X4)|~r2_hidden(X5,k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_5122, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|v1_xboole_0(X3)|~r2_hidden(X2,k1_tarski(np__1))|~r2_hidden(X1,k9_xtuple_0(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|~v1_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_funct_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_relat_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_5837, plain, (r2_hidden(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),k4_finseq_1(k4_matrix_1(X5,k9_finseq_1(X3))))|k4_matrix_1(X2,k9_finseq_1(X3))=X4|v1_xboole_0(X5)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X5,k9_finseq_1(X3)))|~v1_funct_1(k4_matrix_1(X5,k9_finseq_1(X3)))|~v1_relat_1(k4_matrix_1(X5,k9_finseq_1(X3)))|~v1_matrix_1(k4_matrix_1(X5,k9_finseq_1(X3)))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k1_matrix_1(X4)|~r2_hidden(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),k1_tarski(np__1))|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X3))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_5847, plain, (k4_matrix_1(X1,k9_finseq_1(X2))=X3|r2_hidden(k2_tarski(k1_tarski(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3)),k2_tarski(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3),X5)),k2_matrix_1(k4_matrix_1(X6,k9_finseq_1(X2))))|v1_xboole_0(X6)|v1_xboole_0(X1)|v1_xboole_0(X4)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k1_matrix_1(X3)|~r2_hidden(X5,k1_tarski(np__1))|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3))|~v7_ordinal1(X5)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X6))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_5836, plain, (r2_hidden(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),k9_xtuple_0(k4_matrix_1(X5,k9_finseq_1(X3))))|k4_matrix_1(X2,k9_finseq_1(X3))=X4|v1_xboole_0(X5)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X5,k9_finseq_1(X3)))|~v1_funct_1(k4_matrix_1(X5,k9_finseq_1(X3)))|~v1_relat_1(k4_matrix_1(X5,k9_finseq_1(X3)))|~v1_matrix_1(k4_matrix_1(X5,k9_finseq_1(X3)))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k1_matrix_1(X4)|~r2_hidden(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),k1_tarski(np__1))|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X3))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_5845, plain, (k4_matrix_1(X1,k9_finseq_1(X2))=X3|r2_hidden(k2_tarski(k2_tarski(X5,esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3)),k1_tarski(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3))),k2_matrix_1(k4_matrix_1(X6,k9_finseq_1(X2))))|v1_xboole_0(X6)|v1_xboole_0(X1)|v1_xboole_0(X4)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k1_matrix_1(X3)|~r2_hidden(X5,k1_tarski(np__1))|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3))|~v7_ordinal1(X5)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X6))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_5850, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~m1_finseq_1(X3,k3_finseq_2(k1_tarski(X4)))|~m1_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X4)))|~v1_matrix_1(X3)|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X4)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(k1_tarski(X4))))).#trainneg
cnf(i_0_6159, plain, (v1_finset_1(esk14_2(X1,X2)))).#trainneg
cnf(i_0_6101, plain, (k3_finseq_1(np__0)!=np__1)).#trainneg
cnf(i_0_718, plain, (X1=k4_matrix_1(X2,X3)|v7_ordinal1(esk6_3(X2,X3,X1))|v1_xboole_0(X2)|v1_xboole_0(k2_matrix_1(X1))|k1_matrix_1(X3)!=k3_finseq_1(X1)|~r2_hidden(k2_tarski(k1_tarski(esk5_3(X2,X3,X1)),k2_tarski(esk4_3(X2,X3,X1),esk5_3(X2,X3,X1))),k2_matrix_1(X3))|~m2_finseq_1(X1,k3_finseq_2(X2))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X1)|~v1_matrix_1(X3)|~m1_subset_1(k2_tarski(k1_tarski(esk4_3(X2,X3,X1)),k2_tarski(esk4_3(X2,X3,X1),esk5_3(X2,X3,X1))),k2_matrix_1(X1)))).#trainneg
cnf(i_0_720, plain, (m1_subset_1(k2_tarski(k1_tarski(esk7_3(X1,X2,X3)),k2_tarski(esk7_3(X1,X2,X3),esk6_3(X1,X2,X3))),k2_matrix_1(X2))|X3=k4_matrix_1(X1,X2)|v7_ordinal1(esk5_3(X1,X2,X3))|v1_xboole_0(X1)|k1_matrix_1(X2)!=k3_finseq_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X3)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_721, plain, (k2_tarski(k1_tarski(esk7_3(X1,X2,X3)),k2_tarski(esk7_3(X1,X2,X3),esk6_3(X1,X2,X3)))=X4|X3=k4_matrix_1(X1,X2)|v7_ordinal1(esk5_3(X1,X2,X3))|v1_xboole_0(X1)|k2_matrix_1(X2)!=k1_tarski(X4)|k1_matrix_1(X2)!=k3_finseq_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X3)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_6059, plain, (esk9_3(X1,k9_finseq_1(k9_finseq_1(X2)),X3)=np__1|k9_finseq_1(k9_finseq_1(X2))=X3|v1_xboole_0(X1)|k1_matrix_1(X3)!=np__1|k3_finseq_1(X3)!=np__1|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_5820, plain, (k2_tarski(k2_tarski(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),X5),k1_tarski(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4)))=X6|k4_matrix_1(X2,k9_finseq_1(X3))=X4|v1_xboole_0(X7)|v1_xboole_0(X2)|v1_xboole_0(X1)|k2_matrix_1(k4_matrix_1(X7,k9_finseq_1(X3)))!=k1_tarski(X6)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k1_matrix_1(X4)|~r2_hidden(X5,k1_tarski(np__1))|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4))|~v7_ordinal1(X5)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X7))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X3))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_5849, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|v1_xboole_0(X1)|k3_matrix_1(k1_tarski(X4),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),np__1,esk9_3(k1_tarski(X4),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3))!=X4|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~m1_finseq_1(X3,k3_finseq_2(k1_tarski(X4)))|~v7_ordinal1(esk9_3(k1_tarski(X4),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X4)))|~v1_matrix_1(X3)|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(k1_tarski(X4))))).#trainneg
cnf(i_0_6281, plain, (k4_matrix_1(X1,k9_finseq_1(X2))=X3|r2_hidden(k2_tarski(k1_tarski(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3)),k2_tarski(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3),esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3))),k2_matrix_1(k4_matrix_1(X5,k9_finseq_1(X2))))|v1_xboole_0(X4)|v1_xboole_0(X1)|v1_xboole_0(X5)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k1_matrix_1(X3)|~r2_hidden(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3),k1_tarski(np__1))|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X5,k9_finseq_1(X2)),k3_finseq_2(X5))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(k4_matrix_1(X5,k9_finseq_1(X2)))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_5852, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|r2_hidden(k2_tarski(k1_tarski(np__1),k2_tarski(np__1,esk9_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3))),k2_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))))|v1_xboole_0(X4)|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1)))).#trainneg
cnf(i_0_6135, plain, (k4_matrix_1(X1,k9_finseq_1(X2))=X3|r2_hidden(k2_tarski(k2_tarski(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3),X5),k1_tarski(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3))),k2_matrix_1(X6))|v1_xboole_0(X4)|v1_xboole_0(X1)|v1_xboole_0(X7)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k1_matrix_1(X3)|k9_finseq_1(X2)!=k4_matrix_1(X7,X6)|~r2_hidden(X5,k1_tarski(np__1))|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3))|~v7_ordinal1(X5)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)),k3_finseq_2(X4))|~m2_finseq_1(X6,k3_finseq_2(X7))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(X6)|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3)|~m1_subset_1(X2,k3_finseq_2(X7)))).#trainneg
cnf(i_0_5853, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|v1_xboole_0(X4)|v1_xboole_0(X1)|k3_matrix_1(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),np__1,esk9_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3))!=k3_matrix_1(X4,X3,np__1,esk9_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3))|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1)))).#trainneg
cnf(i_0_6064, plain, (k2_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))=k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1))|v1_xboole_0(X1)|~v1_funct_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_relat_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_5826, plain, (r2_hidden(k2_tarski(k2_tarski(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4)),k1_tarski(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4))),k2_matrix_1(X5))|v1_xboole_0(X6)|k4_matrix_1(X2,k9_finseq_1(X3))=X4|v1_xboole_0(X7)|v1_xboole_0(X2)|v1_xboole_0(X1)|X5!=k4_matrix_1(X6,k4_matrix_1(X7,k9_finseq_1(X3)))|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4))|~m2_finseq_1(X5,k3_finseq_2(X6))|~m2_finseq_1(k4_matrix_1(X7,k9_finseq_1(X3)),k3_finseq_2(X6))|~v1_matrix_1(X5)|~v1_matrix_1(k4_matrix_1(X7,k9_finseq_1(X3)))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k1_matrix_1(X4)|~r2_hidden(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X7))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X3))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_5830, plain, (r2_hidden(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),k4_finseq_1(k4_matrix_1(X5,k9_finseq_1(X3))))|k4_matrix_1(X2,k9_finseq_1(X3))=X4|v1_xboole_0(X5)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X5,k9_finseq_1(X3)))|~v1_funct_1(k4_matrix_1(X5,k9_finseq_1(X3)))|~v1_relat_1(k4_matrix_1(X5,k9_finseq_1(X3)))|~v1_matrix_1(k4_matrix_1(X5,k9_finseq_1(X3)))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k1_matrix_1(X4)|~r2_hidden(X6,k1_tarski(np__1))|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4))|~v7_ordinal1(X6)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X3))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_2761, plain, (m1_subset_1(X1,k1_zfmisc_1(np__0))|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_6303, plain, (v7_ordinal1(esk14_2(X1,X2)))).#trainneg
cnf(i_0_6079, plain, (esk11_1(k1_zfmisc_1(np__0))=np__0)).#trainneg
cnf(i_0_722, plain, (X1=k4_matrix_1(X2,X3)|v7_ordinal1(esk7_3(X2,X3,X1))|v1_xboole_0(X2)|v1_xboole_0(k2_matrix_1(X1))|k1_matrix_1(X3)!=k3_finseq_1(X1)|~r2_hidden(k2_tarski(k1_tarski(esk5_3(X2,X3,X1)),k2_tarski(esk4_3(X2,X3,X1),esk5_3(X2,X3,X1))),k2_matrix_1(X3))|~m2_finseq_1(X1,k3_finseq_2(X2))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X1)|~v1_matrix_1(X3)|~m1_subset_1(k2_tarski(k1_tarski(esk4_3(X2,X3,X1)),k2_tarski(esk4_3(X2,X3,X1),esk5_3(X2,X3,X1))),k2_matrix_1(X1)))).#trainneg
cnf(i_0_2116, plain, (r2_hidden(np__1,k2_tarski(np__1,np__2))|np__2!=np__1)).#trainneg
cnf(i_0_2737, plain, (v3_card_1(k2_tarski(np__1,np__2),np__1)|np__2!=np__1)).#trainneg
cnf(i_0_5085, plain, (X1=esk11_1(k3_finseq_2(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4496, plain, (esk14_2(X4,X5)=esk14_2(X2,X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_723, plain, (X1=k4_matrix_1(X2,X3)|v1_xboole_0(X2)|v1_xboole_0(k2_matrix_1(X1))|k3_matrix_1(X2,X1,esk6_3(X2,X3,X1),esk7_3(X2,X3,X1))!=k3_matrix_1(X2,X3,esk7_3(X2,X3,X1),esk6_3(X2,X3,X1))|k1_matrix_1(X3)!=k3_finseq_1(X1)|~r2_hidden(k2_tarski(k1_tarski(esk5_3(X2,X3,X1)),k2_tarski(esk4_3(X2,X3,X1),esk5_3(X2,X3,X1))),k2_matrix_1(X3))|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(X3)|~v1_matrix_1(X1)|~m1_subset_1(k2_tarski(k1_tarski(esk4_3(X2,X3,X1)),k2_tarski(esk4_3(X2,X3,X1),esk5_3(X2,X3,X1))),k2_matrix_1(X1)))).#trainneg
cnf(i_0_724, plain, (X1=k4_matrix_1(X2,X3)|r2_hidden(k2_tarski(k1_tarski(esk7_3(X2,X3,X1)),k2_tarski(esk7_3(X2,X3,X1),esk6_3(X2,X3,X1))),k2_matrix_1(X3))|v1_xboole_0(X2)|v1_xboole_0(k2_matrix_1(X1))|k1_matrix_1(X3)!=k3_finseq_1(X1)|~r2_hidden(k2_tarski(k1_tarski(esk5_3(X2,X3,X1)),k2_tarski(esk4_3(X2,X3,X1),esk5_3(X2,X3,X1))),k2_matrix_1(X3))|~m2_finseq_1(X1,k3_finseq_2(X2))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X1)|~v1_matrix_1(X3)|~m1_subset_1(k2_tarski(k1_tarski(esk4_3(X2,X3,X1)),k2_tarski(esk4_3(X2,X3,X1),esk5_3(X2,X3,X1))),k2_matrix_1(X1)))).#trainneg
cnf(i_0_5861, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|r2_hidden(k2_tarski(k1_tarski(X4),k2_tarski(np__1,X4)),k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1)))|v1_xboole_0(X1)|v1_xboole_0(X5)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~r2_hidden(X4,k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X5))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_5833, plain, (r2_hidden(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),k2_finseq_1(k1_matrix_1(k4_matrix_1(X5,k9_finseq_1(X3)))))|k4_matrix_1(X2,k9_finseq_1(X3))=X4|v1_xboole_0(X5)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X5,k9_finseq_1(X3)))|~v1_funct_1(k4_matrix_1(X5,k9_finseq_1(X3)))|~v1_relat_1(k4_matrix_1(X5,k9_finseq_1(X3)))|~v1_matrix_1(k4_matrix_1(X5,k9_finseq_1(X3)))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k1_matrix_1(X4)|~r2_hidden(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),k1_tarski(np__1))|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X3))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_5842, plain, (k4_matrix_1(X1,X2)=X3|r2_hidden(k2_tarski(k2_tarski(esk8_3(X4,k4_matrix_1(X1,X2),X3),X5),k1_tarski(esk8_3(X4,k4_matrix_1(X1,X2),X3))),k2_matrix_1(k4_matrix_1(X6,X2)))|v1_xboole_0(X6)|v1_xboole_0(X1)|v1_xboole_0(X4)|k3_finseq_1(k4_matrix_1(X1,X2))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,X2))!=k1_matrix_1(X3)|~r2_hidden(X5,k1_tarski(np__1))|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X1,X2),X3))|~v7_ordinal1(X5)|~m2_finseq_1(k4_matrix_1(X1,X2),k3_finseq_2(X4))|~m2_finseq_1(X2,k3_finseq_2(X6))|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(X2)|~v1_matrix_1(X3)|k3_finseq_1(X2)!=np__1|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6197, plain, (esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4)=X5|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|k4_finseq_1(k9_finseq_1(k9_finseq_1(X3)))!=k1_tarski(X5)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_6130, plain, (k9_finseq_1(X1)=X2|r2_hidden(esk9_3(X3,k9_finseq_1(X1),X2),k1_tarski(np__1))|v1_xboole_0(X3)|k1_matrix_1(X2)!=np__1|k3_finseq_1(X2)!=np__1|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X3))|~m2_finseq_1(X2,k3_finseq_2(X3))|~v1_matrix_1(X2)|k3_finseq_1(X1)!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6333, plain, (k4_matrix_1(X1,X2)=X3|r2_hidden(k2_tarski(k2_tarski(esk8_3(X4,k4_matrix_1(X1,X2),X3),X5),k1_tarski(esk8_3(X4,k4_matrix_1(X1,X2),X3))),k2_matrix_1(X6))|v1_xboole_0(X4)|v1_xboole_0(X1)|v1_xboole_0(X7)|k3_finseq_1(k4_matrix_1(X1,X2))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,X2))!=k1_matrix_1(X3)|X2!=k4_matrix_1(X7,X6)|~r2_hidden(X5,k1_tarski(np__1))|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X1,X2),X3))|~v7_ordinal1(X5)|~m2_finseq_1(k4_matrix_1(X1,X2),k3_finseq_2(X4))|~m2_finseq_1(X2,k3_finseq_2(X7))|~m2_finseq_1(X2,k3_finseq_2(X1))|~m2_finseq_1(X6,k3_finseq_2(X7))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(X2)|~v1_matrix_1(X6)|~v1_matrix_1(X3)|k3_finseq_1(X2)!=np__1|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5829, plain, (k3_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(X3)),esk8_3(X4,k4_matrix_1(X5,k9_finseq_1(X3)),X6),esk8_3(X4,k4_matrix_1(X5,k9_finseq_1(X3)),X6))=k3_matrix_1(X1,X7,esk8_3(X4,k4_matrix_1(X5,k9_finseq_1(X3)),X6),esk8_3(X4,k4_matrix_1(X5,k9_finseq_1(X3)),X6))|v1_xboole_0(X1)|k4_matrix_1(X5,k9_finseq_1(X3))=X6|v1_xboole_0(X2)|v1_xboole_0(X5)|v1_xboole_0(X4)|X7!=k4_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(X3)))|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X5,k9_finseq_1(X3)),X6))|~m2_finseq_1(X7,k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~v1_matrix_1(X7)|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))|k3_finseq_1(k4_matrix_1(X5,k9_finseq_1(X3)))!=k3_finseq_1(X6)|k1_matrix_1(k4_matrix_1(X5,k9_finseq_1(X3)))!=k1_matrix_1(X6)|~r2_hidden(esk8_3(X4,k4_matrix_1(X5,k9_finseq_1(X3)),X6),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X5,k9_finseq_1(X3)),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X5))|~m2_finseq_1(X6,k3_finseq_2(X4))|~v1_matrix_1(k9_finseq_1(X3))|~v1_matrix_1(X6))).#trainneg
cnf(i_0_6137, plain, (k4_matrix_1(X1,k9_finseq_1(X2))=X3|r2_hidden(k2_tarski(k2_tarski(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3),X5),k1_tarski(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3))),k2_matrix_1(X6))|v1_xboole_0(X4)|v1_xboole_0(X1)|v1_xboole_0(X7)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k1_matrix_1(X3)|k9_finseq_1(X2)!=k4_matrix_1(X7,X6)|~r2_hidden(X5,k1_tarski(np__1))|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3))|~v7_ordinal1(X5)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)),k3_finseq_2(X4))|~m2_finseq_1(X6,k3_finseq_2(X7))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(X6)|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3)|~m1_finseq_1(X2,X7)|~v7_ordinal1(k3_finseq_1(X2)))).#trainneg
cnf(i_0_4597, plain, (v1_relat_1(k1_card_1(X3))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_4186, plain, (r1_xxreal_0(esk25_0,k3_finseq_1(np__0))|r1_xxreal_0(k3_finseq_1(np__0),esk25_0))).#trainneg
cnf(i_0_734, plain, (m1_subset_1(esk3_2(X1,esk16_1(X2)),X2)|esk3_2(X1,esk16_1(X2))=X1|esk16_1(X2)=k1_tarski(X1))).#trainneg
cnf(i_0_5749, plain, (m1_subset_1(np__1,X1)|X1!=k2_tarski(np__1,np__2)|np__2!=np__1)).#trainneg
cnf(i_0_742, plain, (esk14_2(X1,X2)=esk14_2(X3,X4))).#trainneg
cnf(i_0_746, plain, (esk14_2(X1,X2)=np__0)).#trainneg
cnf(i_0_6532, plain, (r2_hidden(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),k2_finseq_1(k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))))|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_funct_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_relat_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_6632, plain, (k4_matrix_1(X1,k9_finseq_1(X2))=X3|r2_hidden(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3),k4_finseq_1(k4_matrix_1(X5,k9_finseq_1(X2))))|v1_xboole_0(X4)|v1_xboole_0(X1)|v1_xboole_0(X5)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k1_matrix_1(X3)|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_6454, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~m1_finseq_1(X3,k3_finseq_2(k1_tarski(X4)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(X3)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X4))))).#trainneg
cnf(i_0_6815, plain, (k3_matrix_1(X1,k4_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(X3))),esk8_3(X4,k4_matrix_1(X5,k9_finseq_1(X3)),X6),esk8_3(X4,k4_matrix_1(X5,k9_finseq_1(X3)),X6))=k3_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(X3)),esk8_3(X4,k4_matrix_1(X5,k9_finseq_1(X3)),X6),esk8_3(X4,k4_matrix_1(X5,k9_finseq_1(X3)),X6))|k4_matrix_1(X5,k9_finseq_1(X3))=X6|v1_xboole_0(X4)|v1_xboole_0(X5)|v1_xboole_0(X2)|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X5,k9_finseq_1(X3)))!=k3_finseq_1(X6)|k1_matrix_1(k4_matrix_1(X5,k9_finseq_1(X3)))!=k1_matrix_1(X6)|~r2_hidden(esk8_3(X4,k4_matrix_1(X5,k9_finseq_1(X3)),X6),k1_tarski(np__1))|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X5,k9_finseq_1(X3)),X6))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X5,k9_finseq_1(X3)),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X5))|~m2_finseq_1(k4_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(X6,k3_finseq_2(X4))|~v1_matrix_1(k9_finseq_1(X3))|~v1_matrix_1(k4_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(X3))))|~v1_matrix_1(X6))).#trainneg
cnf(i_0_6731, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|r2_hidden(k2_tarski(k2_tarski(np__1,X4),k1_tarski(X4)),k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1)))|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~r2_hidden(X4,k1_tarski(np__1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_6865, plain, (~v1_xboole_0(X1)|k1_zfmisc_1(np__0)!=k2_tarski(np__1,np__2)|np__2!=np__1)).#trainneg
cnf(i_0_768, plain, (m1_subset_1(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_1891, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_754, plain, (k4_matrix_1(X1,X2)=k4_matrix_1(X1,X3)|r2_hidden(k2_tarski(k1_tarski(esk5_3(X1,X3,k4_matrix_1(X1,X2))),k2_tarski(esk4_3(X1,X3,k4_matrix_1(X1,X2)),esk5_3(X1,X3,k4_matrix_1(X1,X2)))),k2_matrix_1(X3))|r2_hidden(k2_tarski(k1_tarski(esk4_3(X1,X3,k4_matrix_1(X1,X2))),k2_tarski(esk4_3(X1,X3,k4_matrix_1(X1,X2)),esk5_3(X1,X3,k4_matrix_1(X1,X2)))),k2_matrix_1(k4_matrix_1(X1,X2)))|v7_ordinal1(esk6_3(X1,X3,k4_matrix_1(X1,X2)))|v1_xboole_0(X1)|k1_matrix_1(X3)!=k3_finseq_1(k4_matrix_1(X1,X2))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,X2))|~v1_matrix_1(X3)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_755, plain, (k4_matrix_1(X1,X2)=k4_matrix_1(X1,X3)|r2_hidden(k2_tarski(k1_tarski(esk5_3(X1,X3,k4_matrix_1(X1,X2))),k2_tarski(esk4_3(X1,X3,k4_matrix_1(X1,X2)),esk5_3(X1,X3,k4_matrix_1(X1,X2)))),k2_matrix_1(X3))|r2_hidden(k2_tarski(k1_tarski(esk4_3(X1,X3,k4_matrix_1(X1,X2))),k2_tarski(esk4_3(X1,X3,k4_matrix_1(X1,X2)),esk5_3(X1,X3,k4_matrix_1(X1,X2)))),k2_matrix_1(k4_matrix_1(X1,X2)))|v7_ordinal1(esk7_3(X1,X3,k4_matrix_1(X1,X2)))|v1_xboole_0(X1)|k1_matrix_1(X3)!=k3_finseq_1(k4_matrix_1(X1,X2))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,X2))|~v1_matrix_1(X3)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_6544, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|r2_hidden(k2_tarski(k2_tarski(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3),X5),k1_tarski(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3))),k2_matrix_1(X6))|v1_xboole_0(X7)|v1_xboole_0(X1)|v1_xboole_0(X4)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|k9_finseq_1(k9_finseq_1(X2))!=k4_matrix_1(X7,X6)|~r2_hidden(X5,k1_tarski(np__1))|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3))|~v7_ordinal1(X5)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X6,k3_finseq_2(X7))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(X6)|~v1_matrix_1(X3)|~m1_subset_1(X2,X7))).#trainneg
cnf(i_0_6531, plain, (k3_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),np__1,esk9_3(X4,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X5))=k3_matrix_1(X1,X6,esk9_3(X4,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X5),np__1)|v1_xboole_0(X1)|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X5|v1_xboole_0(X2)|v1_xboole_0(X4)|X6!=k4_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v7_ordinal1(esk9_3(X4,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X5))|~m2_finseq_1(X6,k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~v1_matrix_1(X6)|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X5)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_6530, plain, (r2_hidden(k2_tarski(k2_tarski(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),np__1),k1_tarski(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4))),k2_matrix_1(X5))|v1_xboole_0(X6)|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))!=k4_matrix_1(X6,X5)|~v7_ordinal1(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4))|~m2_finseq_1(X5,k3_finseq_2(X6))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X6))|~v1_matrix_1(X5)|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_4343, plain, (v4_relat_1(esk11_1(k3_finseq_2(X1)),k5_numbers))).#trainneg
cnf(i_0_4188, plain, (r1_xxreal_0(np__0,k3_finseq_1(np__0))|r1_xxreal_0(k3_finseq_1(np__0),np__0))).#trainneg
cnf(i_0_769, plain, (k2_tarski(k1_tarski(X1),k2_tarski(X1,X2))=X3|k2_zfmisc_1(X4,X5)!=k1_tarski(X3)|~r2_hidden(X2,X5)|~r2_hidden(X1,X4))).#trainneg
cnf(i_0_6827, plain, (r1_xxreal_0(esk25_0,k3_finseq_1(np__0))|~v2_xxreal_0(k3_finseq_1(np__0))|~v1_xreal_0(k3_finseq_1(np__0)))).#trainneg
cnf(i_0_3500, plain, (r1_xxreal_0(X1,np__0)|r1_xxreal_0(np__0,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_756, plain, (k4_matrix_1(X1,X2)=k4_matrix_1(X1,X3)|r2_hidden(k2_tarski(k1_tarski(esk7_3(X1,X3,k4_matrix_1(X1,X2))),k2_tarski(esk7_3(X1,X3,k4_matrix_1(X1,X2)),esk6_3(X1,X3,k4_matrix_1(X1,X2)))),k2_matrix_1(X3))|r2_hidden(k2_tarski(k1_tarski(esk5_3(X1,X3,k4_matrix_1(X1,X2))),k2_tarski(esk4_3(X1,X3,k4_matrix_1(X1,X2)),esk5_3(X1,X3,k4_matrix_1(X1,X2)))),k2_matrix_1(X3))|r2_hidden(k2_tarski(k1_tarski(esk4_3(X1,X3,k4_matrix_1(X1,X2))),k2_tarski(esk4_3(X1,X3,k4_matrix_1(X1,X2)),esk5_3(X1,X3,k4_matrix_1(X1,X2)))),k2_matrix_1(k4_matrix_1(X1,X2)))|v1_xboole_0(X1)|k1_matrix_1(X3)!=k3_finseq_1(k4_matrix_1(X1,X2))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,X2))|~v1_matrix_1(X3)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_786, plain, (v3_ordinal1(k3_pre_poly(X1,X2))|v1_xboole_0(X1)|~v3_ordinal1(k3_finseq_2(X1))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_6613, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|r2_hidden(k2_tarski(k2_tarski(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3),X5),k1_tarski(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3))),k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1)))|v1_xboole_0(X6)|v1_xboole_0(X1)|v1_xboole_0(X4)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~r2_hidden(X5,k1_tarski(np__1))|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3))|~v7_ordinal1(X5)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X6))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_funct_1(k4_matrix_1(X6,k9_finseq_1(k9_finseq_1(X2))))|~v1_relat_1(k4_matrix_1(X6,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X6,k9_finseq_1(k9_finseq_1(X2))))|~m1_subset_1(X2,X6))).#trainneg
cnf(i_0_6741, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5))),k2_matrix_1(k4_matrix_1(X6,k9_finseq_1(X4))))|k4_matrix_1(X3,k9_finseq_1(X4))=X5|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X6)|~r2_hidden(X1,k4_finseq_1(k4_matrix_1(X6,k9_finseq_1(X4))))|~v1_finseq_1(k4_matrix_1(X6,k9_finseq_1(X4)))|~v1_funct_1(k4_matrix_1(X6,k9_finseq_1(X4)))|~v1_relat_1(k4_matrix_1(X6,k9_finseq_1(X4)))|~v1_matrix_1(k4_matrix_1(X6,k9_finseq_1(X4)))|k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(X4)))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(X4)))!=k1_matrix_1(X5)|~r2_hidden(esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5),k1_tarski(np__1))|~v7_ordinal1(esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(X4)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X4),k3_finseq_2(X6))|~m2_finseq_1(k9_finseq_1(X4),k3_finseq_2(X3))|~m2_finseq_1(X5,k3_finseq_2(X2))|~v1_matrix_1(k9_finseq_1(X4))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_6421, plain, (k2_tarski(k2_tarski(X1,esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5)),k1_tarski(esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5)))=X6|k4_matrix_1(X3,k9_finseq_1(X4))=X5|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X7)|k2_matrix_1(k4_matrix_1(X7,k9_finseq_1(X4)))!=k1_tarski(X6)|k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(X4)))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(X4)))!=k1_matrix_1(X5)|~r2_hidden(X1,k1_tarski(np__1))|~v7_ordinal1(esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5))|~v7_ordinal1(X1)|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(X4)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X4),k3_finseq_2(X7))|~m2_finseq_1(k9_finseq_1(X4),k3_finseq_2(X3))|~m2_finseq_1(X5,k3_finseq_2(X2))|~v1_matrix_1(k9_finseq_1(X4))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_6402, plain, (k4_matrix_1(X1,k9_finseq_1(X2))=X3|r2_hidden(k2_tarski(k1_tarski(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3)),k2_tarski(X5,esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3))),k2_matrix_1(k4_matrix_1(X6,k9_finseq_1(X2))))|v1_xboole_0(X4)|v1_xboole_0(X1)|v1_xboole_0(X6)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k1_matrix_1(X3)|~r2_hidden(X5,k1_tarski(np__1))|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3))|~v7_ordinal1(X5)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X6))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_4189, plain, (r1_xxreal_0(esk17_0,k3_finseq_1(np__0))|r1_xxreal_0(k3_finseq_1(np__0),esk17_0))).#trainneg
cnf(i_0_6876, plain, (~v1_xboole_0(X1)|k1_zfmisc_1(X1)!=k2_tarski(np__1,np__2)|np__2!=np__1)).#trainneg
cnf(i_0_793, plain, (r2_hidden(X1,X2)|v1_xboole_0(k2_zfmisc_1(X4,X2))|~m1_subset_1(k2_tarski(k1_tarski(X3),k2_tarski(X3,X1)),k2_zfmisc_1(X4,X2)))).#trainneg
cnf(i_0_6890, plain, (v1_relat_1(np__1)|~v4_funct_1(X1)|X1!=k2_tarski(np__1,np__2)|np__2!=np__1)).#trainneg
cnf(i_0_2258, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X1)|v1_xboole_0(X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_6615, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|r2_hidden(k2_tarski(k1_tarski(np__1),k2_tarski(np__1,esk9_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3))),k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1)))|v1_xboole_0(X1)|v1_xboole_0(X4)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(X3)|~v1_funct_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_relat_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_6388, plain, (k3_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(X3)),esk8_3(X4,k4_matrix_1(X5,k9_finseq_1(X3)),X6),X7)=k3_matrix_1(X1,X8,X7,esk8_3(X4,k4_matrix_1(X5,k9_finseq_1(X3)),X6))|v1_xboole_0(X1)|k4_matrix_1(X5,k9_finseq_1(X3))=X6|v1_xboole_0(X4)|v1_xboole_0(X5)|v1_xboole_0(X2)|X8!=k4_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(X3)))|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X5,k9_finseq_1(X3)),X6))|~v7_ordinal1(X7)|~m2_finseq_1(X8,k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~v1_matrix_1(X8)|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))|k3_finseq_1(k4_matrix_1(X5,k9_finseq_1(X3)))!=k3_finseq_1(X6)|k1_matrix_1(k4_matrix_1(X5,k9_finseq_1(X3)))!=k1_matrix_1(X6)|~r2_hidden(X7,k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X5,k9_finseq_1(X3)),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X5))|~m2_finseq_1(X6,k3_finseq_2(X4))|~v1_matrix_1(k9_finseq_1(X3))|~v1_matrix_1(X6))).#trainneg
cnf(i_0_6488, plain, (esk9_3(X1,k9_finseq_1(X2),X3)=np__1|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|k1_matrix_1(X3)!=np__1|k3_finseq_1(X3)!=np__1|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(X3)|k3_finseq_1(X2)!=np__1|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6610, plain, (k2_tarski(k2_tarski(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),X5),k1_tarski(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4)))=X6|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X7)|k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1))!=k1_tarski(X6)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~r2_hidden(X5,k1_tarski(np__1))|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4))|~v7_ordinal1(X5)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X7))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4)|~v1_funct_1(k4_matrix_1(X7,k9_finseq_1(k9_finseq_1(X3))))|~v1_relat_1(k4_matrix_1(X7,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k4_matrix_1(X7,k9_finseq_1(k9_finseq_1(X3))))|~m1_subset_1(X3,X7))).#trainneg
cnf(i_0_6393, plain, (r2_hidden(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),k9_xtuple_0(k4_matrix_1(X5,k9_finseq_1(X3))))|k4_matrix_1(X2,k9_finseq_1(X3))=X4|v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X5)|~v1_finseq_1(k4_matrix_1(X5,k9_finseq_1(X3)))|~v1_funct_1(k4_matrix_1(X5,k9_finseq_1(X3)))|~v1_relat_1(k4_matrix_1(X5,k9_finseq_1(X3)))|~v1_matrix_1(k4_matrix_1(X5,k9_finseq_1(X3)))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k1_matrix_1(X4)|~r2_hidden(X6,k1_tarski(np__1))|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4))|~v7_ordinal1(X6)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X3))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_4906, plain, (v5_relat_1(esk11_1(k3_finseq_2(X1)),X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4518, plain, (r1_xxreal_0(esk22_0,k3_finseq_1(np__0))|r1_xxreal_0(k3_finseq_1(np__0),esk22_0))).#trainneg
cnf(i_0_799, plain, (m1_subset_1(esk3_2(X1,esk11_1(k1_zfmisc_1(X2))),X2)|esk3_2(X1,esk11_1(k1_zfmisc_1(X2)))=X1|esk11_1(k1_zfmisc_1(X2))=k1_tarski(X1))).#trainneg
cnf(i_0_6891, plain, (v1_funct_1(np__1)|~v4_funct_1(X1)|X1!=k2_tarski(np__1,np__2)|np__2!=np__1)).#trainneg
cnf(i_0_1385, plain, (v3_xxreal_0(np__0)|v2_xxreal_0(X1)|r1_xxreal_0(X1,np__0)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_800, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk11_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,esk11_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_802, plain, (k9_xtuple_0(k3_finseq_1(X1))=k2_finseq_1(k3_finseq_1(X1))|~m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_finseq_1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(k3_finseq_1(X1))|~v1_funct_1(X1)|~v1_relat_1(k3_finseq_1(X1))|~v1_relat_1(X1))).#trainneg
cnf(i_0_6742, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5))),k2_matrix_1(k4_matrix_1(X6,k9_finseq_1(X4))))|k4_matrix_1(X3,k9_finseq_1(X4))=X5|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X6)|~r2_hidden(X1,k9_xtuple_0(k4_matrix_1(X6,k9_finseq_1(X4))))|~v1_finseq_1(k4_matrix_1(X6,k9_finseq_1(X4)))|~v1_funct_1(k4_matrix_1(X6,k9_finseq_1(X4)))|~v1_relat_1(k4_matrix_1(X6,k9_finseq_1(X4)))|~v1_matrix_1(k4_matrix_1(X6,k9_finseq_1(X4)))|k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(X4)))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(X4)))!=k1_matrix_1(X5)|~r2_hidden(esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5),k1_tarski(np__1))|~v7_ordinal1(esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(X4)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X4),k3_finseq_2(X6))|~m2_finseq_1(k9_finseq_1(X4),k3_finseq_2(X3))|~m2_finseq_1(X5,k3_finseq_2(X2))|~v1_matrix_1(k9_finseq_1(X4))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_6959, plain, (k4_matrix_1(X1,k9_finseq_1(X2))=X3|r2_hidden(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3),k9_xtuple_0(k4_matrix_1(X5,k9_finseq_1(X2))))|v1_xboole_0(X5)|v1_xboole_0(X1)|v1_xboole_0(X4)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k1_matrix_1(X3)|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3)|~v1_finseq_1(k4_matrix_1(X5,k9_finseq_1(X2)))|~v1_funct_1(k4_matrix_1(X5,k9_finseq_1(X2)))|~v1_relat_1(k4_matrix_1(X5,k9_finseq_1(X2))))).#trainneg
cnf(i_0_7320, plain, (r2_hidden(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),k1_tarski(np__1))|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X1)|v1_xboole_0(X2)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_6964, plain, (k4_matrix_1(k1_tarski(X1),k9_finseq_1(k9_finseq_1(X2)))=X3|k3_finseq_1(k4_matrix_1(k1_tarski(X1),k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(k1_tarski(X1),k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~m1_finseq_1(X3,k3_finseq_2(k1_tarski(X1)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(k1_tarski(X1)))|~v1_matrix_1(k4_matrix_1(k1_tarski(X1),k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_7403, plain, (v3_xxreal_0(np__0)|r1_xxreal_0(esk25_0,np__0))).#trainneg
cnf(i_0_747, plain, (k1_card_1(X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7202, plain, (r1_xxreal_0(esk25_0,np__0)|~v2_xxreal_0(np__0))).#trainneg
cnf(i_0_829, plain, (k1_finseq_1(X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7145, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|r2_hidden(k2_tarski(k2_tarski(np__1,esk9_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3)),k1_tarski(esk9_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3))),k2_matrix_1(X5))|v1_xboole_0(X4)|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))!=k4_matrix_1(X1,X5)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X5,k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(X5)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_6819, plain, (k4_matrix_1(X1,k9_finseq_1(X2))=X3|r2_hidden(k2_tarski(k2_tarski(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3),X5),k1_tarski(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3))),k2_matrix_1(X6))|v1_xboole_0(X7)|v1_xboole_0(X1)|v1_xboole_0(X4)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k1_matrix_1(X3)|k9_finseq_1(X2)!=k4_matrix_1(X7,X6)|~r2_hidden(X5,k1_tarski(np__1))|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3))|~v7_ordinal1(k3_finseq_1(X2))|~v7_ordinal1(X5)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X6,k3_finseq_2(X7))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(X6)|~v1_matrix_1(X3)|~m2_finseq_1(X2,X7))).#trainneg
cnf(i_0_4757, plain, (k1_matrix_1(k9_finseq_1(X1))=np__1|k3_finseq_1(X1)!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6940, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,esk9_3(X2,k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),X5))),k2_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))=X5|v1_xboole_0(X2)|v1_xboole_0(X3)|~r2_hidden(X1,k4_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|~v1_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_funct_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_relat_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))!=k1_matrix_1(X5)|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m2_finseq_1(X5,k3_finseq_2(X2))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_5593, plain, (esk8_3(X1,k4_matrix_1(X2,X3),X4)=np__1|k4_matrix_1(X2,X3)=X4|v1_xboole_0(X1)|v1_xboole_0(X2)|k3_finseq_1(k4_matrix_1(X2,X3))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,X3))!=k1_matrix_1(X4)|k1_matrix_1(X3)!=np__1|~m2_finseq_1(k4_matrix_1(X2,X3),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X3)|~v1_matrix_1(X4))).#trainneg
cnf(i_0_6739, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5),X1)),k2_matrix_1(k4_matrix_1(X6,k9_finseq_1(X4))))|k4_matrix_1(X3,k9_finseq_1(X4))=X5|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X6)|~r2_hidden(X1,k4_finseq_1(k4_matrix_1(X6,k9_finseq_1(X4))))|~v1_finseq_1(k4_matrix_1(X6,k9_finseq_1(X4)))|~v1_funct_1(k4_matrix_1(X6,k9_finseq_1(X4)))|~v1_relat_1(k4_matrix_1(X6,k9_finseq_1(X4)))|~v1_matrix_1(k4_matrix_1(X6,k9_finseq_1(X4)))|k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(X4)))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(X4)))!=k1_matrix_1(X5)|~r2_hidden(esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5),k1_tarski(np__1))|~v7_ordinal1(esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(X4)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X4),k3_finseq_2(X6))|~m2_finseq_1(k9_finseq_1(X4),k3_finseq_2(X3))|~m2_finseq_1(X5,k3_finseq_2(X2))|~v1_matrix_1(k9_finseq_1(X4))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_831, plain, (k9_xtuple_0(X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7288, plain, (r1_xxreal_0(esk17_0,k3_finseq_1(np__0))|~v2_xxreal_0(k3_finseq_1(np__0))|~v1_xreal_0(esk17_0)|~v1_xreal_0(k3_finseq_1(np__0)))).#trainneg
cnf(i_0_1584, plain, (v3_xxreal_0(esk22_0)|v2_xxreal_0(X1)|r1_xxreal_0(X1,esk22_0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_839, plain, (r2_hidden(k2_tarski(k2_tarski(X1,esk3_2(X2,k2_finseq_1(k1_matrix_1(X3)))),k1_tarski(X1)),k2_matrix_1(X3))|esk3_2(X2,k2_finseq_1(k1_matrix_1(X3)))=X2|k2_finseq_1(k1_matrix_1(X3))=k1_tarski(X2)|~r2_hidden(X1,k4_finseq_1(X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_840, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3))|v1_xboole_0(k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X1,k4_finseq_1(X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v1_matrix_1(X3)|~m1_subset_1(X2,k2_finseq_1(k1_matrix_1(X3))))).#trainneg
cnf(i_0_5011, plain, (X1=X2|k3_finseq_1(X1)!=k3_finseq_1(X2)|k1_matrix_1(X1)!=k1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(k1_tarski(X3)))|~m1_finseq_1(X1,k3_finseq_2(k1_tarski(X3)))|~v1_matrix_1(X2)|~v1_matrix_1(X1)|~m1_subset_1(esk8_3(k1_tarski(X3),X1,X2),k5_numbers))).#trainneg
cnf(i_0_6196, plain, (esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4)=X5|k4_matrix_1(X2,k9_finseq_1(X3))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k1_matrix_1(X4)|k4_finseq_1(k9_finseq_1(X3))!=k1_tarski(X5)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X3))|~v1_matrix_1(X4)|~m1_subset_1(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),k1_tarski(np__1)))).#trainneg
cnf(i_0_4735, plain, (k2_zfmisc_1(k1_tarski(np__1),k2_finseq_1(k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X1))))))=k2_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X1))))|v1_xboole_0(X2)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X1)),k3_finseq_2(X2)))).#trainneg
cnf(i_0_4891, plain, (k9_finseq_1(X1)=X2|r2_hidden(k2_tarski(k1_tarski(esk9_3(X3,k9_finseq_1(X1),X2)),k2_tarski(esk9_3(X3,k9_finseq_1(X1),X2),X4)),k2_matrix_1(k4_matrix_1(X5,k9_finseq_1(X1))))|v1_xboole_0(X5)|v1_xboole_0(X3)|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(X2)|k3_finseq_1(X2)!=np__1|~r2_hidden(X4,k1_tarski(np__1))|~v7_ordinal1(esk9_3(X3,k9_finseq_1(X1),X2))|~v7_ordinal1(X4)|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X3))|~m2_finseq_1(X2,k3_finseq_2(X3))|~v1_matrix_1(k9_finseq_1(X1))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_5663, plain, (r2_hidden(X1,k1_tarski(np__1))|v1_xboole_0(X3)|~r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X2)),k2_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|~v1_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_funct_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_relat_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_6634, plain, (m1_subset_1(esk11_1(k3_finseq_2(X1)),k1_zfmisc_1(np__0))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2947, plain, (m2_subset_1(k1_matrix_1(np__0),k1_numbers,k5_numbers)|v1_xboole_0(X1)|~v1_matrix_1(np__0))).#trainneg
cnf(i_0_7149, plain, (v2_xxreal_0(np__0)|r1_xxreal_0(np__0,k3_finseq_1(np__0))|~v2_xxreal_0(k3_finseq_1(np__0))|~v1_xreal_0(k3_finseq_1(np__0)))).#trainneg
cnf(i_0_1884, plain, (v3_xxreal_0(np__1)|v2_xxreal_0(X1)|r1_xxreal_0(X1,np__1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1890, plain, (v3_xxreal_0(np__0)|v2_xxreal_0(X1)|r1_xxreal_0(X1,np__0)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_843, plain, (v1_finset_1(k9_xtuple_0(X1))|~v7_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4763, plain, (esk8_3(X1,k4_matrix_1(X2,k2_matrlin(X3,k9_finseq_1(X4))),X5)=X6|k4_matrix_1(X2,k2_matrlin(X3,k9_finseq_1(X4)))=X5|v1_xboole_0(X2)|v1_xboole_0(X1)|v1_xboole_0(X3)|k3_finseq_1(k4_matrix_1(X2,k2_matrlin(X3,k9_finseq_1(X4))))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X2,k2_matrlin(X3,k9_finseq_1(X4))))!=k1_matrix_1(X5)|k1_tarski(np__1)!=k1_tarski(X6)|~m2_finseq_1(k4_matrix_1(X2,k2_matrlin(X3,k9_finseq_1(X4))),k3_finseq_2(X1))|~m2_finseq_1(k2_matrlin(X3,k9_finseq_1(X4)),k3_finseq_2(X2))|~m2_finseq_1(X5,k3_finseq_2(X1))|~v1_matrix_1(k2_matrlin(X3,k9_finseq_1(X4)))|~v1_matrix_1(X5)|k3_finseq_1(k2_matrlin(X3,k9_finseq_1(X4)))!=np__1|~m1_finseq_1(k9_finseq_1(X4),X3))).#trainneg
cnf(i_0_6745, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5))),k2_matrix_1(k4_matrix_1(X6,k9_finseq_1(X4))))|v1_xboole_0(X6)|k4_matrix_1(X3,k9_finseq_1(X4))=X5|v1_xboole_0(X2)|v1_xboole_0(X3)|~r2_hidden(X1,k2_finseq_1(k1_matrix_1(k9_finseq_1(X4))))|~m2_finseq_1(k9_finseq_1(X4),k3_finseq_2(X6))|~v1_matrix_1(k9_finseq_1(X4))|k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(X4)))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(X4)))!=k1_matrix_1(X5)|~r2_hidden(esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5),k1_tarski(np__1))|~v7_ordinal1(esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(X4)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X4),k3_finseq_2(X3))|~m2_finseq_1(X5,k3_finseq_2(X2))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_8003, plain, (k9_xtuple_0(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))=k1_tarski(np__1)|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_funct_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_relat_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1)))).#trainneg
cnf(i_0_7721, plain, (v3_xxreal_0(esk22_0)|v2_xxreal_0(np__0)|r1_xxreal_0(np__0,esk22_0))).#trainneg
cnf(i_0_849, plain, (m2_finseq_2(X1,X2,k3_finseq_2(X2))|v1_xboole_0(X2)|~m1_subset_1(k1_funct_1(X1,np__1),X2)|k3_finseq_1(X1)!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7940, plain, (v3_xxreal_0(np__1)|v2_xxreal_0(esk22_0)|r1_xxreal_0(esk22_0,np__1))).#trainneg
cnf(i_0_857, plain, (esk11_1(X1)=X2|v1_xboole_0(X1)|X1!=k1_tarski(X2))).#trainneg
cnf(i_0_7852, plain, (r2_hidden(X1,k1_tarski(np__1))|v1_xboole_0(X3)|~r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3)))).#trainneg
cnf(i_0_7533, plain, (r2_hidden(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),k4_finseq_1(X5))|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_finseq_1(X5)|~v1_funct_1(X5)|~v1_relat_1(X5)|~v1_matrix_1(X5)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))!=k4_matrix_1(X2,X5)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X5,k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_7613, plain, (r2_hidden(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),k1_tarski(np__1))|v1_xboole_0(X2)|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X1)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~r2_hidden(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),k4_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_7573, plain, (X1=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))|v1_xboole_0(X2)|v1_xboole_0(X4)|~m1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)),X4,np__1,X1)|~v7_ordinal1(X1)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~v1_matrix_1(k9_finseq_1(X3))|k3_finseq_1(X3)!=np__1|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_7870, plain, (k2_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))=k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1))|v1_xboole_0(X1)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_7525, plain, (r2_hidden(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),k1_tarski(np__1))|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_matrix_1(k9_finseq_1(X5))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))!=k4_matrix_1(X2,k9_finseq_1(X5))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X5),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_4779, plain, (r2_hidden(k2_tarski(k2_tarski(esk8_3(X1,k4_matrix_1(X2,X3),X4),X5),k1_tarski(esk8_3(X1,k4_matrix_1(X2,X3),X4))),k2_matrix_1(X6))|v1_xboole_0(X7)|k4_matrix_1(X2,X3)=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|X3!=k4_matrix_1(X7,X6)|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,X3),X4))|~v7_ordinal1(X5)|~m2_finseq_1(X6,k3_finseq_2(X7))|~m2_finseq_1(X3,k3_finseq_2(X7))|~v1_matrix_1(X6)|~v1_matrix_1(X3)|k3_finseq_1(k4_matrix_1(X2,X3))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,X3))!=k1_matrix_1(X4)|~r2_hidden(X5,k4_finseq_1(X3))|~m2_finseq_1(k4_matrix_1(X2,X3),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_5672, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k2_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|v1_xboole_0(X3)|~r2_hidden(X1,k1_tarski(np__1))|~r2_hidden(X2,k4_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|~v1_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_funct_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_relat_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_7934, plain, (m1_subset_1(X1,np__0)|~r2_hidden(X1,esk11_1(k3_finseq_2(X2)))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_7958, plain, (v3_xxreal_0(np__0)|v2_xxreal_0(esk22_0)|r1_xxreal_0(esk22_0,np__0))).#trainneg
cnf(i_0_6872, plain, (~v1_xboole_0(X2)|k1_zfmisc_1(k2_zfmisc_1(X1,X2))!=k2_tarski(np__1,np__2)|np__2!=np__1)).#trainneg
cnf(i_0_877, plain, (r2_hidden(X1,k4_finseq_1(X2))|v1_xboole_0(k2_matrix_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_matrix_1(X2)|~m1_subset_1(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_7499, plain, (k1_card_1(np__0)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7713, plain, (r1_xxreal_0(esk17_0,np__0)|~v2_xxreal_0(np__0)|~v1_xreal_0(esk17_0))).#trainneg
cnf(i_0_7033, plain, (r1_xxreal_0(np__2,k3_finseq_1(np__0))|r1_xxreal_0(k3_finseq_1(np__0),np__2))).#trainneg
cnf(i_0_7703, plain, (v3_card_1(np__0,X2)|~v3_card_1(X1,X2)|~v1_card_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_893, plain, (k1_finseq_1(X1)=k1_card_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6930, plain, (esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4)=X5|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X1)|v1_xboole_0(X2)|k2_finseq_1(k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))!=k1_tarski(X5)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_6790, plain, (esk9_3(X1,k9_finseq_1(X2),X3)=X4|k9_finseq_1(X2)=X3|v1_xboole_0(X1)|k1_tarski(np__1)!=k1_tarski(X4)|k1_matrix_1(X3)!=np__1|k3_finseq_1(X3)!=np__1|k3_finseq_1(X2)!=np__1|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_6938, plain, (r2_hidden(k2_tarski(k2_tarski(X1,esk9_3(X2,k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),X5)),k1_tarski(X1)),k2_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))=X5|v1_xboole_0(X2)|v1_xboole_0(X3)|~r2_hidden(X1,k9_xtuple_0(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|~v1_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_funct_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_relat_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))!=k1_matrix_1(X5)|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m2_finseq_1(X5,k3_finseq_2(X2))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_5140, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|v1_xboole_0(X3)|~r2_hidden(X2,k1_tarski(np__1))|~r2_hidden(X1,k1_tarski(np__1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_6218, plain, (r2_hidden(esk8_3(X1,k4_matrix_1(X2,X3),X4),k4_finseq_1(k2_matrlin(X5,k9_finseq_1(X6))))|k4_matrix_1(X2,X3)=X4|v1_xboole_0(X5)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_finseq_1(k2_matrlin(X5,k9_finseq_1(X6)))|~v1_funct_1(k2_matrlin(X5,k9_finseq_1(X6)))|~v1_relat_1(k2_matrlin(X5,k9_finseq_1(X6)))|~v1_matrix_1(k2_matrlin(X5,k9_finseq_1(X6)))|k2_matrlin(X5,k9_finseq_1(X6))!=k4_matrix_1(X5,X3)|k3_finseq_1(k4_matrix_1(X2,X3))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,X3))!=k1_matrix_1(X4)|~m1_finseq_1(k9_finseq_1(X6),X5)|~r2_hidden(esk8_3(X1,k4_matrix_1(X2,X3),X4),k4_finseq_1(X3))|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,X3),X4))|~m2_finseq_1(k4_matrix_1(X2,X3),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X5))|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X3)|~v1_matrix_1(X4))).#trainneg
cnf(i_0_7532, plain, (r2_hidden(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),k9_xtuple_0(X5))|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_finseq_1(X5)|~v1_funct_1(X5)|~v1_relat_1(X5)|~v1_matrix_1(X5)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))!=k4_matrix_1(X2,X5)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X5,k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_7726, plain, (v3_xxreal_0(esk22_0)|v2_xxreal_0(k3_finseq_1(np__0))|r1_xxreal_0(k3_finseq_1(np__0),esk22_0))).#trainneg
cnf(i_0_6873, plain, (~v1_xboole_0(X1)|k1_zfmisc_1(k2_zfmisc_1(X1,X2))!=k2_tarski(np__1,np__2)|np__2!=np__1)).#trainneg
cnf(i_0_895, plain, (k1_card_1(X1)=k1_card_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7510, plain, (k3_finseq_1(np__0)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7365, plain, (v2_xxreal_0(esk22_0)|r1_xxreal_0(esk22_0,k3_finseq_1(np__0))|~v2_xxreal_0(k3_finseq_1(np__0))|~v1_xreal_0(esk22_0)|~v1_xreal_0(k3_finseq_1(np__0)))).#trainneg
cnf(i_0_7730, plain, (v3_xxreal_0(esk22_0)|v2_xxreal_0(X1)|r1_xxreal_0(X1,esk22_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7956, plain, (v3_xxreal_0(np__1)|v2_xxreal_0(X1)|r1_xxreal_0(X1,np__1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_896, plain, (k9_xtuple_0(X1)=k1_card_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_903, plain, (k1_finseq_1(X1)=k9_xtuple_0(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8181, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|r2_hidden(esk9_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3),k1_tarski(np__1))|v1_xboole_0(X4)|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_8083, plain, (esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4)=X5|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X1)|v1_xboole_0(X2)|k4_finseq_1(X6)!=k1_tarski(X5)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))!=k4_matrix_1(X2,X6)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X6,k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X6)|~v1_matrix_1(X4))).#trainneg
cnf(i_0_8378, plain, (k4_matrix_1(X1,k9_finseq_1(X2))=X3|r2_hidden(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3),k4_finseq_1(k2_matrlin(X5,k9_finseq_1(X6))))|v1_xboole_0(X4)|v1_xboole_0(X1)|v1_xboole_0(X5)|k2_matrlin(X5,k9_finseq_1(X6))!=k4_matrix_1(X5,k9_finseq_1(X2))|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k1_matrix_1(X3)|~m1_finseq_1(k9_finseq_1(X6),X5)|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3)|~r2_hidden(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3),k1_tarski(np__1)))).#trainneg
cnf(i_0_8244, plain, (k1_card_1(np__0)=np__0)).#trainneg
cnf(i_0_906, plain, (k9_xtuple_0(X1)=k9_xtuple_0(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8325, plain, (k3_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),esk9_3(X4,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X5),esk9_3(X4,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X5))=k3_matrix_1(X1,X6,esk9_3(X4,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X5),esk9_3(X4,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X5))|v1_xboole_0(X1)|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X5|v1_xboole_0(X2)|v1_xboole_0(X4)|X6!=k4_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v7_ordinal1(esk9_3(X4,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X5))|~m2_finseq_1(X6,k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~v1_matrix_1(X6)|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X5)|~r2_hidden(esk9_3(X4,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X5),k9_xtuple_0(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_8382, plain, (esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4)=X5|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X1)|v1_xboole_0(X2)|k9_xtuple_0(X6)!=k1_tarski(X5)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))!=k4_matrix_1(X2,X6)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X6,k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X6)|~v1_matrix_1(X4))).#trainneg
cnf(i_0_8300, plain, (esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4)=X5|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|k1_tarski(np__1)!=k1_tarski(X5)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=np__1|~m2_subset_1(k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5266, plain, (k2_zfmisc_1(k9_xtuple_0(k2_matrlin(X1,X2)),k2_finseq_1(k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k2_matrlin(X1,X2))))))=k2_matrix_1(k4_matrix_1(X3,k9_finseq_1(k2_matrlin(X1,X2))))|v1_xboole_0(X1)|v1_xboole_0(X3)|~m1_finseq_1(X2,X1)|~v7_ordinal1(k3_finseq_1(X2))|~m2_finseq_1(k9_finseq_1(k2_matrlin(X1,X2)),k3_finseq_2(X3))|~v1_matrix_1(k9_finseq_1(k2_matrlin(X1,X2)))|~v1_finseq_1(k2_matrlin(X1,X2))|~v1_funct_1(k2_matrlin(X1,X2))|~v1_relat_1(k2_matrlin(X1,X2)))).#trainneg
cnf(i_0_5132, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|v1_xboole_0(X3)|~r2_hidden(X2,k1_tarski(np__1))|~r2_hidden(X1,k4_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|~v1_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_funct_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_relat_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_7556, plain, (k2_zfmisc_1(k4_finseq_1(k9_finseq_1(X1)),k1_tarski(np__1))=k2_matrix_1(k9_finseq_1(X1))|~v1_matrix_1(k9_finseq_1(X1))|k3_finseq_1(X1)!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8467, plain, (k3_finseq_1(np__0)=np__0)).#trainneg
cnf(i_0_911, plain, (esk16_1(X1)=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_912, plain, (esk16_1(X1)=k1_card_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_913, plain, (esk16_1(X1)=k9_xtuple_0(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6941, plain, (r2_hidden(k2_tarski(k2_tarski(X1,esk9_3(X2,k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),X5)),k1_tarski(X1)),k2_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|v1_xboole_0(X3)|k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))=X5|v1_xboole_0(X2)|~r2_hidden(X1,k1_tarski(np__1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))!=k1_matrix_1(X5)|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X2))|~m2_finseq_1(X5,k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_6529, plain, (r2_hidden(k2_tarski(k2_tarski(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),np__1),k1_tarski(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4))),k2_matrix_1(X5))|v1_xboole_0(X6)|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|X5!=k4_matrix_1(X6,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v7_ordinal1(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4))|~m2_finseq_1(X5,k3_finseq_2(X6))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X6))|~v1_matrix_1(X5)|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_1053, plain, (esk16_1(X1)=k1_finseq_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8064, plain, (v2_xxreal_0(esk22_0)|r1_xxreal_0(esk22_0,np__1)|~v1_xxreal_0(np__1))).#trainneg
cnf(i_0_915, plain, (r2_hidden(X1,X2)|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X3,X1)),k2_zfmisc_1(X2,X4)))).#trainneg
cnf(i_0_919, plain, (r2_hidden(X1,X2)|v1_xboole_0(k2_zfmisc_1(X2,X4))|~m1_subset_1(k2_tarski(k1_tarski(X1),k2_tarski(X1,X3)),k2_zfmisc_1(X2,X4)))).#trainneg
cnf(i_0_8853, plain, (X1=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,X1,esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),esk9_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),esk9_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|k3_finseq_1(X1)!=k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(X1)!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X1)|~r2_hidden(esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~r2_hidden(esk9_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~v7_ordinal1(esk9_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~v7_ordinal1(esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2)))).#trainneg
cnf(i_0_8618, plain, (esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4)=X5|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|k4_finseq_1(k9_finseq_1(k9_finseq_1(X3)))!=k1_tarski(X5)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_8862, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|v1_xboole_0(X1)|k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3),esk8_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3))!=k3_matrix_1(X1,X3,esk8_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3),esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3))|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~v1_matrix_1(X3)|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~r2_hidden(esk8_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3),k1_tarski(np__1))|~r2_hidden(esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3),k1_tarski(np__1))|~v7_ordinal1(esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3))|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1)))).#trainneg
cnf(i_0_8858, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|v1_xboole_0(X3)|k3_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),np__1,esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))!=k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)),esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~r2_hidden(esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),k1_tarski(np__1))|~v7_ordinal1(esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3)))).#trainneg
cnf(i_0_8610, plain, (esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4)=X5|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|k1_tarski(np__1)!=k1_tarski(X5)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_8863, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|v1_xboole_0(X1)|k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3),np__1)!=k3_matrix_1(X1,X3,np__1,esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3))|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(X3)|~r2_hidden(esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3),k1_tarski(np__1))|~v7_ordinal1(esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3)))).#trainneg
cnf(i_0_9134, plain, (r2_hidden(X1,k1_tarski(np__1))|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X2,X1)),k2_matrix_1(k9_finseq_1(X3)))|~v1_matrix_1(k9_finseq_1(X3)))).#trainneg
cnf(i_0_9135, plain, (r2_hidden(X1,k1_tarski(np__1))|v1_xboole_0(X3)|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X2,X1)),k2_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3)))).#trainneg
cnf(i_0_9009, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_card_1(X3))|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_922, plain, (m1_subset_1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k2_zfmisc_1(X3,X4))|~r2_hidden(X1,X4)|~r2_hidden(X2,X3))).#trainneg
cnf(i_0_923, plain, (k2_tarski(k2_tarski(X1,X2),k1_tarski(X2))=X3|k2_zfmisc_1(X4,X5)!=k1_tarski(X3)|~r2_hidden(X1,X5)|~r2_hidden(X2,X4))).#trainneg
cnf(i_0_931, plain, (v1_xboole_0(esk21_1(X1))|v1_xboole_0(X1)|~v1_funct_1(esk21_1(X1))|~v1_relat_1(esk21_1(X1))|~v1_xboole_0(k3_finseq_1(esk21_1(X1))))).#trainneg
cnf(i_0_9091, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|r2_hidden(k2_tarski(k2_tarski(np__1,esk9_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3)),k1_tarski(esk9_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3))),k2_matrix_1(k4_matrix_1(X5,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))))|v1_xboole_0(X4)|v1_xboole_0(X1)|v1_xboole_0(X5)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X5))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X5,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))),k3_finseq_2(X5))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X5,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_8630, plain, (k4_matrix_1(X1,k9_finseq_1(X2))=X3|r2_hidden(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3),k9_xtuple_0(k2_matrlin(X5,k9_finseq_1(X6))))|v1_xboole_0(X5)|v1_xboole_0(X1)|v1_xboole_0(X4)|k2_matrlin(X5,k9_finseq_1(X6))!=k4_matrix_1(X5,k9_finseq_1(X2))|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k1_matrix_1(X3)|~m1_finseq_1(k9_finseq_1(X6),X5)|~r2_hidden(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3),k1_tarski(np__1))|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(X2)),X3))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3)|~v1_finseq_1(k2_matrlin(X5,k9_finseq_1(X6)))|~v1_funct_1(k2_matrlin(X5,k9_finseq_1(X6)))|~v1_relat_1(k2_matrlin(X5,k9_finseq_1(X6))))).#trainneg
cnf(i_0_8912, plain, (r2_hidden(X1,k1_tarski(np__1))|~r2_hidden(k2_tarski(k2_tarski(X2,X1),k1_tarski(X2)),k2_matrix_1(k9_finseq_1(X3)))|k3_finseq_1(X3)!=np__1|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_8864, plain, (k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X4),np__1)=k3_matrix_1(X1,X5,esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X4),np__1)|k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X4|v1_xboole_0(X3)|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X4)|X5!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X5,k3_finseq_2(X1))|~m2_finseq_1(X4,k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(X5)|~v1_matrix_1(X4)|~r2_hidden(esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X4),k1_tarski(np__1))|~v7_ordinal1(esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X4)))).#trainneg
cnf(i_0_5520, plain, (esk8_3(X1,k4_matrix_1(X2,k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),X5)=np__1|k4_matrix_1(X2,k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))=X5|v1_xboole_0(X3)|v1_xboole_0(X1)|v1_xboole_0(X2)|k3_finseq_1(k4_matrix_1(X2,k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X2,k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))!=k1_matrix_1(X5)|~m2_finseq_1(k4_matrix_1(X2,k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X2))|~m2_finseq_1(X5,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(X5)|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_4813, plain, (r2_hidden(k2_tarski(k2_tarski(esk8_3(X1,k4_matrix_1(X2,X3),X4),esk8_3(X1,k4_matrix_1(X2,X3),X4)),k1_tarski(esk8_3(X1,k4_matrix_1(X2,X3),X4))),k2_matrix_1(X5))|v1_xboole_0(X6)|k4_matrix_1(X2,X3)=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|X5!=k4_matrix_1(X6,X3)|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,X3),X4))|~m2_finseq_1(X5,k3_finseq_2(X6))|~m2_finseq_1(X3,k3_finseq_2(X6))|~v1_matrix_1(X5)|~v1_matrix_1(X3)|k3_finseq_1(k4_matrix_1(X2,X3))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,X3))!=k1_matrix_1(X4)|~r2_hidden(esk8_3(X1,k4_matrix_1(X2,X3),X4),k9_xtuple_0(X3))|~m2_finseq_1(k4_matrix_1(X2,X3),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_9064, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k9_xtuple_0(X3))|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_9124, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_finseq_1(X3))|~v1_xboole_0(X3)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_8945, plain, (k9_xtuple_0(np__0)=np__0)).#trainneg
cnf(i_0_7365, plain, (v2_xxreal_0(esk22_0)|r1_xxreal_0(esk22_0,np__0)|~v2_xxreal_0(np__0)|~v1_xreal_0(esk22_0))).#trainneg
cnf(i_0_932, plain, (v1_xboole_0(esk19_1(X1))|~v1_funct_1(esk19_1(X1))|~v1_relat_1(esk19_1(X1))|~v1_xboole_0(k3_finseq_1(esk19_1(X1)))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_937, plain, (v1_xboole_0(k9_xtuple_0(X1))|v1_xboole_0(k2_finseq_1(k1_matrix_1(X1)))|~v1_xboole_0(k2_matrix_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_7563, plain, (k2_zfmisc_1(k1_tarski(np__1),k2_finseq_1(k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X1)))))=k2_matrix_1(k4_matrix_1(X2,k9_finseq_1(X1)))|v1_xboole_0(X2)|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~v1_matrix_1(k9_finseq_1(X1))|k3_finseq_1(X1)!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8865, plain, (k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X4),esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X4))=k3_matrix_1(X1,X5,esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X4),esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X4))|k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X4|v1_xboole_0(X3)|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X4)|X5!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~r2_hidden(esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X4),k9_xtuple_0(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X5,k3_finseq_2(X1))|~m2_finseq_1(X4,k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(X5)|~v1_matrix_1(X4)|~r2_hidden(esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X4),k1_tarski(np__1))|~v7_ordinal1(esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X4)))).#trainneg
cnf(i_0_9228, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),np__1)!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)),esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~r2_hidden(esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~v7_ordinal1(esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1)))).#trainneg
cnf(i_0_9460, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),np__1,esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k9_finseq_1(k9_finseq_1(X3))!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))))).#trainneg
cnf(i_0_9223, plain, (esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4)=np__1|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X1)|v1_xboole_0(X2)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_9205, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),esk8_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)),esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),esk8_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~r2_hidden(esk8_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~r2_hidden(esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~v7_ordinal1(esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1)))).#trainneg
cnf(i_0_9216, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X3)|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k9_finseq_1(k9_finseq_1(X4))!=k4_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))))).#trainneg
cnf(i_0_9571, plain, (~v3_card_1(np__0,np__1))).#trainneg
cnf(i_0_941, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_matrix_1(X2)))|~v1_xboole_0(k9_xtuple_0(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_9405, plain, (r2_hidden(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),k4_finseq_1(k4_matrix_1(X5,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))))|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X5)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X5,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~v1_funct_1(k4_matrix_1(X5,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~v1_relat_1(k4_matrix_1(X5,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~v1_matrix_1(k4_matrix_1(X5,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X5))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_9404, plain, (r2_hidden(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),k9_xtuple_0(k4_matrix_1(X5,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))))|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X5)|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X5,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~v1_funct_1(k4_matrix_1(X5,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~v1_relat_1(k4_matrix_1(X5,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~v1_matrix_1(k4_matrix_1(X5,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X5))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_9610, plain, (~v1_xboole_0(k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1))))).#trainneg
cnf(i_0_9222, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X3)|v1_xboole_0(X1)|k3_matrix_1(X3,X5,esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),np__1)!=k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)),esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|X5!=k4_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(X5,k3_finseq_2(X3))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_8210, plain, (v2_xxreal_0(esk22_0)|r1_xxreal_0(esk22_0,np__0)|~v2_xxreal_0(np__0))).#trainneg
cnf(i_0_942, plain, (v4_relat_1(X1,k9_xtuple_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(k2_matrix_1(X2)))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_2276, plain, (r1_xxreal_0(np__0,X1)|r1_xxreal_0(X1,np__0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_8059, plain, (v2_xxreal_0(np__0)|r1_xxreal_0(np__0,esk22_0)|~v2_xxreal_0(esk22_0)|~v1_xxreal_0(esk22_0))).#trainneg
cnf(i_0_9127, plain, (r1_xxreal_0(esk25_0,esk22_0)|r1_xxreal_0(esk22_0,np__1)|~v1_xreal_0(esk22_0)|~v1_xxreal_0(np__1))).#trainneg
cnf(i_0_8271, plain, (v1_card_1(k1_finseq_1(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_951, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(X2)))|v1_xboole_0(k2_matrix_1(X2))|~m1_subset_1(k2_tarski(k2_tarski(X3,X1),k1_tarski(X3)),k2_matrix_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_9256, plain, (r2_hidden(X1,k1_tarski(np__1))|v1_xboole_0(X2)|~r2_hidden(k2_tarski(k2_tarski(X3,X1),k1_tarski(X1)),k2_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X4)))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X2)))).#trainneg
cnf(i_0_9253, plain, (r2_hidden(X1,k1_tarski(np__1))|v1_xboole_0(X2)|~r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X3)),k2_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X4)))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X2)))).#trainneg
cnf(i_0_6740, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5),X1)),k2_matrix_1(k4_matrix_1(X6,k9_finseq_1(X4))))|k4_matrix_1(X3,k9_finseq_1(X4))=X5|v1_xboole_0(X2)|v1_xboole_0(X3)|v1_xboole_0(X6)|~r2_hidden(X1,k9_xtuple_0(k4_matrix_1(X6,k9_finseq_1(X4))))|~v1_finseq_1(k4_matrix_1(X6,k9_finseq_1(X4)))|~v1_funct_1(k4_matrix_1(X6,k9_finseq_1(X4)))|~v1_relat_1(k4_matrix_1(X6,k9_finseq_1(X4)))|~v1_matrix_1(k4_matrix_1(X6,k9_finseq_1(X4)))|k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(X4)))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(X4)))!=k1_matrix_1(X5)|~r2_hidden(esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5),k1_tarski(np__1))|~v7_ordinal1(esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(X4)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X4),k3_finseq_2(X6))|~m2_finseq_1(k9_finseq_1(X4),k3_finseq_2(X3))|~m2_finseq_1(X5,k3_finseq_2(X2))|~v1_matrix_1(k9_finseq_1(X4))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_9126, plain, (v2_xxreal_0(np__0)|r1_xxreal_0(np__0,esk22_0)|r1_xxreal_0(esk22_0,np__1)|~v1_xreal_0(esk22_0)|~v1_xxreal_0(np__1))).#trainneg
cnf(i_0_2324, plain, (r1_xxreal_0(esk25_0,X1)|r1_xxreal_0(X1,esk25_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_9570, plain, (k2_zfmisc_1(np__0,k2_finseq_1(k1_matrix_1(np__0)))=k2_matrix_1(np__0)|~v1_matrix_1(np__0))).#trainneg
cnf(i_0_2306, plain, (k2_zfmisc_1(k1_tarski(np__1),k2_finseq_1(k1_matrix_1(np__1)))=k2_matrix_1(np__1)|~v1_finseq_1(np__1)|~v1_funct_1(np__1)|~v1_relat_1(np__1)|~v1_matrix_1(np__1))).#trainneg
cnf(i_0_953, plain, (r2_hidden(X1,X2)|v1_xboole_0(k2_zfmisc_1(X3,X2))|~m1_subset_1(k2_tarski(k2_tarski(X1,X4),k1_tarski(X4)),k2_zfmisc_1(X3,X2)))).#trainneg
cnf(i_0_4969, plain, (r2_hidden(esk9_3(X1,k9_finseq_1(X2),X3),k4_finseq_1(k4_matrix_1(X4,k9_finseq_1(X2))))|k9_finseq_1(X2)=X3|v1_xboole_0(X4)|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X4,k9_finseq_1(X2)))|~v1_funct_1(k4_matrix_1(X4,k9_finseq_1(X2)))|~v1_relat_1(k4_matrix_1(X4,k9_finseq_1(X2)))|~v1_matrix_1(k4_matrix_1(X4,k9_finseq_1(X2)))|k1_matrix_1(k9_finseq_1(X2))!=k1_matrix_1(X3)|k3_finseq_1(X3)!=np__1|~r2_hidden(esk9_3(X1,k9_finseq_1(X2),X3),k4_finseq_1(k9_finseq_1(X2)))|~v7_ordinal1(esk9_3(X1,k9_finseq_1(X2),X3))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_4773, plain, (r2_hidden(X1,k1_tarski(np__1))|k4_matrix_1(X3,k9_finseq_1(X4))=X5|v1_xboole_0(X3)|v1_xboole_0(X2)|~v1_matrix_1(k9_finseq_1(X4))|k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(X4)))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(X4)))!=k1_matrix_1(X5)|~r2_hidden(X1,k4_finseq_1(k9_finseq_1(X4)))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(X4)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X4),k3_finseq_2(X3))|~m2_finseq_1(X5,k3_finseq_2(X2))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_5898, plain, (v1_relat_1(esk26_1(k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1))))|v1_xboole_0(k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1))))).#trainneg
cnf(i_0_8997, plain, (k1_card_1(X2)!=k2_tarski(np__1,np__2)|np__2!=np__1|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2337, plain, (r1_xxreal_0(esk17_0,X1)|r1_xxreal_0(X1,esk17_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5899, plain, (v1_relat_1(esk24_1(k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1))))|v1_xboole_0(k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1))))).#trainneg
cnf(i_0_5900, plain, (v1_relat_1(esk20_1(k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1))))|v1_xboole_0(k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1))))).#trainneg
cnf(i_0_980, plain, (esk11_1(k1_zfmisc_1(X1))=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5832, plain, (r2_hidden(X1,k2_finseq_1(k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))))|k4_matrix_1(X5,k9_finseq_1(X3))=X6|v1_xboole_0(X2)|v1_xboole_0(X5)|v1_xboole_0(X4)|~v1_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)))|~v1_funct_1(k4_matrix_1(X2,k9_finseq_1(X3)))|~v1_relat_1(k4_matrix_1(X2,k9_finseq_1(X3)))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))|k3_finseq_1(k4_matrix_1(X5,k9_finseq_1(X3)))!=k3_finseq_1(X6)|k1_matrix_1(k4_matrix_1(X5,k9_finseq_1(X3)))!=k1_matrix_1(X6)|~r2_hidden(X1,k1_tarski(np__1))|~v7_ordinal1(esk8_3(X4,k4_matrix_1(X5,k9_finseq_1(X3)),X6))|~v7_ordinal1(X1)|~m2_finseq_1(k4_matrix_1(X5,k9_finseq_1(X3)),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X5))|~m2_finseq_1(X6,k3_finseq_2(X4))|~v1_matrix_1(k9_finseq_1(X3))|~v1_matrix_1(X6))).#trainneg
cnf(i_0_6942, plain, (r2_hidden(k2_tarski(k2_tarski(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),X5),k1_tarski(X5)),k2_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X1)|v1_xboole_0(X2)|~r2_hidden(X5,k4_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~v1_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_funct_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_relat_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_9760, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),np__1,np__1)!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))).#trainneg
cnf(i_0_9757, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X3)|v1_xboole_0(X1)|k3_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),np__1,np__1)!=k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))).#trainneg
cnf(i_0_9761, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),np__1,np__1)!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k9_finseq_1(k9_finseq_1(X3))!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))).#trainneg
cnf(i_0_9784, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),np__1,esk8_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)),np__1,esk8_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~r2_hidden(esk8_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))).#trainneg
cnf(i_0_9747, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k9_finseq_1(k9_finseq_1(X3))!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~r2_hidden(esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~v7_ordinal1(esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))))).#trainneg
cnf(i_0_9750, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X3)|v1_xboole_0(X1)|k3_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),esk8_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),np__1)!=k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)),np__1,esk8_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~r2_hidden(esk8_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1)))).#trainneg
cnf(i_0_9912, plain, (v1_card_1(k1_finseq_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_10047, plain, (k2_matrix_1(np__0)=np__0|~v1_matrix_1(np__0))).#trainneg
cnf(i_0_981, plain, (esk11_1(k1_zfmisc_1(X1))=k1_card_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_8520, plain, (v1_card_1(k9_xtuple_0(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_8535, plain, (v1_relat_1(k9_xtuple_0(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_982, plain, (esk11_1(k1_zfmisc_1(X1))=k9_xtuple_0(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_990, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_9753, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|v1_xboole_0(X4)|v1_xboole_0(X1)|k3_matrix_1(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3),np__1)!=k3_matrix_1(X4,X3,esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3),np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1)))).#trainneg
cnf(i_0_9858, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|v1_xboole_0(X3)|k3_matrix_1(X3,X5,np__1,np__1)!=k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|X5!=k4_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m2_finseq_1(X5,k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_10076, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_matrix_1(np__0)))|~v1_matrix_1(np__0))).#trainneg
cnf(i_0_10037, plain, (v1_relat_1(k2_matrix_1(np__0))|~v1_matrix_1(np__0))).#trainneg
cnf(i_0_8969, plain, (v1_card_1(esk16_1(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_9742, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,X4,esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),np__1)!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)),esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~r2_hidden(esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|X4!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_9758, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|v1_xboole_0(X1)|k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),np__1,esk8_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3))!=k3_matrix_1(X1,X3,esk8_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3),np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~r2_hidden(esk8_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_9759, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|v1_xboole_0(X1)|k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),np__1,np__1)!=k3_matrix_1(X1,X3,np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_9797, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|v1_xboole_0(X3)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k9_finseq_1(k9_finseq_1(X4))!=k4_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))).#trainneg
cnf(i_0_9962, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),np__1,esk8_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)),np__1,esk8_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=np__1|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=np__1|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~r2_hidden(esk8_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1)))).#trainneg
cnf(i_0_9764, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|v1_xboole_0(X1)|v1_xboole_0(X4)|k3_matrix_1(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),np__1,np__1)!=k3_matrix_1(X4,X3,np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_10215, plain, (k4_matrix_1(X1,k9_finseq_1(X2))=X3|r2_hidden(X4,k1_tarski(np__1))|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k1_matrix_1(X3)|~r2_hidden(X4,k4_finseq_1(k9_finseq_1(X2)))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X2))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_10072, plain, (v1_xboole_0(k2_matrix_1(np__0))|~v1_matrix_1(np__0))).#trainneg
cnf(i_0_9848, plain, (~v1_xboole_0(k2_zfmisc_1(k2_tarski(np__1,np__2),k2_tarski(np__1,np__2)))|np__2!=np__1)).#trainneg
cnf(i_0_8984, plain, (v1_relat_1(esk16_1(X2))|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1011, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3))|v1_xboole_0(X4)|v1_xboole_0(k2_matrix_1(X5))|X5!=k4_matrix_1(X4,X3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(X5)|~m1_subset_1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k2_matrix_1(X5)))).#trainneg
cnf(i_0_9748, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,X4,esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),np__1)!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k9_finseq_1(k9_finseq_1(X3))!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|X4!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_9868, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|v1_xboole_0(X3)|k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X5)),np__1,esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))!=k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)),esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X5)))!=k4_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X5))),k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X5))))|~r2_hidden(esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),k1_tarski(np__1))|~v7_ordinal1(esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X5)),k3_finseq_2(X3)))).#trainneg
cnf(i_0_10045, plain, (k1_zfmisc_1(k2_matrix_1(np__0))!=k2_tarski(np__1,np__2)|np__2!=np__1|~v1_matrix_1(np__0))).#trainneg
cnf(i_0_8274, plain, (v1_finset_1(k1_finseq_1(X2))|~v1_finset_1(X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_8279, plain, (v7_ordinal1(k1_finseq_1(X2))|~v1_finset_1(X1)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_10290, plain, (r2_hidden(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),k9_xtuple_0(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_funct_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_relat_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~r2_hidden(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),k4_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_7802, plain, (k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),X3,X3)=k3_matrix_1(X1,X4,X3,X3)|v1_xboole_0(X1)|X4!=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))|~v7_ordinal1(X3)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4)|~r2_hidden(X3,k1_tarski(np__1)))).#trainneg
cnf(i_0_6943, plain, (r2_hidden(k2_tarski(k2_tarski(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),X5),k1_tarski(X5)),k2_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X1)|v1_xboole_0(X2)|~r2_hidden(X5,k9_xtuple_0(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~v1_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_funct_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_relat_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_4837, plain, (r2_hidden(k2_tarski(k2_tarski(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4)),k1_tarski(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4))),k2_matrix_1(X5))|v1_xboole_0(X6)|k4_matrix_1(X2,k9_finseq_1(X3))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|k9_finseq_1(X3)!=k4_matrix_1(X6,X5)|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4))|~m2_finseq_1(X5,k3_finseq_2(X6))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X6))|~v1_matrix_1(X5)|~v1_matrix_1(k9_finseq_1(X3))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k1_matrix_1(X4)|~r2_hidden(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_8290, plain, (v1_finset_1(k1_card_1(X2))|~v7_ordinal1(X1)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1037, plain, (r2_hidden(X1,X2)|v1_xboole_0(k2_zfmisc_1(X2,X4))|~m1_subset_1(k2_tarski(k2_tarski(X3,X1),k1_tarski(X1)),k2_zfmisc_1(X2,X4)))).#trainneg
cnf(i_0_1044, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3))|v1_xboole_0(X4)|v1_xboole_0(k2_matrix_1(X5))|X3!=k4_matrix_1(X4,X5)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(X5)|~m1_subset_1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X2)),k2_matrix_1(X5)))).#trainneg
cnf(i_0_10310, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),np__1,np__1)!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k9_finseq_1(k9_finseq_1(X3))!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))).#trainneg
cnf(i_0_10306, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),np__1,np__1)!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))).#trainneg
cnf(i_0_10327, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k9_finseq_1(k9_finseq_1(X3))!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))).#trainneg
cnf(i_0_10379, plain, (v1_xboole_0(X1)|~v1_matrix_1(np__0)|~m1_subset_1(X1,k1_zfmisc_1(np__0)))).#trainneg
cnf(i_0_10460, plain, (v1_card_1(k9_xtuple_0(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_10756, plain, (k1_zfmisc_1(np__0)!=k2_tarski(np__1,np__2)|np__2!=np__1|~v1_matrix_1(np__0))).#trainneg
cnf(i_0_1045, plain, (m1_subset_1(esk3_2(X1,esk28_0),k1_numbers)|esk3_2(X1,esk28_0)=X1|k1_tarski(X1)=esk28_0)).#trainneg
cnf(i_0_1052, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X2)|r1_xxreal_0(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_1056, plain, (esk11_1(k1_zfmisc_1(X1))=k1_finseq_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_10744, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X3)|v1_xboole_0(X1)|k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X5)),np__1,np__1)!=k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X5)))!=k4_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X5))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X5)),k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X5)))))).#trainneg
cnf(i_0_10617, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X4)),np__1,esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)),esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X4)))!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~r2_hidden(esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X4))))|~v7_ordinal1(esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X1)))).#trainneg
cnf(i_0_10480, plain, (v1_relat_1(k9_xtuple_0(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1058, plain, (k1_finseq_1(X1)=k1_finseq_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4261, plain, (r1_xxreal_0(X1,esk25_0)|r1_xxreal_0(esk25_0,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_10607, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,X4,np__1,np__1)!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|X4!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_10823, plain, (r2_hidden(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),k4_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_funct_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_relat_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~r2_hidden(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),k9_xtuple_0(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_10568, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X3)|v1_xboole_0(X1)|k3_matrix_1(X3,X5,np__1,np__1)!=k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|X5!=k4_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m2_finseq_1(X5,k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_10592, plain, (v1_card_1(esk16_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_9051, plain, (k9_xtuple_0(X2)!=k2_tarski(np__1,np__2)|np__2!=np__1|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_4352, plain, (r1_xxreal_0(X1,esk17_0)|r1_xxreal_0(esk17_0,X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1062, plain, (v1_relat_1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)))|~v4_funct_1(k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_1063, plain, (v1_funct_1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)))|~v4_funct_1(k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_10805, plain, (k3_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3,X3)=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),X3,X3)|v1_xboole_0(X1)|~r2_hidden(X3,k1_tarski(np__1))|~v7_ordinal1(X3)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))))).#trainneg
cnf(i_0_10631, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|v1_xboole_0(X1)|k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),np__1,np__1)!=k3_matrix_1(X1,X3,np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_10335, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|k3_matrix_1(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),esk8_3(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)))),np__1)!=X3|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~m2_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m1_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(k1_tarski(X3)))|~v7_ordinal1(esk8_3(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))))).#trainneg
cnf(i_0_10703, plain, (v1_relat_1(esk16_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_9902, plain, (r1_xxreal_0(esk22_0,np__1)|r1_xxreal_0(esk25_0,esk22_0)|~v1_xreal_0(esk22_0))).#trainneg
cnf(i_0_9110, plain, (k1_finseq_1(X2)!=k2_tarski(np__1,np__2)|np__2!=np__1|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1064, plain, (v3_ordinal1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)))|~v3_ordinal1(k2_zfmisc_1(X3,X4))|~r2_hidden(X2,X4)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_4523, plain, (r1_xxreal_0(esk22_0,X1)|r1_xxreal_0(X1,esk22_0)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1067, plain, (m1_subset_1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3))|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X1,k4_finseq_1(X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_1068, plain, (m1_subset_1(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3))|~r2_hidden(X2,k2_finseq_1(k1_matrix_1(X3)))|~r2_hidden(X1,k9_xtuple_0(X3))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_6476, plain, (k9_finseq_1(k9_finseq_1(X1))=X2|r2_hidden(k2_tarski(k1_tarski(esk8_3(X3,k9_finseq_1(k9_finseq_1(X1)),X2)),k2_tarski(esk8_3(X3,k9_finseq_1(k9_finseq_1(X1)),X2),np__1)),k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1)))|v1_xboole_0(X3)|k3_finseq_1(X2)!=np__1|k1_matrix_1(X2)!=np__1|~m2_finseq_1(X2,k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X1)),k3_finseq_2(X3))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_4981, plain, (k9_finseq_1(X1)=X2|r2_hidden(k2_tarski(k1_tarski(esk9_3(X3,k9_finseq_1(X1),X2)),k2_tarski(esk9_3(X3,k9_finseq_1(X1),X2),X4)),k2_matrix_1(k4_matrix_1(X5,k9_finseq_1(X1))))|v1_xboole_0(X5)|v1_xboole_0(X3)|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(X2)|k3_finseq_1(X2)!=np__1|~r2_hidden(X4,k4_finseq_1(k9_finseq_1(X1)))|~v7_ordinal1(esk9_3(X3,k9_finseq_1(X1),X2))|~v7_ordinal1(X4)|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X3))|~m2_finseq_1(X2,k3_finseq_2(X3))|~v1_matrix_1(k9_finseq_1(X1))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_8331, plain, (r2_hidden(k2_tarski(k2_tarski(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4)),k1_tarski(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4))),k2_matrix_1(X5))|v1_xboole_0(X6)|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))!=k4_matrix_1(X6,X5)|~v7_ordinal1(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4))|~m2_finseq_1(X5,k3_finseq_2(X6))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X6))|~v1_matrix_1(X5)|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~r2_hidden(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),k9_xtuple_0(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_10765, plain, (v1_finset_1(k1_finseq_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_10051, plain, (v1_xboole_0(esk16_1(k2_matrix_1(np__0)))|~v1_matrix_1(np__0))).#trainneg
cnf(i_0_7038, plain, (r1_xxreal_0(np__2,X1)|r1_xxreal_0(X1,np__2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_10056, plain, (v1_relat_1(esk16_1(k2_matrix_1(np__0)))|~v1_matrix_1(np__0))).#trainneg
cnf(i_0_10073, plain, (v4_relat_1(k2_matrix_1(np__0),np__0)|~v1_matrix_1(np__0))).#trainneg
cnf(i_0_9940, plain, (r2_hidden(X1,k1_tarski(np__1))|v1_xboole_0(X2)|v1_xboole_0(k2_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X4)))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X2))|~m1_subset_1(k2_tarski(k2_tarski(X3,X1),k1_tarski(X1)),k2_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X4))))))).#trainneg
cnf(i_0_4788, plain, (r2_hidden(k2_tarski(k2_tarski(esk8_3(X1,k4_matrix_1(X2,X3),X4),esk8_3(X1,k4_matrix_1(X2,X3),X4)),k1_tarski(esk8_3(X1,k4_matrix_1(X2,X3),X4))),k2_matrix_1(X5))|v1_xboole_0(X6)|k4_matrix_1(X2,X3)=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|X3!=k4_matrix_1(X6,X5)|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,X3),X4))|~m2_finseq_1(X5,k3_finseq_2(X6))|~m2_finseq_1(X3,k3_finseq_2(X6))|~v1_matrix_1(X5)|~v1_matrix_1(X3)|k3_finseq_1(k4_matrix_1(X2,X3))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,X3))!=k1_matrix_1(X4)|~r2_hidden(esk8_3(X1,k4_matrix_1(X2,X3),X4),k4_finseq_1(X3))|~m2_finseq_1(k4_matrix_1(X2,X3),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_10904, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),np__1,np__1)!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))).#trainneg
cnf(i_0_11087, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|v1_xboole_0(X3)|k3_matrix_1(X3,k4_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))),np__1,np__1)!=k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))),k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))))).#trainneg
cnf(i_0_10906, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k9_finseq_1(k9_finseq_1(X3))!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))).#trainneg
cnf(i_0_11060, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,X4,np__1,np__1)!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|X4!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_10992, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X4)),np__1,np__1)!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X4)))!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X4)))))).#trainneg
cnf(i_0_10989, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|v1_xboole_0(X3)|k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X5)),np__1,np__1)!=k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X5)))!=k4_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X5))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X5)),k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X5)))))).#trainneg
cnf(i_0_11080, plain, (X1=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,X1,np__1,np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k1_matrix_1(X1)!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(X1)!=np__1|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X1)|k3_finseq_1(X1)!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11081, plain, (k9_finseq_1(k9_finseq_1(X1))=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X1)),np__1,np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=np__1|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X1)),k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))).#trainneg
cnf(i_0_10786, plain, (v7_ordinal1(k1_finseq_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1076, plain, (m1_matrix_1(k9_finseq_1(k3_finseq_1(X1)),X2,np__1,k3_finseq_1(X1))|v1_xboole_0(X2)|~m1_finseq_1(k3_finseq_1(X1),X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1077, plain, (m1_matrix_1(X1,X2,np__1,k3_finseq_1(k1_funct_1(X1,np__1)))|v1_xboole_0(X2)|~m1_finseq_1(k1_funct_1(X1,np__1),X2)|k3_finseq_1(X1)!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11203, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~m1_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(k1_tarski(X3)))|~v7_ordinal1(esk8_3(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)))))|~m2_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m1_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X3))))).#trainneg
cnf(i_0_11086, plain, (k9_finseq_1(X1)=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,k9_finseq_1(X1),np__1,np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=np__1|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k9_finseq_1(X1))|k3_finseq_1(X1)!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11298, plain, (r2_hidden(esk8_3(X1,k9_finseq_1(k9_finseq_1(X2)),X3),k1_tarski(np__1))|k9_finseq_1(k9_finseq_1(X2))=X3|v1_xboole_0(X1)|k3_finseq_1(X3)!=np__1|k1_matrix_1(X3)!=np__1|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_11198, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|k3_matrix_1(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),np__1,np__1)!=X3|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~m1_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))))).#trainneg
cnf(i_0_8058, plain, (k9_xtuple_0(k4_matrix_1(X1,k9_finseq_1(X2)))=k1_tarski(np__1)|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)))|~v1_funct_1(k4_matrix_1(X1,k9_finseq_1(X2)))|~v1_relat_1(k4_matrix_1(X1,k9_finseq_1(X2)))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|k3_finseq_1(X2)!=np__1|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_10850, plain, (v1_finset_1(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1082, plain, (k3_matrix_1(X1,X2,X3,X4)=k3_matrix_1(X1,X5,X4,X3)|v1_xboole_0(X1)|v1_xboole_0(k2_matrix_1(X2))|X5!=k4_matrix_1(X1,X2)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m2_finseq_1(X5,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X5)|~v1_matrix_1(X2)|~m1_subset_1(k2_tarski(k1_tarski(X3),k2_tarski(X3,X4)),k2_matrix_1(X2)))).#trainneg
cnf(i_0_1083, plain, (k3_matrix_1(X1,X2,esk8_3(X3,X2,X4),esk9_3(X3,X2,X4))=k3_matrix_1(X1,X5,esk9_3(X3,X2,X4),esk8_3(X3,X2,X4))|v1_xboole_0(X1)|X2=X4|v1_xboole_0(X3)|X5!=k4_matrix_1(X1,X2)|~v7_ordinal1(esk8_3(X3,X2,X4))|~v7_ordinal1(esk9_3(X3,X2,X4))|~m2_finseq_1(X5,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X5)|~v1_matrix_1(X2)|k3_finseq_1(X2)!=k3_finseq_1(X4)|k1_matrix_1(X2)!=k1_matrix_1(X4)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X2,k3_finseq_2(X3))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_1084, plain, (k3_matrix_1(X1,X2,esk7_3(X3,X2,X4),esk6_3(X3,X2,X4))=k3_matrix_1(X1,X5,esk6_3(X3,X2,X4),esk7_3(X3,X2,X4))|v1_xboole_0(X1)|X4=k4_matrix_1(X3,X2)|v7_ordinal1(esk4_3(X3,X2,X4))|v1_xboole_0(X3)|X5!=k4_matrix_1(X1,X2)|~v7_ordinal1(esk7_3(X3,X2,X4))|~v7_ordinal1(esk6_3(X3,X2,X4))|~m2_finseq_1(X5,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X5)|~v1_matrix_1(X2)|k1_matrix_1(X2)!=k3_finseq_1(X4)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X2,k3_finseq_2(X3))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_5824, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5))),k2_matrix_1(X6))|v1_xboole_0(X7)|k4_matrix_1(X3,k9_finseq_1(X4))=X5|v1_xboole_0(X8)|v1_xboole_0(X3)|v1_xboole_0(X2)|X6!=k4_matrix_1(X7,k4_matrix_1(X8,k9_finseq_1(X4)))|~v7_ordinal1(X1)|~v7_ordinal1(esk8_3(X2,k4_matrix_1(X3,k9_finseq_1(X4)),X5))|~m2_finseq_1(X6,k3_finseq_2(X7))|~m2_finseq_1(k4_matrix_1(X8,k9_finseq_1(X4)),k3_finseq_2(X7))|~v1_matrix_1(X6)|~v1_matrix_1(k4_matrix_1(X8,k9_finseq_1(X4)))|k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(X4)))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(X4)))!=k1_matrix_1(X5)|~r2_hidden(X1,k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(X4)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X4),k3_finseq_2(X8))|~m2_finseq_1(k9_finseq_1(X4),k3_finseq_2(X3))|~m2_finseq_1(X5,k3_finseq_2(X2))|~v1_matrix_1(k9_finseq_1(X4))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_5119, plain, (k2_zfmisc_1(k1_tarski(np__1),k2_finseq_1(k1_matrix_1(k4_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))))))=k2_matrix_1(k4_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))))|v1_xboole_0(X3)|v1_xboole_0(X1)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_6957, plain, (r2_hidden(k2_tarski(k2_tarski(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),X5),k1_tarski(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4))),k2_matrix_1(k4_matrix_1(X6,k9_finseq_1(k9_finseq_1(X3)))))|v1_xboole_0(X6)|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|~r2_hidden(X5,k1_tarski(np__1))|~v1_finseq_1(k4_matrix_1(X6,k9_finseq_1(k9_finseq_1(X3))))|~v1_funct_1(k4_matrix_1(X6,k9_finseq_1(k9_finseq_1(X3))))|~v1_relat_1(k4_matrix_1(X6,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k4_matrix_1(X6,k9_finseq_1(k9_finseq_1(X3))))|~m1_subset_1(X3,X6)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X6))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_1086, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k9_xtuple_0(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1085, plain, (k3_matrix_1(X1,X2,esk7_3(X3,X2,X4),esk6_3(X3,X2,X4))=k3_matrix_1(X1,X5,esk6_3(X3,X2,X4),esk7_3(X3,X2,X4))|v1_xboole_0(X1)|X4=k4_matrix_1(X3,X2)|v7_ordinal1(esk5_3(X3,X2,X4))|v1_xboole_0(X3)|X5!=k4_matrix_1(X1,X2)|~v7_ordinal1(esk7_3(X3,X2,X4))|~v7_ordinal1(esk6_3(X3,X2,X4))|~m2_finseq_1(X5,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X5)|~v1_matrix_1(X2)|k1_matrix_1(X2)!=k3_finseq_1(X4)|~m2_finseq_1(X4,k3_finseq_2(X3))|~m2_finseq_1(X2,k3_finseq_2(X3))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_1087, plain, (m1_subset_1(esk3_2(X1,k4_finseq_1(X2)),k5_numbers)|esk3_2(X1,k4_finseq_1(X2))=X1|k4_finseq_1(X2)=k1_tarski(X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5362, plain, (k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),X3,X4)=k3_matrix_1(X1,X5,X4,X3)|v1_xboole_0(X1)|X5!=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))|~r2_hidden(k2_tarski(k2_tarski(X4,X3),k1_tarski(X3)),k2_zfmisc_1(k1_tarski(np__1),k1_tarski(np__1)))|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m2_finseq_1(X5,k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~v1_matrix_1(X5))).#trainneg
cnf(i_0_7914, plain, (r2_hidden(X1,k1_tarski(np__1))|v1_xboole_0(X2)|~r2_hidden(k2_tarski(k1_tarski(X3),k2_tarski(X3,X1)),k2_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X4)))))|~v1_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X4))))|~v1_funct_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X4))))|~v1_relat_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X4))))|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_6401, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|r2_hidden(k2_tarski(k1_tarski(np__1),k2_tarski(np__1,X5)),k2_matrix_1(k4_matrix_1(X6,k9_finseq_1(k9_finseq_1(X2)))))|v1_xboole_0(X4)|v1_xboole_0(X1)|v1_xboole_0(X6)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~r2_hidden(X5,k1_tarski(np__1))|~v7_ordinal1(X5)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X4))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X6))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X4))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_6382, plain, (k2_tarski(k1_tarski(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4)),k2_tarski(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),X5))=X6|k4_matrix_1(X2,k9_finseq_1(X3))=X4|v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X7)|k2_matrix_1(k4_matrix_1(X7,k9_finseq_1(X3)))!=k1_tarski(X6)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k1_matrix_1(X4)|~r2_hidden(X5,k1_tarski(np__1))|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4))|~v7_ordinal1(X5)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X7))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X3))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_5136, plain, (k3_finseq_1(k4_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))=np__1|v1_xboole_0(X1)|v1_xboole_0(X2)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_8622, plain, (esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4)=X5|k4_matrix_1(X2,k9_finseq_1(X3))=X4|v1_xboole_0(X6)|v1_xboole_0(X2)|v1_xboole_0(X1)|k4_finseq_1(k2_matrlin(X6,k9_finseq_1(X7)))!=k1_tarski(X5)|k2_matrlin(X6,k9_finseq_1(X7))!=k4_matrix_1(X6,k9_finseq_1(X3))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(X3)))!=k1_matrix_1(X4)|~m1_finseq_1(k9_finseq_1(X7),X6)|~r2_hidden(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4),k1_tarski(np__1))|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(X3)),X4))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X6))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k9_finseq_1(X3))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_5123, plain, (r2_hidden(k2_tarski(k1_tarski(X1),k2_tarski(X1,X2)),k2_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|v1_xboole_0(X3)|~r2_hidden(X2,k1_tarski(np__1))|~r2_hidden(X1,k9_xtuple_0(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|~v1_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_funct_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_relat_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_11439, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))),np__1,np__1)!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))))).#trainneg
cnf(i_0_11209, plain, (r1_xxreal_0(esk25_0,esk22_0)|r1_xxreal_0(esk22_0,np__1))).#trainneg
cnf(i_0_1088, plain, (m1_subset_1(X1,k5_numbers)|v1_xboole_0(k4_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(X1,k4_finseq_1(X2)))).#trainneg
cnf(i_0_10052, plain, (v1_xboole_0(esk11_1(k1_zfmisc_1(k2_matrix_1(np__0))))|~v1_matrix_1(np__0))).#trainneg
cnf(i_0_1099, plain, (k2_tarski(k2_tarski(X1,esk3_2(X2,X3)),k1_tarski(X1))=X4|esk3_2(X2,X3)=X2|X3=k1_tarski(X2)|k2_zfmisc_1(X5,X3)!=k1_tarski(X4)|~r2_hidden(X1,X5))).#trainneg
cnf(i_0_1100, plain, (k2_tarski(k2_tarski(X1,k2_tarski(k1_tarski(esk8_3(X2,X3,X4)),k2_tarski(esk8_3(X2,X3,X4),esk9_3(X2,X3,X4)))),k1_tarski(X1))=X5|X3=X4|v1_xboole_0(X2)|k2_zfmisc_1(X6,k2_matrix_1(X3))!=k1_tarski(X5)|~r2_hidden(X1,X6)|k3_finseq_1(X3)!=k3_finseq_1(X4)|k1_matrix_1(X3)!=k1_matrix_1(X4)|~m2_finseq_1(X4,k3_finseq_2(X2))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X4)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_11442, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X4)),np__1,np__1)!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X4)))!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X4)))))).#trainneg
cnf(i_0_11444, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X3)|v1_xboole_0(X1)|k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X5)),np__1,np__1)!=k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X5)))!=k4_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X5)),k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X5)))))).#trainneg
cnf(i_0_11607, plain, (X1=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_finseq_1(X1)!=k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(X1)!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~r2_hidden(esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~r2_hidden(esk9_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X1)|k9_finseq_1(k9_finseq_1(X3))!=k4_matrix_1(X2,X1))).#trainneg
cnf(i_0_11431, plain, (k4_matrix_1(X1,k9_finseq_1(X2))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,k9_finseq_1(X2),np__1,np__1)!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k3_finseq_1(X2)!=np__1|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_11455, plain, (k2_matrlin(X1,k9_finseq_1(X2))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X3)|v1_xboole_0(X1)|k3_matrix_1(X3,k2_matrlin(X1,k9_finseq_1(X2)),np__1,np__1)!=k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)),np__1,np__1)|k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))!=np__1|k3_finseq_1(k2_matrlin(X1,k9_finseq_1(X2)))!=np__1|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m2_finseq_1(k2_matrlin(X1,k9_finseq_1(X2)),k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k2_matrlin(X1,k9_finseq_1(X2)))|~m1_finseq_1(k9_finseq_1(X2),X1))).#trainneg
cnf(i_0_11438, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(X3))|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X3)))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X3)))|k9_finseq_1(X3)!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X3)))|k3_finseq_1(X3)!=np__1|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_11632, plain, (X1=X2|v1_xboole_0(X3)|k3_matrix_1(X3,X4,esk9_3(X3,X1,X2),esk8_3(X3,X1,X2))!=k3_matrix_1(X3,X2,esk8_3(X3,X1,X2),esk9_3(X3,X1,X2))|k3_finseq_1(X1)!=k3_finseq_1(X2)|k1_matrix_1(X1)!=k1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X3))|~m2_finseq_1(X1,k3_finseq_2(X3))|~v1_matrix_1(X2)|~v1_matrix_1(X1)|X4!=k4_matrix_1(X3,X1)|~m2_finseq_1(X4,k3_finseq_2(X3))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_11512, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~m1_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X3)))|~m1_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X3))))).#trainneg
cnf(i_0_11570, plain, (v7_ordinal1(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1101, plain, (k2_tarski(k2_tarski(X1,k2_tarski(k1_tarski(esk7_3(X2,X3,X4)),k2_tarski(esk7_3(X2,X3,X4),esk6_3(X2,X3,X4)))),k1_tarski(X1))=X5|X4=k4_matrix_1(X2,X3)|v7_ordinal1(esk4_3(X2,X3,X4))|v1_xboole_0(X2)|k2_zfmisc_1(X6,k2_matrix_1(X3))!=k1_tarski(X5)|~r2_hidden(X1,X6)|k1_matrix_1(X3)!=k3_finseq_1(X4)|~m2_finseq_1(X4,k3_finseq_2(X2))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X4)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_10058, plain, (v1_relat_1(esk11_1(k1_zfmisc_1(k2_matrix_1(np__0))))|~v1_matrix_1(np__0))).#trainneg
cnf(i_0_10068, plain, (v4_relat_1(esk16_1(k2_matrix_1(np__0)),np__0)|~v1_matrix_1(np__0))).#trainneg
cnf(i_0_1102, plain, (k2_tarski(k2_tarski(X1,k2_tarski(k1_tarski(esk7_3(X2,X3,X4)),k2_tarski(esk7_3(X2,X3,X4),esk6_3(X2,X3,X4)))),k1_tarski(X1))=X5|X4=k4_matrix_1(X2,X3)|v7_ordinal1(esk5_3(X2,X3,X4))|v1_xboole_0(X2)|k2_zfmisc_1(X6,k2_matrix_1(X3))!=k1_tarski(X5)|~r2_hidden(X1,X6)|k1_matrix_1(X3)!=k3_finseq_1(X4)|~m2_finseq_1(X4,k3_finseq_2(X2))|~m2_finseq_1(X3,k3_finseq_2(X2))|~v1_matrix_1(X4)|~v1_matrix_1(X3))).#trainneg
cnf(i_0_1103, plain, (k2_tarski(k2_tarski(X1,k2_tarski(k2_tarski(X2,X3),k1_tarski(X2))),k1_tarski(X1))=X4|k2_zfmisc_1(X5,k2_zfmisc_1(X6,X7))!=k1_tarski(X4)|~r2_hidden(X1,X5)|~r2_hidden(X3,X7)|~r2_hidden(X2,X6))).#trainneg
cnf(i_0_11436, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|k3_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)),np__1,np__1)!=X3|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(k1_tarski(X3)))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~m1_finseq_1(k4_matrix_1(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))),k3_finseq_2(k1_tarski(X3)))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))))).#trainneg
cnf(i_0_11631, plain, (X1=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,X4,esk9_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),esk9_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|k3_finseq_1(X1)!=k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(X1)!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~r2_hidden(esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~r2_hidden(esk9_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X1)|X4!=k4_matrix_1(X2,X1)|~m2_finseq_1(X4,k3_finseq_2(X2))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_11933, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|r2_hidden(k2_tarski(k1_tarski(np__1),k2_tarski(np__1,X4)),k2_matrix_1(k4_matrix_1(X5,k9_finseq_1(k9_finseq_1(X2)))))|v1_xboole_0(X5)|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~r2_hidden(X4,k1_tarski(np__1))|~v7_ordinal1(X4)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X5))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(X3,k3_finseq_2(X1))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_11907, plain, (r2_hidden(esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4),k1_tarski(np__1))|v1_xboole_0(X2)|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X1)|~v1_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_funct_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_relat_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~m1_subset_1(X3,X2)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~m2_finseq_1(X4,k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_11368, plain, (esk16_1(k2_matrix_1(np__0))=np__0|~v1_matrix_1(np__0))).#trainneg
cnf(i_0_1104, plain, (k2_tarski(k2_tarski(X1,k2_tarski(k2_tarski(X2,X3),k1_tarski(X3))),k1_tarski(X1))=X4|k2_zfmisc_1(X5,k2_zfmisc_1(X6,X7))!=k1_tarski(X4)|~r2_hidden(X1,X5)|~r2_hidden(X2,X7)|~r2_hidden(X3,X6))).#trainneg
cnf(i_0_10109, plain, (k1_zfmisc_1(k2_matrix_1(np__1))!=k2_tarski(np__1,np__2)|np__2!=np__1|~v1_xboole_0(k2_finseq_1(k1_matrix_1(np__1)))|~v1_finseq_1(np__1)|~v1_funct_1(np__1)|~v1_relat_1(np__1)|~v1_matrix_1(np__1))).#trainneg
cnf(i_0_1105, plain, (k2_tarski(k2_tarski(X1,k2_tarski(k1_tarski(X2),k2_tarski(X2,X3))),k1_tarski(X1))=X4|k2_zfmisc_1(X5,k2_zfmisc_1(X6,X7))!=k1_tarski(X4)|~r2_hidden(X1,X5)|~r2_hidden(X3,X7)|~r2_hidden(X2,X6))).#trainneg
cnf(i_0_11434, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|k3_matrix_1(k1_tarski(X3),k4_matrix_1(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))),np__1,np__1)!=X3|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(k1_tarski(X3)))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~m1_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(k1_tarski(X3))))).#trainneg
cnf(i_0_11965, plain, (esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4)=X5|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X1)|v1_xboole_0(X2)|v1_xboole_0(X6)|k2_matrlin(X6,k9_finseq_1(X7))!=k4_matrix_1(X6,k9_finseq_1(k9_finseq_1(X3)))|k4_finseq_1(k2_matrlin(X6,k9_finseq_1(X7)))!=k1_tarski(X5)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~m1_finseq_1(k9_finseq_1(X7),X6)|~v7_ordinal1(esk8_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X6))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_11437, plain, (k4_matrix_1(X1,k9_finseq_1(X2))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k9_finseq_1(k9_finseq_1(X3))!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(X2)))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k3_finseq_1(X2)!=np__1|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_11641, plain, (k3_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),esk8_3(X4,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X5),np__1)=k3_matrix_1(X1,X6,np__1,esk8_3(X4,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X5))|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X5|v1_xboole_0(X4)|v1_xboole_0(X1)|v1_xboole_0(X2)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X5)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X5)|X6!=k4_matrix_1(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~m2_finseq_1(X6,k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(X5,k3_finseq_2(X4))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X4))|~v1_matrix_1(X6)|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X5)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2)))).#trainneg
cnf(i_0_11440, plain, (k4_matrix_1(X1,k9_finseq_1(X2))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,X4,np__1,np__1)!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|X4!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(X2)))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4)|k3_finseq_1(X2)!=np__1|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_9900, plain, (v2_xxreal_0(np__0)|r1_xxreal_0(np__0,esk22_0)|r1_xxreal_0(esk22_0,np__1)|~v1_xxreal_0(esk22_0)|~v1_xxreal_0(np__1))).#trainneg
cnf(i_0_1112, plain, (k3_finseq_1(X1)=k1_card_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_10115, plain, (v1_xboole_0(k2_matrix_1(np__1))|~v1_xboole_0(k2_finseq_1(k1_matrix_1(np__1)))|~v1_finseq_1(np__1)|~v1_funct_1(np__1)|~v1_relat_1(np__1)|~v1_matrix_1(np__1))).#trainneg
cnf(i_0_1113, plain, (k3_finseq_1(X1)=k9_xtuple_0(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1114, plain, (k3_finseq_1(X1)=k1_finseq_1(X2)|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_12136, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X4)),np__1,np__1)!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X4)))!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X4)))))).#trainneg
cnf(i_0_12261, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|k3_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)),np__1,np__1)!=X3|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))),k3_finseq_2(k1_tarski(X3))))).#trainneg
cnf(i_0_12304, plain, (esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4)=X5|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X1)|v1_xboole_0(X2)|k1_tarski(np__1)!=k1_tarski(X5)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_12272, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X3)|v1_xboole_0(X1)|k3_matrix_1(X3,X5,np__1,esk8_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))!=k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)),np__1,esk8_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|X5!=k4_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~r2_hidden(esk8_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~m2_finseq_1(X5,k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(X5)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1)))).#trainneg
cnf(i_0_9889, plain, (r1_xxreal_0(esk11_1(k5_numbers),np__0)|r1_xxreal_0(np__0,esk11_1(k5_numbers)))).#trainneg
cnf(i_0_12362, plain, (k1_zfmisc_1(k2_matrix_1(np__1))!=k2_tarski(np__1,np__2)|np__2!=np__1|~v1_finseq_1(np__1)|~v1_funct_1(np__1)|~v1_relat_1(np__1)|~v1_matrix_1(np__1)|~v1_xboole_0(k1_matrix_1(np__1)))).#trainneg
cnf(i_0_12317, plain, (~v1_subset_1(np__0,k2_matrix_1(np__0))|~v1_matrix_1(np__0))).#trainneg
cnf(i_0_1118, plain, (r2_hidden(X1,k9_xtuple_0(X2))|v1_xboole_0(k2_matrix_1(X2))|~m1_subset_1(k2_tarski(k2_tarski(X1,X3),k1_tarski(X1)),k2_matrix_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_matrix_1(X2))).#trainneg
cnf(i_0_12167, plain, (X1=X2|k3_matrix_1(k1_tarski(X3),X4,esk9_3(k1_tarski(X3),X1,X2),esk8_3(k1_tarski(X3),X1,X2))!=X3|k3_finseq_1(X1)!=k3_finseq_1(X2)|k1_matrix_1(X1)!=k1_matrix_1(X2)|X4!=k4_matrix_1(k1_tarski(X3),X1)|~m2_finseq_1(X2,k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(X1,k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(X4,k3_finseq_2(k1_tarski(X3)))|~v1_matrix_1(X2)|~v1_matrix_1(X1)|~v1_matrix_1(X4)|~m1_finseq_1(X2,k3_finseq_2(k1_tarski(X3)))|~v7_ordinal1(esk9_3(k1_tarski(X3),X1,X2))|~v7_ordinal1(esk8_3(k1_tarski(X3),X1,X2)))).#trainneg
cnf(i_0_12145, plain, (X1=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_finseq_1(X1)!=k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(X1)!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k9_finseq_1(k9_finseq_1(X3))!=k4_matrix_1(X2,X1)|~r2_hidden(esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X1)|~m1_subset_1(esk9_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1)))).#trainneg
cnf(i_0_12312, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X3)|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k9_finseq_1(k9_finseq_1(X4))!=k4_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~r2_hidden(esk8_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_10026, plain, (r1_xxreal_0(esk11_1(k5_numbers),esk25_0)|r1_xxreal_0(esk25_0,esk11_1(k5_numbers)))).#trainneg
cnf(i_0_7206, plain, (r1_xxreal_0(np__0,X1)|r1_xxreal_0(X1,np__0)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_12523, plain, (v1_xboole_0(k2_matrix_1(np__1))|~v1_finseq_1(np__1)|~v1_funct_1(np__1)|~v1_relat_1(np__1)|~v1_matrix_1(np__1)|~v1_xboole_0(k1_matrix_1(np__1)))).#trainneg
cnf(i_0_12002, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(X3))|v1_xboole_0(X1)|k3_matrix_1(X1,k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))),np__1,np__1)!=k3_matrix_1(X1,k9_finseq_1(X3),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X3)))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X3)))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X3)))|k3_finseq_1(X3)!=np__1|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_12267, plain, (k4_matrix_1(X1,X2)=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X3)|v1_xboole_0(X1)|k3_matrix_1(X3,X5,esk9_3(X3,k4_matrix_1(X1,X2),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),np__1)!=k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)),esk9_3(X3,k4_matrix_1(X1,X2),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),np__1)|k3_finseq_1(k4_matrix_1(X1,X2))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,X2))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|X5!=k4_matrix_1(X3,k4_matrix_1(X1,X2))|~r2_hidden(esk9_3(X3,k4_matrix_1(X1,X2),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X1,X2),k3_finseq_2(X3))|~m2_finseq_1(X5,k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X1,X2))|~v1_matrix_1(X5)|k1_matrix_1(X2)!=np__1|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_12287, plain, (X1=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X4)),esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),esk9_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),esk9_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|k3_finseq_1(X1)!=k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(X1)!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X4)))!=k4_matrix_1(X2,X1)|~r2_hidden(esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~r2_hidden(esk9_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X1,k3_finseq_2(X2))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X1)|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X4))))|~v7_ordinal1(esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~v7_ordinal1(esk9_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X2)))).#trainneg
cnf(i_0_12264, plain, (k9_finseq_1(X1)=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,X4,esk9_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),esk9_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),np__1)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=np__1|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|X4!=k4_matrix_1(X2,k9_finseq_1(X1))|~r2_hidden(esk9_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k9_finseq_1(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_12162, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|v1_xboole_0(X4)|v1_xboole_0(X1)|k3_matrix_1(X4,X5,np__1,esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3))!=k3_matrix_1(X4,X3,esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3),np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|X5!=k4_matrix_1(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X4))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(X5)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1)))).#trainneg
cnf(i_0_12395, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~m1_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m1_finseq_1(k4_matrix_1(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))),k3_finseq_2(k1_tarski(X3)))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))))).#trainneg
cnf(i_0_12151, plain, (k4_matrix_1(X1,k9_finseq_1(X2))=k4_matrix_1(X1,k9_finseq_1(X3))|v1_xboole_0(X1)|k3_matrix_1(X1,k9_finseq_1(X2),np__1,np__1)!=k3_matrix_1(X1,k9_finseq_1(X3),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X3)))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X3)))|k3_finseq_1(X2)!=np__1|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m2_finseq_1(k9_finseq_1(X3),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X3)))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))|k3_finseq_1(X3)!=np__1|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_10232, plain, (r1_xxreal_0(esk11_1(k5_numbers),esk17_0)|r1_xxreal_0(esk17_0,esk11_1(k5_numbers)))).#trainneg
cnf(i_0_1128, plain, (r2_hidden(k2_tarski(k2_tarski(X1,X2),k1_tarski(X1)),k2_matrix_1(X3))|v1_xboole_0(X4)|v1_xboole_0(k2_matrix_1(X5))|X3!=k4_matrix_1(X4,X5)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(X5)|~m1_subset_1(k2_tarski(k1_tarski(X2),k2_tarski(X2,X1)),k2_matrix_1(X5)))).#trainneg
cnf(i_0_10121, plain, (v1_relat_1(esk16_1(k2_matrix_1(np__1)))|~v1_finseq_1(np__1)|~v1_funct_1(np__1)|~v1_relat_1(np__1)|~v1_matrix_1(np__1))).#trainneg
cnf(i_0_1129, plain, (r2_hidden(k2_tarski(k2_tarski(esk9_3(X1,X2,X3),esk8_3(X1,X2,X3)),k1_tarski(esk9_3(X1,X2,X3))),k2_matrix_1(X4))|v1_xboole_0(X5)|X2=X3|v1_xboole_0(X1)|X4!=k4_matrix_1(X5,X2)|~v7_ordinal1(esk9_3(X1,X2,X3))|~v7_ordinal1(esk8_3(X1,X2,X3))|~m2_finseq_1(X4,k3_finseq_2(X5))|~m2_finseq_1(X2,k3_finseq_2(X5))|~v1_matrix_1(X4)|~v1_matrix_1(X2)|k3_finseq_1(X2)!=k3_finseq_1(X3)|k1_matrix_1(X2)!=k1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_1130, plain, (r2_hidden(k2_tarski(k2_tarski(esk6_3(X1,X2,X3),esk7_3(X1,X2,X3)),k1_tarski(esk6_3(X1,X2,X3))),k2_matrix_1(X4))|v1_xboole_0(X5)|X3=k4_matrix_1(X1,X2)|v7_ordinal1(esk4_3(X1,X2,X3))|v1_xboole_0(X1)|X4!=k4_matrix_1(X5,X2)|~v7_ordinal1(esk6_3(X1,X2,X3))|~v7_ordinal1(esk7_3(X1,X2,X3))|~m2_finseq_1(X4,k3_finseq_2(X5))|~m2_finseq_1(X2,k3_finseq_2(X5))|~v1_matrix_1(X4)|~v1_matrix_1(X2)|k1_matrix_1(X2)!=k3_finseq_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_12452, plain, (k4_matrix_1(X1,k9_finseq_1(X2))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(X2))),np__1,np__1)!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k3_finseq_1(X2)!=np__1|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(X2),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(X2))),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(X2)))|~v1_matrix_1(k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(X2)))))).#trainneg
cnf(i_0_13069, plain, (k9_finseq_1(X1)=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k9_finseq_1(k9_finseq_1(X3))!=k4_matrix_1(X2,k9_finseq_1(X1))|~r2_hidden(esk9_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k9_finseq_1(X1)))).#trainneg
cnf(i_0_11049, plain, (r1_xxreal_0(esk25_0,esk11_1(k1_numbers))|r1_xxreal_0(esk11_1(k1_numbers),esk25_0))).#trainneg
cnf(i_0_1132, plain, (m1_subset_1(esk3_2(X1,k5_numbers),k1_numbers)|esk3_2(X1,k5_numbers)=X1|k1_tarski(X1)=k5_numbers)).#trainneg
cnf(i_0_13004, plain, (r1_xxreal_0(esk25_0,esk11_1(k5_numbers))|~v2_xxreal_0(esk11_1(k5_numbers))|~v1_xreal_0(esk11_1(k5_numbers)))).#trainneg
cnf(i_0_1131, plain, (r2_hidden(k2_tarski(k2_tarski(esk6_3(X1,X2,X3),esk7_3(X1,X2,X3)),k1_tarski(esk6_3(X1,X2,X3))),k2_matrix_1(X4))|v1_xboole_0(X5)|X3=k4_matrix_1(X1,X2)|v7_ordinal1(esk5_3(X1,X2,X3))|v1_xboole_0(X1)|X4!=k4_matrix_1(X5,X2)|~v7_ordinal1(esk6_3(X1,X2,X3))|~v7_ordinal1(esk7_3(X1,X2,X3))|~m2_finseq_1(X4,k3_finseq_2(X5))|~m2_finseq_1(X2,k3_finseq_2(X5))|~v1_matrix_1(X4)|~v1_matrix_1(X2)|k1_matrix_1(X2)!=k3_finseq_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X1))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X3))).#trainneg
cnf(i_0_1138, plain, (esk16_1(k2_zfmisc_1(X1,X2))=np__0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_12676, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|k3_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)),np__1,np__1)!=X3|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))))).#trainneg
cnf(i_0_13103, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~m1_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))),k3_finseq_2(k1_tarski(X3))))).#trainneg
cnf(i_0_13073, plain, (k9_finseq_1(k9_finseq_1(X1))=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,X4,np__1,np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=np__1|X4!=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X1)))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X1)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=np__1)).#trainneg
cnf(i_0_13050, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X3)|v1_xboole_0(X1)|k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X5)),esk8_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),np__1)!=k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)),np__1,esk8_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X5)))!=k4_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~r2_hidden(esk8_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X5))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X5)),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X5))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1)))).#trainneg
cnf(i_0_12891, plain, (k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),X3,X4)=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)),X4,X3)|v1_xboole_0(X1)|~r2_hidden(X3,k1_tarski(np__1))|~r2_hidden(X4,k1_tarski(np__1))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_12945, plain, (k2_matrlin(X1,k9_finseq_1(X2))=k9_finseq_1(k9_finseq_1(X4))|v1_xboole_0(X1)|v1_xboole_0(X3)|k3_matrix_1(X3,k2_matrlin(X1,k9_finseq_1(X2)),np__1,np__1)!=k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)),np__1,np__1)|k3_finseq_1(k2_matrlin(X1,k9_finseq_1(X2)))!=np__1|~m1_finseq_1(k9_finseq_1(X2),X1)|~m2_finseq_1(k2_matrlin(X1,k9_finseq_1(X2)),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_12677, plain, (esk9_3(X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),X4)=np__1|k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))=X4|v1_xboole_0(X2)|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k3_finseq_1(X4)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=k1_matrix_1(X4)|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_13075, plain, (k9_finseq_1(X1)=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,X4,np__1,np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|X4!=k4_matrix_1(X2,k9_finseq_1(X1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k9_finseq_1(X1))|~v1_matrix_1(X4)|k1_matrix_1(k9_finseq_1(X1))!=np__1|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=np__1)).#trainneg
cnf(i_0_12486, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_finseq_1(X4))|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_11110, plain, (r1_xxreal_0(esk17_0,esk11_1(k1_numbers))|r1_xxreal_0(esk11_1(k1_numbers),esk17_0))).#trainneg
cnf(i_0_1139, plain, (esk16_1(k2_zfmisc_1(X1,X2))=k1_card_1(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_13113, plain, (r1_xxreal_0(esk17_0,esk11_1(k5_numbers))|~v2_xxreal_0(esk11_1(k5_numbers))|~v1_xreal_0(esk17_0)|~v1_xreal_0(esk11_1(k5_numbers)))).#trainneg
cnf(i_0_1140, plain, (esk16_1(k2_zfmisc_1(X1,X2))=k9_xtuple_0(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1141, plain, (esk16_1(k2_zfmisc_1(X1,X2))=k1_finseq_1(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_13166, negated_conjecture, (~m1_finseq_1(k9_finseq_1(esk2_0),esk1_0)|~m1_subset_1(k9_finseq_1(esk2_0),k3_finseq_2(esk1_0)))).#trainneg
cnf(i_0_13177, plain, (k9_finseq_1(X1)=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k9_finseq_1(k9_finseq_1(X3))!=k4_matrix_1(X2,k9_finseq_1(X1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k9_finseq_1(X1))|~m1_subset_1(esk9_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1)))).#trainneg
cnf(i_0_13264, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~m1_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))))).#trainneg
cnf(i_0_13259, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|k3_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)),np__1,np__1)!=X3|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))))).#trainneg
cnf(i_0_13267, plain, (k9_finseq_1(k9_finseq_1(X1))=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,X4,np__1,np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=np__1|X4!=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X1)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X1)),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_1142, plain, (v1_xboole_0(esk16_1(k2_matrix_1(X1)))|~v1_xboole_0(k4_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_13183, plain, (r1_xxreal_0(esk25_0,esk11_1(k1_numbers))|~v2_xxreal_0(esk11_1(k1_numbers))|~v1_xreal_0(esk11_1(k1_numbers)))).#trainneg
cnf(i_0_13488, plain, (r1_xxreal_0(esk17_0,esk11_1(k1_numbers))|~v2_xxreal_0(esk11_1(k1_numbers))|~v1_xreal_0(esk17_0)|~v1_xreal_0(esk11_1(k1_numbers)))).#trainneg
cnf(i_0_1143, plain, (v1_xboole_0(esk16_1(k2_matrix_1(X1)))|~v1_xboole_0(k9_xtuple_0(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_matrix_1(X1))).#trainneg
cnf(i_0_1146, plain, (v4_relat_1(k4_matrix_1(X1,X2),k5_numbers)|v1_xboole_0(X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v1_matrix_1(X2))).#trainneg
cnf(i_0_13328, plain, (X1=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,X1,esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),esk9_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),esk9_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|k3_finseq_1(X1)!=k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(X1)!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~r2_hidden(esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~r2_hidden(esk9_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(X1,k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X1)|~v7_ordinal1(esk9_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~v7_ordinal1(esk8_3(X2,X1,k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_13284, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|k3_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)),np__1,esk8_3(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)))))!=X3|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X5)))!=k4_matrix_1(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X5))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X5)),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X5))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m1_finseq_1(k9_finseq_1(k9_finseq_1(X5)),k3_finseq_2(k1_tarski(X3)))|~v7_ordinal1(esk8_3(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))))).#trainneg
cnf(i_0_11242, plain, (r1_xxreal_0(esk11_1(k5_numbers),esk22_0)|r1_xxreal_0(esk22_0,esk11_1(k5_numbers)))).#trainneg
cnf(i_0_1148, plain, (v4_relat_1(k3_pre_poly(X1,X2),k5_numbers)|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_7311, plain, (v1_relat_1(np__1)|k3_finseq_2(X1)!=k2_tarski(np__1,np__2)|np__2!=np__1)).#trainneg
cnf(i_0_10070, plain, (v4_relat_1(esk11_1(k1_zfmisc_1(k2_matrix_1(np__0))),np__0)|~v1_matrix_1(np__0))).#trainneg
cnf(i_0_1156, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk26_1(X2))|v1_xboole_0(X2)|~m1_subset_1(X1,esk26_1(X2)))).#trainneg
cnf(i_0_1157, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk24_1(X2))|v1_xboole_0(X2)|~m1_subset_1(X1,esk24_1(X2)))).#trainneg
cnf(i_0_13174, plain, (k9_finseq_1(k9_finseq_1(X1))=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=np__1|k9_finseq_1(k9_finseq_1(X3))!=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X1)))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X1)),k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=np__1)).#trainneg
cnf(i_0_13275, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|k3_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X5)),esk8_3(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)))),np__1)!=X3|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X5)))!=k4_matrix_1(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X5))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X5)),k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X5))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m1_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(k1_tarski(X3)))|~v7_ordinal1(esk8_3(k1_tarski(X3),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4))))))).#trainneg
cnf(i_0_13424, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X3)|v1_xboole_0(X1)|k3_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),esk8_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),np__1)!=k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)),np__1,esk8_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~r2_hidden(esk8_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_13471, plain, (k9_finseq_1(X1)=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,X4,np__1,np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|X4!=k4_matrix_1(X2,k9_finseq_1(X1))|k1_matrix_1(k9_finseq_1(X1))!=np__1|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k9_finseq_1(X1))|~v1_matrix_1(X4))).#trainneg
cnf(i_0_13176, plain, (k9_finseq_1(X1)=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|k9_finseq_1(k9_finseq_1(X3))!=k4_matrix_1(X2,k9_finseq_1(X1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k9_finseq_1(X1))|k1_matrix_1(k9_finseq_1(X1))!=np__1|k3_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))!=np__1)).#trainneg
cnf(i_0_13430, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=X3|v1_xboole_0(X4)|v1_xboole_0(X1)|k3_matrix_1(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3),np__1)!=k3_matrix_1(X4,X3,esk8_3(X4,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),X3),np__1)|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(X3)|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X4))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X4))|~v1_matrix_1(X3)|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_10432, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk11_1(k1_zfmisc_1(X4)))|~v1_xboole_0(X3)|~v1_xboole_0(X2)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_11388, plain, (r1_xxreal_0(esk11_1(k5_numbers),np__2)|r1_xxreal_0(np__2,esk11_1(k5_numbers)))).#trainneg
cnf(i_0_1158, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk20_1(X2))|v1_xboole_0(X2)|~m1_subset_1(X1,esk20_1(X2)))).#trainneg
cnf(i_0_7402, plain, (v1_funct_1(np__1)|k3_finseq_2(X1)!=k2_tarski(np__1,np__2)|np__2!=np__1)).#trainneg
cnf(i_0_10071, plain, (v5_relat_1(k2_matrix_1(np__0),k2_finseq_1(k1_matrix_1(np__0)))|~v1_matrix_1(np__0))).#trainneg
cnf(i_0_12026, plain, (esk11_1(k1_zfmisc_1(k2_matrix_1(np__0)))=np__0|~v1_matrix_1(np__0))).#trainneg
cnf(i_0_1161, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk13_1(X2))|v1_xboole_0(X2)|~m1_subset_1(X1,esk13_1(X2)))).#trainneg
cnf(i_0_13295, plain, (X1=X2|k3_matrix_1(k1_tarski(X3),k9_finseq_1(k9_finseq_1(X4)),esk8_3(k1_tarski(X3),X1,X2),esk9_3(k1_tarski(X3),X1,X2))!=X3|k9_finseq_1(k9_finseq_1(X4))!=k4_matrix_1(k1_tarski(X3),X1)|k3_finseq_1(X1)!=k3_finseq_1(X2)|k1_matrix_1(X1)!=k1_matrix_1(X2)|~m1_finseq_1(X2,k3_finseq_2(k1_tarski(X3)))|~v7_ordinal1(esk9_3(k1_tarski(X3),X1,X2))|~v7_ordinal1(esk8_3(k1_tarski(X3),X1,X2))|~m2_finseq_1(X1,k3_finseq_2(k1_tarski(X3)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(k1_tarski(X3)))|~v1_matrix_1(X2)|~v1_matrix_1(X1)|~r2_hidden(esk8_3(k1_tarski(X3),X1,X2),k1_tarski(np__1))|~r2_hidden(esk9_3(k1_tarski(X3),X1,X2),k1_tarski(np__1))|~m1_subset_1(X4,k1_tarski(X3)))).#trainneg
cnf(i_0_13337, plain, (k9_finseq_1(X1)=k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X2)|k3_matrix_1(X2,X4,esk9_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),np__1)!=k3_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)),np__1,esk9_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|k1_matrix_1(k9_finseq_1(X1))!=k1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|X4!=k4_matrix_1(X2,k9_finseq_1(X1))|~r2_hidden(esk9_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X2))|~m2_finseq_1(k9_finseq_1(X1),k3_finseq_2(X2))|~m2_finseq_1(X4,k3_finseq_2(X2))|~v1_matrix_1(k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(k9_finseq_1(X1))|~v1_matrix_1(X4)|~v7_ordinal1(esk9_3(X2,k9_finseq_1(X1),k4_matrix_1(X2,k9_finseq_1(k9_finseq_1(X3)))))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_13343, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))|v1_xboole_0(X1)|k3_matrix_1(X1,X4,esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),np__1)!=k3_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)),np__1,esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|X4!=k4_matrix_1(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~r2_hidden(esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X1))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(X1))|~m2_finseq_1(X4,k3_finseq_2(X1))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3))))|~v1_matrix_1(X4)|~v7_ordinal1(esk9_3(X1,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X3)))))|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_13332, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X1)|v1_xboole_0(X3)|k3_matrix_1(X3,X5,esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),np__1)!=k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)),np__1,esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|X5!=k4_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(X1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m2_finseq_1(X5,k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(X5)|~r2_hidden(esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),k1_tarski(np__1))|~v7_ordinal1(esk9_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_13426, plain, (k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2)))=k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))|v1_xboole_0(X3)|v1_xboole_0(X1)|k3_matrix_1(X3,X5,np__1,esk8_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))!=k3_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)),np__1,esk8_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))))|k3_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k3_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|k1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))!=k1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|X5!=k4_matrix_1(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~r2_hidden(esk8_3(X3,k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4)))),k1_tarski(np__1))|~m2_finseq_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))),k3_finseq_2(X3))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X4)),k3_finseq_2(X3))|~m2_finseq_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))),k3_finseq_2(X3))|~m2_finseq_1(X5,k3_finseq_2(X3))|~v1_matrix_1(k4_matrix_1(X3,k9_finseq_1(k9_finseq_1(X4))))|~v1_matrix_1(k4_matrix_1(X1,k9_finseq_1(k9_finseq_1(X2))))|~v1_matrix_1(X5)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_13360, plain, (k3_matrix_1(k1_tarski(X1),k9_finseq_1(k9_finseq_1(X2)),X4,X3)=X1|~m1_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(k1_tarski(X1)))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~r2_hidden(X3,k1_tarski(np__1))|~r2_hidden(X4,k1_tarski(np__1))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X2)),k3_finseq_2(k1_tarski(X1)))|~m1_subset_1(X2,k1_tarski(X1)))).#trainneg
cnf(i_0_13421, plain, (k2_matrlin(X1,k9_finseq_1(X2))=k9_finseq_1(k9_finseq_1(X3))|v1_xboole_0(X1)|k3_matrix_1(k1_tarski(X4),k9_finseq_1(k9_finseq_1(X3)),np__1,np__1)!=X4|k3_finseq_1(k2_matrlin(X1,k9_finseq_1(X2)))!=np__1|~m1_finseq_1(k9_finseq_1(X2),X1)|~m2_finseq_1(k2_matrlin(X1,k9_finseq_1(X2)),k3_finseq_2(k1_tarski(X4)))|~m2_finseq_1(k9_finseq_1(k9_finseq_1(X3)),k3_finseq_2(k1_tarski(X4)))|~m1_subset_1(X3,k1_tarski(X4))|~m1_finseq_1(k2_matrlin(X1,k9_finseq_1(X2)),k3_finseq_2(k1_tarski(X4)))|~v1_matrix_1(k2_matrlin(X1,k9_finseq_1(X2))))).#trainneg
cnf(i_0_13005, plain, (r1_xxreal_0(esk11_1(k1_numbers),np__0)|r1_xxreal_0(np__0,esk11_1(k1_numbers)))).#trainneg
cnf(i_0_6889, plain, (m2_finseq_1(np__1,X1)|k3_finseq_2(X1)!=k2_tarski(np__1,np__2)|np__2!=np__1)).#trainneg
cnf(i_0_10103, plain, (v1_xboole_0(k2_finseq_1(k1_matrix_1(np__1)))|~v1_xboole_0(k2_matrix_1(np__1))|~v1_finseq_1(np__1)|~v1_funct_1(np__1)|~v1_relat_1(np__1)|~v1_matrix_1(np__1))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 156
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 278
# Removed in clause preprocessing      : 15
# Initial clauses in saturation        : 263
# Processed clauses                    : 2636
# ...of these trivial                  : 49
# ...subsumed                          : 1080
# ...remaining for further processing  : 1507
# Other redundant clauses eliminated   : 4
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 60
# Backward-rewritten                   : 139
# Generated clauses                    : 13998
# ...of the previous two non-trivial   : 13003
# Contextual simplify-reflections      : 361
# Paramodulations                      : 13922
# Factorizations                       : 0
# Equation resolutions                 : 76
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1307
#    Positive orientable unit clauses  : 127
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 32
#    Non-unit-clauses                  : 1147
# Current number of unprocessed clauses: 9987
# ...number of literals in the above   : 96554
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1709
# Proof object given clauses           : 145
# Proof search given clauses           : 1507
# Clause-clause subsumption calls (NU) : 1550788
# Rec. Clause-clause subsumption calls : 87660
# Non-unit clause-clause subsumptions  : 1371
# Unit Clause-clause subsumption calls : 21485
# Rewrite failures with RHS unbound    : 36
# BW rewrite match attempts            : 94
# BW rewrite match successes           : 63
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 532288

# -------------------------------------------------
# User time                : 36.119 s
# System time              : 1.502 s
# Total time               : 37.620 s
# Maximum resident set size: 471744 pages
36.11user 1.59system 0:38.07elapsed 99%CPU (0avgtext+0avgdata 471744maxresident)k
0inputs+944outputs (0major+443081minor)pagefaults 0swaps
