# ENIGMA: Connected to the TF server '127.0.0.1:8888'.
# Preprocessing time       : 0.017 s
#TF#SERVER: Sending data to the TF server (conj=5, context=104, query=162, total=271).
#TF#SERVER: Sent 78079 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1068 bytes expected).
#TF#SERVER: Received 1068 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=195, query=512, total=712).
#TF#SERVER: Sent 350840 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2832 bytes expected).
#TF#SERVER: Received 2832 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=195, query=13, total=213).
#TF#SERVER: Sent 60161 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (836 bytes expected).
#TF#SERVER: Received 836 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=245, query=512, total=762).
#TF#SERVER: Sent 433979 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3032 bytes expected).
#TF#SERVER: Received 3032 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=245, query=17, total=267).
#TF#SERVER: Sent 84064 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1052 bytes expected).
#TF#SERVER: Received 1052 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=288, query=512, total=805).
#TF#SERVER: Sent 460573 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3204 bytes expected).
#TF#SERVER: Received 3204 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=288, query=91, total=384).
#TF#SERVER: Sent 162373 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1520 bytes expected).
#TF#SERVER: Received 1520 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=324, query=512, total=841).
#TF#SERVER: Sent 441626 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3348 bytes expected).
#TF#SERVER: Received 3348 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=324, query=56, total=385).
#TF#SERVER: Sent 147078 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1524 bytes expected).
#TF#SERVER: Received 1524 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=358, query=512, total=875).
#TF#SERVER: Sent 419024 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3484 bytes expected).
#TF#SERVER: Received 3484 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=358, query=35, total=398).
#TF#SERVER: Sent 129219 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1576 bytes expected).
#TF#SERVER: Received 1576 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=384, query=512, total=901).
#TF#SERVER: Sent 504032 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3588 bytes expected).
#TF#SERVER: Received 3588 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=405, query=512, total=922).
#TF#SERVER: Sent 499226 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3672 bytes expected).
#TF#SERVER: Received 3672 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=405, query=15, total=425).
#TF#SERVER: Sent 125766 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1684 bytes expected).
#TF#SERVER: Received 1684 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=418, query=512, total=935).
#TF#SERVER: Sent 514510 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3724 bytes expected).
#TF#SERVER: Received 3724 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=418, query=86, total=509).
#TF#SERVER: Sent 184597 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2020 bytes expected).
#TF#SERVER: Received 2020 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=430, query=512, total=947).
#TF#SERVER: Sent 553279 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3772 bytes expected).
#TF#SERVER: Received 3772 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=430, query=6, total=441).
#TF#SERVER: Sent 133560 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1748 bytes expected).
#TF#SERVER: Received 1748 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=439, query=512, total=956).
#TF#SERVER: Sent 501244 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3808 bytes expected).
#TF#SERVER: Received 3808 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=439, query=3, total=447).
#TF#SERVER: Sent 134504 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1772 bytes expected).
#TF#SERVER: Received 1772 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=455, query=512, total=972).
#TF#SERVER: Sent 549711 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3872 bytes expected).
#TF#SERVER: Received 3872 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=455, query=105, total=565).
#TF#SERVER: Sent 239096 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2244 bytes expected).
#TF#SERVER: Received 2244 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=464, query=512, total=981).
#TF#SERVER: Sent 519330 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3908 bytes expected).
#TF#SERVER: Received 3908 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=464, query=18, total=487).
#TF#SERVER: Sent 158603 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (1932 bytes expected).
#TF#SERVER: Received 1932 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=491, query=512, total=1008).
#TF#SERVER: Sent 449289 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4016 bytes expected).
#TF#SERVER: Received 4016 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=491, query=53, total=549).
#TF#SERVER: Sent 189870 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2180 bytes expected).
#TF#SERVER: Received 2180 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=499, query=512, total=1016).
#TF#SERVER: Sent 563554 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4048 bytes expected).
#TF#SERVER: Received 4048 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=513, query=512, total=1030).
#TF#SERVER: Sent 591498 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4104 bytes expected).
#TF#SERVER: Received 4104 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=513, query=127, total=645).
#TF#SERVER: Sent 262135 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2564 bytes expected).
#TF#SERVER: Received 2564 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=517, query=512, total=1034).
#TF#SERVER: Sent 563117 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4120 bytes expected).
#TF#SERVER: Received 4120 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=517, query=252, total=774).
#TF#SERVER: Sent 421586 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3080 bytes expected).
#TF#SERVER: Received 3080 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=532, query=512, total=1049).
#TF#SERVER: Sent 612732 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4180 bytes expected).
#TF#SERVER: Received 4180 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=532, query=100, total=637).
#TF#SERVER: Sent 268870 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2532 bytes expected).
#TF#SERVER: Received 2532 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=538, query=512, total=1055).
#TF#SERVER: Sent 683971 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4204 bytes expected).
#TF#SERVER: Received 4204 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=538, query=163, total=706).
#TF#SERVER: Sent 337606 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2808 bytes expected).
#TF#SERVER: Received 2808 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=548, query=512, total=1065).
#TF#SERVER: Sent 606513 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4244 bytes expected).
#TF#SERVER: Received 4244 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=548, query=101, total=654).
#TF#SERVER: Sent 253665 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2600 bytes expected).
#TF#SERVER: Received 2600 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=554, query=512, total=1071).
#TF#SERVER: Sent 685405 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4268 bytes expected).
#TF#SERVER: Received 4268 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=554, query=87, total=646).
#TF#SERVER: Sent 262145 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2568 bytes expected).
#TF#SERVER: Received 2568 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=556, query=512, total=1073).
#TF#SERVER: Sent 743812 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4276 bytes expected).
#TF#SERVER: Received 4276 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=556, query=58, total=619).
#TF#SERVER: Sent 201401 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2460 bytes expected).
#TF#SERVER: Received 2460 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=565, query=512, total=1082).
#TF#SERVER: Sent 621448 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4312 bytes expected).
#TF#SERVER: Received 4312 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=565, query=29, total=599).
#TF#SERVER: Sent 202669 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2380 bytes expected).
#TF#SERVER: Received 2380 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=569, query=512, total=1086).
#TF#SERVER: Sent 688745 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4328 bytes expected).
#TF#SERVER: Received 4328 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=569, query=47, total=621).
#TF#SERVER: Sent 225459 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2468 bytes expected).
#TF#SERVER: Received 2468 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=575, query=512, total=1092).
#TF#SERVER: Sent 722272 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4352 bytes expected).
#TF#SERVER: Received 4352 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=575, query=7, total=587).
#TF#SERVER: Sent 186815 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2332 bytes expected).
#TF#SERVER: Received 2332 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=583, query=512, total=1100).
#TF#SERVER: Sent 704619 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4384 bytes expected).
#TF#SERVER: Received 4384 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=583, query=165, total=753).
#TF#SERVER: Sent 365892 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2996 bytes expected).
#TF#SERVER: Received 2996 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=598, query=512, total=1115).
#TF#SERVER: Sent 628601 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4444 bytes expected).
#TF#SERVER: Received 4444 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=598, query=94, total=697).
#TF#SERVER: Sent 268726 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2772 bytes expected).
#TF#SERVER: Received 2772 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=604, query=512, total=1121).
#TF#SERVER: Sent 771854 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4468 bytes expected).
#TF#SERVER: Received 4468 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=604, query=107, total=716).
#TF#SERVER: Sent 307643 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2848 bytes expected).
#TF#SERVER: Received 2848 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=614, query=512, total=1131).
#TF#SERVER: Sent 785731 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4508 bytes expected).
#TF#SERVER: Received 4508 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=614, query=426, total=1045).
#TF#SERVER: Sent 686881 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4164 bytes expected).
#TF#SERVER: Received 4164 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=619, query=512, total=1136).
#TF#SERVER: Sent 657029 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4528 bytes expected).
#TF#SERVER: Received 4528 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=619, query=1, total=625).
#TF#SERVER: Sent 196238 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2484 bytes expected).
#TF#SERVER: Received 2484 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=622, query=512, total=1139).
#TF#SERVER: Sent 780491 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4540 bytes expected).
#TF#SERVER: Received 4540 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=622, query=158, total=785).
#TF#SERVER: Sent 367759 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3124 bytes expected).
#TF#SERVER: Received 3124 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=632, query=512, total=1149).
#TF#SERVER: Sent 585526 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4580 bytes expected).
#TF#SERVER: Received 4580 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=632, query=90, total=727).
#TF#SERVER: Sent 282290 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2892 bytes expected).
#TF#SERVER: Received 2892 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=643, query=512, total=1160).
#TF#SERVER: Sent 613005 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4624 bytes expected).
#TF#SERVER: Received 4624 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=643, query=54, total=702).
#TF#SERVER: Sent 247532 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2792 bytes expected).
#TF#SERVER: Received 2792 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=652, query=512, total=1169).
#TF#SERVER: Sent 707428 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4660 bytes expected).
#TF#SERVER: Received 4660 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=652, query=117, total=774).
#TF#SERVER: Sent 286280 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3080 bytes expected).
#TF#SERVER: Received 3080 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=664, query=512, total=1181).
#TF#SERVER: Sent 665700 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4708 bytes expected).
#TF#SERVER: Received 4708 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=664, query=20, total=689).
#TF#SERVER: Sent 221865 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2740 bytes expected).
#TF#SERVER: Received 2740 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=667, query=512, total=1184).
#TF#SERVER: Sent 797001 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4720 bytes expected).
#TF#SERVER: Received 4720 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=667, query=218, total=890).
#TF#SERVER: Sent 456168 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3544 bytes expected).
#TF#SERVER: Received 3544 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=668, query=512, total=1185).
#TF#SERVER: Sent 805348 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4724 bytes expected).
#TF#SERVER: Received 4724 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=668, query=218, total=891).
#TF#SERVER: Sent 450866 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3548 bytes expected).
#TF#SERVER: Received 3548 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=681, query=512, total=1198).
#TF#SERVER: Sent 629159 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4776 bytes expected).
#TF#SERVER: Received 4776 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=681, query=49, total=735).
#TF#SERVER: Sent 254826 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2924 bytes expected).
#TF#SERVER: Received 2924 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=690, query=512, total=1207).
#TF#SERVER: Sent 646080 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4812 bytes expected).
#TF#SERVER: Received 4812 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=690, query=86, total=781).
#TF#SERVER: Sent 304642 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3108 bytes expected).
#TF#SERVER: Received 3108 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=699, query=512, total=1216).
#TF#SERVER: Sent 772666 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4848 bytes expected).
#TF#SERVER: Received 4848 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=699, query=109, total=813).
#TF#SERVER: Sent 341364 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3236 bytes expected).
#TF#SERVER: Received 3236 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=703, query=512, total=1220).
#TF#SERVER: Sent 787081 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4864 bytes expected).
#TF#SERVER: Received 4864 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=703, query=348, total=1056).
#TF#SERVER: Sent 621205 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4208 bytes expected).
#TF#SERVER: Received 4208 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=711, query=512, total=1228).
#TF#SERVER: Sent 793929 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4896 bytes expected).
#TF#SERVER: Received 4896 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=711, query=135, total=851).
#TF#SERVER: Sent 348570 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3388 bytes expected).
#TF#SERVER: Received 3388 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=721, query=512, total=1238).
#TF#SERVER: Sent 664553 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4936 bytes expected).
#TF#SERVER: Received 4936 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=721, query=50, total=776).
#TF#SERVER: Sent 282095 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3088 bytes expected).
#TF#SERVER: Received 3088 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=724, query=512, total=1241).
#TF#SERVER: Sent 805461 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4948 bytes expected).
#TF#SERVER: Received 4948 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=724, query=61, total=790).
#TF#SERVER: Sent 299358 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3144 bytes expected).
#TF#SERVER: Received 3144 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=728, query=512, total=1245).
#TF#SERVER: Sent 790559 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4964 bytes expected).
#TF#SERVER: Received 4964 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=728, query=384, total=1117).
#TF#SERVER: Sent 665814 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4452 bytes expected).
#TF#SERVER: Received 4452 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=730, query=512, total=1247).
#TF#SERVER: Sent 871319 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4972 bytes expected).
#TF#SERVER: Received 4972 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=730, query=184, total=919).
#TF#SERVER: Sent 408454 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3660 bytes expected).
#TF#SERVER: Received 3660 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=731, query=512, total=1248).
#TF#SERVER: Sent 841094 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4976 bytes expected).
#TF#SERVER: Received 4976 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=731, query=348, total=1084).
#TF#SERVER: Sent 628878 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4320 bytes expected).
#TF#SERVER: Received 4320 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=733, query=512, total=1250).
#TF#SERVER: Sent 771529 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4984 bytes expected).
#TF#SERVER: Received 4984 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=733, query=512, total=1250).
#TF#SERVER: Sent 839123 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4984 bytes expected).
#TF#SERVER: Received 4984 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=733, query=217, total=955).
#TF#SERVER: Sent 479868 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3804 bytes expected).
#TF#SERVER: Received 3804 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=738, query=512, total=1255).
#TF#SERVER: Sent 796830 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5004 bytes expected).
#TF#SERVER: Received 5004 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=738, query=59, total=802).
#TF#SERVER: Sent 302343 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3192 bytes expected).
#TF#SERVER: Received 3192 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=748, query=512, total=1265).
#TF#SERVER: Sent 674294 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5044 bytes expected).
#TF#SERVER: Received 5044 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=748, query=92, total=845).
#TF#SERVER: Sent 330523 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3364 bytes expected).
#TF#SERVER: Received 3364 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=765, query=512, total=1282).
#TF#SERVER: Sent 636320 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5112 bytes expected).
#TF#SERVER: Received 5112 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=765, query=132, total=902).
#TF#SERVER: Sent 386389 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3592 bytes expected).
#TF#SERVER: Received 3592 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=780, query=512, total=1297).
#TF#SERVER: Sent 601049 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5172 bytes expected).
#TF#SERVER: Received 5172 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=780, query=101, total=886).
#TF#SERVER: Sent 357165 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3528 bytes expected).
#TF#SERVER: Received 3528 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=797, query=512, total=1314).
#TF#SERVER: Sent 705585 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5240 bytes expected).
#TF#SERVER: Received 5240 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=797, query=199, total=1001).
#TF#SERVER: Sent 436007 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3988 bytes expected).
#TF#SERVER: Received 3988 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=820, query=512, total=1337).
#TF#SERVER: Sent 699813 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5332 bytes expected).
#TF#SERVER: Received 5332 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=820, query=49, total=874).
#TF#SERVER: Sent 300131 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3480 bytes expected).
#TF#SERVER: Received 3480 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=847, query=512, total=1364).
#TF#SERVER: Sent 593836 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5440 bytes expected).
#TF#SERVER: Received 5440 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=847, query=121, total=973).
#TF#SERVER: Sent 375551 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3876 bytes expected).
#TF#SERVER: Received 3876 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=864, query=512, total=1381).
#TF#SERVER: Sent 718819 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5508 bytes expected).
#TF#SERVER: Received 5508 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=864, query=74, total=943).
#TF#SERVER: Sent 354504 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3756 bytes expected).
#TF#SERVER: Received 3756 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=875, query=512, total=1392).
#TF#SERVER: Sent 770345 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5552 bytes expected).
#TF#SERVER: Received 5552 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=875, query=371, total=1251).
#TF#SERVER: Sent 634464 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4988 bytes expected).
#TF#SERVER: Received 4988 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=879, query=512, total=1396).
#TF#SERVER: Sent 771820 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5568 bytes expected).
#TF#SERVER: Received 5568 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=879, query=88, total=972).
#TF#SERVER: Sent 358146 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3872 bytes expected).
#TF#SERVER: Received 3872 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=902, query=512, total=1419).
#TF#SERVER: Sent 644803 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5660 bytes expected).
#TF#SERVER: Received 5660 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=902, query=228, total=1135).
#TF#SERVER: Sent 557910 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4524 bytes expected).
#TF#SERVER: Received 4524 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=915, query=512, total=1432).
#TF#SERVER: Sent 752592 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5712 bytes expected).
#TF#SERVER: Received 5712 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=915, query=346, total=1266).
#TF#SERVER: Sent 619131 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5048 bytes expected).
#TF#SERVER: Received 5048 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=923, query=512, total=1440).
#TF#SERVER: Sent 900710 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5744 bytes expected).
#TF#SERVER: Received 5744 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=931, query=512, total=1448).
#TF#SERVER: Sent 873333 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5776 bytes expected).
#TF#SERVER: Received 5776 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=931, query=243, total=1179).
#TF#SERVER: Sent 545674 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4700 bytes expected).
#TF#SERVER: Received 4700 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=935, query=512, total=1452).
#TF#SERVER: Sent 785136 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5792 bytes expected).
#TF#SERVER: Received 5792 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=935, query=42, total=982).
#TF#SERVER: Sent 321487 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3912 bytes expected).
#TF#SERVER: Received 3912 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=940, query=512, total=1457).
#TF#SERVER: Sent 837522 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5812 bytes expected).
#TF#SERVER: Received 5812 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=940, query=97, total=1042).
#TF#SERVER: Sent 398896 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4152 bytes expected).
#TF#SERVER: Received 4152 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=945, query=512, total=1462).
#TF#SERVER: Sent 712966 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5832 bytes expected).
#TF#SERVER: Received 5832 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=945, query=75, total=1025).
#TF#SERVER: Sent 369386 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4084 bytes expected).
#TF#SERVER: Received 4084 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=949, query=512, total=1466).
#TF#SERVER: Sent 703291 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5848 bytes expected).
#TF#SERVER: Received 5848 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=949, query=146, total=1100).
#TF#SERVER: Sent 454589 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4384 bytes expected).
#TF#SERVER: Received 4384 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=957, query=512, total=1474).
#TF#SERVER: Sent 714879 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5880 bytes expected).
#TF#SERVER: Received 5880 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=957, query=125, total=1087).
#TF#SERVER: Sent 408910 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4332 bytes expected).
#TF#SERVER: Received 4332 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=964, query=512, total=1481).
#TF#SERVER: Sent 780752 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5908 bytes expected).
#TF#SERVER: Received 5908 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=964, query=96, total=1065).
#TF#SERVER: Sent 416637 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4244 bytes expected).
#TF#SERVER: Received 4244 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=968, query=512, total=1485).
#TF#SERVER: Sent 908530 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5924 bytes expected).
#TF#SERVER: Received 5924 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=968, query=61, total=1034).
#TF#SERVER: Sent 386041 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4120 bytes expected).
#TF#SERVER: Received 4120 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=972, query=512, total=1489).
#TF#SERVER: Sent 915654 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5940 bytes expected).
#TF#SERVER: Received 5940 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=972, query=25, total=1002).
#TF#SERVER: Sent 335548 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3992 bytes expected).
#TF#SERVER: Received 3992 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=977, query=512, total=1494).
#TF#SERVER: Sent 783267 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5960 bytes expected).
#TF#SERVER: Received 5960 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=977, query=131, total=1113).
#TF#SERVER: Sent 452580 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4436 bytes expected).
#TF#SERVER: Received 4436 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=981, query=512, total=1498).
#TF#SERVER: Sent 937234 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5976 bytes expected).
#TF#SERVER: Received 5976 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=981, query=88, total=1074).
#TF#SERVER: Sent 411725 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4280 bytes expected).
#TF#SERVER: Received 4280 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=985, query=512, total=1502).
#TF#SERVER: Sent 887796 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5992 bytes expected).
#TF#SERVER: Received 5992 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=985, query=512, total=1502).
#TF#SERVER: Sent 928778 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5992 bytes expected).
#TF#SERVER: Received 5992 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=985, query=324, total=1314).
#TF#SERVER: Sent 709810 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5240 bytes expected).
#TF#SERVER: Received 5240 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=989, query=512, total=1506).
#TF#SERVER: Sent 891385 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6008 bytes expected).
#TF#SERVER: Received 6008 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=989, query=220, total=1214).
#TF#SERVER: Sent 578358 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4840 bytes expected).
#TF#SERVER: Received 4840 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=994, query=512, total=1511).
#TF#SERVER: Sent 801021 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6028 bytes expected).
#TF#SERVER: Received 6028 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=994, query=364, total=1363).
#TF#SERVER: Sent 619566 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5436 bytes expected).
#TF#SERVER: Received 5436 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1001, query=512, total=1518).
#TF#SERVER: Sent 662312 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6056 bytes expected).
#TF#SERVER: Received 6056 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1001, query=27, total=1033).
#TF#SERVER: Sent 338029 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4116 bytes expected).
#TF#SERVER: Received 4116 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1006, query=512, total=1523).
#TF#SERVER: Sent 728313 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6076 bytes expected).
#TF#SERVER: Received 6076 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1006, query=36, total=1047).
#TF#SERVER: Sent 358335 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4172 bytes expected).
#TF#SERVER: Received 4172 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1011, query=512, total=1528).
#TF#SERVER: Sent 701412 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6096 bytes expected).
#TF#SERVER: Received 6096 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1011, query=96, total=1112).
#TF#SERVER: Sent 401476 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4432 bytes expected).
#TF#SERVER: Received 4432 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1017, query=512, total=1534).
#TF#SERVER: Sent 740527 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6120 bytes expected).
#TF#SERVER: Received 6120 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1017, query=54, total=1076).
#TF#SERVER: Sent 386432 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4288 bytes expected).
#TF#SERVER: Received 4288 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1021, query=512, total=1538).
#TF#SERVER: Sent 1026887 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6136 bytes expected).
#TF#SERVER: Received 6136 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1021, query=48, total=1074).
#TF#SERVER: Sent 390749 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4280 bytes expected).
#TF#SERVER: Received 4280 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 622805 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=109, total=1138).
#TF#SERVER: Sent 463652 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4536 bytes expected).
#TF#SERVER: Received 4536 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 652563 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=47, total=1076).
#TF#SERVER: Sent 362600 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4288 bytes expected).
#TF#SERVER: Received 4288 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 827641 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=90, total=1119).
#TF#SERVER: Sent 429574 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4460 bytes expected).
#TF#SERVER: Received 4460 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 974636 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=163, total=1192).
#TF#SERVER: Sent 554073 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4752 bytes expected).
#TF#SERVER: Received 4752 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 999097 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1101010 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=465, total=1494).
#TF#SERVER: Sent 1011611 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5960 bytes expected).
#TF#SERVER: Received 5960 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 880401 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=24, total=1053).
#TF#SERVER: Sent 354251 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4196 bytes expected).
#TF#SERVER: Received 4196 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 945136 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1082372 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1086240 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=72, total=1101).
#TF#SERVER: Sent 427166 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4388 bytes expected).
#TF#SERVER: Received 4388 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 772515 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=220, total=1249).
#TF#SERVER: Sent 522536 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4980 bytes expected).
#TF#SERVER: Received 4980 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 853936 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=214, total=1243).
#TF#SERVER: Sent 583676 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4956 bytes expected).
#TF#SERVER: Received 4956 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 903865 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=19, total=1048).
#TF#SERVER: Sent 358871 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4176 bytes expected).
#TF#SERVER: Received 4176 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1022419 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1142511 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=188, total=1217).
#TF#SERVER: Sent 598872 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4852 bytes expected).
#TF#SERVER: Received 4852 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1076735 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1114063 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=476, total=1505).
#TF#SERVER: Sent 1037085 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6004 bytes expected).
#TF#SERVER: Received 6004 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1145344 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=285, total=1314).
#TF#SERVER: Sent 753276 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5240 bytes expected).
#TF#SERVER: Received 5240 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 742218 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=313, total=1342).
#TF#SERVER: Sent 724033 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5352 bytes expected).
#TF#SERVER: Received 5352 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 856240 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=303, total=1332).
#TF#SERVER: Sent 782577 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5312 bytes expected).
#TF#SERVER: Received 5312 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 973039 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1096060 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1104050 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=470, total=1499).
#TF#SERVER: Sent 1019931 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5980 bytes expected).
#TF#SERVER: Received 5980 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 840460 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=139, total=1168).
#TF#SERVER: Sent 475685 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4656 bytes expected).
#TF#SERVER: Received 4656 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 990162 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=196, total=1225).
#TF#SERVER: Sent 580392 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4884 bytes expected).
#TF#SERVER: Received 4884 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 881515 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 991710 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 995644 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=374, total=1403).
#TF#SERVER: Sent 815773 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5596 bytes expected).
#TF#SERVER: Received 5596 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1150511 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=342, total=1371).
#TF#SERVER: Sent 823997 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5468 bytes expected).
#TF#SERVER: Received 5468 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 957140 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1119308 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=125, total=1154).
#TF#SERVER: Sent 499995 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4600 bytes expected).
#TF#SERVER: Received 4600 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1156074 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=357, total=1386).
#TF#SERVER: Sent 849783 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5528 bytes expected).
#TF#SERVER: Received 5528 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 930648 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1136567 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=189, total=1218).
#TF#SERVER: Sent 589882 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4856 bytes expected).
#TF#SERVER: Received 4856 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1033171 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=272, total=1301).
#TF#SERVER: Sent 667699 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5188 bytes expected).
#TF#SERVER: Received 5188 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 792688 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=23, total=1052).
#TF#SERVER: Sent 355761 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4192 bytes expected).
#TF#SERVER: Received 4192 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1017637 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1121806 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=24, total=1053).
#TF#SERVER: Sent 363949 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4196 bytes expected).
#TF#SERVER: Received 4196 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 766020 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1153336 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=313, total=1342).
#TF#SERVER: Sent 791543 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5352 bytes expected).
#TF#SERVER: Received 5352 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1033355 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=286, total=1315).
#TF#SERVER: Sent 704037 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5244 bytes expected).
#TF#SERVER: Received 5244 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 925224 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=409, total=1438).
#TF#SERVER: Sent 844576 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5736 bytes expected).
#TF#SERVER: Received 5736 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 717279 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=245, total=1274).
#TF#SERVER: Sent 628918 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5080 bytes expected).
#TF#SERVER: Received 5080 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 767936 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=97, total=1126).
#TF#SERVER: Sent 423172 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4488 bytes expected).
#TF#SERVER: Received 4488 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 837654 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=45, total=1074).
#TF#SERVER: Sent 375787 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4280 bytes expected).
#TF#SERVER: Received 4280 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 836931 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=72, total=1101).
#TF#SERVER: Sent 388958 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4388 bytes expected).
#TF#SERVER: Received 4388 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 822375 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=43, total=1072).
#TF#SERVER: Sent 377426 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4272 bytes expected).
#TF#SERVER: Received 4272 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 811914 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=44, total=1073).
#TF#SERVER: Sent 383951 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4276 bytes expected).
#TF#SERVER: Received 4276 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 781273 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=27, total=1056).
#TF#SERVER: Sent 361427 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4208 bytes expected).
#TF#SERVER: Received 4208 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 893291 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=98, total=1127).
#TF#SERVER: Sent 424998 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4492 bytes expected).
#TF#SERVER: Received 4492 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 823083 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=222, total=1251).
#TF#SERVER: Sent 529463 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4988 bytes expected).
#TF#SERVER: Received 4988 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 761191 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=148, total=1177).
#TF#SERVER: Sent 494811 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4692 bytes expected).
#TF#SERVER: Received 4692 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 863054 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=133, total=1162).
#TF#SERVER: Sent 507788 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4632 bytes expected).
#TF#SERVER: Received 4632 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 907394 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=186, total=1215).
#TF#SERVER: Sent 545137 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4844 bytes expected).
#TF#SERVER: Received 4844 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 846278 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 917736 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=57, total=1086).
#TF#SERVER: Sent 392548 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4328 bytes expected).
#TF#SERVER: Received 4328 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 738782 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=183, total=1212).
#TF#SERVER: Sent 468957 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4832 bytes expected).
#TF#SERVER: Received 4832 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 877570 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=108, total=1137).
#TF#SERVER: Sent 428777 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4532 bytes expected).
#TF#SERVER: Received 4532 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 847362 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=23, total=1052).
#TF#SERVER: Sent 362876 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4192 bytes expected).
#TF#SERVER: Received 4192 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 882026 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=295, total=1324).
#TF#SERVER: Sent 654357 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5280 bytes expected).
#TF#SERVER: Received 5280 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 959874 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=33, total=1062).
#TF#SERVER: Sent 356367 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4232 bytes expected).
#TF#SERVER: Received 4232 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 953274 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=198, total=1227).
#TF#SERVER: Sent 572316 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4892 bytes expected).
#TF#SERVER: Received 4892 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 985638 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=150, total=1179).
#TF#SERVER: Sent 512607 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4700 bytes expected).
#TF#SERVER: Received 4700 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 814237 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=83, total=1112).
#TF#SERVER: Sent 458440 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4432 bytes expected).
#TF#SERVER: Received 4432 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 940291 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=201, total=1230).
#TF#SERVER: Sent 573947 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4904 bytes expected).
#TF#SERVER: Received 4904 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 783398 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=13, total=1042).
#TF#SERVER: Sent 342934 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4152 bytes expected).
#TF#SERVER: Received 4152 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 776702 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=28, total=1057).
#TF#SERVER: Sent 361264 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4212 bytes expected).
#TF#SERVER: Received 4212 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 905902 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=297, total=1326).
#TF#SERVER: Sent 693641 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5288 bytes expected).
#TF#SERVER: Received 5288 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 757306 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=267, total=1296).
#TF#SERVER: Sent 572094 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (5168 bytes expected).
#TF#SERVER: Received 5168 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 834416 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=106, total=1135).
#TF#SERVER: Sent 428884 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4524 bytes expected).
#TF#SERVER: Received 4524 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 1053401 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=489, total=1518).
#TF#SERVER: Sent 1061124 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6056 bytes expected).
#TF#SERVER: Received 6056 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 910210 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=111, total=1140).
#TF#SERVER: Sent 471382 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4544 bytes expected).
#TF#SERVER: Received 4544 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 841403 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=16, total=1045).
#TF#SERVER: Sent 356468 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4164 bytes expected).
#TF#SERVER: Received 4164 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 906197 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=29, total=1058).
#TF#SERVER: Sent 358508 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4216 bytes expected).
#TF#SERVER: Received 4216 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 731786 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=107, total=1136).
#TF#SERVER: Sent 421038 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4528 bytes expected).
#TF#SERVER: Received 4528 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=512, total=1541).
#TF#SERVER: Sent 744064 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (6148 bytes expected).
#TF#SERVER: Received 6148 bytes from the TF server.
#TF#SERVER: Sending data to the TF server (conj=5, context=1024, query=153, total=1182).
#TF#SERVER: Sent 468064 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (4712 bytes expected).
#TF#SERVER: Received 4712 bytes from the TF server.

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(commutativity_k7_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k7_real_1(X1,X2)=k7_real_1(X2,X1)), file('/scratch/mptp/minsubpreds4_different/t20_complex2', commutativity_k7_real_1)).
fof(t8_complex1, axiom, ![X1]:(v1_xcmplx_0(X1)=>![X2]:(v1_xcmplx_0(X2)=>(k3_complex1(k2_xcmplx_0(X1,X2))=k7_real_1(k3_complex1(X1),k3_complex1(X2))&k4_complex1(k2_xcmplx_0(X1,X2))=k7_real_1(k4_complex1(X1),k4_complex1(X2))))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', t8_complex1)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/scratch/mptp/minsubpreds4_different/t20_complex2', t6_boole)).
fof(dt_k7_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>m1_subset_1(k7_real_1(X1,X2),k1_numbers)), file('/scratch/mptp/minsubpreds4_different/t20_complex2', dt_k7_real_1)).
fof(redefinition_k7_real_1, axiom, ![X1, X2]:((m1_subset_1(X1,k1_numbers)&v1_xreal_0(X2))=>k7_real_1(X1,X2)=k2_xcmplx_0(X1,X2)), file('/scratch/mptp/minsubpreds4_different/t20_complex2', redefinition_k7_real_1)).
fof(dt_k4_complex1, axiom, ![X1]:(v1_xcmplx_0(X1)=>m1_subset_1(k4_complex1(X1),k1_numbers)), file('/scratch/mptp/minsubpreds4_different/t20_complex2', dt_k4_complex1)).
fof(t12_complex1, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>(k3_complex1(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)))=X1&k4_complex1(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)))=X2))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', t12_complex1)).
fof(commutativity_k2_xcmplx_0, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)), file('/scratch/mptp/minsubpreds4_different/t20_complex2', commutativity_k2_xcmplx_0)).
fof(redefinition_k6_numbers, axiom, k6_numbers=k1_xboole_0, file('/scratch/mptp/minsubpreds4_different/t20_complex2', redefinition_k6_numbers)).
fof(rc4_xreal_0, axiom, ?[X1]:(((v1_xboole_0(X1)&v1_xcmplx_0(X1))&v1_xxreal_0(X1))&v1_xreal_0(X1)), file('/scratch/mptp/minsubpreds4_different/t20_complex2', rc4_xreal_0)).
fof(cc1_xreal_0, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>v1_xreal_0(X1)), file('/scratch/mptp/minsubpreds4_different/t20_complex2', cc1_xreal_0)).
fof(cc3_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xcmplx_0(X1)), file('/scratch/mptp/minsubpreds4_different/t20_complex2', cc3_xreal_0)).
fof(rqRealMult__k3_xcmplx_0__r0_cr1r0_r0, axiom, k3_xcmplx_0(np__0,k1_xcmplx_0)=np__0, file('/scratch/mptp/minsubpreds4_different/t20_complex2', rqRealMult__k3_xcmplx_0__r0_cr1r0_r0)).
fof(redefinition_k7_complex1, axiom, k7_complex1=k1_xcmplx_0, file('/scratch/mptp/minsubpreds4_different/t20_complex2', redefinition_k7_complex1)).
fof(spc0_boole, axiom, v1_xboole_0(np__0), file('/scratch/mptp/minsubpreds4_different/t20_complex2', spc0_boole)).
fof(t1_arithm, axiom, ![X1]:(v1_xcmplx_0(X1)=>k2_xcmplx_0(X1,k6_numbers)=X1), file('/scratch/mptp/minsubpreds4_different/t20_complex2', t1_arithm)).
fof(cc1_xcmplx_0, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>v1_xcmplx_0(X1)), file('/scratch/mptp/minsubpreds4_different/t20_complex2', cc1_xcmplx_0)).
fof(t4_xxreal_1, axiom, ![X1]:(v1_xxreal_0(X1)=>![X2]:(v1_xxreal_0(X2)=>![X3]:(v1_xxreal_0(X3)=>(r2_hidden(X1,k4_xxreal_1(X2,X3))<=>(~(r1_xxreal_0(X1,X2))&~(r1_xxreal_0(X3,X1))))))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', t4_xxreal_1)).
fof(redefinition_k2_rcomp_1, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>k2_rcomp_1(X1,X2)=k4_xxreal_1(X1,X2)), file('/scratch/mptp/minsubpreds4_different/t20_complex2', redefinition_k2_rcomp_1)).
fof(rc4_xxreal_0, axiom, ?[X1]:(v1_xboole_0(X1)&v1_xxreal_0(X1)), file('/scratch/mptp/minsubpreds4_different/t20_complex2', rc4_xxreal_0)).
fof(cc13_membered, axiom, ![X1]:(v2_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xxreal_0(X2))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', cc13_membered)).
fof(dt_k1_comptrig, axiom, ![X1]:(v1_xcmplx_0(X1)=>m1_subset_1(k1_comptrig(X1),k1_numbers)), file('/scratch/mptp/minsubpreds4_different/t20_complex2', dt_k1_comptrig)).
fof(spc0_numerals, axiom, ((m2_subset_1(np__0,k1_numbers,k5_numbers)&m1_subset_1(np__0,k5_numbers))&m1_subset_1(np__0,k1_numbers)), file('/scratch/mptp/minsubpreds4_different/t20_complex2', spc0_numerals)).
fof(t18_complex2, axiom, ![X1]:(v1_xcmplx_0(X1)=>(r2_hidden(k1_comptrig(X1),k2_rcomp_1(k6_numbers,k32_sin_cos))<=>~(r1_xxreal_0(k4_complex1(X1),k6_numbers)))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', t18_complex2)).
fof(dt_k32_sin_cos, axiom, m1_subset_1(k32_sin_cos,k1_numbers), file('/scratch/mptp/minsubpreds4_different/t20_complex2', dt_k32_sin_cos)).
fof(cc4_membered, axiom, ![X1]:(v3_membered(X1)=>v2_membered(X1)), file('/scratch/mptp/minsubpreds4_different/t20_complex2', cc4_membered)).
fof(fc12_xreal_0, axiom, ![X1, X2]:((((v2_xxreal_0(X1)&v1_xreal_0(X1))&~(v3_xxreal_0(X2)))&v1_xreal_0(X2))=>v2_xxreal_0(k2_xcmplx_0(X2,X1))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', fc12_xreal_0)).
fof(fc3_membered, axiom, v3_membered(k1_numbers), file('/scratch/mptp/minsubpreds4_different/t20_complex2', fc3_membered)).
fof(cc3_xxreal_0, axiom, ![X1]:((v1_xxreal_0(X1)&v2_xxreal_0(X1))=>((~(v1_xboole_0(X1))&v1_xxreal_0(X1))&~(v3_xxreal_0(X1)))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', cc3_xxreal_0)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('/scratch/mptp/minsubpreds4_different/t20_complex2', fc1_xboole_0)).
fof(t1_real, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>((r1_xxreal_0(X1,X2)&v2_xxreal_0(X1))=>v2_xxreal_0(X2)))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', t1_real)).
fof(spc6_arithm, axiom, ![X1, X2, X3]:(((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))&v1_xcmplx_0(X3))=>k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', spc6_arithm)).
fof(cc5_xxreal_0, axiom, ![X1]:((v1_xxreal_0(X1)&v3_xxreal_0(X1))=>((~(v1_xboole_0(X1))&v1_xxreal_0(X1))&~(v2_xxreal_0(X1)))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', cc5_xxreal_0)).
fof(fc2_xcmplx_0, axiom, ![X1, X2]:((v1_xcmplx_0(X1)&v1_xcmplx_0(X2))=>v1_xcmplx_0(k2_xcmplx_0(X1,X2))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', fc2_xcmplx_0)).
fof(t20_complex2, conjecture, ![X1]:(v1_xcmplx_0(X1)=>![X2]:(v1_xcmplx_0(X2)=>~(((~(r1_xxreal_0(k32_sin_cos,k1_comptrig(X1)))&~(r1_xxreal_0(k32_sin_cos,k1_comptrig(X2))))&r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(X1,X2))))))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', t20_complex2)).
fof(t8_real, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>~(((~(r1_xxreal_0(X1,X2))&~(v3_xxreal_0(X2)))&~(v2_xxreal_0(X1)))))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', t8_real)).
fof(fc5_xreal_0, axiom, ![X1, X2]:((v1_xreal_0(X1)&v1_xreal_0(X2))=>v1_xreal_0(k2_xcmplx_0(X1,X2))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', fc5_xreal_0)).
fof(reflexivity_r1_xxreal_0, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>r1_xxreal_0(X1,X1)), file('/scratch/mptp/minsubpreds4_different/t20_complex2', reflexivity_r1_xxreal_0)).
fof(cc4_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xxreal_0(X1)), file('/scratch/mptp/minsubpreds4_different/t20_complex2', cc4_xreal_0)).
fof(t2_real, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>((r1_xxreal_0(X1,X2)&v3_xxreal_0(X2))=>v3_xxreal_0(X1)))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', t2_real)).
fof(t46_complex1, axiom, ![X1]:(v1_xcmplx_0(X1)=>r1_xxreal_0(k6_numbers,k17_complex1(X1))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', t46_complex1)).
fof(dt_k16_complex1, axiom, ![X1]:(v1_xcmplx_0(X1)=>v1_xreal_0(k16_complex1(X1))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', dt_k16_complex1)).
fof(redefinition_k17_complex1, axiom, ![X1]:(v1_xcmplx_0(X1)=>k17_complex1(X1)=k16_complex1(X1)), file('/scratch/mptp/minsubpreds4_different/t20_complex2', redefinition_k17_complex1)).
fof(dt_k31_sin_cos, axiom, v1_xreal_0(k31_sin_cos), file('/scratch/mptp/minsubpreds4_different/t20_complex2', dt_k31_sin_cos)).
fof(redefinition_k32_sin_cos, axiom, k32_sin_cos=k31_sin_cos, file('/scratch/mptp/minsubpreds4_different/t20_complex2', redefinition_k32_sin_cos)).
fof(t5_real, axiom, ![X1]:(v1_xreal_0(X1)=>![X2]:(v1_xreal_0(X2)=>(r1_xxreal_0(X1,X2)=>(v1_xboole_0(X2)|v3_xxreal_0(X1)|v2_xxreal_0(X2))))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', t5_real)).
fof(t34_comptrig, axiom, ![X1]:(v1_xcmplx_0(X1)=>(r1_xxreal_0(k6_numbers,k1_comptrig(X1))&~(r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k1_comptrig(X1))))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', t34_comptrig)).
fof(connectedness_r1_xxreal_0, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>(r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', connectedness_r1_xxreal_0)).
fof(t15_complex2, axiom, ![X1]:(v1_xcmplx_0(X1)=>(k1_comptrig(X1)=k6_numbers<=>X1=k17_complex1(X1))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', t15_complex2)).
fof(t35_comptrig, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>(r1_xxreal_0(k6_numbers,X1)=>k1_comptrig(X1)=k6_numbers)), file('/scratch/mptp/minsubpreds4_different/t20_complex2', t35_comptrig)).
fof(t5_comptrig, axiom, ((((((((((~(r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k6_numbers))&~(r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__2))))&~(r1_xxreal_0(k32_sin_cos,k6_numbers)))&~(r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k1_real_1(k10_real_1(k32_sin_cos,np__2)))))&~(r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k32_sin_cos)))&~(r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k10_real_1(k32_sin_cos,np__2))))&~(r1_xxreal_0(k6_numbers,k1_real_1(k10_real_1(k32_sin_cos,np__2)))))&~(r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k6_numbers)))&~(r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k32_sin_cos)))&~(r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos))))&~(r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k6_numbers))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', t5_comptrig)).
fof(fc9_xreal_0, axiom, ![X1, X2]:((((~(v3_xxreal_0(X1))&v1_xreal_0(X1))&~(v3_xxreal_0(X2)))&v1_xreal_0(X2))=>~(v3_xxreal_0(k2_xcmplx_0(X1,X2)))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', fc9_xreal_0)).
fof(cc14_membered, axiom, ![X1]:(v3_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v1_xreal_0(X2))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', cc14_membered)).
fof(dt_k17_complex1, axiom, ![X1]:(v1_xcmplx_0(X1)=>m1_subset_1(k17_complex1(X1),k1_numbers)), file('/scratch/mptp/minsubpreds4_different/t20_complex2', dt_k17_complex1)).
fof(cc4_xxreal_0, axiom, ![X1]:(((~(v1_xboole_0(X1))&v1_xxreal_0(X1))&~(v3_xxreal_0(X1)))=>(v1_xxreal_0(X1)&v2_xxreal_0(X1))), file('/scratch/mptp/minsubpreds4_different/t20_complex2', cc4_xxreal_0)).
fof(rqLessOrEqual__r1_xxreal_0__r0_r0, axiom, r1_xxreal_0(np__0,np__0), file('/scratch/mptp/minsubpreds4_different/t20_complex2', rqLessOrEqual__r1_xxreal_0__r0_r0)).
fof(c_0_56, plain, ![X26, X27]:(~m1_subset_1(X26,k1_numbers)|~v1_xreal_0(X27)|k7_real_1(X26,X27)=k7_real_1(X27,X26)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k7_real_1])])).
fof(c_0_57, plain, ![X92, X93]:((k3_complex1(k2_xcmplx_0(X92,X93))=k7_real_1(k3_complex1(X92),k3_complex1(X93))|~v1_xcmplx_0(X93)|~v1_xcmplx_0(X92))&(k4_complex1(k2_xcmplx_0(X92,X93))=k7_real_1(k4_complex1(X92),k4_complex1(X93))|~v1_xcmplx_0(X93)|~v1_xcmplx_0(X92))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t8_complex1])])])])).
fof(c_0_58, plain, ![X89]:(~v1_xboole_0(X89)|X89=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
fof(c_0_59, plain, ![X36, X37]:(~m1_subset_1(X36,k1_numbers)|~v1_xreal_0(X37)|m1_subset_1(k7_real_1(X36,X37),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k7_real_1])])).
fof(c_0_60, plain, ![X58, X59]:(~m1_subset_1(X58,k1_numbers)|~v1_xreal_0(X59)|k7_real_1(X58,X59)=k2_xcmplx_0(X58,X59)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k7_real_1])])).
cnf(c_0_61, plain, (k7_real_1(X1,X2)=k7_real_1(X2,X1)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_62, plain, (k4_complex1(k2_xcmplx_0(X1,X2))=k7_real_1(k4_complex1(X1),k4_complex1(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
fof(c_0_63, plain, ![X35]:(~v1_xcmplx_0(X35)|m1_subset_1(k4_complex1(X35),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_complex1])])).
fof(c_0_64, plain, ![X68, X69]:((k3_complex1(k2_xcmplx_0(X68,k3_xcmplx_0(X69,k7_complex1)))=X68|~v1_xreal_0(X69)|~v1_xreal_0(X68))&(k4_complex1(k2_xcmplx_0(X68,k3_xcmplx_0(X69,k7_complex1)))=X69|~v1_xreal_0(X69)|~v1_xreal_0(X68))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t12_complex1])])])])).
fof(c_0_65, plain, ![X24, X25]:(~v1_xcmplx_0(X24)|~v1_xcmplx_0(X25)|k2_xcmplx_0(X24,X25)=k2_xcmplx_0(X25,X24)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k2_xcmplx_0])])).
cnf(c_0_66, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_67, plain, (k6_numbers=k1_xboole_0), inference(split_conjunct,[status(thm)],[redefinition_k6_numbers])).
fof(c_0_68, plain, (((v1_xboole_0(esk4_0)&v1_xcmplx_0(esk4_0))&v1_xxreal_0(esk4_0))&v1_xreal_0(esk4_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc4_xreal_0])])).
cnf(c_0_69, plain, (m1_subset_1(k7_real_1(X1,X2),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_70, plain, (k7_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_71, plain, (k7_real_1(X1,X2)=k7_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_61])).
cnf(c_0_72, plain, (k7_real_1(k4_complex1(X1),k4_complex1(X2))=k4_complex1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_62])).
cnf(c_0_73, plain, (m1_subset_1(k4_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_63])).
fof(c_0_74, plain, ![X16]:(~m1_subset_1(X16,k1_numbers)|v1_xreal_0(X16)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_xreal_0])])).
cnf(c_0_75, plain, (k3_complex1(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)))=X1|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_76, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_65])).
fof(c_0_77, plain, ![X18]:(~v1_xreal_0(X18)|v1_xcmplx_0(X18)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc3_xreal_0])])).
cnf(c_0_78, plain, (k3_xcmplx_0(np__0,k1_xcmplx_0)=np__0), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r0_cr1r0_r0])).
cnf(c_0_79, plain, (k7_complex1=k1_xcmplx_0), inference(split_conjunct,[status(thm)],[redefinition_k7_complex1])).
cnf(c_0_80, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_66])).
cnf(c_0_81, plain, (k1_xboole_0=k6_numbers), inference(evalgc,[status(thm)],[c_0_67])).
cnf(c_0_82, plain, (v1_xboole_0(np__0)), inference(split_conjunct,[status(thm)],[spc0_boole])).
cnf(c_0_83, plain, (v1_xboole_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_68])).
fof(c_0_84, plain, ![X72]:(~v1_xcmplx_0(X72)|k2_xcmplx_0(X72,k6_numbers)=X72), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_arithm])])).
cnf(c_0_85, plain, (m1_subset_1(k7_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_69])).
cnf(c_0_86, plain, (k7_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_70])).
fof(c_0_87, plain, ![X15]:(~m1_subset_1(X15,k1_numbers)|v1_xcmplx_0(X15)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_xcmplx_0])])).
fof(c_0_88, plain, ![X84, X85, X86]:(((~r1_xxreal_0(X84,X85)|~r2_hidden(X84,k4_xxreal_1(X85,X86))|~v1_xxreal_0(X86)|~v1_xxreal_0(X85)|~v1_xxreal_0(X84))&(~r1_xxreal_0(X86,X84)|~r2_hidden(X84,k4_xxreal_1(X85,X86))|~v1_xxreal_0(X86)|~v1_xxreal_0(X85)|~v1_xxreal_0(X84)))&(r1_xxreal_0(X84,X85)|r1_xxreal_0(X86,X84)|r2_hidden(X84,k4_xxreal_1(X85,X86))|~v1_xxreal_0(X86)|~v1_xxreal_0(X85)|~v1_xxreal_0(X84))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t4_xxreal_1])])])])])).
fof(c_0_89, plain, ![X56, X57]:(~v1_xxreal_0(X56)|~v1_xxreal_0(X57)|k2_rcomp_1(X56,X57)=k4_xxreal_1(X56,X57)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_rcomp_1])])).
fof(c_0_90, plain, (v1_xboole_0(esk5_0)&v1_xxreal_0(esk5_0)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc4_xxreal_0])])).
fof(c_0_91, plain, ![X6, X7]:(~v2_membered(X6)|(~m1_subset_1(X7,X6)|v1_xxreal_0(X7))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc13_membered])])])).
fof(c_0_92, plain, ![X32]:(~v1_xcmplx_0(X32)|m1_subset_1(k1_comptrig(X32),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_comptrig])])).
cnf(c_0_93, plain, (k7_real_1(k4_complex1(X2),k4_complex1(X1))=k4_complex1(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_complex1(X2))|~m1_subset_1(k4_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72])])).
cnf(c_0_94, plain, (m1_subset_1(k4_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_73])).
cnf(c_0_95, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_96, plain, (k3_complex1(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)))=X1|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_75])).
cnf(c_0_97, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_76])).
cnf(c_0_98, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_77])).
cnf(c_0_99, plain, (k3_xcmplx_0(np__0,k1_xcmplx_0)=np__0), inference(evalgc,[status(thm)],[c_0_78])).
cnf(c_0_100, plain, (k1_xcmplx_0=k7_complex1), inference(evalgc,[status(thm)],[c_0_79])).
cnf(c_0_101, plain, (X1=k6_numbers|~v1_xboole_0(X1)), inference(rw,[status(thm)],[c_0_80, c_0_81])).
cnf(c_0_102, plain, (v1_xboole_0(np__0)), inference(evalgc,[status(thm)],[c_0_82])).
cnf(c_0_103, plain, (v1_xreal_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_104, plain, (v1_xboole_0(esk4_0)), inference(evalgc,[status(thm)],[c_0_83])).
cnf(c_0_105, plain, (v1_xcmplx_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_106, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_107, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_86])])).
cnf(c_0_108, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_87])).
cnf(c_0_109, plain, (m1_subset_1(np__0,k1_numbers)), inference(split_conjunct,[status(thm)],[spc0_numerals])).
cnf(c_0_110, plain, (~r1_xxreal_0(X1,X2)|~r2_hidden(X2,k4_xxreal_1(X3,X1))|~v1_xxreal_0(X1)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_111, plain, (k2_rcomp_1(X1,X2)=k4_xxreal_1(X1,X2)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_89])).
fof(c_0_112, plain, ![X71]:((~r2_hidden(k1_comptrig(X71),k2_rcomp_1(k6_numbers,k32_sin_cos))|~r1_xxreal_0(k4_complex1(X71),k6_numbers)|~v1_xcmplx_0(X71))&(r1_xxreal_0(k4_complex1(X71),k6_numbers)|r2_hidden(k1_comptrig(X71),k2_rcomp_1(k6_numbers,k32_sin_cos))|~v1_xcmplx_0(X71))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t18_complex2])])])])).
cnf(c_0_113, plain, (v1_xboole_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_114, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_115, plain, (m1_subset_1(k32_sin_cos,k1_numbers)), inference(split_conjunct,[status(thm)],[dt_k32_sin_cos])).
fof(c_0_116, plain, ![X20]:(~v3_membered(X20)|v2_membered(X20)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_membered])])).
cnf(c_0_117, plain, (m1_subset_1(k1_comptrig(X1),k1_numbers)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_92])).
fof(c_0_118, plain, ![X38, X39]:(~v2_xxreal_0(X38)|~v1_xreal_0(X38)|v3_xxreal_0(X39)|~v1_xreal_0(X39)|v2_xxreal_0(k2_xcmplx_0(X39,X38))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc12_xreal_0])])])).
cnf(c_0_119, plain, (k7_real_1(k4_complex1(X1),k4_complex1(X2))=k4_complex1(k2_xcmplx_0(X2,X1))|~v1_xreal_0(k4_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(csr,[status(thm)],[c_0_93, c_0_94])).
cnf(c_0_120, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_95])).
cnf(c_0_121, plain, (k7_real_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_71])])).
cnf(c_0_122, plain, (k3_complex1(k2_xcmplx_0(k3_xcmplx_0(X2,k7_complex1),X1))=X1|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,k7_complex1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_97])])).
cnf(c_0_123, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_124, plain, (k3_xcmplx_0(np__0,k7_complex1)=np__0), inference(rw,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_125, plain, (np__0=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102])])).
cnf(c_0_126, plain, (v1_xreal_0(esk4_0)), inference(evalgc,[status(thm)],[c_0_103])).
cnf(c_0_127, plain, (esk4_0=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_104])])).
cnf(c_0_128, plain, (v1_xcmplx_0(esk4_0)), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_129, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_130, plain, (k3_complex1(k2_xcmplx_0(X1,X2))=k7_real_1(k3_complex1(X1),k3_complex1(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_131, plain, (m1_subset_1(k2_xcmplx_0(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_97])])).
cnf(c_0_132, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_108])).
cnf(c_0_133, plain, (m1_subset_1(np__0,k1_numbers)), inference(evalgc,[status(thm)],[c_0_109])).
cnf(c_0_134, plain, (~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~r2_hidden(X2,k4_xxreal_1(X3,X1))), inference(evalgc,[status(thm)],[c_0_110])).
cnf(c_0_135, plain, (k4_xxreal_1(X1,X2)=k2_rcomp_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_111])).
cnf(c_0_136, plain, (r1_xxreal_0(k4_complex1(X1),k6_numbers)|r2_hidden(k1_comptrig(X1),k2_rcomp_1(k6_numbers,k32_sin_cos))|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_137, plain, (v1_xxreal_0(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_90])).
cnf(c_0_138, plain, (v1_xboole_0(esk5_0)), inference(evalgc,[status(thm)],[c_0_113])).
cnf(c_0_139, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_114])).
cnf(c_0_140, plain, (m1_subset_1(k32_sin_cos,k1_numbers)), inference(evalgc,[status(thm)],[c_0_115])).
cnf(c_0_141, plain, (v2_membered(X1)|~v3_membered(X1)), inference(split_conjunct,[status(thm)],[c_0_116])).
cnf(c_0_142, plain, (v3_membered(k1_numbers)), inference(split_conjunct,[status(thm)],[fc3_membered])).
cnf(c_0_143, plain, (m1_subset_1(k1_comptrig(X1),k1_numbers)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_117])).
fof(c_0_144, plain, ![X19]:(((~v1_xboole_0(X19)|(~v1_xxreal_0(X19)|~v2_xxreal_0(X19)))&(v1_xxreal_0(X19)|(~v1_xxreal_0(X19)|~v2_xxreal_0(X19))))&(~v3_xxreal_0(X19)|(~v1_xxreal_0(X19)|~v2_xxreal_0(X19)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc3_xxreal_0])])])])).
cnf(c_0_145, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_146, plain, (v3_xxreal_0(X2)|v2_xxreal_0(k2_xcmplx_0(X2,X1))|~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_118])).
cnf(c_0_147, plain, (k2_xcmplx_0(k4_complex1(X1),k4_complex1(X2))=k4_complex1(k2_xcmplx_0(X2,X1))|~v1_xreal_0(k4_complex1(X2))|~m1_subset_1(k4_complex1(X1),k1_numbers)|~v1_xreal_0(k4_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_119])])).
cnf(c_0_148, plain, (v1_xreal_0(k4_complex1(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_94])])).
cnf(c_0_149, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_121])])).
cnf(c_0_150, plain, (k3_complex1(k2_xcmplx_0(k3_xcmplx_0(X1,k7_complex1),X2))=X2|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,k7_complex1))), inference(csr,[status(thm)],[c_0_122, c_0_123])).
cnf(c_0_151, plain, (k3_xcmplx_0(k6_numbers,k7_complex1)=k6_numbers), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_124, c_0_125]), c_0_125])])).
cnf(c_0_152, plain, (v1_xreal_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_126, c_0_127])])).
cnf(c_0_153, plain, (v1_xcmplx_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_128, c_0_127])])).
cnf(c_0_154, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_97])])).
cnf(c_0_155, plain, (k7_real_1(k3_complex1(X1),k3_complex1(X2))=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_130])).
cnf(c_0_156, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_131, c_0_132]), c_0_123])).
cnf(c_0_157, plain, (m1_subset_1(k6_numbers,k1_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_133, c_0_125])])).
fof(c_0_158, plain, ![X73, X74]:(~v1_xreal_0(X73)|(~v1_xreal_0(X74)|(~r1_xxreal_0(X73,X74)|~v2_xxreal_0(X73)|v2_xxreal_0(X74)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_real])])])).
cnf(c_0_159, plain, (~r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v1_xxreal_0(X3)|~r1_xxreal_0(X3,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135])])).
cnf(c_0_160, plain, (r1_xxreal_0(k4_complex1(X1),k6_numbers)|r2_hidden(k1_comptrig(X1),k2_rcomp_1(k6_numbers,k32_sin_cos))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_136])).
cnf(c_0_161, plain, (v1_xxreal_0(esk5_0)), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_162, plain, (esk5_0=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_138])])).
cnf(c_0_163, plain, (v1_xxreal_0(k32_sin_cos)|~v2_membered(k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_140])])).
cnf(c_0_164, plain, (v2_membered(X1)|~v3_membered(X1)), inference(evalgc,[status(thm)],[c_0_141])).
cnf(c_0_165, plain, (v3_membered(k1_numbers)), inference(evalgc,[status(thm)],[c_0_142])).
cnf(c_0_166, plain, (v1_xxreal_0(k1_comptrig(X1))|~v2_membered(k1_numbers)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_143])])).
cnf(c_0_167, plain, (~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v2_xxreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_144])).
cnf(c_0_168, plain, (v1_xboole_0(k1_xboole_0)), inference(evalgc,[status(thm)],[c_0_145])).
cnf(c_0_169, plain, (v3_xxreal_0(X2)|v2_xxreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_146])).
cnf(c_0_170, plain, (k2_xcmplx_0(k4_complex1(X1),k4_complex1(X2))=k4_complex1(k2_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_147, c_0_94]), c_0_148]), c_0_148])).
cnf(c_0_171, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_149, c_0_120]), c_0_120])).
cnf(c_0_172, plain, (k3_complex1(k2_xcmplx_0(k6_numbers,X1))=X1|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_151]), c_0_152]), c_0_153])])])).
cnf(c_0_173, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_154, c_0_153])])])).
cnf(c_0_174, plain, (k2_xcmplx_0(k3_complex1(X1),k3_complex1(X2))=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_155])])).
cnf(c_0_175, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_129]), c_0_157])])])).
cnf(c_0_176, plain, (v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~v2_xxreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_158])).
cnf(c_0_177, plain, (r1_xxreal_0(k4_complex1(X1),k6_numbers)|~v1_xxreal_0(k6_numbers)|~v1_xxreal_0(k1_comptrig(X1))|~v1_xxreal_0(k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k1_comptrig(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_160])])).
cnf(c_0_178, plain, (v1_xxreal_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_161, c_0_162])])).
cnf(c_0_179, plain, (v1_xxreal_0(k32_sin_cos)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_164]), c_0_165])])])).
cnf(c_0_180, plain, (v1_xxreal_0(k1_comptrig(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_164]), c_0_165])])])).
cnf(c_0_181, plain, (~v1_xxreal_0(X1)|~v2_xxreal_0(X1)|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_167])).
cnf(c_0_182, plain, (v1_xboole_0(k6_numbers)), inference(rw,[status(thm)],[c_0_168, c_0_81])).
cnf(c_0_183, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k2_xcmplx_0(X2,X1))|~v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_97])])).
cnf(c_0_184, plain, (k2_xcmplx_0(k4_complex1(X2),k4_complex1(X1))=k4_complex1(k2_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(k4_complex1(X2),k1_numbers)|~m1_subset_1(k4_complex1(X1),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_171])])).
fof(c_0_185, plain, ![X62, X63, X64]:(~v1_xcmplx_0(X62)|~v1_xcmplx_0(X63)|~v1_xcmplx_0(X64)|k2_xcmplx_0(k2_xcmplx_0(X62,X63),X64)=k2_xcmplx_0(X62,k2_xcmplx_0(X63,X64))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[spc6_arithm])])).
cnf(c_0_186, plain, (k3_complex1(X1)=X1|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_173])])).
cnf(c_0_187, plain, (k2_xcmplx_0(k3_complex1(X2),k3_complex1(X1))=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k3_complex1(X2),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_171])])).
cnf(c_0_188, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_175, c_0_123])).
cnf(c_0_189, plain, (v2_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_176])).
cnf(c_0_190, plain, (r1_xxreal_0(k4_complex1(X1),k6_numbers)|~r1_xxreal_0(k32_sin_cos,k1_comptrig(X1))|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_177, c_0_178]), c_0_179])]), c_0_180])).
cnf(c_0_191, plain, (~v2_xxreal_0(k6_numbers)|~v1_xxreal_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_182])])).
cnf(c_0_192, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k2_xcmplx_0(X2,X1))|~v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_183, c_0_123]), c_0_123])).
cnf(c_0_193, plain, (k2_xcmplx_0(k4_complex1(X1),k4_complex1(X2))=k4_complex1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_184, c_0_94]), c_0_94])).
fof(c_0_194, plain, ![X23]:(((~v1_xboole_0(X23)|(~v1_xxreal_0(X23)|~v3_xxreal_0(X23)))&(v1_xxreal_0(X23)|(~v1_xxreal_0(X23)|~v3_xxreal_0(X23))))&(~v2_xxreal_0(X23)|(~v1_xxreal_0(X23)|~v3_xxreal_0(X23)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc5_xxreal_0])])])])).
cnf(c_0_195, plain, (k4_complex1(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)))=X2|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_64])).
cnf(c_0_196, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)), inference(split_conjunct,[status(thm)],[c_0_185])).
cnf(c_0_197, plain, (k3_complex1(X1)=X1|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_186, c_0_123])).
cnf(c_0_198, plain, (k2_xcmplx_0(k3_complex1(X1),k3_complex1(X2))=k3_complex1(k2_xcmplx_0(X2,X1))|~v1_xreal_0(k3_complex1(X1))|~m1_subset_1(k3_complex1(X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(csr,[status(thm)],[c_0_187, c_0_188])).
cnf(c_0_199, plain, (v2_xxreal_0(k6_numbers)|~v2_xxreal_0(k4_complex1(X1))|~v1_xreal_0(k4_complex1(X1))|~r1_xxreal_0(k32_sin_cos,k1_comptrig(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_190]), c_0_152])])])).
cnf(c_0_200, plain, (~v2_xxreal_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_191, c_0_178])])])).
cnf(c_0_201, plain, (v3_xxreal_0(k4_complex1(X1))|v2_xxreal_0(k4_complex1(k2_xcmplx_0(X2,X1)))|~v2_xxreal_0(k4_complex1(X2))|~v1_xreal_0(k4_complex1(X1))|~v1_xreal_0(k4_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_193])])).
fof(c_0_202, plain, ![X43, X44]:(~v1_xcmplx_0(X43)|~v1_xcmplx_0(X44)|v1_xcmplx_0(k2_xcmplx_0(X43,X44))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc2_xcmplx_0])])).
fof(c_0_203, negated_conjecture, ~(![X1]:(v1_xcmplx_0(X1)=>![X2]:(v1_xcmplx_0(X2)=>~(((~(r1_xxreal_0(k32_sin_cos,k1_comptrig(X1)))&~(r1_xxreal_0(k32_sin_cos,k1_comptrig(X2))))&r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(X1,X2)))))))), inference(assume_negation,[status(cth)],[t20_complex2])).
fof(c_0_204, plain, ![X94, X95]:(~v1_xreal_0(X94)|(~v1_xreal_0(X95)|(r1_xxreal_0(X94,X95)|v3_xxreal_0(X95)|v2_xxreal_0(X94)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t8_real])])])])).
cnf(c_0_205, plain, (~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v3_xxreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_194])).
cnf(c_0_206, plain, (k4_complex1(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)))=X2|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_195])).
cnf(c_0_207, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_196])).
cnf(c_0_208, plain, (k7_real_1(k3_complex1(X1),X2)=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_197])])).
fof(c_0_209, plain, ![X48, X49]:(~v1_xreal_0(X48)|~v1_xreal_0(X49)|v1_xreal_0(k2_xcmplx_0(X48,X49))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc5_xreal_0])])).
cnf(c_0_210, plain, (k3_complex1(k2_xcmplx_0(X2,X1))=k2_xcmplx_0(X1,k3_complex1(X2))|~v1_xreal_0(X1)|~m1_subset_1(k3_complex1(X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_197])])).
cnf(c_0_211, plain, (~v2_xxreal_0(k4_complex1(X1))|~r1_xxreal_0(k32_sin_cos,k1_comptrig(X1))|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[inference(sr,[status(thm)],[c_0_199, c_0_200]), c_0_148])).
cnf(c_0_212, plain, (v3_xxreal_0(k4_complex1(X1))|v2_xxreal_0(k4_complex1(k2_xcmplx_0(X2,X1)))|~v2_xxreal_0(k4_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_201, c_0_148]), c_0_148])).
cnf(c_0_213, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(split_conjunct,[status(thm)],[c_0_202])).
fof(c_0_214, negated_conjecture, (v1_xcmplx_0(esk1_0)&(v1_xcmplx_0(esk2_0)&((~r1_xxreal_0(k32_sin_cos,k1_comptrig(esk1_0))&~r1_xxreal_0(k32_sin_cos,k1_comptrig(esk2_0)))&r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_203])])])])).
cnf(c_0_215, plain, (r1_xxreal_0(X1,X2)|v3_xxreal_0(X2)|v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_204])).
cnf(c_0_216, plain, (~v1_xxreal_0(X1)|~v1_xboole_0(X1)|~v3_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_205])).
cnf(c_0_217, plain, (k4_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1))))=X3|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X3,k7_complex1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_207])])).
cnf(c_0_218, plain, (k7_real_1(k3_complex1(X1),X2)=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_208, c_0_123])).
cnf(c_0_219, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_209])).
cnf(c_0_220, plain, (k3_complex1(k2_xcmplx_0(X1,X2))=k2_xcmplx_0(X2,k3_complex1(X1))|~v1_xreal_0(X2)|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_210, c_0_123])).
cnf(c_0_221, plain, (k3_complex1(k6_numbers)=k6_numbers), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_129]), c_0_152]), c_0_153])])])).
cnf(c_0_222, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X3,X1)|r2_hidden(X1,k4_xxreal_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_223, plain, (v3_xxreal_0(k4_complex1(X2))|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v2_xxreal_0(k4_complex1(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_212])])).
cnf(c_0_224, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_213])).
cnf(c_0_225, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0)))), inference(split_conjunct,[status(thm)],[c_0_214])).
cnf(c_0_226, negated_conjecture, (v1_xcmplx_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_214])).
cnf(c_0_227, negated_conjecture, (v1_xcmplx_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_214])).
cnf(c_0_228, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_215])).
cnf(c_0_229, plain, (~v3_xxreal_0(k6_numbers)|~v1_xxreal_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_182])])).
cnf(c_0_230, plain, (k4_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k6_numbers)))=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_151]), c_0_152]), c_0_153])])])).
fof(c_0_231, plain, ![X60, X61]:(~v1_xxreal_0(X60)|~v1_xxreal_0(X61)|r1_xxreal_0(X60,X60)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[reflexivity_r1_xxreal_0])])).
fof(c_0_232, plain, ![X21]:(~v1_xreal_0(X21)|v1_xxreal_0(X21)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_xreal_0])])).
fof(c_0_233, plain, ![X76, X77]:(~v1_xreal_0(X76)|(~v1_xreal_0(X77)|(~r1_xxreal_0(X76,X77)|~v3_xxreal_0(X77)|v3_xxreal_0(X76)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_real])])])).
fof(c_0_234, plain, ![X80]:(~v1_xcmplx_0(X80)|r1_xxreal_0(k6_numbers,k17_complex1(X80))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t46_complex1])])).
fof(c_0_235, plain, ![X30]:(~v1_xcmplx_0(X30)|v1_xreal_0(k16_complex1(X30))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k16_complex1])])).
fof(c_0_236, plain, ![X55]:(~v1_xcmplx_0(X55)|k17_complex1(X55)=k16_complex1(X55)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k17_complex1])])).
cnf(c_0_237, plain, (k3_complex1(k2_xcmplx_0(X1,X2))=k2_xcmplx_0(k3_complex1(X1),X2)|~v1_xreal_0(X2)|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_218])])).
cnf(c_0_238, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_219])).
cnf(c_0_239, plain, (k3_complex1(X1)=k2_xcmplx_0(X1,k6_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_173]), c_0_221]), c_0_221]), c_0_157]), c_0_153])])])).
cnf(c_0_240, plain, (~r2_hidden(k1_comptrig(X1),k2_rcomp_1(k6_numbers,k32_sin_cos))|~r1_xxreal_0(k4_complex1(X1),k6_numbers)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_241, plain, (r1_xxreal_0(X3,X1)|r1_xxreal_0(X1,X2)|r2_hidden(X1,k4_xxreal_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_222])).
cnf(c_0_242, plain, (v3_xxreal_0(k4_complex1(X1))|~v2_xxreal_0(k4_complex1(X2))|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(X2,X1)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(csr,[status(thm)],[c_0_223, c_0_224])).
cnf(c_0_243, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[c_0_225])).
cnf(c_0_244, negated_conjecture, (v1_xcmplx_0(esk2_0)), inference(evalgc,[status(thm)],[c_0_226])).
cnf(c_0_245, negated_conjecture, (v1_xcmplx_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_227])).
cnf(c_0_246, plain, (v3_xxreal_0(k6_numbers)|v2_xxreal_0(X1)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_152])])).
cnf(c_0_247, plain, (~v3_xxreal_0(k6_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_229, c_0_178])])])).
cnf(c_0_248, plain, (k4_complex1(k2_xcmplx_0(X1,X2))=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_129])])).
cnf(c_0_249, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_231])).
cnf(c_0_250, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_232])).
cnf(c_0_251, plain, (v1_xreal_0(k31_sin_cos)), inference(split_conjunct,[status(thm)],[dt_k31_sin_cos])).
cnf(c_0_252, plain, (k32_sin_cos=k31_sin_cos), inference(split_conjunct,[status(thm)],[redefinition_k32_sin_cos])).
cnf(c_0_253, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)|~v3_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_233])).
cnf(c_0_254, plain, (r1_xxreal_0(k6_numbers,k17_complex1(X1))|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_234])).
cnf(c_0_255, plain, (v1_xreal_0(k16_complex1(X1))|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_235])).
cnf(c_0_256, plain, (k17_complex1(X1)=k16_complex1(X1)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_236])).
cnf(c_0_257, plain, (k3_complex1(X1)=k2_xcmplx_0(k6_numbers,X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_173]), c_0_221]), c_0_221]), c_0_157]), c_0_153])])])).
cnf(c_0_258, plain, (~v3_xxreal_0(X1)|~v1_xxreal_0(X1)|~v2_xxreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_144])).
cnf(c_0_259, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(X3)|r1_xxreal_0(X3,k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_228, c_0_238])])).
cnf(c_0_260, plain, (k3_complex1(X1)=k2_xcmplx_0(X1,k6_numbers)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_239, c_0_123])).
fof(c_0_261, plain, ![X87, X88]:(~v1_xreal_0(X87)|(~v1_xreal_0(X88)|(~r1_xxreal_0(X87,X88)|(v1_xboole_0(X88)|v3_xxreal_0(X87)|v2_xxreal_0(X88))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t5_real])])])).
fof(c_0_262, plain, ![X78]:((r1_xxreal_0(k6_numbers,k1_comptrig(X78))|~v1_xcmplx_0(X78))&(~r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k1_comptrig(X78))|~v1_xcmplx_0(X78))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t34_comptrig])])])])).
cnf(c_0_263, plain, (~v1_xcmplx_0(X1)|~r1_xxreal_0(k4_complex1(X1),k6_numbers)|~r2_hidden(k1_comptrig(X1),k2_rcomp_1(k6_numbers,k32_sin_cos))), inference(evalgc,[status(thm)],[c_0_240])).
cnf(c_0_264, plain, (r2_hidden(X1,k2_rcomp_1(X2,X3))|r1_xxreal_0(X1,X2)|r1_xxreal_0(X3,X1)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_135])])).
cnf(c_0_265, negated_conjecture, (v3_xxreal_0(k4_complex1(esk2_0))|~v2_xxreal_0(k4_complex1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_243]), c_0_244]), c_0_245])])])).
cnf(c_0_266, plain, (v2_xxreal_0(X1)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1)), inference(sr,[status(thm)],[c_0_246, c_0_247])).
cnf(c_0_267, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_248, c_0_173]), c_0_153])])])).
cnf(c_0_268, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_249])).
cnf(c_0_269, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_250])).
cnf(c_0_270, plain, (v1_xreal_0(k31_sin_cos)), inference(evalgc,[status(thm)],[c_0_251])).
cnf(c_0_271, plain, (k31_sin_cos=k32_sin_cos), inference(evalgc,[status(thm)],[c_0_252])).
cnf(c_0_272, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X2)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_253])).
cnf(c_0_273, plain, (r1_xxreal_0(k6_numbers,k17_complex1(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_254])).
cnf(c_0_274, plain, (v1_xreal_0(k16_complex1(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_255])).
cnf(c_0_275, plain, (k16_complex1(X1)=k17_complex1(X1)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_256])).
cnf(c_0_276, plain, (k3_complex1(X1)=k2_xcmplx_0(k6_numbers,X1)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_257, c_0_123])).
cnf(c_0_277, plain, (~v1_xxreal_0(X1)|~v2_xxreal_0(X1)|~v3_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_258])).
cnf(c_0_278, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(k4_complex1(X3))|r1_xxreal_0(k4_complex1(X3),k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xcmplx_0(X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_148])])).
cnf(c_0_279, plain, (k2_xcmplx_0(k6_numbers,k6_numbers)=k6_numbers), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_260]), c_0_152])])])).
fof(c_0_280, plain, ![X28, X29]:(~v1_xxreal_0(X28)|~v1_xxreal_0(X29)|(r1_xxreal_0(X28,X29)|r1_xxreal_0(X29,X28))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[connectedness_r1_xxreal_0])])).
cnf(c_0_281, plain, (v3_xxreal_0(k4_complex1(X1))|~v2_xxreal_0(k4_complex1(X2))|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_97])])).
cnf(c_0_282, plain, (v1_xboole_0(X2)|v3_xxreal_0(X1)|v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~r1_xxreal_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_261])).
cnf(c_0_283, plain, (r1_xxreal_0(k6_numbers,k1_comptrig(X1))|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_262])).
cnf(c_0_284, plain, (r1_xxreal_0(k32_sin_cos,k1_comptrig(X1))|r1_xxreal_0(k1_comptrig(X1),k6_numbers)|~r1_xxreal_0(k4_complex1(X1),k6_numbers)|~v1_xcmplx_0(X1)|~v1_xxreal_0(k32_sin_cos)|~v1_xxreal_0(k6_numbers)|~v1_xxreal_0(k1_comptrig(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_264])])).
cnf(c_0_285, negated_conjecture, (v3_xxreal_0(k4_complex1(esk2_0))|r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_266])])).
cnf(c_0_286, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_267, c_0_123])).
cnf(c_0_287, negated_conjecture, (~r1_xxreal_0(k32_sin_cos,k1_comptrig(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_214])).
cnf(c_0_288, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_269])])).
cnf(c_0_289, plain, (v1_xreal_0(k32_sin_cos)), inference(rw,[status(thm)],[c_0_270, c_0_271])).
cnf(c_0_290, plain, (v3_xxreal_0(k6_numbers)|~v3_xxreal_0(k17_complex1(X1))|~v1_xreal_0(k17_complex1(X1))|~v1_xreal_0(k6_numbers)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_273])])).
cnf(c_0_291, plain, (v1_xreal_0(k17_complex1(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_274, c_0_275])])).
cnf(c_0_292, plain, (k4_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1))))=k7_real_1(k4_complex1(X1),X3)|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_206])])).
cnf(c_0_293, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_238])])).
cnf(c_0_294, plain, (k2_xcmplx_0(k6_numbers,X1)=k2_xcmplx_0(X1,k6_numbers)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_276])])).
cnf(c_0_295, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_xxreal_0(X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_277, c_0_266])])).
cnf(c_0_296, plain, (v2_xxreal_0(k4_complex1(X1))|r1_xxreal_0(k4_complex1(X1),k6_numbers)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_279]), c_0_152])]), c_0_247])])).
cnf(c_0_297, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_280])).
cnf(c_0_298, negated_conjecture, (v3_xxreal_0(k4_complex1(esk1_0))|~v2_xxreal_0(k4_complex1(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_281, c_0_243]), c_0_245]), c_0_244])])])).
cnf(c_0_299, plain, (v2_xxreal_0(X2)|v1_xboole_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_282])).
cnf(c_0_300, plain, (r1_xxreal_0(k6_numbers,k1_comptrig(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_283])).
cnf(c_0_301, plain, (r1_xxreal_0(k1_comptrig(X1),k6_numbers)|r1_xxreal_0(k32_sin_cos,k1_comptrig(X1))|~r1_xxreal_0(k4_complex1(X1),k6_numbers)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_284, c_0_179]), c_0_178])]), c_0_180])).
cnf(c_0_302, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_286]), c_0_247])])).
cnf(c_0_303, negated_conjecture, (~r1_xxreal_0(k32_sin_cos,k1_comptrig(esk1_0))), inference(evalgc,[status(thm)],[c_0_287])).
cnf(c_0_304, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_289])])).
cnf(c_0_305, plain, (v3_xxreal_0(k6_numbers)|~v3_xxreal_0(k17_complex1(X1))|~v1_xreal_0(k6_numbers)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_290, c_0_291])).
fof(c_0_306, plain, ![X70]:((k1_comptrig(X70)!=k6_numbers|X70=k17_complex1(X70)|~v1_xcmplx_0(X70))&(X70!=k17_complex1(X70)|k1_comptrig(X70)=k6_numbers|~v1_xcmplx_0(X70))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t15_complex2])])])).
cnf(c_0_307, plain, (k3_complex1(k2_xcmplx_0(X1,X2))=k7_real_1(X1,X2)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_197])])).
cnf(c_0_308, plain, (k4_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k6_numbers)))=k7_real_1(k4_complex1(X1),k6_numbers)|~v1_xreal_0(X2)|~v1_xcmplx_0(k2_xcmplx_0(X2,k6_numbers))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_292, c_0_151]), c_0_152])])])).
cnf(c_0_309, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,k6_numbers))|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_293, c_0_294]), c_0_152])])])).
cnf(c_0_310, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_xxreal_0(X1)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_295, c_0_269])).
cnf(c_0_311, negated_conjecture, (v3_xxreal_0(k4_complex1(esk2_0))|r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_296]), c_0_245])])])).
cnf(c_0_312, negated_conjecture, (~r1_xxreal_0(k32_sin_cos,k1_comptrig(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_214])).
cnf(c_0_313, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_297])).
cnf(c_0_314, negated_conjecture, (v3_xxreal_0(k4_complex1(esk1_0))|r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_266])])).
cnf(c_0_315, plain, (v3_xxreal_0(k6_numbers)|v1_xboole_0(k1_comptrig(X1))|v2_xxreal_0(k1_comptrig(X1))|~v1_xreal_0(k1_comptrig(X1))|~v1_xreal_0(k6_numbers)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_299, c_0_300])])).
cnf(c_0_316, plain, (v1_xreal_0(k1_comptrig(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_143])])).
cnf(c_0_317, plain, (v3_xxreal_0(k6_numbers)|~v3_xxreal_0(k1_comptrig(X1))|~v1_xreal_0(k1_comptrig(X1))|~v1_xreal_0(k6_numbers)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_300])])).
cnf(c_0_318, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_301, c_0_302]), c_0_245])]), c_0_303])])).
cnf(c_0_319, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_304, c_0_269])])).
cnf(c_0_320, plain, (v3_xxreal_0(k6_numbers)|~v3_xxreal_0(k17_complex1(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_305, c_0_152])])])).
cnf(c_0_321, plain, (X1=k17_complex1(X1)|k1_comptrig(X1)!=k6_numbers|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_306])).
fof(c_0_322, plain, ![X79]:(~m1_subset_1(X79,k1_numbers)|(~r1_xxreal_0(k6_numbers,X79)|k1_comptrig(X79)=k6_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t35_comptrig])])).
cnf(c_0_323, plain, (k3_complex1(k2_xcmplx_0(X1,k6_numbers))=X1|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_151]), c_0_152])])])).
cnf(c_0_324, plain, (k3_complex1(k2_xcmplx_0(X1,X2))=k7_real_1(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(csr,[status(thm)],[c_0_307, c_0_123])).
cnf(c_0_325, plain, (k4_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k6_numbers)))=k7_real_1(k4_complex1(X1),k6_numbers)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_308, c_0_309])).
cnf(c_0_326, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_220]), c_0_221]), c_0_221]), c_0_157]), c_0_153])])])).
cnf(c_0_327, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_311])])).
cnf(c_0_328, negated_conjecture, (~r1_xxreal_0(k32_sin_cos,k1_comptrig(esk2_0))), inference(evalgc,[status(thm)],[c_0_312])).
cnf(c_0_329, plain, (r1_xxreal_0(X1,k1_comptrig(X2))|r1_xxreal_0(k1_comptrig(X2),X1)|~v1_xxreal_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_313, c_0_180])])).
fof(c_0_330, plain, ((((((((((~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k6_numbers)&~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__2)))&~r1_xxreal_0(k32_sin_cos,k6_numbers))&~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k1_real_1(k10_real_1(k32_sin_cos,np__2))))&~r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k32_sin_cos))&~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k10_real_1(k32_sin_cos,np__2)))&~r1_xxreal_0(k6_numbers,k1_real_1(k10_real_1(k32_sin_cos,np__2))))&~r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k6_numbers))&~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k32_sin_cos))&~r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos)))&~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k6_numbers)), inference(fof_simplification,[status(thm)],[t5_comptrig])).
cnf(c_0_331, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_314, c_0_286]), c_0_247])])).
cnf(c_0_332, plain, (v3_xxreal_0(k6_numbers)|v1_xboole_0(k1_comptrig(X1))|v2_xxreal_0(k1_comptrig(X1))|~v1_xreal_0(k6_numbers)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_315, c_0_316])).
fof(c_0_333, plain, ![X50, X51]:(v3_xxreal_0(X50)|~v1_xreal_0(X50)|v3_xxreal_0(X51)|~v1_xreal_0(X51)|~v3_xxreal_0(k2_xcmplx_0(X50,X51))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc9_xreal_0])])])).
cnf(c_0_334, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(k6_numbers)|r1_xxreal_0(k6_numbers,k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_259, c_0_152])])).
cnf(c_0_335, plain, (v3_xxreal_0(k6_numbers)|~v3_xxreal_0(k1_comptrig(X1))|~v1_xreal_0(k6_numbers)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_317, c_0_316])).
cnf(c_0_336, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k6_numbers)|~v1_xreal_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_148]), c_0_245])])])).
cnf(c_0_337, plain, (v3_xxreal_0(X1)|v1_xboole_0(X1)|v2_xxreal_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_299, c_0_319])])).
cnf(c_0_338, plain, (~v3_xxreal_0(k17_complex1(X1))|~v1_xcmplx_0(X1)), inference(sr,[status(thm)],[c_0_320, c_0_247])).
cnf(c_0_339, plain, (k17_complex1(X1)=X1|k1_comptrig(X1)!=k6_numbers|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_321])).
cnf(c_0_340, plain, (k1_comptrig(X1)=k6_numbers|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k6_numbers,X1)), inference(split_conjunct,[status(thm)],[c_0_322])).
cnf(c_0_341, plain, (k7_real_1(X1,k6_numbers)=X1|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_324]), c_0_152])])])).
cnf(c_0_342, plain, (k4_complex1(k2_xcmplx_0(X1,X2))=k7_real_1(k4_complex1(X1),k6_numbers)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_326])])).
fof(c_0_343, plain, ![X8, X9]:(~v3_membered(X8)|(~m1_subset_1(X9,X8)|v1_xreal_0(X9))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc14_membered])])])).
cnf(c_0_344, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_301, c_0_327]), c_0_244])]), c_0_328])])).
cnf(c_0_345, plain, (r1_xxreal_0(k1_comptrig(X1),k1_comptrig(X2))|r1_xxreal_0(k1_comptrig(X2),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_180])])).
cnf(c_0_346, plain, (~r1_xxreal_0(k32_sin_cos,k6_numbers)), inference(split_conjunct,[status(thm)],[c_0_330])).
cnf(c_0_347, plain, (k4_complex1(k2_xcmplx_0(X1,X2))=k4_complex1(k2_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_193])])).
cnf(c_0_348, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_301, c_0_331]), c_0_244])]), c_0_328])])).
cnf(c_0_349, plain, (v3_xxreal_0(k6_numbers)|v1_xboole_0(k1_comptrig(X1))|v2_xxreal_0(k1_comptrig(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_332, c_0_152])])])).
cnf(c_0_350, plain, (v3_xxreal_0(X1)|v3_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v3_xxreal_0(k2_xcmplx_0(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_333])).
cnf(c_0_351, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|r1_xxreal_0(k6_numbers,k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)), inference(sr,[status(thm)],[c_0_334, c_0_200])).
cnf(c_0_352, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_237]), c_0_221]), c_0_221]), c_0_157]), c_0_153])])])).
fof(c_0_353, plain, ![X31]:(~v1_xcmplx_0(X31)|m1_subset_1(k17_complex1(X31),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k17_complex1])])).
cnf(c_0_354, plain, (v3_xxreal_0(k6_numbers)|~v3_xxreal_0(k1_comptrig(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_335, c_0_152])])])).
cnf(c_0_355, negated_conjecture, (~v2_xxreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_336]), c_0_152])]), c_0_200])])).
cnf(c_0_356, plain, (X1=k6_numbers|v3_xxreal_0(X1)|v2_xxreal_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_337])])).
cnf(c_0_357, plain, (~v3_xxreal_0(X1)|~v1_xcmplx_0(X1)|k1_comptrig(X1)!=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_338, c_0_339])])).
cnf(c_0_358, plain, (k1_comptrig(X1)=k6_numbers|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k6_numbers,X1)), inference(evalgc,[status(thm)],[c_0_340])).
cnf(c_0_359, plain, (k4_complex1(k2_xcmplx_0(X1,X2))=k4_complex1(X1)|~v1_xreal_0(k4_complex1(X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_341, c_0_342])])).
cnf(c_0_360, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_343])).
cnf(c_0_361, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k6_numbers)|r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_301, c_0_344]), c_0_245])]), c_0_303])])).
cnf(c_0_362, negated_conjecture, (r1_xxreal_0(k1_comptrig(X1),k1_comptrig(esk1_0))|r1_xxreal_0(k1_comptrig(esk1_0),k1_comptrig(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_345, c_0_245])])).
cnf(c_0_363, plain, (v1_xcmplx_0(k32_sin_cos)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_289])])).
cnf(c_0_364, plain, (r1_xxreal_0(X1,k32_sin_cos)|r1_xxreal_0(k32_sin_cos,X1)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_313, c_0_179])])).
cnf(c_0_365, plain, (~r1_xxreal_0(k32_sin_cos,k6_numbers)), inference(evalgc,[status(thm)],[c_0_346])).
cnf(c_0_366, plain, (k4_complex1(k2_xcmplx_0(X2,X1))=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_347])])).
cnf(c_0_367, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|~v1_xreal_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_348, c_0_148]), c_0_244])])])).
cnf(c_0_368, plain, (v1_xboole_0(k1_comptrig(X1))|v2_xxreal_0(k1_comptrig(X1))|~v1_xcmplx_0(X1)), inference(sr,[status(thm)],[c_0_349, c_0_247])).
cnf(c_0_369, plain, (v3_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(k2_xcmplx_0(X1,X2))), inference(evalgc,[status(thm)],[c_0_350])).
cnf(c_0_370, plain, (v3_xxreal_0(X1)|r1_xxreal_0(k6_numbers,X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_351, c_0_352]), c_0_152])])])).
cnf(c_0_371, plain, (m1_subset_1(k17_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1)), inference(split_conjunct,[status(thm)],[c_0_353])).
cnf(c_0_372, plain, (~v3_xxreal_0(k1_comptrig(X1))|~v1_xcmplx_0(X1)), inference(sr,[status(thm)],[c_0_354, c_0_247])).
cnf(c_0_373, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers|v3_xxreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_355, c_0_356])])).
cnf(c_0_374, plain, (~v3_xxreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k6_numbers,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_357, c_0_358])])).
cnf(c_0_375, plain, (k4_complex1(k2_xcmplx_0(X1,X2))=k4_complex1(X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)), inference(csr,[status(thm)],[c_0_359, c_0_148])).
cnf(c_0_376, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_360])).
cnf(c_0_377, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|r1_xxreal_0(k1_comptrig(esk1_0),k6_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_148]), c_0_244])])])).
fof(c_0_378, plain, ![X22]:((v1_xxreal_0(X22)|(v1_xboole_0(X22)|~v1_xxreal_0(X22)|v3_xxreal_0(X22)))&(v2_xxreal_0(X22)|(v1_xboole_0(X22)|~v1_xxreal_0(X22)|v3_xxreal_0(X22)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc4_xxreal_0])])])])).
cnf(c_0_379, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k1_comptrig(k32_sin_cos))|r1_xxreal_0(k1_comptrig(k32_sin_cos),k1_comptrig(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_362, c_0_363])])).
cnf(c_0_380, plain, (r1_xxreal_0(k6_numbers,k32_sin_cos)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_364, c_0_178]), c_0_365])])).
cnf(c_0_381, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,k6_numbers))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_173]), c_0_153])])])).
cnf(c_0_382, negated_conjecture, (~v2_xxreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_367]), c_0_152])]), c_0_200])])).
cnf(c_0_383, plain, (k1_comptrig(X1)=k6_numbers|v2_xxreal_0(k1_comptrig(X1))|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_368])])).
cnf(c_0_384, negated_conjecture, (~m1_subset_1(k2_xcmplx_0(esk1_0,esk2_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k2_xcmplx_0(esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_358]), c_0_365])])).
cnf(c_0_385, plain, (v3_xxreal_0(X1)|v3_xxreal_0(X2)|r1_xxreal_0(k6_numbers,k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_369, c_0_370])])).
cnf(c_0_386, plain, (m1_subset_1(k17_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[c_0_371])).
cnf(c_0_387, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers|~v1_xreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_372, c_0_373]), c_0_245])])])).
cnf(c_0_388, plain, (~v3_xxreal_0(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k6_numbers,X1)), inference(csr,[status(thm)],[c_0_374, c_0_132])).
cnf(c_0_389, negated_conjecture, (v3_xxreal_0(k4_complex1(esk1_0))|r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_296]), c_0_244])])])).
cnf(c_0_390, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_375])])).
cnf(c_0_391, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_376, c_0_107]), c_0_165])])])).
cnf(c_0_392, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k6_numbers)|~v2_xxreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(k1_comptrig(esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_377]), c_0_152])]), c_0_200])])).
cnf(c_0_393, plain, (v2_xxreal_0(X1)|v1_xboole_0(X1)|v3_xxreal_0(X1)|~v1_xxreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_378])).
cnf(c_0_394, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_comptrig(esk1_0))|r1_xxreal_0(k1_comptrig(esk1_0),k6_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_379, c_0_358]), c_0_140]), c_0_380])])])).
cnf(c_0_395, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_381]), c_0_244])]), c_0_247])])).
cnf(c_0_396, negated_conjecture, (k1_comptrig(esk2_0)=k6_numbers|~v1_xreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_383]), c_0_244])])])).
cnf(c_0_397, negated_conjecture, (~r1_xxreal_0(k6_numbers,k2_xcmplx_0(esk1_0,esk2_0))|~v1_xreal_0(esk2_0)|~m1_subset_1(esk1_0,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_384, c_0_107])])).
cnf(c_0_398, plain, (v3_xxreal_0(X1)|v3_xxreal_0(X2)|r1_xxreal_0(k6_numbers,k2_xcmplx_0(X2,X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)), inference(csr,[status(thm)],[c_0_385, c_0_238])).
cnf(c_0_399, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|k1_comptrig(X1)!=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_386, c_0_339])])).
cnf(c_0_400, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers|~v1_xreal_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_316]), c_0_245])])])).
cnf(c_0_401, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|k1_comptrig(X1)!=k6_numbers), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_339])])).
cnf(c_0_402, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|~m1_subset_1(k4_complex1(esk1_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_388, c_0_389])])).
cnf(c_0_403, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_390, c_0_391])])).
cnf(c_0_404, plain, (r1_xxreal_0(np__0,np__0)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r0_r0])).
cnf(c_0_405, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k6_numbers)|k1_comptrig(esk2_0)=k6_numbers|~v1_xreal_0(k1_comptrig(esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_383]), c_0_244])])])).
cnf(c_0_406, plain, (v2_xxreal_0(X1)|v1_xboole_0(X1)|v3_xxreal_0(X1)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_393])).
cnf(c_0_407, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_comptrig(esk1_0))|~v2_xxreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_394]), c_0_152])]), c_0_200])])).
cnf(c_0_408, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k6_numbers)|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_301, c_0_395]), c_0_245])]), c_0_303])])).
cnf(c_0_409, negated_conjecture, (k1_comptrig(esk2_0)=k6_numbers|~v1_xreal_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_316]), c_0_244])])])).
cnf(c_0_410, negated_conjecture, (v3_xxreal_0(esk2_0)|v3_xxreal_0(esk1_0)|~v1_xreal_0(esk2_0)|~m1_subset_1(esk1_0,k1_numbers)|~v1_xreal_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_397, c_0_398])])).
cnf(c_0_411, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers)|~v1_xreal_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_399, c_0_400]), c_0_245])])])).
cnf(c_0_412, negated_conjecture, (v1_xreal_0(esk1_0)|~v1_xreal_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_400]), c_0_245])])])).
cnf(c_0_413, negated_conjecture, (~v3_xxreal_0(esk1_0)|~v1_xreal_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_357, c_0_400]), c_0_245])])])).
cnf(c_0_414, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_94]), c_0_245])])])).
cnf(c_0_415, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)), inference(csr,[status(thm)],[c_0_403, c_0_132])).
cnf(c_0_416, plain, (r1_xxreal_0(np__0,np__0)), inference(evalgc,[status(thm)],[c_0_404])).
cnf(c_0_417, negated_conjecture, (k1_comptrig(esk2_0)=k6_numbers|~v2_xxreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_405]), c_0_152])]), c_0_200])])).
cnf(c_0_418, plain, (X1=k6_numbers|v3_xxreal_0(X1)|v2_xxreal_0(X1)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_406])])).
cnf(c_0_419, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_comptrig(esk1_0))|k1_comptrig(esk1_0)=k6_numbers|~v1_xreal_0(k1_comptrig(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_407, c_0_383]), c_0_245])])])).
cnf(c_0_420, negated_conjecture, (~v2_xxreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_408]), c_0_152])]), c_0_200])])).
cnf(c_0_421, negated_conjecture, (~v3_xxreal_0(esk2_0)|~v1_xreal_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_357, c_0_409]), c_0_244])])])).
cnf(c_0_422, negated_conjecture, (v3_xxreal_0(esk2_0)|~v1_xreal_0(esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_410, c_0_411]), c_0_412]), c_0_413])).
cnf(c_0_423, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_301, c_0_414]), c_0_244])]), c_0_328])])).
cnf(c_0_424, plain, (k4_complex1(X1)=k6_numbers|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_415, c_0_152])])).
cnf(c_0_425, plain, (r1_xxreal_0(k6_numbers,k6_numbers)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_416, c_0_125]), c_0_125])])).
cnf(c_0_426, negated_conjecture, (k1_comptrig(esk2_0)=k6_numbers|k1_comptrig(esk1_0)=k6_numbers|v3_xxreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk2_0))|~v1_xxreal_0(k1_comptrig(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_417, c_0_418])])).
cnf(c_0_427, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers|~v3_xxreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_419]), c_0_152])]), c_0_247])])).
cnf(c_0_428, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers))|~v1_xreal_0(k1_comptrig(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_420, c_0_383]), c_0_245])])])).
cnf(c_0_429, negated_conjecture, (~v1_xreal_0(esk1_0)|~v1_xreal_0(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_421, c_0_422])])).
cnf(c_0_430, negated_conjecture, (v1_xreal_0(esk2_0)|~v1_xreal_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_409]), c_0_244])])])).
cnf(c_0_431, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|~m1_subset_1(esk1_0,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_423, c_0_424]), c_0_425])])])).
cnf(c_0_432, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers|k1_comptrig(esk2_0)=k6_numbers|~v1_xreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk2_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_426, c_0_269]), c_0_427])).
cnf(c_0_433, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_428, c_0_316]), c_0_245])])])).
cnf(c_0_434, negated_conjecture, (~v1_xreal_0(esk1_0)), inference(csr,[status(thm)],[c_0_429, c_0_430])).
cnf(c_0_435, negated_conjecture, (~v2_xxreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(k1_comptrig(esk2_0))|~m1_subset_1(esk1_0,k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_431]), c_0_152])]), c_0_200])])).
cnf(c_0_436, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers|k1_comptrig(esk2_0)=k6_numbers|~v1_xreal_0(k1_comptrig(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_432, c_0_316]), c_0_244])])])).
cnf(c_0_437, negated_conjecture, (~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_433]), c_0_245])]), c_0_434])])).
cnf(c_0_438, negated_conjecture, (k1_comptrig(esk2_0)=k6_numbers|~v1_xreal_0(k1_comptrig(esk2_0))|~m1_subset_1(esk1_0,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_435, c_0_383]), c_0_244])])])).
cnf(c_0_439, negated_conjecture, (k1_comptrig(esk2_0)=k6_numbers|k1_comptrig(esk1_0)=k6_numbers), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_436, c_0_316]), c_0_245])])])).
cnf(c_0_440, negated_conjecture, (~m1_subset_1(esk2_0,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_437, c_0_391]), c_0_152])])])).
cnf(c_0_441, negated_conjecture, (k1_comptrig(esk2_0)=k6_numbers|~m1_subset_1(esk1_0,k1_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_438, c_0_316]), c_0_244])])])).
cnf(c_0_442, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_399, c_0_439]), c_0_244])]), c_0_440])])).
cnf(c_0_443, negated_conjecture, (~m1_subset_1(esk1_0,k1_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_399, c_0_441]), c_0_244])]), c_0_440])])).
cnf(c_0_444, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_399, c_0_442]), c_0_245])]), c_0_443]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 445
# Proof object clause steps            : 343
# Proof object formula steps           : 102
# Proof object conjectures             : 77
# Proof object clause conjectures      : 74
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 68
# Proof object initial formulas used   : 56
# Proof object generating inferences   : 153
# Proof object simplifying inferences  : 257
# Training examples: 232 positive, 1091 negative
# Training: Positive examples begin
cnf(c_0_271, plain, (k31_sin_cos=k32_sin_cos)).# trainpos
cnf(c_0_81, plain, (k1_xboole_0=k6_numbers)).# trainpos
cnf(c_0_100, plain, (k1_xcmplx_0=k7_complex1)).# trainpos
cnf(c_0_245, negated_conjecture, (v1_xcmplx_0(esk1_0))).# trainpos
cnf(c_0_244, negated_conjecture, (v1_xcmplx_0(esk2_0))).# trainpos
cnf(c_0_128, plain, (v1_xcmplx_0(esk4_0))).# trainpos
cnf(c_0_161, plain, (v1_xxreal_0(esk5_0))).# trainpos
cnf(c_0_165, plain, (v3_membered(k1_numbers))).# trainpos
cnf(c_0_303, negated_conjecture, (~r1_xxreal_0(k32_sin_cos,k1_comptrig(esk1_0)))).# trainpos
cnf(c_0_328, negated_conjecture, (~r1_xxreal_0(k32_sin_cos,k1_comptrig(esk2_0)))).# trainpos
cnf(c_0_270, plain, (v1_xreal_0(k31_sin_cos))).# trainpos
cnf(c_0_126, plain, (v1_xreal_0(esk4_0))).# trainpos
cnf(c_0_365, plain, (~r1_xxreal_0(k32_sin_cos,k6_numbers))).# trainpos
cnf(c_0_164, plain, (v2_membered(X1)|~v3_membered(X1))).# trainpos
cnf(c_0_269, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_168, plain, (v1_xboole_0(k1_xboole_0))).# trainpos
cnf(c_0_102, plain, (v1_xboole_0(np__0))).# trainpos
cnf(c_0_104, plain, (v1_xboole_0(esk4_0))).# trainpos
cnf(c_0_138, plain, (v1_xboole_0(esk5_0))).# trainpos
cnf(c_0_99, plain, (k3_xcmplx_0(np__0,k1_xcmplx_0)=np__0)).# trainpos
cnf(c_0_416, plain, (r1_xxreal_0(np__0,np__0))).# trainpos
cnf(c_0_140, plain, (m1_subset_1(k32_sin_cos,k1_numbers))).# trainpos
cnf(c_0_133, plain, (m1_subset_1(np__0,k1_numbers))).# trainpos
cnf(c_0_80, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_123, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_243, negated_conjecture, (r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0))))).# trainpos
cnf(c_0_132, plain, (v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_406, plain, (v2_xxreal_0(X1)|v1_xboole_0(X1)|v3_xxreal_0(X1)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_224, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_275, plain, (k16_complex1(X1)=k17_complex1(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_274, plain, (v1_xreal_0(k16_complex1(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_129, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_120, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_300, plain, (r1_xxreal_0(k6_numbers,k1_comptrig(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_273, plain, (r1_xxreal_0(k6_numbers,k17_complex1(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_143, plain, (m1_subset_1(k1_comptrig(X1),k1_numbers)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_386, plain, (m1_subset_1(k17_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_94, plain, (m1_subset_1(k4_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_339, plain, (k17_complex1(X1)=X1|k1_comptrig(X1)!=k6_numbers|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_181, plain, (~v1_xxreal_0(X1)|~v2_xxreal_0(X1)|~v1_xboole_0(X1))).# trainpos
cnf(c_0_277, plain, (~v1_xxreal_0(X1)|~v2_xxreal_0(X1)|~v3_xxreal_0(X1))).# trainpos
cnf(c_0_216, plain, (~v1_xxreal_0(X1)|~v1_xboole_0(X1)|~v3_xxreal_0(X1))).# trainpos
cnf(c_0_238, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_268, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_139, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_376, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_358, plain, (k1_comptrig(X1)=k6_numbers|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k6_numbers,X1))).# trainpos
cnf(c_0_97, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_135, plain, (k4_xxreal_1(X1,X2)=k2_rcomp_1(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_86, plain, (k7_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_71, plain, (k7_real_1(X1,X2)=k7_real_1(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_263, plain, (~v1_xcmplx_0(X1)|~r1_xxreal_0(k4_complex1(X1),k6_numbers)|~r2_hidden(k1_comptrig(X1),k2_rcomp_1(k6_numbers,k32_sin_cos)))).# trainpos
cnf(c_0_313, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_160, plain, (r1_xxreal_0(k4_complex1(X1),k6_numbers)|r2_hidden(k1_comptrig(X1),k2_rcomp_1(k6_numbers,k32_sin_cos))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_134, plain, (~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~r2_hidden(X2,k4_xxreal_1(X3,X1)))).# trainpos
cnf(c_0_228, plain, (v2_xxreal_0(X1)|v3_xxreal_0(X2)|r1_xxreal_0(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_189, plain, (v2_xxreal_0(X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_272, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(X2)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_169, plain, (v3_xxreal_0(X2)|v2_xxreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v2_xxreal_0(X1))).# trainpos
cnf(c_0_207, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_299, plain, (v2_xxreal_0(X2)|v1_xboole_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_85, plain, (m1_subset_1(k7_real_1(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_206, plain, (k4_complex1(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)))=X2|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_96, plain, (k3_complex1(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)))=X1|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_72, plain, (k7_real_1(k4_complex1(X1),k4_complex1(X2))=k4_complex1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_155, plain, (k7_real_1(k3_complex1(X1),k3_complex1(X2))=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_369, plain, (v3_xxreal_0(X2)|v3_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(k2_xcmplx_0(X1,X2)))).# trainpos
cnf(c_0_241, plain, (r1_xxreal_0(X3,X1)|r1_xxreal_0(X1,X2)|r2_hidden(X1,k4_xxreal_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_125, plain, (np__0=k6_numbers)).# trainpos
cnf(c_0_363, plain, (v1_xcmplx_0(k32_sin_cos))).# trainpos
cnf(c_0_162, plain, (esk5_0=k6_numbers)).# trainpos
cnf(c_0_127, plain, (esk4_0=k6_numbers)).# trainpos
cnf(c_0_191, plain, (~v2_xxreal_0(k6_numbers)|~v1_xxreal_0(k6_numbers))).# trainpos
cnf(c_0_229, plain, (~v3_xxreal_0(k6_numbers)|~v1_xxreal_0(k6_numbers))).# trainpos
cnf(c_0_316, plain, (v1_xreal_0(k1_comptrig(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_418, plain, (X1=k6_numbers|v3_xxreal_0(X1)|v2_xxreal_0(X1)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_384, negated_conjecture, (~m1_subset_1(k2_xcmplx_0(esk1_0,esk2_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k2_xcmplx_0(esk1_0,esk2_0)))).# trainpos
cnf(c_0_399, plain, (m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|k1_comptrig(X1)!=k6_numbers)).# trainpos
cnf(c_0_107, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_174, plain, (k2_xcmplx_0(k3_complex1(X1),k3_complex1(X2))=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_121, plain, (k7_real_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_154, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k6_numbers))).# trainpos
cnf(c_0_291, plain, (v1_xreal_0(k17_complex1(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_163, plain, (v1_xxreal_0(k32_sin_cos)|~v2_membered(k1_numbers))).# trainpos
cnf(c_0_93, plain, (k7_real_1(k4_complex1(X2),k4_complex1(X1))=k4_complex1(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_complex1(X2))|~m1_subset_1(k4_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_264, plain, (r2_hidden(X1,k2_rcomp_1(X2,X3))|r1_xxreal_0(X1,X2)|r1_xxreal_0(X3,X1)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_148, plain, (v1_xreal_0(k4_complex1(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_159, plain, (~r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v1_xxreal_0(X3)|~r1_xxreal_0(X3,X1))).# trainpos
cnf(c_0_293, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_288, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_315, plain, (v3_xxreal_0(k6_numbers)|v1_xboole_0(k1_comptrig(X1))|v2_xxreal_0(k1_comptrig(X1))|~v1_xreal_0(k1_comptrig(X1))|~v1_xreal_0(k6_numbers)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_166, plain, (v1_xxreal_0(k1_comptrig(X1))|~v2_membered(k1_numbers)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_290, plain, (v3_xxreal_0(k6_numbers)|~v3_xxreal_0(k17_complex1(X1))|~v1_xreal_0(k17_complex1(X1))|~v1_xreal_0(k6_numbers)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_317, plain, (v3_xxreal_0(k6_numbers)|~v3_xxreal_0(k1_comptrig(X1))|~v1_xreal_0(k1_comptrig(X1))|~v1_xreal_0(k6_numbers)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_259, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(X3)|r1_xxreal_0(X3,k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_183, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k2_xcmplx_0(X2,X1))|~v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_292, plain, (k4_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1))))=k7_real_1(k4_complex1(X1),X3)|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_152, plain, (v1_xreal_0(k6_numbers))).# trainpos
cnf(c_0_179, plain, (v1_xxreal_0(k32_sin_cos))).# trainpos
cnf(c_0_178, plain, (v1_xxreal_0(k6_numbers))).# trainpos
cnf(c_0_397, negated_conjecture, (~r1_xxreal_0(k6_numbers,k2_xcmplx_0(esk1_0,esk2_0))|~v1_xreal_0(esk2_0)|~m1_subset_1(esk1_0,k1_numbers))).# trainpos
cnf(c_0_425, plain, (r1_xxreal_0(k6_numbers,k6_numbers))).# trainpos
cnf(c_0_180, plain, (v1_xxreal_0(k1_comptrig(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_153, plain, (v1_xcmplx_0(k6_numbers))).# trainpos
cnf(c_0_157, plain, (m1_subset_1(k6_numbers,k1_numbers))).# trainpos
cnf(c_0_151, plain, (k3_xcmplx_0(k6_numbers,k7_complex1)=k6_numbers)).# trainpos
cnf(c_0_304, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_177, plain, (r1_xxreal_0(k4_complex1(X1),k6_numbers)|~v1_xxreal_0(k6_numbers)|~v1_xxreal_0(k1_comptrig(X1))|~v1_xxreal_0(k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k1_comptrig(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_217, plain, (k4_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1))))=X3|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X3,k7_complex1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_401, plain, (v1_xreal_0(X1)|~v1_xcmplx_0(X1)|k1_comptrig(X1)!=k6_numbers)).# trainpos
cnf(c_0_147, plain, (k2_xcmplx_0(k4_complex1(X1),k4_complex1(X2))=k4_complex1(k2_xcmplx_0(X2,X1))|~v1_xreal_0(k4_complex1(X2))|~m1_subset_1(k4_complex1(X1),k1_numbers)|~v1_xreal_0(k4_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_149, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).# trainpos
cnf(c_0_122, plain, (k3_complex1(k2_xcmplx_0(k3_xcmplx_0(X2,k7_complex1),X1))=X1|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,k7_complex1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_284, plain, (r1_xxreal_0(k32_sin_cos,k1_comptrig(X1))|r1_xxreal_0(k1_comptrig(X1),k6_numbers)|~r1_xxreal_0(k4_complex1(X1),k6_numbers)|~v1_xcmplx_0(X1)|~v1_xxreal_0(k32_sin_cos)|~v1_xxreal_0(k6_numbers)|~v1_xxreal_0(k1_comptrig(X1)))).# trainpos
cnf(c_0_131, plain, (m1_subset_1(k2_xcmplx_0(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_175, plain, (m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_184, plain, (k2_xcmplx_0(k4_complex1(X2),k4_complex1(X1))=k4_complex1(k2_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(k4_complex1(X2),k1_numbers)|~m1_subset_1(k4_complex1(X1),k1_numbers))).# trainpos
cnf(c_0_173, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_319, plain, (r1_xxreal_0(X1,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_247, plain, (~v3_xxreal_0(k6_numbers))).# trainpos
cnf(c_0_200, plain, (~v2_xxreal_0(k6_numbers))).# trainpos
cnf(c_0_364, plain, (r1_xxreal_0(X1,k32_sin_cos)|r1_xxreal_0(k32_sin_cos,X1)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_230, plain, (k4_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k6_numbers)))=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_187, plain, (k2_xcmplx_0(k3_complex1(X2),k3_complex1(X1))=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(k3_complex1(X2),k1_numbers))).# trainpos
cnf(c_0_172, plain, (k3_complex1(k2_xcmplx_0(k6_numbers,X1))=X1|~v1_xreal_0(X1))).# trainpos
cnf(c_0_354, plain, (v3_xxreal_0(k6_numbers)|~v3_xxreal_0(k1_comptrig(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_323, plain, (k3_complex1(k2_xcmplx_0(X1,k6_numbers))=X1|~v1_xreal_0(X1))).# trainpos
cnf(c_0_320, plain, (v3_xxreal_0(k6_numbers)|~v3_xxreal_0(k17_complex1(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_199, plain, (v2_xxreal_0(k6_numbers)|~v2_xxreal_0(k4_complex1(X1))|~v1_xreal_0(k4_complex1(X1))|~r1_xxreal_0(k32_sin_cos,k1_comptrig(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_329, plain, (r1_xxreal_0(X1,k1_comptrig(X2))|r1_xxreal_0(k1_comptrig(X2),X1)|~v1_xxreal_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_186, plain, (k3_complex1(X1)=X1|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_337, plain, (v3_xxreal_0(X1)|v1_xboole_0(X1)|v2_xxreal_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_221, plain, (k3_complex1(k6_numbers)=k6_numbers)).# trainpos
cnf(c_0_380, plain, (r1_xxreal_0(k6_numbers,k32_sin_cos))).# trainpos
cnf(c_0_248, plain, (k4_complex1(k2_xcmplx_0(X1,X2))=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_391, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_357, plain, (~v3_xxreal_0(X1)|~v1_xcmplx_0(X1)|k1_comptrig(X1)!=k6_numbers)).# trainpos
cnf(c_0_345, plain, (r1_xxreal_0(k1_comptrig(X1),k1_comptrig(X2))|r1_xxreal_0(k1_comptrig(X2),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_267, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_208, plain, (k7_real_1(k3_complex1(X1),X2)=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).# trainpos
cnf(c_0_356, plain, (X1=k6_numbers|v3_xxreal_0(X1)|v2_xxreal_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_374, plain, (~v3_xxreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k6_numbers,X1))).# trainpos
cnf(c_0_210, plain, (k3_complex1(k2_xcmplx_0(X2,X1))=k2_xcmplx_0(X1,k3_complex1(X2))|~v1_xreal_0(X1)|~m1_subset_1(k3_complex1(X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_362, negated_conjecture, (r1_xxreal_0(k1_comptrig(X1),k1_comptrig(esk1_0))|r1_xxreal_0(k1_comptrig(esk1_0),k1_comptrig(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_201, plain, (v3_xxreal_0(k4_complex1(X1))|v2_xxreal_0(k4_complex1(k2_xcmplx_0(X2,X1)))|~v2_xxreal_0(k4_complex1(X2))|~v1_xreal_0(k4_complex1(X1))|~v1_xreal_0(k4_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_347, plain, (k4_complex1(k2_xcmplx_0(X1,X2))=k4_complex1(k2_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_307, plain, (k3_complex1(k2_xcmplx_0(X1,X2))=k7_real_1(X1,X2)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_326, plain, (k2_xcmplx_0(X1,k6_numbers)=X1|~v1_xreal_0(X1))).# trainpos
cnf(c_0_239, plain, (k3_complex1(X1)=k2_xcmplx_0(X1,k6_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_223, plain, (v3_xxreal_0(k4_complex1(X2))|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v2_xxreal_0(k4_complex1(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_349, plain, (v3_xxreal_0(k6_numbers)|v1_xboole_0(k1_comptrig(X1))|v2_xxreal_0(k1_comptrig(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_237, plain, (k3_complex1(k2_xcmplx_0(X1,X2))=k2_xcmplx_0(k3_complex1(X1),X2)|~v1_xreal_0(X2)|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_366, plain, (k4_complex1(k2_xcmplx_0(X2,X1))=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_279, plain, (k2_xcmplx_0(k6_numbers,k6_numbers)=k6_numbers)).# trainpos
cnf(c_0_383, plain, (k1_comptrig(X1)=k6_numbers|v2_xxreal_0(k1_comptrig(X1))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_257, plain, (k3_complex1(X1)=k2_xcmplx_0(k6_numbers,X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_341, plain, (k7_real_1(X1,k6_numbers)=X1|~v1_xreal_0(X1))).# trainpos
cnf(c_0_379, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k1_comptrig(k32_sin_cos))|r1_xxreal_0(k1_comptrig(k32_sin_cos),k1_comptrig(esk1_0)))).# trainpos
cnf(c_0_352, plain, (k2_xcmplx_0(k6_numbers,X1)=X1|~v1_xreal_0(X1))).# trainpos
cnf(c_0_381, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,k6_numbers))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_394, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_comptrig(esk1_0))|r1_xxreal_0(k1_comptrig(esk1_0),k6_numbers))).# trainpos
cnf(c_0_265, negated_conjecture, (v3_xxreal_0(k4_complex1(esk2_0))|~v2_xxreal_0(k4_complex1(esk1_0)))).# trainpos
cnf(c_0_281, plain, (v3_xxreal_0(k4_complex1(X1))|~v2_xxreal_0(k4_complex1(X2))|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).# trainpos
cnf(c_0_298, negated_conjecture, (v3_xxreal_0(k4_complex1(esk1_0))|~v2_xxreal_0(k4_complex1(esk2_0)))).# trainpos
cnf(c_0_407, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_comptrig(esk1_0))|~v2_xxreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk1_0)))).# trainpos
cnf(c_0_246, plain, (v3_xxreal_0(k6_numbers)|v2_xxreal_0(X1)|r1_xxreal_0(X1,k6_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_295, plain, (r1_xxreal_0(X1,k6_numbers)|~v3_xxreal_0(X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_285, negated_conjecture, (v3_xxreal_0(k4_complex1(esk2_0))|r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk1_0)))).# trainpos
cnf(c_0_314, negated_conjecture, (v3_xxreal_0(k4_complex1(esk1_0))|r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0)))).# trainpos
cnf(c_0_302, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(esk2_0))).# trainpos
cnf(c_0_331, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(esk1_0))).# trainpos
cnf(c_0_294, plain, (k2_xcmplx_0(k6_numbers,X1)=k2_xcmplx_0(X1,k6_numbers)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_318, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(esk2_0))).# trainpos
cnf(c_0_348, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(esk1_0))).# trainpos
cnf(c_0_336, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k6_numbers)|~v1_xreal_0(esk2_0))).# trainpos
cnf(c_0_309, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,k6_numbers))|~v1_xreal_0(X1))).# trainpos
cnf(c_0_367, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|~v1_xreal_0(esk1_0))).# trainpos
cnf(c_0_355, negated_conjecture, (~v2_xxreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(esk2_0))).# trainpos
cnf(c_0_382, negated_conjecture, (~v2_xxreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(esk1_0))).# trainpos
cnf(c_0_373, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers|v3_xxreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(esk2_0))).# trainpos
cnf(c_0_387, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers|~v1_xreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(esk2_0))).# trainpos
cnf(c_0_400, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers|~v1_xreal_0(esk2_0))).# trainpos
cnf(c_0_413, negated_conjecture, (~v3_xxreal_0(esk1_0)|~v1_xreal_0(esk2_0))).# trainpos
cnf(c_0_412, negated_conjecture, (v1_xreal_0(esk1_0)|~v1_xreal_0(esk2_0))).# trainpos
cnf(c_0_411, negated_conjecture, (m1_subset_1(esk1_0,k1_numbers)|~v1_xreal_0(esk2_0))).# trainpos
cnf(c_0_278, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(k4_complex1(X3))|r1_xxreal_0(k4_complex1(X3),k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xcmplx_0(X3))).# trainpos
cnf(c_0_296, plain, (v2_xxreal_0(k4_complex1(X1))|r1_xxreal_0(k4_complex1(X1),k6_numbers)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_311, negated_conjecture, (v3_xxreal_0(k4_complex1(esk2_0))|r1_xxreal_0(k4_complex1(esk1_0),k6_numbers))).# trainpos
cnf(c_0_389, negated_conjecture, (v3_xxreal_0(k4_complex1(esk1_0))|r1_xxreal_0(k4_complex1(esk2_0),k6_numbers))).# trainpos
cnf(c_0_327, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0)))).# trainpos
cnf(c_0_344, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0)))).# trainpos
cnf(c_0_334, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(k6_numbers)|r1_xxreal_0(k6_numbers,k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_370, plain, (v3_xxreal_0(X1)|r1_xxreal_0(k6_numbers,X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_361, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k6_numbers)|r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0)))).# trainpos
cnf(c_0_395, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers)))).# trainpos
cnf(c_0_377, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|r1_xxreal_0(k1_comptrig(esk1_0),k6_numbers))).# trainpos
cnf(c_0_408, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k6_numbers)|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers)))).# trainpos
cnf(c_0_385, plain, (v3_xxreal_0(X1)|v3_xxreal_0(X2)|r1_xxreal_0(k6_numbers,k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2)))).# trainpos
cnf(c_0_410, negated_conjecture, (v3_xxreal_0(esk2_0)|v3_xxreal_0(esk1_0)|~v1_xreal_0(esk2_0)|~m1_subset_1(esk1_0,k1_numbers)|~v1_xreal_0(esk1_0))).# trainpos
cnf(c_0_420, negated_conjecture, (~v2_xxreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers)))).# trainpos
cnf(c_0_428, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers))|~v1_xreal_0(k1_comptrig(esk1_0)))).# trainpos
cnf(c_0_402, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|~m1_subset_1(k4_complex1(esk1_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0)))).# trainpos
cnf(c_0_396, negated_conjecture, (k1_comptrig(esk2_0)=k6_numbers|~v1_xreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(esk1_0))).# trainpos
cnf(c_0_409, negated_conjecture, (k1_comptrig(esk2_0)=k6_numbers|~v1_xreal_0(esk1_0))).# trainpos
cnf(c_0_414, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0)))).# trainpos
cnf(c_0_430, negated_conjecture, (v1_xreal_0(esk2_0)|~v1_xreal_0(esk1_0))).# trainpos
cnf(c_0_421, negated_conjecture, (~v3_xxreal_0(esk2_0)|~v1_xreal_0(esk1_0))).# trainpos
cnf(c_0_423, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0)))).# trainpos
cnf(c_0_429, negated_conjecture, (~v1_xreal_0(esk1_0)|~v1_xreal_0(esk2_0))).# trainpos
cnf(c_0_433, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers)))).# trainpos
cnf(c_0_437, negated_conjecture, (~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers)))).# trainpos
cnf(c_0_440, negated_conjecture, (~m1_subset_1(esk2_0,k1_numbers))).# trainpos
cnf(c_0_419, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_comptrig(esk1_0))|k1_comptrig(esk1_0)=k6_numbers|~v1_xreal_0(k1_comptrig(esk1_0)))).# trainpos
cnf(c_0_427, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers|~v3_xxreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk1_0)))).# trainpos
cnf(c_0_308, plain, (k4_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k6_numbers)))=k7_real_1(k4_complex1(X1),k6_numbers)|~v1_xreal_0(X2)|~v1_xcmplx_0(k2_xcmplx_0(X2,k6_numbers))|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_342, plain, (k4_complex1(k2_xcmplx_0(X1,X2))=k7_real_1(k4_complex1(X1),k6_numbers)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_359, plain, (k4_complex1(k2_xcmplx_0(X1,X2))=k4_complex1(X1)|~v1_xreal_0(k4_complex1(X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_390, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).# trainpos
cnf(c_0_403, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_424, plain, (k4_complex1(X1)=k6_numbers|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_431, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|~m1_subset_1(esk1_0,k1_numbers))).# trainpos
cnf(c_0_435, negated_conjecture, (~v2_xxreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(k1_comptrig(esk2_0))|~m1_subset_1(esk1_0,k1_numbers))).# trainpos
cnf(c_0_392, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k6_numbers)|~v2_xxreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(k1_comptrig(esk2_0)))).# trainpos
cnf(c_0_405, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k6_numbers)|k1_comptrig(esk2_0)=k6_numbers|~v1_xreal_0(k1_comptrig(esk2_0)))).# trainpos
cnf(c_0_438, negated_conjecture, (k1_comptrig(esk2_0)=k6_numbers|~v1_xreal_0(k1_comptrig(esk2_0))|~m1_subset_1(esk1_0,k1_numbers))).# trainpos
cnf(c_0_441, negated_conjecture, (k1_comptrig(esk2_0)=k6_numbers|~m1_subset_1(esk1_0,k1_numbers))).# trainpos
cnf(c_0_443, negated_conjecture, (~m1_subset_1(esk1_0,k1_numbers))).# trainpos
cnf(c_0_417, negated_conjecture, (k1_comptrig(esk2_0)=k6_numbers|~v2_xxreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk2_0)))).# trainpos
cnf(c_0_426, negated_conjecture, (k1_comptrig(esk2_0)=k6_numbers|k1_comptrig(esk1_0)=k6_numbers|v3_xxreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk2_0))|~v1_xxreal_0(k1_comptrig(esk1_0)))).# trainpos
cnf(c_0_436, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers|k1_comptrig(esk2_0)=k6_numbers|~v1_xreal_0(k1_comptrig(esk1_0)))).# trainpos
cnf(c_0_439, negated_conjecture, (k1_comptrig(esk2_0)=k6_numbers|k1_comptrig(esk1_0)=k6_numbers)).# trainpos
cnf(c_0_442, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers)).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_59, plain, (k4_ordinal1=k5_numbers)).#trainneg
cnf(i_0_66, plain, (k4_xcmplx_0(np__0)=np__0)).#trainneg
cnf(i_0_49, plain, (v1_xcmplx_0(esk3_0))).#trainneg
cnf(i_0_51, plain, (v1_xxreal_0(esk4_0))).#trainneg
cnf(i_0_75, plain, (k1_comptrig(X1)=k6_numbers|k17_complex1(X1)!=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_70, plain, (m2_subset_1(np__0,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_101, plain, (~r1_xxreal_0(k32_sin_cos,k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_47, plain, (v6_membered(k4_ordinal1))).#trainneg
cnf(i_0_80, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_69, plain, (m1_subset_1(np__0,k5_numbers))).#trainneg
cnf(i_0_13, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_98, plain, (~r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_95, plain, (~r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k6_numbers))).#trainneg
cnf(i_0_10, plain, (v5_membered(X1)|~v6_membered(X1))).#trainneg
cnf(i_0_102, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k6_numbers))).#trainneg
cnf(i_0_96, plain, (~r1_xxreal_0(k6_numbers,k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_94, plain, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k32_sin_cos))).#trainneg
cnf(i_0_105, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_43, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_42, plain, (v1_int_1(k4_xcmplx_0(X1))|~v1_int_1(X1))).#trainneg
cnf(i_0_82, plain, (k3_xcmplx_0(X1,k6_numbers)=k6_numbers|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_92, plain, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k6_numbers))).#trainneg
cnf(i_0_84, plain, (~v1_xcmplx_0(X1)|~r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k1_comptrig(X1)))).#trainneg
cnf(i_0_99, plain, (~r1_xxreal_0(k10_real_1(k32_sin_cos,np__2),k1_real_1(k10_real_1(k32_sin_cos,np__2))))).#trainneg
cnf(i_0_93, plain, (~r1_xxreal_0(k8_real_1(np__2,k32_sin_cos),k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos)))).#trainneg
cnf(i_0_97, plain, (~r1_xxreal_0(k8_real_1(k10_real_1(np__3,np__2),k32_sin_cos),k10_real_1(k32_sin_cos,np__2)))).#trainneg
cnf(i_0_38, plain, (v3_xxreal_0(X1)|v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_37, plain, (v3_xxreal_0(X1)|~v1_xreal_0(X1)|~v2_xxreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_8, plain, (v1_int_1(X2)|~v5_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_9, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_45, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_31, plain, (m1_subset_1(k2_rcomp_1(X1,X2),k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_91, plain, (~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2)|~r2_hidden(X1,k4_xxreal_1(X2,X3)))).#trainneg
cnf(i_0_88, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_40, plain, (v2_xxreal_0(X2)|v2_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v3_xxreal_0(k3_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_72, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,X2),X3)=k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_157, plain, (v1_xxreal_0(k6_numbers)|~v2_membered(k5_numbers))).#trainneg
cnf(i_0_178, plain, (v7_ordinal1(k6_numbers))).#trainneg
cnf(i_0_129, plain, (v1_xcmplx_0(k1_comptrig(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_121, plain, (X1=X2|v3_xxreal_0(X2)|v2_xxreal_0(X2)|~v1_xboole_0(X1)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_260, plain, (k2_xcmplx_0(k4_complex1(X1),k4_complex1(X2))=k4_complex1(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k4_complex1(X2))|~m1_subset_1(k4_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_136, plain, (r1_xxreal_0(k6_numbers,X1)|~v1_xcmplx_0(X1)|k1_comptrig(X1)!=k6_numbers)).#trainneg
cnf(i_0_122, plain, (v1_xcmplx_0(k16_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_131, plain, (v1_xcmplx_0(k17_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_182, plain, (v7_ordinal1(k32_sin_cos)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_203, plain, (~r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|~r1_xxreal_0(k4_complex1(X1),k6_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k6_numbers,X1))).#trainneg
cnf(i_0_164, plain, (v1_xreal_0(k6_numbers)|~v3_membered(k5_numbers))).#trainneg
cnf(i_0_171, plain, (v1_int_1(k6_numbers)|~v5_membered(k5_numbers))).#trainneg
cnf(i_0_213, plain, (v3_xxreal_0(k32_sin_cos)|v2_xxreal_0(X1)|r1_xxreal_0(X1,k32_sin_cos)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_124, plain, (v1_xcmplx_0(k4_xcmplx_0(k4_xcmplx_0(X1)))|~v1_int_1(X1))).#trainneg
cnf(i_0_218, plain, (~r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_269, plain, (k7_real_1(k3_complex1(X1),k3_complex1(X2))=k3_complex1(k2_xcmplx_0(X2,X1))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_208, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(X1),k6_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k6_numbers,X1))).#trainneg
cnf(i_0_263, plain, (k7_real_1(k4_complex1(X1),k4_complex1(X2))=k4_complex1(k2_xcmplx_0(X2,X1))|~v1_xreal_0(k4_complex1(X2))|~m1_subset_1(k4_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_189, plain, (m1_subset_1(k6_numbers,k1_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k6_numbers,X1))).#trainneg
cnf(i_0_133, plain, (v1_xcmplx_0(k4_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_175, plain, (v1_int_1(k32_sin_cos)|~v5_membered(k1_numbers))).#trainneg
cnf(i_0_188, plain, (r1_xxreal_0(k6_numbers,k6_numbers)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k6_numbers,X1))).#trainneg
cnf(i_0_273, plain, (v2_xxreal_0(X1)|v2_xxreal_0(k6_numbers)|~v3_xxreal_0(k6_numbers)|~v1_xreal_0(k6_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_222, negated_conjecture, (v3_xxreal_0(k32_sin_cos)|~v3_xxreal_0(k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0)))|~v1_xreal_0(k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0))))).#trainneg
cnf(i_0_217, negated_conjecture, (v2_xxreal_0(k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0)))|~v2_xxreal_0(k32_sin_cos)|~v1_xreal_0(k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0))))).#trainneg
cnf(i_0_206, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_158, plain, (v1_xxreal_0(k4_complex1(X1))|~v2_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_159, plain, (v1_xxreal_0(k17_complex1(X1))|~v2_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_248, plain, (m1_subset_1(k7_real_1(X2,X1),k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_268, plain, (m1_subset_1(k3_complex1(k2_xcmplx_0(X1,X2)),k1_numbers)|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_228, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_262, plain, (m1_subset_1(k4_complex1(k2_xcmplx_0(X1,X2)),k1_numbers)|~v1_xreal_0(k4_complex1(X2))|~m1_subset_1(k4_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_267, plain, (k7_real_1(k3_complex1(X2),k3_complex1(X1))=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_235, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_172, plain, (v1_int_1(k4_complex1(X1))|~v5_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_240, negated_conjecture, (v3_xxreal_0(k32_sin_cos)|v1_xboole_0(k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0)))|v2_xxreal_0(k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0)))|~v1_xreal_0(k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0))))).#trainneg
cnf(i_0_251, plain, (m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_173, plain, (v1_int_1(k17_complex1(X1))|~v5_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_224, plain, (v3_xxreal_0(X1)|v2_xxreal_0(X1)|~v2_xxreal_0(k6_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_174, plain, (v1_int_1(k1_comptrig(X1))|~v5_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_179, plain, (v7_ordinal1(k4_complex1(X1))|~v6_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_250, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k2_rcomp_1(X2,X3))|~v1_xxreal_0(X3)|~v1_xxreal_0(X2))).#trainneg
cnf(i_0_242, plain, (v1_xreal_0(k7_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_234, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(k6_numbers,X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_227, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k6_numbers))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k6_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_238, plain, (v3_xxreal_0(k6_numbers)|v1_xboole_0(k17_complex1(X1))|v2_xxreal_0(k17_complex1(X1))|~v1_xreal_0(k17_complex1(X1))|~v1_xreal_0(k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_241, plain, (v1_xxreal_0(k7_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_245, plain, (v1_xcmplx_0(k7_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_215, plain, (v2_xxreal_0(k17_complex1(X1))|~v2_xxreal_0(k6_numbers)|~v1_xreal_0(k17_complex1(X1))|~v1_xreal_0(k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_280, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,k6_numbers))=k6_numbers|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(k6_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_216, plain, (v2_xxreal_0(k1_comptrig(X1))|~v2_xxreal_0(k6_numbers)|~v1_xreal_0(k1_comptrig(X1))|~v1_xreal_0(k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_180, plain, (v7_ordinal1(k17_complex1(X1))|~v6_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_181, plain, (v7_ordinal1(k1_comptrig(X1))|~v6_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_199, plain, (v1_xxreal_0(k2_rcomp_1(X1,X2))|~v2_membered(k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_223, plain, (v3_xxreal_0(X1)|~v3_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v2_xxreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_200, plain, (v1_xreal_0(k2_rcomp_1(X1,X2))|~v3_membered(k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_201, plain, (v1_int_1(k2_rcomp_1(X1,X2))|~v5_membered(k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_202, plain, (v7_ordinal1(k2_rcomp_1(X1,X2))|~v6_membered(k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_243, plain, (v1_int_1(k7_real_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_229, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_211, plain, (v3_xxreal_0(k16_complex1(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k16_complex1(X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_230, plain, (v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_66, plain, (k4_xcmplx_0(k6_numbers)=k6_numbers)).#trainneg
cnf(i_0_355, plain, (v1_int_1(k6_numbers))).#trainneg
cnf(i_0_244, plain, (v7_ordinal1(k7_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_70, plain, (m2_subset_1(k6_numbers,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_405, plain, (v1_int_1(k32_sin_cos)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_231, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v2_xxreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_232, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X4)=k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_293, plain, (v3_xxreal_0(k1_comptrig(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k1_comptrig(X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_429, plain, (v1_xxreal_0(k4_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_358, plain, (v3_xxreal_0(k32_sin_cos)|r1_xxreal_0(X1,k32_sin_cos)|~v3_xxreal_0(X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_384, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k32_sin_cos),k6_numbers)|~r1_xxreal_0(k6_numbers,k32_sin_cos))).#trainneg
cnf(i_0_254, plain, (k4_complex1(k2_xcmplx_0(k3_xcmplx_0(X2,k7_complex1),X1))=X2|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X2,k7_complex1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_420, negated_conjecture, (~v3_xxreal_0(k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0)))|~v1_xxreal_0(k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0)))|~v2_xxreal_0(k32_sin_cos)|~v1_xreal_0(k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0))))).#trainneg
cnf(i_0_428, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_383, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k1_comptrig(X1)),k6_numbers)|~r1_xxreal_0(k6_numbers,k1_comptrig(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_523, plain, (v1_xcmplx_0(k4_complex1(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_579, negated_conjecture, (k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0))=k6_numbers|v3_xxreal_0(k32_sin_cos)|v2_xxreal_0(k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0)))|~v1_xreal_0(k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0))))).#trainneg
cnf(i_0_716, plain, (v1_xcmplx_0(k7_real_1(k4_complex1(X1),X3))|~v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1))))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_433, plain, (v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~v1_xcmplx_0(X1)|k1_comptrig(X1)!=k6_numbers)).#trainneg
cnf(i_0_298, plain, (X1=X2|v3_xxreal_0(X2)|v2_xxreal_0(X2)|v3_xxreal_0(X1)|v2_xxreal_0(X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_257, plain, (k3_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1))))=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X3,k7_complex1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_620, plain, (m1_subset_1(X1,k1_numbers)|r1_xxreal_0(X3,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X3)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_367, plain, (r1_xxreal_0(k4_complex1(X1),k6_numbers)|~v1_xxreal_0(k32_sin_cos)|~v1_xxreal_0(k6_numbers)|~v1_xxreal_0(k1_comptrig(X1))|~r1_xxreal_0(k1_comptrig(X1),k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_394, plain, (v3_xxreal_0(k4_complex1(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k4_complex1(X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_379, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k4_complex1(X1)),k6_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_750, plain, (v1_xxreal_0(k17_complex1(X1))|~v6_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_537, plain, (m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(k3_xcmplx_0(X2,k7_complex1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_281, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_594, plain, (v1_int_1(X1)|~v5_membered(k1_numbers)|~v1_xcmplx_0(X1)|k1_comptrig(X1)!=k6_numbers)).#trainneg
cnf(i_0_577, plain, (v1_xcmplx_0(k4_xcmplx_0(k4_complex1(X1)))|~v5_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_593, plain, (v1_xcmplx_0(k4_xcmplx_0(k17_complex1(X1)))|~v5_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_265, plain, (k4_complex1(k2_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)),X3))=k7_real_1(X2,k4_complex1(X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_270, plain, (k3_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1))))=k7_real_1(k3_complex1(X1),X2)|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1279, negated_conjecture, (~r1_xxreal_0(k6_numbers,k2_xcmplx_0(esk1_0,esk2_0))|~v1_xreal_0(esk1_0)|~m1_subset_1(esk2_0,k1_numbers))).#trainneg
cnf(i_0_973, plain, (r1_xxreal_0(k4_complex1(k32_sin_cos),k6_numbers)|~r1_xxreal_0(k6_numbers,k32_sin_cos))).#trainneg
cnf(i_0_886, plain, (r1_xxreal_0(k32_sin_cos,k32_sin_cos))).#trainneg
cnf(i_0_854, plain, (v3_xxreal_0(k6_numbers)|~v2_xxreal_0(k6_numbers))).#trainneg
cnf(i_0_285, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(k6_numbers,X2))=k3_xcmplx_0(k6_numbers,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_841, plain, (r1_xxreal_0(X1,k32_sin_cos)|r1_xxreal_0(k32_sin_cos,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_271, plain, (k3_complex1(k2_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)),X3))=k7_real_1(X1,k3_complex1(X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_274, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v3_xxreal_0(X3)|~v3_xxreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1045, negated_conjecture, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k1_comptrig(esk1_0)),k6_numbers))).#trainneg
cnf(i_0_1044, negated_conjecture, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k1_comptrig(esk2_0)),k6_numbers))).#trainneg
cnf(i_0_885, plain, (r1_xxreal_0(k1_comptrig(X1),k1_comptrig(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_863, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k6_numbers),k6_numbers))).#trainneg
cnf(i_0_295, plain, (v3_xxreal_0(X1)|k4_xcmplx_0(X1)=k6_numbers|v3_xxreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xxreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_1043, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k1_comptrig(esk3_0)),k6_numbers))).#trainneg
cnf(i_0_1046, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k1_comptrig(k6_numbers)),k6_numbers))).#trainneg
cnf(i_0_961, plain, (v3_xxreal_0(k1_comptrig(X1))|v2_xxreal_0(k32_sin_cos)|r1_xxreal_0(k32_sin_cos,k1_comptrig(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1222, plain, (X1=k32_sin_cos|v3_xxreal_0(X1)|v3_xxreal_0(k32_sin_cos)|v2_xxreal_0(X1)|v2_xxreal_0(k32_sin_cos)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_844, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1121, plain, (m1_subset_1(k3_complex1(k2_xcmplx_0(X2,X1)),k1_numbers)|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_843, plain, (r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k6_numbers,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_969, plain, (r1_xxreal_0(k4_complex1(X1),k4_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1058, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k1_comptrig(k32_sin_cos)),k6_numbers))).#trainneg
cnf(i_0_1311, plain, (v1_xxreal_0(X1)|~v6_membered(k1_numbers)|~v1_xcmplx_0(X1)|k1_comptrig(X1)!=k6_numbers)).#trainneg
cnf(i_0_850, plain, (r1_xxreal_0(X1,k1_comptrig(X2))|r1_xxreal_0(k1_comptrig(X2),X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_278, plain, (k4_complex1(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1))))=k3_xcmplx_0(X2,X3)|~v1_xreal_0(k3_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xcmplx_0(k7_complex1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_279, plain, (k3_complex1(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1))))=X1|~v1_xreal_0(k3_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xcmplx_0(k7_complex1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1123, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_1711, negated_conjecture, (v2_xxreal_0(k32_sin_cos))).#trainneg
cnf(i_0_1584, negated_conjecture, (r1_xxreal_0(k4_complex1(k1_comptrig(esk2_0)),k6_numbers))).#trainneg
cnf(i_0_1580, negated_conjecture, (r1_xxreal_0(k4_complex1(k1_comptrig(esk1_0)),k6_numbers))).#trainneg
cnf(i_0_280, plain, (k3_xcmplx_0(X1,k3_xcmplx_0(X2,k6_numbers))=k6_numbers|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_300, plain, (v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1493, plain, (v3_xxreal_0(k6_numbers)|v3_xxreal_0(k32_sin_cos)|r1_xxreal_0(k6_numbers,k32_sin_cos))).#trainneg
cnf(i_0_1401, negated_conjecture, (~r1_xxreal_0(k6_numbers,k2_xcmplx_0(esk1_0,esk2_0))|~v1_xreal_0(k2_xcmplx_0(esk1_0,esk2_0)))).#trainneg
cnf(i_0_282, plain, (v2_xxreal_0(k3_xcmplx_0(X1,X2))|v2_xxreal_0(X3)|~v3_xxreal_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)))|~v1_xreal_0(X3)|~v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_283, plain, (k3_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,X3)),X4)=k3_xcmplx_0(k3_xcmplx_0(X1,X2),k3_xcmplx_0(X3,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1503, plain, (v1_xcmplx_0(k3_xcmplx_0(k6_numbers,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1815, plain, (k4_complex1(X1)=k6_numbers|v3_xxreal_0(k4_complex1(X1))|~r1_xxreal_0(k32_sin_cos,k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xxreal_0(k4_complex1(X1)))).#trainneg
cnf(i_0_1512, plain, (r1_xxreal_0(k32_sin_cos,k1_comptrig(X1))|r1_xxreal_0(k1_comptrig(X1),k32_sin_cos)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_234, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(k6_numbers,X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1695, plain, (r1_xxreal_0(k4_complex1(k6_numbers),k6_numbers))).#trainneg
cnf(i_0_1701, plain, (r1_xxreal_0(k4_complex1(k1_comptrig(esk3_0)),k6_numbers))).#trainneg
cnf(i_0_302, plain, (v1_int_1(k2_xcmplx_0(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_303, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1408, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(X1),k6_numbers)|~r1_xxreal_0(k6_numbers,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1816, plain, (v3_xxreal_0(k32_sin_cos)|r1_xxreal_0(k4_complex1(X1),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_complex1(X1)))).#trainneg
cnf(i_0_1634, plain, (v3_xxreal_0(k3_complex1(X1))|v2_xxreal_0(k3_complex1(k2_xcmplx_0(X2,X1)))|~v2_xxreal_0(k3_complex1(X2))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1404, plain, (v1_int_1(X1)|~v5_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1705, plain, (r1_xxreal_0(k4_complex1(k1_comptrig(k6_numbers)),k6_numbers))).#trainneg
cnf(i_0_227, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k6_numbers))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_304, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1405, plain, (v7_ordinal1(X1)|~v6_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1806, plain, (r1_xxreal_0(k4_complex1(k1_comptrig(k32_sin_cos)),k6_numbers))).#trainneg
cnf(i_0_307, plain, (m1_subset_1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),k1_numbers)|~v1_xreal_0(X3)|~m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1923, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))|~m1_subset_1(k2_xcmplx_0(X2,X3),k1_numbers)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k1_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2038, plain, (k7_real_1(X1,k3_complex1(X2))=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2048, negated_conjecture, (~v3_xxreal_0(k32_sin_cos))).#trainneg
cnf(i_0_2056, negated_conjecture, (~v2_xxreal_0(k4_complex1(k1_comptrig(esk2_0)))|~v1_xreal_0(k4_complex1(k1_comptrig(esk2_0))))).#trainneg
cnf(i_0_2061, negated_conjecture, (~v2_xxreal_0(k4_complex1(k1_comptrig(esk1_0)))|~v1_xreal_0(k4_complex1(k1_comptrig(esk1_0))))).#trainneg
cnf(i_0_2053, plain, (X1=X2|v3_xxreal_0(X2)|v2_xxreal_0(X2)|~v1_xboole_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_2182, plain, (v1_xcmplx_0(k3_xcmplx_0(k6_numbers,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1316, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,k6_numbers))|~v1_xcmplx_0(k7_complex1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_973, plain, (r1_xxreal_0(k4_complex1(k32_sin_cos),k6_numbers))).#trainneg
cnf(i_0_420, negated_conjecture, (~v3_xxreal_0(k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0)))|~v1_xreal_0(k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0))))).#trainneg
cnf(i_0_352, plain, (v3_xxreal_0(k17_complex1(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k17_complex1(X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2392, plain, (v1_int_1(X1)|~v1_xreal_0(X1)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_2302, plain, (~v2_xxreal_0(k4_complex1(k6_numbers))|~v1_xreal_0(k4_complex1(k6_numbers)))).#trainneg
cnf(i_0_2306, plain, (~v2_xxreal_0(k4_complex1(k1_comptrig(esk3_0)))|~v1_xreal_0(k4_complex1(k1_comptrig(esk3_0))))).#trainneg
cnf(i_0_2487, negated_conjecture, (r1_xxreal_0(k1_comptrig(X1),k1_comptrig(esk2_0))|r1_xxreal_0(k1_comptrig(esk2_0),k1_comptrig(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_312, plain, (v1_xcmplx_0(k3_complex1(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(k3_complex1(X2))|~v1_xcmplx_0(k3_complex1(X1))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2202, plain, (v3_xxreal_0(k32_sin_cos)|r1_xxreal_0(k4_complex1(X1),k32_sin_cos)|k4_complex1(X1)=k6_numbers|~v1_xreal_0(k4_complex1(X1))|~r1_xxreal_0(k32_sin_cos,k1_comptrig(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2040, plain, (k7_real_1(X1,k3_complex1(X2))=k3_complex1(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X1)|~m1_subset_1(k3_complex1(X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_217, negated_conjecture, (v2_xxreal_0(k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0)))|~v1_xreal_0(k1_comptrig(k2_xcmplx_0(esk1_0,esk2_0))))).#trainneg
cnf(i_0_357, plain, (v3_xxreal_0(X1)|v3_xxreal_0(k32_sin_cos)|r1_xxreal_0(k4_xcmplx_0(X1),k32_sin_cos)|~v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_2080, plain, (k3_xcmplx_0(k6_numbers,k6_numbers)=k6_numbers|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2395, plain, (~v2_xxreal_0(k4_complex1(k1_comptrig(k6_numbers)))|~v1_xreal_0(k4_complex1(k1_comptrig(k6_numbers))))).#trainneg
cnf(i_0_2483, plain, (~v2_xxreal_0(k4_complex1(k1_comptrig(k32_sin_cos)))|~v1_xreal_0(k4_complex1(k1_comptrig(k32_sin_cos))))).#trainneg
cnf(i_0_2506, plain, (r1_xxreal_0(k1_comptrig(X1),k1_comptrig(k32_sin_cos))|r1_xxreal_0(k1_comptrig(k32_sin_cos),k1_comptrig(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_313, plain, (v1_xreal_0(k3_complex1(k2_xcmplx_0(X1,X2)))|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(k3_complex1(X1))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_314, plain, (v3_xxreal_0(k3_complex1(X1))|v2_xxreal_0(k3_complex1(k2_xcmplx_0(X1,X2)))|~v2_xxreal_0(k3_complex1(X2))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1444, plain, (k4_complex1(k3_xcmplx_0(X1,k7_complex1))=X1|~v1_xreal_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X1,k7_complex1)))).#trainneg
cnf(i_0_2733, negated_conjecture, (k4_complex1(k1_comptrig(esk2_0))=k6_numbers|v3_xxreal_0(k4_complex1(k1_comptrig(esk2_0)))|~v1_xreal_0(k4_complex1(k1_comptrig(esk2_0))))).#trainneg
cnf(i_0_2732, negated_conjecture, (r1_xxreal_0(k4_complex1(k1_comptrig(esk2_0)),k32_sin_cos)|~v1_xreal_0(k4_complex1(k1_comptrig(esk2_0))))).#trainneg
cnf(i_0_2775, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k1_comptrig(esk2_0)))).#trainneg
cnf(i_0_3128, plain, (k4_complex1(k6_numbers)=k6_numbers)).#trainneg
cnf(i_0_2954, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k1_comptrig(esk1_0)))).#trainneg
cnf(i_0_2999, plain, (r1_xxreal_0(k1_comptrig(k32_sin_cos),k1_comptrig(k32_sin_cos)))).#trainneg
cnf(i_0_315, plain, (v3_xxreal_0(k3_complex1(X1))|v3_xxreal_0(k3_complex1(X2))|~v3_xxreal_0(k3_complex1(k2_xcmplx_0(X1,X2)))|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(k3_complex1(X1))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_316, plain, (k2_xcmplx_0(k3_complex1(X1),k2_xcmplx_0(k3_complex1(X2),X3))=k2_xcmplx_0(k3_complex1(k2_xcmplx_0(X1,X2)),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_complex1(X2))|~v1_xcmplx_0(k3_complex1(X1))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2737, negated_conjecture, (r1_xxreal_0(k4_complex1(k1_comptrig(esk1_0)),k32_sin_cos)|~v1_xreal_0(k4_complex1(k1_comptrig(esk1_0))))).#trainneg
cnf(i_0_2736, negated_conjecture, (k4_complex1(k1_comptrig(esk1_0))=k6_numbers|v3_xxreal_0(k4_complex1(k1_comptrig(esk1_0)))|~v1_xreal_0(k4_complex1(k1_comptrig(esk1_0)))|~v1_xxreal_0(k4_complex1(k1_comptrig(esk1_0))))).#trainneg
cnf(i_0_2776, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k1_comptrig(esk1_0))|r1_xxreal_0(k1_comptrig(esk1_0),k1_comptrig(esk2_0)))).#trainneg
cnf(i_0_3117, plain, (~v2_xxreal_0(X1)|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k3_xcmplx_0(X1,k7_complex1)))|~v1_xcmplx_0(k3_xcmplx_0(X1,k7_complex1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2918, plain, (k3_xcmplx_0(k6_numbers,k6_numbers)=k6_numbers)).#trainneg
cnf(i_0_2796, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k1_comptrig(k32_sin_cos))|r1_xxreal_0(k1_comptrig(k32_sin_cos),k1_comptrig(esk2_0)))).#trainneg
cnf(i_0_2883, plain, (k2_xcmplx_0(k6_numbers,k3_complex1(X1))=X1|~v1_xreal_0(X1)|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2839, plain, (k4_complex1(X1)=k6_numbers|~m1_subset_1(k4_complex1(X1),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(X1))|~r1_xxreal_0(k32_sin_cos,k1_comptrig(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3093, plain, (v3_xxreal_0(k4_complex1(X1))|~v2_xxreal_0(k4_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X1)))).#trainneg
cnf(i_0_2726, plain, (v3_xxreal_0(X1)|k4_xcmplx_0(X1)=k6_numbers|v3_xxreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_431, plain, (r1_xxreal_0(X1,k17_complex1(X2))|r1_xxreal_0(k17_complex1(X2),X1)|~v1_xxreal_0(X1)|~v2_membered(k1_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_319, plain, (k3_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1))))=k2_xcmplx_0(k3_complex1(X1),X2)|~v1_xreal_0(X2)|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_320, plain, (k3_complex1(k2_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)),X3))=k2_xcmplx_0(X1,k3_complex1(X3))|~v1_xreal_0(k3_complex1(X3))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3268, plain, (k3_complex1(X1)=k7_real_1(k6_numbers,X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3746, plain, (~v2_xxreal_0(k4_complex1(X1))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3322, plain, (k3_complex1(X1)=k7_real_1(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3596, plain, (v1_xcmplx_0(k3_xcmplx_0(X1,k6_numbers))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_238, plain, (v3_xxreal_0(k6_numbers)|v1_xboole_0(k17_complex1(X1))|v2_xxreal_0(k17_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_432, plain, (r1_xxreal_0(X1,k17_complex1(X2))|r1_xxreal_0(k17_complex1(X2),X1)|~v1_xreal_0(X1)|~v2_membered(k1_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2756, plain, (~v2_xxreal_0(k4_complex1(k32_sin_cos))|~v1_xreal_0(k4_complex1(k32_sin_cos)))).#trainneg
cnf(i_0_332, plain, (k2_xcmplx_0(k3_complex1(X2),k3_complex1(X1))=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_complex1(X1))|~m1_subset_1(k3_complex1(X2),k1_numbers))).#trainneg
cnf(i_0_3242, plain, (k7_real_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3251, plain, (k7_real_1(k6_numbers,X1)=X1|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3608, plain, (X1=k1_comptrig(X2)|v2_xxreal_0(k1_comptrig(X2))|~v1_xboole_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2757, plain, (r1_xxreal_0(k1_comptrig(k32_sin_cos),k6_numbers)|r1_xxreal_0(k32_sin_cos,k1_comptrig(k32_sin_cos)))).#trainneg
cnf(i_0_434, plain, (v1_xxreal_0(k7_real_1(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_365, plain, (k4_complex1(k2_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)),X1))=k7_real_1(k4_complex1(X1),X3)|~v1_xreal_0(k4_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_366, plain, (k4_complex1(k2_xcmplx_0(X3,k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1))))=k7_real_1(X2,k4_complex1(X3))|~v1_xreal_0(X2)|~v1_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)))|~v1_xcmplx_0(X3)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3862, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,k6_numbers))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_435, plain, (v1_xreal_0(k7_real_1(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_370, plain, (m1_subset_1(k3_complex1(k2_xcmplx_0(X2,X1)),k1_numbers)|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4262, plain, (k3_complex1(k2_xcmplx_0(X1,X2))=k2_xcmplx_0(k3_complex1(X1),X2)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_complex1(X1)))).#trainneg
cnf(i_0_3850, plain, (k7_real_1(k6_numbers,k6_numbers)=k6_numbers)).#trainneg
cnf(i_0_436, plain, (v1_int_1(k7_real_1(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3807, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(k6_numbers,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3606, plain, (k1_comptrig(X1)=X2|v3_xxreal_0(X2)|v2_xxreal_0(X2)|v2_xxreal_0(k1_comptrig(X1))|~v1_xxreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3607, plain, (k1_comptrig(X1)=X2|v3_xxreal_0(X2)|v2_xxreal_0(X2)|v2_xxreal_0(k1_comptrig(X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4271, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4558, plain, (v1_xcmplx_0(k2_xcmplx_0(k6_numbers,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_437, plain, (v7_ordinal1(k7_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_373, plain, (k3_complex1(k2_xcmplx_0(X2,X1))=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers))).#trainneg
cnf(i_0_374, plain, (k3_complex1(k2_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)),X1))=k7_real_1(k3_complex1(X1),X2)|~v1_xreal_0(X2)|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_438, plain, (v1_xcmplx_0(k7_real_1(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_375, plain, (k3_complex1(k2_xcmplx_0(X3,k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1))))=k7_real_1(X1,k3_complex1(X3))|~v1_xreal_0(k3_complex1(X3))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_380, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k17_complex1(X1)),k6_numbers)|~r1_xxreal_0(k6_numbers,k17_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4668, plain, (k2_xcmplx_0(k3_complex1(X1),k6_numbers)=X1|~v1_xreal_0(X1)|~v1_xreal_0(k3_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4165, plain, (k4_complex1(k32_sin_cos)=k6_numbers|v3_xxreal_0(k4_complex1(k32_sin_cos))|~v1_xreal_0(k4_complex1(k32_sin_cos)))).#trainneg
cnf(i_0_2505, plain, (r1_xxreal_0(k1_comptrig(X1),k1_comptrig(k1_comptrig(X2)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X2)),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_381, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k7_real_1(X1,X2)),k6_numbers)|~r1_xxreal_0(k6_numbers,k7_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_382, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k2_xcmplx_0(X1,X2)),k6_numbers)|~r1_xxreal_0(k6_numbers,k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5459, plain, (k2_xcmplx_0(k6_numbers,k3_complex1(X1))=X1|~v1_xcmplx_0(k3_complex1(X1))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5516, plain, (k2_xcmplx_0(X1,k3_complex1(X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4278, plain, (k7_real_1(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3365, plain, (r1_xxreal_0(k1_comptrig(k6_numbers),k6_numbers)|r1_xxreal_0(k32_sin_cos,k1_comptrig(k6_numbers)))).#trainneg
cnf(i_0_4708, plain, (k3_complex1(k2_xcmplx_0(X2,X1))=k2_xcmplx_0(k3_complex1(X1),X2)|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4650, plain, (k2_xcmplx_0(k3_complex1(X1),X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_3635, plain, (k2_xcmplx_0(k3_complex1(X1),k6_numbers)=X1|~v1_xreal_0(X1)|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4164, plain, (r1_xxreal_0(k4_complex1(k32_sin_cos),k32_sin_cos)|~v1_xreal_0(k4_complex1(k32_sin_cos)))).#trainneg
cnf(i_0_4001, plain, (k7_real_1(k6_numbers,k2_xcmplx_0(X1,k6_numbers))=X1|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,k6_numbers)))).#trainneg
cnf(i_0_1811, plain, (~v2_xxreal_0(X2)|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1))))|~v1_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1223, plain, (X1=X2|v3_xxreal_0(X1)|v3_xxreal_0(X2)|v2_xxreal_0(X1)|v2_xxreal_0(X2)|~v1_xxreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_3613, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_comptrig(esk2_0))|r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers))).#trainneg
cnf(i_0_1809, plain, (~v2_xxreal_0(X1)|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(k3_xcmplx_0(X1,k7_complex1),X2)))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,k7_complex1),X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,k7_complex1)))).#trainneg
cnf(i_0_5620, plain, (v3_xxreal_0(k3_complex1(X1))|~v3_xxreal_0(X1)|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4681, plain, (k2_xcmplx_0(k3_complex1(X1),X2)=k7_real_1(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(k3_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5612, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k2_xcmplx_0(X2,X1)),k6_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k6_numbers,k2_xcmplx_0(X2,X1))|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3726, plain, (X1=k6_numbers|~r1_xxreal_0(k6_numbers,X1)|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k3_xcmplx_0(X1,k7_complex1)))|~v1_xcmplx_0(k3_xcmplx_0(X1,k7_complex1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2742, plain, (X1=X2|v3_xxreal_0(X2)|v2_xxreal_0(X2)|v3_xxreal_0(X1)|v2_xxreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5569, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(k1_comptrig(X2)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X2)),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_442, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k7_real_1(X1,X2)),k6_numbers)|~r1_xxreal_0(k6_numbers,k7_real_1(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_3622, plain, (k2_xcmplx_0(k3_complex1(X1),X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4652, plain, (k2_xcmplx_0(k3_complex1(X1),X2)=k7_real_1(k6_numbers,k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_509, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(k6_numbers,X2))=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6437, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(k1_comptrig(esk2_0)))|r1_xxreal_0(k1_comptrig(k1_comptrig(esk2_0)),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6438, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(k1_comptrig(esk1_0)))|r1_xxreal_0(k1_comptrig(k1_comptrig(esk1_0)),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6422, negated_conjecture, (v3_xxreal_0(k4_complex1(esk2_0))|r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_6421, negated_conjecture, (v3_xxreal_0(k4_complex1(esk2_0))|k4_complex1(esk1_0)=k6_numbers|v3_xxreal_0(k4_complex1(esk1_0))|~v1_xxreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_6423, negated_conjecture, (v3_xxreal_0(k4_complex1(esk2_0))|k4_complex1(esk1_0)=k6_numbers|v3_xxreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_6319, plain, (k3_complex1(k7_real_1(X1,X2))=k7_real_1(k3_complex1(X1),X2)|~v1_xreal_0(X2)|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_510, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))=k2_xcmplx_0(X3,k2_xcmplx_0(X2,X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6462, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(k1_comptrig(k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k1_comptrig(k32_sin_cos)),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6981, negated_conjecture, (k4_complex1(esk1_0)=k6_numbers|v3_xxreal_0(k4_complex1(esk1_0))|~v1_xxreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers)))).#trainneg
cnf(i_0_6979, negated_conjecture, (k4_complex1(esk1_0)=k6_numbers|v3_xxreal_0(k4_complex1(esk1_0))|~m1_subset_1(k4_complex1(esk2_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0))|~v1_xxreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_6986, negated_conjecture, (k4_complex1(esk1_0)=k6_numbers|v3_xxreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers)))).#trainneg
cnf(i_0_6984, negated_conjecture, (k4_complex1(esk1_0)=k6_numbers|v3_xxreal_0(k4_complex1(esk1_0))|~m1_subset_1(k4_complex1(esk2_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_7036, plain, (k7_real_1(k3_complex1(X1),X2)=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6975, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_6976, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers)))).#trainneg
cnf(i_0_6973, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_6873, plain, (k2_xcmplx_0(k6_numbers,X1)=k2_xcmplx_0(X1,k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6917, negated_conjecture, (v3_xxreal_0(k4_complex1(esk1_0))|r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_6920, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(esk2_0)),k1_comptrig(k1_comptrig(esk2_0))))).#trainneg
cnf(i_0_6948, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(esk1_0)),k1_comptrig(k1_comptrig(esk1_0))))).#trainneg
cnf(i_0_458, plain, (v1_xxreal_0(k3_complex1(k2_xcmplx_0(X1,X2)))|~v2_membered(k1_numbers)|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_459, plain, (v1_xreal_0(k3_complex1(k2_xcmplx_0(X1,X2)))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6974, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~m1_subset_1(k4_complex1(esk2_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_6982, negated_conjecture, (k4_complex1(esk1_0)=k6_numbers|v3_xxreal_0(k4_complex1(esk1_0))|~v1_xxreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk2_0)))).#trainneg
cnf(i_0_6985, negated_conjecture, (k4_complex1(esk1_0)=k6_numbers|v3_xxreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_6916, negated_conjecture, (v3_xxreal_0(k4_complex1(esk1_0))|k4_complex1(esk2_0)=k6_numbers|v3_xxreal_0(k4_complex1(esk2_0))|~v1_xxreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_6980, negated_conjecture, (k4_complex1(esk1_0)=k6_numbers|v3_xxreal_0(k4_complex1(esk1_0))|~v1_xxreal_0(k4_complex1(esk1_0))|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_6987, negated_conjecture, (k4_complex1(esk1_0)=k6_numbers|v3_xxreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk2_0)))).#trainneg
cnf(i_0_6977, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk2_0)))).#trainneg
cnf(i_0_4160, plain, (v1_xboole_0(X1)|v2_xxreal_0(X1)|~v1_xcmplx_0(X1)|k1_comptrig(X1)!=k6_numbers)).#trainneg
cnf(i_0_6918, negated_conjecture, (v3_xxreal_0(k4_complex1(esk1_0))|k4_complex1(esk2_0)=k6_numbers|v3_xxreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_516, plain, (k2_xcmplx_0(k6_numbers,k2_xcmplx_0(X1,X2))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(k6_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1320, plain, (v1_int_1(X1)|k1_comptrig(X1)!=k6_numbers|~v1_xcmplx_0(X1)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_460, plain, (v1_int_1(k3_complex1(k2_xcmplx_0(X1,X2)))|~v5_membered(k1_numbers)|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_461, plain, (v7_ordinal1(k3_complex1(k2_xcmplx_0(X1,X2)))|~v6_membered(k1_numbers)|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7751, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk1_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0)))).#trainneg
cnf(i_0_7361, plain, (k3_complex1(k2_xcmplx_0(X1,X2))=k2_xcmplx_0(k3_complex1(X1),X2)|~v1_xreal_0(X2)|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7357, negated_conjecture, (k4_complex1(esk1_0)=k6_numbers|~m1_subset_1(k4_complex1(esk1_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk1_0))|~m1_subset_1(k4_complex1(esk2_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0)))).#trainneg
cnf(i_0_517, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2))=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7578, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(esk1_0))).#trainneg
cnf(i_0_7349, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(k32_sin_cos)),k1_comptrig(k1_comptrig(k32_sin_cos))))).#trainneg
cnf(i_0_7548, plain, (~v2_xxreal_0(k4_complex1(X1))|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(X1,k6_numbers)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8460, plain, (r1_xxreal_0(k1_comptrig(k6_numbers),k1_comptrig(k6_numbers)))).#trainneg
cnf(i_0_960, plain, (v3_xxreal_0(k1_comptrig(X1))|v2_xxreal_0(k1_comptrig(X2))|r1_xxreal_0(k1_comptrig(X2),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_466, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k3_complex1(k2_xcmplx_0(X1,X2))),k6_numbers)|~r1_xxreal_0(k6_numbers,k3_complex1(k2_xcmplx_0(X1,X2)))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_470, plain, (m1_subset_1(k3_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))),k1_numbers)|~v1_xreal_0(k3_complex1(X3))|~m1_subset_1(k3_complex1(k2_xcmplx_0(X1,X2)),k1_numbers)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7770, negated_conjecture, (k4_complex1(esk2_0)=k6_numbers|v3_xxreal_0(k4_complex1(esk1_0))|~m1_subset_1(k4_complex1(esk2_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_7353, negated_conjecture, (k4_complex1(esk1_0)=k6_numbers|~m1_subset_1(k4_complex1(esk1_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0))|~v1_xxreal_0(k4_complex1(esk1_0))|~m1_subset_1(k4_complex1(esk2_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0)))).#trainneg
cnf(i_0_7577, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|~m1_subset_1(k4_complex1(esk1_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_7757, negated_conjecture, (k4_complex1(esk2_0)=k6_numbers|v3_xxreal_0(k4_complex1(esk1_0))|~m1_subset_1(k4_complex1(esk2_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0))|~v1_xxreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_5742, plain, (k7_real_1(X1,k3_complex1(X2))=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_8138, negated_conjecture, (k4_complex1(esk1_0)=k6_numbers|~v1_xreal_0(k4_complex1(esk1_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0)))).#trainneg
cnf(i_0_519, plain, (v1_xxreal_0(k4_complex1(k2_xcmplx_0(X1,X2)))|~v2_membered(k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7355, negated_conjecture, (k4_complex1(esk1_0)=k6_numbers|~m1_subset_1(k4_complex1(esk1_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0))|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers))|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_4654, plain, (k2_xcmplx_0(k3_complex1(X1),X2)=k2_xcmplx_0(k6_numbers,k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6908, plain, (k2_xcmplx_0(k3_complex1(X2),X1)=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(k3_complex1(X2),k1_numbers))).#trainneg
cnf(i_0_9140, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_8675, negated_conjecture, (k4_complex1(esk2_0)=k6_numbers|~v1_xreal_0(k4_complex1(esk2_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0))|~v1_xreal_0(esk1_0))).#trainneg
cnf(i_0_520, plain, (v1_xreal_0(k4_complex1(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_476, plain, (k2_xcmplx_0(X3,k2_xcmplx_0(X4,k2_xcmplx_0(X1,X2)))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4)))|~v1_xcmplx_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_8748, negated_conjecture, (k4_complex1(esk1_0)=k6_numbers|~v1_xreal_0(k4_complex1(esk1_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0))|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_521, plain, (v1_int_1(k4_complex1(k2_xcmplx_0(X1,X2)))|~v5_membered(k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_4382, plain, (r1_xxreal_0(k32_sin_cos,k1_comptrig(k32_sin_cos))|~v2_xxreal_0(k1_comptrig(k32_sin_cos))|~v1_xreal_0(k1_comptrig(k32_sin_cos)))).#trainneg
cnf(i_0_8676, negated_conjecture, (k4_complex1(esk2_0)=k6_numbers|~v1_xreal_0(k4_complex1(esk2_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0))|~v1_xreal_0(k2_xcmplx_0(esk1_0,k6_numbers)))).#trainneg
cnf(i_0_522, plain, (v7_ordinal1(k4_complex1(k2_xcmplx_0(X1,X2)))|~v6_membered(k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5819, plain, (r1_xxreal_0(k32_sin_cos,k1_comptrig(k6_numbers))|~v2_xxreal_0(k1_comptrig(k6_numbers))|~v1_xreal_0(k1_comptrig(k6_numbers)))).#trainneg
cnf(i_0_6193, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_comptrig(esk2_0))|~v2_xxreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(k1_comptrig(esk2_0)))).#trainneg
cnf(i_0_480, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1)))|~v2_xxreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_8693, plain, (k3_complex1(k2_xcmplx_0(k3_complex1(X1),X2))=k3_complex1(k2_xcmplx_0(X1,k3_complex1(X2)))|~v1_xreal_0(k3_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_5494, plain, (v2_xxreal_0(X1)|~v2_xxreal_0(k3_complex1(X1))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7351, negated_conjecture, (k4_complex1(esk1_0)=k6_numbers|~m1_subset_1(k4_complex1(esk1_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0))|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers))|~v1_xxreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_8677, negated_conjecture, (k4_complex1(esk2_0)=k6_numbers|~v1_xreal_0(k4_complex1(esk2_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0))|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk1_0)))).#trainneg
cnf(i_0_7579, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k2_xcmplx_0(esk1_0,k6_numbers)))).#trainneg
cnf(i_0_482, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1)),X4)=k2_xcmplx_0(X1,k2_xcmplx_0(k2_xcmplx_0(X2,X3),X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_10313, plain, (v2_xxreal_0(X1)|k3_complex1(X1)=k6_numbers|v3_xxreal_0(k3_complex1(X1))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_10248, plain, (k3_complex1(k3_complex1(X1))=k3_complex1(k2_xcmplx_0(X1,k6_numbers))|~v1_xreal_0(k3_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_5162, plain, (k3_complex1(k2_xcmplx_0(X2,X1))=k2_xcmplx_0(X2,k3_complex1(X1))|~v1_xreal_0(X2)|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_complex1(X2))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_9989, plain, (r1_xxreal_0(k32_sin_cos,k1_comptrig(k32_sin_cos))|k1_comptrig(k32_sin_cos)=k6_numbers|v3_xxreal_0(k1_comptrig(k32_sin_cos))|~v1_xreal_0(k1_comptrig(k32_sin_cos)))).#trainneg
cnf(i_0_7580, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk1_0)))).#trainneg
cnf(i_0_484, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(k2_xcmplx_0(X2,X3),X4))=k2_xcmplx_0(X4,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1)))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_485, plain, (m1_subset_1(k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1)),k1_numbers)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_8686, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk2_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_8687, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk2_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0)))).#trainneg
cnf(i_0_8699, plain, (k3_complex1(k2_xcmplx_0(X1,X2))=k2_xcmplx_0(X1,k3_complex1(X2))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(X1,k1_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_4957, plain, (k3_complex1(k2_xcmplx_0(X2,X1))=k7_real_1(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_9990, plain, (r1_xxreal_0(k32_sin_cos,k1_comptrig(k32_sin_cos))|k1_comptrig(k32_sin_cos)=k6_numbers|~v1_xreal_0(k1_comptrig(k32_sin_cos)))).#trainneg
cnf(i_0_8750, negated_conjecture, (k4_complex1(esk1_0)=k6_numbers|~v1_xreal_0(k4_complex1(esk1_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0))|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk2_0)))).#trainneg
cnf(i_0_2772, plain, (r1_xxreal_0(k4_complex1(k1_comptrig(esk3_0)),k32_sin_cos)|~v1_xreal_0(k4_complex1(k1_comptrig(esk3_0))))).#trainneg
cnf(i_0_486, plain, (m1_subset_1(k3_complex1(k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))),k1_numbers)|~v1_xreal_0(k3_complex1(k2_xcmplx_0(X2,X3)))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_9262, plain, (k3_complex1(k2_xcmplx_0(X2,X1))=k2_xcmplx_0(X2,k3_complex1(X1))|~v1_xreal_0(X2)|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_7477, plain, (v1_xcmplx_0(k4_complex1(k2_xcmplx_0(k6_numbers,X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11609, plain, (v2_xxreal_0(k1_comptrig(k32_sin_cos))|k1_comptrig(k32_sin_cos)=k6_numbers|~v1_xreal_0(k1_comptrig(k32_sin_cos)))).#trainneg
cnf(i_0_11607, plain, (k1_comptrig(k32_sin_cos)=k6_numbers|~v3_xxreal_0(k1_comptrig(k32_sin_cos))|~v1_xreal_0(k1_comptrig(k32_sin_cos)))).#trainneg
cnf(i_0_1831, plain, (r1_xxreal_0(k1_comptrig(X1),k2_xcmplx_0(X2,X3))|r1_xxreal_0(k2_xcmplx_0(X2,X3),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1521, plain, (r1_xxreal_0(k32_sin_cos,k2_xcmplx_0(X1,X2))|r1_xxreal_0(k2_xcmplx_0(X1,X2),k32_sin_cos)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_488, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,k2_xcmplx_0(X4,X2)))=k2_xcmplx_0(k2_xcmplx_0(X3,X4),k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_491, plain, (k2_xcmplx_0(X4,k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4)))|~v1_xcmplx_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_11545, plain, (k7_real_1(X2,X1)=k7_real_1(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7541, plain, (v1_xcmplx_0(k4_complex1(k2_xcmplx_0(X1,k6_numbers)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2945, plain, (r1_xxreal_0(k4_complex1(k1_comptrig(k6_numbers)),k32_sin_cos)|~v1_xreal_0(k4_complex1(k1_comptrig(k6_numbers))))).#trainneg
cnf(i_0_492, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1)=k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_494, plain, (v1_xreal_0(k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2)))|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_11534, plain, (v1_xreal_0(k7_real_1(X2,X1))|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(k3_complex1(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_10060, plain, (r1_xxreal_0(k32_sin_cos,k1_comptrig(k6_numbers))|k1_comptrig(k6_numbers)=k6_numbers|v3_xxreal_0(k1_comptrig(k6_numbers))|~v1_xreal_0(k1_comptrig(k6_numbers)))).#trainneg
cnf(i_0_2950, plain, (r1_xxreal_0(k4_complex1(k1_comptrig(k32_sin_cos)),k32_sin_cos)|~v1_xreal_0(k4_complex1(k1_comptrig(k32_sin_cos))))).#trainneg
cnf(i_0_11943, plain, (r1_xxreal_0(k2_xcmplx_0(X1,k32_sin_cos),k1_comptrig(X2))|r1_xxreal_0(k1_comptrig(X2),k2_xcmplx_0(X1,k32_sin_cos))|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_495, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2)))|~v2_xxreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_496, plain, (v3_xxreal_0(X1)|v3_xxreal_0(k2_xcmplx_0(X2,X3))|~v3_xxreal_0(k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2)))|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12915, plain, (k2_xcmplx_0(k6_numbers,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1)))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_11430, plain, (k2_xcmplx_0(k3_complex1(X1),k3_complex1(X2))=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6623, plain, (k3_complex1(k2_xcmplx_0(X1,X2))=k2_xcmplx_0(X1,k3_complex1(X2))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,k3_complex1(X2))))).#trainneg
cnf(i_0_10061, plain, (r1_xxreal_0(k32_sin_cos,k1_comptrig(k6_numbers))|k1_comptrig(k6_numbers)=k6_numbers|~v1_xreal_0(k1_comptrig(k6_numbers)))).#trainneg
cnf(i_0_13516, plain, (r1_xxreal_0(k1_comptrig(X1),k2_xcmplx_0(k1_comptrig(X2),k32_sin_cos))|r1_xxreal_0(k2_xcmplx_0(k1_comptrig(X2),k32_sin_cos),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_497, plain, (k2_xcmplx_0(k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2)),X4)=k2_xcmplx_0(X1,k2_xcmplx_0(k2_xcmplx_0(X2,X3),X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4192, plain, (k2_xcmplx_0(k3_complex1(X2),k3_complex1(X1))=k3_complex1(k2_xcmplx_0(X2,X1))|~m1_subset_1(k3_complex1(X2),k1_numbers)|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xreal_0(k3_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4049, plain, (r1_xxreal_0(k4_complex1(X1),k32_sin_cos)|~v1_xreal_0(X1)|~v1_xreal_0(k4_complex1(X1)))).#trainneg
cnf(i_0_14196, plain, (r1_xxreal_0(k4_complex1(X1),k6_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(k4_complex1(X1)))).#trainneg
cnf(i_0_14182, plain, (v2_xxreal_0(k1_comptrig(k6_numbers))|k1_comptrig(k6_numbers)=k6_numbers|~v1_xreal_0(k1_comptrig(k6_numbers)))).#trainneg
cnf(i_0_615, plain, (v1_xcmplx_0(k4_xcmplx_0(k1_comptrig(X1)))|~v5_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_14180, plain, (k1_comptrig(k6_numbers)=k6_numbers|~v3_xxreal_0(k1_comptrig(k6_numbers))|~v1_xreal_0(k1_comptrig(k6_numbers)))).#trainneg
cnf(i_0_2502, plain, (r1_xxreal_0(k1_comptrig(X1),k1_comptrig(k2_xcmplx_0(X2,X3)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X2,X3)),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_6461, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(k1_comptrig(k1_comptrig(X2))))|r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(X2))),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_499, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(k2_xcmplx_0(X2,X3),X4))=k2_xcmplx_0(X4,k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_500, plain, (m1_subset_1(k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2)),k1_numbers)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_14746, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_14739, plain, (r1_xxreal_0(k4_xcmplx_0(X1),k6_numbers)|k4_xcmplx_0(X1)=k6_numbers|v3_xxreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_14736, plain, (r1_xxreal_0(k3_complex1(X1),k6_numbers)|~v1_xreal_0(k3_complex1(X1))|~v3_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_14749, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|k4_complex1(esk1_0)=k6_numbers|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers))|~v1_xxreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_14741, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_8465, plain, (~v2_xxreal_0(X1)|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(k3_xcmplx_0(X1,k7_complex1),k6_numbers)))|~v1_xcmplx_0(k3_xcmplx_0(X1,k7_complex1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_14230, plain, (r1_xxreal_0(k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos),k1_comptrig(k1_comptrig(X2)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X2)),k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_2795, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k1_comptrig(k1_comptrig(X1)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(esk2_0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_501, plain, (m1_subset_1(k3_complex1(k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2))),k1_numbers)|~v1_xreal_0(k3_complex1(k2_xcmplx_0(X2,X3)))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_3809, plain, (k4_complex1(k2_xcmplx_0(X1,X2))=k6_numbers|~v1_xreal_0(k2_xcmplx_0(k2_xcmplx_0(k6_numbers,X2),X1))|~v1_xcmplx_0(k2_xcmplx_0(k6_numbers,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_7759, negated_conjecture, (k4_complex1(esk1_0)=k6_numbers|~m1_subset_1(k4_complex1(esk1_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0))|~v1_xreal_0(esk2_0)|~v1_xxreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_11475, plain, (k3_complex1(k2_xcmplx_0(X2,X1))=k2_xcmplx_0(X1,k3_complex1(X2))|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_10065, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_comptrig(esk2_0))|k1_comptrig(esk2_0)=k6_numbers|v3_xxreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(k1_comptrig(esk2_0)))).#trainneg
cnf(i_0_2974, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k1_comptrig(k1_comptrig(X1)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(esk1_0))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_2998, plain, (r1_xxreal_0(k1_comptrig(k32_sin_cos),k1_comptrig(k1_comptrig(X1)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(k32_sin_cos))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_503, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X4,k2_xcmplx_0(X2,X3)))=k2_xcmplx_0(k2_xcmplx_0(X3,X4),k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_505, plain, (k2_xcmplx_0(k3_complex1(X1),k2_xcmplx_0(k3_complex1(X2),X3))=k2_xcmplx_0(X3,k3_complex1(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_complex1(X2))|~v1_xcmplx_0(k3_complex1(X1))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_15532, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_15584, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_15583, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~v2_xxreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_15531, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~m1_subset_1(k4_complex1(esk2_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_15534, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers)))).#trainneg
cnf(i_0_751, plain, (v7_ordinal1(X1)|~v6_membered(k1_numbers)|~v1_xcmplx_0(X1)|k1_comptrig(X1)!=k6_numbers)).#trainneg
cnf(i_0_16052, negated_conjecture, (~v2_xxreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_16059, negated_conjecture, (~v2_xxreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(esk1_0))).#trainneg
cnf(i_0_508, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))=k2_xcmplx_0(X4,k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_512, plain, (k2_xcmplx_0(k3_complex1(X3),k2_xcmplx_0(X1,k3_complex1(X2)))=k2_xcmplx_0(X1,k3_complex1(k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(k3_complex1(X3))|~v1_xcmplx_0(k3_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_complex1(X3))|~m1_subset_1(k3_complex1(X2),k1_numbers)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_15572, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|~m1_subset_1(k4_complex1(esk1_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_15895, negated_conjecture, (k4_complex1(esk1_0)=k6_numbers|~v1_xreal_0(esk2_0)|~v1_xxreal_0(k4_complex1(esk1_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0)))).#trainneg
cnf(i_0_16977, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk1_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_17889, negated_conjecture, (k4_complex1(esk1_0)=k6_numbers|~v1_xreal_0(esk2_0)|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0)))).#trainneg
cnf(i_0_13517, plain, (r1_xxreal_0(k1_comptrig(X1),k2_xcmplx_0(k32_sin_cos,k32_sin_cos))|r1_xxreal_0(k2_xcmplx_0(k32_sin_cos,k32_sin_cos),k1_comptrig(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_515, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4)))=k2_xcmplx_0(X4,k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_527, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k4_complex1(k2_xcmplx_0(X1,X2))),k6_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_16978, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk1_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0)))).#trainneg
cnf(i_0_17940, plain, (v1_xcmplx_0(k2_xcmplx_0(k6_numbers,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_528, plain, (m1_subset_1(k4_complex1(k3_complex1(k2_xcmplx_0(X1,X2))),k1_numbers)|~v1_xcmplx_0(k3_complex1(X2))|~v1_xcmplx_0(k3_complex1(X1))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_531, plain, (m1_subset_1(k4_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))),k1_numbers)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_17000, negated_conjecture, (k4_complex1(esk2_0)=k6_numbers|v3_xxreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(esk1_0))).#trainneg
cnf(i_0_15576, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk1_0)))).#trainneg
cnf(i_0_17886, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_16957, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_15975, plain, (k2_xcmplx_0(X2,k3_complex1(X1))=k7_real_1(X2,X1)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(k3_complex1(X1)))).#trainneg
cnf(i_0_18926, negated_conjecture, (~v2_xxreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk1_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0)))).#trainneg
cnf(i_0_15591, plain, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(k6_numbers,k3_xcmplx_0(X1,k7_complex1))))|~v1_xcmplx_0(k3_xcmplx_0(X1,k7_complex1)))).#trainneg
cnf(i_0_16981, negated_conjecture, (~v2_xxreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers)))).#trainneg
cnf(i_0_19259, negated_conjecture, (~v2_xxreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk1_0)))).#trainneg
cnf(i_0_19266, negated_conjecture, (~v2_xxreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk1_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0)))).#trainneg
cnf(i_0_15575, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k2_xcmplx_0(esk1_0,k6_numbers)))).#trainneg
cnf(i_0_10066, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_comptrig(esk2_0))|k1_comptrig(esk2_0)=k6_numbers|~v1_xreal_0(k1_comptrig(esk2_0)))).#trainneg
cnf(i_0_8040, plain, (k2_xcmplx_0(k3_complex1(X1),X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_6909, plain, (k2_xcmplx_0(k3_complex1(X2),X1)=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(k3_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_19509, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|~v1_xreal_0(k2_xcmplx_0(esk1_0,k6_numbers))|~v1_xreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_1008, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,X3))|r1_xxreal_0(k2_xcmplx_0(X2,X3),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_545, plain, (k3_complex1(k2_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)),X1))=k7_real_1(k3_complex1(X1),X2)|~v1_xreal_0(k3_complex1(X1))|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1216, plain, (v1_xxreal_0(X1)|k1_comptrig(X1)!=k6_numbers|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_546, plain, (k3_complex1(k2_xcmplx_0(X3,k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1))))=k7_real_1(X1,k3_complex1(X3))|~v1_xreal_0(X1)|~m1_subset_1(k3_complex1(X3),k1_numbers)|~v1_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_547, plain, (k2_xcmplx_0(X2,k2_xcmplx_0(X1,k6_numbers))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k6_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_20048, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|~v1_xreal_0(k2_xcmplx_0(esk1_0,k6_numbers))|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers)))).#trainneg
cnf(i_0_20050, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|~v1_xreal_0(k2_xcmplx_0(esk1_0,k6_numbers)))).#trainneg
cnf(i_0_19514, negated_conjecture, (v1_xboole_0(k1_comptrig(esk2_0))|v2_xxreal_0(k1_comptrig(esk2_0))|k1_comptrig(esk2_0)=k6_numbers|~v1_xreal_0(k1_comptrig(esk2_0)))).#trainneg
cnf(i_0_20284, plain, (v1_xxreal_0(X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k6_numbers,X1))).#trainneg
cnf(i_0_548, plain, (k4_complex1(k2_xcmplx_0(X2,k2_xcmplx_0(X1,k3_xcmplx_0(X3,k7_complex1))))=X3|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X3,k7_complex1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_549, plain, (k3_complex1(k2_xcmplx_0(X2,k2_xcmplx_0(X1,k3_xcmplx_0(X3,k7_complex1))))=k2_xcmplx_0(X1,X2)|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X3,k7_complex1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_14750, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|k4_complex1(esk1_0)=k6_numbers|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk2_0))|~v1_xxreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_4122, plain, (k7_real_1(k2_xcmplx_0(X1,X2),k6_numbers)=k7_real_1(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_11561, plain, (k7_real_1(X2,k3_complex1(X1))=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_10312, plain, (v2_xxreal_0(X1)|r1_xxreal_0(k3_complex1(X1),k32_sin_cos)|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_2503, plain, (r1_xxreal_0(k1_comptrig(X1),k1_comptrig(k2_xcmplx_0(X2,X3)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X2,X3)),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_550, plain, (k2_xcmplx_0(X2,k2_xcmplx_0(X1,k2_xcmplx_0(X3,X4)))=k2_xcmplx_0(X4,k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_551, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4)))=k2_xcmplx_0(X4,k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4))).#trainneg
cnf(i_0_21146, plain, (k3_complex1(k2_xcmplx_0(X2,X1))=k2_xcmplx_0(X1,k3_complex1(X2))|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_20717, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k32_sin_cos)|~v1_xreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_21225, plain, (v2_xxreal_0(X1)|r1_xxreal_0(X1,k32_sin_cos)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_21086, plain, (k7_real_1(X1,k3_complex1(X2))=k7_real_1(k2_xcmplx_0(X1,X2),k6_numbers)|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_20725, negated_conjecture, (k1_comptrig(esk2_0)=k6_numbers|v3_xxreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(esk1_0))).#trainneg
cnf(i_0_754, plain, (v1_xxreal_0(k2_rcomp_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_19508, negated_conjecture, (~v2_xxreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k2_xcmplx_0(esk1_0,k6_numbers)))).#trainneg
cnf(i_0_20689, negated_conjecture, (~v2_xxreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(k2_xcmplx_0(esk1_0,k6_numbers)))).#trainneg
cnf(i_0_21247, plain, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,k32_sin_cos)),k1_comptrig(X2))|r1_xxreal_0(k1_comptrig(X2),k1_comptrig(k2_xcmplx_0(X1,k32_sin_cos)))|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_554, plain, (v1_xreal_0(k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3)))|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_15535, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk2_0)))).#trainneg
cnf(i_0_3754, plain, (v3_xxreal_0(k4_complex1(X1))|~v2_xxreal_0(k4_complex1(X2))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_22894, plain, (r1_xxreal_0(k4_complex1(X1),k32_sin_cos)|~r1_xxreal_0(k32_sin_cos,k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k4_complex1(X1)))).#trainneg
cnf(i_0_22935, plain, (k3_complex1(k2_xcmplx_0(X1,X2))=k7_real_1(k2_xcmplx_0(X1,X2),k6_numbers)|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_767, plain, (v1_int_1(k2_rcomp_1(X1,X2))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v6_membered(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_4023, plain, (k7_real_1(k6_numbers,X1)=k2_xcmplx_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23006, plain, (r1_xxreal_0(k1_comptrig(X1),k1_comptrig(k2_xcmplx_0(k32_sin_cos,k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k32_sin_cos,k32_sin_cos)),k1_comptrig(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_23005, plain, (r1_xxreal_0(k1_comptrig(X1),k1_comptrig(k2_xcmplx_0(k1_comptrig(X2),k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k1_comptrig(X2),k32_sin_cos)),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_555, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3)))|~v2_xxreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_556, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v3_xxreal_0(X3)|~v3_xxreal_0(k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3)))|~v1_xreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_22899, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_comptrig(esk1_0))|r1_xxreal_0(k1_comptrig(esk1_0),k32_sin_cos)|~v1_xreal_0(k1_comptrig(esk1_0)))).#trainneg
cnf(i_0_22870, plain, (k3_complex1(k2_xcmplx_0(X2,X1))=k2_xcmplx_0(X2,k3_complex1(X1))|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_22897, negated_conjecture, (r1_xxreal_0(k6_numbers,k1_comptrig(esk2_0))|r1_xxreal_0(k1_comptrig(esk2_0),k32_sin_cos)|~v1_xreal_0(k1_comptrig(esk2_0)))).#trainneg
cnf(i_0_23538, plain, (k7_real_1(X1,k6_numbers)=k2_xcmplx_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23489, negated_conjecture, (v1_xxreal_0(esk1_0)|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_768, plain, (v1_xxreal_0(k2_rcomp_1(X1,X2))|~v6_membered(k1_zfmisc_1(k1_numbers))|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_23521, plain, (v1_xxreal_0(k2_xcmplx_0(X1,k6_numbers))|~v1_xreal_0(X1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_557, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3)),X4)=k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_23494, negated_conjecture, (v1_xboole_0(esk1_0)|v2_xxreal_0(esk1_0)|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_769, plain, (v1_xcmplx_0(k4_xcmplx_0(k7_real_1(X1,X2)))|~v5_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_5566, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(k2_xcmplx_0(X2,X3)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X2,X3)),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_559, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))=k2_xcmplx_0(X4,k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3)))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_560, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3)),X4)=k2_xcmplx_0(X3,k2_xcmplx_0(k2_xcmplx_0(X1,X2),X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_22983, negated_conjecture, (k4_complex1(esk2_0)=k6_numbers|v3_xxreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k2_xcmplx_0(esk1_0,k6_numbers))|~v1_xreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_22989, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k32_sin_cos)|~v1_xreal_0(k2_xcmplx_0(esk1_0,k6_numbers))|~v1_xreal_0(k1_comptrig(esk2_0)))).#trainneg
cnf(i_0_24557, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v5_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_24526, negated_conjecture, (esk1_0=k6_numbers|v2_xxreal_0(esk1_0)|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_24579, negated_conjecture, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,X2)),k1_comptrig(k1_comptrig(esk2_0)))|r1_xxreal_0(k1_comptrig(k1_comptrig(esk2_0)),k1_comptrig(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_24580, negated_conjecture, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,X2)),k1_comptrig(k1_comptrig(esk1_0)))|r1_xxreal_0(k1_comptrig(k1_comptrig(esk1_0)),k1_comptrig(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_561, plain, (m1_subset_1(k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3)),k1_numbers)|~v1_xreal_0(X3)|~m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_23912, plain, (k2_xcmplx_0(X1,k3_complex1(X2))=k2_xcmplx_0(X1,X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_16956, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v2_xxreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_14183, plain, (v2_xxreal_0(X1)|r1_xxreal_0(k3_complex1(X1),k6_numbers)|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_22903, plain, (r1_xxreal_0(X1,k32_sin_cos)|~v3_xxreal_0(X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_25811, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v1_xreal_0(X1)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_23172, negated_conjecture, (~v2_xxreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk2_0)))).#trainneg
cnf(i_0_25814, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(esk2_0)),k1_comptrig(k2_xcmplx_0(X1,esk2_0)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,esk2_0)),k1_comptrig(k1_comptrig(esk2_0)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_25815, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(esk2_0)),k1_comptrig(k2_xcmplx_0(X1,esk1_0)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,esk1_0)),k1_comptrig(k1_comptrig(esk2_0)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_563, plain, (m1_subset_1(k3_complex1(k2_xcmplx_0(X2,k2_xcmplx_0(X1,X3))),k1_numbers)|~v1_xreal_0(k3_complex1(X3))|~m1_subset_1(k3_complex1(k2_xcmplx_0(X1,X2)),k1_numbers)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_26051, plain, (k3_complex1(k2_xcmplx_0(X1,X2))=k2_xcmplx_0(k3_complex1(X1),X2)|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(k3_complex1(X1),X2))|~v1_xreal_0(k3_complex1(X1)))).#trainneg
cnf(i_0_26143, plain, (k2_xcmplx_0(k3_complex1(X1),k3_complex1(X2))=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(k3_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_26153, plain, (k2_xcmplx_0(k3_complex1(X2),k3_complex1(X1))=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(k3_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_23477, negated_conjecture, (r1_xxreal_0(k6_numbers,esk1_0)|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_25845, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(esk2_0)),k1_comptrig(k2_xcmplx_0(X1,k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,k32_sin_cos)),k1_comptrig(k1_comptrig(esk2_0)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_25847, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(esk1_0)),k1_comptrig(k2_xcmplx_0(X1,esk2_0)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,esk2_0)),k1_comptrig(k1_comptrig(esk1_0)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_565, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,k2_xcmplx_0(X2,X4)))=k2_xcmplx_0(X4,k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_567, plain, (k2_xcmplx_0(k3_complex1(X2),k2_xcmplx_0(k3_complex1(X1),X3))=k2_xcmplx_0(k3_complex1(k2_xcmplx_0(X1,X2)),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k3_complex1(X1))|~v1_xcmplx_0(k3_complex1(X2))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_26156, plain, (k2_xcmplx_0(k3_complex1(X1),X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k3_complex1(X1))|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_18823, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X3,k2_xcmplx_0(X4,X2)))=k2_xcmplx_0(X4,k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_29582, plain, (k2_xcmplx_0(X2,k2_xcmplx_0(X1,k2_xcmplx_0(X3,k6_numbers)))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_29309, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4)))=k2_xcmplx_0(X4,k2_xcmplx_0(X3,k2_xcmplx_0(X2,X1)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X4))).#trainneg
cnf(i_0_4083, plain, (k7_real_1(k6_numbers,X1)=k2_xcmplx_0(k6_numbers,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_25848, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(esk1_0)),k1_comptrig(k2_xcmplx_0(X1,esk1_0)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,esk1_0)),k1_comptrig(k1_comptrig(esk1_0)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_25878, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(esk1_0)),k1_comptrig(k2_xcmplx_0(X1,k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,k32_sin_cos)),k1_comptrig(k1_comptrig(esk1_0)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_572, plain, (k2_xcmplx_0(k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2)),X4)=k2_xcmplx_0(k2_xcmplx_0(X2,X3),k2_xcmplx_0(X1,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_5905, plain, (k3_complex1(X1)=k2_xcmplx_0(X1,k6_numbers)|~v1_xcmplx_0(k3_complex1(X1))|~v1_xreal_0(k2_xcmplx_0(X1,k6_numbers))|~v1_xreal_0(k3_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_17887, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0)))).#trainneg
cnf(i_0_30894, plain, (v1_xxreal_0(k2_xcmplx_0(k6_numbers,X1))|~v1_xreal_0(X1)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_23474, negated_conjecture, (~r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|~r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_24610, plain, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,X2)),k1_comptrig(k1_comptrig(k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k1_comptrig(k32_sin_cos)),k1_comptrig(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_20539, plain, (k4_complex1(k2_xcmplx_0(X2,X1))=k6_numbers|~v1_xreal_0(k2_xcmplx_0(k2_xcmplx_0(X2,k6_numbers),X1))|~v1_xcmplx_0(k2_xcmplx_0(X2,k6_numbers))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_9455, plain, (v1_xreal_0(k4_complex1(X1))|~v1_xcmplx_0(k3_complex1(X1))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_14186, plain, (v3_xxreal_0(X1)|r1_xxreal_0(k4_xcmplx_0(X1),k6_numbers)|~v1_xreal_0(X1)|~v1_xreal_0(k4_xcmplx_0(X1)))).#trainneg
cnf(i_0_19671, plain, (k2_xcmplx_0(X2,k3_complex1(X1))=k2_xcmplx_0(X1,X2)|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_31839, plain, (v1_xreal_0(k4_complex1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_30880, negated_conjecture, (v7_ordinal1(esk1_0)|~v6_membered(k1_numbers)|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_801, plain, (v3_xxreal_0(k16_complex1(X1))|v2_xxreal_0(k4_complex1(X2))|r1_xxreal_0(k4_complex1(X2),k16_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_31493, negated_conjecture, (~v2_xxreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk2_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0)))).#trainneg
cnf(i_0_31501, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(k32_sin_cos)),k1_comptrig(k2_xcmplx_0(X1,esk2_0)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,esk2_0)),k1_comptrig(k1_comptrig(k32_sin_cos)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_31502, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(k32_sin_cos)),k1_comptrig(k2_xcmplx_0(X1,esk1_0)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,esk1_0)),k1_comptrig(k1_comptrig(k32_sin_cos)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_599, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(k17_complex1(X3))|r1_xxreal_0(k17_complex1(X3),k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_600, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(k16_complex1(X3))|r1_xxreal_0(k16_complex1(X3),k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_31494, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0)))).#trainneg
cnf(i_0_31564, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k4_complex1(X2))|r1_xxreal_0(k4_complex1(X2),X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_24120, negated_conjecture, (r1_xxreal_0(esk1_0,esk1_0)|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_802, plain, (v3_xxreal_0(k16_complex1(X1))|v2_xxreal_0(k17_complex1(X2))|r1_xxreal_0(k17_complex1(X2),k16_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_30879, negated_conjecture, (v1_int_1(esk1_0)|~v5_membered(k1_numbers)|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_31532, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(k32_sin_cos)),k1_comptrig(k2_xcmplx_0(X1,k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,k32_sin_cos)),k1_comptrig(k1_comptrig(k32_sin_cos)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_14833, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(X1))),k1_comptrig(k1_comptrig(k1_comptrig(X2))))|r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(X2))),k1_comptrig(k1_comptrig(k1_comptrig(X1))))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_601, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(k2_xcmplx_0(X3,X4))|r1_xxreal_0(k2_xcmplx_0(X3,X4),k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X4)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_602, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(k1_comptrig(X3))|r1_xxreal_0(k1_comptrig(X3),k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_32475, plain, (r1_xxreal_0(k4_complex1(X1),k6_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_32454, plain, (v2_xxreal_0(k4_complex1(X1))|r1_xxreal_0(k4_complex1(X1),k32_sin_cos)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_32157, negated_conjecture, (k4_complex1(esk2_0)=k6_numbers|v3_xxreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk2_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0))|~v1_xxreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_32103, plain, (v1_xcmplx_0(k4_complex1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_9469, plain, (v1_xreal_0(k4_complex1(k2_xcmplx_0(X1,k6_numbers)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_803, plain, (v3_xxreal_0(k16_complex1(X1))|v2_xxreal_0(k16_complex1(X2))|r1_xxreal_0(k16_complex1(X2),k16_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_4136, plain, (k7_real_1(X1,k6_numbers)=k2_xcmplx_0(k6_numbers,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_32540, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(X1))),k1_comptrig(k1_comptrig(k1_comptrig(esk2_0))))|r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(esk2_0))),k1_comptrig(k1_comptrig(k1_comptrig(X1))))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_32541, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(X1))),k1_comptrig(k1_comptrig(k1_comptrig(esk1_0))))|r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(esk1_0))),k1_comptrig(k1_comptrig(k1_comptrig(X1))))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_32282, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k17_complex1(X2))|r1_xxreal_0(k17_complex1(X2),X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_32104, plain, (v3_xxreal_0(k4_complex1(X1))|v2_xxreal_0(X2)|r1_xxreal_0(X2,k4_complex1(X1))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_31865, plain, (k3_complex1(k2_xcmplx_0(X2,X1))=k2_xcmplx_0(X2,k3_complex1(X1))|~v1_xreal_0(k3_complex1(X1))|~m1_subset_1(k3_complex1(X2),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(X2))).#trainneg
cnf(i_0_806, plain, (v3_xxreal_0(k16_complex1(X1))|v2_xxreal_0(k1_comptrig(X2))|r1_xxreal_0(k1_comptrig(X2),k16_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_32505, negated_conjecture, (v1_int_1(esk1_0)|~v1_xreal_0(esk2_0)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_32806, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k1_comptrig(X2))|r1_xxreal_0(k1_comptrig(X2),X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_609, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k2_xcmplx_0(X3,k2_xcmplx_0(X2,X1)))|~v2_xxreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_611, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1)))|~v2_xxreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_32930, plain, (r1_xxreal_0(k4_complex1(X1),k32_sin_cos)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_32911, negated_conjecture, (v3_xxreal_0(k4_complex1(esk1_0))|r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos))).#trainneg
cnf(i_0_32918, negated_conjecture, (v3_xxreal_0(k4_complex1(esk2_0))|r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos))).#trainneg
cnf(i_0_32906, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_32954, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_9479, plain, (v1_xreal_0(k4_complex1(k2_xcmplx_0(k6_numbers,X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_32953, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_32802, plain, (v2_xxreal_0(k1_comptrig(X1))|r1_xxreal_0(k1_comptrig(X1),k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_32955, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|~v1_xreal_0(esk1_0))).#trainneg
cnf(i_0_32907, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_810, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_32571, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(X1))),k1_comptrig(k1_comptrig(k1_comptrig(k32_sin_cos))))|r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(k32_sin_cos))),k1_comptrig(k1_comptrig(k1_comptrig(X1))))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_623, plain, (v3_xxreal_0(k7_real_1(X1,X2))|v2_xxreal_0(X3)|r1_xxreal_0(X3,k7_real_1(X1,X2))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_33456, plain, (v2_xxreal_0(k1_comptrig(X1))|r1_xxreal_0(k1_comptrig(X1),k32_sin_cos)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_33842, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_33843, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_33844, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers)))).#trainneg
cnf(i_0_33837, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|~v1_xreal_0(esk1_0))).#trainneg
cnf(i_0_33836, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_33841, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_24121, negated_conjecture, (r1_xxreal_0(k32_sin_cos,esk1_0)|r1_xxreal_0(esk1_0,k32_sin_cos)|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_835, plain, (v3_xxreal_0(k16_complex1(X1))|v2_xxreal_0(k6_numbers)|r1_xxreal_0(k6_numbers,k16_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_32603, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k2_xcmplx_0(X2,X3))|r1_xxreal_0(k2_xcmplx_0(X2,X3),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_624, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(k7_real_1(X3,X4))|r1_xxreal_0(k7_real_1(X3,X4),k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(X4)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_33855, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_33838, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|~v1_xreal_0(k2_xcmplx_0(esk1_0,k6_numbers)))).#trainneg
cnf(i_0_33444, plain, (v3_xxreal_0(k4_complex1(X1))|v2_xxreal_0(k1_comptrig(X2))|r1_xxreal_0(k1_comptrig(X2),k4_complex1(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_33861, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v2_xxreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_33840, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~m1_subset_1(k4_complex1(esk2_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0)))).#trainneg
cnf(i_0_33804, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2)))|~v2_xxreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_33845, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk2_0)))).#trainneg
cnf(i_0_33834, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|~m1_subset_1(k4_complex1(esk1_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0)))).#trainneg
cnf(i_0_627, plain, (v1_xreal_0(k3_complex1(k2_xcmplx_0(X2,X1)))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_34304, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_34524, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0)))).#trainneg
cnf(i_0_34712, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0)))).#trainneg
cnf(i_0_866, plain, (k3_xcmplx_0(k6_numbers,k3_xcmplx_0(k7_complex1,X1))=k3_xcmplx_0(k6_numbers,X1)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k7_complex1))).#trainneg
cnf(i_0_15631, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k2_xcmplx_0(k1_comptrig(k1_comptrig(X2)),k32_sin_cos))|r1_xxreal_0(k2_xcmplx_0(k1_comptrig(k1_comptrig(X2)),k32_sin_cos),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_34303, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|~v2_xxreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_867, plain, (k4_complex1(k2_xcmplx_0(X1,k6_numbers))=k6_numbers|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23569, plain, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k32_sin_cos,k32_sin_cos)),k1_comptrig(k2_xcmplx_0(X1,X2)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,X2)),k1_comptrig(k2_xcmplx_0(k32_sin_cos,k32_sin_cos)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_23604, plain, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos)),k1_comptrig(k1_comptrig(X2)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X2)),k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_690, plain, (r1_xxreal_0(X1,k7_real_1(X2,X3))|r1_xxreal_0(k7_real_1(X2,X3),X1)|~v1_xxreal_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_32958, negated_conjecture, (k4_complex1(esk2_0)=k6_numbers|~m1_subset_1(k4_complex1(esk2_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk2_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk1_0)))).#trainneg
cnf(i_0_33964, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_33445, plain, (v3_xxreal_0(k4_complex1(X1))|v2_xxreal_0(k1_comptrig(X2))|r1_xxreal_0(k1_comptrig(X2),k4_complex1(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_871, plain, (v1_xxreal_0(k7_real_1(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_691, plain, (r1_xxreal_0(X1,k7_real_1(X2,X3))|r1_xxreal_0(k7_real_1(X2,X3),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_32904, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~m1_subset_1(k4_complex1(esk2_0),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0)))).#trainneg
cnf(i_0_32910, plain, (r1_xxreal_0(k4_complex1(X1),k32_sin_cos)|~v3_xxreal_0(k4_complex1(X1))|~v1_xxreal_0(k4_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_33835, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_33839, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk1_0)))).#trainneg
cnf(i_0_32437, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|~v1_xreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk1_0)))).#trainneg
cnf(i_0_32909, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk2_0)))).#trainneg
cnf(i_0_6921, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(esk2_0)),k1_comptrig(k1_comptrig(esk1_0)))|r1_xxreal_0(k1_comptrig(k1_comptrig(esk1_0)),k1_comptrig(k1_comptrig(esk2_0))))).#trainneg
cnf(i_0_694, plain, (v1_xxreal_0(k3_complex1(k2_xcmplx_0(X2,X1)))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v2_membered(k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_35360, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos))).#trainneg
cnf(i_0_35175, negated_conjecture, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,esk2_0)),k1_comptrig(k2_xcmplx_0(k32_sin_cos,k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k32_sin_cos,k32_sin_cos)),k1_comptrig(k2_xcmplx_0(X1,esk2_0)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_35176, negated_conjecture, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,esk1_0)),k1_comptrig(k2_xcmplx_0(k32_sin_cos,k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k32_sin_cos,k32_sin_cos)),k1_comptrig(k2_xcmplx_0(X1,esk1_0)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_35001, plain, (r1_xxreal_0(X1,k32_sin_cos)|r1_xxreal_0(k6_numbers,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_35371, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos)|r1_xxreal_0(k4_complex1(esk2_0),k6_numbers))).#trainneg
cnf(i_0_35359, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|r1_xxreal_0(k4_complex1(esk1_0),k6_numbers))).#trainneg
cnf(i_0_35329, negated_conjecture, (r1_xxreal_0(k4_complex1(esk1_0),k6_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0)))).#trainneg
cnf(i_0_32956, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|~v1_xreal_0(k2_xcmplx_0(esk1_0,k6_numbers)))).#trainneg
cnf(i_0_883, plain, (r1_xxreal_0(k17_complex1(X1),k17_complex1(X1))|~v2_membered(k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_35302, plain, (v1_xxreal_0(k2_xcmplx_0(X1,k6_numbers))|~v6_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_32957, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk1_0)))).#trainneg
cnf(i_0_35208, plain, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,k32_sin_cos)),k1_comptrig(k2_xcmplx_0(k32_sin_cos,k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k32_sin_cos,k32_sin_cos)),k1_comptrig(k2_xcmplx_0(X1,k32_sin_cos)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_35210, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(k2_xcmplx_0(k1_comptrig(esk2_0),k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k1_comptrig(esk2_0),k32_sin_cos)),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_34185, plain, (v1_xcmplx_0(k2_xcmplx_0(k2_xcmplx_0(X2,X3),k2_xcmplx_0(X4,X1)))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_35784, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k6_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0)))).#trainneg
cnf(i_0_35771, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|r1_xxreal_0(k4_complex1(esk1_0),k32_sin_cos))).#trainneg
cnf(i_0_884, plain, (r1_xxreal_0(k7_real_1(X1,X2),k7_real_1(X1,X2))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v2_membered(k1_numbers))).#trainneg
cnf(i_0_35303, plain, (v1_xxreal_0(k2_xcmplx_0(k6_numbers,X1))|~v6_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_35211, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(k2_xcmplx_0(k1_comptrig(esk1_0),k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k1_comptrig(esk1_0),k32_sin_cos)),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_35243, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(k2_xcmplx_0(k1_comptrig(k32_sin_cos),k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k1_comptrig(k32_sin_cos),k32_sin_cos)),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_35803, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk1_0)))).#trainneg
cnf(i_0_36010, plain, (v1_xcmplx_0(k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X4,X3)))|~v1_xreal_0(k2_xcmplx_0(X4,X1))|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_942, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(k6_numbers,X3))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_36186, negated_conjecture, (~m1_subset_1(esk2_0,k1_numbers)|~r1_xxreal_0(k6_numbers,esk2_0)|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_6945, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(esk2_0)),k1_comptrig(k1_comptrig(k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k1_comptrig(k32_sin_cos)),k1_comptrig(k1_comptrig(esk2_0))))).#trainneg
cnf(i_0_704, plain, (v1_xcmplx_0(k3_complex1(k2_xcmplx_0(X2,X1)))|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_36921, plain, (k2_xcmplx_0(k6_numbers,k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_36188, negated_conjecture, (r1_xxreal_0(esk2_0,k32_sin_cos)|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_36256, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers|v3_xxreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers))|~v1_xreal_0(k1_comptrig(esk1_0)))).#trainneg
cnf(i_0_36257, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k32_sin_cos)|~v1_xreal_0(k2_xcmplx_0(esk2_0,k6_numbers))|~v1_xreal_0(k1_comptrig(esk1_0)))).#trainneg
cnf(i_0_36187, negated_conjecture, (r1_xxreal_0(esk2_0,k6_numbers)|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_6972, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(esk1_0)),k1_comptrig(k1_comptrig(k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k1_comptrig(k32_sin_cos)),k1_comptrig(k1_comptrig(esk1_0))))).#trainneg
cnf(i_0_36179, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k6_numbers)|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk2_0)))).#trainneg
cnf(i_0_32476, plain, (r1_xxreal_0(k4_complex1(X1),k6_numbers)|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(X1,k6_numbers)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_954, plain, (v3_xxreal_0(k1_comptrig(X1))|v2_xxreal_0(k4_complex1(X2))|r1_xxreal_0(k4_complex1(X2),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_36335, negated_conjecture, (~v2_xxreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk1_0)))).#trainneg
cnf(i_0_23543, negated_conjecture, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k32_sin_cos,k32_sin_cos)),k1_comptrig(esk2_0))|r1_xxreal_0(k1_comptrig(esk2_0),k1_comptrig(k2_xcmplx_0(k32_sin_cos,k32_sin_cos))))).#trainneg
cnf(i_0_714, plain, (m1_subset_1(k7_real_1(k4_complex1(X1),X3),k1_numbers)|~v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1))))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_34263, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k32_sin_cos)|~v1_xreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(esk1_0))).#trainneg
cnf(i_0_9236, plain, (v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_955, plain, (v3_xxreal_0(k1_comptrig(X1))|v2_xxreal_0(k17_complex1(X2))|r1_xxreal_0(k17_complex1(X2),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_23544, negated_conjecture, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k32_sin_cos,k32_sin_cos)),k1_comptrig(esk1_0))|r1_xxreal_0(k1_comptrig(esk1_0),k1_comptrig(k2_xcmplx_0(k32_sin_cos,k32_sin_cos))))).#trainneg
cnf(i_0_23573, plain, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k32_sin_cos,k32_sin_cos)),k1_comptrig(k32_sin_cos))|r1_xxreal_0(k1_comptrig(k32_sin_cos),k1_comptrig(k2_xcmplx_0(k32_sin_cos,k32_sin_cos))))).#trainneg
cnf(i_0_717, plain, (v1_xreal_0(k7_real_1(k4_complex1(X1),X3))|~v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1))))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_718, plain, (v1_int_1(k7_real_1(k4_complex1(X1),X3))|~v5_membered(k1_numbers)|~v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1))))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_33888, plain, (v1_xreal_0(k4_complex1(k2_xcmplx_0(k6_numbers,X1)))|~v1_xcmplx_0(k3_complex1(X1))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_23151, plain, (v1_xreal_0(k2_xcmplx_0(k2_xcmplx_0(X4,X2),k2_xcmplx_0(X1,X3)))|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4))).#trainneg
cnf(i_0_6201, plain, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(k6_numbers,X2))|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(k3_xcmplx_0(X1,k7_complex1),X2)))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,k7_complex1),X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,k7_complex1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_956, plain, (v3_xxreal_0(k1_comptrig(X1))|v2_xxreal_0(k16_complex1(X2))|r1_xxreal_0(k16_complex1(X2),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_37698, negated_conjecture, (~v2_xxreal_0(esk2_0)|~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_37709, negated_conjecture, (~v2_xxreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k1_comptrig(esk1_0))|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk2_0)))).#trainneg
cnf(i_0_719, plain, (v7_ordinal1(k7_real_1(k4_complex1(X1),X3))|~v6_membered(k1_numbers)|~v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1))))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_6202, plain, (~v2_xxreal_0(X1)|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,k6_numbers))|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(k3_xcmplx_0(X1,k7_complex1),X2)))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,k7_complex1),X2))|~v1_xcmplx_0(k3_xcmplx_0(X1,k7_complex1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_967, plain, (r1_xxreal_0(X1,k4_complex1(X2))|r1_xxreal_0(k4_complex1(X2),X1)|~v1_xxreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_732, plain, (k4_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(k3_xcmplx_0(X3,k7_complex1),X2)))=k7_real_1(k4_complex1(X1),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X3,k7_complex1),X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X3,k7_complex1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_37743, plain, (r1_xxreal_0(k4_complex1(X1),k6_numbers)|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(k6_numbers,X1)))|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_37981, plain, (v2_xxreal_0(X1)|r1_xxreal_0(X1,k1_comptrig(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|k1_comptrig(X1)!=k6_numbers)).#trainneg
cnf(i_0_3646, plain, (k2_xcmplx_0(k3_complex1(X1),X2)=k7_real_1(X1,X2)|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_968, plain, (r1_xxreal_0(X1,k4_complex1(X2))|r1_xxreal_0(k4_complex1(X2),X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_35798, plain, (r1_xxreal_0(X1,X1)|~v2_membered(k1_numbers)|~v1_xcmplx_0(X1)|k1_comptrig(X1)!=k6_numbers)).#trainneg
cnf(i_0_735, plain, (k4_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,k3_xcmplx_0(X4,k7_complex1)))))=k7_real_1(k4_complex1(k2_xcmplx_0(X1,X2)),X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X3,k3_xcmplx_0(X4,k7_complex1)))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_736, plain, (k4_complex1(k2_xcmplx_0(X2,k2_xcmplx_0(X1,k2_xcmplx_0(X3,k3_xcmplx_0(X4,k7_complex1)))))=k7_real_1(k4_complex1(k2_xcmplx_0(X1,X2)),X4)|~v1_xreal_0(X4)|~v1_xreal_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X3,k3_xcmplx_0(X4,k7_complex1)))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_38480, negated_conjecture, (m1_subset_1(esk2_0,k1_numbers)|~v1_xreal_0(esk1_0))).#trainneg
cnf(i_0_39430, plain, (r1_xxreal_0(X1,X1)|k1_comptrig(X1)!=k6_numbers|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_38705, plain, (r1_xxreal_0(k4_complex1(X1),k6_numbers)|r1_xxreal_0(k6_numbers,k4_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_38713, plain, (r1_xxreal_0(k4_complex1(X1),k1_comptrig(X2))|r1_xxreal_0(k1_comptrig(X2),k4_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_38605, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk2_0))|~v1_xreal_0(k1_comptrig(esk1_0)))).#trainneg
cnf(i_0_740, plain, (k4_complex1(k2_xcmplx_0(k3_xcmplx_0(X3,k7_complex1),k2_xcmplx_0(X1,X2)))=k7_real_1(k4_complex1(X1),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X3,k7_complex1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_741, plain, (k4_complex1(k2_xcmplx_0(X2,k2_xcmplx_0(k3_xcmplx_0(X3,k7_complex1),X1)))=k7_real_1(k4_complex1(X1),X3)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X3,k7_complex1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_8872, plain, (k2_xcmplx_0(X2,k3_complex1(X1))=k2_xcmplx_0(k6_numbers,k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k3_complex1(X1))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_12886, plain, (k2_xcmplx_0(k6_numbers,k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_34310, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|~v2_xxreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_39898, plain, (r1_xxreal_0(X1,X1)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers)|~r1_xxreal_0(k6_numbers,X1))).#trainneg
cnf(i_0_994, plain, (k4_complex1(k2_xcmplx_0(k6_numbers,X1))=k6_numbers|~v1_xreal_0(X1))).#trainneg
cnf(i_0_41100, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|k4_complex1(esk1_0)=k6_numbers|v3_xxreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk2_0)))).#trainneg
cnf(i_0_41171, plain, (k6_numbers=X1|~v1_xreal_0(X1)|~v1_xreal_0(k3_xcmplx_0(X1,k7_complex1)))).#trainneg
cnf(i_0_772, plain, (v1_int_1(k3_complex1(k2_xcmplx_0(X2,X1)))|~v5_membered(k1_numbers)|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_33202, plain, (v3_xxreal_0(k4_complex1(X1))|r1_xxreal_0(k6_numbers,k4_complex1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_31390, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,k6_numbers),k2_xcmplx_0(X1,X3))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2)))).#trainneg
cnf(i_0_37869, plain, (k3_complex1(k2_xcmplx_0(X1,X2))=k2_xcmplx_0(k3_complex1(X1),X2)|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(k3_complex1(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_17986, plain, (k3_complex1(k2_xcmplx_0(k6_numbers,X1))=k2_xcmplx_0(X1,k6_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_33038, plain, (v1_xreal_0(k4_complex1(k2_xcmplx_0(X1,k6_numbers)))|~v1_xcmplx_0(k3_complex1(X1))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_19924, plain, (k4_complex1(k2_xcmplx_0(X2,X1))=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X2,k3_complex1(X1)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k3_complex1(X1))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_7125, plain, (k2_xcmplx_0(k6_numbers,k2_xcmplx_0(X2,X1))=k2_xcmplx_0(X1,X2)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1004, plain, (r1_xxreal_0(X1,k17_complex1(X2))|r1_xxreal_0(k17_complex1(X2),X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_780, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,k3_complex1(k2_xcmplx_0(X2,X3))))|~v1_xcmplx_0(k3_complex1(X3))|~v1_xcmplx_0(k3_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_complex1(X3))|~m1_subset_1(k3_complex1(X2),k1_numbers)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_783, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4))))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_42521, negated_conjecture, (~v1_xreal_0(esk2_0))).#trainneg
cnf(i_0_39965, negated_conjecture, (k1_comptrig(esk1_0)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk2_0)))).#trainneg
cnf(i_0_1005, plain, (r1_xxreal_0(X1,k16_complex1(X2))|r1_xxreal_0(k16_complex1(X2),X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_42513, negated_conjecture, (~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk2_0)))).#trainneg
cnf(i_0_1006, plain, (r1_xxreal_0(X1,k7_real_1(X2,X3))|r1_xxreal_0(k7_real_1(X2,X3),X1)|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_42794, plain, (k3_complex1(k2_xcmplx_0(X2,X1))=k2_xcmplx_0(X1,X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1048, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k1_comptrig(k4_xcmplx_0(k4_xcmplx_0(X1)))),k6_numbers)|~v1_int_1(X1))).#trainneg
cnf(i_0_804, plain, (v3_xxreal_0(k16_complex1(X1))|v2_xxreal_0(k7_real_1(X2,X3))|r1_xxreal_0(k7_real_1(X2,X3),k16_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_43790, plain, (k3_complex1(X1)=k2_xcmplx_0(k6_numbers,X1)|~v1_xreal_0(k2_xcmplx_0(k6_numbers,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_42508, plain, (k4_complex1(k2_xcmplx_0(X1,X2))=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1049, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k1_comptrig(k4_xcmplx_0(X1))),k6_numbers)|v3_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_805, plain, (v3_xxreal_0(k16_complex1(X1))|v2_xxreal_0(k2_xcmplx_0(X2,X3))|r1_xxreal_0(k2_xcmplx_0(X2,X3),k16_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_809, plain, (v3_xxreal_0(k16_complex1(X1))|v2_xxreal_0(k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4)))|r1_xxreal_0(k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4)),k16_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_44186, plain, (k4_complex1(k2_xcmplx_0(X2,X1))=k6_numbers|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1050, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k1_comptrig(k4_complex1(X1))),k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_811, plain, (v3_xxreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|v2_xxreal_0(X4)|r1_xxreal_0(X4,k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xreal_0(X4)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_812, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(k2_xcmplx_0(X3,k2_xcmplx_0(X4,X5)))|r1_xxreal_0(k2_xcmplx_0(X3,k2_xcmplx_0(X4,X5)),k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~v1_xreal_0(X5)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_44751, plain, (v3_xxreal_0(k2_xcmplx_0(k6_numbers,k2_xcmplx_0(X1,X2)))|v2_xxreal_0(X3)|r1_xxreal_0(X3,k2_xcmplx_0(k6_numbers,k2_xcmplx_0(X1,X2)))|~v1_xreal_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1051, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k1_comptrig(k17_complex1(X1))),k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_813, plain, (v1_xreal_0(k2_xcmplx_0(X1,k3_complex1(k2_xcmplx_0(X2,X3))))|~v1_xreal_0(k2_xcmplx_0(X1,k3_complex1(X2)))|~v1_xreal_0(k3_complex1(X3))|~v1_xcmplx_0(k3_complex1(X2))|~v1_xcmplx_0(X1)|~m1_subset_1(k3_complex1(X2),k1_numbers)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_816, plain, (v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4))))|~v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xreal_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_44103, plain, (k4_complex1(k7_real_1(X1,X2))=k6_numbers|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(X2)|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_42734, plain, (v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(k2_xcmplx_0(k6_numbers,X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1052, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k1_comptrig(k16_complex1(X1))),k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_43424, plain, (v1_xcmplx_0(k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_817, plain, (v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X3,k2_xcmplx_0(X2,X4))))|~v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xreal_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_821, plain, (v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X4,k2_xcmplx_0(X2,X3))))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_1057, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k1_comptrig(k1_comptrig(X1))),k6_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_822, plain, (v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X3,k2_xcmplx_0(X4,X2))))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_823, plain, (v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2)))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_46672, negated_conjecture, (v1_xboole_0(k1_comptrig(esk1_0))|v2_xxreal_0(k1_comptrig(esk1_0))|k1_comptrig(esk1_0)=k6_numbers|~v1_xreal_0(k1_comptrig(esk1_0)))).#trainneg
cnf(i_0_10926, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4)))=k2_xcmplx_0(X2,k2_xcmplx_0(k2_xcmplx_0(X4,X1),X3))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_826, plain, (v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4))))|~v1_xreal_0(k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3))|~v1_xreal_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_47148, plain, (v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(k2_xcmplx_0(k6_numbers,X2))|~v1_xreal_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_45747, plain, (k4_complex1(k7_real_1(X1,X2))=k6_numbers|~v1_xreal_0(k3_complex1(k2_xcmplx_0(X1,X2)))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_827, plain, (v1_xreal_0(k2_xcmplx_0(X2,k2_xcmplx_0(X1,k2_xcmplx_0(X3,X4))))|~v1_xreal_0(k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3))|~v1_xreal_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_828, plain, (v1_xreal_0(k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3)))).#trainneg
cnf(i_0_48660, plain, (v1_xreal_0(k2_xcmplx_0(k6_numbers,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_28455, plain, (k2_xcmplx_0(k3_complex1(X2),X1)=k2_xcmplx_0(k3_complex1(X1),X2)|~v1_xreal_0(k2_xcmplx_0(k3_complex1(X1),X2))|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(k3_complex1(X1))|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(k3_complex1(X2),X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1272, plain, (v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v2_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_21027, plain, (k2_xcmplx_0(k6_numbers,k2_xcmplx_0(X1,X2))=k7_real_1(X1,X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_38579, plain, (~v2_xxreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(k6_numbers,X2))|~v1_xreal_0(X1)|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(X2,k3_xcmplx_0(X1,k7_complex1))))|~v1_xcmplx_0(k3_xcmplx_0(X1,k7_complex1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_49403, plain, (v1_xreal_0(k2_xcmplx_0(X1,k6_numbers))|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_49560, plain, (v1_xxreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_49578, plain, (m1_subset_1(k7_real_1(X1,X2),k1_numbers)|~m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_49882, plain, (v1_xreal_0(k7_real_1(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_4135, plain, (k7_real_1(X1,k6_numbers)=k7_real_1(k6_numbers,X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_50325, plain, (v1_xxreal_0(X1)|~m1_subset_1(X1,k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_50395, plain, (v1_xxreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_34994, plain, (k1_comptrig(k6_numbers)=k6_numbers|r1_xxreal_0(k6_numbers,k1_comptrig(k6_numbers))|~v1_xreal_0(k1_comptrig(k6_numbers)))).#trainneg
cnf(i_0_874, plain, (v7_ordinal1(k3_complex1(k2_xcmplx_0(X2,X1)))|~v6_membered(k1_numbers)|~v1_xreal_0(k3_complex1(X2))|~m1_subset_1(k3_complex1(X1),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_50611, plain, (m1_subset_1(k4_complex1(k2_xcmplx_0(X1,X2)),k1_numbers)|~v1_xreal_0(k4_complex1(X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_50589, plain, (v1_xxreal_0(k17_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_50896, plain, (v1_xxreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_50750, plain, (k4_complex1(k2_xcmplx_0(X1,X2))=k6_numbers|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_51483, plain, (r1_xxreal_0(k17_complex1(X1),k17_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_51147, plain, (~v2_xxreal_0(k4_complex1(X1))|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_51325, plain, (k4_complex1(k2_xcmplx_0(X2,X1))=k4_complex1(X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_51370, plain, (m1_subset_1(k4_complex1(X1),k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_50655, plain, (~v2_xxreal_0(k7_real_1(k4_complex1(X1),k6_numbers))|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_52108, plain, (v1_xxreal_0(k4_complex1(X1))|~v2_membered(k1_numbers)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_894, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X4,k2_xcmplx_0(X2,X3))))|~v2_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_34996, plain, (k1_comptrig(k32_sin_cos)=k6_numbers|r1_xxreal_0(k6_numbers,k1_comptrig(k32_sin_cos))|~v1_xreal_0(k1_comptrig(k32_sin_cos)))).#trainneg
cnf(i_0_895, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X3,k2_xcmplx_0(X4,X2))))|~v2_xxreal_0(k2_xcmplx_0(X3,X4))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_51998, plain, (k4_complex1(X2)=k4_complex1(X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_33115, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(esk2_0))),k1_comptrig(k1_comptrig(k1_comptrig(esk2_0)))))).#trainneg
cnf(i_0_900, plain, (v3_xxreal_0(k2_xcmplx_0(X1,X2))|v2_xxreal_0(k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1))|~v2_xxreal_0(X3)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3)))).#trainneg
cnf(i_0_51441, plain, (m1_subset_1(k4_complex1(k2_xcmplx_0(X2,X1)),k1_numbers)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_51274, plain, (k4_complex1(k7_real_1(X1,X2))=k4_complex1(X1)|~v1_xreal_0(k3_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_51530, plain, (v1_xxreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,k6_numbers))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1186, plain, (v1_xcmplx_0(k7_real_1(k4_complex1(X1),k6_numbers))|~v1_xreal_0(X2)|~v1_xcmplx_0(k2_xcmplx_0(X2,k6_numbers))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_33150, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(esk1_0))),k1_comptrig(k1_comptrig(k1_comptrig(esk1_0)))))).#trainneg
cnf(i_0_34235, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(k32_sin_cos))),k1_comptrig(k1_comptrig(k1_comptrig(k32_sin_cos)))))).#trainneg
cnf(i_0_51226, plain, (k4_complex1(k3_complex1(X1))=k4_complex1(X1)|~v1_xcmplx_0(k3_complex1(X1))|~v1_xreal_0(X1)|~m1_subset_1(k3_complex1(X1),k1_numbers))).#trainneg
cnf(i_0_51557, plain, (v1_xxreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(k6_numbers,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_54028, plain, (v1_xcmplx_0(k7_real_1(k4_complex1(X1),k6_numbers))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1196, plain, (v1_xcmplx_0(k7_real_1(k4_complex1(X1),X2))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xcmplx_0(X1)|~v1_xreal_0(k3_xcmplx_0(X2,k7_complex1)))).#trainneg
cnf(i_0_904, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,k6_numbers))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1209, plain, (k3_complex1(k3_xcmplx_0(X1,k7_complex1))=k6_numbers|~v1_xreal_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X1,k7_complex1)))).#trainneg
cnf(i_0_35841, plain, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k32_sin_cos,k32_sin_cos)),k1_comptrig(k2_xcmplx_0(k32_sin_cos,k32_sin_cos))))).#trainneg
cnf(i_0_11942, plain, (r1_xxreal_0(k2_xcmplx_0(X1,k1_comptrig(X2)),k1_comptrig(X3))|r1_xxreal_0(k1_comptrig(X3),k2_xcmplx_0(X1,k1_comptrig(X2)))|~v1_xreal_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_905, plain, (k4_complex1(k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,k3_xcmplx_0(X4,k7_complex1))))=X4|~v1_xreal_0(X4)|~v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(k3_xcmplx_0(X4,k7_complex1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_906, plain, (k3_complex1(k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,k3_xcmplx_0(X4,k7_complex1))))=k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))|~v1_xreal_0(X4)|~v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(k3_xcmplx_0(X4,k7_complex1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_53619, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k32_sin_cos)|~m1_subset_1(esk1_0,k1_numbers))).#trainneg
cnf(i_0_53618, negated_conjecture, (r1_xxreal_0(k4_complex1(esk2_0),k6_numbers)|~m1_subset_1(esk1_0,k1_numbers))).#trainneg
cnf(i_0_52118, plain, (v1_xxreal_0(k4_complex1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_55010, plain, (r1_xxreal_0(k1_comptrig(X1),k2_xcmplx_0(k32_sin_cos,k1_comptrig(X2)))|r1_xxreal_0(k2_xcmplx_0(k32_sin_cos,k1_comptrig(X2)),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_55009, plain, (r1_xxreal_0(k1_comptrig(X1),k2_xcmplx_0(k1_comptrig(X2),k1_comptrig(X3)))|r1_xxreal_0(k2_xcmplx_0(k1_comptrig(X2),k1_comptrig(X3)),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_910, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,k2_xcmplx_0(X4,X5)))=k2_xcmplx_0(X5,k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X4))|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(k2_xcmplx_0(X4,X5))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_911, plain, (v1_xcmplx_0(k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,k2_xcmplx_0(X4,X5))))|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(k2_xcmplx_0(X4,X5))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_53669, plain, (k4_complex1(X2)=k6_numbers|~v1_xreal_0(X1)|~v1_xcmplx_0(X2)|~m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers))).#trainneg
cnf(i_0_57630, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(X2)|~m1_subset_1(k2_xcmplx_0(X1,X2),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_50512, plain, (v1_xcmplx_0(k7_real_1(X1,k6_numbers))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_53612, negated_conjecture, (k4_complex1(esk2_0)=k6_numbers|~v1_xreal_0(k4_complex1(esk2_0))|~r1_xxreal_0(k6_numbers,k4_complex1(esk2_0))|~m1_subset_1(esk1_0,k1_numbers))).#trainneg
cnf(i_0_55696, plain, (r1_xxreal_0(k2_xcmplx_0(k32_sin_cos,k1_comptrig(X1)),k1_comptrig(k1_comptrig(X2)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X2)),k2_xcmplx_0(k32_sin_cos,k1_comptrig(X1)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_915, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,k2_xcmplx_0(X4,X5))))=k2_xcmplx_0(k2_xcmplx_0(X3,X4),k2_xcmplx_0(X5,k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(k2_xcmplx_0(X3,k2_xcmplx_0(X4,X5)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_52192, plain, (k4_complex1(k4_complex1(X1))=k6_numbers|~v1_xreal_0(k4_complex1(k2_xcmplx_0(X1,X2)))|~v1_xreal_0(k4_complex1(X2))|~v1_xcmplx_0(k4_complex1(X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_52273, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X3))|~v1_xreal_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_57757, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~m1_subset_1(X3,k1_numbers)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_51154, plain, (~v2_xxreal_0(k4_complex1(X1))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_50558, plain, (v1_xcmplx_0(k7_real_1(k6_numbers,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1234, plain, (k3_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k6_numbers)))=k2_xcmplx_0(X1,X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_57873, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k2_xcmplx_0(k32_sin_cos,k1_comptrig(k1_comptrig(X2))))|r1_xxreal_0(k2_xcmplx_0(k32_sin_cos,k1_comptrig(k1_comptrig(X2))),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_917, plain, (v1_xcmplx_0(k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4)))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_918, plain, (v1_xreal_0(k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4)))|~v1_xreal_0(X4)|~v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_59278, plain, (k4_complex1(k4_complex1(X1))=k6_numbers|~v1_xcmplx_0(X1)|~v1_int_1(X2))).#trainneg
cnf(i_0_59268, plain, (k4_complex1(k4_complex1(X1))=k6_numbers|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_59422, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X3,X2))|~v1_xreal_0(k2_xcmplx_0(X1,X3))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_1252, plain, (m1_subset_1(X1,k1_numbers)|r1_xxreal_0(X1,X2)|r1_xxreal_0(k6_numbers,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_55598, negated_conjecture, (~v2_xxreal_0(k4_complex1(esk2_0))|~v1_xreal_0(k4_complex1(esk2_0))|~m1_subset_1(esk1_0,k1_numbers))).#trainneg
cnf(i_0_919, plain, (v3_xxreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|v2_xxreal_0(k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4)))|~v2_xxreal_0(X4)|~v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xreal_0(X4)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_920, plain, (v3_xxreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|v3_xxreal_0(X4)|~v3_xxreal_0(k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4)))|~v1_xreal_0(X4)|~v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_60956, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_1257, plain, (m1_subset_1(X1,k1_numbers)|r1_xxreal_0(X1,X2)|r1_xxreal_0(k32_sin_cos,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_921, plain, (k2_xcmplx_0(k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4)),X5)=k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),k2_xcmplx_0(X4,X5))|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1258, plain, (m1_subset_1(X1,k1_numbers)|r1_xxreal_0(X1,X2)|r1_xxreal_0(X3,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_60984, plain, (m1_subset_1(X1,k1_numbers)|r1_xxreal_0(k6_numbers,X1)|r1_xxreal_0(X1,k32_sin_cos)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_924, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),k2_xcmplx_0(X4,X5))=k2_xcmplx_0(X5,k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4)))|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_925, plain, (k2_xcmplx_0(k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4)),X5)=k2_xcmplx_0(X4,k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),X5))|~v1_xcmplx_0(X5)|~v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_62550, plain, (m1_subset_1(X1,k1_numbers)|r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X3)|~v1_xreal_0(X2)|~v1_xxreal_0(X1)|~v1_xreal_0(X3))).#trainneg
cnf(i_0_53730, plain, (k4_complex1(k2_xcmplx_0(X1,X2))=k6_numbers|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_50569, plain, (v1_xreal_0(k7_real_1(k6_numbers,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_51163, plain, (r1_xxreal_0(k4_complex1(X1),k6_numbers)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_52257, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_62551, negated_conjecture, (r1_xxreal_0(esk2_0,k32_sin_cos)|r1_xxreal_0(k6_numbers,esk2_0)|~v1_xxreal_0(esk2_0))).#trainneg
cnf(i_0_62559, plain, (v1_int_1(X1)|r1_xxreal_0(X1,k32_sin_cos)|r1_xxreal_0(k6_numbers,X1)|~v5_membered(k1_numbers)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_926, plain, (m1_subset_1(k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4)),k1_numbers)|~v1_xreal_0(X4)|~m1_subset_1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),k1_numbers)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_927, plain, (m1_subset_1(k4_complex1(k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))),k1_numbers)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1274, plain, (v1_int_1(k2_xcmplx_0(X1,X2))|~v5_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_65412, plain, (v1_int_1(X1)|r1_xxreal_0(k6_numbers,X1)|r1_xxreal_0(X1,k32_sin_cos)|~v1_xxreal_0(X1)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_928, plain, (m1_subset_1(k3_complex1(k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))),k1_numbers)|~v1_xreal_0(k3_complex1(X4))|~m1_subset_1(k3_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3))),k1_numbers)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1275, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v6_membered(k1_numbers)|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_930, plain, (v3_xxreal_0(X1)|v2_xxreal_0(k2_xcmplx_0(k2_xcmplx_0(X2,X3),k2_xcmplx_0(X4,X1)))|~v2_xxreal_0(k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4)))|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_931, plain, (k2_xcmplx_0(X1,k2_xcmplx_0(k2_xcmplx_0(X2,X3),k2_xcmplx_0(X4,X5)))=k2_xcmplx_0(X5,k2_xcmplx_0(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4))))|~v1_xcmplx_0(X5)|~v1_xcmplx_0(k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_60800, plain, (k4_complex1(k4_complex1(X1))=k6_numbers|~v1_xcmplx_0(k7_real_1(X1,X2))|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1276, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_932, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(k2_xcmplx_0(X2,X3),k2_xcmplx_0(X4,X5))))|~v1_xcmplx_0(X5)|~v1_xcmplx_0(k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_936, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))=k2_xcmplx_0(X1,k2_xcmplx_0(k2_xcmplx_0(X2,X3),X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_68606, plain, (k4_complex1(k4_complex1(X1))=k6_numbers|~v1_xreal_0(X1))).#trainneg
cnf(i_0_52272, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_60959, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X4))|~v1_xreal_0(X3)|~v1_xcmplx_0(X4))).#trainneg
cnf(i_0_51906, plain, (~v2_xxreal_0(k4_complex1(X1))|~v1_xreal_0(X2)|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(X2,X1)))|~v1_xcmplx_0(k2_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_19513, negated_conjecture, (k1_comptrig(esk2_0)=k6_numbers|~v3_xxreal_0(k1_comptrig(esk2_0))|~v1_xreal_0(k1_comptrig(esk2_0)))).#trainneg
cnf(i_0_5567, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(k2_xcmplx_0(X2,X3)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X2,X3)),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_937, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))=k2_xcmplx_0(k2_xcmplx_0(X2,X3),k2_xcmplx_0(X1,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_940, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4))),X5)=k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),k2_xcmplx_0(X4,X5))|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_51280, plain, (k4_complex1(k2_xcmplx_0(X1,X2))=k4_complex1(X1)|~v1_xreal_0(k2_xcmplx_0(k6_numbers,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1293, plain, (v3_xxreal_0(k4_complex1(X1))|v2_xxreal_0(k6_numbers)|r1_xxreal_0(k6_numbers,k4_complex1(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_70479, plain, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,k32_sin_cos)),k1_comptrig(k1_comptrig(X2)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X2)),k1_comptrig(k2_xcmplx_0(X1,k32_sin_cos)))|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_941, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X3,k2_xcmplx_0(X2,X4))),X5)=k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)),k2_xcmplx_0(X4,X5))|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_943, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X4,k2_xcmplx_0(X2,X3))),X5)=k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(k2_xcmplx_0(X3,X4),X5))|~v1_xcmplx_0(X5)|~v1_xcmplx_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_53978, plain, (k4_complex1(k7_real_1(X2,X1))=k4_complex1(X1)|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_51151, plain, (k4_complex1(X1)=k6_numbers|v3_xxreal_0(k4_complex1(X1))|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1294, plain, (v3_xxreal_0(k4_complex1(X1))|v2_xxreal_0(k4_complex1(X2))|r1_xxreal_0(k4_complex1(X2),k4_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_13514, plain, (r1_xxreal_0(k1_comptrig(X1),k2_xcmplx_0(k2_xcmplx_0(X2,X3),k32_sin_cos))|r1_xxreal_0(k2_xcmplx_0(k2_xcmplx_0(X2,X3),k32_sin_cos),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_944, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X3,k2_xcmplx_0(X4,X2))),X5)=k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(k2_xcmplx_0(X3,X4),X5))|~v1_xcmplx_0(X5)|~v1_xcmplx_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_945, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X3,X2)),X4)=k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_74887, plain, (k4_complex1(k2_xcmplx_0(X1,X2))=k4_complex1(X2)|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_74889, plain, (k4_complex1(k2_xcmplx_0(X2,X1))=k4_complex1(X2)|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_72854, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(k6_numbers,X2)))).#trainneg
cnf(i_0_1295, plain, (v3_xxreal_0(k4_complex1(X1))|v2_xxreal_0(k17_complex1(X2))|r1_xxreal_0(k17_complex1(X2),k4_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_35247, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|~v2_xxreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_75156, plain, (r1_xxreal_0(k2_xcmplx_0(k2_xcmplx_0(X1,k32_sin_cos),k32_sin_cos),k1_comptrig(X2))|r1_xxreal_0(k1_comptrig(X2),k2_xcmplx_0(k2_xcmplx_0(X1,k32_sin_cos),k32_sin_cos))|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_6458, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(k1_comptrig(k2_xcmplx_0(X2,X3))))|r1_xxreal_0(k1_comptrig(k1_comptrig(k2_xcmplx_0(X2,X3))),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_947, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4))),X5)=k2_xcmplx_0(k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3),k2_xcmplx_0(X4,X5))|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_77360, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_77225, plain, (k4_complex1(k2_xcmplx_0(k6_numbers,X1))=k4_complex1(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1296, plain, (v3_xxreal_0(k4_complex1(X1))|v2_xxreal_0(k16_complex1(X2))|r1_xxreal_0(k16_complex1(X2),k4_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_77473, plain, (r1_xxreal_0(k1_comptrig(X1),k2_xcmplx_0(k2_xcmplx_0(k32_sin_cos,k32_sin_cos),k32_sin_cos))|r1_xxreal_0(k2_xcmplx_0(k2_xcmplx_0(k32_sin_cos,k32_sin_cos),k32_sin_cos),k1_comptrig(X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_77472, plain, (r1_xxreal_0(k1_comptrig(X1),k2_xcmplx_0(k2_xcmplx_0(k1_comptrig(X2),k32_sin_cos),k32_sin_cos))|r1_xxreal_0(k2_xcmplx_0(k2_xcmplx_0(k1_comptrig(X2),k32_sin_cos),k32_sin_cos),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_948, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,k2_xcmplx_0(X1,k2_xcmplx_0(X3,X4))),X5)=k2_xcmplx_0(k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3),k2_xcmplx_0(X4,X5))|~v1_xcmplx_0(X5)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_949, plain, (k2_xcmplx_0(k2_xcmplx_0(k2_xcmplx_0(X2,X3),X1),X4)=k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(X2,X3)))).#trainneg
cnf(i_0_77421, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k6_numbers)|k4_complex1(esk1_0)=k6_numbers|v3_xxreal_0(k4_complex1(esk1_0))|~v1_xreal_0(k4_complex1(esk1_0)))).#trainneg
cnf(i_0_62558, plain, (v1_xreal_0(X1)|r1_xxreal_0(X1,k32_sin_cos)|r1_xxreal_0(k6_numbers,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_14227, plain, (r1_xxreal_0(k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos),k1_comptrig(k2_xcmplx_0(X2,X3)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X2,X3)),k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_24609, plain, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,X2)),k1_comptrig(k1_comptrig(k1_comptrig(X3))))|r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(X3))),k1_comptrig(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_951, plain, (k2_xcmplx_0(k2_xcmplx_0(X3,k2_xcmplx_0(X1,X2)),X4)=k2_xcmplx_0(k2_xcmplx_0(X1,X2),k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_78711, plain, (k4_complex1(k2_xcmplx_0(X1,k6_numbers))=k4_complex1(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_78534, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_61744, plain, (m1_subset_1(X1,k1_numbers)|r1_xxreal_0(k32_sin_cos,X1)|r1_xxreal_0(X1,k6_numbers)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_60991, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk2_0),k32_sin_cos)|~v1_xreal_0(k1_comptrig(esk2_0))|~m1_subset_1(esk1_0,k1_numbers))).#trainneg
cnf(i_0_80633, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(X1))),k1_comptrig(k2_xcmplx_0(X2,esk2_0)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X2,esk2_0)),k1_comptrig(k1_comptrig(k1_comptrig(X1))))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_80634, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(X1))),k1_comptrig(k2_xcmplx_0(X2,esk1_0)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X2,esk1_0)),k1_comptrig(k1_comptrig(k1_comptrig(X1))))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_957, plain, (v3_xxreal_0(k1_comptrig(X1))|v2_xxreal_0(k7_real_1(X2,X3))|r1_xxreal_0(k7_real_1(X2,X3),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_80677, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(X1))),k1_comptrig(k2_xcmplx_0(X2,k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X2,k32_sin_cos)),k1_comptrig(k1_comptrig(k1_comptrig(X1))))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_55693, plain, (r1_xxreal_0(k2_xcmplx_0(k32_sin_cos,k1_comptrig(X1)),k1_comptrig(k2_xcmplx_0(X2,X3)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X2,X3)),k2_xcmplx_0(k32_sin_cos,k1_comptrig(X1)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_958, plain, (v3_xxreal_0(k1_comptrig(X1))|v2_xxreal_0(k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4)))|r1_xxreal_0(k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4)),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_959, plain, (v3_xxreal_0(k1_comptrig(X1))|v2_xxreal_0(k2_xcmplx_0(X2,X3))|r1_xxreal_0(k2_xcmplx_0(X2,X3),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_81783, plain, (v1_xcmplx_0(X1)|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k32_sin_cos,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1327, plain, (v1_xcmplx_0(k4_xcmplx_0(X1))|~v5_membered(k1_numbers)|~v1_xcmplx_0(X1)|k1_comptrig(X1)!=k6_numbers)).#trainneg
cnf(i_0_81781, plain, (v1_int_1(X1)|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k32_sin_cos,X1)|~v5_membered(k1_numbers)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_55738, plain, (r1_xxreal_0(k2_xcmplx_0(k1_comptrig(X1),k1_comptrig(X2)),k1_comptrig(k1_comptrig(X3)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X3)),k2_xcmplx_0(k1_comptrig(X1),k1_comptrig(X2)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_32570, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(X1))),k1_comptrig(k1_comptrig(k1_comptrig(k1_comptrig(X2)))))|r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(k1_comptrig(X2)))),k1_comptrig(k1_comptrig(k1_comptrig(X1))))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_971, plain, (v1_xxreal_0(k7_real_1(k4_complex1(X1),X3))|~v1_xcmplx_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1))))|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_81902, plain, (v2_xxreal_0(X1)|r1_xxreal_0(X1,k1_comptrig(X2))|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_81773, negated_conjecture, (r1_xxreal_0(esk2_0,k6_numbers)|r1_xxreal_0(k32_sin_cos,esk2_0)|~v1_xxreal_0(esk2_0))).#trainneg
cnf(i_0_81776, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(esk1_0,esk2_0),k6_numbers)|r1_xxreal_0(k32_sin_cos,k2_xcmplx_0(esk1_0,esk2_0))|~r1_xxreal_0(k6_numbers,k2_xcmplx_0(esk1_0,esk2_0))|~v1_xxreal_0(k2_xcmplx_0(esk1_0,esk2_0)))).#trainneg
cnf(i_0_1383, plain, (k3_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k6_numbers)))=k7_real_1(k3_complex1(X1),X2)|~v1_xreal_0(X2)|~v1_xcmplx_0(k2_xcmplx_0(X2,k6_numbers))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_82507, plain, (v1_int_1(X1)|r1_xxreal_0(k32_sin_cos,X1)|r1_xxreal_0(X1,k6_numbers)|~v1_xxreal_0(X1)|~v6_membered(k1_numbers))).#trainneg
cnf(i_0_82509, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k2_xcmplx_0(k1_comptrig(X2),k1_comptrig(esk2_0)))|r1_xxreal_0(k2_xcmplx_0(k1_comptrig(X2),k1_comptrig(esk2_0)),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_82510, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k2_xcmplx_0(k1_comptrig(X2),k1_comptrig(esk1_0)))|r1_xxreal_0(k2_xcmplx_0(k1_comptrig(X2),k1_comptrig(esk1_0)),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_977, plain, (k4_complex1(k2_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,k7_complex1),X2),X3))=k7_real_1(X1,k4_complex1(X3))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,k7_complex1),X2))|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,k7_complex1)))).#trainneg
cnf(i_0_82811, plain, (k7_real_1(k3_complex1(X1),X2)=k2_xcmplx_0(X1,X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_81780, plain, (v1_xreal_0(X1)|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k32_sin_cos,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_82553, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k2_xcmplx_0(k1_comptrig(X2),k1_comptrig(k32_sin_cos)))|r1_xxreal_0(k2_xcmplx_0(k1_comptrig(X2),k1_comptrig(k32_sin_cos)),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_35073, plain, (r1_xxreal_0(k2_xcmplx_0(k1_comptrig(k1_comptrig(X1)),k32_sin_cos),k1_comptrig(k1_comptrig(k1_comptrig(X2))))|r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(X2))),k2_xcmplx_0(k1_comptrig(k1_comptrig(X1)),k32_sin_cos))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_984, plain, (k4_complex1(k2_xcmplx_0(X3,k2_xcmplx_0(k3_xcmplx_0(X1,k7_complex1),X2)))=k7_real_1(X1,k4_complex1(X3))|~v1_xreal_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,k7_complex1),X2))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,k7_complex1)))).#trainneg
cnf(i_0_82823, plain, (k7_real_1(k3_complex1(X1),k3_complex1(X2))=k3_complex1(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k3_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_82943, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(k1_comptrig(esk2_0),k1_comptrig(esk2_0)),k1_comptrig(k1_comptrig(X1)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k2_xcmplx_0(k1_comptrig(esk2_0),k1_comptrig(esk2_0)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_82944, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(k1_comptrig(esk1_0),k1_comptrig(esk2_0)),k1_comptrig(k1_comptrig(X1)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k2_xcmplx_0(k1_comptrig(esk1_0),k1_comptrig(esk2_0)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_82826, plain, (k7_real_1(k3_complex1(X1),k6_numbers)=k3_complex1(k2_xcmplx_0(X1,k6_numbers))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_82651, plain, (v2_xxreal_0(k8_real_1(np__2,k32_sin_cos))|~v1_xcmplx_0(X1)|~v1_xreal_0(k8_real_1(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_82940, negated_conjecture, (r1_xxreal_0(esk1_0,k6_numbers)|r1_xxreal_0(k32_sin_cos,esk1_0)|~v1_xxreal_0(esk1_0))).#trainneg
cnf(i_0_82987, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(k1_comptrig(k32_sin_cos),k1_comptrig(esk2_0)),k1_comptrig(k1_comptrig(X1)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k2_xcmplx_0(k1_comptrig(k32_sin_cos),k1_comptrig(esk2_0)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_82989, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(k1_comptrig(esk2_0),k1_comptrig(esk1_0)),k1_comptrig(k1_comptrig(X1)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k2_xcmplx_0(k1_comptrig(esk2_0),k1_comptrig(esk1_0)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_989, plain, (k4_complex1(k2_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X2,k7_complex1),X3),X1))=k7_real_1(k4_complex1(X1),X2)|~v1_xreal_0(k4_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X2,k7_complex1),X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X2,k7_complex1)))).#trainneg
cnf(i_0_53309, plain, (k4_complex1(X3)=k4_complex1(X2)|~v1_xreal_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3))).#trainneg
cnf(i_0_84087, plain, (v2_xxreal_0(k8_real_1(np__2,k32_sin_cos))|~v1_xreal_0(k8_real_1(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_82990, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(k1_comptrig(esk1_0),k1_comptrig(esk1_0)),k1_comptrig(k1_comptrig(X1)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k2_xcmplx_0(k1_comptrig(esk1_0),k1_comptrig(esk1_0)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_83033, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(k1_comptrig(k32_sin_cos),k1_comptrig(esk1_0)),k1_comptrig(k1_comptrig(X1)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k2_xcmplx_0(k1_comptrig(k32_sin_cos),k1_comptrig(esk1_0)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_991, plain, (X1=X2|~v1_xreal_0(X2)|~v1_xreal_0(k3_xcmplx_0(X1,k7_complex1))|~v1_xreal_0(X1)|~v1_xreal_0(k3_xcmplx_0(X2,k7_complex1))|~v1_xcmplx_0(k3_xcmplx_0(X1,k7_complex1)))).#trainneg
cnf(i_0_70161, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_60983, plain, (m1_subset_1(X1,k1_numbers)|r1_xxreal_0(k6_numbers,X1)|r1_xxreal_0(X1,k1_comptrig(X2))|~v1_xxreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_51165, plain, (r1_xxreal_0(k4_complex1(X1),k32_sin_cos)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1446, plain, (k4_complex1(k2_xcmplx_0(k3_xcmplx_0(X1,k7_complex1),X2))=k7_real_1(X1,k4_complex1(X2))|~v1_xreal_0(X1)|~v1_xcmplx_0(k3_xcmplx_0(X1,k7_complex1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_84642, plain, (~v3_xxreal_0(k8_real_1(np__2,k32_sin_cos))|~v1_xxreal_0(k8_real_1(np__2,k32_sin_cos))|~v1_xreal_0(k8_real_1(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_83411, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(k1_comptrig(esk2_0),k1_comptrig(k32_sin_cos)),k1_comptrig(k1_comptrig(X1)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k2_xcmplx_0(k1_comptrig(esk2_0),k1_comptrig(k32_sin_cos)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_83412, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(k1_comptrig(esk1_0),k1_comptrig(k32_sin_cos)),k1_comptrig(k1_comptrig(X1)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k2_xcmplx_0(k1_comptrig(esk1_0),k1_comptrig(k32_sin_cos)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_992, plain, (k7_real_1(k4_complex1(k3_xcmplx_0(X1,k7_complex1)),X3)=X1|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)))|~v1_xcmplx_0(k3_xcmplx_0(X1,k7_complex1))|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1))))).#trainneg
cnf(i_0_995, plain, (k4_complex1(k2_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)),X3))=k3_xcmplx_0(X1,X2)|~v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)))|~v1_xcmplx_0(k7_complex1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1449, plain, (k4_complex1(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)))=k7_real_1(k4_complex1(X1),X2)|~v1_xreal_0(X2)|~v1_xcmplx_0(k3_xcmplx_0(X2,k7_complex1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_83455, plain, (r1_xxreal_0(k2_xcmplx_0(k1_comptrig(k32_sin_cos),k1_comptrig(k32_sin_cos)),k1_comptrig(k1_comptrig(X1)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k2_xcmplx_0(k1_comptrig(k32_sin_cos),k1_comptrig(k32_sin_cos)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_35242, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(k2_xcmplx_0(k1_comptrig(k1_comptrig(X2)),k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k1_comptrig(k1_comptrig(X2)),k32_sin_cos)),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_997, plain, (k4_complex1(k2_xcmplx_0(X3,k2_xcmplx_0(k3_xcmplx_0(X1,k7_complex1),X2)))=X1|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(k3_xcmplx_0(X1,k7_complex1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_85239, plain, (r1_xxreal_0(k6_numbers,k8_real_1(np__2,k32_sin_cos))|~v1_xreal_0(k8_real_1(np__2,k32_sin_cos)))).#trainneg
cnf(i_0_57758, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X3,X1))|~m1_subset_1(X2,k1_numbers)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_70026, plain, (k4_complex1(k4_complex1(k2_xcmplx_0(X2,X1)))=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1451, plain, (v1_xcmplx_0(k7_real_1(k4_complex1(X1),X2))|~v1_xreal_0(X2)|~v1_xcmplx_0(k3_xcmplx_0(X2,k7_complex1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_81782, plain, (v7_ordinal1(X1)|r1_xxreal_0(X1,k6_numbers)|r1_xxreal_0(k32_sin_cos,X1)|~v6_membered(k1_numbers)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_60051, plain, (r1_xxreal_0(k2_xcmplx_0(k32_sin_cos,k1_comptrig(k1_comptrig(X1))),k1_comptrig(k1_comptrig(k1_comptrig(X2))))|r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(X2))),k2_xcmplx_0(k32_sin_cos,k1_comptrig(k1_comptrig(X1))))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_78863, plain, (r1_xxreal_0(k2_xcmplx_0(k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos),k32_sin_cos),k1_comptrig(k1_comptrig(X2)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X2)),k2_xcmplx_0(k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos),k32_sin_cos))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1007, plain, (r1_xxreal_0(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4)))|r1_xxreal_0(k2_xcmplx_0(X2,k2_xcmplx_0(X3,X4)),X1)|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X3))|~v1_xreal_0(X4)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_81839, negated_conjecture, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),esk2_0)),k1_comptrig(k1_comptrig(k1_comptrig(X2))))|r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(X2))),k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),esk2_0)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_81885, negated_conjecture, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),esk1_0)),k1_comptrig(k1_comptrig(k1_comptrig(X2))))|r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(X2))),k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),esk1_0)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_72979, plain, (k4_complex1(k2_xcmplx_0(X2,X1))=k4_complex1(X1)|~v1_xreal_0(k2_xcmplx_0(k6_numbers,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_86627, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(X1))),k2_xcmplx_0(k32_sin_cos,k1_comptrig(k1_comptrig(k2_xcmplx_0(X2,k6_numbers)))))|r1_xxreal_0(k2_xcmplx_0(k32_sin_cos,k1_comptrig(k1_comptrig(k2_xcmplx_0(X2,k6_numbers)))),k1_comptrig(k1_comptrig(k1_comptrig(X1))))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_86640, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k2_xcmplx_0(k2_xcmplx_0(k1_comptrig(esk2_0),k32_sin_cos),k32_sin_cos))|r1_xxreal_0(k2_xcmplx_0(k2_xcmplx_0(k1_comptrig(esk2_0),k32_sin_cos),k32_sin_cos),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_86641, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k2_xcmplx_0(k2_xcmplx_0(k1_comptrig(esk1_0),k32_sin_cos),k32_sin_cos))|r1_xxreal_0(k2_xcmplx_0(k2_xcmplx_0(k1_comptrig(esk1_0),k32_sin_cos),k32_sin_cos),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_51372, plain, (m1_subset_1(k4_complex1(k2_xcmplx_0(X1,k6_numbers)),k1_numbers)|~v1_xreal_0(X1)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_22990, negated_conjecture, (k1_comptrig(esk2_0)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(esk1_0,k6_numbers))|~v1_xreal_0(k1_comptrig(esk2_0)))).#trainneg
cnf(i_0_86685, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k2_xcmplx_0(k2_xcmplx_0(k1_comptrig(k32_sin_cos),k32_sin_cos),k32_sin_cos))|r1_xxreal_0(k2_xcmplx_0(k2_xcmplx_0(k1_comptrig(k32_sin_cos),k32_sin_cos),k32_sin_cos),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_86839, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(esk2_0))),k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),esk2_0)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),esk2_0)),k1_comptrig(k1_comptrig(k1_comptrig(esk2_0))))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1026, plain, (X3=X1|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)))|~v1_xcmplx_0(k3_xcmplx_0(X1,k7_complex1))|~v1_xreal_0(k2_xcmplx_0(k3_xcmplx_0(X1,k7_complex1),X2))|~v1_xreal_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X3,k7_complex1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_59552, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X3,X2))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~m1_subset_1(X3,k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_70154, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X3,X2))|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_87411, negated_conjecture, (k1_comptrig(esk2_0)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(esk1_0,k6_numbers)))).#trainneg
cnf(i_0_52392, plain, (~v2_xxreal_0(k7_real_1(k4_complex1(X1),k6_numbers))|~v1_xreal_0(X2)|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(X2,X1)))|~v1_xcmplx_0(k2_xcmplx_0(X2,X1))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_86840, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(esk1_0))),k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),esk2_0)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),esk2_0)),k1_comptrig(k1_comptrig(k1_comptrig(esk1_0))))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_86884, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(k32_sin_cos))),k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),esk2_0)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),esk2_0)),k1_comptrig(k1_comptrig(k1_comptrig(k32_sin_cos))))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1028, plain, (k7_real_1(k4_complex1(X1),X3)=X3|~v1_xreal_0(X3)|~v1_xreal_0(X2)|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)))|~v1_xcmplx_0(X1)|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(k3_xcmplx_0(X3,k7_complex1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1032, plain, (k4_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X2,k2_xcmplx_0(X3,k3_xcmplx_0(X4,k7_complex1)))))=X4|~v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xreal_0(X4)|~v1_xcmplx_0(k3_xcmplx_0(X4,k7_complex1))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_87867, negated_conjecture, (~v1_xreal_0(k2_xcmplx_0(esk1_0,k6_numbers)))).#trainneg
cnf(i_0_86886, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(esk2_0))),k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),esk1_0)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),esk1_0)),k1_comptrig(k1_comptrig(k1_comptrig(esk2_0))))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_86887, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(esk1_0))),k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),esk1_0)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),esk1_0)),k1_comptrig(k1_comptrig(k1_comptrig(esk1_0))))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1033, plain, (k4_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(X3,k2_xcmplx_0(X2,k3_xcmplx_0(X4,k7_complex1)))))=X4|~v1_xreal_0(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))|~v1_xreal_0(X4)|~v1_xcmplx_0(k3_xcmplx_0(X4,k7_complex1))|~v1_xcmplx_0(k2_xcmplx_0(X2,X3))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_1034, plain, (k4_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(k3_xcmplx_0(X3,k7_complex1),X2)))=X3|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X3,k7_complex1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1483, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X1,k6_numbers))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_88541, negated_conjecture, (~v1_xreal_0(k2_xcmplx_0(k6_numbers,esk1_0)))).#trainneg
cnf(i_0_86931, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(k32_sin_cos))),k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),esk1_0)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),esk1_0)),k1_comptrig(k1_comptrig(k1_comptrig(k32_sin_cos))))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_81980, plain, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos)),k1_comptrig(k1_comptrig(k1_comptrig(X2))))|r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(X2))),k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1039, plain, (k4_complex1(k2_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)),X1))=X3|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X3,k7_complex1))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1))))).#trainneg
cnf(i_0_1485, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,k6_numbers),k2_xcmplx_0(X2,X3))=k2_xcmplx_0(k2_xcmplx_0(X1,X2),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_89359, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(esk2_0))),k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos)),k1_comptrig(k1_comptrig(k1_comptrig(esk2_0))))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1047, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k1_comptrig(k3_xcmplx_0(X1,X2))),k6_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_89360, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(esk1_0))),k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos)),k1_comptrig(k1_comptrig(k1_comptrig(esk1_0))))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_89404, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(k1_comptrig(k32_sin_cos))),k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos)))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos)),k1_comptrig(k1_comptrig(k1_comptrig(k32_sin_cos))))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1053, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k1_comptrig(k7_real_1(X1,X2))),k6_numbers)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1054, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k1_comptrig(k2_xcmplx_0(X1,k2_xcmplx_0(X2,X3)))),k6_numbers)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_51996, plain, (k4_complex1(X2)=k7_real_1(k4_complex1(X1),k6_numbers)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~v1_xreal_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_90956, plain, (~v2_xxreal_0(k7_real_1(k4_complex1(X2),k6_numbers))|~v1_xreal_0(X1)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_60940, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X3,k2_xcmplx_0(X2,X4)))|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_1508, plain, (r1_xxreal_0(k32_sin_cos,k4_complex1(X1))|r1_xxreal_0(k4_complex1(X1),k32_sin_cos)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_82986, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(esk2_0)),k1_comptrig(k1_comptrig(X2)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X2)),k2_xcmplx_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(esk2_0)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_83032, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(esk1_0)),k1_comptrig(k1_comptrig(X2)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X2)),k2_xcmplx_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(esk1_0)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1055, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k1_comptrig(k2_xcmplx_0(X1,X2))),k6_numbers)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1056, plain, (r2_hidden(k6_numbers,k2_rcomp_1(k6_numbers,k32_sin_cos))|r1_xxreal_0(k4_complex1(k1_comptrig(k2_xcmplx_0(X1,X2))),k6_numbers)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_51287, plain, (k4_complex1(k2_xcmplx_0(X1,X2))=k4_complex1(X1)|~v1_xreal_0(k2_xcmplx_0(X2,k6_numbers))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_59567, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~m1_subset_1(k2_xcmplx_0(X3,X4),k1_numbers)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X3,X2))|~v1_xreal_0(X4)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_83454, plain, (r1_xxreal_0(k2_xcmplx_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(k32_sin_cos)),k1_comptrig(k1_comptrig(X2)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X2)),k2_xcmplx_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(k32_sin_cos)))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_13335, plain, (k2_xcmplx_0(k2_xcmplx_0(X2,X1),X3)=k2_xcmplx_0(X2,k2_xcmplx_0(X3,X1))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1064, plain, (v3_xxreal_0(k4_complex1(X1))|v3_xxreal_0(k4_complex1(X2))|~v3_xxreal_0(k4_complex1(k2_xcmplx_0(X2,X1)))|~v1_xreal_0(k4_complex1(X2))|~v1_xreal_0(k4_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_92242, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X3,X2))|~v1_xreal_0(k2_xcmplx_0(X4,X5))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3)|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~m1_subset_1(X5,k1_numbers)|~v1_xcmplx_0(X4))).#trainneg
cnf(i_0_65271, plain, (r1_xxreal_0(k4_complex1(X1),k6_numbers)|~v1_xreal_0(X2)|~v1_xcmplx_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_60957, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X4,X2))|~v1_xreal_0(X3)|~v1_xcmplx_0(X4))).#trainneg
cnf(i_0_50527, plain, (k4_complex1(k7_real_1(X1,k6_numbers))=k6_numbers|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_93405, plain, (r1_xxreal_0(k4_complex1(X1),k6_numbers)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1517, plain, (r1_xxreal_0(k32_sin_cos,k17_complex1(X1))|r1_xxreal_0(k17_complex1(X1),k32_sin_cos)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1065, plain, (k2_xcmplx_0(k4_complex1(X1),k2_xcmplx_0(k4_complex1(X2),X3))=k2_xcmplx_0(k4_complex1(k2_xcmplx_0(X2,X1)),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k4_complex1(X2))|~v1_xcmplx_0(k4_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1068, plain, (k2_xcmplx_0(k4_complex1(X1),k2_xcmplx_0(k4_complex1(X2),X3))=k2_xcmplx_0(X3,k4_complex1(k2_xcmplx_0(X2,X1)))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k4_complex1(X2))|~v1_xcmplx_0(k4_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_93681, plain, (~v2_xxreal_0(k4_complex1(X1))|~v1_xreal_0(k4_complex1(X1))|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_93088, plain, (v3_xxreal_0(k4_complex1(X1))|~v3_xxreal_0(k4_complex1(k2_xcmplx_0(X1,k6_numbers)))|~v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_77430, negated_conjecture, (r1_xxreal_0(k1_comptrig(esk1_0),k6_numbers)|r1_xxreal_0(k1_comptrig(esk2_0),k32_sin_cos)|~v1_xreal_0(k1_comptrig(esk2_0)))).#trainneg
cnf(i_0_87961, plain, (~v2_xxreal_0(k7_real_1(k4_complex1(k2_xcmplx_0(X1,X2)),k6_numbers))|~r1_xxreal_0(k32_sin_cos,k1_comptrig(k2_xcmplx_0(X1,X2)))|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_72785, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(k2_xcmplx_0(k6_numbers,X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_84879, plain, (k4_complex1(k2_xcmplx_0(X1,X2))=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X4,X5))|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~v1_xreal_0(X5)|~v1_xcmplx_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X4)|~v1_xreal_0(k2_xcmplx_0(X1,X3))|~v1_xreal_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1518, plain, (r1_xxreal_0(k32_sin_cos,k16_complex1(X1))|r1_xxreal_0(k16_complex1(X1),k32_sin_cos)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_14204, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos),k1_comptrig(esk2_0))|r1_xxreal_0(k1_comptrig(esk2_0),k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_14205, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos),k1_comptrig(esk1_0))|r1_xxreal_0(k1_comptrig(esk1_0),k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1069, plain, (k2_xcmplx_0(k4_complex1(X2),k2_xcmplx_0(k4_complex1(X1),X3))=k2_xcmplx_0(k4_complex1(k2_xcmplx_0(X2,X1)),X3)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(k4_complex1(X1))|~v1_xcmplx_0(k4_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1519, plain, (r1_xxreal_0(k32_sin_cos,k7_real_1(X1,X2))|r1_xxreal_0(k7_real_1(X1,X2),k32_sin_cos)|~v1_xreal_0(X2)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_14231, plain, (r1_xxreal_0(k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos),k1_comptrig(k32_sin_cos))|r1_xxreal_0(k1_comptrig(k32_sin_cos),k2_xcmplx_0(k1_comptrig(X1),k32_sin_cos))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_55657, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(k32_sin_cos,k1_comptrig(X1)),k1_comptrig(esk2_0))|r1_xxreal_0(k1_comptrig(esk2_0),k2_xcmplx_0(k32_sin_cos,k1_comptrig(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1071, plain, (m1_subset_1(k4_complex1(k4_complex1(k2_xcmplx_0(X2,X1))),k1_numbers)|~v1_xcmplx_0(k4_complex1(X2))|~v1_xcmplx_0(k4_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1072, plain, (m1_subset_1(k3_complex1(k4_complex1(k2_xcmplx_0(X2,X1))),k1_numbers)|~v1_xreal_0(k3_complex1(k4_complex1(X2)))|~m1_subset_1(k3_complex1(k4_complex1(X1)),k1_numbers)|~v1_xcmplx_0(k4_complex1(X2))|~v1_xcmplx_0(k4_complex1(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_51716, plain, (k4_complex1(k7_real_1(X1,X2))=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_96477, plain, (v1_xreal_0(k4_complex1(k4_complex1(k2_xcmplx_0(X1,X2))))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_96518, plain, (m1_subset_1(k4_complex1(k4_complex1(X1)),k1_numbers)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1536, plain, (k3_complex1(k2_xcmplx_0(k2_xcmplx_0(X1,k6_numbers),X2))=k7_real_1(X1,k3_complex1(X2))|~v1_xreal_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(X1,k6_numbers))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_55658, negated_conjecture, (r1_xxreal_0(k2_xcmplx_0(k32_sin_cos,k1_comptrig(X1)),k1_comptrig(esk1_0))|r1_xxreal_0(k1_comptrig(esk1_0),k2_xcmplx_0(k32_sin_cos,k1_comptrig(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_55697, plain, (r1_xxreal_0(k2_xcmplx_0(k32_sin_cos,k1_comptrig(X1)),k1_comptrig(k32_sin_cos))|r1_xxreal_0(k1_comptrig(k32_sin_cos),k2_xcmplx_0(k32_sin_cos,k1_comptrig(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1074, plain, (v3_xxreal_0(k4_complex1(X1))|v2_xxreal_0(k4_complex1(k2_xcmplx_0(X1,X2)))|~v2_xxreal_0(k4_complex1(X2))|~v1_xreal_0(k4_complex1(X1))|~v1_xreal_0(k4_complex1(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1075, plain, (k2_xcmplx_0(k4_complex1(X3),k2_xcmplx_0(X1,k4_complex1(X2)))=k2_xcmplx_0(X1,k4_complex1(k2_xcmplx_0(X3,X2)))|~v1_xcmplx_0(k4_complex1(X3))|~v1_xcmplx_0(k4_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_97243, plain, (v1_xreal_0(k4_complex1(k4_complex1(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_96988, plain, (k4_complex1(k3_complex1(k2_xcmplx_0(X2,X1)))=k6_numbers|~v1_xreal_0(k3_complex1(X2))|~v1_xreal_0(X1)|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_97242, plain, (v1_xcmplx_0(k4_complex1(k4_complex1(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1076, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,k4_complex1(k2_xcmplx_0(X3,X2))))|~v1_xcmplx_0(k4_complex1(X3))|~v1_xcmplx_0(k4_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_1077, plain, (v1_xreal_0(k2_xcmplx_0(X1,k4_complex1(k2_xcmplx_0(X3,X2))))|~v1_xreal_0(k2_xcmplx_0(X1,k4_complex1(X2)))|~v1_xreal_0(k4_complex1(X3))|~v1_xcmplx_0(k4_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_98620, plain, (k4_complex1(k3_complex1(X1))=k6_numbers|~v1_xreal_0(X1))).#trainneg
cnf(i_0_1616, plain, (k4_complex1(k2_xcmplx_0(X1,k2_xcmplx_0(k6_numbers,X2)))=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1078, plain, (v3_xxreal_0(k2_xcmplx_0(X1,k4_complex1(X2)))|v2_xxreal_0(k2_xcmplx_0(X1,k4_complex1(k2_xcmplx_0(X3,X2))))|~v2_xxreal_0(k4_complex1(X3))|~v1_xreal_0(k2_xcmplx_0(X1,k4_complex1(X2)))|~v1_xreal_0(k4_complex1(X3))|~v1_xcmplx_0(k4_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_1079, plain, (k2_xcmplx_0(k2_xcmplx_0(X1,k4_complex1(X2)),k2_xcmplx_0(k4_complex1(X3),X4))=k2_xcmplx_0(k2_xcmplx_0(X1,k4_complex1(k2_xcmplx_0(X3,X2))),X4)|~v1_xcmplx_0(X4)|~v1_xcmplx_0(k4_complex1(X3))|~v1_xcmplx_0(k4_complex1(X2))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X3))).#trainneg
cnf(i_0_21246, plain, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(X1,k1_comptrig(X2))),k1_comptrig(X3))|r1_xxreal_0(k1_comptrig(X3),k1_comptrig(k2_xcmplx_0(X1,k1_comptrig(X2))))|~v1_xreal_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1081, plain, (k4_complex1(k2_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X2,k7_complex1),X3),X1))=k2_xcmplx_0(k4_complex1(X1),X2)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X2,k7_complex1),X3))|~v1_xreal_0(X2)|~v1_xreal_0(X3)|~v1_xcmplx_0(k3_xcmplx_0(X2,k7_complex1)))).#trainneg
cnf(i_0_1082, plain, (k4_complex1(k2_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)),X1))=k2_xcmplx_0(k4_complex1(X1),X3)|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(X2,k3_xcmplx_0(X3,k7_complex1)))|~v1_xreal_0(X3)|~v1_xreal_0(X2))).#trainneg
cnf(i_0_1621, plain, (k4_complex1(k2_xcmplx_0(X1,k6_numbers))=k6_numbers|~v1_xreal_0(k2_xcmplx_0(k6_numbers,X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(k2_xcmplx_0(X1,k6_numbers)))).#trainneg
cnf(i_0_101032, plain, (r1_xxreal_0(k1_comptrig(X1),k1_comptrig(k2_xcmplx_0(k32_sin_cos,k1_comptrig(X2))))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k32_sin_cos,k1_comptrig(X2))),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_101031, plain, (r1_xxreal_0(k1_comptrig(X1),k1_comptrig(k2_xcmplx_0(k1_comptrig(X2),k1_comptrig(X3))))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k1_comptrig(X2),k1_comptrig(X3))),k1_comptrig(X1))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X3)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1085, plain, (k4_complex1(k2_xcmplx_0(X3,k2_xcmplx_0(k3_xcmplx_0(X1,k7_complex1),X2)))=k2_xcmplx_0(X1,k4_complex1(X3))|~v1_xcmplx_0(k2_xcmplx_0(k3_xcmplx_0(X1,k7_complex1),X2))|~v1_xcmplx_0(X3)|~v1_xreal_0(X1)|~v1_xreal_0(X2)|~v1_xcmplx_0(k3_xcmplx_0(X1,k7_complex1)))).#trainneg
cnf(i_0_1086, plain, (k4_complex1(k2_xcmplx_0(X3,k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1))))=k2_xcmplx_0(X2,k4_complex1(X3))|~v1_xcmplx_0(k2_xcmplx_0(X1,k3_xcmplx_0(X2,k7_complex1)))|~v1_xcmplx_0(X3)|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_101847, plain, (r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k32_sin_cos,k1_comptrig(X1))),k1_comptrig(k1_comptrig(X2)))|r1_xxreal_0(k1_comptrig(k1_comptrig(X2)),k1_comptrig(k2_xcmplx_0(k32_sin_cos,k1_comptrig(X1))))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_70162, plain, (k4_complex1(X1)=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X2,X1))|~v1_xreal_0(k2_xcmplx_0(X3,X4))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2)|~v1_xreal_0(k2_xcmplx_0(X4,X2))|~v1_xreal_0(X3)|~v1_xcmplx_0(X4))).#trainneg
cnf(i_0_97247, plain, (v1_xxreal_0(k4_complex1(k4_complex1(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_1630, plain, (k4_complex1(k2_xcmplx_0(X2,k2_xcmplx_0(k6_numbers,X1)))=k6_numbers|~v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_102479, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(k2_xcmplx_0(k32_sin_cos,k1_comptrig(esk2_0))))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k32_sin_cos,k1_comptrig(esk2_0))),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_102480, negated_conjecture, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(k2_xcmplx_0(k32_sin_cos,k1_comptrig(esk1_0))))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k32_sin_cos,k1_comptrig(esk1_0))),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_95156, plain, (v3_xxreal_0(k4_complex1(k3_complex1(X1)))|~v3_xxreal_0(k4_complex1(X1))|~v1_xreal_0(k3_complex1(X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_99066, plain, (v1_xcmplx_0(k2_xcmplx_0(X1,k4_complex1(X2)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1667, plain, (k3_complex1(k2_xcmplx_0(k2_xcmplx_0(k6_numbers,X1),X2))=k7_real_1(X1,k3_complex1(X2))|~v1_xcmplx_0(X2)|~v1_xcmplx_0(k2_xcmplx_0(k6_numbers,X1))|~v1_xreal_0(X1))).#trainneg
cnf(i_0_5543, plain, (k4_complex1(k32_sin_cos)=k6_numbers|~m1_subset_1(k4_complex1(k32_sin_cos),k1_numbers)|~r1_xxreal_0(k6_numbers,k4_complex1(k32_sin_cos))|~v1_xreal_0(k4_complex1(k32_sin_cos)))).#trainneg
cnf(i_0_102526, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(k2_xcmplx_0(k32_sin_cos,k1_comptrig(k32_sin_cos))))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k32_sin_cos,k1_comptrig(k32_sin_cos))),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_102525, plain, (r1_xxreal_0(k1_comptrig(k1_comptrig(X1)),k1_comptrig(k2_xcmplx_0(k32_sin_cos,k1_comptrig(k1_comptrig(X2)))))|r1_xxreal_0(k1_comptrig(k2_xcmplx_0(k32_sin_cos,k1_comptrig(k1_comptrig(X2)))),k1_comptrig(k1_comptrig(X1)))|~v1_xcmplx_0(X1)|~v1_xcmplx_0(X2))).#trainneg
cnf(i_0_1097, plain, (k4_complex1(k2_xcmplx_0(k3_xcmplx_0(X2,k7_complex1),X1))=X2|~v1_xreal_0(X2)|~v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)|~m1_subset_1(k3_xcmplx_0(X2,k7_complex1),k1_numbers))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 74
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 108
# Removed in clause preprocessing      : 3
# Initial clauses in saturation        : 105
# Processed clauses                    : 4177
# ...of these trivial                  : 33
# ...subsumed                          : 2821
# ...remaining for further processing  : 1323
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 140
# Backward-rewritten                   : 104
# Generated clauses                    : 104097
# ...of the previous two non-trivial   : 100665
# Contextual simplify-reflections      : 356
# Paramodulations                      : 104097
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1079
#    Positive orientable unit clauses  : 51
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 23
#    Non-unit-clauses                  : 1005
# Current number of unprocessed clauses: 96080
# ...number of literals in the above   : 730491
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1567
# Proof object given clauses           : 232
# Proof search given clauses           : 1323
# Clause-clause subsumption calls (NU) : 187772
# Rec. Clause-clause subsumption calls : 55756
# Non-unit clause-clause subsumptions  : 2789
# Unit Clause-clause subsumption calls : 2294
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 165
# BW rewrite match successes           : 21
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 2352126

# -------------------------------------------------
# User time                : 23.393 s
# System time              : 0.631 s
# Total time               : 24.025 s
# Maximum resident set size: 258148 pages
