2020-06-01 09:35:39.715974: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_8_bb1/premsel_enigma_01_2020_T30_loop02_epoch_389
2020-06-01 09:35:39.779170: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-06-01 09:35:40.005559: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-06-01 09:35:40.925392: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-06-01 09:35:44.200051: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 4484070 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_8_bb1/premsel_enigma_01_2020_T30_loop02_epoch_389' loaded (query=256; context=768; weigths=real; len_mult=0.000000).
# Preprocessing time       : 5.909 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d5_finseq_1, axiom, ![X1]:k5_finseq_1(X1)=k1_tarski(k4_tarski(np__1,X1)), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', d5_finseq_1)).
fof(redefinition_k9_finseq_1, axiom, ![X1]:k9_finseq_1(X1)=k5_finseq_1(X1), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', redefinition_k9_finseq_1)).
fof(d10_finseq_1, axiom, ![X1, X2, X3]:k11_finseq_1(X1,X2,X3)=k7_finseq_1(k7_finseq_1(k9_finseq_1(X1),k9_finseq_1(X2)),k9_finseq_1(X3)), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', d10_finseq_1)).
fof(t45_finseq_1, axiom, ![X1, X2, X3, X4]:(((v1_relat_1(X4)&v1_funct_1(X4))&v1_finseq_1(X4))=>(X4=k11_finseq_1(X1,X2,X3)<=>(((k3_finseq_1(X4)=np__3&k1_funct_1(X4,np__1)=X1)&k1_funct_1(X4,np__2)=X2)&k1_funct_1(X4,np__3)=X3))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', t45_finseq_1)).
fof(fc11_finseq_1, axiom, ![X1, X2, X3]:v1_finseq_1(k11_finseq_1(X1,X2,X3)), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', fc11_finseq_1)).
fof(fc9_finseq_1, axiom, ![X1, X2, X3]:(v1_relat_1(k11_finseq_1(X1,X2,X3))&v1_funct_1(k11_finseq_1(X1,X2,X3))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', fc9_finseq_1)).
fof(redefinition_k3_finseq_4, axiom, ![X1, X2, X3, X4]:((((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))&m1_subset_1(X3,X1))&m1_subset_1(X4,X1))=>k3_finseq_4(X1,X2,X3,X4)=k11_finseq_1(X2,X3,X4)), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', redefinition_k3_finseq_4)).
fof(t20_matrix_9, conjecture, ![X1]:((((((((((((~(v2_struct_0(X1))&~(v6_struct_0(X1)))&v13_algstr_0(X1))&v33_algstr_0(X1))&v3_group_1(X1))&v5_group_1(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v4_vectsp_1(X1))&v5_vectsp_1(X1))&l6_algstr_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>![X6]:(m1_subset_1(X6,u1_struct_0(X1))=>![X7]:(m1_subset_1(X7,u1_struct_0(X1))=>![X8]:(m1_subset_1(X8,u1_struct_0(X1))=>![X9]:(m1_subset_1(X9,u1_struct_0(X1))=>![X10]:(m1_subset_1(X10,u1_struct_0(X1))=>![X11]:(m1_matrix_1(X11,u1_struct_0(X1),np__3,np__3)=>(X11=k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k3_finseq_4(u1_struct_0(X1),X2,X3,X4),k3_finseq_4(u1_struct_0(X1),X5,X6,X7),k3_finseq_4(u1_struct_0(X1),X8,X9,X10))=>![X12]:(m1_matrix_2(X12,k12_matrix_2(np__3))=>(X12=k3_finseq_4(k5_numbers,np__1,np__2,np__3)=>k10_matrix_3(np__3,X1,X11,X12)=k3_finseq_4(u1_struct_0(X1),X2,X6,X10))))))))))))))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', t20_matrix_9)).
fof(fc6_ordinal1, axiom, (~(v1_xboole_0(k4_ordinal1))&v3_ordinal1(k4_ordinal1)), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', fc6_ordinal1)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', redefinition_m2_subset_1)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', cc1_subset_1)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', fc1_subset_1)).
fof(dt_l2_struct_0, axiom, ![X1]:(l2_struct_0(X1)=>l1_struct_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', dt_l2_struct_0)).
fof(dt_l2_algstr_0, axiom, ![X1]:(l2_algstr_0(X1)=>(l2_struct_0(X1)&l1_algstr_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', dt_l2_algstr_0)).
fof(cc5_ordinal1, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:(m1_subset_1(X2,X1)=>v3_ordinal1(X2))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', cc5_ordinal1)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', redefinition_k5_numbers)).
fof(d3_finseq_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>![X2]:(m2_subset_1(X2,k1_numbers,k5_numbers)=>(X2=k3_finseq_1(X1)<=>k2_finseq_1(X2)=k9_xtuple_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', d3_finseq_1)).
fof(t37_matrix_1, axiom, ![X1]:(v7_ordinal1(X1)=>![X2]:(v7_ordinal1(X2)=>![X3]:(v7_ordinal1(X3)=>![X4]:(v7_ordinal1(X4)=>![X5]:(~(v1_xboole_0(X5))=>![X6]:(m1_matrix_1(X6,X5,X1,X2)=>((((r1_xxreal_0(np__1,X3)&r1_xxreal_0(X3,X1))&r1_xxreal_0(np__1,X4))&r1_xxreal_0(X4,X2))=>r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X6))))))))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', t37_matrix_1)).
fof(dt_m2_finseq_1, axiom, ![X1, X2]:(m2_finseq_1(X2,X1)=>((v1_funct_1(X2)&v1_finseq_1(X2))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X1))))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', dt_m2_finseq_1)).
fof(dt_l6_algstr_0, axiom, ![X1]:(l6_algstr_0(X1)=>(l2_algstr_0(X1)&l5_algstr_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', dt_l6_algstr_0)).
fof(d7_matrix_3, axiom, ![X1]:(v7_ordinal1(X1)=>![X2]:((((((((((((~(v2_struct_0(X2))&~(v6_struct_0(X2)))&v13_algstr_0(X2))&v33_algstr_0(X2))&v3_group_1(X2))&v5_group_1(X2))&v2_rlvect_1(X2))&v3_rlvect_1(X2))&v4_rlvect_1(X2))&v4_vectsp_1(X2))&v5_vectsp_1(X2))&l6_algstr_0(X2))=>![X3]:(m1_matrix_1(X3,u1_struct_0(X2),X1,X1)=>![X4]:(m1_matrix_2(X4,k12_matrix_2(X1))=>![X5]:(m2_finseq_1(X5,u1_struct_0(X2))=>(X5=k10_matrix_3(X1,X2,X3,X4)<=>(k3_finseq_1(X5)=X1&![X6]:(v7_ordinal1(X6)=>![X7]:(v7_ordinal1(X7)=>((r2_hidden(X6,k4_finseq_1(X5))&X7=k1_funct_1(X4,X6))=>k1_funct_1(X5,X6)=k3_matrix_1(u1_struct_0(X2),X3,X6,X7))))))))))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', d7_matrix_3)).
fof(cc6_card_1, axiom, ![X1]:((v3_ordinal1(X1)&v1_finset_1(X1))=>v7_ordinal1(X1)), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', cc6_card_1)).
fof(spc3_numerals, axiom, (((v2_xxreal_0(np__3)&m2_subset_1(np__3,k1_numbers,k5_numbers))&m1_subset_1(np__3,k5_numbers))&m1_subset_1(np__3,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', spc3_numerals)).
fof(cc4_card_1, axiom, ![X1]:(m1_subset_1(X1,k4_ordinal1)=>v1_finset_1(X1)), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', cc4_card_1)).
fof(fraenkel_a_1_0_finseq_1, axiom, ![X1, X2]:(v7_ordinal1(X2)=>(r2_hidden(X1,a_1_0_finseq_1(X2))<=>?[X3]:(((m2_subset_1(X3,k1_numbers,k5_numbers)&X1=X3)&r1_xxreal_0(np__1,X3))&r1_xxreal_0(X3,X2)))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', fraenkel_a_1_0_finseq_1)).
fof(redefinition_k2_finseq_1, axiom, ![X1]:(v7_ordinal1(X1)=>k2_finseq_1(X1)=k1_finseq_1(X1)), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', redefinition_k2_finseq_1)).
fof(d1_finseq_1, axiom, ![X1]:(v7_ordinal1(X1)=>k1_finseq_1(X1)=a_1_0_finseq_1(X1)), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', d1_finseq_1)).
fof(dt_k5_numbers, axiom, m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', dt_k5_numbers)).
fof(dt_k3_finseq_4, axiom, ![X1, X2, X3, X4]:((((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))&m1_subset_1(X3,X1))&m1_subset_1(X4,X1))=>m2_finseq_1(k3_finseq_4(X1,X2,X3,X4),X1)), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', dt_k3_finseq_4)).
fof(d5_matrix_1, axiom, ![X1, X2]:((v1_matrix_1(X2)&m2_finseq_1(X2,k3_finseq_2(X1)))=>![X3]:(v7_ordinal1(X3)=>![X4]:(v7_ordinal1(X4)=>(r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X2))=>![X5]:(m1_subset_1(X5,X1)=>(X5=k3_matrix_1(X1,X2,X3,X4)<=>?[X6]:((m2_finseq_1(X6,X1)&X6=k1_funct_1(X2,X3))&X5=k1_funct_1(X6,X4)))))))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', d5_matrix_1)).
fof(dt_m1_matrix_1, axiom, ![X1, X2, X3]:(((~(v1_xboole_0(X1))&v7_ordinal1(X2))&v7_ordinal1(X3))=>![X4]:(m1_matrix_1(X4,X1,X2,X3)=>(v1_matrix_1(X4)&m2_finseq_1(X4,k3_finseq_2(X1))))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', dt_m1_matrix_1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', fc2_struct_0)).
fof(redefinition_k4_finseq_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>k4_finseq_1(X1)=k9_xtuple_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', redefinition_k4_finseq_1)).
fof(cc8_ordinal1, axiom, ![X1]:(m1_subset_1(X1,k4_ordinal1)=>v7_ordinal1(X1)), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', cc8_ordinal1)).
fof(dt_k3_finseq_1, axiom, ![X1]:(((v1_relat_1(X1)&v1_funct_1(X1))&v1_finseq_1(X1))=>m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', dt_k3_finseq_1)).
fof(dt_m1_finseq_1, axiom, ![X1, X2]:(m1_finseq_1(X2,X1)=>((v1_relat_1(X2)&v1_funct_1(X2))&v1_finseq_1(X2))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', dt_m1_finseq_1)).
fof(redefinition_m2_finseq_1, axiom, ![X1, X2]:(m2_finseq_1(X2,X1)<=>m1_finseq_1(X2,X1)), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', redefinition_m2_finseq_1)).
fof(dt_k10_matrix_3, axiom, ![X1, X2, X3, X4]:(((((((((((((((v7_ordinal1(X1)&~(v2_struct_0(X2)))&~(v6_struct_0(X2)))&v13_algstr_0(X2))&v33_algstr_0(X2))&v3_group_1(X2))&v5_group_1(X2))&v2_rlvect_1(X2))&v3_rlvect_1(X2))&v4_rlvect_1(X2))&v4_vectsp_1(X2))&v5_vectsp_1(X2))&l6_algstr_0(X2))&m1_matrix_1(X3,u1_struct_0(X2),X1,X1))&m1_subset_1(X4,k12_matrix_2(X1)))=>m2_finseq_1(k10_matrix_3(X1,X2,X3,X4),u1_struct_0(X2))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', dt_k10_matrix_3)).
fof(spc2_numerals, axiom, (((v2_xxreal_0(np__2)&m2_subset_1(np__2,k1_numbers,k5_numbers))&m1_subset_1(np__2,k5_numbers))&m1_subset_1(np__2,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', spc2_numerals)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', spc1_numerals)).
fof(redefinition_m1_matrix_2, axiom, ![X1]:((~(v1_xboole_0(X1))&v3_matrix_2(X1))=>![X2]:(m1_matrix_2(X2,X1)<=>m1_subset_1(X2,X1))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', redefinition_m1_matrix_2)).
fof(rqLessOrEqual__r1_xxreal_0__r3_r3, axiom, r1_xxreal_0(np__3,np__3), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', rqLessOrEqual__r1_xxreal_0__r3_r3)).
fof(rqLessOrEqual__r1_xxreal_0__r1_r3, axiom, r1_xxreal_0(np__1,np__3), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', rqLessOrEqual__r1_xxreal_0__r1_r3)).
fof(fc2_matrix_2, axiom, ![X1]:(v7_ordinal1(X1)=>(~(v1_xboole_0(k12_matrix_2(X1)))&v3_matrix_2(k12_matrix_2(X1)))), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', fc2_matrix_2)).
fof(rqLessOrEqual__r1_xxreal_0__r1_r1, axiom, r1_xxreal_0(np__1,np__1), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', rqLessOrEqual__r1_xxreal_0__r1_r1)).
fof(rqLessOrEqual__r1_xxreal_0__r2_r3, axiom, r1_xxreal_0(np__2,np__3), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', rqLessOrEqual__r1_xxreal_0__r2_r3)).
fof(rqLessOrEqual__r1_xxreal_0__r1_r2, axiom, r1_xxreal_0(np__1,np__2), file('/scratch/mptp/allpred/preds__0.005/t20_matrix_9', rqLessOrEqual__r1_xxreal_0__r1_r2)).
fof(c_0_47, plain, ![X49]:k5_finseq_1(X49)=k1_tarski(k4_tarski(np__1,X49)), inference(variable_rename,[status(thm)],[d5_finseq_1])).
fof(c_0_48, plain, ![X48]:k9_finseq_1(X48)=k5_finseq_1(X48), inference(variable_rename,[status(thm)],[redefinition_k9_finseq_1])).
fof(c_0_49, plain, ![X56, X57, X58]:k11_finseq_1(X56,X57,X58)=k7_finseq_1(k7_finseq_1(k9_finseq_1(X56),k9_finseq_1(X57)),k9_finseq_1(X58)), inference(variable_rename,[status(thm)],[d10_finseq_1])).
cnf(c_0_50, plain, (k5_finseq_1(X1)=k1_tarski(k4_tarski(np__1,X1))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_51, plain, (k9_finseq_1(X1)=k5_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_52, plain, ![X25, X26, X27, X28]:(((((k3_finseq_1(X28)=np__3|X28!=k11_finseq_1(X25,X26,X27)|(~v1_relat_1(X28)|~v1_funct_1(X28)|~v1_finseq_1(X28)))&(k1_funct_1(X28,np__1)=X25|X28!=k11_finseq_1(X25,X26,X27)|(~v1_relat_1(X28)|~v1_funct_1(X28)|~v1_finseq_1(X28))))&(k1_funct_1(X28,np__2)=X26|X28!=k11_finseq_1(X25,X26,X27)|(~v1_relat_1(X28)|~v1_funct_1(X28)|~v1_finseq_1(X28))))&(k1_funct_1(X28,np__3)=X27|X28!=k11_finseq_1(X25,X26,X27)|(~v1_relat_1(X28)|~v1_funct_1(X28)|~v1_finseq_1(X28))))&(k3_finseq_1(X28)!=np__3|k1_funct_1(X28,np__1)!=X25|k1_funct_1(X28,np__2)!=X26|k1_funct_1(X28,np__3)!=X27|X28=k11_finseq_1(X25,X26,X27)|(~v1_relat_1(X28)|~v1_funct_1(X28)|~v1_finseq_1(X28)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t45_finseq_1])])])).
cnf(c_0_53, plain, (k11_finseq_1(X1,X2,X3)=k7_finseq_1(k7_finseq_1(k9_finseq_1(X1),k9_finseq_1(X2)),k9_finseq_1(X3))), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_54, plain, (k9_finseq_1(X1)=k1_tarski(k4_tarski(np__1,X1))), inference(rw,[status(thm)],[c_0_50, c_0_51])).
fof(c_0_55, plain, ![X95, X96, X97]:v1_finseq_1(k11_finseq_1(X95,X96,X97)), inference(variable_rename,[status(thm)],[fc11_finseq_1])).
fof(c_0_56, plain, ![X85, X86, X87]:(v1_relat_1(k11_finseq_1(X85,X86,X87))&v1_funct_1(k11_finseq_1(X85,X86,X87))), inference(variable_rename,[status(thm)],[fc9_finseq_1])).
cnf(c_0_57, plain, (k1_funct_1(X1,np__1)=X2|X1!=k11_finseq_1(X2,X3,X4)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_58, plain, (k11_finseq_1(X1,X2,X3)=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_53, c_0_54]), c_0_54]), c_0_54])).
cnf(c_0_59, plain, (v1_finseq_1(k11_finseq_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_60, plain, (v1_funct_1(k11_finseq_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_56])).
cnf(c_0_61, plain, (v1_relat_1(k11_finseq_1(X1,X2,X3))), inference(split_conjunct,[status(thm)],[c_0_56])).
fof(c_0_62, plain, ![X52, X53, X54, X55]:(v1_xboole_0(X52)|~m1_subset_1(X53,X52)|~m1_subset_1(X54,X52)|~m1_subset_1(X55,X52)|k3_finseq_4(X52,X53,X54,X55)=k11_finseq_1(X53,X54,X55)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_finseq_4])])])).
cnf(c_0_63, plain, (k3_finseq_1(X1)=np__3|X1!=k11_finseq_1(X2,X3,X4)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_64, plain, (k1_funct_1(X1,np__3)=X2|X1!=k11_finseq_1(X3,X4,X2)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_65, plain, (k1_funct_1(X1,np__1)=X2|X1!=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X2)),k1_tarski(k4_tarski(np__1,X3))),k1_tarski(k4_tarski(np__1,X4)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(rw,[status(thm)],[c_0_57, c_0_58])).
cnf(c_0_66, plain, (v1_finseq_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))), inference(rw,[status(thm)],[c_0_59, c_0_58])).
cnf(c_0_67, plain, (v1_funct_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))), inference(rw,[status(thm)],[c_0_60, c_0_58])).
cnf(c_0_68, plain, (v1_relat_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))), inference(rw,[status(thm)],[c_0_61, c_0_58])).
cnf(c_0_69, plain, (v1_xboole_0(X1)|k3_finseq_4(X1,X2,X3,X4)=k11_finseq_1(X2,X3,X4)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_62])).
fof(c_0_70, negated_conjecture, ~(![X1]:((((((((((((~(v2_struct_0(X1))&~(v6_struct_0(X1)))&v13_algstr_0(X1))&v33_algstr_0(X1))&v3_group_1(X1))&v5_group_1(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v4_vectsp_1(X1))&v5_vectsp_1(X1))&l6_algstr_0(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>![X6]:(m1_subset_1(X6,u1_struct_0(X1))=>![X7]:(m1_subset_1(X7,u1_struct_0(X1))=>![X8]:(m1_subset_1(X8,u1_struct_0(X1))=>![X9]:(m1_subset_1(X9,u1_struct_0(X1))=>![X10]:(m1_subset_1(X10,u1_struct_0(X1))=>![X11]:(m1_matrix_1(X11,u1_struct_0(X1),np__3,np__3)=>(X11=k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(X1))),k3_finseq_4(u1_struct_0(X1),X2,X3,X4),k3_finseq_4(u1_struct_0(X1),X5,X6,X7),k3_finseq_4(u1_struct_0(X1),X8,X9,X10))=>![X12]:(m1_matrix_2(X12,k12_matrix_2(np__3))=>(X12=k3_finseq_4(k5_numbers,np__1,np__2,np__3)=>k10_matrix_3(np__3,X1,X11,X12)=k3_finseq_4(u1_struct_0(X1),X2,X6,X10)))))))))))))))), inference(assume_negation,[status(cth)],[t20_matrix_9])).
fof(c_0_71, plain, (~v1_xboole_0(k4_ordinal1)&v3_ordinal1(k4_ordinal1)), inference(fof_simplification,[status(thm)],[fc6_ordinal1])).
fof(c_0_72, plain, ![X161, X162, X163]:((~m2_subset_1(X163,X161,X162)|m1_subset_1(X163,X162)|(v1_xboole_0(X161)|v1_xboole_0(X162)|~m1_subset_1(X162,k1_zfmisc_1(X161))))&(~m1_subset_1(X163,X162)|m2_subset_1(X163,X161,X162)|(v1_xboole_0(X161)|v1_xboole_0(X162)|~m1_subset_1(X162,k1_zfmisc_1(X161))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])).
fof(c_0_73, plain, ![X100, X101]:(~v1_xboole_0(X100)|(~m1_subset_1(X101,k1_zfmisc_1(X100))|v1_xboole_0(X101))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
cnf(c_0_74, plain, (k3_finseq_1(X1)=np__3|X1!=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X2)),k1_tarski(k4_tarski(np__1,X3))),k1_tarski(k4_tarski(np__1,X4)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(rw,[status(thm)],[c_0_63, c_0_58])).
cnf(c_0_75, plain, (k1_funct_1(X1,np__3)=X2|X1!=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X3)),k1_tarski(k4_tarski(np__1,X4))),k1_tarski(k4_tarski(np__1,X2)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(rw,[status(thm)],[c_0_64, c_0_58])).
cnf(c_0_76, plain, (k1_funct_1(X1,np__1)=X2|X1!=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X2)),k1_tarski(k4_tarski(np__1,X3))),k1_tarski(k4_tarski(np__1,X4)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_65])).
cnf(c_0_77, plain, (v1_finseq_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))), inference(evalgc,[status(thm)],[c_0_66])).
cnf(c_0_78, plain, (v1_funct_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))), inference(evalgc,[status(thm)],[c_0_67])).
cnf(c_0_79, plain, (v1_relat_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))), inference(evalgc,[status(thm)],[c_0_68])).
cnf(c_0_80, plain, (k3_finseq_4(X1,X2,X3,X4)=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X2)),k1_tarski(k4_tarski(np__1,X3))),k1_tarski(k4_tarski(np__1,X4)))|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(rw,[status(thm)],[c_0_69, c_0_58])).
fof(c_0_81, negated_conjecture, ((((((((((((~v2_struct_0(esk1_0)&~v6_struct_0(esk1_0))&v13_algstr_0(esk1_0))&v33_algstr_0(esk1_0))&v3_group_1(esk1_0))&v5_group_1(esk1_0))&v2_rlvect_1(esk1_0))&v3_rlvect_1(esk1_0))&v4_rlvect_1(esk1_0))&v4_vectsp_1(esk1_0))&v5_vectsp_1(esk1_0))&l6_algstr_0(esk1_0))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(m1_subset_1(esk3_0,u1_struct_0(esk1_0))&(m1_subset_1(esk4_0,u1_struct_0(esk1_0))&(m1_subset_1(esk5_0,u1_struct_0(esk1_0))&(m1_subset_1(esk6_0,u1_struct_0(esk1_0))&(m1_subset_1(esk7_0,u1_struct_0(esk1_0))&(m1_subset_1(esk8_0,u1_struct_0(esk1_0))&(m1_subset_1(esk9_0,u1_struct_0(esk1_0))&(m1_subset_1(esk10_0,u1_struct_0(esk1_0))&(m1_matrix_1(esk11_0,u1_struct_0(esk1_0),np__3,np__3)&(esk11_0=k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0))&(m1_matrix_2(esk12_0,k12_matrix_2(np__3))&(esk12_0=k3_finseq_4(k5_numbers,np__1,np__2,np__3)&k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk6_0,esk10_0))))))))))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_70])])])])).
fof(c_0_82, plain, ![X207]:~v1_xboole_0(k1_zfmisc_1(X207)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
fof(c_0_83, plain, ![X102]:(~l2_struct_0(X102)|l1_struct_0(X102)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l2_struct_0])])).
fof(c_0_84, plain, ![X50]:((l2_struct_0(X50)|~l2_algstr_0(X50))&(l1_algstr_0(X50)|~l2_algstr_0(X50))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l2_algstr_0])])])).
fof(c_0_85, plain, ![X127, X128]:(~v3_ordinal1(X127)|(~m1_subset_1(X128,X127)|v3_ordinal1(X128))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc5_ordinal1])])])).
cnf(c_0_86, plain, (v3_ordinal1(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_87, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
fof(c_0_88, plain, ![X29, X30]:((X30!=k3_finseq_1(X29)|k2_finseq_1(X30)=k9_xtuple_0(X29)|~m2_subset_1(X30,k1_numbers,k5_numbers)|(~v1_relat_1(X29)|~v1_funct_1(X29)|~v1_finseq_1(X29)))&(k2_finseq_1(X30)!=k9_xtuple_0(X29)|X30=k3_finseq_1(X29)|~m2_subset_1(X30,k1_numbers,k5_numbers)|(~v1_relat_1(X29)|~v1_funct_1(X29)|~v1_finseq_1(X29)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_finseq_1])])])])).
cnf(c_0_89, plain, (m2_subset_1(X1,X3,X2)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_72])).
cnf(c_0_90, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_73])).
cnf(c_0_91, plain, (~v1_xboole_0(k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_92, plain, (k3_finseq_1(X1)=np__3|X1!=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X2)),k1_tarski(k4_tarski(np__1,X3))),k1_tarski(k4_tarski(np__1,X4)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_74])).
fof(c_0_93, plain, ![X63, X64, X65, X66, X67, X68]:(~v7_ordinal1(X63)|(~v7_ordinal1(X64)|(~v7_ordinal1(X65)|(~v7_ordinal1(X66)|(v1_xboole_0(X67)|(~m1_matrix_1(X68,X67,X63,X64)|(~r1_xxreal_0(np__1,X65)|~r1_xxreal_0(X65,X63)|~r1_xxreal_0(np__1,X66)|~r1_xxreal_0(X66,X64)|r2_hidden(k4_tarski(X65,X66),k2_matrix_1(X68))))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t37_matrix_1])])])])).
cnf(c_0_94, plain, (k1_funct_1(X1,np__3)=X2|X1!=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X3)),k1_tarski(k4_tarski(np__1,X4))),k1_tarski(k4_tarski(np__1,X2)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_75])).
cnf(c_0_95, plain, (k1_funct_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),np__1)=X1), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_76]), c_0_77]), c_0_78]), c_0_79])])])).
cnf(c_0_96, plain, (k3_finseq_4(X1,X2,X3,X4)=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X2)),k1_tarski(k4_tarski(np__1,X3))),k1_tarski(k4_tarski(np__1,X4)))|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_80])).
cnf(c_0_97, negated_conjecture, (esk11_0=k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_98, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_82])).
fof(c_0_99, plain, ![X148, X149]:(((v1_funct_1(X149)|~m2_finseq_1(X149,X148))&(v1_finseq_1(X149)|~m2_finseq_1(X149,X148)))&(m1_subset_1(X149,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X148)))|~m2_finseq_1(X149,X148))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m2_finseq_1])])])).
cnf(c_0_100, plain, (l1_struct_0(X1)|~l2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_83])).
cnf(c_0_101, plain, (l2_struct_0(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_84])).
fof(c_0_102, plain, ![X47]:((l2_algstr_0(X47)|~l6_algstr_0(X47))&(l5_algstr_0(X47)|~l6_algstr_0(X47))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l6_algstr_0])])])).
fof(c_0_103, plain, ![X31, X32, X33, X34, X35, X36, X37]:(((k3_finseq_1(X35)=X31|X35!=k10_matrix_3(X31,X32,X33,X34)|~m2_finseq_1(X35,u1_struct_0(X32))|~m1_matrix_2(X34,k12_matrix_2(X31))|~m1_matrix_1(X33,u1_struct_0(X32),X31,X31)|(v2_struct_0(X32)|v6_struct_0(X32)|~v13_algstr_0(X32)|~v33_algstr_0(X32)|~v3_group_1(X32)|~v5_group_1(X32)|~v2_rlvect_1(X32)|~v3_rlvect_1(X32)|~v4_rlvect_1(X32)|~v4_vectsp_1(X32)|~v5_vectsp_1(X32)|~l6_algstr_0(X32))|~v7_ordinal1(X31))&(~v7_ordinal1(X36)|(~v7_ordinal1(X37)|(~r2_hidden(X36,k4_finseq_1(X35))|X37!=k1_funct_1(X34,X36)|k1_funct_1(X35,X36)=k3_matrix_1(u1_struct_0(X32),X33,X36,X37)))|X35!=k10_matrix_3(X31,X32,X33,X34)|~m2_finseq_1(X35,u1_struct_0(X32))|~m1_matrix_2(X34,k12_matrix_2(X31))|~m1_matrix_1(X33,u1_struct_0(X32),X31,X31)|(v2_struct_0(X32)|v6_struct_0(X32)|~v13_algstr_0(X32)|~v33_algstr_0(X32)|~v3_group_1(X32)|~v5_group_1(X32)|~v2_rlvect_1(X32)|~v3_rlvect_1(X32)|~v4_rlvect_1(X32)|~v4_vectsp_1(X32)|~v5_vectsp_1(X32)|~l6_algstr_0(X32))|~v7_ordinal1(X31)))&((v7_ordinal1(esk13_5(X31,X32,X33,X34,X35))|k3_finseq_1(X35)!=X31|X35=k10_matrix_3(X31,X32,X33,X34)|~m2_finseq_1(X35,u1_struct_0(X32))|~m1_matrix_2(X34,k12_matrix_2(X31))|~m1_matrix_1(X33,u1_struct_0(X32),X31,X31)|(v2_struct_0(X32)|v6_struct_0(X32)|~v13_algstr_0(X32)|~v33_algstr_0(X32)|~v3_group_1(X32)|~v5_group_1(X32)|~v2_rlvect_1(X32)|~v3_rlvect_1(X32)|~v4_rlvect_1(X32)|~v4_vectsp_1(X32)|~v5_vectsp_1(X32)|~l6_algstr_0(X32))|~v7_ordinal1(X31))&((v7_ordinal1(esk14_5(X31,X32,X33,X34,X35))|k3_finseq_1(X35)!=X31|X35=k10_matrix_3(X31,X32,X33,X34)|~m2_finseq_1(X35,u1_struct_0(X32))|~m1_matrix_2(X34,k12_matrix_2(X31))|~m1_matrix_1(X33,u1_struct_0(X32),X31,X31)|(v2_struct_0(X32)|v6_struct_0(X32)|~v13_algstr_0(X32)|~v33_algstr_0(X32)|~v3_group_1(X32)|~v5_group_1(X32)|~v2_rlvect_1(X32)|~v3_rlvect_1(X32)|~v4_rlvect_1(X32)|~v4_vectsp_1(X32)|~v5_vectsp_1(X32)|~l6_algstr_0(X32))|~v7_ordinal1(X31))&(((r2_hidden(esk13_5(X31,X32,X33,X34,X35),k4_finseq_1(X35))|k3_finseq_1(X35)!=X31|X35=k10_matrix_3(X31,X32,X33,X34)|~m2_finseq_1(X35,u1_struct_0(X32))|~m1_matrix_2(X34,k12_matrix_2(X31))|~m1_matrix_1(X33,u1_struct_0(X32),X31,X31)|(v2_struct_0(X32)|v6_struct_0(X32)|~v13_algstr_0(X32)|~v33_algstr_0(X32)|~v3_group_1(X32)|~v5_group_1(X32)|~v2_rlvect_1(X32)|~v3_rlvect_1(X32)|~v4_rlvect_1(X32)|~v4_vectsp_1(X32)|~v5_vectsp_1(X32)|~l6_algstr_0(X32))|~v7_ordinal1(X31))&(esk14_5(X31,X32,X33,X34,X35)=k1_funct_1(X34,esk13_5(X31,X32,X33,X34,X35))|k3_finseq_1(X35)!=X31|X35=k10_matrix_3(X31,X32,X33,X34)|~m2_finseq_1(X35,u1_struct_0(X32))|~m1_matrix_2(X34,k12_matrix_2(X31))|~m1_matrix_1(X33,u1_struct_0(X32),X31,X31)|(v2_struct_0(X32)|v6_struct_0(X32)|~v13_algstr_0(X32)|~v33_algstr_0(X32)|~v3_group_1(X32)|~v5_group_1(X32)|~v2_rlvect_1(X32)|~v3_rlvect_1(X32)|~v4_rlvect_1(X32)|~v4_vectsp_1(X32)|~v5_vectsp_1(X32)|~l6_algstr_0(X32))|~v7_ordinal1(X31)))&(k1_funct_1(X35,esk13_5(X31,X32,X33,X34,X35))!=k3_matrix_1(u1_struct_0(X32),X33,esk13_5(X31,X32,X33,X34,X35),esk14_5(X31,X32,X33,X34,X35))|k3_finseq_1(X35)!=X31|X35=k10_matrix_3(X31,X32,X33,X34)|~m2_finseq_1(X35,u1_struct_0(X32))|~m1_matrix_2(X34,k12_matrix_2(X31))|~m1_matrix_1(X33,u1_struct_0(X32),X31,X31)|(v2_struct_0(X32)|v6_struct_0(X32)|~v13_algstr_0(X32)|~v33_algstr_0(X32)|~v3_group_1(X32)|~v5_group_1(X32)|~v2_rlvect_1(X32)|~v3_rlvect_1(X32)|~v4_rlvect_1(X32)|~v4_vectsp_1(X32)|~v5_vectsp_1(X32)|~l6_algstr_0(X32))|~v7_ordinal1(X31)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d7_matrix_3])])])])])])).
fof(c_0_104, plain, ![X115]:(~v3_ordinal1(X115)|~v1_finset_1(X115)|v7_ordinal1(X115)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc6_card_1])])).
cnf(c_0_105, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_85])).
cnf(c_0_106, plain, (m1_subset_1(np__3,k5_numbers)), inference(split_conjunct,[status(thm)],[spc3_numerals])).
cnf(c_0_107, plain, (v3_ordinal1(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_108, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_87])).
fof(c_0_109, plain, ![X107]:(~m1_subset_1(X107,k4_ordinal1)|v1_finset_1(X107)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_card_1])])).
fof(c_0_110, plain, ![X59, X60, X62]:(((((m2_subset_1(esk16_2(X59,X60),k1_numbers,k5_numbers)|~r2_hidden(X59,a_1_0_finseq_1(X60))|~v7_ordinal1(X60))&(X59=esk16_2(X59,X60)|~r2_hidden(X59,a_1_0_finseq_1(X60))|~v7_ordinal1(X60)))&(r1_xxreal_0(np__1,esk16_2(X59,X60))|~r2_hidden(X59,a_1_0_finseq_1(X60))|~v7_ordinal1(X60)))&(r1_xxreal_0(esk16_2(X59,X60),X60)|~r2_hidden(X59,a_1_0_finseq_1(X60))|~v7_ordinal1(X60)))&(~m2_subset_1(X62,k1_numbers,k5_numbers)|X59!=X62|~r1_xxreal_0(np__1,X62)|~r1_xxreal_0(X62,X60)|r2_hidden(X59,a_1_0_finseq_1(X60))|~v7_ordinal1(X60))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fraenkel_a_1_0_finseq_1])])])])])).
fof(c_0_111, plain, ![X104]:(~v7_ordinal1(X104)|k2_finseq_1(X104)=k1_finseq_1(X104)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_finseq_1])])).
fof(c_0_112, plain, ![X88]:(~v7_ordinal1(X88)|k1_finseq_1(X88)=a_1_0_finseq_1(X88)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_finseq_1])])).
cnf(c_0_113, plain, (k2_finseq_1(X1)=k9_xtuple_0(X2)|X1!=k3_finseq_1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_114, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_115, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_116, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(split_conjunct,[status(thm)],[dt_k5_numbers])).
cnf(c_0_117, plain, (~v1_xboole_0(k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_91])).
cnf(c_0_118, plain, (k3_finseq_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))=np__3|~v1_finseq_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_92]), c_0_78]), c_0_79])])])).
fof(c_0_119, plain, ![X91, X92, X93, X94]:(v1_xboole_0(X91)|~m1_subset_1(X92,X91)|~m1_subset_1(X93,X91)|~m1_subset_1(X94,X91)|m2_finseq_1(k3_finseq_4(X91,X92,X93,X94),X91)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_finseq_4])])])).
fof(c_0_120, plain, ![X40, X41, X42, X43, X44, X46]:((((m2_finseq_1(esk15_5(X40,X41,X42,X43,X44),X40)|X44!=k3_matrix_1(X40,X41,X42,X43)|~m1_subset_1(X44,X40)|~r2_hidden(k4_tarski(X42,X43),k2_matrix_1(X41))|~v7_ordinal1(X43)|~v7_ordinal1(X42)|(~v1_matrix_1(X41)|~m2_finseq_1(X41,k3_finseq_2(X40))))&(esk15_5(X40,X41,X42,X43,X44)=k1_funct_1(X41,X42)|X44!=k3_matrix_1(X40,X41,X42,X43)|~m1_subset_1(X44,X40)|~r2_hidden(k4_tarski(X42,X43),k2_matrix_1(X41))|~v7_ordinal1(X43)|~v7_ordinal1(X42)|(~v1_matrix_1(X41)|~m2_finseq_1(X41,k3_finseq_2(X40)))))&(X44=k1_funct_1(esk15_5(X40,X41,X42,X43,X44),X43)|X44!=k3_matrix_1(X40,X41,X42,X43)|~m1_subset_1(X44,X40)|~r2_hidden(k4_tarski(X42,X43),k2_matrix_1(X41))|~v7_ordinal1(X43)|~v7_ordinal1(X42)|(~v1_matrix_1(X41)|~m2_finseq_1(X41,k3_finseq_2(X40)))))&(~m2_finseq_1(X46,X40)|X46!=k1_funct_1(X41,X42)|X44!=k1_funct_1(X46,X43)|X44=k3_matrix_1(X40,X41,X42,X43)|~m1_subset_1(X44,X40)|~r2_hidden(k4_tarski(X42,X43),k2_matrix_1(X41))|~v7_ordinal1(X43)|~v7_ordinal1(X42)|(~v1_matrix_1(X41)|~m2_finseq_1(X41,k3_finseq_2(X40))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d5_matrix_1])])])])])).
cnf(c_0_121, plain, (v1_xboole_0(X5)|r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X6))|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~v7_ordinal1(X4)|~m1_matrix_1(X6,X5,X1,X2)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_122, negated_conjecture, (m1_matrix_1(esk11_0,u1_struct_0(esk1_0),np__3,np__3)), inference(split_conjunct,[status(thm)],[c_0_81])).
fof(c_0_123, plain, ![X165, X166, X167, X168]:((v1_matrix_1(X168)|~m1_matrix_1(X168,X165,X166,X167)|(v1_xboole_0(X165)|~v7_ordinal1(X166)|~v7_ordinal1(X167)))&(m2_finseq_1(X168,k3_finseq_2(X165))|~m1_matrix_1(X168,X165,X166,X167)|(v1_xboole_0(X165)|~v7_ordinal1(X166)|~v7_ordinal1(X167)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_matrix_1])])])])])).
cnf(c_0_124, plain, (k1_funct_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),np__3)=X3), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_94]), c_0_77]), c_0_78]), c_0_79])])])).
cnf(c_0_125, plain, (k1_funct_1(k3_finseq_4(X4,X1,X2,X3),np__1)=X1|v1_xboole_0(X4)|~m1_subset_1(X3,X4)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_96])])).
cnf(c_0_126, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0))=esk11_0), inference(evalgc,[status(thm)],[c_0_97])).
cnf(c_0_127, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_128, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_99])).
fof(c_0_129, plain, ![X71]:(v2_struct_0(X71)|~l1_struct_0(X71)|~v1_xboole_0(u1_struct_0(X71))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
cnf(c_0_130, plain, (l1_struct_0(X1)|~l2_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_100])).
cnf(c_0_131, plain, (l2_struct_0(X1)|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_101])).
cnf(c_0_132, plain, (l2_algstr_0(X1)|~l6_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_102])).
cnf(c_0_133, plain, (k1_funct_1(X3,X1)=k3_matrix_1(u1_struct_0(X5),X6,X1,X2)|v2_struct_0(X5)|v6_struct_0(X5)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~r2_hidden(X1,k4_finseq_1(X3))|X2!=k1_funct_1(X4,X1)|X3!=k10_matrix_3(X7,X5,X6,X4)|~m2_finseq_1(X3,u1_struct_0(X5))|~m1_matrix_2(X4,k12_matrix_2(X7))|~m1_matrix_1(X6,u1_struct_0(X5),X7,X7)|~v13_algstr_0(X5)|~v33_algstr_0(X5)|~v3_group_1(X5)|~v5_group_1(X5)|~v2_rlvect_1(X5)|~v3_rlvect_1(X5)|~v4_rlvect_1(X5)|~v4_vectsp_1(X5)|~v5_vectsp_1(X5)|~l6_algstr_0(X5)|~v7_ordinal1(X7)), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_134, negated_conjecture, (l6_algstr_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_135, negated_conjecture, (v5_vectsp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_136, negated_conjecture, (v4_vectsp_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_137, negated_conjecture, (v4_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_138, negated_conjecture, (v3_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_139, negated_conjecture, (v2_rlvect_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_140, negated_conjecture, (v5_group_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_141, negated_conjecture, (v3_group_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_142, negated_conjecture, (v33_algstr_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_143, negated_conjecture, (v13_algstr_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_144, negated_conjecture, (~v6_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_145, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_146, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1)), inference(split_conjunct,[status(thm)],[c_0_104])).
cnf(c_0_147, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_148, plain, (m1_subset_1(np__3,k5_numbers)), inference(evalgc,[status(thm)],[c_0_106])).
cnf(c_0_149, plain, (v3_ordinal1(k5_numbers)), inference(rw,[status(thm)],[c_0_107, c_0_108])).
cnf(c_0_150, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_109])).
cnf(c_0_151, plain, (r2_hidden(X2,a_1_0_finseq_1(X3))|~m2_subset_1(X1,k1_numbers,k5_numbers)|X2!=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X3)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_152, plain, (k2_finseq_1(X1)=k1_finseq_1(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_111])).
cnf(c_0_153, plain, (k1_finseq_1(X1)=a_1_0_finseq_1(X1)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_112])).
cnf(c_0_154, plain, (k9_xtuple_0(X2)=k2_finseq_1(X1)|X1!=k3_finseq_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_113])).
cnf(c_0_155, plain, (v1_xboole_0(X1)|m2_subset_1(X2,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m1_subset_1(X2,X1)), inference(csr,[status(thm)],[c_0_114, c_0_115])).
cnf(c_0_156, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers))), inference(evalgc,[status(thm)],[c_0_116])).
cnf(c_0_157, plain, (~v1_xboole_0(k5_numbers)), inference(rw,[status(thm)],[c_0_117, c_0_108])).
cnf(c_0_158, plain, (k3_finseq_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))=np__3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_118, c_0_77])])).
cnf(c_0_159, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_160, plain, (v1_xboole_0(X1)|m2_finseq_1(k3_finseq_4(X1,X2,X3,X4),X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_161, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_99])).
cnf(c_0_162, plain, (X5=k3_matrix_1(X2,X3,X4,X6)|~m2_finseq_1(X1,X2)|X1!=k1_funct_1(X3,X4)|X5!=k1_funct_1(X1,X6)|~m1_subset_1(X5,X2)|~r2_hidden(k4_tarski(X4,X6),k2_matrix_1(X3))|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X2))), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_163, plain, (v1_xboole_0(X5)|r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X6))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~m1_matrix_1(X6,X5,X1,X2)), inference(evalgc,[status(thm)],[c_0_121])).
cnf(c_0_164, negated_conjecture, (m1_matrix_1(esk11_0,u1_struct_0(esk1_0),np__3,np__3)), inference(evalgc,[status(thm)],[c_0_122])).
cnf(c_0_165, plain, (v1_matrix_1(X1)|v1_xboole_0(X2)|~m1_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X4)), inference(split_conjunct,[status(thm)],[c_0_123])).
cnf(c_0_166, plain, (X1=k11_finseq_1(X2,X3,X4)|k3_finseq_1(X1)!=np__3|k1_funct_1(X1,np__1)!=X2|k1_funct_1(X1,np__2)!=X3|k1_funct_1(X1,np__3)!=X4|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_167, plain, (k1_funct_1(k3_finseq_4(X4,X1,X2,X3),np__3)=X3|v1_xboole_0(X4)|~m1_subset_1(X3,X4)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_96])])).
cnf(c_0_168, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126]), c_0_127])])).
cnf(c_0_169, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_128])).
cnf(c_0_170, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_171, plain, (l1_struct_0(X1)|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_131])])).
cnf(c_0_172, plain, (l2_algstr_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_132])).
cnf(c_0_173, plain, (k1_funct_1(X3,X1)=k3_matrix_1(u1_struct_0(X5),X6,X1,X2)|v2_struct_0(X5)|v6_struct_0(X5)|X2!=k1_funct_1(X4,X1)|X3!=k10_matrix_3(X7,X5,X6,X4)|~v13_algstr_0(X5)|~v33_algstr_0(X5)|~v3_group_1(X5)|~v5_group_1(X5)|~v2_rlvect_1(X5)|~v3_rlvect_1(X5)|~v4_rlvect_1(X5)|~v4_vectsp_1(X5)|~v5_vectsp_1(X5)|~l6_algstr_0(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_matrix_2(X4,k12_matrix_2(X7))|~m2_finseq_1(X3,u1_struct_0(X5))|~r2_hidden(X1,k4_finseq_1(X3))|~m1_matrix_1(X6,u1_struct_0(X5),X7,X7)), inference(evalgc,[status(thm)],[c_0_133])).
cnf(c_0_174, negated_conjecture, (l6_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_134])).
cnf(c_0_175, negated_conjecture, (v5_vectsp_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_135])).
cnf(c_0_176, negated_conjecture, (v4_vectsp_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_136])).
cnf(c_0_177, negated_conjecture, (v4_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_178, negated_conjecture, (v3_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_138])).
cnf(c_0_179, negated_conjecture, (v2_rlvect_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_139])).
cnf(c_0_180, negated_conjecture, (v5_group_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_140])).
cnf(c_0_181, negated_conjecture, (v3_group_1(esk1_0)), inference(evalgc,[status(thm)],[c_0_141])).
cnf(c_0_182, negated_conjecture, (v33_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_142])).
cnf(c_0_183, negated_conjecture, (v13_algstr_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_143])).
cnf(c_0_184, negated_conjecture, (~v6_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_144])).
cnf(c_0_185, negated_conjecture, (~v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[c_0_145])).
cnf(c_0_186, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1)), inference(evalgc,[status(thm)],[c_0_146])).
cnf(c_0_187, plain, (v3_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_148]), c_0_149])])])).
cnf(c_0_188, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_150])).
fof(c_0_189, plain, ![X103]:(~v1_relat_1(X103)|~v1_funct_1(X103)|~v1_finseq_1(X103)|k4_finseq_1(X103)=k9_xtuple_0(X103)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_finseq_1])])).
cnf(c_0_190, plain, (m2_subset_1(np__3,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[spc3_numerals])).
cnf(c_0_191, plain, (k3_finseq_1(X1)=X2|v2_struct_0(X3)|v6_struct_0(X3)|X1!=k10_matrix_3(X2,X3,X4,X5)|~m2_finseq_1(X1,u1_struct_0(X3))|~m1_matrix_2(X5,k12_matrix_2(X2))|~m1_matrix_1(X4,u1_struct_0(X3),X2,X2)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_103])).
cnf(c_0_192, plain, (r2_hidden(X2,a_1_0_finseq_1(X3))|X1!=X2|~v7_ordinal1(X3)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_151])).
cnf(c_0_193, plain, (k1_finseq_1(X1)=k2_finseq_1(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_152])).
cnf(c_0_194, plain, (k1_finseq_1(X1)=a_1_0_finseq_1(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_153])).
cnf(c_0_195, plain, (k9_xtuple_0(X1)=k2_finseq_1(X2)|X2!=k3_finseq_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_155]), c_0_156])]), c_0_157])])).
cnf(c_0_196, plain, (k3_finseq_1(k3_finseq_4(X4,X1,X2,X3))=np__3|v1_xboole_0(X4)|~m1_subset_1(X3,X4)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_96])])).
cnf(c_0_197, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_159])).
cnf(c_0_198, plain, (v1_xboole_0(X1)|m2_finseq_1(k3_finseq_4(X1,X2,X3,X4),X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_160])).
cnf(c_0_199, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_161])).
fof(c_0_200, plain, ![X79]:(~m1_subset_1(X79,k4_ordinal1)|v7_ordinal1(X79)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_ordinal1])])).
fof(c_0_201, plain, ![X169]:(~v1_relat_1(X169)|~v1_funct_1(X169)|~v1_finseq_1(X169)|m2_subset_1(k3_finseq_1(X169),k1_numbers,k5_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_finseq_1])])).
fof(c_0_202, plain, ![X130, X131]:(((v1_relat_1(X131)|~m1_finseq_1(X131,X130))&(v1_funct_1(X131)|~m1_finseq_1(X131,X130)))&(v1_finseq_1(X131)|~m1_finseq_1(X131,X130))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_finseq_1])])])).
fof(c_0_203, plain, ![X98, X99]:((~m2_finseq_1(X99,X98)|m1_finseq_1(X99,X98))&(~m1_finseq_1(X99,X98)|m2_finseq_1(X99,X98))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_m2_finseq_1])])).
cnf(c_0_204, plain, (X5=k3_matrix_1(X2,X3,X4,X6)|X5!=k1_funct_1(X1,X6)|X1!=k1_funct_1(X3,X4)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v1_matrix_1(X3)|~m1_subset_1(X5,X2)|~m2_finseq_1(X1,X2)|~m2_finseq_1(X3,k3_finseq_2(X2))|~r2_hidden(k4_tarski(X4,X6),k2_matrix_1(X3))), inference(evalgc,[status(thm)],[c_0_162])).
cnf(c_0_205, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(k4_tarski(X1,X2),k2_matrix_1(esk11_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_164])])).
cnf(c_0_206, plain, (v1_matrix_1(X1)|v1_xboole_0(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4)), inference(evalgc,[status(thm)],[c_0_165])).
cnf(c_0_207, plain, (m2_finseq_1(X1,k3_finseq_2(X2))|v1_xboole_0(X2)|~m1_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X3)|~v7_ordinal1(X4)), inference(split_conjunct,[status(thm)],[c_0_123])).
cnf(c_0_208, plain, (X1=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X2)),k1_tarski(k4_tarski(np__1,X3))),k1_tarski(k4_tarski(np__1,X4)))|k3_finseq_1(X1)!=np__3|k1_funct_1(X1,np__3)!=X4|k1_funct_1(X1,np__1)!=X2|k1_funct_1(X1,np__2)!=X3|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(rw,[status(thm)],[c_0_166, c_0_58])).
cnf(c_0_209, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_126]), c_0_127])])).
cnf(c_0_210, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_169])])).
cnf(c_0_211, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_170])).
cnf(c_0_212, plain, (l1_struct_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_172])])).
cnf(c_0_213, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X1!=k10_matrix_3(np__3,esk1_0,esk11_0,X4)|X3!=k1_funct_1(X4,X2)|~r2_hidden(X2,k4_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(esk1_0))|~v7_ordinal1(np__3)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X4,k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_164]), c_0_174]), c_0_175]), c_0_176]), c_0_177]), c_0_178]), c_0_179]), c_0_180]), c_0_181]), c_0_182]), c_0_183])]), c_0_184]), c_0_185])])).
cnf(c_0_214, plain, (v7_ordinal1(np__3)|~v1_finset_1(np__3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_187])])).
cnf(c_0_215, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_188, c_0_108])).
cnf(c_0_216, plain, (k4_finseq_1(X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_189])).
cnf(c_0_217, plain, (m2_subset_1(np__3,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_190])).
cnf(c_0_218, plain, (k3_finseq_1(X1)=X2|v2_struct_0(X3)|v6_struct_0(X3)|X1!=k10_matrix_3(X2,X3,X4,X5)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X5,k12_matrix_2(X2))|~m2_finseq_1(X1,u1_struct_0(X3))|~m1_matrix_1(X4,u1_struct_0(X3),X2,X2)), inference(evalgc,[status(thm)],[c_0_191])).
fof(c_0_219, plain, ![X74, X75, X76, X77]:(~v7_ordinal1(X74)|v2_struct_0(X75)|v6_struct_0(X75)|~v13_algstr_0(X75)|~v33_algstr_0(X75)|~v3_group_1(X75)|~v5_group_1(X75)|~v2_rlvect_1(X75)|~v3_rlvect_1(X75)|~v4_rlvect_1(X75)|~v4_vectsp_1(X75)|~v5_vectsp_1(X75)|~l6_algstr_0(X75)|~m1_matrix_1(X76,u1_struct_0(X75),X74,X74)|~m1_subset_1(X77,k12_matrix_2(X74))|m2_finseq_1(k10_matrix_3(X74,X75,X76,X77),u1_struct_0(X75))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k10_matrix_3])])])).
cnf(c_0_220, plain, (r2_hidden(X1,a_1_0_finseq_1(X2))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_192])])).
cnf(c_0_221, plain, (a_1_0_finseq_1(X1)=k2_finseq_1(X1)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_194])])).
cnf(c_0_222, plain, (k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4))=k2_finseq_1(X5)|v1_xboole_0(X1)|X5!=np__3|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X5,k5_numbers)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_196])])).
cnf(c_0_223, plain, (v1_relat_1(k3_finseq_4(X4,X1,X2,X3))|v1_xboole_0(X4)|~m1_subset_1(X3,X4)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_79, c_0_96])])).
cnf(c_0_224, plain, (v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_198])])).
cnf(c_0_225, plain, (v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_198])])).
cnf(c_0_226, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_200])).
cnf(c_0_227, plain, (m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_201])).
cnf(c_0_228, negated_conjecture, (esk12_0=k3_finseq_4(k5_numbers,np__1,np__2,np__3)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_229, plain, (m1_subset_1(np__2,k5_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
cnf(c_0_230, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_231, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_202])).
cnf(c_0_232, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_203])).
cnf(c_0_233, negated_conjecture, (X1=k3_matrix_1(X2,esk11_0,X3,X4)|v1_xboole_0(u1_struct_0(esk1_0))|X1!=k1_funct_1(X5,X4)|X5!=k1_funct_1(esk11_0,X3)|~v1_matrix_1(esk11_0)|~m2_finseq_1(esk11_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,X2)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,np__3)|~r1_xxreal_0(X3,np__3)|~v7_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_205])])).
cnf(c_0_234, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_matrix_1(esk11_0)|~v7_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_164])])).
cnf(c_0_235, plain, (v1_xboole_0(X2)|m2_finseq_1(X1,k3_finseq_2(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4)), inference(evalgc,[status(thm)],[c_0_207])).
cnf(c_0_236, plain, (X1=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X2)),k1_tarski(k4_tarski(np__1,X3))),k1_tarski(k4_tarski(np__1,X4)))|k3_finseq_1(X1)!=np__3|k1_funct_1(X1,np__3)!=X4|k1_funct_1(X1,np__1)!=X2|k1_funct_1(X1,np__2)!=X3|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_208])).
cnf(c_0_237, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_169])])).
cnf(c_0_238, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_169])])).
cnf(c_0_239, negated_conjecture, (m1_subset_1(esk10_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_240, negated_conjecture, (m1_subset_1(esk9_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_241, negated_conjecture, (m1_subset_1(esk8_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_242, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_212])])).
cnf(c_0_243, plain, (k1_funct_1(X1,np__2)=X2|X1!=k11_finseq_1(X3,X2,X4)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(split_conjunct,[status(thm)],[c_0_52])).
cnf(c_0_244, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X3!=k1_funct_1(X1,X2)|~r2_hidden(X2,k4_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1)))|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(np__3)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X1,k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_213])])).
cnf(c_0_245, plain, (v7_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_215]), c_0_148])])])).
cnf(c_0_246, plain, (k9_xtuple_0(X1)=k4_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_216])).
cnf(c_0_247, plain, (k9_xtuple_0(X1)=k2_finseq_1(np__3)|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_217])])).
cnf(c_0_248, negated_conjecture, (k3_finseq_1(X1)=np__3|X1!=k10_matrix_3(np__3,esk1_0,esk11_0,X2)|~m2_finseq_1(X1,u1_struct_0(esk1_0))|~v7_ordinal1(np__3)|~m1_matrix_2(X2,k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_164]), c_0_174]), c_0_175]), c_0_176]), c_0_177]), c_0_178]), c_0_179]), c_0_180]), c_0_181]), c_0_182]), c_0_183])]), c_0_184]), c_0_185])])).
cnf(c_0_249, plain, (v2_struct_0(X2)|v6_struct_0(X2)|m2_finseq_1(k10_matrix_3(X1,X2,X3,X4),u1_struct_0(X2))|~v7_ordinal1(X1)|~v13_algstr_0(X2)|~v33_algstr_0(X2)|~v3_group_1(X2)|~v5_group_1(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~m1_matrix_1(X3,u1_struct_0(X2),X1,X1)|~m1_subset_1(X4,k12_matrix_2(X1))), inference(split_conjunct,[status(thm)],[c_0_219])).
cnf(c_0_250, plain, (r2_hidden(X1,k2_finseq_1(X2))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_221])])).
cnf(c_0_251, plain, (k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4))=k2_finseq_1(X5)|v1_xboole_0(X1)|X5!=np__3|~m1_subset_1(X5,k5_numbers)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_222, c_0_223]), c_0_224]), c_0_225])).
cnf(c_0_252, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_226])).
cnf(c_0_253, plain, (m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_227])).
cnf(c_0_254, negated_conjecture, (k3_finseq_4(k5_numbers,np__1,np__2,np__3)=esk12_0), inference(evalgc,[status(thm)],[c_0_228])).
cnf(c_0_255, plain, (m1_subset_1(np__2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_229])).
cnf(c_0_256, plain, (m1_subset_1(np__1,k5_numbers)), inference(evalgc,[status(thm)],[c_0_230])).
cnf(c_0_257, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_231])).
cnf(c_0_258, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_232])).
fof(c_0_259, plain, ![X69, X70]:((~m1_matrix_2(X70,X69)|m1_subset_1(X70,X69)|(v1_xboole_0(X69)|~v3_matrix_2(X69)))&(~m1_subset_1(X70,X69)|m1_matrix_2(X70,X69)|(v1_xboole_0(X69)|~v3_matrix_2(X69)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m1_matrix_2])])])])])).
cnf(c_0_260, negated_conjecture, (X1=k3_matrix_1(X2,esk11_0,X3,X4)|v1_xboole_0(u1_struct_0(esk1_0))|X5!=k1_funct_1(esk11_0,X3)|X1!=k1_funct_1(X5,X4)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,np__3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(esk11_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~v7_ordinal1(np__3)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,X2)), inference(csr,[status(thm)],[c_0_233, c_0_234])).
cnf(c_0_261, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|m2_finseq_1(esk11_0,k3_finseq_2(u1_struct_0(esk1_0)))|~v7_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_164])])).
cnf(c_0_262, plain, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X2)),k1_tarski(k4_tarski(np__1,X3))),k1_tarski(k4_tarski(np__1,k1_funct_1(X1,np__3))))=X1|k1_funct_1(X1,np__1)!=X2|k1_funct_1(X1,np__2)!=X3|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_236])])).
cnf(c_0_263, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_169])])).
cnf(c_0_264, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_169])])).
cnf(c_0_265, negated_conjecture, (m1_subset_1(esk10_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_239])).
cnf(c_0_266, negated_conjecture, (m1_subset_1(esk9_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_240])).
cnf(c_0_267, negated_conjecture, (m1_subset_1(esk8_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_241])).
cnf(c_0_268, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_174]), c_0_185])])).
cnf(c_0_269, negated_conjecture, (m1_subset_1(esk7_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_270, negated_conjecture, (m1_subset_1(esk6_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_271, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_272, plain, (k1_funct_1(X1,np__2)=X2|X1!=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X3)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X4)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(rw,[status(thm)],[c_0_243, c_0_58])).
cnf(c_0_273, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X3!=k1_funct_1(X1,X2)|~r2_hidden(X2,k4_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1)))|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X1,k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_244, c_0_245])])])).
cnf(c_0_274, plain, (k4_finseq_1(X1)=k2_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__3), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_247])])).
cnf(c_0_275, negated_conjecture, (k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))=np__3|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(np__3)|~m1_matrix_2(X1,k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_248])])).
cnf(c_0_276, plain, (v2_struct_0(X2)|v6_struct_0(X2)|m2_finseq_1(k10_matrix_3(X1,X2,X3,X4),u1_struct_0(X2))|~v13_algstr_0(X2)|~v33_algstr_0(X2)|~v3_group_1(X2)|~v5_group_1(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v7_ordinal1(X1)|~m1_subset_1(X4,k12_matrix_2(X1))|~m1_matrix_1(X3,u1_struct_0(X2),X1,X1)), inference(evalgc,[status(thm)],[c_0_249])).
cnf(c_0_277, plain, (r2_hidden(X1,k9_xtuple_0(k3_finseq_4(X3,X4,X5,X6)))|v1_xboole_0(X3)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|X2!=np__3|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X6,X3)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_251])])).
cnf(c_0_278, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_252, c_0_108])).
cnf(c_0_279, plain, (r1_xxreal_0(np__3,np__3)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r3_r3])).
cnf(c_0_280, plain, (r1_xxreal_0(np__1,np__3)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r1_r3])).
cnf(c_0_281, plain, (k9_xtuple_0(X1)=k2_finseq_1(k3_finseq_1(X2))|k3_finseq_1(X2)!=k3_finseq_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_253])])).
cnf(c_0_282, negated_conjecture, (m2_finseq_1(esk12_0,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_254]), c_0_148]), c_0_255]), c_0_256])]), c_0_157])])).
cnf(c_0_283, plain, (v1_relat_1(X1)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_258])])).
cnf(c_0_284, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~m1_matrix_2(X1,X2)|~v3_matrix_2(X2)), inference(split_conjunct,[status(thm)],[c_0_259])).
cnf(c_0_285, negated_conjecture, (m1_matrix_2(esk12_0,k12_matrix_2(np__3))), inference(split_conjunct,[status(thm)],[c_0_81])).
fof(c_0_286, plain, ![X83]:((~v1_xboole_0(k12_matrix_2(X83))|~v7_ordinal1(X83))&(v3_matrix_2(k12_matrix_2(X83))|~v7_ordinal1(X83))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_matrix_2])])])])).
cnf(c_0_287, negated_conjecture, (X1=k3_matrix_1(X2,esk11_0,X3,X4)|v1_xboole_0(u1_struct_0(esk1_0))|X5!=k1_funct_1(esk11_0,X3)|X1!=k1_funct_1(X5,X4)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,np__3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(esk11_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_260, c_0_245])])])).
cnf(c_0_288, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|m2_finseq_1(esk11_0,k3_finseq_2(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_261, c_0_245])])])).
cnf(c_0_289, plain, (k3_finseq_4(X1,X2,X3,k1_funct_1(X4,np__3))=X4|v1_xboole_0(X1)|~m1_subset_1(k1_funct_1(X4,np__3),X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|k1_funct_1(X4,np__1)!=X2|k1_funct_1(X4,np__2)!=X3|k3_finseq_1(X4)!=np__3|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_262])])).
cnf(c_0_290, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_169])])).
cnf(c_0_291, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_264, c_0_198]), c_0_265]), c_0_266]), c_0_267])]), c_0_268])])).
cnf(c_0_292, negated_conjecture, (m1_subset_1(esk7_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_269])).
cnf(c_0_293, negated_conjecture, (m1_subset_1(esk6_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_270])).
cnf(c_0_294, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_271])).
cnf(c_0_295, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_296, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_297, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_298, plain, (k1_funct_1(X1,np__2)=X2|X1!=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X3)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X4)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1)), inference(evalgc,[status(thm)],[c_0_272])).
cnf(c_0_299, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X3!=k1_funct_1(X1,X2)|~r2_hidden(X2,k2_finseq_1(np__3))|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X1,k12_matrix_2(np__3))|k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))!=np__3|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_274])])).
cnf(c_0_300, negated_conjecture, (k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))=np__3|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~m1_matrix_2(X1,k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_275, c_0_245])])])).
cnf(c_0_301, negated_conjecture, (m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(np__3)|~m1_subset_1(X1,k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_276, c_0_164]), c_0_174]), c_0_175]), c_0_176]), c_0_177]), c_0_178]), c_0_179]), c_0_180]), c_0_181]), c_0_182]), c_0_183])]), c_0_184]), c_0_185])])).
cnf(c_0_302, plain, (v1_xboole_0(X1)|r2_hidden(X2,k9_xtuple_0(k3_finseq_4(X1,X3,X4,X5)))|X6!=np__3|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,X6)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X6,k5_numbers)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)), inference(csr,[status(thm)],[c_0_277, c_0_278])).
cnf(c_0_303, plain, (r1_xxreal_0(np__3,np__3)), inference(evalgc,[status(thm)],[c_0_279])).
cnf(c_0_304, plain, (r1_xxreal_0(np__1,np__3)), inference(evalgc,[status(thm)],[c_0_280])).
cnf(c_0_305, plain, (k2_finseq_1(k3_finseq_1(X1))=k9_xtuple_0(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_281])])).
cnf(c_0_306, negated_conjecture, (k3_finseq_1(esk12_0)=np__3), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_254]), c_0_148]), c_0_255]), c_0_256])]), c_0_157])])).
cnf(c_0_307, negated_conjecture, (v1_finseq_1(esk12_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_282])])).
cnf(c_0_308, negated_conjecture, (v1_funct_1(esk12_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_282])])).
cnf(c_0_309, negated_conjecture, (v1_relat_1(esk12_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_282])])).
cnf(c_0_310, plain, (r1_xxreal_0(np__1,np__1)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r1_r1])).
cnf(c_0_311, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_312, plain, (v1_xboole_0(X2)|m1_subset_1(X1,X2)|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2)), inference(evalgc,[status(thm)],[c_0_284])).
cnf(c_0_313, negated_conjecture, (m1_matrix_2(esk12_0,k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[c_0_285])).
cnf(c_0_314, plain, (v3_matrix_2(k12_matrix_2(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_286])).
cnf(c_0_315, negated_conjecture, (X1=k3_matrix_1(X2,esk11_0,X3,X4)|X5!=k1_funct_1(esk11_0,X3)|X1!=k1_funct_1(X5,X4)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,np__3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(esk11_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,X2)), inference(sr,[status(thm)],[c_0_287, c_0_268])).
cnf(c_0_316, negated_conjecture, (m2_finseq_1(esk11_0,k3_finseq_2(u1_struct_0(esk1_0)))), inference(sr,[status(thm)],[c_0_288, c_0_268])).
cnf(c_0_317, plain, (k3_finseq_4(X1,X2,X3,X7)=k3_finseq_4(X4,X5,X6,X7)|v1_xboole_0(X1)|v1_xboole_0(X4)|k1_funct_1(k3_finseq_4(X4,X5,X6,X7),np__1)!=X2|k1_funct_1(k3_finseq_4(X4,X5,X6,X7),np__2)!=X3|k3_finseq_1(k3_finseq_4(X4,X5,X6,X7))!=np__3|~v1_finseq_1(k3_finseq_4(X4,X5,X6,X7))|~v1_funct_1(k3_finseq_4(X4,X5,X6,X7))|~v1_relat_1(k3_finseq_4(X4,X5,X6,X7))|~m1_subset_1(X7,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X7,X4)|~m1_subset_1(X6,X4)|~m1_subset_1(X5,X4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_289, c_0_167])])).
cnf(c_0_318, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_198]), c_0_265]), c_0_266]), c_0_267])]), c_0_268])])).
cnf(c_0_319, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_198]), c_0_292]), c_0_293]), c_0_294])]), c_0_268])])).
cnf(c_0_320, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_295])).
cnf(c_0_321, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_296])).
cnf(c_0_322, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_297])).
cnf(c_0_323, plain, (k1_funct_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),np__2)=X2), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_298]), c_0_77]), c_0_78]), c_0_79])])])).
cnf(c_0_324, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X3!=k1_funct_1(X1,X2)|~r2_hidden(X2,k2_finseq_1(np__3))|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X1,k12_matrix_2(np__3))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_299, c_0_283]), c_0_197]), c_0_199]), c_0_300])).
cnf(c_0_325, negated_conjecture, (m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_301, c_0_245])])])).
cnf(c_0_326, plain, (v1_xboole_0(X1)|r2_hidden(np__3,k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4)))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_303]), c_0_304]), c_0_217]), c_0_148])])])).
cnf(c_0_327, negated_conjecture, (k9_xtuple_0(esk12_0)=k2_finseq_1(np__3)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_306]), c_0_307]), c_0_308]), c_0_309])])])).
cnf(c_0_328, plain, (r1_xxreal_0(np__1,np__1)), inference(evalgc,[status(thm)],[c_0_310])).
cnf(c_0_329, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_311])).
cnf(c_0_330, plain, (v3_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_256]), c_0_149])])])).
cnf(c_0_331, negated_conjecture, (v1_xboole_0(k12_matrix_2(np__3))|m1_subset_1(esk12_0,k12_matrix_2(np__3))|~v3_matrix_2(k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_312, c_0_313])])).
cnf(c_0_332, plain, (v3_matrix_2(k12_matrix_2(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_314])).
cnf(c_0_333, plain, (~v1_xboole_0(k12_matrix_2(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_286])).
cnf(c_0_334, plain, (r1_xxreal_0(np__2,np__3)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r2_r3])).
cnf(c_0_335, plain, (r1_xxreal_0(np__1,np__2)), inference(split_conjunct,[status(thm)],[rqLessOrEqual__r1_xxreal_0__r1_r2])).
cnf(c_0_336, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
cnf(c_0_337, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X4!=k1_funct_1(esk11_0,X2)|X1!=k1_funct_1(X4,X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,np__3)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(X4,u1_struct_0(esk1_0))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_316])])).
cnf(c_0_338, plain, (k3_finseq_4(X1,X2,X3,X4)=k3_finseq_4(X5,X6,X7,X4)|v1_xboole_0(X5)|v1_xboole_0(X1)|k1_funct_1(k3_finseq_4(X5,X6,X7,X4),np__1)!=X2|k1_funct_1(k3_finseq_4(X5,X6,X7,X4),np__2)!=X3|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X4,X5)|~m1_subset_1(X7,X5)|~m1_subset_1(X6,X5)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_317, c_0_223]), c_0_224]), c_0_225]), c_0_196])).
cnf(c_0_339, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_198]), c_0_292]), c_0_293]), c_0_294])]), c_0_268])])).
cnf(c_0_340, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_319, c_0_198]), c_0_320]), c_0_321]), c_0_322])]), c_0_268])])).
cnf(c_0_341, plain, (k1_funct_1(k3_finseq_4(X4,X1,X2,X3),np__2)=X2|v1_xboole_0(X4)|~m1_subset_1(X3,X4)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_96])])).
cnf(c_0_342, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X3!=k1_funct_1(X1,X2)|~r2_hidden(X2,k2_finseq_1(np__3))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_324, c_0_325])])).
cnf(c_0_343, negated_conjecture, (r2_hidden(np__3,k2_finseq_1(np__3))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_326, c_0_254]), c_0_327]), c_0_148]), c_0_255]), c_0_256])]), c_0_157])])).
cnf(c_0_344, plain, (v1_xboole_0(X1)|r2_hidden(np__1,k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4)))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_304]), c_0_328]), c_0_329]), c_0_148])])])).
cnf(c_0_345, plain, (v7_ordinal1(np__1)|~v1_finset_1(np__1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_330])])).
cnf(c_0_346, negated_conjecture, (v1_xboole_0(k12_matrix_2(np__3))|m1_subset_1(esk12_0,k12_matrix_2(np__3))|~v7_ordinal1(np__3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_331, c_0_332])])).
cnf(c_0_347, plain, (~v7_ordinal1(X1)|~v1_xboole_0(k12_matrix_2(X1))), inference(evalgc,[status(thm)],[c_0_333])).
cnf(c_0_348, plain, (r1_xxreal_0(np__2,np__3)), inference(evalgc,[status(thm)],[c_0_334])).
cnf(c_0_349, plain, (r1_xxreal_0(np__1,np__2)), inference(evalgc,[status(thm)],[c_0_335])).
cnf(c_0_350, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_336])).
cnf(c_0_351, plain, (v3_ordinal1(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_147, c_0_255]), c_0_149])])])).
cnf(c_0_352, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|k3_finseq_4(u1_struct_0(esk1_0),X4,X5,X6)!=k1_funct_1(esk11_0,X2)|X1!=k1_funct_1(k3_finseq_4(u1_struct_0(esk1_0),X4,X5,X6),X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X6,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_198]), c_0_268])])).
cnf(c_0_353, plain, (k3_finseq_4(X1,X2,X3,X4)=k3_finseq_4(X5,X2,X6,X4)|v1_xboole_0(X5)|v1_xboole_0(X1)|k1_funct_1(k3_finseq_4(X5,X2,X6,X4),np__2)!=X3|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X4,X5)|~m1_subset_1(X6,X5)|~m1_subset_1(X2,X5)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(spm,[status(thm)],[c_0_338, c_0_125])])])).
cnf(c_0_354, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|~m2_finseq_1(k1_funct_1(esk11_0,np__1),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_339, c_0_340])])).
cnf(c_0_355, negated_conjecture, (m2_finseq_1(k1_funct_1(esk11_0,np__1),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_340]), c_0_320]), c_0_321]), c_0_322])]), c_0_268])])).
cnf(c_0_356, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_341, c_0_126]), c_0_127])])).
cnf(c_0_357, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X1,X2,X3,X4)),np__3)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,X4)|v1_xboole_0(X1)|~v7_ordinal1(X4)|~m1_matrix_2(k3_finseq_4(X1,X2,X3,X4),k12_matrix_2(np__3))|~m1_subset_1(k3_finseq_4(X1,X2,X3,X4),k12_matrix_2(np__3))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_342, c_0_167]), c_0_343]), c_0_245])])])])).
cnf(c_0_358, negated_conjecture, (v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v7_ordinal1(np__3)|~m1_subset_1(X1,k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_301])])).
cnf(c_0_359, negated_conjecture, (v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v7_ordinal1(np__3)|~m1_subset_1(X1,k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_301])])).
cnf(c_0_360, negated_conjecture, (v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v7_ordinal1(np__3)|~m1_subset_1(X1,k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_301])])).
cnf(c_0_361, negated_conjecture, (r2_hidden(np__1,k2_finseq_1(np__3))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_254]), c_0_327]), c_0_148]), c_0_255]), c_0_256])]), c_0_157])])).
cnf(c_0_362, plain, (v7_ordinal1(np__1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_345, c_0_215]), c_0_256])])])).
cnf(c_0_363, negated_conjecture, (m1_subset_1(esk12_0,k12_matrix_2(np__3))|~v7_ordinal1(np__3)), inference(csr,[status(thm)],[c_0_346, c_0_347])).
cnf(c_0_364, plain, (v1_xboole_0(X1)|r2_hidden(np__2,k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4)))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_348]), c_0_349]), c_0_350]), c_0_148])])])).
cnf(c_0_365, plain, (v7_ordinal1(np__2)|~v1_finset_1(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_351])])).
cnf(c_0_366, negated_conjecture, (k1_funct_1(k3_finseq_4(u1_struct_0(esk1_0),X1,X2,X3),X4)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X5,X4)|k3_finseq_4(u1_struct_0(esk1_0),X1,X2,X3)!=k1_funct_1(esk11_0,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X4,np__3)|~r1_xxreal_0(X5,np__3)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~m1_subset_1(k1_funct_1(k3_finseq_4(u1_struct_0(esk1_0),X1,X2,X3),X4),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_352])])).
cnf(c_0_367, plain, (k3_finseq_4(X1,X2,k1_funct_1(k3_finseq_4(X3,X2,X4,X5),np__2),X5)=k3_finseq_4(X3,X2,X4,X5)|v1_xboole_0(X1)|v1_xboole_0(X3)|~m1_subset_1(X5,X1)|~m1_subset_1(k1_funct_1(k3_finseq_4(X3,X2,X4,X5),np__2),X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~m1_subset_1(X2,X3)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_353])])).
cnf(c_0_368, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_354, c_0_355])])).
cnf(c_0_369, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_169])])).
cnf(c_0_370, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,X6))))=k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X3,X4,X5,X6))|v1_xboole_0(X3)|k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X3,X4,X5,X6)),np__1)!=X1|k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X3,X4,X5,X6)),np__2)!=X2|k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X3,X4,X5,X6)))!=np__3|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X3,X4,X5,X6)))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X3,X4,X5,X6)))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X3,X4,X5,X6)))|~v7_ordinal1(X6)|~m1_matrix_2(k3_finseq_4(X3,X4,X5,X6),k12_matrix_2(np__3))|~m1_subset_1(k3_finseq_4(X3,X4,X5,X6),k12_matrix_2(np__3))|~m1_subset_1(X6,X3)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_357])])).
cnf(c_0_371, negated_conjecture, (v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~m1_subset_1(X1,k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_358, c_0_245])])])).
cnf(c_0_372, negated_conjecture, (v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~m1_subset_1(X1,k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_359, c_0_245])])])).
cnf(c_0_373, negated_conjecture, (v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~m1_subset_1(X1,k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_360, c_0_245])])])).
cnf(c_0_374, negated_conjecture, (k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))=np__3|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_325])])).
cnf(c_0_375, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X1,X2,X3,X4)),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X2)|v1_xboole_0(X1)|~v7_ordinal1(X2)|~m1_matrix_2(k3_finseq_4(X1,X2,X3,X4),k12_matrix_2(np__3))|~m1_subset_1(k3_finseq_4(X1,X2,X3,X4),k12_matrix_2(np__3))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_342, c_0_125]), c_0_361]), c_0_362])])])])).
cnf(c_0_376, negated_conjecture, (m1_subset_1(esk12_0,k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_363, c_0_245])])])).
cnf(c_0_377, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,k1_funct_1(X1,X2))=k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)|~r2_hidden(X2,k2_finseq_1(np__3))|~v7_ordinal1(k1_funct_1(X1,X2))|~v7_ordinal1(X2)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_342])])).
cnf(c_0_378, negated_conjecture, (k1_funct_1(esk12_0,np__2)=np__2), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_341, c_0_254]), c_0_148]), c_0_255]), c_0_256])]), c_0_157])])).
cnf(c_0_379, negated_conjecture, (r2_hidden(np__2,k2_finseq_1(np__3))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_364, c_0_254]), c_0_327]), c_0_148]), c_0_255]), c_0_256])]), c_0_157])])).
cnf(c_0_380, plain, (v7_ordinal1(np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_365, c_0_215]), c_0_255])])])).
cnf(c_0_381, negated_conjecture, (k1_funct_1(k3_finseq_4(X2,X1,X3,X4),X5)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X6,X5)|v1_xboole_0(X2)|k3_finseq_4(X2,X1,X3,X4)!=k1_funct_1(esk11_0,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X5,np__3)|~r1_xxreal_0(X6,np__3)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~m1_subset_1(k1_funct_1(k3_finseq_4(X2,X1,X3,X4),X5),u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(k3_finseq_4(X2,X1,X3,X4),np__2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_367]), c_0_268])])).
cnf(c_0_382, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__3),np__2)=esk9_0), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_341, c_0_368]), c_0_265]), c_0_266]), c_0_267])]), c_0_268])])).
cnf(c_0_383, plain, (k3_finseq_4(X1,X2,X3,X4)=k3_finseq_4(X5,X2,X3,X4)|v1_xboole_0(X1)|v1_xboole_0(X5)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X4,X5)|~m1_subset_1(X3,X5)|~m1_subset_1(X2,X5)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_96])])).
cnf(c_0_384, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_369, c_0_169])])).
cnf(c_0_385, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,X3))))=k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X4,X5,X6,X3))|v1_xboole_0(X4)|k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X4,X5,X6,X3)),np__1)!=X1|k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X4,X5,X6,X3)),np__2)!=X2|~v7_ordinal1(X3)|~m1_matrix_2(k3_finseq_4(X4,X5,X6,X3),k12_matrix_2(np__3))|~m1_subset_1(k3_finseq_4(X4,X5,X6,X3),k12_matrix_2(np__3))|~m1_subset_1(X3,X4)|~m1_subset_1(X6,X4)|~m1_subset_1(X5,X4)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_370, c_0_371]), c_0_372]), c_0_373]), c_0_374])).
cnf(c_0_386, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_375, c_0_254]), c_0_362]), c_0_313]), c_0_376]), c_0_148]), c_0_255]), c_0_256])]), c_0_157])])).
cnf(c_0_387, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),np__2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_377, c_0_378]), c_0_379]), c_0_380]), c_0_313]), c_0_376])])])).
cnf(c_0_388, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__3),X1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X1)|k1_funct_1(esk11_0,np__3)!=k1_funct_1(esk11_0,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k1_funct_1(k1_funct_1(esk11_0,np__3),X1),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_381, c_0_368]), c_0_382]), c_0_266]), c_0_265]), c_0_267]), c_0_266])]), c_0_268])])).
cnf(c_0_389, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__3),np__3)=esk10_0), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_368]), c_0_265]), c_0_266]), c_0_267])]), c_0_268])])).
cnf(c_0_390, plain, (v1_xboole_0(X1)|m2_finseq_1(k3_finseq_4(X5,X2,X3,X4),X1)|v1_xboole_0(X5)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X4,X5)|~m1_subset_1(X3,X5)|~m1_subset_1(X2,X5)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_383])])).
cnf(c_0_391, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_384, c_0_340])])).
cnf(c_0_392, negated_conjecture, (m2_finseq_1(k1_funct_1(esk11_0,np__3),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_368]), c_0_265]), c_0_266]), c_0_267])]), c_0_268])])).
cnf(c_0_393, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,np__3))))=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)!=X1|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X2), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_385, c_0_254]), c_0_386]), c_0_387]), c_0_245]), c_0_313]), c_0_376]), c_0_148]), c_0_255]), c_0_256])]), c_0_157])])).
cnf(c_0_394, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,np__3)=esk10_0|k1_funct_1(esk11_0,np__3)!=k1_funct_1(esk11_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_388, c_0_389]), c_0_304]), c_0_303]), c_0_245]), c_0_265])])])).
cnf(c_0_395, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|v1_xboole_0(X4)|k3_finseq_4(X4,X5,X6,X7)!=k1_funct_1(esk11_0,X2)|X1!=k1_funct_1(k3_finseq_4(X4,X5,X6,X7),X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X7,u1_struct_0(esk1_0))|~m1_subset_1(X6,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk1_0))|~m1_subset_1(X7,X4)|~m1_subset_1(X6,X4)|~m1_subset_1(X5,X4)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_390]), c_0_268])])).
cnf(c_0_396, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_391, c_0_368]), c_0_392])])).
cnf(c_0_397, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,esk10_0)))=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)!=X1|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X2), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_394]), c_0_304]), c_0_303]), c_0_245])])])).
cnf(c_0_398, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,np__1)|v1_xboole_0(X3)|k3_finseq_4(X3,X1,X4,X5)!=k1_funct_1(esk11_0,X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~m1_subset_1(X1,X3)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_125]), c_0_328]), c_0_304]), c_0_362])])])])).
cnf(c_0_399, negated_conjecture, (k1_funct_1(k3_finseq_4(X1,X2,X3,X4),X5)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X6,X5)|v1_xboole_0(X1)|k3_finseq_4(X1,X2,X3,X4)!=k1_funct_1(esk11_0,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X5,np__3)|~r1_xxreal_0(X6,np__3)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~m1_subset_1(k1_funct_1(k3_finseq_4(X1,X2,X3,X4),X5),u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_395])])).
cnf(c_0_400, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_169]), c_0_355])])])).
cnf(c_0_401, negated_conjecture, (k3_finseq_4(X1,X2,X3,esk10_0)=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|v1_xboole_0(X1)|~m1_subset_1(esk10_0,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)!=X2|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X3), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_397])])).
cnf(c_0_402, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,np__1)=esk2_0|k1_funct_1(esk11_0,np__1)!=k1_funct_1(esk11_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_398, c_0_340]), c_0_322]), c_0_320]), c_0_321])]), c_0_268])])).
cnf(c_0_403, negated_conjecture, (X3=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X5,np__2)|v1_xboole_0(X1)|k3_finseq_4(X1,X2,X3,X4)!=k1_funct_1(esk11_0,X5)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,np__3)|~v7_ordinal1(X5)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_399, c_0_341]), c_0_349]), c_0_348]), c_0_380])])])).
cnf(c_0_404, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_198]), c_0_292]), c_0_293]), c_0_294])]), c_0_268])])).
cnf(c_0_405, negated_conjecture, (k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk6_0,esk10_0)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_406, negated_conjecture, (k3_finseq_4(X1,X2,X3,esk10_0)=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|v1_xboole_0(X1)|esk2_0!=X2|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X3|~m1_subset_1(esk10_0,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_401, c_0_402]), c_0_328]), c_0_304]), c_0_362])])])).
cnf(c_0_407, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,np__2)=esk6_0|k1_funct_1(esk11_0,np__2)!=k1_funct_1(esk11_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_403, c_0_404]), c_0_293]), c_0_292]), c_0_294])]), c_0_268])])).
cnf(c_0_408, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk6_0,esk10_0)!=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)), inference(evalgc,[status(thm)],[c_0_405])).
cnf(c_0_409, negated_conjecture, (k3_finseq_4(X1,X2,X3,esk10_0)=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|v1_xboole_0(X1)|esk6_0!=X3|esk2_0!=X2|~m1_subset_1(esk10_0,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_407]), c_0_349]), c_0_348]), c_0_380])])])).
cnf(c_0_410, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_408, c_0_409]), c_0_265]), c_0_293]), c_0_322])]), c_0_268]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 411
# Proof object clause steps            : 326
# Proof object formula steps           : 85
# Proof object conjectures             : 148
# Proof object clause conjectures      : 145
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 86
# Proof object initial formulas used   : 47
# Proof object generating inferences   : 117
# Proof object simplifying inferences  : 323
# Training examples: 211 positive, 1392 negative
# Training: Positive examples begin
cnf(c_0_108, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_183, negated_conjecture, (v13_algstr_0(esk1_0))).# trainpos
cnf(c_0_182, negated_conjecture, (v33_algstr_0(esk1_0))).# trainpos
cnf(c_0_181, negated_conjecture, (v3_group_1(esk1_0))).# trainpos
cnf(c_0_180, negated_conjecture, (v5_group_1(esk1_0))).# trainpos
cnf(c_0_179, negated_conjecture, (v2_rlvect_1(esk1_0))).# trainpos
cnf(c_0_178, negated_conjecture, (v3_rlvect_1(esk1_0))).# trainpos
cnf(c_0_177, negated_conjecture, (v4_rlvect_1(esk1_0))).# trainpos
cnf(c_0_176, negated_conjecture, (v4_vectsp_1(esk1_0))).# trainpos
cnf(c_0_175, negated_conjecture, (v5_vectsp_1(esk1_0))).# trainpos
cnf(c_0_174, negated_conjecture, (l6_algstr_0(esk1_0))).# trainpos
cnf(c_0_185, negated_conjecture, (~v2_struct_0(esk1_0))).# trainpos
cnf(c_0_184, negated_conjecture, (~v6_struct_0(esk1_0))).# trainpos
cnf(c_0_107, plain, (v3_ordinal1(k4_ordinal1))).# trainpos
cnf(c_0_408, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk6_0,esk10_0)!=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))).# trainpos
cnf(c_0_217, plain, (m2_subset_1(np__3,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_117, plain, (~v1_xboole_0(k4_ordinal1))).# trainpos
cnf(c_0_148, plain, (m1_subset_1(np__3,k5_numbers))).# trainpos
cnf(c_0_256, plain, (m1_subset_1(np__1,k5_numbers))).# trainpos
cnf(c_0_255, plain, (m1_subset_1(np__2,k5_numbers))).# trainpos
cnf(c_0_322, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_321, negated_conjecture, (m1_subset_1(esk3_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_303, plain, (r1_xxreal_0(np__3,np__3))).# trainpos
cnf(c_0_304, plain, (r1_xxreal_0(np__1,np__3))).# trainpos
cnf(c_0_329, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_172, plain, (l2_algstr_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_328, plain, (r1_xxreal_0(np__1,np__1))).# trainpos
cnf(c_0_349, plain, (r1_xxreal_0(np__1,np__2))).# trainpos
cnf(c_0_348, plain, (r1_xxreal_0(np__2,np__3))).# trainpos
cnf(c_0_127, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_156, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).# trainpos
cnf(c_0_320, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_294, negated_conjecture, (m1_subset_1(esk5_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_293, negated_conjecture, (m1_subset_1(esk6_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_292, negated_conjecture, (m1_subset_1(esk7_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_267, negated_conjecture, (m1_subset_1(esk8_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_266, negated_conjecture, (m1_subset_1(esk9_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_265, negated_conjecture, (m1_subset_1(esk10_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_347, plain, (~v7_ordinal1(X1)|~v1_xboole_0(k12_matrix_2(X1)))).# trainpos
cnf(c_0_350, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_313, negated_conjecture, (m1_matrix_2(esk12_0,k12_matrix_2(np__3)))).# trainpos
cnf(c_0_164, negated_conjecture, (m1_matrix_1(esk11_0,u1_struct_0(esk1_0),np__3,np__3))).# trainpos
cnf(c_0_131, plain, (l2_struct_0(X1)|~l2_algstr_0(X1))).# trainpos
cnf(c_0_130, plain, (l1_struct_0(X1)|~l2_struct_0(X1))).# trainpos
cnf(c_0_254, negated_conjecture, (k3_finseq_4(k5_numbers,np__1,np__2,np__3)=esk12_0)).# trainpos
cnf(c_0_126, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0))=esk11_0)).# trainpos
cnf(c_0_193, plain, (k1_finseq_1(X1)=k2_finseq_1(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_194, plain, (k1_finseq_1(X1)=a_1_0_finseq_1(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_332, plain, (v3_matrix_2(k12_matrix_2(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_252, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).# trainpos
cnf(c_0_188, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1))).# trainpos
cnf(c_0_186, plain, (v7_ordinal1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1))).# trainpos
cnf(c_0_257, plain, (v1_relat_1(X1)|~m1_finseq_1(X1,X2))).# trainpos
cnf(c_0_197, plain, (v1_funct_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_199, plain, (v1_finseq_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_258, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_211, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).# trainpos
cnf(c_0_147, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_115, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_246, plain, (k9_xtuple_0(X1)=k4_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_253, plain, (m2_subset_1(k3_finseq_1(X1),k1_numbers,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_169, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_312, plain, (v1_xboole_0(X2)|m1_subset_1(X1,X2)|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2))).# trainpos
cnf(c_0_79, plain, (v1_relat_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))))).# trainpos
cnf(c_0_192, plain, (r2_hidden(X2,a_1_0_finseq_1(X3))|X1!=X2|~v7_ordinal1(X3)|~r1_xxreal_0(X1,X3)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_78, plain, (v1_funct_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))))).# trainpos
cnf(c_0_114, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_206, plain, (v1_matrix_1(X1)|v1_xboole_0(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4))).# trainpos
cnf(c_0_92, plain, (k3_finseq_1(X1)=np__3|X1!=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X2)),k1_tarski(k4_tarski(np__1,X3))),k1_tarski(k4_tarski(np__1,X4)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_77, plain, (v1_finseq_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))))).# trainpos
cnf(c_0_198, plain, (v1_xboole_0(X1)|m2_finseq_1(k3_finseq_4(X1,X2,X3,X4),X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_154, plain, (k9_xtuple_0(X2)=k2_finseq_1(X1)|X1!=k3_finseq_1(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_235, plain, (v1_xboole_0(X2)|m2_finseq_1(X1,k3_finseq_2(X2))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_matrix_1(X1,X2,X3,X4))).# trainpos
cnf(c_0_94, plain, (k1_funct_1(X1,np__3)=X2|X1!=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X3)),k1_tarski(k4_tarski(np__1,X4))),k1_tarski(k4_tarski(np__1,X2)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_76, plain, (k1_funct_1(X1,np__1)=X2|X1!=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X2)),k1_tarski(k4_tarski(np__1,X3))),k1_tarski(k4_tarski(np__1,X4)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_298, plain, (k1_funct_1(X1,np__2)=X2|X1!=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X3)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X4)))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_96, plain, (k3_finseq_4(X1,X2,X3,X4)=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X2)),k1_tarski(k4_tarski(np__1,X3))),k1_tarski(k4_tarski(np__1,X4)))|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_236, plain, (X1=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X2)),k1_tarski(k4_tarski(np__1,X3))),k1_tarski(k4_tarski(np__1,X4)))|k3_finseq_1(X1)!=np__3|k1_funct_1(X1,np__3)!=X4|k1_funct_1(X1,np__1)!=X2|k1_funct_1(X1,np__2)!=X3|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).# trainpos
cnf(c_0_218, plain, (k3_finseq_1(X1)=X2|v2_struct_0(X3)|v6_struct_0(X3)|X1!=k10_matrix_3(X2,X3,X4,X5)|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X5,k12_matrix_2(X2))|~m2_finseq_1(X1,u1_struct_0(X3))|~m1_matrix_1(X4,u1_struct_0(X3),X2,X2))).# trainpos
cnf(c_0_204, plain, (X5=k3_matrix_1(X2,X3,X4,X6)|X5!=k1_funct_1(X1,X6)|X1!=k1_funct_1(X3,X4)|~v7_ordinal1(X6)|~v7_ordinal1(X4)|~v1_matrix_1(X3)|~m1_subset_1(X5,X2)|~m2_finseq_1(X1,X2)|~m2_finseq_1(X3,k3_finseq_2(X2))|~r2_hidden(k4_tarski(X4,X6),k2_matrix_1(X3)))).# trainpos
cnf(c_0_163, plain, (v1_xboole_0(X5)|r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X6))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~r1_xxreal_0(X4,X2)|~r1_xxreal_0(X3,X1)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~m1_matrix_1(X6,X5,X1,X2))).# trainpos
cnf(c_0_276, plain, (v2_struct_0(X2)|v6_struct_0(X2)|m2_finseq_1(k10_matrix_3(X1,X2,X3,X4),u1_struct_0(X2))|~v13_algstr_0(X2)|~v33_algstr_0(X2)|~v3_group_1(X2)|~v5_group_1(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v7_ordinal1(X1)|~m1_subset_1(X4,k12_matrix_2(X1))|~m1_matrix_1(X3,u1_struct_0(X2),X1,X1))).# trainpos
cnf(c_0_173, plain, (k1_funct_1(X3,X1)=k3_matrix_1(u1_struct_0(X5),X6,X1,X2)|v2_struct_0(X5)|v6_struct_0(X5)|X2!=k1_funct_1(X4,X1)|X3!=k10_matrix_3(X7,X5,X6,X4)|~v13_algstr_0(X5)|~v33_algstr_0(X5)|~v3_group_1(X5)|~v5_group_1(X5)|~v2_rlvect_1(X5)|~v3_rlvect_1(X5)|~v4_rlvect_1(X5)|~v4_vectsp_1(X5)|~v5_vectsp_1(X5)|~l6_algstr_0(X5)|~v7_ordinal1(X7)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_matrix_2(X4,k12_matrix_2(X7))|~m2_finseq_1(X3,u1_struct_0(X5))|~r2_hidden(X1,k4_finseq_1(X3))|~m1_matrix_1(X6,u1_struct_0(X5),X7,X7))).# trainpos
cnf(c_0_95, plain, (k1_funct_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),np__1)=X1)).# trainpos
cnf(c_0_323, plain, (k1_funct_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),np__2)=X2)).# trainpos
cnf(c_0_118, plain, (k3_finseq_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))=np__3|~v1_finseq_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))))).# trainpos
cnf(c_0_262, plain, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X2)),k1_tarski(k4_tarski(np__1,X3))),k1_tarski(k4_tarski(np__1,k1_funct_1(X1,np__3))))=X1|k1_funct_1(X1,np__1)!=X2|k1_funct_1(X1,np__2)!=X3|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_383, plain, (k3_finseq_4(X1,X2,X3,X4)=k3_finseq_4(X5,X2,X3,X4)|v1_xboole_0(X1)|v1_xboole_0(X5)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X4,X5)|~m1_subset_1(X3,X5)|~m1_subset_1(X2,X5))).# trainpos
cnf(c_0_248, negated_conjecture, (k3_finseq_1(X1)=np__3|X1!=k10_matrix_3(np__3,esk1_0,esk11_0,X2)|~m2_finseq_1(X1,u1_struct_0(esk1_0))|~v7_ordinal1(np__3)|~m1_matrix_2(X2,k12_matrix_2(np__3)))).# trainpos
cnf(c_0_261, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|m2_finseq_1(esk11_0,k3_finseq_2(u1_struct_0(esk1_0)))|~v7_ordinal1(np__3))).# trainpos
cnf(c_0_124, plain, (k1_funct_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),np__3)=X3)).# trainpos
cnf(c_0_282, negated_conjecture, (m2_finseq_1(esk12_0,k5_numbers))).# trainpos
cnf(c_0_220, plain, (r2_hidden(X1,a_1_0_finseq_1(X2))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_351, plain, (v3_ordinal1(np__2))).# trainpos
cnf(c_0_330, plain, (v3_ordinal1(np__1))).# trainpos
cnf(c_0_187, plain, (v3_ordinal1(np__3))).# trainpos
cnf(c_0_171, plain, (l1_struct_0(X1)|~l2_algstr_0(X1))).# trainpos
cnf(c_0_221, plain, (a_1_0_finseq_1(X1)=k2_finseq_1(X1)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_195, plain, (k9_xtuple_0(X1)=k2_finseq_1(X2)|X2!=k3_finseq_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_234, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_matrix_1(esk11_0)|~v7_ordinal1(np__3))).# trainpos
cnf(c_0_283, plain, (v1_relat_1(X1)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_281, plain, (k9_xtuple_0(X1)=k2_finseq_1(k3_finseq_1(X2))|k3_finseq_1(X2)!=k3_finseq_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).# trainpos
cnf(c_0_301, negated_conjecture, (m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(np__3)|~m1_subset_1(X1,k12_matrix_2(np__3)))).# trainpos
cnf(c_0_205, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(k4_tarski(X1,X2),k2_matrix_1(esk11_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~v7_ordinal1(np__3))).# trainpos
cnf(c_0_331, negated_conjecture, (v1_xboole_0(k12_matrix_2(np__3))|m1_subset_1(esk12_0,k12_matrix_2(np__3))|~v3_matrix_2(k12_matrix_2(np__3)))).# trainpos
cnf(c_0_225, plain, (v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_213, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X1!=k10_matrix_3(np__3,esk1_0,esk11_0,X4)|X3!=k1_funct_1(X4,X2)|~r2_hidden(X2,k4_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(esk1_0))|~v7_ordinal1(np__3)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X4,k12_matrix_2(np__3)))).# trainpos
cnf(c_0_223, plain, (v1_relat_1(k3_finseq_4(X4,X1,X2,X3))|v1_xboole_0(X4)|~m1_subset_1(X3,X4)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4))).# trainpos
cnf(c_0_247, plain, (k9_xtuple_0(X1)=k2_finseq_1(np__3)|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_224, plain, (v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_212, plain, (l1_struct_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_125, plain, (k1_funct_1(k3_finseq_4(X4,X1,X2,X3),np__1)=X1|v1_xboole_0(X4)|~m1_subset_1(X3,X4)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4))).# trainpos
cnf(c_0_341, plain, (k1_funct_1(k3_finseq_4(X4,X1,X2,X3),np__2)=X2|v1_xboole_0(X4)|~m1_subset_1(X3,X4)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4))).# trainpos
cnf(c_0_250, plain, (r2_hidden(X1,k2_finseq_1(X2))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_275, negated_conjecture, (k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))=np__3|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(np__3)|~m1_matrix_2(X1,k12_matrix_2(np__3)))).# trainpos
cnf(c_0_305, plain, (k2_finseq_1(k3_finseq_1(X1))=k9_xtuple_0(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).# trainpos
cnf(c_0_308, negated_conjecture, (v1_funct_1(esk12_0))).# trainpos
cnf(c_0_196, plain, (k3_finseq_1(k3_finseq_4(X4,X1,X2,X3))=np__3|v1_xboole_0(X4)|~m1_subset_1(X3,X4)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4))).# trainpos
cnf(c_0_167, plain, (k1_funct_1(k3_finseq_4(X4,X1,X2,X3),np__3)=X3|v1_xboole_0(X4)|~m1_subset_1(X3,X4)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4))).# trainpos
cnf(c_0_346, negated_conjecture, (v1_xboole_0(k12_matrix_2(np__3))|m1_subset_1(esk12_0,k12_matrix_2(np__3))|~v7_ordinal1(np__3))).# trainpos
cnf(c_0_214, plain, (v7_ordinal1(np__3)|~v1_finset_1(np__3))).# trainpos
cnf(c_0_307, negated_conjecture, (v1_finseq_1(esk12_0))).# trainpos
cnf(c_0_345, plain, (v7_ordinal1(np__1)|~v1_finset_1(np__1))).# trainpos
cnf(c_0_244, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X3!=k1_funct_1(X1,X2)|~r2_hidden(X2,k4_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1)))|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(np__3)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X1,k12_matrix_2(np__3)))).# trainpos
cnf(c_0_365, plain, (v7_ordinal1(np__2)|~v1_finset_1(np__2))).# trainpos
cnf(c_0_309, negated_conjecture, (v1_relat_1(esk12_0))).# trainpos
cnf(c_0_233, negated_conjecture, (X1=k3_matrix_1(X2,esk11_0,X3,X4)|v1_xboole_0(u1_struct_0(esk1_0))|X1!=k1_funct_1(X5,X4)|X5!=k1_funct_1(esk11_0,X3)|~v1_matrix_1(esk11_0)|~m2_finseq_1(esk11_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,X2)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,np__3)|~r1_xxreal_0(X3,np__3)|~v7_ordinal1(np__3))).# trainpos
cnf(c_0_359, negated_conjecture, (v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v7_ordinal1(np__3)|~m1_subset_1(X1,k12_matrix_2(np__3)))).# trainpos
cnf(c_0_390, plain, (v1_xboole_0(X1)|m2_finseq_1(k3_finseq_4(X5,X2,X3,X4),X1)|v1_xboole_0(X5)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X4,X5)|~m1_subset_1(X3,X5)|~m1_subset_1(X2,X5))).# trainpos
cnf(c_0_358, negated_conjecture, (v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v7_ordinal1(np__3)|~m1_subset_1(X1,k12_matrix_2(np__3)))).# trainpos
cnf(c_0_274, plain, (k4_finseq_1(X1)=k2_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__3)).# trainpos
cnf(c_0_360, negated_conjecture, (v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v7_ordinal1(np__3)|~m1_subset_1(X1,k12_matrix_2(np__3)))).# trainpos
cnf(c_0_242, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l6_algstr_0(X1))).# trainpos
cnf(c_0_362, plain, (v7_ordinal1(np__1))).# trainpos
cnf(c_0_245, plain, (v7_ordinal1(np__3))).# trainpos
cnf(c_0_380, plain, (v7_ordinal1(np__2))).# trainpos
cnf(c_0_306, negated_conjecture, (k3_finseq_1(esk12_0)=np__3)).# trainpos
cnf(c_0_378, negated_conjecture, (k1_funct_1(esk12_0,np__2)=np__2)).# trainpos
cnf(c_0_168, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).# trainpos
cnf(c_0_356, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).# trainpos
cnf(c_0_209, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).# trainpos
cnf(c_0_222, plain, (k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4))=k2_finseq_1(X5)|v1_xboole_0(X1)|X5!=np__3|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X5,k5_numbers)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_327, negated_conjecture, (k9_xtuple_0(esk12_0)=k2_finseq_1(np__3))).# trainpos
cnf(c_0_268, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_210, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_288, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|m2_finseq_1(esk11_0,k3_finseq_2(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_369, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_376, negated_conjecture, (m1_subset_1(esk12_0,k12_matrix_2(np__3)))).# trainpos
cnf(c_0_325, negated_conjecture, (m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k12_matrix_2(np__3)))).# trainpos
cnf(c_0_300, negated_conjecture, (k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))=np__3|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~m1_matrix_2(X1,k12_matrix_2(np__3)))).# trainpos
cnf(c_0_237, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_287, negated_conjecture, (X1=k3_matrix_1(X2,esk11_0,X3,X4)|v1_xboole_0(u1_struct_0(esk1_0))|X5!=k1_funct_1(esk11_0,X3)|X1!=k1_funct_1(X5,X4)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X4,np__3)|~r1_xxreal_0(X3,np__3)|~m2_finseq_1(esk11_0,k3_finseq_2(X2))|~m2_finseq_1(X5,X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_273, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X3!=k1_funct_1(X1,X2)|~r2_hidden(X2,k4_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1)))|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X1,k12_matrix_2(np__3)))).# trainpos
cnf(c_0_371, negated_conjecture, (v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~m1_subset_1(X1,k12_matrix_2(np__3)))).# trainpos
cnf(c_0_372, negated_conjecture, (v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~m1_subset_1(X1,k12_matrix_2(np__3)))).# trainpos
cnf(c_0_373, negated_conjecture, (v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~m1_subset_1(X1,k12_matrix_2(np__3)))).# trainpos
cnf(c_0_277, plain, (r2_hidden(X1,k9_xtuple_0(k3_finseq_4(X3,X4,X5,X6)))|v1_xboole_0(X3)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|X2!=np__3|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X6,X3)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3))).# trainpos
cnf(c_0_238, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_263, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_384, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_374, negated_conjecture, (k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))=np__3|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).# trainpos
cnf(c_0_344, plain, (v1_xboole_0(X1)|r2_hidden(np__1,k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4)))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_364, plain, (v1_xboole_0(X1)|r2_hidden(np__2,k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4)))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_326, plain, (v1_xboole_0(X1)|r2_hidden(np__3,k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4)))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_299, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X3!=k1_funct_1(X1,X2)|~r2_hidden(X2,k2_finseq_1(np__3))|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X1,k12_matrix_2(np__3))|k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))!=np__3|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1)))).# trainpos
cnf(c_0_264, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_290, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_343, negated_conjecture, (r2_hidden(np__3,k2_finseq_1(np__3)))).# trainpos
cnf(c_0_379, negated_conjecture, (r2_hidden(np__2,k2_finseq_1(np__3)))).# trainpos
cnf(c_0_361, negated_conjecture, (r2_hidden(np__1,k2_finseq_1(np__3)))).# trainpos
cnf(c_0_289, plain, (k3_finseq_4(X1,X2,X3,k1_funct_1(X4,np__3))=X4|v1_xboole_0(X1)|~m1_subset_1(k1_funct_1(X4,np__3),X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|k1_funct_1(X4,np__1)!=X2|k1_funct_1(X4,np__2)!=X3|k3_finseq_1(X4)!=np__3|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4))).# trainpos
cnf(c_0_291, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_318, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_342, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X3!=k1_funct_1(X1,X2)|~r2_hidden(X2,k2_finseq_1(np__3))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).# trainpos
cnf(c_0_339, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_375, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X1,X2,X3,X4)),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X2)|v1_xboole_0(X1)|~v7_ordinal1(X2)|~m1_matrix_2(k3_finseq_4(X1,X2,X3,X4),k12_matrix_2(np__3))|~m1_subset_1(k3_finseq_4(X1,X2,X3,X4),k12_matrix_2(np__3))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_377, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,k1_funct_1(X1,X2))=k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)|~r2_hidden(X2,k2_finseq_1(np__3))|~v7_ordinal1(k1_funct_1(X1,X2))|~v7_ordinal1(X2)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).# trainpos
cnf(c_0_357, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X1,X2,X3,X4)),np__3)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,X4)|v1_xboole_0(X1)|~v7_ordinal1(X4)|~m1_matrix_2(k3_finseq_4(X1,X2,X3,X4),k12_matrix_2(np__3))|~m1_subset_1(k3_finseq_4(X1,X2,X3,X4),k12_matrix_2(np__3))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_317, plain, (k3_finseq_4(X1,X2,X3,X7)=k3_finseq_4(X4,X5,X6,X7)|v1_xboole_0(X1)|v1_xboole_0(X4)|k1_funct_1(k3_finseq_4(X4,X5,X6,X7),np__1)!=X2|k1_funct_1(k3_finseq_4(X4,X5,X6,X7),np__2)!=X3|k3_finseq_1(k3_finseq_4(X4,X5,X6,X7))!=np__3|~v1_finseq_1(k3_finseq_4(X4,X5,X6,X7))|~v1_funct_1(k3_finseq_4(X4,X5,X6,X7))|~v1_relat_1(k3_finseq_4(X4,X5,X6,X7))|~m1_subset_1(X7,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X7,X4)|~m1_subset_1(X6,X4)|~m1_subset_1(X5,X4))).# trainpos
cnf(c_0_319, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_353, plain, (k3_finseq_4(X1,X2,X3,X4)=k3_finseq_4(X5,X2,X6,X4)|v1_xboole_0(X5)|v1_xboole_0(X1)|k1_funct_1(k3_finseq_4(X5,X2,X6,X4),np__2)!=X3|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X4,X5)|~m1_subset_1(X6,X5)|~m1_subset_1(X2,X5))).# trainpos
cnf(c_0_340, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1))).# trainpos
cnf(c_0_387, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),np__2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2))).# trainpos
cnf(c_0_386, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1))).# trainpos
cnf(c_0_355, negated_conjecture, (m2_finseq_1(k1_funct_1(esk11_0,np__1),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_354, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|~m2_finseq_1(k1_funct_1(esk11_0,np__1),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_392, negated_conjecture, (m2_finseq_1(k1_funct_1(esk11_0,np__3),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_382, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__3),np__2)=esk9_0)).# trainpos
cnf(c_0_389, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__3),np__3)=esk10_0)).# trainpos
cnf(c_0_370, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,X6))))=k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X3,X4,X5,X6))|v1_xboole_0(X3)|k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X3,X4,X5,X6)),np__1)!=X1|k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X3,X4,X5,X6)),np__2)!=X2|k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X3,X4,X5,X6)))!=np__3|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X3,X4,X5,X6)))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X3,X4,X5,X6)))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X3,X4,X5,X6)))|~v7_ordinal1(X6)|~m1_matrix_2(k3_finseq_4(X3,X4,X5,X6),k12_matrix_2(np__3))|~m1_subset_1(k3_finseq_4(X3,X4,X5,X6),k12_matrix_2(np__3))|~m1_subset_1(X6,X3)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3))).# trainpos
cnf(c_0_391, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).# trainpos
cnf(c_0_400, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_404, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2))).# trainpos
cnf(c_0_393, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,np__3))))=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)!=X1|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X2)).# trainpos
cnf(c_0_367, plain, (k3_finseq_4(X1,X2,k1_funct_1(k3_finseq_4(X3,X2,X4,X5),np__2),X5)=k3_finseq_4(X3,X2,X4,X5)|v1_xboole_0(X1)|v1_xboole_0(X3)|~m1_subset_1(X5,X1)|~m1_subset_1(k1_funct_1(k3_finseq_4(X3,X2,X4,X5),np__2),X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~m1_subset_1(X2,X3))).# trainpos
cnf(c_0_337, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X4!=k1_funct_1(esk11_0,X2)|X1!=k1_funct_1(X4,X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,np__3)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(X4,u1_struct_0(esk1_0))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_352, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|k3_finseq_4(u1_struct_0(esk1_0),X4,X5,X6)!=k1_funct_1(esk11_0,X2)|X1!=k1_funct_1(k3_finseq_4(u1_struct_0(esk1_0),X4,X5,X6),X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X6,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_395, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|v1_xboole_0(X4)|k3_finseq_4(X4,X5,X6,X7)!=k1_funct_1(esk11_0,X2)|X1!=k1_funct_1(k3_finseq_4(X4,X5,X6,X7),X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X7,u1_struct_0(esk1_0))|~m1_subset_1(X6,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk1_0))|~m1_subset_1(X7,X4)|~m1_subset_1(X6,X4)|~m1_subset_1(X5,X4))).# trainpos
cnf(c_0_399, negated_conjecture, (k1_funct_1(k3_finseq_4(X1,X2,X3,X4),X5)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X6,X5)|v1_xboole_0(X1)|k3_finseq_4(X1,X2,X3,X4)!=k1_funct_1(esk11_0,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X5,np__3)|~r1_xxreal_0(X6,np__3)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~m1_subset_1(k1_funct_1(k3_finseq_4(X1,X2,X3,X4),X5),u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_366, negated_conjecture, (k1_funct_1(k3_finseq_4(u1_struct_0(esk1_0),X1,X2,X3),X4)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X5,X4)|k3_finseq_4(u1_struct_0(esk1_0),X1,X2,X3)!=k1_funct_1(esk11_0,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X4,np__3)|~r1_xxreal_0(X5,np__3)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~m1_subset_1(k1_funct_1(k3_finseq_4(u1_struct_0(esk1_0),X1,X2,X3),X4),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_398, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,np__1)|v1_xboole_0(X3)|k3_finseq_4(X3,X1,X4,X5)!=k1_funct_1(esk11_0,X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~m1_subset_1(X1,X3))).# trainpos
cnf(c_0_381, negated_conjecture, (k1_funct_1(k3_finseq_4(X2,X1,X3,X4),X5)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X6,X5)|v1_xboole_0(X2)|k3_finseq_4(X2,X1,X3,X4)!=k1_funct_1(esk11_0,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X5,np__3)|~r1_xxreal_0(X6,np__3)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~m1_subset_1(k1_funct_1(k3_finseq_4(X2,X1,X3,X4),X5),u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(k3_finseq_4(X2,X1,X3,X4),np__2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_402, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,np__1)=esk2_0|k1_funct_1(esk11_0,np__1)!=k1_funct_1(esk11_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_403, negated_conjecture, (X3=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X5,np__2)|v1_xboole_0(X1)|k3_finseq_4(X1,X2,X3,X4)!=k1_funct_1(esk11_0,X5)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,np__3)|~v7_ordinal1(X5)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_407, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,np__2)=esk6_0|k1_funct_1(esk11_0,np__2)!=k1_funct_1(esk11_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_388, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__3),X1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X1)|k1_funct_1(esk11_0,np__3)!=k1_funct_1(esk11_0,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k1_funct_1(k1_funct_1(esk11_0,np__3),X1),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_394, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,np__3)=esk10_0|k1_funct_1(esk11_0,np__3)!=k1_funct_1(esk11_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_397, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,esk10_0)))=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)!=X1|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X2)).# trainpos
cnf(c_0_401, negated_conjecture, (k3_finseq_4(X1,X2,X3,esk10_0)=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|v1_xboole_0(X1)|~m1_subset_1(esk10_0,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)!=X2|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X3)).# trainpos
cnf(c_0_406, negated_conjecture, (k3_finseq_4(X1,X2,X3,esk10_0)=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|v1_xboole_0(X1)|esk2_0!=X2|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X3|~m1_subset_1(esk10_0,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_409, negated_conjecture, (k3_finseq_4(X1,X2,X3,esk10_0)=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|v1_xboole_0(X1)|esk6_0!=X3|esk2_0!=X2|~m1_subset_1(esk10_0,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_137, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_67, plain, (v2_xxreal_0(np__3))).#trainneg
cnf(i_0_178, plain, (~v1_xboole_0(np__3))).#trainneg
cnf(i_0_94, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_71, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_193, plain, (~v1_finset_1(k4_ordinal1))).#trainneg
cnf(i_0_77, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_210, plain, (v1_xboole_0(esk18_1(X1)))).#trainneg
cnf(i_0_129, plain, (v1_card_1(k1_card_1(X1)))).#trainneg
cnf(i_0_145, plain, (k2_xcmplx_0(np__1,np__1)=np__2)).#trainneg
cnf(i_0_194, plain, (k2_xcmplx_0(np__1,np__2)=np__3)).#trainneg
cnf(i_0_153, plain, (k2_xcmplx_0(np__2,np__1)=np__3)).#trainneg
cnf(i_0_64, plain, (m1_subset_1(np__3,k1_numbers))).#trainneg
cnf(i_0_91, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_68, plain, (m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_204, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_200, plain, (~r1_xxreal_0(np__3,np__1))).#trainneg
cnf(i_0_180, plain, (~r1_xxreal_0(np__3,np__2))).#trainneg
cnf(i_0_201, plain, (~r1_xxreal_0(np__2,np__1))).#trainneg
cnf(i_0_45, plain, (l5_algstr_0(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_186, plain, (r1_xxreal_0(np__2,np__2))).#trainneg
cnf(i_0_136, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_219, plain, (~v1_xboole_0(k3_finseq_2(X1)))).#trainneg
cnf(i_0_175, plain, (~v1_xboole_0(k1_tarski(X1)))).#trainneg
cnf(i_0_161, plain, (k1_card_1(k1_card_1(X1))=k1_card_1(X1))).#trainneg
cnf(i_0_222, plain, (v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_85, plain, (l2_struct_0(X1)|~l4_struct_0(X1))).#trainneg
cnf(i_0_50, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_53, plain, (l4_algstr_0(X1)|~l5_algstr_0(X1))).#trainneg
cnf(i_0_113, plain, (m1_subset_1(esk17_1(X1),X1))).#trainneg
cnf(i_0_176, plain, (m1_finseq_2(k3_finseq_2(X1),X1))).#trainneg
cnf(i_0_52, plain, (l4_struct_0(X1)|~l5_algstr_0(X1))).#trainneg
cnf(i_0_220, plain, (v3_ordinal1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_90, plain, (l1_struct_0(X1)|~l1_algstr_0(X1))).#trainneg
cnf(i_0_151, plain, (l1_struct_0(X1)|~l3_struct_0(X1))).#trainneg
cnf(i_0_109, plain, (l1_struct_0(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_80, plain, (l3_struct_0(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_84, plain, (l3_struct_0(X1)|~l4_struct_0(X1))).#trainneg
cnf(i_0_79, plain, (l3_algstr_0(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_199, plain, (v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_75, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_226, plain, (~v1_xboole_0(X1)|~v3_card_1(X1,np__1))).#trainneg
cnf(i_0_225, plain, (v1_zfmisc_1(X1)|~v3_card_1(X1,np__1))).#trainneg
cnf(i_0_146, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_132, plain, (v7_struct_0(X1)|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_211, plain, (m1_subset_1(esk18_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_105, plain, (v1_relat_1(k1_tarski(k4_tarski(np__1,X1))))).#trainneg
cnf(i_0_104, plain, (v1_funct_1(k1_tarski(k4_tarski(np__1,X1))))).#trainneg
cnf(i_0_177, plain, (v1_finseq_1(k1_tarski(k4_tarski(np__1,X1))))).#trainneg
cnf(i_0_227, plain, (v1_zfmisc_1(k1_tarski(k4_tarski(np__1,X1))))).#trainneg
cnf(i_0_158, plain, (v1_xboole_0(k1_card_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_124, plain, (v1_finset_1(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_89, plain, (v6_struct_0(X1)|~l4_struct_0(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_191, plain, (v2_struct_0(X1)|v1_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_221, plain, (~v1_xboole_0(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))))).#trainneg
cnf(i_0_197, plain, (~v1_zfmisc_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))))).#trainneg
cnf(i_0_165, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_190, plain, (v2_struct_0(X1)|v2_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_117, plain, (k10_matrix_2(X1)=k11_matrix_2(X1)|v1_xboole_0(X1)|~v3_matrix_2(X1))).#trainneg
cnf(i_0_116, plain, (k3_xcmplx_0(np__1,X1)=X1|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_208, plain, (v1_xboole_0(X1)|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_203, plain, (v1_finset_1(X1)|~v1_finset_1(k1_card_1(X1)))).#trainneg
cnf(i_0_173, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_126, plain, (v1_funct_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_125, plain, (v1_finseq_1(X1)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_98, plain, (m2_finseq_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_128, plain, (v5_relat_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_155, plain, (v1_relat_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_131, plain, (m1_subset_1(k2_finseq_1(X1),k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_59, plain, (esk16_2(X1,X2)=X1|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_170, plain, (v1_xboole_0(X1)|v7_ordinal1(k10_matrix_2(X1))|~v3_matrix_2(X1))).#trainneg
cnf(i_0_217, plain, (v4_relat_1(X1,k5_numbers)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_205, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v1_zfmisc_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_147, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_143, plain, (v2_struct_0(X1)|v1_group_1(X1)|~v4_vectsp_1(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_149, plain, (v2_struct_0(X1)|v3_vectsp_1(X1)|~v4_vectsp_1(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_148, plain, (v2_struct_0(X1)|v6_vectsp_1(X1)|~v4_vectsp_1(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_179, plain, (r1_xxreal_0(X1,X1)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_114, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_115, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_154, plain, (v1_funct_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_214, plain, (v1_finset_1(X2)|~v5_finset_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_135, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_152, plain, (v1_xboole_0(X1)|m1_subset_1(k11_matrix_2(X1),k5_numbers)|~v3_matrix_2(X1))).#trainneg
cnf(i_0_81, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_213, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_181, plain, (v1_xreal_0(k2_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_212, plain, (v1_xreal_0(k3_xcmplx_0(X1,X2))|~v1_xreal_0(X2)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_112, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_209, plain, (k2_xcmplx_0(X1,X2)=k2_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_174, plain, (m1_subset_1(k4_finseq_1(X1),k1_zfmisc_1(k5_numbers))|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_32, plain, (X1=k3_finseq_1(X2)|k9_xtuple_0(X2)!=k2_finseq_1(X1)|~v1_relat_1(X2)|~v1_funct_1(X2)|~v1_finseq_1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_121, plain, (v1_funct_1(X1)|v1_xboole_0(X2)|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2))).#trainneg
cnf(i_0_223, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1))))).#trainneg
cnf(i_0_72, plain, (v1_xboole_0(X2)|m1_matrix_2(X1,X2)|~v3_matrix_2(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_106, plain, (k3_finseq_1(X1)=k1_card_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_195, plain, (k3_xcmplx_0(X1,X2)=k3_xcmplx_0(X2,X1)|~v1_xcmplx_0(X2)|~v1_xcmplx_0(X1))).#trainneg
cnf(i_0_138, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_139, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_130, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_110, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_224, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_58, plain, (r1_xxreal_0(np__1,esk16_2(X1,X2))|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_57, plain, (r1_xxreal_0(esk16_2(X1,X2),X2)|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_189, plain, (m2_subset_1(k1_nat_1(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_60, plain, (m2_subset_1(esk16_2(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~r2_hidden(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_196, plain, (v1_funct_2(X1,X2,X3)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_184, plain, (v2_funct_1(X1)|~v1_funct_1(X1)|~v3_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_183, plain, (v2_funct_2(X1,X2)|~v1_funct_1(X1)|~v3_funct_2(X1,X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_160, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_120, plain, (v1_xboole_0(X2)|v1_funct_2(X1,k2_finseq_1(k11_matrix_2(X2)),k2_finseq_1(k11_matrix_2(X2)))|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2))).#trainneg
cnf(i_0_182, plain, (v3_card_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),np__3))).#trainneg
cnf(i_0_119, plain, (v1_xboole_0(X2)|v3_funct_2(X1,k2_finseq_1(k11_matrix_2(X2)),k2_finseq_1(k11_matrix_2(X2)))|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2))).#trainneg
cnf(i_0_118, plain, (v1_xboole_0(X2)|m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k2_finseq_1(k11_matrix_2(X2)),k2_finseq_1(k11_matrix_2(X2)))))|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2))).#trainneg
cnf(i_0_168, plain, (v3_funct_2(X1,X2,X3)|~v1_funct_1(X1)|~v2_funct_1(X1)|~v2_funct_2(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_108, plain, (m1_subset_1(k3_matrix_1(X2,X1,X3,X4),X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X1)|~m1_finseq_1(X1,k3_finseq_2(X2)))).#trainneg
cnf(i_0_43, plain, (esk15_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m1_subset_1(X5,X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X2)))).#trainneg
cnf(i_0_42, plain, (k1_funct_1(esk15_5(X2,X3,X4,X5,X1),X5)=X1|X1!=k3_matrix_1(X2,X3,X4,X5)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~v1_matrix_1(X3)|~m1_subset_1(X1,X2)|~m2_finseq_1(X3,k3_finseq_2(X2))|~r2_hidden(k4_tarski(X4,X5),k2_matrix_1(X3)))).#trainneg
cnf(i_0_34, plain, (X1=k10_matrix_3(X2,X3,X4,X5)|v2_struct_0(X3)|v6_struct_0(X3)|k3_finseq_1(X1)!=X2|k3_matrix_1(u1_struct_0(X3),X4,esk13_5(X2,X3,X4,X5,X1),esk14_5(X2,X3,X4,X5,X1))!=k1_funct_1(X1,esk13_5(X2,X3,X4,X5,X1))|~v13_algstr_0(X3)|~v33_algstr_0(X3)|~v3_group_1(X3)|~v5_group_1(X3)|~v2_rlvect_1(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X5,k12_matrix_2(X2))|~m2_finseq_1(X1,u1_struct_0(X3))|~m1_matrix_1(X4,u1_struct_0(X3),X2,X2))).#trainneg
cnf(i_0_44, plain, (m2_finseq_1(esk15_5(X1,X2,X3,X4,X5),X1)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_matrix_1(X2)|~m1_subset_1(X5,X1)|~m2_finseq_1(X2,k3_finseq_2(X1))|~r2_hidden(k4_tarski(X3,X4),k2_matrix_1(X2)))).#trainneg
cnf(i_0_38, plain, (X5=k10_matrix_3(X1,X2,X3,X4)|v2_struct_0(X2)|v6_struct_0(X2)|v7_ordinal1(esk13_5(X1,X2,X3,X4,X5))|k3_finseq_1(X5)!=X1|~v13_algstr_0(X2)|~v33_algstr_0(X2)|~v3_group_1(X2)|~v5_group_1(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v7_ordinal1(X1)|~m1_matrix_2(X4,k12_matrix_2(X1))|~m2_finseq_1(X5,u1_struct_0(X2))|~m1_matrix_1(X3,u1_struct_0(X2),X1,X1))).#trainneg
cnf(i_0_37, plain, (X5=k10_matrix_3(X1,X2,X3,X4)|v2_struct_0(X2)|v6_struct_0(X2)|v7_ordinal1(esk14_5(X1,X2,X3,X4,X5))|k3_finseq_1(X5)!=X1|~v13_algstr_0(X2)|~v33_algstr_0(X2)|~v3_group_1(X2)|~v5_group_1(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v7_ordinal1(X1)|~m1_matrix_2(X4,k12_matrix_2(X1))|~m2_finseq_1(X5,u1_struct_0(X2))|~m1_matrix_1(X3,u1_struct_0(X2),X1,X1))).#trainneg
cnf(i_0_36, plain, (X5=k10_matrix_3(X1,X2,X3,X4)|v2_struct_0(X2)|v6_struct_0(X2)|r2_hidden(esk13_5(X1,X2,X3,X4,X5),k4_finseq_1(X5))|k3_finseq_1(X5)!=X1|~v13_algstr_0(X2)|~v33_algstr_0(X2)|~v3_group_1(X2)|~v5_group_1(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v7_ordinal1(X1)|~m1_matrix_2(X4,k12_matrix_2(X1))|~m2_finseq_1(X5,u1_struct_0(X2))|~m1_matrix_1(X3,u1_struct_0(X2),X1,X1))).#trainneg
cnf(i_0_35, plain, (X5=k10_matrix_3(X1,X2,X3,X4)|k1_funct_1(X4,esk13_5(X1,X2,X3,X4,X5))=esk14_5(X1,X2,X3,X4,X5)|v2_struct_0(X2)|v6_struct_0(X2)|k3_finseq_1(X5)!=X1|~v13_algstr_0(X2)|~v33_algstr_0(X2)|~v3_group_1(X2)|~v5_group_1(X2)|~v2_rlvect_1(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v7_ordinal1(X1)|~m1_matrix_2(X4,k12_matrix_2(X1))|~m2_finseq_1(X5,u1_struct_0(X2))|~m1_matrix_1(X3,u1_struct_0(X2),X1,X1))).#trainneg
cnf(i_0_233, plain, (l4_struct_0(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_366, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk18_1(X2)))).#trainneg
cnf(i_0_367, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk17_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_316, plain, (v3_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_229, plain, (esk18_1(X1)=k1_xboole_0)).#trainneg
cnf(i_0_357, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_361, plain, (~v1_finset_1(k1_numbers))).#trainneg
cnf(i_0_232, plain, (l4_algstr_0(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_476, plain, (k1_funct_1(X1,np__2)=X2|v1_xboole_0(X5)|X1!=k3_finseq_4(X5,X3,X2,X4)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X4,X5)|~m1_subset_1(X2,X5)|~m1_subset_1(X3,X5))).#trainneg
cnf(i_0_475, plain, (v1_xboole_0(X4)|~v1_zfmisc_1(k3_finseq_4(X4,X1,X2,X3))|~m1_subset_1(X3,X4)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4))).#trainneg
cnf(i_0_477, plain, (k1_funct_1(X1,np__3)=X2|v1_xboole_0(X5)|X1!=k3_finseq_4(X5,X3,X4,X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,X5)|~m1_subset_1(X4,X5)|~m1_subset_1(X3,X5))).#trainneg
cnf(i_0_478, plain, (k3_finseq_1(X1)=np__3|v1_xboole_0(X5)|X1!=k3_finseq_4(X5,X2,X3,X4)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X4,X5)|~m1_subset_1(X3,X5)|~m1_subset_1(X2,X5))).#trainneg
cnf(i_0_437, negated_conjecture, (m2_finseq_1(esk11_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_481, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0))),k1_tarski(k4_tarski(np__1,k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)))),k1_tarski(k4_tarski(np__1,k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0))))=esk11_0|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_479, plain, (k1_funct_1(X1,np__1)=X2|v1_xboole_0(X5)|X1!=k3_finseq_4(X5,X2,X3,X4)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X4,X5)|~m1_subset_1(X3,X5)|~m1_subset_1(X2,X5))).#trainneg
cnf(i_0_369, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k5_numbers))).#trainneg
cnf(i_0_283, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_364, plain, (v1_xreal_0(np__2)|~v1_xreal_0(np__1))).#trainneg
cnf(i_0_378, plain, (k3_finseq_1(X1)=np__2|k9_xtuple_0(X1)!=k2_finseq_1(np__2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_235, plain, (l1_struct_0(X1)|~l4_struct_0(X1))).#trainneg
cnf(i_0_245, plain, (~m1_subset_1(k5_numbers,k5_numbers))).#trainneg
cnf(i_0_342, plain, (~v5_finset_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_237, plain, (l1_struct_0(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_242, plain, (k1_card_1(X1)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_250, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_380, plain, (k3_finseq_1(X1)=np__3|k9_xtuple_0(X1)!=k2_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_473, plain, (v3_card_1(k3_finseq_4(X4,X1,X2,X3),np__3)|v1_xboole_0(X4)|~m1_subset_1(X3,X4)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4))).#trainneg
cnf(i_0_382, plain, (k3_finseq_1(X1)=k3_finseq_1(X2)|k9_xtuple_0(X2)!=k2_finseq_1(k3_finseq_1(X1))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_379, plain, (k3_finseq_1(X1)=np__1|k9_xtuple_0(X1)!=k2_finseq_1(np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_385, plain, (v1_xboole_0(k12_matrix_2(X1))|m1_matrix_2(X2,k12_matrix_2(X1))|~m1_subset_1(X2,k12_matrix_2(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_430, plain, (m1_subset_1(k3_finseq_1(X1),k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_272, plain, (v1_relat_1(k1_xboole_0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_262, plain, (v1_relat_1(esk17_1(X1))|~v4_funct_1(X1))).#trainneg
cnf(i_0_273, plain, (v1_relat_1(np__2)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_425, plain, (X1=k3_finseq_1(X2)|k9_xtuple_0(X2)!=k2_finseq_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_254, plain, (v1_finset_1(X1)|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_274, plain, (v1_relat_1(np__2)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_275, plain, (v1_relat_1(np__1)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_276, plain, (v1_relat_1(np__1)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_279, plain, (v1_relat_1(np__3)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_299, plain, (v1_funct_1(np__2)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_255, plain, (v1_finset_1(X1)|~m1_subset_1(k1_card_1(X1),k5_numbers))).#trainneg
cnf(i_0_263, negated_conjecture, (v1_relat_1(esk10_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_407, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_416, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_408, plain, (k1_nat_1(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_482, plain, (v1_xboole_0(X1)|m2_finseq_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X2)),k1_tarski(k4_tarski(np__1,X3))),k1_tarski(k4_tarski(np__1,X4))),X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_415, plain, (m2_subset_1(k1_nat_1(X2,X1),k1_numbers,k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_455, plain, (v1_relat_1(X1)|v1_xboole_0(X2)|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2))).#trainneg
cnf(i_0_440, plain, (k9_xtuple_0(X1)=k2_finseq_1(k1_nat_1(X2,X3))|k1_nat_1(X2,X3)!=k3_finseq_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_474, plain, (v1_xboole_0(X4)|~v1_xboole_0(k3_finseq_4(X4,X1,X2,X3))|~m1_subset_1(X3,X4)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4))).#trainneg
cnf(i_0_278, plain, (v1_relat_1(np__3)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_286, plain, (v1_funct_1(esk17_1(X1))|~v4_funct_1(X1))).#trainneg
cnf(i_0_296, plain, (v1_funct_1(k1_xboole_0)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_264, negated_conjecture, (v1_relat_1(esk9_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_298, plain, (v1_funct_1(np__2)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_300, plain, (v1_funct_1(np__1)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_301, plain, (v1_funct_1(np__1)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_304, plain, (v1_funct_1(np__3)|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_339, plain, (v1_finset_1(np__2)|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_265, negated_conjecture, (v1_relat_1(esk8_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_266, negated_conjecture, (v1_relat_1(esk7_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_413, plain, (k1_nat_1(X1,X2)=k3_finseq_1(X3)|k9_xtuple_0(X3)!=k2_finseq_1(k1_nat_1(X1,X2))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_442, plain, (k9_xtuple_0(X1)=k2_finseq_1(np__2)|k3_finseq_1(X1)!=np__2|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_282, plain, (v1_xboole_0(X1)|v7_ordinal1(k11_matrix_2(X1))|~v3_matrix_2(X1))).#trainneg
cnf(i_0_443, plain, (k9_xtuple_0(X1)=k2_finseq_1(np__1)|k3_finseq_1(X1)!=np__1|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_484, plain, (esk15_5(X1,X2,X3,X4,X5)=k1_funct_1(X2,X3)|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X5,X1)|~m1_subset_1(k4_tarski(X3,X4),k2_matrix_1(X2)))).#trainneg
cnf(i_0_303, plain, (v1_funct_1(np__3)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_306, plain, (v3_ordinal1(esk17_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_417, plain, (esk16_2(X1,X2)=k3_finseq_1(X3)|k9_xtuple_0(X3)!=k2_finseq_1(esk16_2(X1,X2))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_267, negated_conjecture, (v1_relat_1(esk6_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_336, plain, (v1_finset_1(k1_xboole_0)|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_338, plain, (v1_finset_1(np__2)|~v5_finset_1(k1_numbers))).#trainneg
cnf(i_0_341, plain, (v1_finset_1(np__1)|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_344, plain, (v1_finset_1(np__3)|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_268, negated_conjecture, (v1_relat_1(esk5_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_269, negated_conjecture, (v1_relat_1(esk4_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_429, plain, (m1_subset_1(k1_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_498, plain, (v6_struct_0(X1)|~v7_struct_0(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_500, plain, (k1_card_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))=np__3)).#trainneg
cnf(i_0_663, negated_conjecture, (v3_card_1(esk12_0,np__3))).#trainneg
cnf(i_0_366, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_211, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_270, negated_conjecture, (v1_relat_1(esk3_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_340, plain, (v1_finset_1(np__1)|~v5_finset_1(k1_numbers))).#trainneg
cnf(i_0_574, negated_conjecture, (~v1_zfmisc_1(esk12_0))).#trainneg
cnf(i_0_271, negated_conjecture, (v1_relat_1(esk2_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_277, plain, (v1_relat_1(k5_numbers)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_621, negated_conjecture, (k1_funct_1(X1,np__1)=np__1|X1!=esk12_0|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_694, plain, (k1_nat_1(X2,X1)=k2_xcmplx_0(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_527, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(X1)|k3_finseq_4(X1,esk2_0,esk6_0,esk10_0)!=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk2_0,X1))).#trainneg
cnf(i_0_343, plain, (v1_finset_1(np__3)|~v5_finset_1(k1_numbers))).#trainneg
cnf(i_0_630, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_753, plain, (k3_finseq_1(X1)=np__2|k2_finseq_1(np__2)!=k2_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__3)).#trainneg
cnf(i_0_280, plain, (v1_relat_1(k2_finseq_1(X1))|~v4_funct_1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_557, plain, (v7_ordinal1(k1_xboole_0)|~v1_finset_1(k1_xboole_0))).#trainneg
cnf(i_0_737, negated_conjecture, (~v1_xboole_0(esk12_0))).#trainneg
cnf(i_0_563, plain, (~m1_subset_1(k1_numbers,k5_numbers))).#trainneg
cnf(i_0_543, negated_conjecture, (v1_funct_1(esk11_0)|v1_xboole_0(u1_struct_0(esk1_0))|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_287, negated_conjecture, (v1_funct_1(esk10_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_288, negated_conjecture, (v1_funct_1(esk9_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_580, negated_conjecture, (k1_funct_1(X1,np__3)=np__3|X1!=esk12_0|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_569, negated_conjecture, (k1_funct_1(X1,np__2)=np__2|X1!=esk12_0|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_695, plain, (k2_xcmplx_0(X2,X1)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_617, plain, (k9_xtuple_0(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))=k2_finseq_1(X4)|X4!=np__3|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_632, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_754, plain, (k3_finseq_1(X1)=np__1|k2_finseq_1(np__3)!=k2_finseq_1(np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__3)).#trainneg
cnf(i_0_289, negated_conjecture, (v1_funct_1(esk8_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_362, plain, (v1_xreal_0(np__3)|~v1_xreal_0(np__1)|~v1_xreal_0(np__2))).#trainneg
cnf(i_0_544, negated_conjecture, (v1_finseq_1(esk11_0)|v1_xboole_0(u1_struct_0(esk1_0))|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_654, negated_conjecture, (v1_relat_1(esk11_0)|v1_xboole_0(u1_struct_0(esk1_0))|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_290, negated_conjecture, (v1_funct_1(esk7_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_291, negated_conjecture, (v1_funct_1(esk6_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_674, plain, (v1_xboole_0(k12_matrix_2(X1))|v1_funct_1(X2)|~v3_matrix_2(k12_matrix_2(X1))|~v7_ordinal1(X1)|~m1_subset_1(X2,k12_matrix_2(X1)))).#trainneg
cnf(i_0_719, plain, (m1_subset_1(k1_nat_1(X1,X2),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_586, negated_conjecture, (k3_finseq_1(X1)=np__3|X1!=esk12_0|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_684, negated_conjecture, (esk15_5(X1,esk11_0,X2,X3,X4)=k1_funct_1(esk11_0,X2)|v1_xboole_0(u1_struct_0(esk1_0))|X4!=k3_matrix_1(X1,esk11_0,X2,X3)|~v1_matrix_1(esk11_0)|~m2_finseq_1(esk11_0,k3_finseq_2(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X4,X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(np__3))).#trainneg
cnf(i_0_646, plain, (v1_card_1(k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_636, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_761, plain, (k3_finseq_1(X1)=np__1|k2_finseq_1(np__2)!=k2_finseq_1(np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__2)).#trainneg
cnf(i_0_292, negated_conjecture, (v1_funct_1(esk5_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_293, negated_conjecture, (v1_funct_1(esk4_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_294, negated_conjecture, (v1_funct_1(esk3_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_403, plain, (v5_relat_1(X1,X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_631, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_295, negated_conjecture, (v1_funct_1(esk2_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_297, plain, (v1_funct_1(k2_finseq_1(X1))|~v4_funct_1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_699, plain, (k9_xtuple_0(X1)=k2_finseq_1(k2_xcmplx_0(X2,X3))|k2_xcmplx_0(X2,X3)!=k3_finseq_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(X2)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_660, plain, (k9_xtuple_0(X1)=k2_finseq_1(k1_card_1(X2))|k1_card_1(X2)!=k3_finseq_1(X1)|~v1_finseq_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_368, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_326, plain, (v1_finset_1(esk17_1(X1))|~v5_finset_1(X1))).#trainneg
cnf(i_0_302, plain, (v1_funct_1(k5_numbers)|~v4_funct_1(k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_762, plain, (k3_finseq_1(X1)=np__3|k2_finseq_1(np__2)!=k2_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__2)).#trainneg
cnf(i_0_307, negated_conjecture, (v3_ordinal1(esk10_0)|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_649, plain, (v1_finset_1(k1_xboole_0)|~v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_308, negated_conjecture, (v3_ordinal1(esk9_0)|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_309, negated_conjecture, (v3_ordinal1(esk8_0)|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_635, plain, (v1_xboole_0(k2_zfmisc_1(X1,X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_310, negated_conjecture, (v3_ordinal1(esk7_0)|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_311, negated_conjecture, (v3_ordinal1(esk6_0)|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_312, negated_conjecture, (v3_ordinal1(esk5_0)|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_618, plain, (k9_xtuple_0(X1)=k2_finseq_1(X2)|X2!=k1_card_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_523, negated_conjecture, (k3_finseq_4(X1,np__1,np__2,np__3)=esk12_0|v1_xboole_0(X1)|~m1_subset_1(np__3,X1)|~m1_subset_1(np__2,X1)|~m1_subset_1(np__1,X1))).#trainneg
cnf(i_0_728, plain, (v1_xboole_0(k12_matrix_2(X1))|v1_relat_1(X2)|~v3_matrix_2(k12_matrix_2(X1))|~v7_ordinal1(X1)|~m1_subset_1(X2,k12_matrix_2(X1)))).#trainneg
cnf(i_0_732, plain, (k9_xtuple_0(X1)=k2_finseq_1(k2_xcmplx_0(X3,X2))|k2_xcmplx_0(X3,X2)!=k3_finseq_1(X1)|~v7_ordinal1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_721, plain, (k9_xtuple_0(X1)=k2_finseq_1(k1_nat_1(X2,X3))|k1_nat_1(X2,X3)!=k3_finseq_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v7_ordinal1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_560, plain, (esk16_2(X1,X2)=X1|~v7_ordinal1(X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_678, plain, (v3_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_313, negated_conjecture, (v3_ordinal1(esk4_0)|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_346, plain, (v1_relat_1(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_314, negated_conjecture, (v3_ordinal1(esk3_0)|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_315, negated_conjecture, (v3_ordinal1(esk2_0)|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_327, negated_conjecture, (v1_finset_1(esk10_0)|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_328, negated_conjecture, (v1_finset_1(esk9_0)|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_329, negated_conjecture, (v1_finset_1(esk8_0)|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_330, negated_conjecture, (v1_finset_1(esk7_0)|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_331, negated_conjecture, (v1_finset_1(esk6_0)|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_317, plain, (v3_ordinal1(k2_finseq_1(X1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_332, negated_conjecture, (v1_finset_1(esk5_0)|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_509, plain, (v3_card_1(X3,np__3)|k1_funct_1(X3,np__1)!=X1|k1_funct_1(X3,np__2)!=X2|k3_finseq_1(X3)!=np__3|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_670, plain, (k3_finseq_1(X1)=np__1|k4_finseq_1(X1)!=k2_finseq_1(np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_333, negated_conjecture, (v1_finset_1(esk4_0)|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_355, plain, (v1_xboole_0(esk17_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_334, negated_conjecture, (v1_finset_1(esk3_0)|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_773, plain, (k3_finseq_1(X1)=np__2|k2_finseq_1(np__2)!=k2_finseq_1(np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__1)).#trainneg
cnf(i_0_335, negated_conjecture, (v1_finset_1(esk2_0)|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_647, plain, (k1_card_1(k1_xboole_0)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_573, negated_conjecture, (~v1_zfmisc_1(esk11_0)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_626, plain, (v1_relat_1(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_627, plain, (v1_funct_1(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_337, plain, (v1_finset_1(k2_finseq_1(X1))|~v5_finset_1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_347, plain, (v1_relat_1(k11_matrix_2(X1))|v1_xboole_0(X1)|~v4_funct_1(k5_numbers)|~v3_matrix_2(X1))).#trainneg
cnf(i_0_681, plain, (m1_subset_1(k1_card_1(X1),k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1090, plain, (v1_finset_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_829, plain, (k2_finseq_1(k1_card_1(X1))=k9_xtuple_0(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_811, plain, (v1_card_1(np__3))).#trainneg
cnf(i_0_808, plain, (v6_struct_0(X1)|~l6_algstr_0(X1)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_812, plain, (k1_card_1(np__3)=np__3)).#trainneg
cnf(i_0_835, plain, (v1_relat_1(k1_xboole_0))).#trainneg
cnf(i_0_921, plain, (np__2=np__3|k2_finseq_1(X4)!=k2_finseq_1(np__2)|X4!=np__3|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_348, plain, (v1_funct_1(k11_matrix_2(X1))|v1_xboole_0(X1)|~v4_funct_1(k5_numbers)|~v3_matrix_2(X1))).#trainneg
cnf(i_0_890, plain, (v7_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_971, plain, (v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_796, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_628, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_349, plain, (v3_ordinal1(k11_matrix_2(X1))|v1_xboole_0(X1)|~v3_matrix_2(X1))).#trainneg
cnf(i_0_350, plain, (v1_finset_1(k11_matrix_2(X1))|v1_xboole_0(X1)|~v5_finset_1(k5_numbers)|~v3_matrix_2(X1))).#trainneg
cnf(i_0_927, plain, (k2_finseq_1(X4)=k4_finseq_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))|X4!=np__3|~m1_subset_1(X4,k5_numbers))).#trainneg
cnf(i_0_799, negated_conjecture, (k1_funct_1(esk12_0,np__1)=np__1)).#trainneg
cnf(i_0_828, plain, (k9_xtuple_0(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))=k2_finseq_1(np__3))).#trainneg
cnf(i_0_996, plain, (v1_finset_1(k1_xboole_0))).#trainneg
cnf(i_0_837, plain, (v4_relat_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_1080, plain, (k1_card_1(k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_775, plain, (k3_finseq_1(X1)=np__3|k2_finseq_1(np__3)!=k2_finseq_1(np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__1)).#trainneg
cnf(i_0_351, plain, (esk16_2(X1,X2)=X1|v1_xboole_0(a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_994, negated_conjecture, (v3_ordinal1(esk10_0)|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_998, negated_conjecture, (v3_ordinal1(esk9_0)|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_629, plain, (v1_finset_1(X1)|~v5_finset_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_999, negated_conjecture, (v3_ordinal1(esk8_0)|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_359, plain, (v1_finset_1(esk17_1(k1_zfmisc_1(X1)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_810, plain, (v1_finset_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))))).#trainneg
cnf(i_0_854, plain, (k1_card_1(k3_finseq_4(X1,X2,X3,X4))=np__3|v1_xboole_0(X1)|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_849, plain, (k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4))=k2_finseq_1(np__3)|v1_xboole_0(X1)|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_856, negated_conjecture, (k3_finseq_1(esk11_0)=np__3|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_864, negated_conjecture, (k1_funct_1(esk12_0,np__3)=np__3)).#trainneg
cnf(i_0_880, plain, (m2_subset_1(k2_xcmplx_0(X1,X2),k1_numbers,k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_841, plain, (v5_relat_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_1003, negated_conjecture, (v3_ordinal1(esk7_0)|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_839, plain, (v2_funct_1(k1_xboole_0)|~v3_funct_2(k1_xboole_0,X1,X2)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_1004, negated_conjecture, (v3_ordinal1(esk6_0)|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1005, negated_conjecture, (v3_ordinal1(esk5_0)|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1061, negated_conjecture, (v3_ordinal1(esk4_0)|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1062, negated_conjecture, (v3_ordinal1(esk3_0)|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_370, plain, (v1_relat_1(k4_finseq_1(X1))|~v4_funct_1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_371, plain, (v1_funct_1(k4_finseq_1(X1))|~v4_funct_1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_826, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_819, negated_conjecture, (k1_card_1(esk11_0)=np__3|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1085, negated_conjecture, (~v1_zfmisc_1(esk11_0)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_832, plain, (v1_funct_1(k1_xboole_0)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1063, negated_conjecture, (v3_ordinal1(esk2_0)|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_372, plain, (v3_ordinal1(k4_finseq_1(X1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_838, plain, (v3_funct_2(k1_xboole_0,X1,X2)|~v2_funct_2(k1_xboole_0,X2)|~v2_funct_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_601, negated_conjecture, (~v1_xboole_0(esk11_0)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_406, plain, (v4_relat_1(X1,k5_numbers)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_395, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_373, plain, (v1_finset_1(k4_finseq_1(X1))|~v5_finset_1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_888, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(X1)|k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk2_0)),k1_tarski(k4_tarski(np__1,esk6_0))),k1_tarski(k4_tarski(np__1,esk10_0)))!=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk2_0,X1))).#trainneg
cnf(i_0_958, negated_conjecture, (k1_card_1(X1)=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|X1!=esk12_0)).#trainneg
cnf(i_0_985, plain, (k9_xtuple_0(X1)=k2_finseq_1(k1_card_1(X2))|k1_card_1(X2)!=k1_card_1(X1)|~v1_finseq_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_597, negated_conjecture, (v1_relat_1(esk11_0)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_842, plain, (v2_funct_2(k1_xboole_0,X1)|~v3_funct_2(k1_xboole_0,X2,X1)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_598, negated_conjecture, (v1_funct_1(esk11_0)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_424, plain, (v2_funct_2(X1,X2)|~v3_funct_2(X1,k5_numbers,X2)|~v1_funct_1(X1)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_392, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~r2_hidden(X1,X3)|~m2_finseq_1(X3,X2))).#trainneg
cnf(i_0_376, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k4_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_384, plain, (v1_xboole_0(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1114, plain, (v1_finset_1(np__3))).#trainneg
cnf(i_0_1139, negated_conjecture, (k1_card_1(esk12_0)=np__3)).#trainneg
cnf(i_0_485, negated_conjecture, (k3_finseq_1(X1)=np__3|X1!=k10_matrix_3(np__3,esk1_0,esk11_0,X2)|~m2_finseq_1(X1,u1_struct_0(esk1_0))|~m1_matrix_2(X2,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_1097, plain, (~v1_finseq_1(k1_numbers)|~v1_funct_1(k1_numbers)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_654, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_relat_1(esk11_0))).#trainneg
cnf(i_0_544, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_finseq_1(esk11_0))).#trainneg
cnf(i_0_388, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_1281, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,np__3)))=esk12_0|np__1!=X1|np__2!=X2)).#trainneg
cnf(i_0_426, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_matrix_1(esk11_0))).#trainneg
cnf(i_0_1178, negated_conjecture, (k9_xtuple_0(esk11_0)=k2_finseq_1(np__3)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1172, plain, (k4_finseq_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))=k2_finseq_1(np__3))).#trainneg
cnf(i_0_1279, negated_conjecture, (k3_finseq_1(esk11_0)=np__3|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1131, plain, (v3_ordinal1(X1)|~v1_card_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1167, plain, (np__2=np__3|k2_finseq_1(np__2)!=k2_finseq_1(np__3))).#trainneg
cnf(i_0_543, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_funct_1(esk11_0))).#trainneg
cnf(i_0_492, negated_conjecture, (k1_funct_1(X1,X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X1!=k10_matrix_3(np__3,esk1_0,esk11_0,X4)|X3!=k1_funct_1(X4,X2)|~r2_hidden(X2,k4_finseq_1(X1))|~m2_finseq_1(X1,u1_struct_0(esk1_0))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X4,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_389, plain, (v1_finset_1(X1)|~v5_finset_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2)))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_1168, plain, (np__3=np__1|k2_finseq_1(np__3)!=k2_finseq_1(np__1))).#trainneg
cnf(i_0_1100, plain, (~v1_finseq_1(k5_numbers)|~v1_funct_1(k5_numbers)|~v1_relat_1(k5_numbers))).#trainneg
cnf(i_0_1314, negated_conjecture, (v4_relat_1(esk12_0,k5_numbers))).#trainneg
cnf(i_0_1366, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|v1_xboole_0(X3)|~m2_finseq_1(X3,X2)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_390, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_391, plain, (v1_finset_1(X1)|~v1_finset_1(k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_396, plain, (m2_subset_1(k1_card_1(X1),k1_numbers,k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1295, plain, (r2_hidden(X1,k2_finseq_1(np__3))|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1280, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))=esk12_0|np__3!=X3|np__1!=X1|np__2!=X2)).#trainneg
cnf(i_0_1231, plain, (k4_finseq_1(k3_finseq_4(X1,X2,X3,X4))=k2_finseq_1(np__3)|v1_xboole_0(X1)|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1198, plain, (v1_finset_1(k3_finseq_4(X4,X1,X2,X3))|v1_xboole_0(X4)|~m1_subset_1(X3,X4)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4))).#trainneg
cnf(i_0_652, plain, (v7_ordinal1(k1_card_1(X1))|~v1_finset_1(k1_card_1(X1)))).#trainneg
cnf(i_0_402, plain, (v5_relat_1(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_684, negated_conjecture, (esk15_5(X1,esk11_0,X2,X3,X4)=k1_funct_1(esk11_0,X2)|v1_xboole_0(u1_struct_0(esk1_0))|X4!=k3_matrix_1(X1,esk11_0,X2,X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,np__3)|~r1_xxreal_0(X2,np__3)|~m2_finseq_1(esk11_0,k3_finseq_2(X1))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_397, plain, (k3_xcmplx_0(X1,np__1)=X1|~v1_xcmplx_0(X1)|~v1_xcmplx_0(np__1))).#trainneg
cnf(i_0_1326, negated_conjecture, (k1_xboole_0=np__3|esk12_0!=k1_xboole_0|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_1185, plain, (k9_xtuple_0(k1_xboole_0)=k2_finseq_1(k1_xboole_0)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_465, plain, (v3_funct_2(X1,k5_numbers,X2)|~v2_funct_2(X1,X2)|~v2_funct_1(X1)|~v1_funct_1(X1)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_420, plain, (v2_funct_1(X1)|~v3_funct_2(X1,k5_numbers,X2)|~v1_funct_1(X1)|~m2_finseq_1(X1,X2))).#trainneg
cnf(i_0_1304, negated_conjecture, (k1_card_1(esk11_0)=np__3|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1181, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(X1)|~m1_subset_1(k3_finseq_4(X1,esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1))).#trainneg
cnf(i_0_1220, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|v1_xboole_0(X1)|v1_xboole_0(u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(X1,esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1))).#trainneg
cnf(i_0_1324, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(X1)|v1_xboole_0(X2)|k3_finseq_4(X2,esk2_0,esk6_0,esk10_0)!=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk2_0,X1)|~m1_subset_1(esk10_0,X2)|~m1_subset_1(esk6_0,X2)|~m1_subset_1(esk2_0,X2))).#trainneg
cnf(i_0_1271, negated_conjecture, (k9_xtuple_0(esk11_0)=k2_finseq_1(X1)|X1!=np__3|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_993, plain, (m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_405, plain, (v4_relat_1(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1))).#trainneg
cnf(i_0_1115, plain, (k9_xtuple_0(np__3)=k2_finseq_1(np__3)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_1362, plain, (v3_funct_2(k1_xboole_0,X1,X2)|~v2_funct_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|~v3_funct_2(k1_xboole_0,k5_numbers,X2)|~m2_finseq_1(k1_xboole_0,X2))).#trainneg
cnf(i_0_1307, negated_conjecture, (~v1_zfmisc_1(esk11_0)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1313, negated_conjecture, (~v1_xboole_0(esk11_0)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_410, plain, (v1_xboole_0(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_419, plain, (v2_funct_1(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v3_funct_2(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2)|~v1_funct_1(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_1329, negated_conjecture, (k9_xtuple_0(X1)=k2_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|X1!=esk12_0)).#trainneg
cnf(i_0_1282, plain, (m1_subset_1(k2_xcmplx_0(X1,X2),k5_numbers)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1277, negated_conjecture, (k3_finseq_1(esk11_0)=np__3|v1_xboole_0(X1)|v1_xboole_0(u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(X1,esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1))).#trainneg
cnf(i_0_1174, plain, (X1=np__3|k2_finseq_1(np__3)!=k2_finseq_1(X1)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_1302, negated_conjecture, (k1_card_1(esk11_0)=np__3|v1_xboole_0(X1)|v1_xboole_0(u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(X1,esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1))).#trainneg
cnf(i_0_1368, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k2_finseq_1(np__3))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|k3_finseq_1(X2)!=np__3)).#trainneg
cnf(i_0_973, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_423, plain, (v2_funct_2(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2)|~v3_funct_2(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2)|~v1_funct_1(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_1474, plain, (m2_subset_1(k1_xboole_0,k1_numbers,k5_numbers)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_439, plain, (k9_xtuple_0(X1)=k2_finseq_1(esk16_2(X2,X3))|esk16_2(X2,X3)!=k3_finseq_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X2,a_1_0_finseq_1(X3))|~v7_ordinal1(X3))).#trainneg
cnf(i_0_1376, negated_conjecture, (k4_finseq_1(esk12_0)=k2_finseq_1(np__3))).#trainneg
cnf(i_0_1448, negated_conjecture, (k4_finseq_1(esk11_0)=k2_finseq_1(np__3)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1446, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k2_finseq_1(np__3)))).#trainneg
cnf(i_0_489, negated_conjecture, (v1_xboole_0(u1_struct_0(esk1_0))|r2_hidden(k4_tarski(X1,X2),k2_matrix_1(esk11_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_1483, plain, (m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1402, negated_conjecture, (v1_finset_1(esk12_0))).#trainneg
cnf(i_0_449, plain, (v3_ordinal1(X1)|v1_xboole_0(X2)|~v3_ordinal1(k1_zfmisc_1(k2_zfmisc_1(k2_finseq_1(k11_matrix_2(X2)),k2_finseq_1(k11_matrix_2(X2)))))|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2))).#trainneg
cnf(i_0_450, plain, (v1_finset_1(X1)|v1_xboole_0(X2)|~v5_finset_1(k1_zfmisc_1(k2_zfmisc_1(k2_finseq_1(k11_matrix_2(X2)),k2_finseq_1(k11_matrix_2(X2)))))|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2))).#trainneg
cnf(i_0_1590, negated_conjecture, (~v1_zfmisc_1(esk11_0)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1440, negated_conjecture, (k9_xtuple_0(esk11_0)=k2_finseq_1(np__3)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1435, negated_conjecture, (k3_finseq_4(X1,X2,X3,np__3)=esk12_0|v1_xboole_0(X1)|~m1_subset_1(np__3,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|np__1!=X2|np__2!=X3)).#trainneg
cnf(i_0_1456, negated_conjecture, (k3_finseq_1(esk11_0)=np__3|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1679, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1444, plain, (m1_subset_1(k2_finseq_1(np__3),k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_451, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(k2_zfmisc_1(k2_finseq_1(k11_matrix_2(X2)),k2_finseq_1(k11_matrix_2(X2))))|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2))).#trainneg
cnf(i_0_1550, plain, (k4_finseq_1(k1_xboole_0)=k2_finseq_1(k1_xboole_0)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_1392, negated_conjecture, (v3_ordinal1(esk12_0)|~v3_ordinal1(k12_matrix_2(np__3)))).#trainneg
cnf(i_0_1466, plain, (v1_finset_1(X1)|~m2_finseq_1(X1,X2)|~v1_xboole_0(k2_zfmisc_1(k5_numbers,X2)))).#trainneg
cnf(i_0_1467, plain, (v1_finset_1(X1)|~m2_finseq_1(X1,X2)|~m1_subset_1(k2_zfmisc_1(k5_numbers,X2),k5_numbers))).#trainneg
cnf(i_0_452, plain, (v1_finset_1(X1)|v1_xboole_0(X2)|~v1_finset_1(k2_zfmisc_1(k2_finseq_1(k11_matrix_2(X2)),k2_finseq_1(k11_matrix_2(X2))))|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2))).#trainneg
cnf(i_0_453, plain, (m1_subset_1(X1,k2_zfmisc_1(k2_finseq_1(k11_matrix_2(X2)),k2_finseq_1(k11_matrix_2(X2))))|v1_xboole_0(X2)|~r2_hidden(X1,X3)|~v3_matrix_2(X2)|~m1_matrix_2(X3,X2))).#trainneg
cnf(i_0_1539, plain, (v7_ordinal1(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_1379, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|v1_xboole_0(X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(X1,esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk7_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk5_0,X1))).#trainneg
cnf(i_0_1451, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|v1_xboole_0(X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(X1,esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk7_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk5_0,X1))).#trainneg
cnf(i_0_1559, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|v1_xboole_0(X1)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(X1,esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1520, negated_conjecture, (v1_finset_1(esk11_0)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1387, negated_conjecture, (k3_finseq_4(X1,esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|v1_xboole_0(X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(X1,esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk7_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk5_0,X1))).#trainneg
cnf(i_0_1680, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_1442, plain, (v3_ordinal1(k2_finseq_1(np__3))|~v3_ordinal1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1384, negated_conjecture, (v4_relat_1(esk11_0,k5_numbers))).#trainneg
cnf(i_0_456, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|~v1_xboole_0(k2_finseq_1(k11_matrix_2(X2)))|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2))).#trainneg
cnf(i_0_1582, plain, (k3_finseq_1(np__3)=np__3|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_1316, plain, (v4_relat_1(k3_finseq_4(X1,X2,X3,X4),k5_numbers)|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1416, plain, (v3_ordinal1(X1)|~m2_finseq_1(X1,X2)|~v1_card_1(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,X2))))).#trainneg
cnf(i_0_457, plain, (v4_relat_1(X1,k2_finseq_1(k11_matrix_2(X2)))|v1_xboole_0(X2)|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2))).#trainneg
cnf(i_0_458, plain, (v2_funct_1(X1)|v1_xboole_0(X2)|~v3_funct_2(X1,k2_finseq_1(k11_matrix_2(X2)),k2_finseq_1(k11_matrix_2(X2)))|~v1_funct_1(X1)|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2))).#trainneg
cnf(i_0_1557, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_subset_1(k3_finseq_4(X2,esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(X1,esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1)|~m1_subset_1(esk7_0,X2)|~m1_subset_1(esk6_0,X2)|~m1_subset_1(esk5_0,X2))).#trainneg
cnf(i_0_1562, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|v1_xboole_0(X1)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(X1,esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1540, negated_conjecture, (esk15_5(u1_struct_0(esk1_0),esk11_0,X1,X2,X3)=k1_funct_1(esk11_0,X1)|X3!=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1564, negated_conjecture, (v1_xboole_0(X1)|v1_xboole_0(X2)|k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk2_0)),k1_tarski(k4_tarski(np__1,esk6_0))),k1_tarski(k4_tarski(np__1,esk10_0)))!=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|~m1_subset_1(esk10_0,X2)|~m1_subset_1(esk6_0,X2)|~m1_subset_1(esk2_0,X2)|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk2_0,X1))).#trainneg
cnf(i_0_1352, plain, (k9_xtuple_0(X1)=k2_finseq_1(np__3)|k1_card_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_finseq_1(np__3)|~v1_funct_1(X1)|~v1_funct_1(np__3)|~v1_relat_1(X1)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_1223, negated_conjecture, (k3_finseq_4(X1,esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|v1_xboole_0(X1)|v1_xboole_0(u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(X1,esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1))).#trainneg
cnf(i_0_1710, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1855, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_1878, negated_conjecture, (v3_ordinal1(esk12_0)|~v1_card_1(k12_matrix_2(np__3)))).#trainneg
cnf(i_0_1951, negated_conjecture, (v2_funct_1(esk12_0)|v1_xboole_0(k12_matrix_2(np__3))|~v3_matrix_2(k12_matrix_2(np__3)))).#trainneg
cnf(i_0_460, plain, (v5_relat_1(X1,k2_finseq_1(k11_matrix_2(X2)))|v1_xboole_0(X2)|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2))).#trainneg
cnf(i_0_461, plain, (v2_funct_2(X1,k2_finseq_1(k11_matrix_2(X2)))|v1_xboole_0(X2)|~v3_funct_2(X1,k2_finseq_1(k11_matrix_2(X2)),k2_finseq_1(k11_matrix_2(X2)))|~v1_funct_1(X1)|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2))).#trainneg
cnf(i_0_1850, negated_conjecture, (k3_finseq_1(esk11_0)=np__3|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1704, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|v1_xboole_0(X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(X1,esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,X1)|~m1_subset_1(esk3_0,X1)|~m1_subset_1(esk2_0,X1))).#trainneg
cnf(i_0_1708, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|v1_xboole_0(X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(X1,esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,X1)|~m1_subset_1(esk3_0,X1)|~m1_subset_1(esk2_0,X1))).#trainneg
cnf(i_0_1914, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|v1_xboole_0(X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(X1,esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk7_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk5_0,X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1910, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|v1_xboole_0(X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(X1,esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk7_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk5_0,X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1960, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|v1_xboole_0(X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(X1,esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1912, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|v1_xboole_0(X1)|v1_xboole_0(X2)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(X2,esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(X1,esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk7_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk5_0,X1)|~m1_subset_1(esk4_0,X2)|~m1_subset_1(esk3_0,X2)|~m1_subset_1(esk2_0,X2))).#trainneg
cnf(i_0_1877, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k2_finseq_1(k1_xboole_0))|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_1580, plain, (k3_finseq_1(np__3)=np__2|k2_finseq_1(np__2)!=k2_finseq_1(np__3)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_463, plain, (v3_funct_2(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2)|~v2_funct_2(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2)|~v2_funct_1(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_funct_1(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))))).#trainneg
cnf(i_0_1441, plain, (v1_relat_1(k2_finseq_1(np__3))|~v4_funct_1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1443, plain, (v1_finset_1(k2_finseq_1(np__3))|~v5_finset_1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1881, plain, (v1_finset_1(X1)|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_466, plain, (m1_subset_1(k3_matrix_1(X1,X2,X3,X4),X1)|~v1_matrix_1(X2)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m2_finseq_1(X2,k3_finseq_2(X1)))).#trainneg
cnf(i_0_1918, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|v1_xboole_0(X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(X1,esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1729, negated_conjecture, (k1_card_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))=np__3|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_1791, plain, (v1_xboole_0(X1)|r2_hidden(np__2,k2_finseq_1(np__3))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1743, plain, (v1_xboole_0(X1)|r2_hidden(np__1,k2_finseq_1(np__3))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1958, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|v1_xboole_0(X1)|v1_xboole_0(X2)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(X2,esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(X1,esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1)|~m1_subset_1(esk4_0,X2)|~m1_subset_1(esk3_0,X2)|~m1_subset_1(esk2_0,X2))).#trainneg
cnf(i_0_1536, plain, (v7_ordinal1(k1_card_1(X1))|~v1_xboole_0(k1_card_1(X1)))).#trainneg
cnf(i_0_1581, plain, (k3_finseq_1(np__3)=np__1|k2_finseq_1(np__3)!=k2_finseq_1(np__1)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_486, plain, (X1=k3_matrix_1(X2,X3,X4,X5)|v1_xboole_0(k2_matrix_1(X3))|X1!=k1_funct_1(X6,X5)|X6!=k1_funct_1(X3,X4)|~v1_matrix_1(X3)|~m2_finseq_1(X3,k3_finseq_2(X2))|~m2_finseq_1(X6,X2)|~v7_ordinal1(X5)|~v7_ordinal1(X4)|~m1_subset_1(X1,X2)|~m1_subset_1(k4_tarski(X4,X5),k2_matrix_1(X3)))).#trainneg
cnf(i_0_1585, plain, (k4_finseq_1(np__3)=k2_finseq_1(np__3)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_487, plain, (esk15_5(X1,X2,X3,np__3,X4)=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X5)),k1_tarski(k4_tarski(np__1,X6))),k1_tarski(k4_tarski(np__1,X4)))|k1_funct_1(esk15_5(X1,X2,X3,np__3,X4),np__1)!=X5|k1_funct_1(esk15_5(X1,X2,X3,np__3,X4),np__2)!=X6|k3_finseq_1(esk15_5(X1,X2,X3,np__3,X4))!=np__3|~v1_finseq_1(esk15_5(X1,X2,X3,np__3,X4))|~v1_funct_1(esk15_5(X1,X2,X3,np__3,X4))|~v1_relat_1(esk15_5(X1,X2,X3,np__3,X4))|X4!=k3_matrix_1(X1,X2,X3,np__3)|~v1_matrix_1(X2)|~r2_hidden(k4_tarski(X3,np__3),k2_matrix_1(X2))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(np__3)|~v7_ordinal1(X3)|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_490, plain, (m2_finseq_1(esk15_5(X1,X2,X3,X4,X5),X1)|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(X5,X1)|~m1_subset_1(k4_tarski(X3,X4),k2_matrix_1(X2)))).#trainneg
cnf(i_0_1700, negated_conjecture, (k3_finseq_4(X1,esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|v1_xboole_0(X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(X1,esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk4_0,X1)|~m1_subset_1(esk3_0,X1)|~m1_subset_1(esk2_0,X1))).#trainneg
cnf(i_0_1956, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_subset_1(k3_finseq_4(X1,esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(X2,esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk10_0,X2)|~m1_subset_1(esk9_0,X2)|~m1_subset_1(esk8_0,X2)|~m1_subset_1(esk7_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk5_0,X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1816, plain, (v1_xboole_0(X1)|r2_hidden(np__3,k4_finseq_1(k3_finseq_4(X1,X2,X3,X4)))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4)))).#trainneg
cnf(i_0_1728, negated_conjecture, (k9_xtuple_0(k10_matrix_3(np__3,esk1_0,esk11_0,X1))=k2_finseq_1(X2)|X2!=np__3|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~m1_subset_1(X2,k5_numbers)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_1926, negated_conjecture, (k3_finseq_4(X1,esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|v1_xboole_0(X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(X1,esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk7_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk5_0,X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1787, negated_conjecture, (r2_hidden(np__2,k9_xtuple_0(esk11_0))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1739, negated_conjecture, (r2_hidden(np__1,k9_xtuple_0(esk11_0))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1807, negated_conjecture, (r2_hidden(np__3,k9_xtuple_0(esk11_0))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1596, negated_conjecture, (k2_finseq_1(k1_xboole_0)=k2_finseq_1(np__3)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|esk12_0!=k1_xboole_0)).#trainneg
cnf(i_0_1996, negated_conjecture, (v2_funct_1(esk12_0)|v1_xboole_0(k12_matrix_2(np__3)))).#trainneg
cnf(i_0_2032, negated_conjecture, (~v1_xboole_0(k2_finseq_1(np__3)))).#trainneg
cnf(i_0_524, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k3_finseq_4(X1,esk8_0,esk9_0,esk10_0))=esk11_0|v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(X1)|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1))).#trainneg
cnf(i_0_1986, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2196, negated_conjecture, (r2_hidden(np__3,k4_finseq_1(esk11_0))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2068, negated_conjecture, (~v1_xboole_0(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_2062, negated_conjecture, (v1_finset_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_2065, negated_conjecture, (k9_xtuple_0(k10_matrix_3(np__3,esk1_0,esk11_0,X1))=k2_finseq_1(np__3)|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_1796, plain, (v1_xboole_0(X1)|r2_hidden(np__2,k4_finseq_1(k3_finseq_4(X1,X2,X3,X4)))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4)))).#trainneg
cnf(i_0_2030, negated_conjecture, (m1_subset_1(np__3,k2_finseq_1(np__3)))).#trainneg
cnf(i_0_525, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k3_finseq_4(X1,esk5_0,esk6_0,esk7_0),k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0))=esk11_0|v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(X1)|~m1_subset_1(esk7_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk5_0,X1))).#trainneg
cnf(i_0_2042, plain, (m1_subset_1(X1,k5_numbers)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,k1_xboole_0)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_526, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k3_finseq_4(X1,esk2_0,esk3_0,esk4_0),k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0))=esk11_0|v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(X1)|~m1_subset_1(esk4_0,X1)|~m1_subset_1(esk3_0,X1)|~m1_subset_1(esk2_0,X1))).#trainneg
cnf(i_0_1720, negated_conjecture, (k4_finseq_1(esk11_0)=k2_finseq_1(np__3)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2024, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|v1_xboole_0(X1)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0))|~m1_subset_1(esk7_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk5_0,X1)|~m2_finseq_1(k3_finseq_4(X1,esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1999, plain, (v3_funct_2(X1,k5_numbers,k2_finseq_1(k11_matrix_2(X2)))|v1_xboole_0(X2)|~v2_funct_1(X1)|~m2_finseq_1(X1,k2_finseq_1(k11_matrix_2(X2)))|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2))).#trainneg
cnf(i_0_2245, negated_conjecture, (m1_subset_1(np__1,k2_finseq_1(np__3)))).#trainneg
cnf(i_0_1748, plain, (v1_xboole_0(X1)|r2_hidden(np__1,k4_finseq_1(k3_finseq_4(X1,X2,X3,X4)))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4)))).#trainneg
cnf(i_0_2003, negated_conjecture, (k3_finseq_1(esk11_0)=np__3|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2236, negated_conjecture, (k2_finseq_1(X2)=k4_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|X2!=np__3|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1175, plain, (k1_nat_1(X1,X2)=np__3|k2_finseq_1(k1_nat_1(X1,X2))!=k2_finseq_1(np__3)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2111, negated_conjecture, (m1_subset_1(np__2,k2_finseq_1(np__3)))).#trainneg
cnf(i_0_1544, plain, (k3_finseq_1(k1_xboole_0)=np__2|k2_finseq_1(k1_xboole_0)!=k2_finseq_1(np__2)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_2367, negated_conjecture, (~v1_xboole_0(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)))).#trainneg
cnf(i_0_2397, negated_conjecture, (v1_relat_1(np__3)|~v4_funct_1(k2_finseq_1(np__3)))).#trainneg
cnf(i_0_556, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk17_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,esk17_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_2368, negated_conjecture, (~v1_xboole_0(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_2508, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X1)|X1!=np__1|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2507, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),np__2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,X1)|X1!=np__2|~v7_ordinal1(X1))).#trainneg
cnf(i_0_2513, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X1,X2,X3,X4)),np__2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,X3)|v1_xboole_0(X1)|~v7_ordinal1(X3)|~m1_matrix_2(k3_finseq_4(X1,X2,X3,X4),k12_matrix_2(np__3))|~m1_subset_1(k3_finseq_4(X1,X2,X3,X4),k12_matrix_2(np__3))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_2344, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk8_0)),k1_tarski(k4_tarski(np__1,esk9_0))),k1_tarski(k4_tarski(np__1,esk10_0))))=esk11_0|v1_xboole_0(X1)|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1))).#trainneg
cnf(i_0_2434, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk5_0)),k1_tarski(k4_tarski(np__1,esk6_0))),k1_tarski(k4_tarski(np__1,esk7_0))),k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0))=esk11_0|v1_xboole_0(X1)|~m1_subset_1(esk7_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk5_0,X1))).#trainneg
cnf(i_0_1001, plain, (k2_zfmisc_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1545, plain, (k3_finseq_1(k1_xboole_0)=np__1|k2_finseq_1(k1_xboole_0)!=k2_finseq_1(np__1)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_558, plain, (m1_subset_1(X1,a_1_0_finseq_1(X2))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1478, plain, (m2_subset_1(k1_xboole_0,k1_numbers,k5_numbers)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2398, negated_conjecture, (v1_funct_1(np__3)|~v4_funct_1(k2_finseq_1(np__3)))).#trainneg
cnf(i_0_2526, negated_conjecture, (v1_relat_1(np__1)|~v4_funct_1(k2_finseq_1(np__3)))).#trainneg
cnf(i_0_2374, negated_conjecture, (k4_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))=k2_finseq_1(np__3)|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_2486, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk2_0)),k1_tarski(k4_tarski(np__1,esk3_0))),k1_tarski(k4_tarski(np__1,esk4_0))),k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0))=esk11_0|v1_xboole_0(X1)|~m1_subset_1(esk4_0,X1)|~m1_subset_1(esk3_0,X1)|~m1_subset_1(esk2_0,X1))).#trainneg
cnf(i_0_2436, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k3_finseq_4(X2,esk2_0,esk3_0,esk4_0),k3_finseq_4(X1,esk5_0,esk6_0,esk7_0),k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0))=esk11_0|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_subset_1(esk7_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk5_0,X1)|~m1_subset_1(esk4_0,X2)|~m1_subset_1(esk3_0,X2)|~m1_subset_1(esk2_0,X2))).#trainneg
cnf(i_0_2386, negated_conjecture, (r2_hidden(np__2,k4_finseq_1(esk11_0))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk2_0,esk3_0,esk4_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1671, plain, (v4_relat_1(esk17_1(k1_zfmisc_1(k1_xboole_0)),X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_2601, negated_conjecture, (k3_finseq_1(k1_xboole_0)=np__2|k2_finseq_1(np__2)!=k2_finseq_1(np__3)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|esk12_0!=k1_xboole_0)).#trainneg
cnf(i_0_2527, negated_conjecture, (v1_funct_1(np__1)|~v4_funct_1(k2_finseq_1(np__3)))).#trainneg
cnf(i_0_2529, negated_conjecture, (v1_finset_1(np__1)|~v5_finset_1(k2_finseq_1(np__3)))).#trainneg
cnf(i_0_2786, plain, (k3_finseq_4(X1,k1_funct_1(k3_finseq_4(X2,X3,X4,X5),np__1),X6,X5)=k3_finseq_4(X2,X3,X4,X5)|v1_xboole_0(X2)|v1_xboole_0(X1)|k1_funct_1(k3_finseq_4(X2,X3,X4,X5),np__2)!=X6|~m1_subset_1(X5,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(k1_funct_1(k3_finseq_4(X2,X3,X4,X5),np__1),X1)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_2753, plain, (v5_relat_1(esk17_1(k1_zfmisc_1(k1_xboole_0)),X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2781, negated_conjecture, (m2_subset_1(k1_xboole_0,k1_numbers,k5_numbers)|~v1_xboole_0(esk11_0))).#trainneg
cnf(i_0_2772, negated_conjecture, (k3_finseq_1(k1_xboole_0)=np__1|k2_finseq_1(np__3)!=k2_finseq_1(np__1)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|esk12_0!=k1_xboole_0)).#trainneg
cnf(i_0_2597, negated_conjecture, (v1_relat_1(np__2)|~v4_funct_1(k2_finseq_1(np__3)))).#trainneg
cnf(i_0_2598, negated_conjecture, (v1_funct_1(np__2)|~v4_funct_1(k2_finseq_1(np__3)))).#trainneg
cnf(i_0_614, plain, (esk16_2(X1,X2)=X1|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_2819, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk2_0)),k1_tarski(k4_tarski(np__1,esk3_0))),k1_tarski(k4_tarski(np__1,esk4_0))),k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0))=esk11_0)).#trainneg
cnf(i_0_1503, negated_conjecture, (k3_finseq_4(X1,X2,X3,X4)=esk12_0|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|np__3!=X4|np__1!=X2|np__2!=X3)).#trainneg
cnf(i_0_947, plain, (m1_subset_1(k2_xcmplx_0(X2,X1),k5_numbers)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_2899, negated_conjecture, (k9_xtuple_0(k1_funct_1(esk11_0,np__1))=k2_finseq_1(np__3))).#trainneg
cnf(i_0_2895, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__1),np__1)=esk2_0)).#trainneg
cnf(i_0_2897, negated_conjecture, (k4_finseq_1(k1_funct_1(esk11_0,np__1))=k2_finseq_1(np__3))).#trainneg
cnf(i_0_2911, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__1),np__2)=esk3_0)).#trainneg
cnf(i_0_2748, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1681, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2998, negated_conjecture, (esk16_2(np__2,np__3)=np__2)).#trainneg
cnf(i_0_1307, negated_conjecture, (~v1_zfmisc_1(esk11_0)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1546, plain, (k3_finseq_1(k1_xboole_0)=np__3|k2_finseq_1(k1_xboole_0)!=k2_finseq_1(np__3)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_2892, negated_conjecture, (~v1_xboole_0(k1_funct_1(esk11_0,np__1)))).#trainneg
cnf(i_0_2893, negated_conjecture, (~v1_zfmisc_1(k1_funct_1(esk11_0,np__1)))).#trainneg
cnf(i_0_633, plain, (v3_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v2_funct_2(k2_zfmisc_1(X1,X2),X2)|~v2_funct_1(k2_zfmisc_1(X1,X2))|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_4, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_funct_1(esk11_0,np__1),k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0))=esk11_0)).#trainneg
cnf(i_0_2888, negated_conjecture, (v1_finseq_1(k1_funct_1(esk11_0,np__1)))).#trainneg
cnf(i_0_2891, negated_conjecture, (v3_card_1(k1_funct_1(esk11_0,np__1),np__3))).#trainneg
cnf(i_0_2887, negated_conjecture, (v1_finset_1(k1_funct_1(esk11_0,np__1)))).#trainneg
cnf(i_0_2770, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2999, negated_conjecture, (esk16_2(np__1,np__3)=np__1)).#trainneg
cnf(i_0_601, negated_conjecture, (~v1_xboole_0(esk11_0)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_634, plain, (v2_funct_1(k2_zfmisc_1(X1,X2))|~v3_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_2600, negated_conjecture, (v1_finset_1(np__2)|~v5_finset_1(k2_finseq_1(np__3)))).#trainneg
cnf(i_0_637, plain, (v2_funct_2(k2_zfmisc_1(X1,X2),X2)|~v3_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v1_funct_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_642, plain, (k3_finseq_1(X1)=np__2|k4_finseq_1(X1)!=k2_finseq_1(np__2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2915, negated_conjecture, (v1_xboole_0(X1)|m2_finseq_1(k1_funct_1(esk11_0,np__1),X1)|~m1_subset_1(esk4_0,X1)|~m1_subset_1(esk3_0,X1)|~m1_subset_1(esk2_0,X1))).#trainneg
cnf(i_0_2917, negated_conjecture, (k3_finseq_4(X1,X2,X3,esk4_0)=k1_funct_1(esk11_0,np__1)|v1_xboole_0(X1)|k1_funct_1(k1_funct_1(esk11_0,np__1),np__1)!=X2|k1_funct_1(k1_funct_1(esk11_0,np__1),np__2)!=X3|~m1_subset_1(esk4_0,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_3012, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk2_0)),k1_tarski(k4_tarski(np__1,esk3_0))),k1_tarski(k4_tarski(np__1,esk4_0)))=k1_funct_1(esk11_0,np__1)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk2_0)),k1_tarski(k4_tarski(np__1,esk3_0))),k1_tarski(k4_tarski(np__1,esk4_0))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2889, negated_conjecture, (v1_funct_1(k1_funct_1(esk11_0,np__1)))).#trainneg
cnf(i_0_2890, negated_conjecture, (v1_relat_1(k1_funct_1(esk11_0,np__1)))).#trainneg
cnf(i_0_3013, negated_conjecture, (m2_finseq_1(esk11_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk2_0)),k1_tarski(k4_tarski(np__1,esk3_0))),k1_tarski(k4_tarski(np__1,esk4_0))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_2894, negated_conjecture, (k3_finseq_4(X1,esk2_0,esk3_0,esk4_0)=k1_funct_1(esk11_0,np__1)|v1_xboole_0(X1)|~m1_subset_1(esk4_0,X1)|~m1_subset_1(esk3_0,X1)|~m1_subset_1(esk2_0,X1))).#trainneg
cnf(i_0_3041, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk2_0)),k1_tarski(k4_tarski(np__1,esk3_0))),k1_tarski(k4_tarski(np__1,esk4_0))),k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k3_finseq_4(X1,esk8_0,esk9_0,esk10_0))=esk11_0|v1_xboole_0(X1)|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1))).#trainneg
cnf(i_0_2876, negated_conjecture, (k3_finseq_4(X1,np__1,X2,np__3)=esk12_0|v1_xboole_0(X1)|np__2!=X2|~m1_subset_1(np__3,X1)|~m1_subset_1(X2,X1)|~m1_subset_1(np__1,X1))).#trainneg
cnf(i_0_3222, negated_conjecture, (v1_finseq_1(np__1)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_573, negated_conjecture, (~v1_zfmisc_1(esk11_0)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3169, negated_conjecture, (~v1_xboole_0(np__1)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_3170, negated_conjecture, (~v1_zfmisc_1(np__1)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_658, plain, (k9_xtuple_0(X1)=k2_finseq_1(k3_finseq_1(X2))|k3_finseq_1(X2)!=k1_card_1(X1)|~v1_finseq_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2))).#trainneg
cnf(i_0_3167, negated_conjecture, (~v1_zfmisc_1(esk11_0)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3017, negated_conjecture, (k9_xtuple_0(esk11_0)=k2_finseq_1(np__3)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk2_0)),k1_tarski(k4_tarski(np__1,esk3_0))),k1_tarski(k4_tarski(np__1,esk4_0))),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3331, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk2_0)),k1_tarski(k4_tarski(np__1,esk3_0))),k1_tarski(k4_tarski(np__1,esk4_0)))=k1_funct_1(esk11_0,np__1)|v1_xboole_0(X1)|~m1_subset_1(esk4_0,X1)|~m1_subset_1(esk3_0,X1)|~m1_subset_1(esk2_0,X1))).#trainneg
cnf(i_0_1859, plain, (v1_funct_1(k2_finseq_1(np__3))|~v4_funct_1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_3224, negated_conjecture, (v1_finset_1(np__1)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1085, negated_conjecture, (~v1_zfmisc_1(esk11_0)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_661, plain, (k3_finseq_1(X1)=np__3|k4_finseq_1(X1)!=k2_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3313, negated_conjecture, (v1_funct_1(np__1)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_3314, negated_conjecture, (v1_relat_1(np__1)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1462, plain, (v1_xboole_0(k3_finseq_4(X1,X2,X3,X4))|m1_subset_1(X5,k2_zfmisc_1(k5_numbers,X1))|v1_xboole_0(X1)|~m1_subset_1(X5,k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_3174, plain, (v3_funct_2(k2_zfmisc_1(X1,X2),X1,X2)|~v2_funct_1(k2_zfmisc_1(X1,X2))|~v1_funct_1(k2_zfmisc_1(X1,X2))|~v3_funct_2(k2_zfmisc_1(X1,X2),k5_numbers,X2)|~m2_finseq_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_668, plain, (k1_card_1(X1)=k3_finseq_1(X2)|k9_xtuple_0(X2)!=k2_finseq_1(k1_card_1(X1))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_669, plain, (k3_finseq_1(X1)=k3_finseq_1(X2)|k4_finseq_1(X2)!=k2_finseq_1(k3_finseq_1(X1))|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3404, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk2_0)),k1_tarski(k4_tarski(np__1,esk3_0))),k1_tarski(k4_tarski(np__1,esk4_0))),k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk8_0)),k1_tarski(k4_tarski(np__1,esk9_0))),k1_tarski(k4_tarski(np__1,esk10_0))))=esk11_0|v1_xboole_0(X1)|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1))).#trainneg
cnf(i_0_1178, negated_conjecture, (k9_xtuple_0(esk11_0)=k2_finseq_1(np__3)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_437, negated_conjecture, (m2_finseq_1(esk11_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_976, plain, (k9_xtuple_0(X1)=k2_finseq_1(k2_xcmplx_0(X2,X3))|k2_xcmplx_0(X2,X3)!=k1_card_1(X1)|~v7_ordinal1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_1874, plain, (m1_subset_1(k2_finseq_1(k1_xboole_0),k1_zfmisc_1(k5_numbers))|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_3134, negated_conjecture, (k9_xtuple_0(np__1)=k2_finseq_1(np__3)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1313, negated_conjecture, (~v1_xboole_0(esk11_0)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3223, negated_conjecture, (v3_card_1(np__1,np__3)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_3058, negated_conjecture, (k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=esk12_0|esk10_0!=np__3|esk2_0!=np__1|esk6_0!=np__2)).#trainneg
cnf(i_0_1463, plain, (v1_xboole_0(k3_finseq_4(X1,X2,X3,X4))|m1_subset_1(X5,k2_zfmisc_1(k5_numbers,X6))|v1_xboole_0(X1)|v1_xboole_0(X6)|~m1_subset_1(X5,k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X4,X6)|~m1_subset_1(X3,X6)|~m1_subset_1(X2,X6)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_698, plain, (k2_xcmplx_0(X1,X2)=k3_finseq_1(X3)|k9_xtuple_0(X3)!=k2_finseq_1(k2_xcmplx_0(X1,X2))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2921, negated_conjecture, (k1_funct_1(X1,np__1)=esk2_0|X1!=k1_funct_1(esk11_0,np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2511, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X1)|~v7_ordinal1(X1)|~m1_matrix_2(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),k12_matrix_2(np__3))|~m1_subset_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),k12_matrix_2(np__3)))).#trainneg
cnf(i_0_3577, negated_conjecture, (k9_xtuple_0(k1_funct_1(esk11_0,np__3))=k2_finseq_1(np__3))).#trainneg
cnf(i_0_3585, negated_conjecture, (r2_hidden(np__3,k4_finseq_1(k1_funct_1(esk11_0,np__3))))).#trainneg
cnf(i_0_3583, negated_conjecture, (r2_hidden(np__1,k4_finseq_1(k1_funct_1(esk11_0,np__3))))).#trainneg
cnf(i_0_3575, negated_conjecture, (k4_finseq_1(k1_funct_1(esk11_0,np__3))=k2_finseq_1(np__3))).#trainneg
cnf(i_0_2819, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk2_0)),k1_tarski(k4_tarski(np__1,esk3_0))),k1_tarski(k4_tarski(np__1,esk4_0))),k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_funct_1(esk11_0,np__3))=esk11_0)).#trainneg
cnf(i_0_4, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_funct_1(esk11_0,np__1),k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_funct_1(esk11_0,np__3))=esk11_0)).#trainneg
cnf(i_0_3146, negated_conjecture, (k4_finseq_1(np__1)=k2_finseq_1(np__3)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_573, negated_conjecture, (~v1_zfmisc_1(esk11_0)|~m1_subset_1(k1_funct_1(esk11_0,np__3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_712, plain, (v1_xboole_0(X1)|m2_finseq_1(X4,X1)|~m1_subset_1(k1_funct_1(X4,np__3),X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|k1_funct_1(X4,np__1)!=X2|k1_funct_1(X4,np__2)!=X3|k3_finseq_1(X4)!=np__3|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4))).#trainneg
cnf(i_0_3569, negated_conjecture, (~v1_xboole_0(k1_funct_1(esk11_0,np__3)))).#trainneg
cnf(i_0_3570, negated_conjecture, (~v1_zfmisc_1(k1_funct_1(esk11_0,np__3)))).#trainneg
cnf(i_0_720, plain, (k1_nat_1(X1,X2)=k3_finseq_1(X3)|k9_xtuple_0(X3)!=k2_finseq_1(k1_nat_1(X1,X2))|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3572, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__3),np__1)=esk8_0)).#trainneg
cnf(i_0_3607, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk2_0)),k1_tarski(k4_tarski(np__1,esk3_0))),k1_tarski(k4_tarski(np__1,esk4_0))),k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk8_0)),k1_tarski(k4_tarski(np__1,esk9_0))),k1_tarski(k4_tarski(np__1,esk10_0))))=esk11_0)).#trainneg
cnf(i_0_3494, negated_conjecture, (k3_finseq_1(k1_funct_1(esk11_0,np__1))=np__3)).#trainneg
cnf(i_0_3568, negated_conjecture, (v3_card_1(k1_funct_1(esk11_0,np__3),np__3))).#trainneg
cnf(i_0_3565, negated_conjecture, (v1_finseq_1(k1_funct_1(esk11_0,np__3)))).#trainneg
cnf(i_0_3484, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk2_0)),k1_tarski(k4_tarski(np__1,esk3_0))),k1_tarski(k4_tarski(np__1,esk4_0)))=k1_funct_1(esk11_0,np__1))).#trainneg
cnf(i_0_3607, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_funct_1(esk11_0,np__1),k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk8_0)),k1_tarski(k4_tarski(np__1,esk9_0))),k1_tarski(k4_tarski(np__1,esk10_0))))=esk11_0)).#trainneg
cnf(i_0_3974, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__1),np__3)=esk4_0)).#trainneg
cnf(i_0_1076, plain, (esk17_1(k1_zfmisc_1(X1))=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3963, negated_conjecture, (v1_finseq_1(np__3)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1313, negated_conjecture, (~v1_xboole_0(esk11_0)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_730, plain, (k9_xtuple_0(X1)=k2_finseq_1(k1_nat_1(X2,X3))|k1_nat_1(X2,X3)!=k1_card_1(X1)|~v7_ordinal1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_3991, negated_conjecture, (k3_finseq_1(X1)=np__3|X1!=k1_funct_1(esk11_0,np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3591, negated_conjecture, (k3_finseq_1(k1_funct_1(esk11_0,np__3))=np__3)).#trainneg
cnf(i_0_3988, negated_conjecture, (k1_funct_1(X1,np__2)=esk3_0|X1!=k1_funct_1(esk11_0,np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3989, negated_conjecture, (k1_funct_1(X1,np__3)=esk4_0|X1!=k1_funct_1(esk11_0,np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3874, plain, (k2_xcmplx_0(X1,X2)=k3_finseq_1(X3)|k9_xtuple_0(X3)!=k2_finseq_1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3567, negated_conjecture, (v1_relat_1(k1_funct_1(esk11_0,np__3)))).#trainneg
cnf(i_0_3881, plain, (k1_nat_1(X1,X2)=k3_finseq_1(k3_finseq_4(X3,X4,X5,X6))|v1_xboole_0(X3)|k2_finseq_1(k1_nat_1(X1,X2))!=k2_finseq_1(np__3)|~v7_ordinal1(X2)|~v1_finseq_1(k3_finseq_4(X3,X4,X5,X6))|~v1_funct_1(k3_finseq_4(X3,X4,X5,X6))|~v1_relat_1(k3_finseq_4(X3,X4,X5,X6))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X6,X3)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_1685, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_3960, negated_conjecture, (k3_finseq_1(np__1)=np__3|esk11_0!=esk12_0)).#trainneg
cnf(i_0_3868, negated_conjecture, (~v1_zfmisc_1(np__3)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_3886, plain, (k1_nat_1(X1,X2)=k3_finseq_1(X3)|k4_finseq_1(X3)!=k2_finseq_1(k1_nat_1(X1,X2))|~v7_ordinal1(X2)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3566, negated_conjecture, (v1_funct_1(k1_funct_1(esk11_0,np__3)))).#trainneg
cnf(i_0_3594, negated_conjecture, (v1_xboole_0(X1)|m2_finseq_1(k1_funct_1(esk11_0,np__3),X1)|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1))).#trainneg
cnf(i_0_3579, negated_conjecture, (k1_card_1(k1_funct_1(esk11_0,np__3))=np__3)).#trainneg
cnf(i_0_3564, negated_conjecture, (v1_finset_1(k1_funct_1(esk11_0,np__3)))).#trainneg
cnf(i_0_2766, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_3711, negated_conjecture, (k9_xtuple_0(np__3)=k2_finseq_1(np__3)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1085, negated_conjecture, (~v1_zfmisc_1(esk11_0)|~m2_finseq_1(k1_funct_1(esk11_0,np__3),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3961, negated_conjecture, (v3_card_1(np__3,np__3)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_752, plain, (k1_nat_1(X1,X2)=k3_finseq_1(X3)|k4_finseq_1(X3)!=k2_finseq_1(k1_nat_1(X1,X2))|~v7_ordinal1(X1)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3596, negated_conjecture, (k3_finseq_4(X1,X2,X3,esk10_0)=k1_funct_1(esk11_0,np__3)|v1_xboole_0(X1)|k1_funct_1(k1_funct_1(esk11_0,np__3),np__1)!=X2|k1_funct_1(k1_funct_1(esk11_0,np__3),np__2)!=X3|~m1_subset_1(esk10_0,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_3597, negated_conjecture, (k1_funct_1(X1,np__2)=esk9_0|X1!=k1_funct_1(esk11_0,np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1872, plain, (v3_ordinal1(k2_finseq_1(k1_xboole_0))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_3735, negated_conjecture, (k4_finseq_1(np__3)=k2_finseq_1(np__3)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1456, negated_conjecture, (k3_finseq_1(esk11_0)=np__3|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_757, plain, (k3_finseq_1(X1)=k3_finseq_1(X2)|k2_finseq_1(k3_finseq_1(X1))!=k2_finseq_1(np__3)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|k3_finseq_1(X2)!=np__3)).#trainneg
cnf(i_0_759, plain, (k1_nat_1(X1,X2)=k3_finseq_1(X3)|k2_finseq_1(k1_nat_1(X1,X2))!=k2_finseq_1(np__3)|~v7_ordinal1(X1)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m1_subset_1(X2,k5_numbers)|k3_finseq_1(X3)!=np__3)).#trainneg
cnf(i_0_764, plain, (k4_finseq_1(X1)=k2_finseq_1(np__2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__2)).#trainneg
cnf(i_0_4043, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))=k1_funct_1(esk11_0,np__1)|esk4_0!=X3|esk2_0!=X1|esk3_0!=X2)).#trainneg
cnf(i_0_4046, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,esk4_0)))=k1_funct_1(esk11_0,np__1)|esk2_0!=X1|esk3_0!=X2)).#trainneg
cnf(i_0_4246, negated_conjecture, (~v1_zfmisc_1(esk11_0)|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4222, negated_conjecture, (k9_xtuple_0(X1)=k2_finseq_1(np__3)|k1_card_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4161, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_4230, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_4147, negated_conjecture, (v1_relat_1(np__3)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_856, negated_conjecture, (k3_finseq_1(esk11_0)=np__3|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_765, plain, (k3_finseq_1(X1)=k3_finseq_1(X2)|k2_finseq_1(k3_finseq_1(X1))!=k2_finseq_1(np__2)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|k3_finseq_1(X2)!=np__2)).#trainneg
cnf(i_0_4190, negated_conjecture, (v1_funct_1(np__3)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_4261, negated_conjecture, (k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,np__3)|esk8_0!=esk2_0|esk9_0!=esk6_0)).#trainneg
cnf(i_0_767, plain, (k1_nat_1(X1,X2)=k3_finseq_1(X3)|k2_finseq_1(k1_nat_1(X1,X2))!=k2_finseq_1(np__2)|~v7_ordinal1(X1)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m1_subset_1(X2,k5_numbers)|k3_finseq_1(X3)!=np__2)).#trainneg
cnf(i_0_778, plain, (k4_finseq_1(X1)=k2_finseq_1(np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__1)).#trainneg
cnf(i_0_4090, negated_conjecture, (k9_xtuple_0(X1)=k2_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|X1!=k1_funct_1(esk11_0,np__1))).#trainneg
cnf(i_0_4082, plain, (k9_xtuple_0(X1)=k2_finseq_1(k2_xcmplx_0(X3,X2))|k2_xcmplx_0(X3,X2)!=k1_card_1(X1)|~v7_ordinal1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_526, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k3_finseq_4(X1,esk2_0,esk3_0,esk4_0),k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_funct_1(esk11_0,np__3))=esk11_0|v1_xboole_0(X1)|~m1_subset_1(esk4_0,X1)|~m1_subset_1(esk3_0,X1)|~m1_subset_1(esk2_0,X1))).#trainneg
cnf(i_0_3601, negated_conjecture, (k1_funct_1(X1,np__1)=esk8_0|X1!=k1_funct_1(esk11_0,np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3595, negated_conjecture, (k3_finseq_4(X1,esk8_0,X2,esk10_0)=k1_funct_1(esk11_0,np__3)|v1_xboole_0(X1)|k1_funct_1(k1_funct_1(esk11_0,np__3),np__2)!=X2|~m1_subset_1(esk10_0,X1)|~m1_subset_1(X2,X1)|~m1_subset_1(esk8_0,X1))).#trainneg
cnf(i_0_4121, negated_conjecture, (k3_finseq_1(np__3)=np__3|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1279, negated_conjecture, (k3_finseq_1(esk11_0)=np__3|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_779, plain, (k3_finseq_1(X1)=k3_finseq_1(X2)|k2_finseq_1(k3_finseq_1(X1))!=k2_finseq_1(np__1)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|k3_finseq_1(X2)!=np__1)).#trainneg
cnf(i_0_4424, negated_conjecture, (~v1_zfmisc_1(esk11_0)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_781, plain, (k1_nat_1(X1,X2)=k3_finseq_1(X3)|k2_finseq_1(k1_nat_1(X1,X2))!=k2_finseq_1(np__1)|~v7_ordinal1(X1)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m1_subset_1(X2,k5_numbers)|k3_finseq_1(X3)!=np__1)).#trainneg
cnf(i_0_782, plain, (v7_ordinal1(esk17_1(X1))|~v1_finset_1(esk17_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_4515, negated_conjecture, (k4_finseq_1(X1)=k2_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|X1!=k1_funct_1(esk11_0,np__1))).#trainneg
cnf(i_0_4390, negated_conjecture, (v1_xboole_0(X1)|m2_finseq_1(k1_funct_1(esk11_0,np__1),X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|esk4_0!=X4|esk2_0!=X2|esk3_0!=X3)).#trainneg
cnf(i_0_4393, negated_conjecture, (k3_finseq_4(X1,X2,X3,X4)=k1_funct_1(esk11_0,np__1)|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|esk4_0!=X4|esk2_0!=X2|esk3_0!=X3)).#trainneg
cnf(i_0_4433, negated_conjecture, (k3_finseq_1(X1)=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k1_card_1(X1)!=np__3)).#trainneg
cnf(i_0_1448, negated_conjecture, (k4_finseq_1(esk11_0)=k2_finseq_1(np__3)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_783, plain, (esk16_2(X1,X2)=k3_finseq_1(X3)|k2_finseq_1(esk16_2(X1,X2))!=k2_finseq_1(np__2)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|k3_finseq_1(X3)!=np__2)).#trainneg
cnf(i_0_784, plain, (esk16_2(X1,X2)=k3_finseq_1(X3)|k2_finseq_1(esk16_2(X1,X2))!=k2_finseq_1(np__1)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|k3_finseq_1(X3)!=np__1)).#trainneg
cnf(i_0_785, plain, (esk16_2(X1,X2)=k3_finseq_1(X3)|k2_finseq_1(esk16_2(X1,X2))!=k2_finseq_1(np__3)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|k3_finseq_1(X3)!=np__3)).#trainneg
cnf(i_0_4153, plain, (k2_xcmplx_0(X1,X2)=k3_finseq_1(k3_finseq_4(X3,X4,X5,X6))|v1_xboole_0(X3)|k2_finseq_1(k2_xcmplx_0(X1,X2))!=k2_finseq_1(np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X6,X3)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2984, negated_conjecture, (v3_card_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),np__3)|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)!=X1|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X2|k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))!=np__3|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)))).#trainneg
cnf(i_0_4529, plain, (k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4))=k2_finseq_1(k2_xcmplx_0(X5,X6))|v1_xboole_0(X1)|k2_xcmplx_0(X5,X6)!=np__3|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X5,k5_numbers)|~m1_subset_1(X6,k5_numbers)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_2743, plain, (v3_ordinal1(X1)|~v1_card_1(k1_zfmisc_1(k1_xboole_0))|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1370, plain, (esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k1_xboole_0|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1720, negated_conjecture, (k4_finseq_1(esk11_0)=k2_finseq_1(np__3)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_786, plain, (esk16_2(X1,X2)=k3_finseq_1(X3)|k4_finseq_1(X3)!=k2_finseq_1(esk16_2(X1,X2))|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_4670, negated_conjecture, (~v1_zfmisc_1(esk11_0))).#trainneg
cnf(i_0_4955, negated_conjecture, (k4_finseq_1(esk11_0)=k2_finseq_1(np__3)|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4686, plain, (v7_ordinal1(esk17_1(X1))|~v3_ordinal1(X1)|~v5_finset_1(X1))).#trainneg
cnf(i_0_4682, plain, (v7_ordinal1(esk17_1(X1))|~v3_ordinal1(X1)|~v1_xboole_0(esk17_1(X1)))).#trainneg
cnf(i_0_1440, negated_conjecture, (k9_xtuple_0(esk11_0)=k2_finseq_1(np__3)|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),u1_struct_0(esk1_0))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_820, plain, (k1_card_1(X3)=np__3|k1_funct_1(X3,np__1)!=X1|k1_funct_1(X3,np__2)!=X2|k3_finseq_1(X3)!=np__3|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_821, plain, (m1_subset_1(X1,k2_finseq_1(X2))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_825, plain, (m1_subset_1(k9_xtuple_0(X1),k1_zfmisc_1(k5_numbers))|~v7_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2509, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),np__3)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,X1)|X1!=np__3|~v7_ordinal1(X1))).#trainneg
cnf(i_0_4150, plain, (k2_xcmplx_0(X2,X1)=k3_finseq_1(k3_finseq_4(X3,X4,X5,X6))|v1_xboole_0(X3)|k2_finseq_1(k2_xcmplx_0(X2,X1))!=k2_finseq_1(np__3)|~v7_ordinal1(X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X6,X3)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_2916, negated_conjecture, (k3_finseq_4(X1,esk2_0,X2,esk4_0)=k1_funct_1(esk11_0,np__1)|v1_xboole_0(X1)|k1_funct_1(k1_funct_1(esk11_0,np__1),np__2)!=X2|~m1_subset_1(esk4_0,X1)|~m1_subset_1(X2,X1)|~m1_subset_1(esk2_0,X1))).#trainneg
cnf(i_0_1056, plain, (v7_ordinal1(k3_finseq_1(X1))|~v1_finset_1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_3864, plain, (m2_finseq_1(X1,k1_zfmisc_1(k1_funct_1(X1,np__3)))|k1_funct_1(X1,np__1)!=X2|k1_funct_1(X1,np__2)!=X3|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_funct_1(X1,np__3)))|~m1_subset_1(X2,k1_zfmisc_1(k1_funct_1(X1,np__3))))).#trainneg
cnf(i_0_2744, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(k1_xboole_0))|~m2_finseq_1(X1,X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_827, plain, (k3_finseq_1(X1)=k3_finseq_1(X2)|k9_xtuple_0(X2)!=k9_xtuple_0(X1)|~v1_finseq_1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4758, negated_conjecture, (k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,np__1)|esk4_0!=esk10_0|esk3_0!=esk6_0)).#trainneg
cnf(i_0_4084, plain, (k9_xtuple_0(X1)=k2_finseq_1(k1_nat_1(X3,X2))|k1_nat_1(X3,X2)!=k1_card_1(X1)|~v7_ordinal1(X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_2664, negated_conjecture, (k3_finseq_4(X1,X2,X3,k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,X7))=k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X4,X5,X6,X7))|v1_xboole_0(X1)|v1_xboole_0(X4)|k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X4,X5,X6,X7)),np__1)!=X2|k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X4,X5,X6,X7)),np__2)!=X3|k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X4,X5,X6,X7)))!=np__3|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X4,X5,X6,X7)))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X4,X5,X6,X7)))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X4,X5,X6,X7)))|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,X7),X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v7_ordinal1(X7)|~m1_matrix_2(k3_finseq_4(X4,X5,X6,X7),k12_matrix_2(np__3))|~m1_subset_1(k3_finseq_4(X4,X5,X6,X7),k12_matrix_2(np__3))|~m1_subset_1(X7,X4)|~m1_subset_1(X6,X4)|~m1_subset_1(X5,X4))).#trainneg
cnf(i_0_4866, negated_conjecture, (v3_card_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),np__3)|k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))!=np__3|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X1|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)))).#trainneg
cnf(i_0_851, plain, (k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4))=k2_finseq_1(k1_nat_1(X5,X6))|v1_xboole_0(X1)|k1_nat_1(X5,X6)!=np__3|~v7_ordinal1(X5)|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X6,k5_numbers)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_4044, negated_conjecture, (esk4_0=np__3|k1_funct_1(esk11_0,np__1)!=esk12_0)).#trainneg
cnf(i_0_1668, plain, (v1_relat_1(esk17_1(k1_zfmisc_1(k1_xboole_0)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1672, plain, (v1_xboole_0(esk17_1(k1_zfmisc_1(k1_xboole_0)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_852, plain, (k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4))=k2_finseq_1(k3_finseq_1(X5))|v1_xboole_0(X1)|k3_finseq_1(X5)!=np__3|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_finseq_1(X5)|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(X5)|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(X5)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_862, plain, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X6)))=k3_finseq_4(X3,X4,X5,X6)|v1_xboole_0(X3)|k1_funct_1(k3_finseq_4(X3,X4,X5,X6),np__1)!=X1|k1_funct_1(k3_finseq_4(X3,X4,X5,X6),np__2)!=X2|k3_finseq_1(k3_finseq_4(X3,X4,X5,X6))!=np__3|~v1_finseq_1(k3_finseq_4(X3,X4,X5,X6))|~v1_funct_1(k3_finseq_4(X3,X4,X5,X6))|~v1_relat_1(k3_finseq_4(X3,X4,X5,X6))|~m1_subset_1(X6,X3)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_5008, negated_conjecture, (k9_xtuple_0(esk11_0)=k2_finseq_1(np__3)|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m2_finseq_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4685, plain, (v7_ordinal1(esk17_1(k1_zfmisc_1(X1)))|~v3_ordinal1(k1_zfmisc_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_1930, plain, (v7_ordinal1(k2_finseq_1(np__3))|~v1_finset_1(k2_finseq_1(np__3))|~v3_ordinal1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_5029, plain, (m1_subset_1(X1,k9_xtuple_0(X2))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,k1_card_1(X2))|~v7_ordinal1(k1_card_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5031, plain, (m1_subset_1(X1,k9_xtuple_0(X2))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5141, plain, (v7_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(k3_finseq_1(X1))|~v1_funct_1(k3_finseq_1(X1))|~v1_relat_1(k3_finseq_1(X1)))).#trainneg
cnf(i_0_4370, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k2_finseq_1(np__2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|k3_finseq_1(X2)!=np__2)).#trainneg
cnf(i_0_1594, plain, (esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3000, negated_conjecture, (esk16_2(np__3,np__3)=np__3)).#trainneg
cnf(i_0_803, negated_conjecture, (k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1547, plain, (k2_finseq_1(k1_xboole_0)=k2_finseq_1(np__2)|k3_finseq_1(k1_xboole_0)!=np__2|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_1871, plain, (v1_relat_1(k2_finseq_1(k1_xboole_0))|~v4_funct_1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_1670, plain, (v5_relat_1(esk17_1(k1_zfmisc_1(k1_xboole_0)),X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_5034, plain, (m1_subset_1(X1,k9_xtuple_0(k3_finseq_4(X3,X4,X5,X6)))|v1_xboole_0(X3)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|X2!=np__3|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X6,X3)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_5033, negated_conjecture, (m1_subset_1(X1,k9_xtuple_0(k10_matrix_3(np__3,esk1_0,esk11_0,X3)))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|X2!=np__3|~m1_matrix_2(X3,k12_matrix_2(np__3))|~m1_subset_1(X3,k12_matrix_2(np__3))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_3599, negated_conjecture, (k3_finseq_1(X1)=np__3|X1!=k1_funct_1(esk11_0,np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4502, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k2_finseq_1(np__1))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|k3_finseq_1(X2)!=np__1)).#trainneg
cnf(i_0_3139, negated_conjecture, (k1_funct_1(np__1,np__1)=esk2_0|esk11_0!=esk12_0)).#trainneg
cnf(i_0_975, plain, (k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4))=k2_finseq_1(k2_xcmplx_0(X5,X6))|v1_xboole_0(X1)|k2_xcmplx_0(X5,X6)!=np__3|~v7_ordinal1(X5)|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X6,k5_numbers)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_5323, plain, (v1_relat_1(esk17_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_5330, plain, (v1_xboole_0(esk17_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_5435, plain, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))=k3_finseq_4(X4,X1,X5,X3)|v1_xboole_0(X4)|k1_funct_1(k3_finseq_4(X4,X1,X5,X3),np__2)!=X2|~m1_subset_1(X3,X4)|~m1_subset_1(X5,X4)|~m1_subset_1(X1,X4))).#trainneg
cnf(i_0_5576, negated_conjecture, (k9_xtuple_0(X1)=k2_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|X1!=k1_funct_1(esk11_0,np__3))).#trainneg
cnf(i_0_5346, plain, (k9_xtuple_0(k3_finseq_4(X2,X3,X4,X5))=k9_xtuple_0(X1)|v1_xboole_0(X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__3|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_2434, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_funct_1(esk11_0,np__1),k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk5_0)),k1_tarski(k4_tarski(np__1,esk6_0))),k1_tarski(k4_tarski(np__1,esk7_0))),k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0))=esk11_0|v1_xboole_0(X1)|~m1_subset_1(esk7_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk5_0,X1))).#trainneg
cnf(i_0_5537, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k2_finseq_1(np__2))|k3_finseq_1(k1_xboole_0)!=np__2|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_5552, plain, (v5_relat_1(esk17_1(k1_zfmisc_1(k1_xboole_0)),X1))).#trainneg
cnf(i_0_5318, negated_conjecture, (esk4_0=np__3|esk12_0!=np__1|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1873, plain, (v1_finset_1(k2_finseq_1(k1_xboole_0))|~v5_finset_1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_2754, plain, (v4_relat_1(esk17_1(k1_zfmisc_1(k1_xboole_0)),X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_984, plain, (k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4))=k2_finseq_1(k1_card_1(X5))|v1_xboole_0(X1)|k1_card_1(X5)!=np__3|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_finseq_1(X5)|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(X5)|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(X5)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_5690, negated_conjecture, (m2_finseq_1(k1_funct_1(esk11_0,np__2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5708, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__2),np__2)=esk6_0)).#trainneg
cnf(i_0_5693, negated_conjecture, (k9_xtuple_0(k1_funct_1(esk11_0,np__2))=k2_finseq_1(np__3))).#trainneg
cnf(i_0_5701, negated_conjecture, (r2_hidden(np__1,k4_finseq_1(k1_funct_1(esk11_0,np__2))))).#trainneg
cnf(i_0_5702, negated_conjecture, (r2_hidden(np__2,k4_finseq_1(k1_funct_1(esk11_0,np__2))))).#trainneg
cnf(i_0_5703, negated_conjecture, (r2_hidden(np__3,k4_finseq_1(k1_funct_1(esk11_0,np__2))))).#trainneg
cnf(i_0_5691, negated_conjecture, (k4_finseq_1(k1_funct_1(esk11_0,np__2))=k2_finseq_1(np__3))).#trainneg
cnf(i_0_5688, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__2),np__1)=esk5_0)).#trainneg
cnf(i_0_5681, negated_conjecture, (v1_finseq_1(k1_funct_1(esk11_0,np__2)))).#trainneg
cnf(i_0_5686, negated_conjecture, (~v1_zfmisc_1(k1_funct_1(esk11_0,np__2)))).#trainneg
cnf(i_0_5684, negated_conjecture, (v3_card_1(k1_funct_1(esk11_0,np__2),np__3))).#trainneg
cnf(i_0_5680, negated_conjecture, (v1_finset_1(k1_funct_1(esk11_0,np__2)))).#trainneg
cnf(i_0_5886, plain, (v4_relat_1(esk17_1(k1_zfmisc_1(k1_xboole_0)),X1))).#trainneg
cnf(i_0_5675, plain, (esk17_1(k1_zfmisc_1(k1_xboole_0))=k1_xboole_0)).#trainneg
cnf(i_0_1313, negated_conjecture, (~v1_xboole_0(esk11_0)|~m1_subset_1(k1_funct_1(esk11_0,np__2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_987, plain, (k9_xtuple_0(X1)=k2_finseq_1(k1_card_1(X2))|k1_card_1(X2)!=k3_finseq_1(X1)|~v1_finseq_1(X1)|~v1_finseq_1(k1_card_1(X2))|~v1_funct_1(X1)|~v1_funct_1(k1_card_1(X2))|~v1_relat_1(X1)|~v1_relat_1(k1_card_1(X2)))).#trainneg
cnf(i_0_5685, negated_conjecture, (~v1_xboole_0(k1_funct_1(esk11_0,np__2)))).#trainneg
cnf(i_0_989, plain, (k9_xtuple_0(X1)=k2_finseq_1(k1_xboole_0)|k3_finseq_1(X1)!=k1_xboole_0|~v1_finseq_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5682, negated_conjecture, (v1_funct_1(k1_funct_1(esk11_0,np__2)))).#trainneg
cnf(i_0_5683, negated_conjecture, (v1_relat_1(k1_funct_1(esk11_0,np__2)))).#trainneg
cnf(i_0_5709, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__2),np__3)=esk7_0)).#trainneg
cnf(i_0_5716, negated_conjecture, (v1_xboole_0(X1)|m2_finseq_1(k1_funct_1(esk11_0,np__2),X1)|~m1_subset_1(esk7_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk5_0,X1))).#trainneg
cnf(i_0_5712, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,esk7_0)))=k1_funct_1(esk11_0,np__2)|k1_funct_1(k1_funct_1(esk11_0,np__2),np__1)!=X1|k1_funct_1(k1_funct_1(esk11_0,np__2),np__2)!=X2)).#trainneg
cnf(i_0_5697, negated_conjecture, (k1_card_1(k1_funct_1(esk11_0,np__2))=np__3)).#trainneg
cnf(i_0_5724, negated_conjecture, (k1_funct_1(X1,np__1)=esk5_0|X1!=k1_funct_1(esk11_0,np__2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5535, plain, (m1_subset_1(k2_finseq_1(np__2),k1_zfmisc_1(k5_numbers))|k3_finseq_1(k1_xboole_0)!=np__2|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_3150, negated_conjecture, (k1_funct_1(np__1,np__2)=esk3_0|esk11_0!=esk12_0)).#trainneg
cnf(i_0_990, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k9_xtuple_0(X2))|~v7_ordinal1(k3_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_1548, plain, (k2_finseq_1(k1_xboole_0)=k2_finseq_1(np__1)|k3_finseq_1(k1_xboole_0)!=np__1|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_3124, negated_conjecture, (m2_finseq_1(np__1,u1_struct_0(esk1_0))|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1682, plain, (v3_funct_2(X1,X2,X3)|~v2_funct_2(X1,X3)|~v2_funct_1(X1)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2767, plain, (v3_funct_2(X1,X2,X3)|~v2_funct_2(X1,X3)|~v2_funct_1(X1)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_992, plain, (m1_subset_1(X1,k5_numbers)|v1_xboole_0(k2_finseq_1(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k2_finseq_1(X2)))).#trainneg
cnf(i_0_5728, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk5_0)),k1_tarski(k4_tarski(np__1,X1))),k1_tarski(k4_tarski(np__1,esk7_0)))=k1_funct_1(esk11_0,np__2)|k1_funct_1(k1_funct_1(esk11_0,np__2),np__2)!=X1)).#trainneg
cnf(i_0_5726, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,np__1)),k1_tarski(k4_tarski(np__1,X1))),k1_tarski(k4_tarski(np__1,np__3)))=esk12_0|np__2!=X1)).#trainneg
cnf(i_0_5718, negated_conjecture, (k3_finseq_4(X1,X2,X3,esk7_0)=k1_funct_1(esk11_0,np__2)|v1_xboole_0(X1)|k1_funct_1(k1_funct_1(esk11_0,np__2),np__1)!=X2|k1_funct_1(k1_funct_1(esk11_0,np__2),np__2)!=X3|~m1_subset_1(esk7_0,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_6045, negated_conjecture, (v1_finseq_1(np__2)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1279, negated_conjecture, (k3_finseq_1(esk11_0)=np__3|~m1_subset_1(k1_funct_1(esk11_0,np__2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_6046, negated_conjecture, (~v1_zfmisc_1(np__2)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1683, plain, (v2_funct_1(X1)|~v3_funct_2(X1,X2,X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2768, plain, (v2_funct_1(X1)|~v3_funct_2(X1,X2,X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_5720, negated_conjecture, (k1_funct_1(X1,np__2)=esk6_0|X1!=k1_funct_1(esk11_0,np__2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5717, negated_conjecture, (k3_finseq_4(X1,esk5_0,X2,esk7_0)=k1_funct_1(esk11_0,np__2)|v1_xboole_0(X1)|k1_funct_1(k1_funct_1(esk11_0,np__2),np__2)!=X2|~m1_subset_1(esk7_0,X1)|~m1_subset_1(X2,X1)|~m1_subset_1(esk5_0,X1))).#trainneg
cnf(i_0_5729, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk2_0)),k1_tarski(k4_tarski(np__1,X1))),k1_tarski(k4_tarski(np__1,esk4_0)))=k1_funct_1(esk11_0,np__1)|esk3_0!=X1)).#trainneg
cnf(i_0_5710, negated_conjecture, (k3_finseq_1(k1_funct_1(esk11_0,np__2))=np__3)).#trainneg
cnf(i_0_5687, negated_conjecture, (k3_finseq_4(X1,esk5_0,esk6_0,esk7_0)=k1_funct_1(esk11_0,np__2)|v1_xboole_0(X1)|~m1_subset_1(esk7_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk5_0,X1))).#trainneg
cnf(i_0_6048, negated_conjecture, (v1_finset_1(np__2)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1456, negated_conjecture, (k3_finseq_1(esk11_0)=np__3|~m2_finseq_1(k1_funct_1(esk11_0,np__2),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1549, plain, (k2_finseq_1(k1_xboole_0)=k2_finseq_1(np__3)|k3_finseq_1(k1_xboole_0)!=np__3|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_6088, negated_conjecture, (~v1_xboole_0(np__2)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_6066, negated_conjecture, (~v1_xboole_0(esk11_0)|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1686, plain, (v2_funct_2(X1,X2)|~v3_funct_2(X1,X3,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X3))).#trainneg
cnf(i_0_2771, plain, (v2_funct_2(X1,X2)|~v3_funct_2(X1,X3,X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1527, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|v1_xboole_0(k4_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1)))|X3!=k1_funct_1(X1,X2)|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X2,k4_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))))).#trainneg
cnf(i_0_2376, negated_conjecture, (X1=k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2))|k2_finseq_1(np__3)!=k2_finseq_1(X1)|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2))|~m1_subset_1(X1,k5_numbers)|~m1_matrix_2(X2,k12_matrix_2(np__3))|~m1_subset_1(X2,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_5725, plain, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,k1_funct_1(k3_finseq_4(X2,X1,X3,X4),np__2)))),k1_tarski(k4_tarski(np__1,X4)))=k3_finseq_4(X2,X1,X3,X4)|v1_xboole_0(X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_2436, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k3_finseq_4(X1,esk2_0,esk3_0,esk4_0),k3_finseq_4(X2,esk5_0,esk6_0,esk7_0),k1_funct_1(esk11_0,np__3))=esk11_0|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_subset_1(esk7_0,X2)|~m1_subset_1(esk6_0,X2)|~m1_subset_1(esk5_0,X2)|~m1_subset_1(esk4_0,X1)|~m1_subset_1(esk3_0,X1)|~m1_subset_1(esk2_0,X1))).#trainneg
cnf(i_0_6224, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k2_finseq_1(np__1))|k3_finseq_1(k1_xboole_0)!=np__1|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_6222, plain, (m1_subset_1(k2_finseq_1(np__1),k1_zfmisc_1(k5_numbers))|k3_finseq_1(k1_xboole_0)!=np__1|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_6529, negated_conjecture, (k3_finseq_1(esk11_0)=np__3)).#trainneg
cnf(i_0_1440, negated_conjecture, (k9_xtuple_0(esk11_0)=k2_finseq_1(np__3)|~m1_subset_1(k1_funct_1(esk11_0,np__2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1043, plain, (k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4))=k2_finseq_1(k1_nat_1(X5,X6))|v1_xboole_0(X1)|k1_nat_1(X5,X6)!=np__3|~v7_ordinal1(X6)|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X5,k5_numbers)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_6550, negated_conjecture, (~v1_xboole_0(esk11_0))).#trainneg
cnf(i_0_6559, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X3!=k1_funct_1(X1,X2)|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X2,k2_finseq_1(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_6608, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk5_0)),k1_tarski(k4_tarski(np__1,esk6_0))),k1_tarski(k4_tarski(np__1,esk7_0)))=k1_funct_1(esk11_0,np__2))).#trainneg
cnf(i_0_6606, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,np__1)),k1_tarski(k4_tarski(np__1,np__2))),k1_tarski(k4_tarski(np__1,np__3)))=esk12_0)).#trainneg
cnf(i_0_6100, negated_conjecture, (v1_funct_1(np__2)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1720, negated_conjecture, (k4_finseq_1(esk11_0)=k2_finseq_1(np__3)|~m1_subset_1(k1_funct_1(esk11_0,np__2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1064, plain, (v7_ordinal1(k2_finseq_1(X1))|~v1_finset_1(k2_finseq_1(X1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_6101, negated_conjecture, (v1_relat_1(np__2)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1065, plain, (v3_ordinal1(k9_xtuple_0(X1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1067, plain, (v3_card_1(X1,np__3)|k1_funct_1(X1,np__2)!=X2|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6607, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk8_0)),k1_tarski(k4_tarski(np__1,esk9_0))),k1_tarski(k4_tarski(np__1,esk10_0)))=k1_funct_1(esk11_0,np__3))).#trainneg
cnf(i_0_6556, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X3!=k1_funct_1(X1,X2)|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X2,k2_finseq_1(np__3))|k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))!=np__3|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1)))).#trainneg
cnf(i_0_6628, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,np__2))),k1_tarski(k4_tarski(np__1,X4)))=k3_finseq_4(X2,X1,X3,X4)|v1_xboole_0(X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~m1_subset_1(X1,X2)|k3_finseq_4(X2,X1,X3,X4)!=esk12_0|~v1_finseq_1(k3_finseq_4(X2,X1,X3,X4))|~v1_funct_1(k3_finseq_4(X2,X1,X3,X4))|~v1_relat_1(k3_finseq_4(X2,X1,X3,X4)))).#trainneg
cnf(i_0_6723, negated_conjecture, (k1_card_1(esk11_0)=np__3)).#trainneg
cnf(i_0_5008, negated_conjecture, (k9_xtuple_0(esk11_0)=k2_finseq_1(np__3)|~m2_finseq_1(k1_funct_1(esk11_0,np__2),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1102, plain, (v3_ordinal1(k9_xtuple_0(X1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(k1_card_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6964, plain, (v3_card_1(X1,np__3)|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6245, plain, (v3_funct_2(X1,X2,X3)|~v2_funct_1(X1)|~v1_xboole_0(X2)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v3_funct_2(X1,k5_numbers,X3)|~m2_finseq_1(X1,X3))).#trainneg
cnf(i_0_1104, plain, (m1_subset_1(k9_xtuple_0(X1),k1_zfmisc_1(k5_numbers))|~v7_ordinal1(k1_card_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1105, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k9_xtuple_0(X2))|~v7_ordinal1(k1_card_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6719, negated_conjecture, (k9_xtuple_0(esk11_0)=k2_finseq_1(np__3))).#trainneg
cnf(i_0_4955, negated_conjecture, (k4_finseq_1(esk11_0)=k2_finseq_1(np__3)|~m2_finseq_1(k1_funct_1(esk11_0,np__2),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1106, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,k1_card_1(X2))|~v7_ordinal1(k1_card_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_5874, plain, (m1_subset_1(X1,k5_numbers)|k3_finseq_1(k1_xboole_0)!=np__2|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6247, plain, (v3_funct_2(X1,X2,X3)|~v2_funct_1(X1)|~v1_xboole_0(X3)|~v1_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0))|~v3_funct_2(X1,k5_numbers,X3)|~m2_finseq_1(X1,X3))).#trainneg
cnf(i_0_1108, plain, (k9_xtuple_0(k1_card_1(X1))=k2_finseq_1(k1_card_1(X1))|~v1_finseq_1(k1_card_1(X1))|~v1_funct_1(k1_card_1(X1))|~v1_relat_1(k1_card_1(X1)))).#trainneg
cnf(i_0_1110, plain, (k9_xtuple_0(X1)=k2_finseq_1(k1_xboole_0)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7159, negated_conjecture, (k4_finseq_1(esk11_0)=k2_finseq_1(np__3))).#trainneg
cnf(i_0_7026, negated_conjecture, (v1_finset_1(esk11_0))).#trainneg
cnf(i_0_1173, plain, (k3_finseq_1(X1)=np__3|k2_finseq_1(k3_finseq_1(X1))!=k2_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1176, plain, (esk16_2(X1,X2)=np__3|k2_finseq_1(esk16_2(X1,X2))!=k2_finseq_1(np__3)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_6847, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X1)|X1!=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)|~v7_ordinal1(X1)|~m1_matrix_2(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),k12_matrix_2(np__3))|~m1_subset_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),k12_matrix_2(np__3)))).#trainneg
cnf(i_0_6845, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)),np__2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,X1)|X1!=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)|~v7_ordinal1(X1)|~m1_matrix_2(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),k12_matrix_2(np__3))|~m1_subset_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),k12_matrix_2(np__3)))).#trainneg
cnf(i_0_6172, negated_conjecture, (k1_card_1(np__2)=np__3|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1194, plain, (X1=k3_finseq_1(X3)|v1_xboole_0(a_1_0_finseq_1(X2))|k9_xtuple_0(X3)!=k2_finseq_1(X1)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m1_subset_1(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_1195, plain, (r1_xxreal_0(np__1,X1)|v1_xboole_0(a_1_0_finseq_1(X2))|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_4095, negated_conjecture, (k9_xtuple_0(X1)=k2_finseq_1(X2)|X2!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k5_numbers)|X1!=k1_funct_1(esk11_0,np__1))).#trainneg
cnf(i_0_5721, negated_conjecture, (k1_funct_1(X1,np__3)=esk7_0|X1!=k1_funct_1(esk11_0,np__2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5153, plain, (m2_finseq_1(X1,k1_zfmisc_1(k1_funct_1(X1,np__3)))|k1_funct_1(X1,np__1)!=X2|k1_funct_1(X1,np__2)!=k1_funct_1(X1,np__3)|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_funct_1(X1,np__3))))).#trainneg
cnf(i_0_2510, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))),np__2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,X2)|~v7_ordinal1(X2)|~m1_matrix_2(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),k12_matrix_2(np__3))|~m1_subset_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),k12_matrix_2(np__3)))).#trainneg
cnf(i_0_6474, negated_conjecture, (k3_finseq_1(np__2)=np__3|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1196, plain, (r1_xxreal_0(X1,X2)|v1_xboole_0(a_1_0_finseq_1(X2))|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_1197, plain, (m2_subset_1(X1,k1_numbers,k5_numbers)|v1_xboole_0(a_1_0_finseq_1(X2))|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_3607, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_funct_1(esk11_0,np__1),k1_funct_1(esk11_0,np__2),k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,esk8_0)),k1_tarski(k4_tarski(np__1,esk9_0))),k1_tarski(k4_tarski(np__1,esk10_0))))=esk11_0)).#trainneg
cnf(i_0_7387, negated_conjecture, (k1_funct_1(X1,np__2)=X2|X1!=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)!=X3|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X2)).#trainneg
cnf(i_0_7381, negated_conjecture, (v3_card_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),np__3)|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)!=X1|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X2)).#trainneg
cnf(i_0_7383, negated_conjecture, (~v1_zfmisc_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)!=X1|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X2)).#trainneg
cnf(i_0_7570, plain, (r1_xxreal_0(X1,X2)|v1_xboole_0(a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_7368, plain, (r1_xxreal_0(np__1,X1)|v1_xboole_0(a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,a_1_0_finseq_1(X2)))).#trainneg
cnf(i_0_7546, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k1_funct_1(esk11_0,np__1)),np__2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,esk3_0)|~v7_ordinal1(esk3_0)|~m1_matrix_2(k1_funct_1(esk11_0,np__1),k12_matrix_2(np__3))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k12_matrix_2(np__3)))).#trainneg
cnf(i_0_1226, plain, (k3_finseq_1(k3_finseq_4(X1,X2,X3,X4))=np__2|v1_xboole_0(X1)|k2_finseq_1(np__2)!=k2_finseq_1(np__3)|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1227, plain, (k3_finseq_1(k3_finseq_4(X1,X2,X3,X4))=np__1|v1_xboole_0(X1)|k2_finseq_1(np__3)!=k2_finseq_1(np__1)|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_7377, negated_conjecture, (k9_xtuple_0(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))=k2_finseq_1(np__3)|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)!=X1|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X2)).#trainneg
cnf(i_0_7371, negated_conjecture, (k4_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))=k2_finseq_1(np__3)|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)!=X1|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X2)).#trainneg
cnf(i_0_7386, negated_conjecture, (k1_card_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))=np__3|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)!=X1|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X2)).#trainneg
cnf(i_0_7372, negated_conjecture, (k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))=np__3|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)!=X1|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X2)).#trainneg
cnf(i_0_3300, negated_conjecture, (v1_xboole_0(X1)|v1_xboole_0(X2)|m2_finseq_1(k1_funct_1(esk11_0,np__1),X2)|~m1_subset_1(esk4_0,X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|~m1_subset_1(esk4_0,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|esk2_0!=X3|esk3_0!=X4)).#trainneg
cnf(i_0_3692, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X4)|X4!=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X1)|~v7_ordinal1(X4)|~m1_matrix_2(k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))),k12_matrix_2(np__3))|~m1_subset_1(k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))),k12_matrix_2(np__3))|~v7_ordinal1(X1)|~m1_matrix_2(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),k12_matrix_2(np__3))|~m1_subset_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),k12_matrix_2(np__3)))).#trainneg
cnf(i_0_481, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,k1_funct_1(esk11_0,np__1))),k1_tarski(k4_tarski(np__1,k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0)))),k1_tarski(k4_tarski(np__1,k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0))))=esk11_0|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_3001, plain, (esk16_2(X1,X2)=X1|v1_xboole_0(k2_finseq_1(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k2_finseq_1(X2)))).#trainneg
cnf(i_0_1229, plain, (k2_finseq_1(np__2)=k2_finseq_1(np__3)|v1_xboole_0(X1)|k3_finseq_1(k3_finseq_4(X1,X2,X3,X4))!=np__2|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1230, plain, (k2_finseq_1(np__3)=k2_finseq_1(np__1)|v1_xboole_0(X1)|k3_finseq_1(k3_finseq_4(X1,X2,X3,X4))!=np__1|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1232, plain, (k3_finseq_1(X1)=k3_finseq_1(k3_finseq_4(X2,X3,X4,X5))|v1_xboole_0(X2)|k2_finseq_1(k3_finseq_1(X1))!=k2_finseq_1(np__3)|~v1_finseq_1(k3_finseq_4(X2,X3,X4,X5))|~v1_finseq_1(X1)|~v1_funct_1(k3_finseq_4(X2,X3,X4,X5))|~v1_funct_1(X1)|~v1_relat_1(k3_finseq_4(X2,X3,X4,X5))|~v1_relat_1(X1)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_7633, plain, (r1_xxreal_0(X1,X2)|v1_xboole_0(k2_finseq_1(X2))|~v7_ordinal1(X2)|~m1_subset_1(X1,k2_finseq_1(X2)))).#trainneg
cnf(i_0_7631, negated_conjecture, (k1_funct_1(X1,np__2)=X2|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X2|X1!=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7803, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,k1_funct_1(X1,np__1))),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X4)|X4!=k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),np__1)|~v7_ordinal1(X4)|~v7_ordinal1(k1_funct_1(X1,np__1))|~m1_matrix_2(k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,k1_funct_1(X1,np__1))),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))),k12_matrix_2(np__3))|~m1_matrix_2(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,k1_funct_1(X1,np__1))),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),k12_matrix_2(np__3))|~m1_subset_1(k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,k1_funct_1(X1,np__1))),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))),k12_matrix_2(np__3))|~m1_subset_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,k1_funct_1(X1,np__1))),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),k12_matrix_2(np__3))|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_5995, negated_conjecture, (k9_xtuple_0(np__2)=k2_finseq_1(np__3)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1234, plain, (k1_nat_1(X1,X2)=k3_finseq_1(k3_finseq_4(X3,X4,X5,X6))|v1_xboole_0(X3)|k2_finseq_1(k1_nat_1(X1,X2))!=k2_finseq_1(np__3)|~v7_ordinal1(X1)|~v1_finseq_1(k3_finseq_4(X3,X4,X5,X6))|~v1_funct_1(k3_finseq_4(X3,X4,X5,X6))|~v1_relat_1(k3_finseq_4(X3,X4,X5,X6))|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X6,X3)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_6704, plain, (m1_subset_1(X1,k5_numbers)|k3_finseq_1(k1_xboole_0)!=np__1|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6047, negated_conjecture, (v3_card_1(np__2,np__3)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_6238, negated_conjecture, (v4_relat_1(np__1,k5_numbers)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1235, plain, (esk16_2(X1,X2)=k3_finseq_1(k3_finseq_4(X3,X4,X5,X6))|v1_xboole_0(X3)|k2_finseq_1(esk16_2(X1,X2))!=k2_finseq_1(np__3)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~v1_finseq_1(k3_finseq_4(X3,X4,X5,X6))|~v1_funct_1(k3_finseq_4(X3,X4,X5,X6))|~v1_relat_1(k3_finseq_4(X3,X4,X5,X6))|~m1_subset_1(X6,X3)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_2176, plain, (v1_funct_1(esk15_5(X1,X2,X3,X4,X5))|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k4_tarski(X3,X4),k2_matrix_1(X2))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_5904, plain, (k9_xtuple_0(k3_finseq_4(X2,X3,X4,X5))=k9_xtuple_0(X1)|v1_xboole_0(X2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k1_card_1(X1)!=np__3|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_5167, plain, (k3_finseq_1(k3_finseq_4(X1,X2,X3,X4))=k3_finseq_1(X5)|v1_xboole_0(X1)|k9_xtuple_0(X5)!=k2_finseq_1(np__3)|~v1_finseq_1(X5)|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(X5)|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(X5)|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_525, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_funct_1(esk11_0,np__1),k3_finseq_4(X1,esk5_0,esk6_0,esk7_0),k3_finseq_4(u1_struct_0(esk1_0),esk8_0,esk9_0,esk10_0))=esk11_0|v1_xboole_0(X1)|~m1_subset_1(esk7_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk5_0,X1))).#trainneg
cnf(i_0_7948, negated_conjecture, (k1_funct_1(X1,np__2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)|X1!=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_6037, negated_conjecture, (k4_finseq_1(np__2)=k2_finseq_1(np__3)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1583, plain, (k2_finseq_1(np__2)=k2_finseq_1(np__3)|k3_finseq_1(np__3)!=np__2|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_1296, plain, (r2_hidden(X1,k2_finseq_1(np__3))|v1_xboole_0(X2)|~r1_xxreal_0(X1,k3_finseq_1(k3_finseq_4(X2,X3,X4,X5)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k3_finseq_1(k3_finseq_4(X2,X3,X4,X5)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(k3_finseq_4(X2,X3,X4,X5))|~v1_funct_1(k3_finseq_4(X2,X3,X4,X5))|~v1_relat_1(k3_finseq_4(X2,X3,X4,X5))|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_1298, plain, (r2_hidden(X1,k2_finseq_1(np__2))|~r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k3_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|k3_finseq_1(X2)!=np__2)).#trainneg
cnf(i_0_1584, plain, (k2_finseq_1(np__3)=k2_finseq_1(np__1)|k3_finseq_1(np__3)!=np__1|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_1299, plain, (r2_hidden(X1,k2_finseq_1(np__1))|~r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k3_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|k3_finseq_1(X2)!=np__1)).#trainneg
cnf(i_0_1300, plain, (r2_hidden(X1,k2_finseq_1(np__3))|~r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k3_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|k3_finseq_1(X2)!=np__3)).#trainneg
cnf(i_0_6595, plain, (v1_xboole_0(X1)|m2_finseq_1(k3_finseq_4(X3,X2,X4,X5),X1)|v1_xboole_0(X3)|~m1_subset_1(X5,X1)|~m1_subset_1(k1_funct_1(k3_finseq_4(X3,X2,X4,X5),np__2),X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_4417, negated_conjecture, (v1_xboole_0(X1)|m2_finseq_1(k1_funct_1(esk11_0,np__1),X1)|~m1_subset_1(esk4_0,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|esk2_0!=X2|esk3_0!=X3)).#trainneg
cnf(i_0_5534, plain, (v3_ordinal1(k2_finseq_1(np__2))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|k3_finseq_1(k1_xboole_0)!=np__2|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_1301, plain, (r2_hidden(X1,k4_finseq_1(X2))|~r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k3_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_6931, plain, (v7_ordinal1(k2_finseq_1(X1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1)|~v1_xboole_0(k2_finseq_1(X1)))).#trainneg
cnf(i_0_1308, plain, (v7_ordinal1(k4_finseq_1(X1))|~v1_finset_1(k4_finseq_1(X1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8481, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X3!=k1_funct_1(X1,X2)|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~r1_xxreal_0(X2,k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1)))|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1)))).#trainneg
cnf(i_0_8264, plain, (np__2=np__3|k3_finseq_1(np__3)!=np__2|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_8333, plain, (np__3=np__1|k3_finseq_1(np__3)!=np__1|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_1315, plain, (v4_relat_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),k5_numbers)|v1_xboole_0(X4)|~m1_subset_1(X3,X4)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4))).#trainneg
cnf(i_0_8352, negated_conjecture, (~v1_xboole_0(k2_finseq_1(np__1))|k3_finseq_1(np__3)!=np__1|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_8523, plain, (v7_ordinal1(k4_finseq_1(X1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v5_finset_1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_1341, plain, (k9_xtuple_0(k1_card_1(X1))=k2_finseq_1(k1_card_1(X2))|k1_card_1(X2)!=k1_card_1(X1)|~v1_finseq_1(k1_card_1(X1))|~v1_finseq_1(X2)|~v1_funct_1(k1_card_1(X1))|~v1_funct_1(X2)|~v1_relat_1(k1_card_1(X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_8324, plain, (m1_subset_1(k2_finseq_1(np__1),k1_zfmisc_1(k5_numbers))|k3_finseq_1(np__3)!=np__1|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_1342, plain, (k2_finseq_1(k1_card_1(X1))=k9_xtuple_0(k1_xboole_0)|k1_card_1(X1)!=k1_xboole_0|~v1_finseq_1(k1_xboole_0)|~v1_finseq_1(X1)|~v1_funct_1(k1_xboole_0)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5541, plain, (k3_finseq_1(k1_xboole_0)=np__3|k2_finseq_1(np__2)!=k2_finseq_1(np__3)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|k3_finseq_1(k1_xboole_0)!=np__2)).#trainneg
cnf(i_0_8501, plain, (v7_ordinal1(k2_finseq_1(X1))|~v1_xboole_0(k2_finseq_1(X1))|~v7_ordinal1(X1)|~v1_card_1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_7391, negated_conjecture, (k1_funct_1(X1,np__1)=X2|X1!=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)!=X2|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X3)).#trainneg
cnf(i_0_6908, negated_conjecture, (v1_xboole_0(X1)|m2_finseq_1(esk12_0,X1)|~m1_subset_1(np__3,X1)|~m1_subset_1(np__2,X1)|~m1_subset_1(np__1,X1))).#trainneg
cnf(i_0_6682, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k3_finseq_4(X1,esk2_0,esk3_0,esk4_0),k1_funct_1(esk11_0,np__2),k1_funct_1(esk11_0,np__3))=esk11_0|v1_xboole_0(X1)|~m1_subset_1(esk4_0,X1)|~m1_subset_1(esk3_0,X1)|~m1_subset_1(esk2_0,X1))).#trainneg
cnf(i_0_5144, plain, (m2_finseq_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_4601, negated_conjecture, (v1_xboole_0(X1)|m2_finseq_1(k1_funct_1(esk11_0,np__3),X1)|~m1_subset_1(esk10_0,X1)|~m1_subset_1(X2,X1)|~m1_subset_1(esk8_0,X1)|esk9_0!=X2)).#trainneg
cnf(i_0_1345, plain, (k2_finseq_1(k1_card_1(X1))=k9_xtuple_0(np__3)|k1_card_1(X1)!=np__3|~v1_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(np__3)|~v1_funct_1(X1)|~v1_relat_1(np__3)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8650, plain, (k9_xtuple_0(k1_xboole_0)=k2_finseq_1(np__3)|k1_xboole_0!=np__3|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_4360, plain, (v1_relat_1(k2_finseq_1(np__2))|~v4_funct_1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__2)).#trainneg
cnf(i_0_1346, plain, (k9_xtuple_0(X1)=k2_finseq_1(k1_card_1(X2))|k1_card_1(X2)!=k1_xboole_0|~v1_finseq_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1348, plain, (k9_xtuple_0(X1)=k2_finseq_1(k1_card_1(X2))|k1_card_1(X2)!=k1_card_1(X1)|~v1_finseq_1(X1)|~v1_finseq_1(k1_card_1(X2))|~v1_funct_1(X1)|~v1_funct_1(k1_card_1(X2))|~v1_relat_1(X1)|~v1_relat_1(k1_card_1(X2)))).#trainneg
cnf(i_0_8529, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X3!=k1_funct_1(X1,X2)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_8527, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X3!=k1_funct_1(X1,X2)|~r1_xxreal_0(X2,k1_card_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1)))|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(k1_card_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1)))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1)))).#trainneg
cnf(i_0_8660, negated_conjecture, (k1_funct_1(X1,np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X2|X1!=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8524, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X3!=k1_funct_1(X1,X2)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_matrix_2(X1,k12_matrix_2(np__3))|k10_matrix_3(np__3,esk1_0,esk11_0,X1)!=k1_funct_1(esk11_0,np__3)|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1)))).#trainneg
cnf(i_0_8753, plain, (m2_finseq_1(k3_finseq_4(X4,X1,X2,X3),k1_zfmisc_1(X3))|v1_xboole_0(X4)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m1_subset_1(X3,X4)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4))).#trainneg
cnf(i_0_8525, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X3!=k1_funct_1(X1,X2)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_matrix_2(X1,k12_matrix_2(np__3))|k10_matrix_3(np__3,esk1_0,esk11_0,X1)!=k1_funct_1(esk11_0,np__1)|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1)))).#trainneg
cnf(i_0_8526, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X3!=k1_funct_1(X1,X2)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_matrix_2(X1,k12_matrix_2(np__3))|k1_card_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))!=np__3|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1)))).#trainneg
cnf(i_0_1349, plain, (k9_xtuple_0(X1)=k2_finseq_1(k1_xboole_0)|k1_card_1(X1)!=k1_xboole_0|~v1_finseq_1(X1)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(X1)|~v1_funct_1(k1_xboole_0)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8849, plain, (k3_finseq_1(k1_xboole_0)=np__3|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|k1_xboole_0!=np__3)).#trainneg
cnf(i_0_4362, plain, (v1_finset_1(k2_finseq_1(np__2))|~v5_finset_1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__2)).#trainneg
cnf(i_0_1353, plain, (k9_xtuple_0(X1)=k2_finseq_1(k1_xboole_0)|k1_card_1(X1)!=k1_xboole_0|~v1_finseq_1(X1)|~v1_finseq_1(X2)|~v1_funct_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X1)|~v1_relat_1(X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_9030, negated_conjecture, (k1_funct_1(X1,np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)|X1!=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8985, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,k1_funct_1(X1,X2))=k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(k1_funct_1(X1,X2))|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_9009, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)),np__3)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,X1)|X1!=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,X2)|~v7_ordinal1(X1)|~m1_matrix_2(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),k12_matrix_2(np__3))|~m1_subset_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),k12_matrix_2(np__3))|X2!=np__3|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4096, negated_conjecture, (k1_card_1(X1)=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|X1!=k1_funct_1(esk11_0,np__1))).#trainneg
cnf(i_0_437, negated_conjecture, (m2_finseq_1(esk11_0,k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__2),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__3),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_8854, plain, (k4_finseq_1(k1_xboole_0)=k2_finseq_1(np__3)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|k1_xboole_0!=np__3)).#trainneg
cnf(i_0_6106, negated_conjecture, (esk7_0=np__3|k1_funct_1(esk11_0,np__2)!=esk12_0)).#trainneg
cnf(i_0_4491, plain, (v1_relat_1(k2_finseq_1(np__1))|~v4_funct_1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__1)).#trainneg
cnf(i_0_4493, plain, (v1_finset_1(k2_finseq_1(np__1))|~v5_finset_1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_finseq_1(X1)!=np__1)).#trainneg
cnf(i_0_1363, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(a_1_0_finseq_1(X3),X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X3)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1364, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(k9_xtuple_0(X3),X2)|~r1_xxreal_0(X1,k3_finseq_1(X3))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k3_finseq_1(X3))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_4178, plain, (k2_xcmplx_0(X2,X1)=k3_finseq_1(X3)|k4_finseq_1(X3)!=k2_finseq_1(k2_xcmplx_0(X2,X1))|~v7_ordinal1(X2)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_3880, negated_conjecture, (k1_nat_1(X1,X2)=k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X3))|k2_finseq_1(k1_nat_1(X1,X2))!=k2_finseq_1(np__3)|~v7_ordinal1(X2)|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X3))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X3))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X3))|~m1_subset_1(X1,k5_numbers)|~m1_matrix_2(X3,k12_matrix_2(np__3))|~m1_subset_1(X3,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_1365, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(k2_finseq_1(X3),X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X3)|~v7_ordinal1(X3)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_8617, plain, (v1_funct_1(k2_finseq_1(np__1))|~v4_funct_1(k1_zfmisc_1(k5_numbers))|k3_finseq_1(np__3)!=np__1|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_1369, plain, (m1_subset_1(X1,k5_numbers)|v1_xboole_0(k4_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~m1_subset_1(X1,k4_finseq_1(X2)))).#trainneg
cnf(i_0_3598, negated_conjecture, (k1_funct_1(X1,np__3)=esk10_0|X1!=k1_funct_1(esk11_0,np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_4436, negated_conjecture, (k4_finseq_1(X1)=k2_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k1_card_1(X1)!=np__3)).#trainneg
cnf(i_0_957, negated_conjecture, (k9_xtuple_0(X1)=k2_finseq_1(X2)|X2!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k5_numbers)|X1!=esk12_0)).#trainneg
cnf(i_0_9170, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),np__3)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,np__3))).#trainneg
cnf(i_0_9260, negated_conjecture, (esk7_0=np__3|esk12_0!=np__2|esk11_0!=esk12_0)).#trainneg
cnf(i_0_8520, plain, (v7_ordinal1(k4_finseq_1(X1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_xboole_0(k4_finseq_1(X1)))).#trainneg
cnf(i_0_1473, plain, (m2_subset_1(k1_card_1(X1),k1_numbers,k5_numbers)|~v1_finseq_1(k1_card_1(X1))|~v1_funct_1(k1_card_1(X1))|~v1_relat_1(k1_card_1(X1)))).#trainneg
cnf(i_0_1480, plain, (m1_subset_1(X1,k2_finseq_1(np__3))|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_9306, plain, (k2_xcmplx_0(X1,X2)=np__3|k2_finseq_1(k2_xcmplx_0(X1,X2))!=k2_finseq_1(np__3)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_9198, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3))|X1!=esk12_0|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9197, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),np__2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3))|X1!=esk12_0|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1481, plain, (m1_subset_1(X1,k2_zfmisc_1(k5_numbers,X2))|~m2_finseq_1(k2_finseq_1(np__3),X2)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1552, plain, (r2_hidden(X1,k2_finseq_1(k1_xboole_0))|~r1_xxreal_0(X1,k3_finseq_1(k1_xboole_0))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k3_finseq_1(k1_xboole_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_9216, negated_conjecture, (~v1_xboole_0(X1)|X1!=k1_funct_1(esk11_0,np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2994, plain, (esk16_2(X1,k3_finseq_1(X2))=X1|~r2_hidden(X1,k9_xtuple_0(X2))|~v7_ordinal1(k3_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_2646, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X1,X2,X3,X4))),np__2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,X5)|v1_xboole_0(X1)|X5!=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,X3)|~v7_ordinal1(X5)|~m1_matrix_2(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X1,X2,X3,X4)),k12_matrix_2(np__3))|~m1_subset_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X1,X2,X3,X4)),k12_matrix_2(np__3))|~v7_ordinal1(X3)|~m1_matrix_2(k3_finseq_4(X1,X2,X3,X4),k12_matrix_2(np__3))|~m1_subset_1(k3_finseq_4(X1,X2,X3,X4),k12_matrix_2(np__3))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_6858, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X2)|X2!=np__1|~v7_ordinal1(X2)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3))|X1!=esk12_0|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5234, negated_conjecture, (k3_finseq_4(X1,X2,X3,k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,X4))=k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X5,X6,X7,X4))|v1_xboole_0(X5)|v1_xboole_0(X1)|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X6)!=X2|k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k3_finseq_4(X5,X6,X7,X4)),np__2)!=X3|~v7_ordinal1(X4)|~m1_matrix_2(k3_finseq_4(X5,X6,X7,X4),k12_matrix_2(np__3))|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,X4),X1)|~m1_subset_1(k3_finseq_4(X5,X6,X7,X4),k12_matrix_2(np__3))|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~m1_subset_1(X4,X5)|~m1_subset_1(X7,X5)|~m1_subset_1(X6,X5)|~v7_ordinal1(X6))).#trainneg
cnf(i_0_6221, plain, (v3_ordinal1(k2_finseq_1(np__1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|k3_finseq_1(k1_xboole_0)!=np__1|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_1569, plain, (m1_subset_1(esk16_2(X1,X2),k5_numbers)|~r1_xxreal_0(np__1,esk16_2(X1,X2))|~r1_xxreal_0(esk16_2(X1,X2),X3)|~v7_ordinal1(X3)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_1233, plain, (X1=k3_finseq_1(k3_finseq_4(X2,X3,X4,X5))|v1_xboole_0(X2)|k2_finseq_1(np__3)!=k2_finseq_1(X1)|~v1_finseq_1(k3_finseq_4(X2,X3,X4,X5))|~v1_funct_1(k3_finseq_4(X2,X3,X4,X5))|~v1_relat_1(k3_finseq_4(X2,X3,X4,X5))|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_9199, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),np__3)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,np__3)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3))|X1!=esk12_0|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5093, negated_conjecture, (v1_xboole_0(X1)|m2_finseq_1(k1_funct_1(esk11_0,np__1),X1)|~m1_subset_1(esk4_0,X1)|~m1_subset_1(X2,X1)|~m1_subset_1(esk2_0,X1)|esk3_0!=X2)).#trainneg
cnf(i_0_2372, negated_conjecture, (k2_finseq_1(np__2)=k2_finseq_1(np__3)|k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))!=np__2|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_1223, negated_conjecture, (k3_finseq_4(X1,esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|v1_xboole_0(X1)|~m1_subset_1(k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(k3_finseq_4(X1,esk8_0,esk9_0,esk10_0),k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))))|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1))).#trainneg
cnf(i_0_9545, plain, (v7_ordinal1(k4_finseq_1(X1))|~v1_xboole_0(k4_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_card_1(k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_9657, plain, (esk16_2(X1,k1_card_1(X2))=X1|~r2_hidden(X1,k9_xtuple_0(X2))|~v7_ordinal1(k1_card_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_9799, negated_conjecture, (m2_finseq_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k2_finseq_1(k11_matrix_2(X1)),k2_finseq_1(k11_matrix_2(X1)))))|v1_xboole_0(X1)|esk3_0!=X2|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k2_finseq_1(k11_matrix_2(X1)),k2_finseq_1(k11_matrix_2(X1)))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k2_finseq_1(k11_matrix_2(X1)),k2_finseq_1(k11_matrix_2(X1)))))|~v3_matrix_2(X1)|~m1_matrix_2(X2,X1))).#trainneg
cnf(i_0_2381, plain, (m1_subset_1(np__2,k4_finseq_1(k3_finseq_4(X1,X2,X3,X4)))|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_8314, plain, (r2_hidden(X1,k2_finseq_1(np__2))|v1_xboole_0(X2)|np__2!=np__3|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(k3_finseq_4(X2,X3,X4,X5))|~v1_funct_1(k3_finseq_4(X2,X3,X4,X5))|~v1_relat_1(k3_finseq_4(X2,X3,X4,X5))|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_8371, plain, (r2_hidden(X1,k2_finseq_1(np__1))|v1_xboole_0(X2)|np__3!=np__1|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(k3_finseq_4(X2,X3,X4,X5))|~v1_funct_1(k3_finseq_4(X2,X3,X4,X5))|~v1_relat_1(k3_finseq_4(X2,X3,X4,X5))|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_1587, plain, (r2_hidden(X1,k2_finseq_1(np__3))|~r1_xxreal_0(X1,k3_finseq_1(np__3))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k3_finseq_1(np__3))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_3693, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X1))|~v7_ordinal1(k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X1))|~m1_matrix_2(k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))),k12_matrix_2(np__3))|~m1_subset_1(k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))),k12_matrix_2(np__3))|~v7_ordinal1(X1)|~m1_matrix_2(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),k12_matrix_2(np__3))|~m1_subset_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),k12_matrix_2(np__3)))).#trainneg
cnf(i_0_3622, negated_conjecture, (k9_xtuple_0(k10_matrix_3(np__3,esk1_0,esk11_0,X1))=k2_finseq_1(k2_xcmplx_0(X2,X3))|k2_xcmplx_0(X2,X3)!=np__3|~v7_ordinal1(X2)|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~m1_subset_1(X3,k5_numbers)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_3882, plain, (k1_nat_1(X1,X2)=k3_finseq_1(k3_finseq_4(X3,X4,X5,X6))|v1_xboole_0(X3)|k2_finseq_1(X7)!=k2_finseq_1(k1_nat_1(X1,X2))|~v7_ordinal1(X2)|~v1_finseq_1(k3_finseq_4(X3,X4,X5,X6))|~v1_funct_1(k3_finseq_4(X3,X4,X5,X6))|~v1_relat_1(k3_finseq_4(X3,X4,X5,X6))|~m1_subset_1(X1,k5_numbers)|X7!=np__3|~m1_subset_1(X7,k5_numbers)|~m1_subset_1(X6,X3)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3))).#trainneg
cnf(i_0_10009, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k10_matrix_3(np__3,esk1_0,esk11_0,k1_funct_1(esk11_0,np__1))),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,esk2_0))|~v7_ordinal1(k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,esk2_0))|~v7_ordinal1(esk2_0)|~m1_matrix_2(k10_matrix_3(np__3,esk1_0,esk11_0,k1_funct_1(esk11_0,np__1)),k12_matrix_2(np__3))|~m1_matrix_2(k1_funct_1(esk11_0,np__1),k12_matrix_2(np__3))|~m1_subset_1(k10_matrix_3(np__3,esk1_0,esk11_0,k1_funct_1(esk11_0,np__1)),k12_matrix_2(np__3))|~m1_subset_1(k1_funct_1(esk11_0,np__1),k12_matrix_2(np__3)))).#trainneg
cnf(i_0_9986, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k10_matrix_3(np__3,esk1_0,esk11_0,k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))))),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X4)|X4!=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X1))|~v7_ordinal1(X4)|~m1_matrix_2(k10_matrix_3(np__3,esk1_0,esk11_0,k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))),k12_matrix_2(np__3))|~m1_subset_1(k10_matrix_3(np__3,esk1_0,esk11_0,k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))),k12_matrix_2(np__3))|~v7_ordinal1(k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X1))|~v7_ordinal1(X1)|~m1_matrix_2(k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))),k12_matrix_2(np__3))|~m1_matrix_2(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),k12_matrix_2(np__3))|~m1_subset_1(k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))),k12_matrix_2(np__3))|~m1_subset_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),k12_matrix_2(np__3)))).#trainneg
cnf(i_0_1639, plain, (v1_xboole_0(X1)|m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0))|~v3_matrix_2(X1)|~m1_matrix_2(X2,X1)|~v1_xboole_0(k2_finseq_1(k11_matrix_2(X1))))).#trainneg
cnf(i_0_9886, plain, (v1_xboole_0(X1)|r2_hidden(X2,k2_finseq_1(np__2))|np__2!=np__3|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,X1)|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_9932, plain, (v1_xboole_0(X1)|r2_hidden(X2,k2_finseq_1(np__1))|np__3!=np__1|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,X1)|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_10011, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k10_matrix_3(np__3,esk1_0,esk11_0,k1_funct_1(esk11_0,np__3))),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,esk8_0))|~v7_ordinal1(k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,esk8_0))|~v7_ordinal1(esk8_0)|~m1_matrix_2(k10_matrix_3(np__3,esk1_0,esk11_0,k1_funct_1(esk11_0,np__3)),k12_matrix_2(np__3))|~m1_matrix_2(k1_funct_1(esk11_0,np__3),k12_matrix_2(np__3))|~m1_subset_1(k10_matrix_3(np__3,esk1_0,esk11_0,k1_funct_1(esk11_0,np__3)),k12_matrix_2(np__3))|~m1_subset_1(k1_funct_1(esk11_0,np__3),k12_matrix_2(np__3)))).#trainneg
cnf(i_0_2377, negated_conjecture, (r2_hidden(X1,k2_finseq_1(np__3))|~r1_xxreal_0(X1,k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2))|~m1_matrix_2(X2,k12_matrix_2(np__3))|~m1_subset_1(X2,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_3571, negated_conjecture, (k3_finseq_4(X1,esk8_0,esk9_0,esk10_0)=k1_funct_1(esk11_0,np__3)|v1_xboole_0(X1)|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1))).#trainneg
cnf(i_0_9843, negated_conjecture, (v1_xboole_0(k12_matrix_2(X1))|m2_finseq_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(k2_zfmisc_1(k2_finseq_1(k11_matrix_2(k12_matrix_2(X1))),k2_finseq_1(k11_matrix_2(k12_matrix_2(X1))))))|esk3_0!=X2|~v3_matrix_2(k12_matrix_2(X1))|~m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(k2_finseq_1(k11_matrix_2(k12_matrix_2(X1))),k2_finseq_1(k11_matrix_2(k12_matrix_2(X1))))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k2_zfmisc_1(k2_finseq_1(k11_matrix_2(k12_matrix_2(X1))),k2_finseq_1(k11_matrix_2(k12_matrix_2(X1))))))|~v7_ordinal1(X1)|~m1_subset_1(X2,k12_matrix_2(X1)))).#trainneg
cnf(i_0_4181, plain, (k2_xcmplx_0(X1,X2)=k3_finseq_1(X3)|k4_finseq_1(X3)!=k2_finseq_1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_1690, plain, (k3_finseq_1(X1)=k1_xboole_0|k9_xtuple_0(X1)!=k2_finseq_1(k1_xboole_0)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_1692, plain, (k9_xtuple_0(X1)=k2_finseq_1(k1_xboole_0)|k3_finseq_1(X1)!=k1_xboole_0|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_10241, negated_conjecture, (r2_hidden(X1,k2_finseq_1(np__3))|~r1_xxreal_0(X1,k1_card_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k1_card_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_matrix_2(X2,k12_matrix_2(np__3))|~m1_subset_1(X2,k12_matrix_2(np__3))|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2)))).#trainneg
cnf(i_0_5379, plain, (v1_xboole_0(X1)|r2_hidden(np__2,k2_finseq_1(k3_finseq_1(X5)))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|k3_finseq_1(X5)!=np__3|~v1_finseq_1(X5)|~v1_funct_1(X5)|~v1_relat_1(X5))).#trainneg
cnf(i_0_7569, plain, (r1_xxreal_0(X1,X2)|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_4350, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X3!=k1_funct_1(X1,X2)|~r2_hidden(X2,k2_finseq_1(np__2))|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_matrix_2(X1,k12_matrix_2(np__3))|k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))!=np__2|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1)))).#trainneg
cnf(i_0_1752, plain, (m1_subset_1(X1,k5_numbers)|v1_xboole_0(k2_finseq_1(np__3))|~m1_subset_1(X1,k2_finseq_1(np__3)))).#trainneg
cnf(i_0_10336, plain, (k3_finseq_1(k1_xboole_0)=k1_xboole_0|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_3795, negated_conjecture, (k1_funct_1(np__3,np__2)=esk9_0|esk11_0!=esk12_0)).#trainneg
cnf(i_0_10338, plain, (k1_xboole_0=np__3|k2_finseq_1(k1_xboole_0)!=k2_finseq_1(np__3)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_10135, plain, (v1_xboole_0(X1)|r2_hidden(X2,k2_finseq_1(np__2))|np__2!=np__3|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_10181, plain, (v1_xboole_0(X1)|r2_hidden(X2,k2_finseq_1(np__1))|np__3!=np__1|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_1694, plain, (k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4))=k2_finseq_1(esk16_2(X5,X6))|v1_xboole_0(X1)|esk16_2(X5,X6)!=np__3|~r2_hidden(X5,a_1_0_finseq_1(X6))|~v7_ordinal1(X6)|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1695, plain, (k9_xtuple_0(X1)=k2_finseq_1(esk16_2(X2,X3))|esk16_2(X2,X3)!=k1_card_1(X1)|~r2_hidden(X2,a_1_0_finseq_1(X3))|~v7_ordinal1(X3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10469, plain, (r2_hidden(np__2,k2_finseq_1(k3_finseq_1(X1)))|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k5_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_10474, plain, (r2_hidden(np__2,k2_finseq_1(k3_finseq_1(X1)))|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_10503, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|~r2_hidden(X1,k9_xtuple_0(X2))|~v7_ordinal1(k3_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_10479, plain, (r2_hidden(np__2,k2_finseq_1(k3_finseq_1(X1)))|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_10461, plain, (r2_hidden(np__2,k2_finseq_1(k3_finseq_1(X1)))|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_10568, plain, (k1_xboole_0=np__3|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|k3_finseq_1(k1_xboole_0)!=np__3)).#trainneg
cnf(i_0_3893, negated_conjecture, (k1_funct_1(np__3,np__1)=esk8_0|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1697, plain, (k9_xtuple_0(X1)=k2_finseq_1(X2)|v1_xboole_0(a_1_0_finseq_1(X3))|X2!=k3_finseq_1(X1)|~r2_hidden(X2,a_1_0_finseq_1(X3))|~v7_ordinal1(X3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,a_1_0_finseq_1(X3)))).#trainneg
cnf(i_0_3632, plain, (v1_funct_1(k2_finseq_1(k1_xboole_0))|~v4_funct_1(k1_zfmisc_1(k5_numbers))|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_10569, plain, (v1_xboole_0(X1)|r2_hidden(X2,k2_finseq_1(np__2))|np__2!=np__3|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_10615, plain, (v1_xboole_0(X1)|r2_hidden(X2,k2_finseq_1(np__1))|np__3!=np__1|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1761, plain, (m1_subset_1(esk16_2(X1,X2),k5_numbers)|~r1_xxreal_0(esk16_2(X1,X2),np__3)|~r1_xxreal_0(np__1,esk16_2(X1,X2))|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_10445, negated_conjecture, (r2_hidden(np__2,k2_finseq_1(k3_finseq_1(X1)))|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3971, negated_conjecture, (k1_funct_1(np__3,np__3)=esk10_0|esk11_0!=esk12_0)).#trainneg
cnf(i_0_10566, plain, (k1_xboole_0=np__3|k2_finseq_1(np__2)!=k2_finseq_1(np__3)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|k3_finseq_1(k1_xboole_0)!=np__2)).#trainneg
cnf(i_0_3744, negated_conjecture, (m2_finseq_1(np__3,u1_struct_0(esk1_0))|esk11_0!=esk12_0)).#trainneg
cnf(i_0_10484, plain, (r2_hidden(np__2,k2_finseq_1(k3_finseq_1(X1)))|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X3,k2_finseq_1(np__3))|~m1_subset_1(X4,k2_finseq_1(np__3))|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_10444, negated_conjecture, (v1_xboole_0(k12_matrix_2(np__3))|r2_hidden(np__2,k2_finseq_1(k3_finseq_1(X1)))|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k12_matrix_2(np__3))|~m1_subset_1(X3,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_4048, negated_conjecture, (k1_funct_1(np__1,np__3)=esk4_0|esk11_0!=esk12_0)).#trainneg
cnf(i_0_524, negated_conjecture, (k3_finseq_4(k1_zfmisc_1(k2_zfmisc_1(k5_numbers,u1_struct_0(esk1_0))),k1_funct_1(esk11_0,np__1),k3_finseq_4(u1_struct_0(esk1_0),esk5_0,esk6_0,esk7_0),k3_finseq_4(X1,esk8_0,esk9_0,esk10_0))=esk11_0|v1_xboole_0(X1)|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1))).#trainneg
cnf(i_0_10567, plain, (k1_xboole_0=np__3|k2_finseq_1(np__3)!=k2_finseq_1(np__1)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|k3_finseq_1(k1_xboole_0)!=np__1)).#trainneg
cnf(i_0_5975, negated_conjecture, (m2_finseq_1(np__2,u1_struct_0(esk1_0))|esk11_0!=esk12_0)).#trainneg
cnf(i_0_9801, negated_conjecture, (m2_finseq_1(k1_funct_1(esk11_0,np__1),k1_zfmisc_1(X1))|esk3_0!=X1|~m1_subset_1(esk4_0,k1_zfmisc_1(X1))|~m1_subset_1(esk2_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4959, plain, (X1=k3_finseq_1(X3)|k4_finseq_1(X3)!=k2_finseq_1(X1)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_3079, negated_conjecture, (v1_xboole_0(X1)|m2_finseq_1(esk12_0,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|np__3!=X4|np__1!=X2|np__2!=X3)).#trainneg
cnf(i_0_10887, plain, (k9_xtuple_0(X1)=k2_finseq_1(X2)|X2!=k3_finseq_1(X1)|~r2_hidden(X2,k2_finseq_1(X3))|~v7_ordinal1(X3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10791, plain, (k9_xtuple_0(X1)=k2_finseq_1(X2)|X2!=k1_card_1(X1)|~r2_hidden(X2,a_1_0_finseq_1(X3))|~v7_ordinal1(X3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,X3)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_10856, plain, (r1_xxreal_0(X1,k3_finseq_1(X2))|~v7_ordinal1(k3_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~r1_xxreal_0(X1,k1_card_1(X2))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k1_card_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_10852, plain, (r1_xxreal_0(np__1,k3_finseq_1(k3_finseq_4(X1,X2,X3,X4)))|v1_xboole_0(X1)|~v7_ordinal1(k3_finseq_1(k3_finseq_4(X1,X2,X3,X4)))|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_10853, plain, (r1_xxreal_0(np__3,k3_finseq_1(k3_finseq_4(X1,X2,X3,X4)))|v1_xboole_0(X1)|~v7_ordinal1(k3_finseq_1(k3_finseq_4(X1,X2,X3,X4)))|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_10851, plain, (r1_xxreal_0(np__2,k3_finseq_1(k3_finseq_4(X1,X2,X3,X4)))|v1_xboole_0(X1)|~v7_ordinal1(k3_finseq_1(k3_finseq_4(X1,X2,X3,X4)))|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_10779, negated_conjecture, (k9_xtuple_0(k10_matrix_3(np__3,esk1_0,esk11_0,X1))=k2_finseq_1(esk16_2(X2,X3))|esk16_2(X2,X3)!=np__3|~r2_hidden(X2,a_1_0_finseq_1(X3))|~v7_ordinal1(X3)|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_10921, plain, (k1_xboole_0=np__3|k3_finseq_1(k1_xboole_0)!=np__2|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|k3_finseq_1(np__3)!=np__2|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_1775, plain, (v3_ordinal1(X1)|v1_xboole_0(X2)|~v3_matrix_2(X2)|~v3_ordinal1(k1_zfmisc_1(k1_xboole_0))|~m1_matrix_2(X1,X2)|~v1_xboole_0(k2_finseq_1(k11_matrix_2(X2))))).#trainneg
cnf(i_0_8954, plain, (k9_xtuple_0(X1)=k2_finseq_1(np__3)|v1_xboole_0(X2)|k1_xboole_0!=np__3|~v1_xboole_0(X1)|~v1_finseq_1(X1)|~v1_finseq_1(k3_finseq_4(X2,X3,X4,X5))|~v1_funct_1(k3_finseq_4(X2,X3,X4,X5))|~v1_relat_1(X1)|~v1_relat_1(k3_finseq_4(X2,X3,X4,X5))|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_1776, plain, (v3_ordinal1(X1)|v1_xboole_0(X2)|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2)|~v1_card_1(k1_zfmisc_1(k2_zfmisc_1(k2_finseq_1(k11_matrix_2(X2)),k2_finseq_1(k11_matrix_2(X2))))))).#trainneg
cnf(i_0_1779, plain, (v1_finset_1(X1)|v1_xboole_0(X2)|~v5_finset_1(k1_zfmisc_1(k1_xboole_0))|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2)|~v1_xboole_0(k2_finseq_1(k11_matrix_2(X2))))).#trainneg
cnf(i_0_10888, plain, (k9_xtuple_0(X1)=k2_finseq_1(X2)|v1_xboole_0(a_1_0_finseq_1(X3))|X2!=k3_finseq_1(X1)|~v7_ordinal1(X3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,a_1_0_finseq_1(X3)))).#trainneg
cnf(i_0_10808, negated_conjecture, (r1_xxreal_0(X1,np__3)|~r2_hidden(X1,k2_finseq_1(np__3)))).#trainneg
cnf(i_0_10804, plain, (r2_hidden(np__2,k2_finseq_1(k3_finseq_1(X1)))|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_11135, plain, (k9_xtuple_0(X1)=k2_finseq_1(X2)|X2!=k1_card_1(X1)|~v7_ordinal1(X3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,X3)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_11387, plain, (k1_xboole_0=np__3|k1_xboole_0!=np__2|k3_finseq_1(np__3)!=np__2|~v1_finseq_1(k1_xboole_0)|~v1_finseq_1(np__3)|~v1_funct_1(k1_xboole_0)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_11396, plain, (k9_xtuple_0(X1)=k2_finseq_1(np__3)|v1_xboole_0(X2)|k1_xboole_0!=np__3|~v1_xboole_0(X1)|~v1_finseq_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X3,X2)|~m1_subset_1(X4,X2))).#trainneg
cnf(i_0_8844, plain, (k9_xtuple_0(np__3)=k2_finseq_1(k1_xboole_0)|k1_xboole_0!=np__3|~v1_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(np__3)|~v1_funct_1(X1)|~v1_relat_1(np__3)|~v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_11383, negated_conjecture, (k9_xtuple_0(k10_matrix_3(np__3,esk1_0,esk11_0,X1))=k2_finseq_1(X2)|X2!=np__3|~r2_hidden(X2,a_1_0_finseq_1(X3))|~v7_ordinal1(X3)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,X3)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_11470, plain, (r2_hidden(np__2,k2_finseq_1(k3_finseq_1(X1)))|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_7758, negated_conjecture, (k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))=np__3|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X1)).#trainneg
cnf(i_0_11034, plain, (k1_xboole_0=np__3|k3_finseq_1(k1_xboole_0)!=np__1|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|k3_finseq_1(np__3)!=np__1|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_6223, plain, (np__2=np__3|k2_finseq_1(np__2)!=k2_finseq_1(np__1)|k1_xboole_0!=np__3|k3_finseq_1(k1_xboole_0)!=np__1|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_1885, plain, (v1_finset_1(X1)|v1_xboole_0(X2)|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2)|~v1_xboole_0(k2_finseq_1(k11_matrix_2(X2))))).#trainneg
cnf(i_0_5946, plain, (k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4))=k2_finseq_1(k1_xboole_0)|v1_xboole_0(X1)|k1_xboole_0!=np__3|~v1_finseq_1(X5)|~v1_funct_1(X5)|~v1_relat_1(X5)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~v1_xboole_0(X5))).#trainneg
cnf(i_0_1886, plain, (v1_finset_1(X1)|v1_xboole_0(X2)|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2)|~v1_xboole_0(k2_zfmisc_1(k2_finseq_1(k11_matrix_2(X2)),k2_finseq_1(k11_matrix_2(X2)))))).#trainneg
cnf(i_0_1887, plain, (v1_finset_1(X1)|v1_xboole_0(X2)|~v3_matrix_2(X2)|~m1_matrix_2(X1,X2)|~m1_subset_1(k2_zfmisc_1(k2_finseq_1(k11_matrix_2(X2)),k2_finseq_1(k11_matrix_2(X2))),k5_numbers))).#trainneg
cnf(i_0_11700, negated_conjecture, (k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))=np__3)).#trainneg
cnf(i_0_11562, plain, (esk16_2(np__2,k3_finseq_1(X1))=np__2|~v7_ordinal1(k3_finseq_1(X1))|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11560, plain, (r1_xxreal_0(np__2,k3_finseq_1(X1))|~v7_ordinal1(k3_finseq_1(X1))|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11590, plain, (r2_hidden(np__2,k9_xtuple_0(X1))|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11583, plain, (r2_hidden(np__2,k2_finseq_1(k1_card_1(X1)))|k1_card_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8156, negated_conjecture, (k3_finseq_1(k3_finseq_4(X1,X2,X3,X4))=k3_finseq_1(X5)|v1_xboole_0(X1)|~v1_finseq_1(X5)|~v1_funct_1(X5)|~v1_relat_1(X5)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|X5!=k1_funct_1(esk11_0,np__3))).#trainneg
cnf(i_0_10924, negated_conjecture, (v4_relat_1(np__3,k5_numbers)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1893, plain, (v1_xboole_0(X1)|m1_subset_1(X2,k2_zfmisc_1(k2_finseq_1(k11_matrix_2(X1)),k2_finseq_1(k11_matrix_2(X1))))|~v3_matrix_2(X1)|~m1_matrix_2(a_1_0_finseq_1(X3),X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,X3)|~v7_ordinal1(X3)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_11037, negated_conjecture, (v4_relat_1(np__2,k5_numbers)|esk11_0!=esk12_0)).#trainneg
cnf(i_0_1894, plain, (v1_xboole_0(X1)|m1_subset_1(X2,k2_zfmisc_1(k2_finseq_1(k11_matrix_2(X1)),k2_finseq_1(k11_matrix_2(X1))))|~v3_matrix_2(X1)|~m1_matrix_2(k9_xtuple_0(X3),X1)|~r1_xxreal_0(X2,k3_finseq_1(X3))|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(k3_finseq_1(X3))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_1895, plain, (v1_xboole_0(X1)|m1_subset_1(X2,k2_zfmisc_1(k2_finseq_1(k11_matrix_2(X1)),k2_finseq_1(k11_matrix_2(X1))))|~v3_matrix_2(X1)|~m1_matrix_2(k2_finseq_1(np__3),X1)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_11785, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),X1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,X2)|X2!=k1_funct_1(esk12_0,X1)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),u1_struct_0(esk1_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_11762, negated_conjecture, (k9_xtuple_0(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))=k2_finseq_1(np__3)|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)))).#trainneg
cnf(i_0_11939, plain, (r2_hidden(np__2,k9_xtuple_0(X1))|k1_card_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2378, negated_conjecture, (k1_nat_1(X1,X2)=k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X3))|k2_finseq_1(k1_nat_1(X1,X2))!=k2_finseq_1(np__3)|~v7_ordinal1(X1)|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X3))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X3))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X3))|~m1_subset_1(X2,k5_numbers)|~m1_matrix_2(X3,k12_matrix_2(np__3))|~m1_subset_1(X3,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_5977, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k1_funct_1(esk11_0,np__2)),np__2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,esk6_0)|~v7_ordinal1(esk6_0)|~m1_matrix_2(k1_funct_1(esk11_0,np__2),k12_matrix_2(np__3))|~m1_subset_1(k1_funct_1(esk11_0,np__2),k12_matrix_2(np__3)))).#trainneg
cnf(i_0_1059, plain, (v3_ordinal1(k1_card_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11704, plain, (k1_xboole_0=np__3|k1_xboole_0!=np__1|k3_finseq_1(np__3)!=np__1|~v1_finseq_1(k1_xboole_0)|~v1_finseq_1(np__3)|~v1_funct_1(k1_xboole_0)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_1896, plain, (v1_xboole_0(X1)|m1_subset_1(X2,k2_zfmisc_1(k2_finseq_1(k11_matrix_2(X1)),k2_finseq_1(k11_matrix_2(X1))))|~v3_matrix_2(X1)|~m1_matrix_2(k2_finseq_1(X3),X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,X3)|~v7_ordinal1(X3)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_11495, plain, (k9_xtuple_0(X1)=k2_finseq_1(np__3)|v1_xboole_0(X2)|k1_xboole_0!=np__3|~v1_xboole_0(X1)|~v1_finseq_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_1897, plain, (v1_xboole_0(X1)|m1_subset_1(X2,k2_zfmisc_1(k2_finseq_1(k11_matrix_2(X1)),k2_finseq_1(k11_matrix_2(X1))))|v1_xboole_0(X3)|~v3_matrix_2(X1)|~m1_matrix_2(X3,X1)|~m1_subset_1(X2,X3))).#trainneg
cnf(i_0_895, negated_conjecture, (X1=k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X2)),k1_tarski(k4_tarski(np__1,X3))),k1_tarski(k4_tarski(np__1,X4)))|np__3!=X4|k1_funct_1(X1,np__1)!=X2|k1_funct_1(X1,np__2)!=X3|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|X1!=esk12_0)).#trainneg
cnf(i_0_8160, plain, (k3_finseq_1(k3_finseq_4(X1,X2,X3,X4))=k3_finseq_1(X5)|v1_xboole_0(X1)|k4_finseq_1(X5)!=k2_finseq_1(np__3)|~v1_finseq_1(X5)|~v1_funct_1(X5)|~v1_relat_1(X5)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_6857, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),np__2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,X2)|X2!=np__2|~v7_ordinal1(X2)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3))|X1!=esk12_0|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2177, plain, (v1_finseq_1(esk15_5(X1,X2,X3,X4,X5))|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k4_tarski(X3,X4),k2_matrix_1(X2))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_8305, plain, (r2_hidden(X1,k2_finseq_1(np__2))|np__2!=np__3|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5979, negated_conjecture, (k1_funct_1(np__2,np__2)=esk6_0|esk11_0!=esk12_0)).#trainneg
cnf(i_0_8820, plain, (np__2=np__3|k9_xtuple_0(np__3)!=k2_finseq_1(np__2)|k1_card_1(X1)!=np__3|~m1_subset_1(k1_card_1(X1),k5_numbers)|~v1_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(np__3)|~v1_funct_1(X1)|~v1_relat_1(np__3)|~v1_relat_1(X1))).#trainneg
cnf(i_0_11840, plain, (r1_xxreal_0(np__2,k1_xboole_0)|k1_xboole_0!=np__3|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_1936, plain, (k2_finseq_1(k3_finseq_1(X1))=k9_xtuple_0(np__3)|k3_finseq_1(X1)!=np__3|~v1_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(np__3)|~v1_funct_1(X1)|~v1_relat_1(np__3)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12135, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),X1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,X2)|X2!=k1_funct_1(esk12_0,X1)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12348, negated_conjecture, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))=k1_funct_1(esk11_0,np__1)|esk2_0!=X1|esk3_0!=X2|np__3!=X3|k1_funct_1(esk11_0,np__1)!=esk12_0)).#trainneg
cnf(i_0_12216, plain, (r2_hidden(np__2,k4_finseq_1(X1))|k1_card_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_12177, plain, (r1_xxreal_0(np__2,k3_finseq_1(X1))|~v7_ordinal1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k1_card_1(X1)!=np__3)).#trainneg
cnf(i_0_10898, plain, (m1_subset_1(esk16_2(X1,np__3),k5_numbers)|~r2_hidden(X1,a_1_0_finseq_1(np__3)))).#trainneg
cnf(i_0_8362, plain, (r2_hidden(X1,k2_finseq_1(np__1))|np__3!=np__1|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6043, negated_conjecture, (k1_funct_1(np__2,np__1)=esk5_0|esk11_0!=esk12_0)).#trainneg
cnf(i_0_5536, plain, (np__2=np__3|k1_xboole_0!=np__3|k3_finseq_1(k1_xboole_0)!=np__2|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_2898, negated_conjecture, (v4_relat_1(k1_funct_1(esk11_0,np__1),k5_numbers))).#trainneg
cnf(i_0_1998, plain, (v3_funct_2(k1_xboole_0,X1,k2_finseq_1(k11_matrix_2(X2)))|v1_xboole_0(X2)|~v2_funct_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|~v3_matrix_2(X2)|~m1_matrix_2(k1_xboole_0,X2))).#trainneg
cnf(i_0_2041, plain, (m1_subset_1(X1,k5_numbers)|v1_xboole_0(k2_finseq_1(k1_xboole_0))|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|~m1_subset_1(X1,k2_finseq_1(k1_xboole_0)))).#trainneg
cnf(i_0_2281, plain, (v1_xboole_0(X1)|~v1_zfmisc_1(X4)|~m1_subset_1(k1_funct_1(X4,np__3),X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|k1_funct_1(X4,np__1)!=X2|k1_funct_1(X4,np__2)!=X3|k3_finseq_1(X4)!=np__3|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4))).#trainneg
cnf(i_0_2241, negated_conjecture, (r2_hidden(X1,k2_finseq_1(X3))|~r1_xxreal_0(X1,k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2)))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2))|X3!=np__3|~m1_matrix_2(X2,k12_matrix_2(np__3))|~m1_subset_1(X2,k12_matrix_2(np__3))|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_8297, plain, (v1_xboole_0(X1)|r2_hidden(X2,k2_finseq_1(np__3))|~r1_xxreal_0(X2,k1_card_1(k3_finseq_4(X1,X3,X4,X5)))|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(k1_card_1(k3_finseq_4(X1,X3,X4,X5)))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~v1_finseq_1(k3_finseq_4(X1,X3,X4,X5))|~v1_funct_1(k3_finseq_4(X1,X3,X4,X5))|~v1_relat_1(k3_finseq_4(X1,X3,X4,X5)))).#trainneg
cnf(i_0_12506, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),X1)=k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2),X1)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k1_funct_1(X2,X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_matrix_2(X2,k12_matrix_2(np__3))|~m1_subset_1(X2,k12_matrix_2(np__3))|k1_funct_1(X2,X1)!=k1_funct_1(esk12_0,X1))).#trainneg
cnf(i_0_2044, plain, (v3_funct_2(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,k2_finseq_1(k11_matrix_2(X2))))),X1,k2_finseq_1(k11_matrix_2(X2)))|v1_xboole_0(X2)|~v2_funct_1(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,k2_finseq_1(k11_matrix_2(X2))))))|~v1_funct_1(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,k2_finseq_1(k11_matrix_2(X2))))))|~v3_matrix_2(X2)|~m1_matrix_2(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,k2_finseq_1(k11_matrix_2(X2))))),X2))).#trainneg
cnf(i_0_2045, plain, (v3_funct_2(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X1,X2)|~v2_funct_1(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v1_funct_1(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))))|~v3_funct_2(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),k5_numbers,X2)|~m2_finseq_1(esk17_1(k1_zfmisc_1(k2_zfmisc_1(X1,X2))),X2))).#trainneg
cnf(i_0_4415, negated_conjecture, (k1_funct_1(X1,np__2)=X2|X1!=k1_funct_1(esk11_0,np__1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|esk2_0!=X3|esk3_0!=X2)).#trainneg
cnf(i_0_1751, plain, (v1_xboole_0(X1)|r2_hidden(np__1,k2_finseq_1(X5))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|X5!=np__3|~m1_subset_1(X5,k5_numbers))).#trainneg
cnf(i_0_12410, plain, (m1_subset_1(X1,k2_finseq_1(np__2))|np__2!=np__3|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12638, plain, (m1_subset_1(X1,k2_finseq_1(np__1))|np__3!=np__1|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_6110, negated_conjecture, (k1_funct_1(np__2,np__3)=esk7_0|esk11_0!=esk12_0)).#trainneg
cnf(i_0_2172, plain, (m1_subset_1(k3_matrix_1(X1,esk15_5(k3_finseq_2(X1),X2,X3,X4,X5),X6,X7),X1)|v1_xboole_0(k2_matrix_1(X2))|~v1_matrix_1(esk15_5(k3_finseq_2(X1),X2,X3,X4,X5))|~v7_ordinal1(X7)|~v7_ordinal1(X6)|X5!=k3_matrix_1(k3_finseq_2(X1),X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(k3_finseq_2(X1)))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k4_tarski(X3,X4),k2_matrix_1(X2))|~m1_subset_1(X5,k3_finseq_2(X1)))).#trainneg
cnf(i_0_3576, negated_conjecture, (v4_relat_1(k1_funct_1(esk11_0,np__3),k5_numbers))).#trainneg
cnf(i_0_2169, plain, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,k1_funct_1(esk15_5(X1,X2,X3,np__3,X4),np__1))),k1_tarski(k4_tarski(np__1,X5))),k1_tarski(k4_tarski(np__1,X4)))=esk15_5(X1,X2,X3,np__3,X4)|k1_funct_1(esk15_5(X1,X2,X3,np__3,X4),np__2)!=X5|k3_finseq_1(esk15_5(X1,X2,X3,np__3,X4))!=np__3|X4!=k3_matrix_1(X1,X2,X3,np__3)|~v1_matrix_1(X2)|~r2_hidden(k4_tarski(X3,np__3),k2_matrix_1(X2))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X3)|~v1_finseq_1(esk15_5(X1,X2,X3,np__3,X4))|~v1_funct_1(esk15_5(X1,X2,X3,np__3,X4))|~v1_relat_1(esk15_5(X1,X2,X3,np__3,X4))|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_2173, plain, (v1_relat_1(esk15_5(X1,X2,X3,X4,X5))|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k4_tarski(X3,X4),k2_matrix_1(X2))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_12821, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X3),X1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,X2)|X2!=k1_funct_1(esk12_0,X1)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|k1_funct_1(X3,X1)!=k1_funct_1(esk12_0,X1)|~v7_ordinal1(k1_funct_1(X3,X1))|~m1_matrix_2(X3,k12_matrix_2(np__3))|~m1_subset_1(X3,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_12822, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X3),X1)=k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2),X1)|k1_funct_1(X2,X1)!=k1_funct_1(esk12_0,X1)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k1_funct_1(X2,X1))|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_matrix_2(X2,k12_matrix_2(np__3))|~m1_subset_1(X2,k12_matrix_2(np__3))|k1_funct_1(X3,X1)!=k1_funct_1(esk12_0,X1)|~v7_ordinal1(k1_funct_1(X3,X1))|~m1_matrix_2(X3,k12_matrix_2(np__3))|~m1_subset_1(X3,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_12819, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),np__3)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,np__3)|k1_funct_1(X1,np__3)!=np__3|~v7_ordinal1(k1_funct_1(X1,np__3))|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_12815, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)|k1_funct_1(X1,np__1)!=np__1|~v7_ordinal1(k1_funct_1(X1,np__1))|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_12816, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X1)|X1!=np__1|~v7_ordinal1(X1)|k1_funct_1(X2,np__1)!=np__1|~v7_ordinal1(k1_funct_1(X2,np__1))|~m1_matrix_2(X2,k12_matrix_2(np__3))|~m1_subset_1(X2,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_12887, negated_conjecture, (r2_hidden(np__1,k2_finseq_1(X1))|X1!=np__3|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k2_finseq_1(np__3))|~m1_subset_1(X3,k2_finseq_1(np__3)))).#trainneg
cnf(i_0_12635, plain, (m1_subset_1(esk16_2(X1,np__3),k5_numbers)|~r2_hidden(X1,k2_finseq_1(np__3)))).#trainneg
cnf(i_0_12637, plain, (m1_subset_1(esk16_2(X1,np__3),k5_numbers)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_12434, plain, (np__2=np__3|k4_finseq_1(np__3)!=k2_finseq_1(np__2)|k1_card_1(X1)!=np__3|~v1_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(np__3)|~v1_funct_1(X1)|~v1_relat_1(np__3)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2174, plain, (v1_xboole_0(esk15_5(X1,X2,X3,X4,X5))|v1_xboole_0(k2_matrix_1(X2))|~v1_xboole_0(X1)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k4_tarski(X3,X4),k2_matrix_1(X2))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_5692, negated_conjecture, (v4_relat_1(k1_funct_1(esk11_0,np__2),k5_numbers))).#trainneg
cnf(i_0_2175, plain, (v4_relat_1(esk15_5(X1,X2,X3,X4,X5),k5_numbers)|v1_xboole_0(k2_matrix_1(X2))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k4_tarski(X3,X4),k2_matrix_1(X2))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_2178, plain, (v1_finset_1(esk15_5(X1,X2,X3,X4,X5))|v1_xboole_0(k2_matrix_1(X2))|~v1_xboole_0(X1)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k4_tarski(X3,X4),k2_matrix_1(X2))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_13000, plain, (k1_funct_1(esk15_5(X1,X2,X3,np__3,X4),np__3)=X4|k1_funct_1(esk15_5(X1,X2,X3,np__3,X4),np__2)!=X5|k3_finseq_1(esk15_5(X1,X2,X3,np__3,X4))!=np__3|X4!=k3_matrix_1(X1,X2,X3,np__3)|~v1_matrix_1(X2)|~r2_hidden(k4_tarski(X3,np__3),k2_matrix_1(X2))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X3)|~v1_finseq_1(esk15_5(X1,X2,X3,np__3,X4))|~v1_funct_1(esk15_5(X1,X2,X3,np__3,X4))|~v1_relat_1(esk15_5(X1,X2,X3,np__3,X4))|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_12798, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),np__1)=np__1|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=esk12_0|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))|k1_funct_1(X1,np__1)!=np__1|~v7_ordinal1(k1_funct_1(X1,np__1))|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_12869, plain, (r2_hidden(np__1,k2_finseq_1(X1))|X1!=np__3|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X5,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_12859, negated_conjecture, (r2_hidden(np__1,k2_finseq_1(X1))|X1!=np__3|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13276, plain, (v3_ordinal1(esk16_2(X1,np__3))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2179, plain, (v1_xboole_0(esk15_5(X1,X2,X3,X4,X5))|m1_subset_1(X6,k2_zfmisc_1(k5_numbers,X1))|v1_xboole_0(k2_matrix_1(X2))|~m1_subset_1(X6,esk15_5(X1,X2,X3,X4,X5))|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k4_tarski(X3,X4),k2_matrix_1(X2))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_12762, negated_conjecture, (k3_finseq_4(X1,X2,X3,k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,np__3))=k10_matrix_3(np__3,esk1_0,esk11_0,X4)|v1_xboole_0(X1)|k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X4),np__1)!=X2|k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X4),np__2)!=X3|k3_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X4))!=np__3|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X4))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X4))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X4))|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,np__3),X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|k1_funct_1(X4,np__3)!=np__3|~v7_ordinal1(k1_funct_1(X4,np__3))|~m1_matrix_2(X4,k12_matrix_2(np__3))|~m1_subset_1(X4,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_12883, plain, (r2_hidden(np__1,k2_finseq_1(X1))|X1!=np__3|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k5_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k5_numbers)))).#trainneg
cnf(i_0_12812, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),np__2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)|k1_funct_1(X1,np__2)!=np__2|~v7_ordinal1(k1_funct_1(X1,np__2))|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_12875, plain, (r2_hidden(np__1,k2_finseq_1(X1))|X1!=np__3|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X3,k5_numbers))).#trainneg
cnf(i_0_8847, plain, (k3_finseq_1(k1_xboole_0)=np__2|k2_finseq_1(np__2)!=k2_finseq_1(np__3)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|k1_xboole_0!=np__3)).#trainneg
cnf(i_0_10522, plain, (m1_subset_1(esk17_1(k2_finseq_1(np__3)),k5_numbers))).#trainneg
cnf(i_0_12888, plain, (r2_hidden(np__1,k2_finseq_1(X1))|X1!=np__3|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k1_numbers)|~m1_subset_1(X3,k1_numbers))).#trainneg
cnf(i_0_12893, plain, (r2_hidden(np__1,k2_finseq_1(X1))|X1!=np__3|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X3,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_10747, plain, (k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4))=k2_finseq_1(X5)|v1_xboole_0(X1)|X5!=np__3|~r2_hidden(X5,a_1_0_finseq_1(X6))|~v7_ordinal1(X6)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X5,X6)|~m2_subset_1(X5,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_1603, negated_conjecture, (k4_finseq_1(X1)=k2_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|X1!=esk12_0)).#trainneg
cnf(i_0_10281, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k1_funct_1(esk11_0,np__3)),np__2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,esk9_0)|v1_xboole_0(X1)|~v7_ordinal1(esk9_0)|~m1_matrix_2(k1_funct_1(esk11_0,np__3),k12_matrix_2(np__3))|~m1_subset_1(k1_funct_1(esk11_0,np__3),k12_matrix_2(np__3))|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk9_0,X1)|~m1_subset_1(esk8_0,X1))).#trainneg
cnf(i_0_8848, plain, (k3_finseq_1(k1_xboole_0)=np__1|k2_finseq_1(np__3)!=k2_finseq_1(np__1)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|k1_xboole_0!=np__3)).#trainneg
cnf(i_0_8335, negated_conjecture, (v1_finset_1(np__2)|~v5_finset_1(k2_finseq_1(np__1))|k3_finseq_1(np__3)!=np__1|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_2303, plain, (k3_finseq_4(X1,X2,X3,X7)=esk15_5(X4,X5,X6,np__3,X7)|v1_xboole_0(X1)|k1_funct_1(esk15_5(X4,X5,X6,np__3,X7),np__1)!=X2|k1_funct_1(esk15_5(X4,X5,X6,np__3,X7),np__2)!=X3|k3_finseq_1(esk15_5(X4,X5,X6,np__3,X7))!=np__3|~v1_finseq_1(esk15_5(X4,X5,X6,np__3,X7))|~v1_funct_1(esk15_5(X4,X5,X6,np__3,X7))|~v1_relat_1(esk15_5(X4,X5,X6,np__3,X7))|~m1_subset_1(X7,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|X7!=k3_matrix_1(X4,X5,X6,np__3)|~v1_matrix_1(X5)|~r2_hidden(k4_tarski(X6,np__3),k2_matrix_1(X5))|~m2_finseq_1(X5,k3_finseq_2(X4))|~v7_ordinal1(X6)|~m1_subset_1(X7,X4))).#trainneg
cnf(i_0_5722, negated_conjecture, (k3_finseq_1(X1)=np__3|X1!=k1_funct_1(esk11_0,np__2)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_5380, plain, (v1_xboole_0(X1)|r2_hidden(np__1,k2_finseq_1(k3_finseq_1(X5)))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|k3_finseq_1(X5)!=np__3|~v1_finseq_1(X5)|~v1_funct_1(X5)|~v1_relat_1(X5))).#trainneg
cnf(i_0_13307, plain, (k1_funct_1(esk15_5(X1,X2,X3,np__3,X4),np__3)=X4|k3_finseq_1(esk15_5(X1,X2,X3,np__3,X4))!=np__3|X4!=k3_matrix_1(X1,X2,X3,np__3)|~v1_matrix_1(X2)|~r2_hidden(k4_tarski(X3,np__3),k2_matrix_1(X2))|~m2_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X3)|~v1_finseq_1(esk15_5(X1,X2,X3,np__3,X4))|~v1_funct_1(esk15_5(X1,X2,X3,np__3,X4))|~v1_relat_1(esk15_5(X1,X2,X3,np__3,X4))|~m1_subset_1(X4,X1))).#trainneg
cnf(i_0_13390, plain, (r2_hidden(np__1,k2_finseq_1(X1))|X1!=np__3|~m1_subset_1(X2,k1_zfmisc_1(k5_numbers))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_13421, plain, (r2_hidden(np__1,k2_finseq_1(X1))|X1!=np__3|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_13350, negated_conjecture, (r2_hidden(np__1,k2_finseq_1(X1))|X1!=np__3|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_13450, plain, (r2_hidden(np__1,k2_finseq_1(X1))|X1!=np__3|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k1_numbers))).#trainneg
cnf(i_0_13367, plain, (v3_ordinal1(X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2439, plain, (m1_subset_1(esk16_2(X1,X2),k5_numbers)|~r1_xxreal_0(np__1,esk16_2(X1,X2))|~r1_xxreal_0(esk16_2(X1,X2),k1_xboole_0)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2))).#trainneg
cnf(i_0_13442, plain, (v3_ordinal1(esk17_1(k2_finseq_1(np__3))))).#trainneg
cnf(i_0_13440, plain, (v1_relat_1(esk17_1(k2_finseq_1(np__3)))|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_8152, plain, (k3_finseq_1(k3_finseq_4(X1,X2,X3,X4))=k3_finseq_1(np__3)|v1_xboole_0(X1)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_10007, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k10_matrix_3(np__3,esk1_0,esk11_0,X3)),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X1))|~v7_ordinal1(k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X1))|~v7_ordinal1(X1)|~m1_matrix_2(k10_matrix_3(np__3,esk1_0,esk11_0,X3),k12_matrix_2(np__3))|~m1_matrix_2(X3,k12_matrix_2(np__3))|~m1_subset_1(k10_matrix_3(np__3,esk1_0,esk11_0,X3),k12_matrix_2(np__3))|~m1_subset_1(X3,k12_matrix_2(np__3))|k1_funct_1(X3,np__1)!=X1|k1_funct_1(X3,np__2)!=X2|k3_finseq_1(X3)!=np__3|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_3422, negated_conjecture, (v1_xboole_0(X1)|m2_finseq_1(esk12_0,X1)|~m1_subset_1(np__3,X1)|~m1_subset_1(X2,X1)|~m1_subset_1(np__1,X1)|np__2!=X2)).#trainneg
cnf(i_0_11133, plain, (k9_xtuple_0(X1)=k2_finseq_1(X2)|X2!=k1_card_1(X1)|~r2_hidden(X2,k2_finseq_1(X3))|~v7_ordinal1(X3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_13845, negated_conjecture, (m2_finseq_1(esk12_0,k1_numbers))).#trainneg
cnf(i_0_13844, negated_conjecture, (m2_finseq_1(esk12_0,k2_finseq_1(np__3)))).#trainneg
cnf(i_0_13898, plain, (k9_xtuple_0(X1)=k2_finseq_1(X2)|X2!=k1_card_1(X1)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13624, plain, (r2_hidden(np__1,k2_finseq_1(X1))|X1!=np__3|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(k3_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_13441, plain, (v1_funct_1(esk17_1(k2_finseq_1(np__3)))|~v4_funct_1(k5_numbers))).#trainneg
cnf(i_0_13625, plain, (r2_hidden(np__1,k2_finseq_1(X1))|X1!=np__3|~m1_subset_1(X1,k5_numbers)|~v7_ordinal1(k1_card_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_13620, plain, (r2_hidden(np__1,k2_finseq_1(X1))|X1!=np__3|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_5705, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k1_funct_1(esk11_0,np__2)),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,esk5_0)|~v7_ordinal1(esk5_0)|~m1_matrix_2(k1_funct_1(esk11_0,np__2),k12_matrix_2(np__3))|~m1_subset_1(k1_funct_1(esk11_0,np__2),k12_matrix_2(np__3)))).#trainneg
cnf(i_0_1961, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,X1),X2)=X3|X3!=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,X2)|~r2_hidden(k4_tarski(X1,X2),k2_matrix_1(esk11_0))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(X1,np__3))).#trainneg
cnf(i_0_2530, plain, (m1_subset_1(np__1,k4_finseq_1(k3_finseq_4(X1,X2,X3,X4)))|v1_xboole_0(X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_3168, negated_conjecture, (k3_finseq_1(k1_xboole_0)=np__3|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|esk12_0!=k1_xboole_0)).#trainneg
cnf(i_0_13443, plain, (v1_finset_1(esk17_1(k2_finseq_1(np__3)))|~v5_finset_1(k5_numbers))).#trainneg
cnf(i_0_3859, plain, (v1_xboole_0(X1)|m2_finseq_1(k3_finseq_4(X2,X3,X4,X5),X1)|v1_xboole_0(X2)|k1_funct_1(k3_finseq_4(X2,X3,X4,X5),np__1)!=X6|k1_funct_1(k3_finseq_4(X2,X3,X4,X5),np__2)!=X7|k3_finseq_1(k3_finseq_4(X2,X3,X4,X5))!=np__3|~v1_finseq_1(k3_finseq_4(X2,X3,X4,X5))|~v1_funct_1(k3_finseq_4(X2,X3,X4,X5))|~v1_relat_1(k3_finseq_4(X2,X3,X4,X5))|~m1_subset_1(X5,X1)|~m1_subset_1(X7,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_4518, negated_conjecture, (r2_hidden(X1,k2_finseq_1(np__3))|~r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k3_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|X2!=k1_funct_1(esk11_0,np__1))).#trainneg
cnf(i_0_7517, plain, (m2_finseq_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X2))),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_5032, negated_conjecture, (m1_subset_1(X1,k4_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X3)))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|X2!=np__3|~m1_matrix_2(X3,k12_matrix_2(np__3))|~m1_subset_1(X3,k12_matrix_2(np__3))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_14218, plain, (m2_finseq_1(k3_finseq_4(X3,X1,X2,X2),k1_zfmisc_1(X2))|v1_xboole_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_13991, plain, (esk16_2(np__1,X1)=np__1|~v7_ordinal1(X1)|X1!=np__3|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_7308, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k2_finseq_1(k1_xboole_0))|~v7_ordinal1(k1_card_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_14041, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,X2)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,X2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13989, plain, (r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)|X1!=np__3|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_14043, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,X1),X2)=X3|X3!=k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|X2!=k1_funct_1(esk12_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_14169, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m2_finseq_1(k3_finseq_4(X1,X3,X4,X5),X2)|k1_funct_1(k3_finseq_4(X1,X3,X4,X5),np__2)!=X6|~m1_subset_1(X5,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(X3,X2)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_14047, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,X1),k1_funct_1(X2,X1))=X3|X3!=k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2),X1)|~r1_xxreal_0(np__1,k1_funct_1(X2,X1))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k1_funct_1(X2,X1),np__3)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(k1_funct_1(X2,X1))|~v7_ordinal1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_matrix_2(X2,k12_matrix_2(np__3))|~m1_subset_1(X2,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_14048, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,X1),k1_funct_1(X2,X1))=X3|X3!=k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2),X1)|~r1_xxreal_0(np__1,k1_funct_1(X2,X1))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k1_funct_1(X2,X1),np__3)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(k1_funct_1(X2,X1))|~v7_ordinal1(X1)|~m1_subset_1(X3,u1_struct_0(esk1_0))|~r2_hidden(X1,k2_finseq_1(np__3))|~m1_matrix_2(X2,k12_matrix_2(np__3))|~m1_subset_1(X2,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_14044, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__2),X1)=X2|X2!=k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X3),np__2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|X1!=np__2|X3!=esk12_0|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3)|~m1_matrix_2(X3,k12_matrix_2(np__3))|~m1_subset_1(X3,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_7310, plain, (m1_subset_1(X1,k5_numbers)|~r2_hidden(X1,k2_finseq_1(k1_xboole_0))|~v7_ordinal1(k3_finseq_1(X2))|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_8836, plain, (k9_xtuple_0(np__3)=k2_finseq_1(k1_xboole_0)|k1_xboole_0!=np__3|~v1_finseq_1(np__3)|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(np__3)|~v1_funct_1(k1_xboole_0)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_2589, plain, (k2_xcmplx_0(X2,X1)=np__3|k2_finseq_1(k2_xcmplx_0(X2,X1))!=k2_finseq_1(np__3)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_14138, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m2_finseq_1(k3_finseq_4(X1,X3,X4,X5),X2)|k1_funct_1(k3_finseq_4(X1,X3,X4,X5),np__2)!=X6|~m1_subset_1(X5,X2)|~m1_subset_1(X6,X2)|~m1_subset_1(k1_funct_1(k3_finseq_4(X1,X3,X4,X5),np__1),X2)|~m1_subset_1(X5,X1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_14391, negated_conjecture, (esk2_0=X1|X1!=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14293, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,X1),X2)=k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),X1),u1_struct_0(esk1_0))|X2!=k1_funct_1(esk12_0,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_14380, negated_conjecture, (esk6_0=X1|X1!=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2901, negated_conjecture, (k1_card_1(k1_funct_1(esk11_0,np__1))=np__3)).#trainneg
cnf(i_0_12432, plain, (np__2=np__3|k2_finseq_1(np__2)!=k2_finseq_1(np__1)|k1_card_1(X1)!=np__3|~v1_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(np__3)|~v1_funct_1(X1)|~v1_relat_1(np__3)|~v1_relat_1(X1)|k3_finseq_1(np__3)!=np__1)).#trainneg
cnf(i_0_14487, plain, (k3_finseq_1(np__3)=k1_xboole_0|~v1_finseq_1(k1_xboole_0)|~v1_finseq_1(np__3)|~v1_funct_1(k1_xboole_0)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|k1_xboole_0!=np__3)).#trainneg
cnf(i_0_13691, plain, (v7_ordinal1(esk17_1(k2_finseq_1(np__3)))|~v1_finset_1(esk17_1(k2_finseq_1(np__3))))).#trainneg
cnf(i_0_8338, negated_conjecture, (v1_relat_1(np__2)|~v4_funct_1(k2_finseq_1(np__1))|k3_finseq_1(np__3)!=np__1|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_2591, plain, (k1_nat_1(X2,X1)=np__3|k2_finseq_1(k1_nat_1(X2,X1))!=k2_finseq_1(np__3)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_14522, negated_conjecture, (v1_xboole_0(X1)|v1_xboole_0(X2)|m2_finseq_1(k1_funct_1(esk11_0,np__2),X1)|esk6_0!=X3|~m1_subset_1(esk5_0,X1)|~m1_subset_1(esk7_0,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(esk7_0,X2)|~m1_subset_1(esk6_0,X2)|~m1_subset_1(esk5_0,X2))).#trainneg
cnf(i_0_5009, plain, (k1_card_1(X1)=np__3|k1_funct_1(X1,np__2)!=X2|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8323, plain, (v3_ordinal1(k2_finseq_1(np__1))|~v3_ordinal1(k1_zfmisc_1(k5_numbers))|k3_finseq_1(np__3)!=np__1|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_8354, negated_conjecture, (v1_funct_1(np__2)|~v4_funct_1(k2_finseq_1(np__1))|k3_finseq_1(np__3)!=np__1|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_8356, negated_conjecture, (v1_relat_1(np__1)|~v4_funct_1(k2_finseq_1(np__1))|k3_finseq_1(np__3)!=np__1|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_2602, plain, (v1_xboole_0(esk17_1(k1_zfmisc_1(X1)))|m1_subset_1(esk17_1(esk17_1(k1_zfmisc_1(X1))),X1))).#trainneg
cnf(i_0_1124, plain, (k9_xtuple_0(np__3)=k2_finseq_1(X1)|X1!=np__3|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_14568, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,X3)|X2!=k1_funct_1(esk12_0,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,X3),u1_struct_0(esk1_0))|X3!=k1_funct_1(esk12_0,X1)|~v7_ordinal1(X3))).#trainneg
cnf(i_0_14574, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)=esk6_0|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14565, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__3),X1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,np__3)|X1!=np__3|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,np__3),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14616, negated_conjecture, (k1_card_1(np__1)=np__3|esk11_0!=esk12_0)).#trainneg
cnf(i_0_14896, plain, (np__2=np__3|k9_xtuple_0(np__3)!=k2_finseq_1(np__2)|X1!=np__3|~m1_subset_1(X1,k5_numbers)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_14665, plain, (v7_ordinal1(esk17_1(k2_finseq_1(np__3))))).#trainneg
cnf(i_0_13719, plain, (k1_card_1(k3_finseq_4(X1,X2,X3,X4))=k3_finseq_1(np__3)|v1_xboole_0(X1)|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_14562, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__1),X1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X2)|X1!=np__1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X2),u1_struct_0(esk1_0))|X2!=np__1|~v7_ordinal1(X2))).#trainneg
cnf(i_0_14569, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,X1),X2)=k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X3),X1)|X2!=k1_funct_1(esk12_0,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X3),X1),u1_struct_0(esk1_0))|k1_funct_1(X3,X1)!=k1_funct_1(esk12_0,X1)|~v7_ordinal1(k1_funct_1(X3,X1))|~m1_matrix_2(X3,k12_matrix_2(np__3))|~m1_subset_1(X3,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_14559, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__1),X1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)|X1!=np__1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14551, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)=esk2_0|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14735, plain, (k1_card_1(X1)=np__3|k3_finseq_1(X1)!=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_8357, negated_conjecture, (v1_funct_1(np__1)|~v4_funct_1(k2_finseq_1(np__1))|k3_finseq_1(np__3)!=np__1|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_8358, negated_conjecture, (v1_finset_1(np__1)|~v5_finset_1(k2_finseq_1(np__1))|k3_finseq_1(np__3)!=np__1|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_12487, plain, (k9_xtuple_0(np__3)=k2_finseq_1(np__1)|v1_xboole_0(X1)|np__3!=np__1|~v1_finseq_1(np__3)|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(np__3)|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(np__3)|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|k2_finseq_1(np__3)!=k2_finseq_1(np__1)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_12489, plain, (k9_xtuple_0(np__3)=k2_finseq_1(np__2)|v1_xboole_0(X1)|np__2!=np__3|~v1_finseq_1(np__3)|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(np__3)|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(np__3)|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|k2_finseq_1(np__2)!=k2_finseq_1(np__3)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_15128, negated_conjecture, (k1_card_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))=np__3|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)))).#trainneg
cnf(i_0_15142, negated_conjecture, (k1_card_1(X1)=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|X1!=k1_funct_1(esk11_0,np__2))).#trainneg
cnf(i_0_15139, negated_conjecture, (k1_card_1(X1)=np__3|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|X1!=k1_funct_1(esk11_0,np__3))).#trainneg
cnf(i_0_2775, plain, (v1_relat_1(X1)|~v4_funct_1(a_1_0_finseq_1(X2))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_2776, plain, (v1_funct_1(X1)|~v4_funct_1(a_1_0_finseq_1(X2))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_14555, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__2),X1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)|X1!=np__2|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5459, negated_conjecture, (m1_subset_1(X1,k9_xtuple_0(k10_matrix_3(np__3,esk1_0,esk11_0,X2)))|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X2))|~m1_matrix_2(X2,k12_matrix_2(np__3))|~m1_subset_1(X2,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_9384, negated_conjecture, (r2_hidden(X1,k2_finseq_1(np__3))|~r1_xxreal_0(X1,k3_finseq_1(X2))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(k3_finseq_1(X2))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~v1_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|k1_card_1(X2)!=np__3)).#trainneg
cnf(i_0_2777, plain, (v3_ordinal1(X1)|~v3_ordinal1(a_1_0_finseq_1(X2))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_14481, plain, (k4_finseq_1(np__3)=k2_finseq_1(k1_xboole_0)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|k1_xboole_0!=np__3|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_15015, plain, (v1_card_1(k3_finseq_1(np__3))|v1_xboole_0(X1)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_15523, plain, (v1_card_1(k3_finseq_1(np__3))|v1_xboole_0(X1)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_2778, plain, (v1_finset_1(X1)|~v5_finset_1(a_1_0_finseq_1(X2))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_15311, negated_conjecture, (k9_xtuple_0(X1)=k2_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|X1!=k1_funct_1(esk11_0,np__2))).#trainneg
cnf(i_0_15300, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|X3!=k1_funct_1(X1,X2)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),u1_struct_0(esk1_0))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_matrix_2(X1,k12_matrix_2(np__3))|k10_matrix_3(np__3,esk1_0,esk11_0,X1)!=k1_funct_1(esk11_0,np__2)|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1)))).#trainneg
cnf(i_0_15401, plain, (v1_relat_1(np__2)|~v4_funct_1(a_1_0_finseq_1(X1))|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15402, plain, (v1_relat_1(np__1)|~v4_funct_1(a_1_0_finseq_1(X1))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14994, plain, (np__2=np__3|k4_finseq_1(np__3)!=k2_finseq_1(np__2)|X1!=np__3|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X1,k5_numbers))).#trainneg
cnf(i_0_9205, negated_conjecture, (k1_xboole_0=np__3|k1_funct_1(esk11_0,np__1)!=k1_xboole_0|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_15012, plain, (v7_ordinal1(k3_finseq_1(np__3))|v1_xboole_0(X1)|~v1_finset_1(k3_finseq_4(X1,X2,X3,X4))|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_15277, negated_conjecture, (k9_xtuple_0(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))=k2_finseq_1(k2_xcmplx_0(X1,X2))|k2_xcmplx_0(X1,X2)!=np__3|~v7_ordinal1(X1)|~v1_finseq_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))|~v1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))|~v1_relat_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0))|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_14557, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__2),X1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,X2)|X1!=np__2|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,X2),u1_struct_0(esk1_0))|X2!=np__2|~v7_ordinal1(X2))).#trainneg
cnf(i_0_15403, plain, (v1_relat_1(np__3)|~v4_funct_1(a_1_0_finseq_1(X1))|~r1_xxreal_0(np__3,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15415, plain, (v1_funct_1(np__2)|~v4_funct_1(a_1_0_finseq_1(X1))|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15578, plain, (v1_card_1(k3_finseq_1(np__3))|v1_xboole_0(X1)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_15720, plain, (v1_xboole_0(X1)|v7_ordinal1(k3_finseq_1(np__3))|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_5557, plain, (v1_xboole_0(X1)|m1_subset_1(np__2,k9_xtuple_0(k3_finseq_4(X1,X2,X3,X4)))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_3535, plain, (k1_card_1(X1)=np__3|k2_finseq_1(k1_card_1(X1))!=k2_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_9639, negated_conjecture, (esk16_2(X1,np__3)=X1|~r2_hidden(X1,k2_finseq_1(np__3)))).#trainneg
cnf(i_0_15416, plain, (v1_funct_1(np__1)|~v4_funct_1(a_1_0_finseq_1(X1))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15417, plain, (v1_funct_1(np__3)|~v4_funct_1(a_1_0_finseq_1(X1))|~r1_xxreal_0(np__3,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15018, plain, (v1_finset_1(k3_finseq_1(np__3))|v1_xboole_0(X1)|~v1_finset_1(k3_finseq_4(X1,X2,X3,X4))|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_10702, plain, (k2_finseq_1(esk16_2(X5,X6))=k4_finseq_1(k3_finseq_4(X1,X2,X3,X4))|v1_xboole_0(X1)|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|esk16_2(X5,X6)!=np__3|~r2_hidden(X5,a_1_0_finseq_1(X6))|~v7_ordinal1(X6)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_16073, negated_conjecture, (esk16_2(X1,np__3)=X1|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_16237, plain, (k2_finseq_1(X1)=k4_finseq_1(k3_finseq_4(X3,X4,X5,X6))|v1_xboole_0(X3)|X1!=np__3|~r2_hidden(X1,a_1_0_finseq_1(X2))|~v7_ordinal1(X2)|~m1_subset_1(X6,X3)|~m1_subset_1(X5,X3)|~m1_subset_1(X4,X3)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_5272, plain, (k2_finseq_1(k1_nat_1(X5,X6))=k4_finseq_1(k3_finseq_4(X1,X2,X3,X4))|v1_xboole_0(X1)|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|k1_nat_1(X5,X6)!=np__3|~v7_ordinal1(X5)|~m1_subset_1(X6,k5_numbers)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_16064, plain, (k1_card_1(X1)=np__3|k9_xtuple_0(X1)!=k2_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_10127, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,k10_matrix_3(np__3,esk1_0,esk11_0,k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,k1_funct_1(X1,np__1))),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))))),np__1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,X4)|X4!=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),np__1))|~v7_ordinal1(k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),np__1))|~v7_ordinal1(X4)|~v7_ordinal1(k1_funct_1(X1,np__1))|~m1_matrix_2(k10_matrix_3(np__3,esk1_0,esk11_0,k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,k1_funct_1(X1,np__1))),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))),k12_matrix_2(np__3))|~m1_matrix_2(k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,k1_funct_1(X1,np__1))),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))),k12_matrix_2(np__3))|~m1_matrix_2(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,k1_funct_1(X1,np__1))),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),k12_matrix_2(np__3))|~m1_subset_1(k10_matrix_3(np__3,esk1_0,esk11_0,k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,k1_funct_1(X1,np__1))),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))))),k12_matrix_2(np__3))|~m1_subset_1(k10_matrix_3(np__3,esk1_0,esk11_0,k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,k1_funct_1(X1,np__1))),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))),k12_matrix_2(np__3))|~m1_subset_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,k1_funct_1(X1,np__1))),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),k12_matrix_2(np__3))|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_15641, plain, (v1_finset_1(np__2)|~v5_finset_1(a_1_0_finseq_1(X1))|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15867, plain, (v1_card_1(k3_finseq_1(np__3))|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_15924, plain, (v1_xboole_0(X1)|v7_ordinal1(k3_finseq_1(np__3))|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_16475, plain, (v1_xboole_0(X1)|v7_ordinal1(k3_finseq_1(np__3))|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_2940, plain, (k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,k1_funct_1(k3_finseq_4(X2,X3,X4,X5),np__1))),k1_tarski(k4_tarski(np__1,X6))),k1_tarski(k4_tarski(np__1,X5)))=k3_finseq_4(X2,X3,X4,X5)|v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_subset_1(X5,X1)|~m1_subset_1(X6,X1)|~m1_subset_1(k1_funct_1(k3_finseq_4(X2,X3,X4,X5),np__1),X1)|k1_funct_1(k3_finseq_4(X2,X3,X4,X5),np__2)!=X6|~m1_subset_1(X5,X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_16374, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X4)),k1_tarski(k4_tarski(np__1,X5))),k1_tarski(k4_tarski(np__1,X6)))!=k1_funct_1(esk11_0,X2)|X1!=k1_funct_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X4)),k1_tarski(k4_tarski(np__1,X5))),k1_tarski(k4_tarski(np__1,X6))),X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X6,u1_struct_0(esk1_0))|~m1_subset_1(X5,u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16426, plain, (k1_card_1(X1)=np__3|k4_finseq_1(X1)!=k2_finseq_1(np__3)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_16381, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|k10_matrix_3(np__3,esk1_0,esk11_0,X4)!=k1_funct_1(esk11_0,X2)|X1!=k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X4),X3)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X4,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_15642, plain, (v1_finset_1(np__1)|~v5_finset_1(a_1_0_finseq_1(X1))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_10535, negated_conjecture, (k1_xboole_0=np__3|k1_funct_1(esk11_0,np__3)!=k1_xboole_0|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0))).#trainneg
cnf(i_0_16080, plain, (v1_finset_1(k3_finseq_1(np__3))|v1_xboole_0(X1)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_16238, negated_conjecture, (X1=np__3|k2_finseq_1(X1)!=k2_finseq_1(np__3)|~r2_hidden(X1,a_1_0_finseq_1(np__3))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_15706, plain, (v1_relat_1(np__2)|~v4_funct_1(k2_finseq_1(X1))|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16695, plain, (v1_finset_1(k3_finseq_1(np__3))|v1_xboole_0(X1)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_16685, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X3)|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X2)|X1!=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X3,X4)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X3,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|X4!=k1_funct_1(esk12_0,X3)|~v7_ordinal1(X4)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_16657, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X3,X2)|k10_matrix_3(np__3,esk1_0,esk11_0,X1)!=k1_funct_1(esk11_0,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(X3,np__3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_16604, negated_conjecture, (k1_funct_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),X4)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X5,X4)|k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3)))!=k1_funct_1(esk11_0,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X4,np__3)|~r1_xxreal_0(X5,np__3)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~m1_subset_1(k1_funct_1(k7_finseq_1(k7_finseq_1(k1_tarski(k4_tarski(np__1,X1)),k1_tarski(k4_tarski(np__1,X2))),k1_tarski(k4_tarski(np__1,X3))),X4),u1_struct_0(esk1_0))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16254, negated_conjecture, (k9_xtuple_0(k10_matrix_3(np__3,esk1_0,esk11_0,X1))=k2_finseq_1(X2)|X2!=np__3|~r2_hidden(X2,a_1_0_finseq_1(np__3))|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_15707, plain, (v1_relat_1(np__1)|~v4_funct_1(k2_finseq_1(X1))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_15865, plain, (v1_relat_1(np__3)|~v4_funct_1(k2_finseq_1(X1))|~r1_xxreal_0(np__3,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16823, plain, (v1_finset_1(k3_finseq_1(np__3))|v1_xboole_0(X1)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_13500, negated_conjecture, (k1_xboole_0=np__3|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|k1_funct_1(esk11_0,np__2)!=k1_xboole_0)).#trainneg
cnf(i_0_15019, plain, (v1_xboole_0(k3_finseq_4(X1,X2,X3,X4))|v1_xboole_0(X1)|~v1_xboole_0(k3_finseq_1(np__3))|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_2960, plain, (k3_finseq_4(X2,X3,X4,k1_funct_1(X5,np__3))=X5|v1_xboole_0(X1)|v1_xboole_0(X2)|k1_funct_1(k3_finseq_4(X2,X3,X4,k1_funct_1(X5,np__3)),np__1)!=k1_funct_1(X5,np__1)|k1_funct_1(X5,np__2)!=X6|k3_finseq_1(X5)!=np__3|~v1_finseq_1(X5)|~v1_funct_1(X5)|~v1_relat_1(X5)|~m1_subset_1(k1_funct_1(X5,np__3),X1)|~m1_subset_1(X6,X1)|~m1_subset_1(k1_funct_1(k3_finseq_4(X2,X3,X4,k1_funct_1(X5,np__3)),np__1),X1)|k1_funct_1(k3_finseq_4(X2,X3,X4,k1_funct_1(X5,np__3)),np__2)!=X6|~m1_subset_1(k1_funct_1(X5,np__3),X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_16980, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X3,X2)|k10_matrix_3(np__3,esk1_0,esk11_0,X1)!=k1_funct_1(esk11_0,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(X3,np__3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,k12_matrix_2(np__3))|k1_funct_1(X1,X2)!=k1_funct_1(esk12_0,X2)|~v7_ordinal1(k1_funct_1(X1,X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_matrix_2(X1,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_16976, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X3),X1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X1)|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X3),X1),u1_struct_0(esk1_0))|k1_funct_1(X3,X1)!=k1_funct_1(esk12_0,X1)|~v7_ordinal1(k1_funct_1(X3,X1))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_matrix_2(X3,k12_matrix_2(np__3))|~m1_subset_1(X3,k12_matrix_2(np__3)))).#trainneg
cnf(i_0_16975, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,X3)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X1)|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,X3),u1_struct_0(esk1_0))|X3!=k1_funct_1(esk12_0,X1)|~v7_ordinal1(X3)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_16928, negated_conjecture, (k1_funct_1(X3,X4)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X5,X4)|X3!=k1_funct_1(esk11_0,X5)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(X4,np__3)|~r1_xxreal_0(X5,np__3)|~v7_ordinal1(X4)|~v7_ordinal1(X5)|~m1_subset_1(k1_funct_1(X3,X4),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(X3,np__3),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|k1_funct_1(X3,np__1)!=X1|k1_funct_1(X3,np__2)!=X2|k3_finseq_1(X3)!=np__3|~v1_finseq_1(X3)|~v1_funct_1(X3)|~v1_relat_1(X3))).#trainneg
cnf(i_0_15866, plain, (v1_funct_1(np__2)|~v4_funct_1(k2_finseq_1(X1))|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16076, plain, (v1_funct_1(np__1)|~v4_funct_1(k2_finseq_1(X1))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_13664, negated_conjecture, (v3_ordinal1(esk12_0)|~r1_xxreal_0(np__1,k12_matrix_2(np__3))|~r1_xxreal_0(k12_matrix_2(np__3),np__3)|~m2_subset_1(k12_matrix_2(np__3),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_13671, negated_conjecture, (v3_ordinal1(esk10_0)|~r1_xxreal_0(np__1,u1_struct_0(esk1_0))|~r1_xxreal_0(u1_struct_0(esk1_0),np__3)|~m2_subset_1(u1_struct_0(esk1_0),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_16885, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__1),X1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X1)|k1_funct_1(esk11_0,np__1)!=k1_funct_1(esk11_0,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k1_funct_1(k1_funct_1(esk11_0,np__1),X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_17232, negated_conjecture, (k1_funct_1(X4,X5)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X6,X5)|v1_xboole_0(X1)|X4!=k1_funct_1(esk11_0,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X5,np__3)|~r1_xxreal_0(X6,np__3)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~m1_subset_1(k1_funct_1(X4,np__2),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(X4,X5),u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(X4,np__3),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(X4,np__3),X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|k1_funct_1(X4,np__1)!=X2|k1_funct_1(X4,np__2)!=X3|k3_finseq_1(X4)!=np__3|~v1_finseq_1(X4)|~v1_funct_1(X4)|~v1_relat_1(X4))).#trainneg
cnf(i_0_17281, negated_conjecture, (k1_funct_1(k3_finseq_4(X1,X2,X3,X4),X5)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X6,X5)|v1_xboole_0(X1)|k3_finseq_4(X1,X2,X3,X4)!=k1_funct_1(esk11_0,X6)|k1_funct_1(k3_finseq_4(X1,X2,X3,X4),np__1)!=X7|k1_funct_1(k3_finseq_4(X1,X2,X3,X4),np__2)!=X8|k3_finseq_1(k3_finseq_4(X1,X2,X3,X4))!=np__3|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X5,np__3)|~r1_xxreal_0(X6,np__3)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(k3_finseq_4(X1,X2,X3,X4),X5),u1_struct_0(esk1_0))|~m1_subset_1(X8,u1_struct_0(esk1_0))|~m1_subset_1(X7,u1_struct_0(esk1_0))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_17352, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,np__2)=esk3_0|k1_funct_1(esk11_0,np__1)!=k1_funct_1(esk11_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_17196, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,np__3)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,np__3)|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,np__3),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16077, plain, (v1_funct_1(np__3)|~v4_funct_1(k2_finseq_1(X1))|~r1_xxreal_0(np__3,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_16462, plain, (v1_finset_1(np__2)|~v5_finset_1(k2_finseq_1(X1))|~r1_xxreal_0(np__2,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_17346, negated_conjecture, (v3_ordinal1(esk12_0)|~r1_xxreal_0(np__1,k12_matrix_2(np__3))|~r1_xxreal_0(k12_matrix_2(np__3),np__3)|~m1_subset_1(k12_matrix_2(np__3),k5_numbers))).#trainneg
cnf(i_0_17347, negated_conjecture, (v3_ordinal1(esk10_0)|~r1_xxreal_0(np__1,u1_struct_0(esk1_0))|~r1_xxreal_0(u1_struct_0(esk1_0),np__3)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_15016, plain, (k1_card_1(k3_finseq_1(np__3))=k3_finseq_1(np__3)|v1_xboole_0(X1)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_17213, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,np__2),X1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X1)|k1_funct_1(esk11_0,np__2)!=k1_funct_1(esk11_0,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m1_subset_1(k1_funct_1(k1_funct_1(esk11_0,np__2),X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_17354, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,np__3)=esk4_0|k1_funct_1(esk11_0,np__1)!=k1_funct_1(esk11_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_17238, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),X1)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X3,X1)|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X3)|X2!=k1_funct_1(esk12_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X3,np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_17235, negated_conjecture, (k1_funct_1(k3_finseq_4(X1,X2,X3,X4),X5)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X6,X5)|v1_xboole_0(X1)|k3_finseq_4(X1,X2,X3,X4)!=k1_funct_1(esk11_0,X6)|~r1_xxreal_0(np__1,X5)|~r1_xxreal_0(np__1,X6)|~r1_xxreal_0(X5,np__3)|~r1_xxreal_0(X6,np__3)|~v7_ordinal1(X5)|~v7_ordinal1(X6)|~m1_subset_1(np__2,u1_struct_0(esk1_0))|~m1_subset_1(k1_funct_1(k3_finseq_4(X1,X2,X3,X4),X5),u1_struct_0(esk1_0))|~m1_subset_1(X4,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X4,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|k3_finseq_4(X1,X2,X3,X4)!=esk12_0|~v1_finseq_1(k3_finseq_4(X1,X2,X3,X4))|~v1_funct_1(k3_finseq_4(X1,X2,X3,X4))|~v1_relat_1(k3_finseq_4(X1,X2,X3,X4)))).#trainneg
cnf(i_0_17452, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,X1),np__2)=X2|X2!=esk3_0|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|k1_funct_1(esk11_0,np__1)!=k1_funct_1(esk11_0,X1))).#trainneg
cnf(i_0_17592, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,X1),np__3)=X2|X2!=esk4_0|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|k1_funct_1(esk11_0,np__1)!=k1_funct_1(esk11_0,X1))).#trainneg
cnf(i_0_17593, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X1)=esk4_0|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X2)|k1_funct_1(esk12_0,X1)!=np__3|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|k1_funct_1(esk11_0,np__1)!=k1_funct_1(esk11_0,X1))).#trainneg
cnf(i_0_16691, plain, (v1_finset_1(np__1)|~v5_finset_1(k2_finseq_1(X1))|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_14992, plain, (np__2=np__3|k2_finseq_1(np__2)!=k2_finseq_1(np__1)|X1!=np__3|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X1,k5_numbers)|k3_finseq_1(np__3)!=np__1)).#trainneg
cnf(i_0_13680, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_17460, negated_conjecture, (esk3_0=esk6_0|k1_funct_1(esk11_0,np__2)!=k1_funct_1(esk11_0,np__1))).#trainneg
cnf(i_0_13672, negated_conjecture, (v3_ordinal1(esk9_0)|~r1_xxreal_0(np__1,u1_struct_0(esk1_0))|~r1_xxreal_0(u1_struct_0(esk1_0),np__3)|~m2_subset_1(u1_struct_0(esk1_0),k1_numbers,k5_numbers))).#trainneg
cnf(i_0_17529, plain, (k1_card_1(k3_finseq_1(np__3))=k3_finseq_1(np__3)|v1_xboole_0(X1)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_2978, plain, (k3_finseq_4(X1,k1_funct_1(X5,np__1),X6,k1_funct_1(X5,np__3))=X5|v1_xboole_0(X1)|v1_xboole_0(X2)|k1_funct_1(X5,np__2)!=X6|~m1_subset_1(k1_funct_1(X5,np__1),X1)|~m1_subset_1(k1_funct_1(X5,np__3),X1)|~m1_subset_1(X6,X1)|~m1_subset_1(k1_funct_1(X5,np__3),X2)|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2)|k1_funct_1(X5,np__1)!=X3|k1_funct_1(X5,np__2)!=X4|k3_finseq_1(X5)!=np__3|~v1_finseq_1(X5)|~v1_funct_1(X5)|~v1_relat_1(X5))).#trainneg
cnf(i_0_17361, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,np__2)=esk9_0|k1_funct_1(esk11_0,np__3)!=k1_funct_1(esk11_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_17351, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,np__3)=np__3|k1_funct_1(esk11_0,np__1)!=k1_funct_1(esk11_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1)|~m1_subset_1(np__3,u1_struct_0(esk1_0))|k1_funct_1(esk11_0,np__1)!=esk12_0)).#trainneg
cnf(i_0_17589, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,np__3)=esk7_0|k1_funct_1(esk11_0,np__2)!=k1_funct_1(esk11_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_17453, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X1)=esk3_0|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X2)|k1_funct_1(esk12_0,X1)!=np__2|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|k1_funct_1(esk11_0,np__1)!=k1_funct_1(esk11_0,X1))).#trainneg
cnf(i_0_17921, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,X1),X2)=esk7_0|X2!=k1_funct_1(esk12_0,X1)|k1_funct_1(esk12_0,X1)!=np__3|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|k1_funct_1(esk11_0,np__2)!=k1_funct_1(esk11_0,X1))).#trainneg
cnf(i_0_17886, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,X1),X2)=esk9_0|X2!=k1_funct_1(esk12_0,X1)|k1_funct_1(esk12_0,X1)!=np__2|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|k1_funct_1(esk11_0,np__3)!=k1_funct_1(esk11_0,X1))).#trainneg
cnf(i_0_17794, negated_conjecture, (esk3_0=esk6_0|k1_funct_1(esk11_0,np__1)!=np__2|esk11_0!=esk12_0)).#trainneg
cnf(i_0_8850, plain, (k2_finseq_1(np__2)=k2_finseq_1(np__3)|k3_finseq_1(k1_xboole_0)!=np__2|~v1_finseq_1(k1_xboole_0)|~v1_funct_1(k1_xboole_0)|k1_xboole_0!=np__3)).#trainneg
cnf(i_0_17897, negated_conjecture, (esk9_0=esk6_0|k1_funct_1(esk11_0,np__2)!=k1_funct_1(esk11_0,np__3))).#trainneg
cnf(i_0_17795, negated_conjecture, (v3_ordinal1(esk9_0)|~r1_xxreal_0(np__1,u1_struct_0(esk1_0))|~r1_xxreal_0(u1_struct_0(esk1_0),np__3)|~m1_subset_1(u1_struct_0(esk1_0),k5_numbers))).#trainneg
cnf(i_0_17798, plain, (k1_card_1(k3_finseq_1(np__3))=k3_finseq_1(np__3)|v1_xboole_0(X1)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_15148, plain, (k9_xtuple_0(np__3)=k2_finseq_1(np__1)|v1_xboole_0(X1)|k2_finseq_1(np__3)!=k2_finseq_1(np__1)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_3151, plain, (m1_subset_1(esk15_5(X1,X2,X3,X4,X5),k1_zfmisc_1(k1_xboole_0))|v1_xboole_0(k2_matrix_1(X2))|~v1_xboole_0(X1)|X5!=k3_matrix_1(X1,X2,X3,X4)|~v1_matrix_1(X2)|~m2_finseq_1(X2,k3_finseq_2(X1))|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~m1_subset_1(k4_tarski(X3,X4),k2_matrix_1(X2))|~m1_subset_1(X5,X1))).#trainneg
cnf(i_0_18000, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,np__3)=esk10_0|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_17708, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X3,X1)=esk4_0|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X3)|X2!=k1_funct_1(esk12_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X3,np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X1)|k1_funct_1(esk11_0,np__1)!=k1_funct_1(esk11_0,X2)|k1_funct_1(esk12_0,X2)!=np__3|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_17705, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),X1)=esk4_0|X2!=k1_funct_1(esk12_0,X1)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X1)|k1_funct_1(esk11_0,np__1)!=k1_funct_1(esk11_0,X2)|k1_funct_1(esk12_0,X2)!=np__3|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_17736, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,np__3)=esk4_0|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1)|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,np__3)|k1_funct_1(esk11_0,np__3)!=k1_funct_1(esk11_0,np__1))).#trainneg
cnf(i_0_17942, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),X1)=esk3_0|X2!=k1_funct_1(esk12_0,X1)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X1)|k1_funct_1(esk11_0,np__1)!=k1_funct_1(esk11_0,X2)|k1_funct_1(esk12_0,X2)!=np__2|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_17936, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X1,np__3)=esk7_0|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1)|k1_funct_1(esk11_0,np__2)!=k1_funct_1(esk11_0,np__3))).#trainneg
cnf(i_0_17924, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X1)=esk7_0|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X2)|k1_funct_1(esk12_0,X1)!=np__3|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|k1_funct_1(esk11_0,np__2)!=k1_funct_1(esk11_0,X1))).#trainneg
cnf(i_0_17986, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X1)=esk10_0|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X2)|k1_funct_1(esk12_0,X1)!=np__3|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|k1_funct_1(esk11_0,np__3)!=k1_funct_1(esk11_0,X1))).#trainneg
cnf(i_0_18023, negated_conjecture, (esk7_0=esk6_0|np__2!=np__3)).#trainneg
cnf(i_0_18120, negated_conjecture, (esk9_0=esk6_0|k1_funct_1(esk11_0,np__3)!=np__2|esk11_0!=esk12_0)).#trainneg
cnf(i_0_3173, plain, (v3_funct_2(k2_zfmisc_1(X1,k2_finseq_1(k11_matrix_2(X2))),X1,k2_finseq_1(k11_matrix_2(X2)))|v1_xboole_0(X2)|~v2_funct_1(k2_zfmisc_1(X1,k2_finseq_1(k11_matrix_2(X2))))|~v1_funct_1(k2_zfmisc_1(X1,k2_finseq_1(k11_matrix_2(X2))))|~v3_matrix_2(X2)|~m1_matrix_2(k2_zfmisc_1(X1,k2_finseq_1(k11_matrix_2(X2))),X2))).#trainneg
cnf(i_0_18065, negated_conjecture, (esk9_0=esk10_0|np__2!=np__3)).#trainneg
cnf(i_0_18026, negated_conjecture, (esk7_0=esk5_0|np__2!=np__3|np__2!=np__1)).#trainneg
cnf(i_0_18124, plain, (k1_card_1(k3_finseq_1(np__3))=k3_finseq_1(np__3)|v1_xboole_0(X1)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_18180, plain, (k9_xtuple_0(np__3)=k2_finseq_1(np__1)|v1_xboole_0(X1)|k2_finseq_1(np__3)!=k2_finseq_1(np__1)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_18295, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X1)=esk10_0|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X2)|k1_funct_1(esk12_0,X1)!=np__3|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X1))).#trainneg
cnf(i_0_18262, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__3,np__3)=esk10_0|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)!=X1|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X2)).#trainneg
cnf(i_0_18280, negated_conjecture, (k1_funct_1(X1,np__3)=esk10_0|X1!=k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)|~v1_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)!=X2|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=X3)).#trainneg
cnf(i_0_18288, negated_conjecture, (v1_xboole_0(X1)|v1_xboole_0(X2)|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk2_0,X1)|~m1_subset_1(esk10_0,X2)|~m1_subset_1(esk6_0,X2)|~m1_subset_1(esk2_0,X2)|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)!=esk2_0|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=esk6_0)).#trainneg
cnf(i_0_18400, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X3,X1)=esk7_0|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X3)|X2!=k1_funct_1(esk12_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X3,np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X1)|k1_funct_1(esk11_0,np__2)!=k1_funct_1(esk11_0,X2)|k1_funct_1(esk12_0,X2)!=np__3|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_17686, negated_conjecture, (k1_funct_1(k1_funct_1(esk11_0,X1),np__2)=esk3_0|k1_funct_1(esk11_0,np__1)!=k1_funct_1(esk11_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(X1,np__3)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_18062, negated_conjecture, (esk9_0=esk8_0|np__2!=np__3|np__3!=np__1)).#trainneg
cnf(i_0_18500, negated_conjecture, (esk6_0=np__3|esk12_0!=np__2|esk11_0!=esk12_0|np__2!=np__3)).#trainneg
cnf(i_0_18073, negated_conjecture, (esk3_0=esk6_0|np__2!=np__1|esk11_0!=esk12_0)).#trainneg
cnf(i_0_18503, negated_conjecture, (esk9_0=esk6_0|np__2!=np__3|esk11_0!=esk12_0)).#trainneg
cnf(i_0_15202, plain, (k9_xtuple_0(np__3)=k2_finseq_1(np__2)|v1_xboole_0(X1)|k2_finseq_1(np__2)!=k2_finseq_1(np__3)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1))).#trainneg
cnf(i_0_18549, plain, (k9_xtuple_0(np__3)=k2_finseq_1(np__1)|v1_xboole_0(X1)|k2_finseq_1(np__3)!=k2_finseq_1(np__1)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3))).#trainneg
cnf(i_0_18397, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0),X1)=esk7_0|X2!=k1_funct_1(esk12_0,X1)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m2_subset_1(X1,k1_numbers,k5_numbers)|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X1)|k1_funct_1(esk11_0,np__2)!=k1_funct_1(esk11_0,X2)|k1_funct_1(esk12_0,X2)!=np__3|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_18287, negated_conjecture, (v1_xboole_0(X1)|~m1_subset_1(esk10_0,X1)|~m1_subset_1(esk6_0,X1)|~m1_subset_1(esk2_0,X1)|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__1,np__1)!=esk2_0|k3_matrix_1(u1_struct_0(esk1_0),esk11_0,np__2,np__2)!=esk6_0)).#trainneg
cnf(i_0_17889, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X2,X1)=esk9_0|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X2)|k1_funct_1(esk12_0,X1)!=np__2|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X2,np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|k1_funct_1(esk11_0,np__3)!=k1_funct_1(esk11_0,X1))).#trainneg
cnf(i_0_17702, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=esk4_0|k1_funct_1(X1,X2)!=k1_funct_1(esk12_0,X2)|X3!=k1_funct_1(esk12_0,X2)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(k1_funct_1(X1,X2))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3))|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X2)|k1_funct_1(esk11_0,np__1)!=k1_funct_1(esk11_0,X3)|k1_funct_1(esk12_0,X3)!=np__3|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_18449, negated_conjecture, (k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X3,X1)=esk10_0|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X3)|X2!=k1_funct_1(esk12_0,X1)|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X1,np__3)|~r1_xxreal_0(X3,np__3)|~v7_ordinal1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers)|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X1)|k1_funct_1(esk11_0,np__3)!=k1_funct_1(esk11_0,X2)|k1_funct_1(esk12_0,X2)!=np__3|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(X2,np__3)|~m2_subset_1(X2,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_17939, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X1),X2)=esk3_0|k1_funct_1(X1,X2)!=k1_funct_1(esk12_0,X2)|X3!=k1_funct_1(esk12_0,X2)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(np__1,X2)|~v7_ordinal1(k1_funct_1(X1,X2))|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_subset_1(X1,k12_matrix_2(np__3))|k10_matrix_3(np__3,esk1_0,esk11_0,esk12_0)!=k1_funct_1(esk11_0,X2)|k1_funct_1(esk11_0,np__1)!=k1_funct_1(esk11_0,X3)|k1_funct_1(esk12_0,X3)!=np__2|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X3,np__3)|~m2_subset_1(X3,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_18545, negated_conjecture, (esk5_0=esk6_0|np__2!=np__3|np__2!=np__1)).#trainneg
cnf(i_0_18734, negated_conjecture, (esk10_0=esk6_0|np__2!=np__3|esk11_0!=esk12_0)).#trainneg
cnf(i_0_18720, negated_conjecture, (esk8_0=esk10_0|np__2!=np__3|np__3!=np__1)).#trainneg
cnf(i_0_18030, negated_conjecture, (esk7_0=esk10_0|k1_funct_1(esk11_0,np__2)!=k1_funct_1(esk11_0,np__3))).#trainneg
cnf(i_0_18737, plain, (k9_xtuple_0(np__3)=k2_finseq_1(np__2)|v1_xboole_0(X1)|k2_finseq_1(np__2)!=k2_finseq_1(np__3)|~v1_finseq_1(np__3)|~v1_funct_1(np__3)|~v1_relat_1(np__3)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_3457, plain, (k9_xtuple_0(k1_card_1(X1))=k2_finseq_1(k3_finseq_1(X2))|k3_finseq_1(X2)!=k1_card_1(X1)|~v1_finseq_1(k1_card_1(X1))|~v1_finseq_1(X2)|~v1_funct_1(k1_card_1(X1))|~v1_funct_1(X2)|~v1_relat_1(k1_card_1(X1))|~v1_relat_1(X2))).#trainneg
cnf(i_0_16977, negated_conjecture, (k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X4),X2)=k3_matrix_1(u1_struct_0(esk1_0),esk11_0,X3,X2)|k10_matrix_3(np__3,esk1_0,esk11_0,X1)!=k1_funct_1(esk11_0,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(X2,np__3)|~r1_xxreal_0(X3,np__3)|~v7_ordinal1(X2)|~v7_ordinal1(X3)|~m1_subset_1(k1_funct_1(k10_matrix_3(np__3,esk1_0,esk11_0,X4),X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,k12_matrix_2(np__3))|k1_funct_1(X1,X2)!=k1_funct_1(esk12_0,X2)|k1_funct_1(X4,X2)!=k1_funct_1(esk12_0,X2)|~v7_ordinal1(k1_funct_1(X1,X2))|~v7_ordinal1(k1_funct_1(X4,X2))|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_matrix_2(X1,k12_matrix_2(np__3))|~m1_matrix_2(X4,k12_matrix_2(np__3))|~m1_subset_1(X4,k12_matrix_2(np__3)))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 135
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 228
# Removed in clause preprocessing      : 12
# Initial clauses in saturation        : 216
# Processed clauses                    : 3270
# ...of these trivial                  : 281
# ...subsumed                          : 1386
# ...remaining for further processing  : 1603
# Other redundant clauses eliminated   : 103
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 101
# Backward-rewritten                   : 150
# Generated clauses                    : 18943
# ...of the previous two non-trivial   : 16718
# Contextual simplify-reflections      : 332
# Paramodulations                      : 18744
# Factorizations                       : 0
# Equation resolutions                 : 199
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1351
#    Positive orientable unit clauses  : 191
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 32
#    Non-unit-clauses                  : 1128
# Current number of unprocessed clauses: 13113
# ...number of literals in the above   : 123262
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1858
# Proof object given clauses           : 211
# Proof search given clauses           : 1603
# Clause-clause subsumption calls (NU) : 976718
# Rec. Clause-clause subsumption calls : 75331
# Non-unit clause-clause subsumptions  : 1681
# Unit Clause-clause subsumption calls : 12730
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 133
# BW rewrite match successes           : 41
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 617867

# -------------------------------------------------
# User time                : 51.856 s
# System time              : 1.280 s
# Total time               : 53.136 s
# Maximum resident set size: 370668 pages
51.85user 1.42system 0:54.52elapsed 97%CPU (0avgtext+0avgdata 370668maxresident)k
0inputs+888outputs (0major+345219minor)pagefaults 0swaps
