2020-05-27 03:16:14.577725: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_75
2020-05-27 03:16:14.635621: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-05-27 03:16:15.170983: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-05-27 03:16:16.504008: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 1926289 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_75' loaded (query=128; context=1024; weigths=real; len_mult=0.000000).
# Preprocessing time       : 2.398 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(dt_m1_alg_1, axiom, ![X1]:(((((~(v2_struct_0(X1))&v2_unialg_1(X1))&v3_unialg_1(X1))&v4_unialg_1(X1))&l1_unialg_1(X1))=>![X2]:(m1_alg_1(X2,X1)=>(((v3_relat_2(X2)&v8_relat_2(X2))&v1_partfun1(X2,u1_struct_0(X1)))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', dt_m1_alg_1)).
fof(dt_k6_alg_1, axiom, ![X1, X2, X3]:(((((((((((((~(v2_struct_0(X1))&v2_unialg_1(X1))&v3_unialg_1(X1))&v4_unialg_1(X1))&l1_unialg_1(X1))&~(v2_struct_0(X2)))&v2_unialg_1(X2))&v3_unialg_1(X2))&v4_unialg_1(X2))&l1_unialg_1(X2))&v1_funct_1(X3))&v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))=>m1_alg_1(k6_alg_1(X1,X2,X3),X1)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', dt_k6_alg_1)).
fof(redefinition_m2_finseq_1, axiom, ![X1, X2]:(m2_finseq_1(X2,X1)<=>m1_finseq_1(X2,X1)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', redefinition_m2_finseq_1)).
fof(dt_u1_unialg_1, axiom, ![X1]:(l1_unialg_1(X1)=>m2_finseq_1(u1_unialg_1(X1),k4_partfun1(k3_finseq_2(u1_struct_0(X1)),u1_struct_0(X1)))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', dt_u1_unialg_1)).
fof(redefinition_k8_eqrel_1, axiom, ![X1, X2]:((((v3_relat_2(X2)&v8_relat_2(X2))&v1_partfun1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))=>k8_eqrel_1(X1,X2)=k7_eqrel_1(X1,X2)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', redefinition_k8_eqrel_1)).
fof(free_g1_unialg_1, axiom, ![X1, X2]:(m1_finseq_1(X2,k4_partfun1(k3_finseq_2(X1),X1))=>![X3, X4]:(g1_unialg_1(X1,X2)=g1_unialg_1(X3,X4)=>(X1=X3&X2=X4))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', free_g1_unialg_1)).
fof(d1_funct_2, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>((X2!=k1_xboole_0=>(v1_funct_2(X3,X1,X2)<=>X1=k1_relset_1(X1,X3)))&(X2=k1_xboole_0=>(v1_funct_2(X3,X1,X2)<=>X3=k1_xboole_0)))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', d1_funct_2)).
fof(dt_k7_alg_1, axiom, ![X1, X2, X3]:(((((((((((((~(v2_struct_0(X1))&v2_unialg_1(X1))&v3_unialg_1(X1))&v4_unialg_1(X1))&l1_unialg_1(X1))&~(v2_struct_0(X2)))&v2_unialg_1(X2))&v3_unialg_1(X2))&v4_unialg_1(X2))&l1_unialg_1(X2))&v1_funct_1(X3))&v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))=>((v1_funct_1(k7_alg_1(X1,X2,X3))&v1_funct_2(k7_alg_1(X1,X2,X3),u1_struct_0(k4_alg_1(X1,k6_alg_1(X1,X2,X3))),u1_struct_0(X2)))&m1_subset_1(k7_alg_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k4_alg_1(X1,k6_alg_1(X1,X2,X3))),u1_struct_0(X2)))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', dt_k7_alg_1)).
fof(cc2_relat_1, axiom, ![X1]:(v1_relat_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_relat_1(X2))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', cc2_relat_1)).
fof(fc6_relat_1, axiom, ![X1, X2]:v1_relat_1(k2_zfmisc_1(X1,X2)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', fc6_relat_1)).
fof(cc2_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>(v4_relat_1(X3,X1)&v5_relat_1(X3,X2))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', cc2_relset_1)).
fof(dt_k8_eqrel_1, axiom, ![X1, X2]:((((v3_relat_2(X2)&v8_relat_2(X2))&v1_partfun1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))=>m1_eqrel_1(k8_eqrel_1(X1,X2),X1)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', dt_k8_eqrel_1)).
fof(abstractness_v1_unialg_1, axiom, ![X1]:(l1_unialg_1(X1)=>(v1_unialg_1(X1)=>X1=g1_unialg_1(u1_struct_0(X1),u1_unialg_1(X1)))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', abstractness_v1_unialg_1)).
fof(redefinition_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>k1_relset_1(X1,X2)=k9_xtuple_0(X2)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', redefinition_k1_relset_1)).
fof(dt_m1_eqrel_1, axiom, ![X1, X2]:(m1_eqrel_1(X2,X1)=>m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', dt_m1_eqrel_1)).
fof(d10_alg_1, axiom, ![X1]:(((((~(v2_struct_0(X1))&v2_unialg_1(X1))&v3_unialg_1(X1))&v4_unialg_1(X1))&l1_unialg_1(X1))=>![X2]:(m1_alg_1(X2,X1)=>k4_alg_1(X1,X2)=g1_unialg_1(k8_eqrel_1(u1_struct_0(X1),X2),k3_alg_1(X1,X2)))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', d10_alg_1)).
fof(d3_funct_1, axiom, ![X1]:((v1_relat_1(X1)&v1_funct_1(X1))=>![X2]:(X2=k10_xtuple_0(X1)<=>![X3]:(r2_hidden(X3,X2)<=>?[X4]:(r2_hidden(X4,k9_xtuple_0(X1))&X3=k1_funct_1(X1,X4))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', d3_funct_1)).
fof(d3_eqrel_1, axiom, ![X1, X2]:((((v3_relat_2(X2)&v8_relat_2(X2))&v1_partfun1(X2,X1))&m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1)))=>(X3=k7_eqrel_1(X1,X2)<=>![X4]:(m1_subset_1(X4,k1_zfmisc_1(X1))=>(r2_hidden(X4,X3)<=>?[X5]:(r2_hidden(X5,X1)&X4=k6_eqrel_1(X1,X1,X2,X5))))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', d3_eqrel_1)).
fof(dt_k4_alg_1, axiom, ![X1, X2]:((((((~(v2_struct_0(X1))&v2_unialg_1(X1))&v3_unialg_1(X1))&v4_unialg_1(X1))&l1_unialg_1(X1))&m1_alg_1(X2,X1))=>(((((~(v2_struct_0(k4_alg_1(X1,X2)))&v1_unialg_1(k4_alg_1(X1,X2)))&v2_unialg_1(k4_alg_1(X1,X2)))&v3_unialg_1(k4_alg_1(X1,X2)))&v4_unialg_1(k4_alg_1(X1,X2)))&l1_unialg_1(k4_alg_1(X1,X2)))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', dt_k4_alg_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', d3_tarski)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', t3_subset)).
fof(fc1_struct_0, axiom, ![X1]:((v2_struct_0(X1)&l1_struct_0(X1))=>v1_xboole_0(u1_struct_0(X1))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', fc1_struct_0)).
fof(dt_l1_unialg_1, axiom, ![X1]:(l1_unialg_1(X1)=>l1_struct_0(X1)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', dt_l1_unialg_1)).
fof(dt_k6_eqrel_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>m1_subset_1(k6_eqrel_1(X1,X2,X3,X4),k1_zfmisc_1(X2))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', dt_k6_eqrel_1)).
fof(d13_alg_1, axiom, ![X1]:(((((~(v2_struct_0(X1))&v2_unialg_1(X1))&v3_unialg_1(X1))&v4_unialg_1(X1))&l1_unialg_1(X1))=>![X2]:(((((~(v2_struct_0(X2))&v2_unialg_1(X2))&v3_unialg_1(X2))&v4_unialg_1(X2))&l1_unialg_1(X2))=>![X3]:(((v1_funct_1(X3)&v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))=>(r1_alg_1(X1,X2,X3)=>![X4]:(((v1_funct_1(X4)&v1_funct_2(X4,u1_struct_0(k4_alg_1(X1,k6_alg_1(X1,X2,X3))),u1_struct_0(X2)))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k4_alg_1(X1,k6_alg_1(X1,X2,X3))),u1_struct_0(X2)))))=>(X4=k7_alg_1(X1,X2,X3)<=>![X5]:(m1_subset_1(X5,u1_struct_0(X1))=>k1_funct_1(X4,k6_eqrel_1(u1_struct_0(X1),u1_struct_0(X1),k6_alg_1(X1,X2,X3),X5))=k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,X5)))))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', d13_alg_1)).
fof(dt_k2_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v5_relat_1(X2,X1))=>m1_subset_1(k2_relset_1(X1,X2),k1_zfmisc_1(X1))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', dt_k2_relset_1)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', t6_boole)).
fof(redefinition_k6_eqrel_1, axiom, ![X1, X2, X3, X4]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>k6_eqrel_1(X1,X2,X3,X4)=k9_relat_1(X3,X4)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', redefinition_k6_eqrel_1)).
fof(redefinition_k2_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v5_relat_1(X2,X1))=>k2_relset_1(X1,X2)=k10_xtuple_0(X2)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', redefinition_k2_relset_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', t1_subset)).
fof(t20_alg_1, conjecture, ![X1]:(((((~(v2_struct_0(X1))&v2_unialg_1(X1))&v3_unialg_1(X1))&v4_unialg_1(X1))&l1_unialg_1(X1))=>![X2]:(((((~(v2_struct_0(X2))&v2_unialg_1(X2))&v3_unialg_1(X2))&v4_unialg_1(X2))&l1_unialg_1(X2))=>![X3]:(((v1_funct_1(X3)&v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))=>(r3_alg_1(X1,X2,X3)=>r4_alg_1(k4_alg_1(X1,k6_alg_1(X1,X2,X3)),X2,k7_alg_1(X1,X2,X3)))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', t20_alg_1)).
fof(d2_alg_1, axiom, ![X1]:(((((~(v2_struct_0(X1))&v2_unialg_1(X1))&v3_unialg_1(X1))&v4_unialg_1(X1))&l1_unialg_1(X1))=>![X2]:(((((~(v2_struct_0(X2))&v2_unialg_1(X2))&v3_unialg_1(X2))&v4_unialg_1(X2))&l1_unialg_1(X2))=>![X3]:(((v1_funct_1(X3)&v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))=>(r2_alg_1(X1,X2,X3)<=>(r1_alg_1(X1,X2,X3)&v2_funct_1(X3)))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', d2_alg_1)).
fof(t19_alg_1, axiom, ![X1]:(((((~(v2_struct_0(X1))&v2_unialg_1(X1))&v3_unialg_1(X1))&v4_unialg_1(X1))&l1_unialg_1(X1))=>![X2]:(((((~(v2_struct_0(X2))&v2_unialg_1(X2))&v3_unialg_1(X2))&v4_unialg_1(X2))&l1_unialg_1(X2))=>![X3]:(((v1_funct_1(X3)&v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))=>(r1_alg_1(X1,X2,X3)=>(r1_alg_1(k4_alg_1(X1,k6_alg_1(X1,X2,X3)),X2,k7_alg_1(X1,X2,X3))&r2_alg_1(k4_alg_1(X1,k6_alg_1(X1,X2,X3)),X2,k7_alg_1(X1,X2,X3))))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', t19_alg_1)).
fof(dt_k1_relset_1, axiom, ![X1, X2]:((v1_relat_1(X2)&v4_relat_1(X2,X1))=>m1_subset_1(k1_relset_1(X1,X2),k1_zfmisc_1(X1))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', dt_k1_relset_1)).
fof(t7_alg_1, axiom, ![X1]:(((((~(v2_struct_0(X1))&v2_unialg_1(X1))&v3_unialg_1(X1))&v4_unialg_1(X1))&l1_unialg_1(X1))=>![X2]:(((((~(v2_struct_0(X2))&v2_unialg_1(X2))&v3_unialg_1(X2))&v4_unialg_1(X2))&l1_unialg_1(X2))=>![X3]:(((v1_funct_1(X3)&v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))=>(r4_alg_1(X1,X2,X3)<=>((r1_alg_1(X1,X2,X3)&k2_relset_1(u1_struct_0(X2),X3)=u1_struct_0(X2))&v2_funct_1(X3)))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', t7_alg_1)).
fof(d3_alg_1, axiom, ![X1]:(((((~(v2_struct_0(X1))&v2_unialg_1(X1))&v3_unialg_1(X1))&v4_unialg_1(X1))&l1_unialg_1(X1))=>![X2]:(((((~(v2_struct_0(X2))&v2_unialg_1(X2))&v3_unialg_1(X2))&v4_unialg_1(X2))&l1_unialg_1(X2))=>![X3]:(((v1_funct_1(X3)&v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))=>(r3_alg_1(X1,X2,X3)<=>(r1_alg_1(X1,X2,X3)&k2_relset_1(u1_struct_0(X2),X3)=u1_struct_0(X2)))))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', d3_alg_1)).
fof(redefinition_k3_funct_2, axiom, ![X1, X2, X3, X4]:(((((~(v1_xboole_0(X1))&v1_funct_1(X3))&v1_funct_2(X3,X1,X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))&m1_subset_1(X4,X1))=>k3_funct_2(X1,X2,X3,X4)=k1_funct_1(X3,X4)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', redefinition_k3_funct_2)).
fof(t7_boole, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', t7_boole)).
fof(rc1_relset_1, axiom, ![X1, X2]:?[X3]:((((m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))&v1_xboole_0(X3))&v1_relat_1(X3))&v4_relat_1(X3,X1))&v5_relat_1(X3,X2)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', rc1_relset_1)).
fof(cc1_relset_1, axiom, ![X1, X2, X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>v1_relat_1(X3)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', cc1_relset_1)).
fof(rc2_funct_2, axiom, ![X1]:?[X2]:(((((((m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))&v1_relat_1(X2))&v4_relat_1(X2,X1))&v5_relat_1(X2,X1))&v1_funct_1(X2))&v1_partfun1(X2,X1))&v1_funct_2(X2,X1,X1))&v3_funct_2(X2,X1,X1)), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', rc2_funct_2)).
fof(cc2_partfun1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&v1_xboole_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))=>~(v1_partfun1(X3,X1)))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', cc2_partfun1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', fc2_struct_0)).
fof(fc11_relat_1, axiom, ![X1]:(v1_xboole_0(X1)=>v1_xboole_0(k10_xtuple_0(X1))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', fc11_relat_1)).
fof(d10_xboole_0, axiom, ![X1, X2]:(X1=X2<=>(r1_tarski(X1,X2)&r1_tarski(X2,X1))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', d10_xboole_0)).
fof(cc4_relset_1, axiom, ![X1, X2]:(v1_xboole_0(X1)=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))=>v1_xboole_0(X3))), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', cc4_relset_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('/home/mptp/local1/bushy_np/hardpred/preds__0.005//t20_alg_1', fc1_xboole_0)).
fof(c_0_47, plain, ![X103, X104]:((((v3_relat_2(X104)|~m1_alg_1(X104,X103)|(v2_struct_0(X103)|~v2_unialg_1(X103)|~v3_unialg_1(X103)|~v4_unialg_1(X103)|~l1_unialg_1(X103)))&(v8_relat_2(X104)|~m1_alg_1(X104,X103)|(v2_struct_0(X103)|~v2_unialg_1(X103)|~v3_unialg_1(X103)|~v4_unialg_1(X103)|~l1_unialg_1(X103))))&(v1_partfun1(X104,u1_struct_0(X103))|~m1_alg_1(X104,X103)|(v2_struct_0(X103)|~v2_unialg_1(X103)|~v3_unialg_1(X103)|~v4_unialg_1(X103)|~l1_unialg_1(X103))))&(m1_subset_1(X104,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X103),u1_struct_0(X103))))|~m1_alg_1(X104,X103)|(v2_struct_0(X103)|~v2_unialg_1(X103)|~v3_unialg_1(X103)|~v4_unialg_1(X103)|~l1_unialg_1(X103)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m1_alg_1])])])])])).
fof(c_0_48, plain, ![X51, X52, X53]:(v2_struct_0(X51)|~v2_unialg_1(X51)|~v3_unialg_1(X51)|~v4_unialg_1(X51)|~l1_unialg_1(X51)|v2_struct_0(X52)|~v2_unialg_1(X52)|~v3_unialg_1(X52)|~v4_unialg_1(X52)|~l1_unialg_1(X52)|~v1_funct_1(X53)|~v1_funct_2(X53,u1_struct_0(X51),u1_struct_0(X52))|~m1_subset_1(X53,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X51),u1_struct_0(X52))))|m1_alg_1(k6_alg_1(X51,X52,X53),X51)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_alg_1])])])).
fof(c_0_49, plain, ![X107, X108]:((~m2_finseq_1(X108,X107)|m1_finseq_1(X108,X107))&(~m1_finseq_1(X108,X107)|m2_finseq_1(X108,X107))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_m2_finseq_1])])).
fof(c_0_50, plain, ![X220]:(~l1_unialg_1(X220)|m2_finseq_1(u1_unialg_1(X220),k4_partfun1(k3_finseq_2(u1_struct_0(X220)),u1_struct_0(X220)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_u1_unialg_1])])).
fof(c_0_51, plain, ![X154, X155]:(~v3_relat_2(X155)|~v8_relat_2(X155)|~v1_partfun1(X155,X154)|~m1_subset_1(X155,k1_zfmisc_1(k2_zfmisc_1(X154,X154)))|k8_eqrel_1(X154,X155)=k7_eqrel_1(X154,X155)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k8_eqrel_1])])).
cnf(c_0_52, plain, (m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2))))|v2_struct_0(X2)|~m1_alg_1(X1,X2)|~v2_unialg_1(X2)|~v3_unialg_1(X2)|~v4_unialg_1(X2)|~l1_unialg_1(X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_53, plain, (v2_struct_0(X1)|v2_struct_0(X2)|m1_alg_1(k6_alg_1(X1,X2,X3),X1)|~v2_unialg_1(X1)|~v3_unialg_1(X1)|~v4_unialg_1(X1)|~l1_unialg_1(X1)|~v2_unialg_1(X2)|~v3_unialg_1(X2)|~v4_unialg_1(X2)|~l1_unialg_1(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_54, plain, ![X198, X199, X200, X201]:((X198=X200|g1_unialg_1(X198,X199)!=g1_unialg_1(X200,X201)|~m1_finseq_1(X199,k4_partfun1(k3_finseq_2(X198),X198)))&(X199=X201|g1_unialg_1(X198,X199)!=g1_unialg_1(X200,X201)|~m1_finseq_1(X199,k4_partfun1(k3_finseq_2(X198),X198)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[free_g1_unialg_1])])])])).
cnf(c_0_55, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_49])).
cnf(c_0_56, plain, (m2_finseq_1(u1_unialg_1(X1),k4_partfun1(k3_finseq_2(u1_struct_0(X1)),u1_struct_0(X1)))|~l1_unialg_1(X1)), inference(split_conjunct,[status(thm)],[c_0_50])).
fof(c_0_57, plain, ![X59, X60, X61]:(((~v1_funct_2(X61,X59,X60)|X59=k1_relset_1(X59,X61)|X60=k1_xboole_0|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(X59,X60))))&(X59!=k1_relset_1(X59,X61)|v1_funct_2(X61,X59,X60)|X60=k1_xboole_0|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(X59,X60)))))&((~v1_funct_2(X61,X59,X60)|X61=k1_xboole_0|X60!=k1_xboole_0|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(X59,X60))))&(X61!=k1_xboole_0|v1_funct_2(X61,X59,X60)|X60!=k1_xboole_0|~m1_subset_1(X61,k1_zfmisc_1(k2_zfmisc_1(X59,X60)))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_funct_2])])])).
fof(c_0_58, plain, ![X38, X39, X40]:(((v1_funct_1(k7_alg_1(X38,X39,X40))|(v2_struct_0(X38)|~v2_unialg_1(X38)|~v3_unialg_1(X38)|~v4_unialg_1(X38)|~l1_unialg_1(X38)|v2_struct_0(X39)|~v2_unialg_1(X39)|~v3_unialg_1(X39)|~v4_unialg_1(X39)|~l1_unialg_1(X39)|~v1_funct_1(X40)|~v1_funct_2(X40,u1_struct_0(X38),u1_struct_0(X39))|~m1_subset_1(X40,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X38),u1_struct_0(X39))))))&(v1_funct_2(k7_alg_1(X38,X39,X40),u1_struct_0(k4_alg_1(X38,k6_alg_1(X38,X39,X40))),u1_struct_0(X39))|(v2_struct_0(X38)|~v2_unialg_1(X38)|~v3_unialg_1(X38)|~v4_unialg_1(X38)|~l1_unialg_1(X38)|v2_struct_0(X39)|~v2_unialg_1(X39)|~v3_unialg_1(X39)|~v4_unialg_1(X39)|~l1_unialg_1(X39)|~v1_funct_1(X40)|~v1_funct_2(X40,u1_struct_0(X38),u1_struct_0(X39))|~m1_subset_1(X40,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X38),u1_struct_0(X39)))))))&(m1_subset_1(k7_alg_1(X38,X39,X40),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k4_alg_1(X38,k6_alg_1(X38,X39,X40))),u1_struct_0(X39))))|(v2_struct_0(X38)|~v2_unialg_1(X38)|~v3_unialg_1(X38)|~v4_unialg_1(X38)|~l1_unialg_1(X38)|v2_struct_0(X39)|~v2_unialg_1(X39)|~v3_unialg_1(X39)|~v4_unialg_1(X39)|~l1_unialg_1(X39)|~v1_funct_1(X40)|~v1_funct_2(X40,u1_struct_0(X38),u1_struct_0(X39))|~m1_subset_1(X40,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X38),u1_struct_0(X39))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_alg_1])])])])).
fof(c_0_59, plain, ![X176, X177]:(~v1_relat_1(X176)|(~m1_subset_1(X177,k1_zfmisc_1(X176))|v1_relat_1(X177))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relat_1])])])).
fof(c_0_60, plain, ![X169, X170]:v1_relat_1(k2_zfmisc_1(X169,X170)), inference(variable_rename,[status(thm)],[fc6_relat_1])).
fof(c_0_61, plain, ![X141, X142, X143]:((v4_relat_1(X143,X141)|~m1_subset_1(X143,k1_zfmisc_1(k2_zfmisc_1(X141,X142))))&(v5_relat_1(X143,X142)|~m1_subset_1(X143,k1_zfmisc_1(k2_zfmisc_1(X141,X142))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_relset_1])])])).
fof(c_0_62, plain, ![X202, X203]:(~v3_relat_2(X203)|~v8_relat_2(X203)|~v1_partfun1(X203,X202)|~m1_subset_1(X203,k1_zfmisc_1(k2_zfmisc_1(X202,X202)))|m1_eqrel_1(k8_eqrel_1(X202,X203),X202)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k8_eqrel_1])])).
cnf(c_0_63, plain, (k8_eqrel_1(X2,X1)=k7_eqrel_1(X2,X1)|~v3_relat_2(X1)|~v8_relat_2(X1)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))), inference(split_conjunct,[status(thm)],[c_0_51])).
cnf(c_0_64, plain, (m1_subset_1(k6_alg_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(X1)|~l1_unialg_1(X2)|~v4_unialg_1(X1)|~v4_unialg_1(X2)|~v3_unialg_1(X1)|~v3_unialg_1(X2)|~v2_unialg_1(X1)|~v2_unialg_1(X2)), inference(spm,[status(thm)],[c_0_52, c_0_53])).
cnf(c_0_65, plain, (X1=X2|g1_unialg_1(X1,X3)!=g1_unialg_1(X2,X4)|~m1_finseq_1(X3,k4_partfun1(k3_finseq_2(X1),X1))), inference(split_conjunct,[status(thm)],[c_0_54])).
cnf(c_0_66, plain, (m1_finseq_1(u1_unialg_1(X1),k4_partfun1(k3_finseq_2(u1_struct_0(X1)),u1_struct_0(X1)))|~l1_unialg_1(X1)), inference(spm,[status(thm)],[c_0_55, c_0_56])).
fof(c_0_67, plain, ![X130]:(~l1_unialg_1(X130)|(~v1_unialg_1(X130)|X130=g1_unialg_1(u1_struct_0(X130),u1_unialg_1(X130)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[abstractness_v1_unialg_1])])).
fof(c_0_68, plain, ![X71, X72]:(~v1_relat_1(X72)|~v4_relat_1(X72,X71)|k1_relset_1(X71,X72)=k9_xtuple_0(X72)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_relset_1])])).
cnf(c_0_69, plain, (X2=k1_relset_1(X2,X1)|X3=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_70, plain, (v1_funct_2(k7_alg_1(X1,X2,X3),u1_struct_0(k4_alg_1(X1,k6_alg_1(X1,X2,X3))),u1_struct_0(X2))|v2_struct_0(X1)|v2_struct_0(X2)|~v2_unialg_1(X1)|~v3_unialg_1(X1)|~v4_unialg_1(X1)|~l1_unialg_1(X1)|~v2_unialg_1(X2)|~v3_unialg_1(X2)|~v4_unialg_1(X2)|~l1_unialg_1(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_71, plain, (m1_subset_1(k7_alg_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k4_alg_1(X1,k6_alg_1(X1,X2,X3))),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X2)|~v2_unialg_1(X1)|~v3_unialg_1(X1)|~v4_unialg_1(X1)|~l1_unialg_1(X1)|~v2_unialg_1(X2)|~v3_unialg_1(X2)|~v4_unialg_1(X2)|~l1_unialg_1(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_72, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_73, plain, (v1_relat_1(k2_zfmisc_1(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_74, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_61])).
fof(c_0_75, plain, ![X119, X120]:(~m1_eqrel_1(X120,X119)|m1_subset_1(X120,k1_zfmisc_1(k1_zfmisc_1(X119)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_m1_eqrel_1])])).
cnf(c_0_76, plain, (m1_eqrel_1(k8_eqrel_1(X2,X1),X2)|~v3_relat_2(X1)|~v8_relat_2(X1)|~v1_partfun1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))), inference(split_conjunct,[status(thm)],[c_0_62])).
cnf(c_0_77, plain, (k8_eqrel_1(u1_struct_0(X1),k6_alg_1(X1,X2,X3))=k7_eqrel_1(u1_struct_0(X1),k6_alg_1(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|~v1_partfun1(k6_alg_1(X1,X2,X3),u1_struct_0(X1))|~v8_relat_2(k6_alg_1(X1,X2,X3))|~v3_relat_2(k6_alg_1(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(X1)|~l1_unialg_1(X2)|~v4_unialg_1(X1)|~v4_unialg_1(X2)|~v3_unialg_1(X1)|~v3_unialg_1(X2)|~v2_unialg_1(X1)|~v2_unialg_1(X2)), inference(spm,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_78, plain, (u1_struct_0(X1)=X2|g1_unialg_1(u1_struct_0(X1),u1_unialg_1(X1))!=g1_unialg_1(X2,X3)|~l1_unialg_1(X1)), inference(spm,[status(thm)],[c_0_65, c_0_66])).
cnf(c_0_79, plain, (X1=g1_unialg_1(u1_struct_0(X1),u1_unialg_1(X1))|~l1_unialg_1(X1)|~v1_unialg_1(X1)), inference(split_conjunct,[status(thm)],[c_0_67])).
fof(c_0_80, plain, ![X85, X86]:(v2_struct_0(X85)|~v2_unialg_1(X85)|~v3_unialg_1(X85)|~v4_unialg_1(X85)|~l1_unialg_1(X85)|(~m1_alg_1(X86,X85)|k4_alg_1(X85,X86)=g1_unialg_1(k8_eqrel_1(u1_struct_0(X85),X86),k3_alg_1(X85,X86)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d10_alg_1])])])])).
fof(c_0_81, plain, ![X41, X42, X43, X45, X46, X47, X49]:((((r2_hidden(esk9_3(X41,X42,X43),k9_xtuple_0(X41))|~r2_hidden(X43,X42)|X42!=k10_xtuple_0(X41)|(~v1_relat_1(X41)|~v1_funct_1(X41)))&(X43=k1_funct_1(X41,esk9_3(X41,X42,X43))|~r2_hidden(X43,X42)|X42!=k10_xtuple_0(X41)|(~v1_relat_1(X41)|~v1_funct_1(X41))))&(~r2_hidden(X46,k9_xtuple_0(X41))|X45!=k1_funct_1(X41,X46)|r2_hidden(X45,X42)|X42!=k10_xtuple_0(X41)|(~v1_relat_1(X41)|~v1_funct_1(X41))))&((~r2_hidden(esk10_2(X41,X47),X47)|(~r2_hidden(X49,k9_xtuple_0(X41))|esk10_2(X41,X47)!=k1_funct_1(X41,X49))|X47=k10_xtuple_0(X41)|(~v1_relat_1(X41)|~v1_funct_1(X41)))&((r2_hidden(esk11_2(X41,X47),k9_xtuple_0(X41))|r2_hidden(esk10_2(X41,X47),X47)|X47=k10_xtuple_0(X41)|(~v1_relat_1(X41)|~v1_funct_1(X41)))&(esk10_2(X41,X47)=k1_funct_1(X41,esk11_2(X41,X47))|r2_hidden(esk10_2(X41,X47),X47)|X47=k10_xtuple_0(X41)|(~v1_relat_1(X41)|~v1_funct_1(X41)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_funct_1])])])])])])).
cnf(c_0_82, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_83, plain, (k1_relset_1(u1_struct_0(k4_alg_1(X1,k6_alg_1(X1,X2,X3))),k7_alg_1(X1,X2,X3))=u1_struct_0(k4_alg_1(X1,k6_alg_1(X1,X2,X3)))|u1_struct_0(X2)=k1_xboole_0|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(X2)|~l1_unialg_1(X1)|~v4_unialg_1(X2)|~v4_unialg_1(X1)|~v3_unialg_1(X2)|~v3_unialg_1(X1)|~v2_unialg_1(X2)|~v2_unialg_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_69, c_0_70]), c_0_71])).
cnf(c_0_84, plain, (v1_relat_1(k7_alg_1(X1,X2,X3))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(X2)|~l1_unialg_1(X1)|~v4_unialg_1(X2)|~v4_unialg_1(X1)|~v3_unialg_1(X2)|~v3_unialg_1(X1)|~v2_unialg_1(X2)|~v2_unialg_1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_72, c_0_71]), c_0_73])])).
cnf(c_0_85, plain, (v4_relat_1(k7_alg_1(X1,X2,X3),u1_struct_0(k4_alg_1(X1,k6_alg_1(X1,X2,X3))))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(X2)|~l1_unialg_1(X1)|~v4_unialg_1(X2)|~v4_unialg_1(X1)|~v3_unialg_1(X2)|~v3_unialg_1(X1)|~v2_unialg_1(X2)|~v2_unialg_1(X1)), inference(spm,[status(thm)],[c_0_74, c_0_71])).
fof(c_0_86, plain, ![X21, X22, X23, X24, X26, X28]:((((r2_hidden(esk5_4(X21,X22,X23,X24),X21)|~r2_hidden(X24,X23)|~m1_subset_1(X24,k1_zfmisc_1(X21))|X23!=k7_eqrel_1(X21,X22)|~m1_subset_1(X23,k1_zfmisc_1(k1_zfmisc_1(X21)))|(~v3_relat_2(X22)|~v8_relat_2(X22)|~v1_partfun1(X22,X21)|~m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(X21,X21)))))&(X24=k6_eqrel_1(X21,X21,X22,esk5_4(X21,X22,X23,X24))|~r2_hidden(X24,X23)|~m1_subset_1(X24,k1_zfmisc_1(X21))|X23!=k7_eqrel_1(X21,X22)|~m1_subset_1(X23,k1_zfmisc_1(k1_zfmisc_1(X21)))|(~v3_relat_2(X22)|~v8_relat_2(X22)|~v1_partfun1(X22,X21)|~m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(X21,X21))))))&(~r2_hidden(X26,X21)|X24!=k6_eqrel_1(X21,X21,X22,X26)|r2_hidden(X24,X23)|~m1_subset_1(X24,k1_zfmisc_1(X21))|X23!=k7_eqrel_1(X21,X22)|~m1_subset_1(X23,k1_zfmisc_1(k1_zfmisc_1(X21)))|(~v3_relat_2(X22)|~v8_relat_2(X22)|~v1_partfun1(X22,X21)|~m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(X21,X21))))))&((m1_subset_1(esk6_3(X21,X22,X23),k1_zfmisc_1(X21))|X23=k7_eqrel_1(X21,X22)|~m1_subset_1(X23,k1_zfmisc_1(k1_zfmisc_1(X21)))|(~v3_relat_2(X22)|~v8_relat_2(X22)|~v1_partfun1(X22,X21)|~m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(X21,X21)))))&((~r2_hidden(esk6_3(X21,X22,X23),X23)|(~r2_hidden(X28,X21)|esk6_3(X21,X22,X23)!=k6_eqrel_1(X21,X21,X22,X28))|X23=k7_eqrel_1(X21,X22)|~m1_subset_1(X23,k1_zfmisc_1(k1_zfmisc_1(X21)))|(~v3_relat_2(X22)|~v8_relat_2(X22)|~v1_partfun1(X22,X21)|~m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(X21,X21)))))&((r2_hidden(esk7_3(X21,X22,X23),X21)|r2_hidden(esk6_3(X21,X22,X23),X23)|X23=k7_eqrel_1(X21,X22)|~m1_subset_1(X23,k1_zfmisc_1(k1_zfmisc_1(X21)))|(~v3_relat_2(X22)|~v8_relat_2(X22)|~v1_partfun1(X22,X21)|~m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(X21,X21)))))&(esk6_3(X21,X22,X23)=k6_eqrel_1(X21,X21,X22,esk7_3(X21,X22,X23))|r2_hidden(esk6_3(X21,X22,X23),X23)|X23=k7_eqrel_1(X21,X22)|~m1_subset_1(X23,k1_zfmisc_1(k1_zfmisc_1(X21)))|(~v3_relat_2(X22)|~v8_relat_2(X22)|~v1_partfun1(X22,X21)|~m1_subset_1(X22,k1_zfmisc_1(k2_zfmisc_1(X21,X21))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_eqrel_1])])])])])).
cnf(c_0_87, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_eqrel_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_88, plain, (m1_eqrel_1(k7_eqrel_1(u1_struct_0(X1),k6_alg_1(X1,X2,X3)),u1_struct_0(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~v1_partfun1(k6_alg_1(X1,X2,X3),u1_struct_0(X1))|~v8_relat_2(k6_alg_1(X1,X2,X3))|~v3_relat_2(k6_alg_1(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(X1)|~l1_unialg_1(X2)|~v4_unialg_1(X1)|~v4_unialg_1(X2)|~v3_unialg_1(X1)|~v3_unialg_1(X2)|~v2_unialg_1(X1)|~v2_unialg_1(X2)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77]), c_0_64])).
cnf(c_0_89, plain, (u1_struct_0(X1)=X2|X1!=g1_unialg_1(X2,X3)|~v1_unialg_1(X1)|~l1_unialg_1(X1)), inference(spm,[status(thm)],[c_0_78, c_0_79])).
cnf(c_0_90, plain, (v2_struct_0(X1)|k4_alg_1(X1,X2)=g1_unialg_1(k8_eqrel_1(u1_struct_0(X1),X2),k3_alg_1(X1,X2))|~v2_unialg_1(X1)|~v3_unialg_1(X1)|~v4_unialg_1(X1)|~l1_unialg_1(X1)|~m1_alg_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_80])).
fof(c_0_91, plain, ![X54, X55]:((((((~v2_struct_0(k4_alg_1(X54,X55))|(v2_struct_0(X54)|~v2_unialg_1(X54)|~v3_unialg_1(X54)|~v4_unialg_1(X54)|~l1_unialg_1(X54)|~m1_alg_1(X55,X54)))&(v1_unialg_1(k4_alg_1(X54,X55))|(v2_struct_0(X54)|~v2_unialg_1(X54)|~v3_unialg_1(X54)|~v4_unialg_1(X54)|~l1_unialg_1(X54)|~m1_alg_1(X55,X54))))&(v2_unialg_1(k4_alg_1(X54,X55))|(v2_struct_0(X54)|~v2_unialg_1(X54)|~v3_unialg_1(X54)|~v4_unialg_1(X54)|~l1_unialg_1(X54)|~m1_alg_1(X55,X54))))&(v3_unialg_1(k4_alg_1(X54,X55))|(v2_struct_0(X54)|~v2_unialg_1(X54)|~v3_unialg_1(X54)|~v4_unialg_1(X54)|~l1_unialg_1(X54)|~m1_alg_1(X55,X54))))&(v4_unialg_1(k4_alg_1(X54,X55))|(v2_struct_0(X54)|~v2_unialg_1(X54)|~v3_unialg_1(X54)|~v4_unialg_1(X54)|~l1_unialg_1(X54)|~m1_alg_1(X55,X54))))&(l1_unialg_1(k4_alg_1(X54,X55))|(v2_struct_0(X54)|~v2_unialg_1(X54)|~v3_unialg_1(X54)|~v4_unialg_1(X54)|~l1_unialg_1(X54)|~m1_alg_1(X55,X54)))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k4_alg_1])])])])).
fof(c_0_92, plain, ![X32, X33, X34, X35, X36]:((~r1_tarski(X32,X33)|(~r2_hidden(X34,X32)|r2_hidden(X34,X33)))&((r2_hidden(esk8_2(X35,X36),X35)|r1_tarski(X35,X36))&(~r2_hidden(esk8_2(X35,X36),X36)|r1_tarski(X35,X36)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_93, plain, ![X79, X80]:((~m1_subset_1(X79,k1_zfmisc_1(X80))|r1_tarski(X79,X80))&(~r1_tarski(X79,X80)|m1_subset_1(X79,k1_zfmisc_1(X80)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_94, plain, (r2_hidden(X3,X4)|~r2_hidden(X1,k9_xtuple_0(X2))|X3!=k1_funct_1(X2,X1)|X4!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_95, plain, (u1_struct_0(k4_alg_1(X1,k6_alg_1(X1,X2,X3)))=k9_xtuple_0(k7_alg_1(X1,X2,X3))|u1_struct_0(X2)=k1_xboole_0|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(X2)|~l1_unialg_1(X1)|~v4_unialg_1(X2)|~v4_unialg_1(X1)|~v3_unialg_1(X2)|~v3_unialg_1(X1)|~v2_unialg_1(X2)|~v2_unialg_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83]), c_0_84]), c_0_85])).
cnf(c_0_96, plain, (v1_funct_1(k7_alg_1(X1,X2,X3))|v2_struct_0(X1)|v2_struct_0(X2)|~v2_unialg_1(X1)|~v3_unialg_1(X1)|~v4_unialg_1(X1)|~l1_unialg_1(X1)|~v2_unialg_1(X2)|~v3_unialg_1(X2)|~v4_unialg_1(X2)|~l1_unialg_1(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_58])).
fof(c_0_97, plain, ![X75]:(~v2_struct_0(X75)|~l1_struct_0(X75)|v1_xboole_0(u1_struct_0(X75))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_struct_0])])).
fof(c_0_98, plain, ![X65]:(~l1_unialg_1(X65)|l1_struct_0(X65)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_unialg_1])])).
cnf(c_0_99, plain, (r2_hidden(X3,X5)|~r2_hidden(X1,X2)|X3!=k6_eqrel_1(X2,X2,X4,X1)|~m1_subset_1(X3,k1_zfmisc_1(X2))|X5!=k7_eqrel_1(X2,X4)|~m1_subset_1(X5,k1_zfmisc_1(k1_zfmisc_1(X2)))|~v3_relat_2(X4)|~v8_relat_2(X4)|~v1_partfun1(X4,X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,X2)))), inference(split_conjunct,[status(thm)],[c_0_86])).
cnf(c_0_100, plain, (m1_subset_1(k7_eqrel_1(u1_struct_0(X1),k6_alg_1(X1,X2,X3)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|v2_struct_0(X1)|v2_struct_0(X2)|~v1_partfun1(k6_alg_1(X1,X2,X3),u1_struct_0(X1))|~v8_relat_2(k6_alg_1(X1,X2,X3))|~v3_relat_2(k6_alg_1(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(X1)|~l1_unialg_1(X2)|~v4_unialg_1(X1)|~v4_unialg_1(X2)|~v3_unialg_1(X1)|~v3_unialg_1(X2)|~v2_unialg_1(X1)|~v2_unialg_1(X2)), inference(spm,[status(thm)],[c_0_87, c_0_88])).
fof(c_0_101, plain, ![X131, X132, X133, X134]:(~m1_subset_1(X133,k1_zfmisc_1(k2_zfmisc_1(X131,X132)))|m1_subset_1(k6_eqrel_1(X131,X132,X133,X134),k1_zfmisc_1(X132))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k6_eqrel_1])])).
cnf(c_0_102, plain, (u1_struct_0(X1)=k8_eqrel_1(u1_struct_0(X2),X3)|v2_struct_0(X2)|X1!=k4_alg_1(X2,X3)|~v1_unialg_1(X1)|~m1_alg_1(X3,X2)|~l1_unialg_1(X1)|~l1_unialg_1(X2)|~v4_unialg_1(X2)|~v3_unialg_1(X2)|~v2_unialg_1(X2)), inference(spm,[status(thm)],[c_0_89, c_0_90])).
cnf(c_0_103, plain, (l1_unialg_1(k4_alg_1(X1,X2))|v2_struct_0(X1)|~v2_unialg_1(X1)|~v3_unialg_1(X1)|~v4_unialg_1(X1)|~l1_unialg_1(X1)|~m1_alg_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_104, plain, (v1_unialg_1(k4_alg_1(X1,X2))|v2_struct_0(X1)|~v2_unialg_1(X1)|~v3_unialg_1(X1)|~v4_unialg_1(X1)|~l1_unialg_1(X1)|~m1_alg_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
fof(c_0_105, plain, ![X9, X10, X11, X12, X13]:((X12!=k7_alg_1(X9,X10,X11)|(~m1_subset_1(X13,u1_struct_0(X9))|k1_funct_1(X12,k6_eqrel_1(u1_struct_0(X9),u1_struct_0(X9),k6_alg_1(X9,X10,X11),X13))=k3_funct_2(u1_struct_0(X9),u1_struct_0(X10),X11,X13))|(~v1_funct_1(X12)|~v1_funct_2(X12,u1_struct_0(k4_alg_1(X9,k6_alg_1(X9,X10,X11))),u1_struct_0(X10))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k4_alg_1(X9,k6_alg_1(X9,X10,X11))),u1_struct_0(X10)))))|~r1_alg_1(X9,X10,X11)|(~v1_funct_1(X11)|~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X10))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X10)))))|(v2_struct_0(X10)|~v2_unialg_1(X10)|~v3_unialg_1(X10)|~v4_unialg_1(X10)|~l1_unialg_1(X10))|(v2_struct_0(X9)|~v2_unialg_1(X9)|~v3_unialg_1(X9)|~v4_unialg_1(X9)|~l1_unialg_1(X9)))&((m1_subset_1(esk4_4(X9,X10,X11,X12),u1_struct_0(X9))|X12=k7_alg_1(X9,X10,X11)|(~v1_funct_1(X12)|~v1_funct_2(X12,u1_struct_0(k4_alg_1(X9,k6_alg_1(X9,X10,X11))),u1_struct_0(X10))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k4_alg_1(X9,k6_alg_1(X9,X10,X11))),u1_struct_0(X10)))))|~r1_alg_1(X9,X10,X11)|(~v1_funct_1(X11)|~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X10))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X10)))))|(v2_struct_0(X10)|~v2_unialg_1(X10)|~v3_unialg_1(X10)|~v4_unialg_1(X10)|~l1_unialg_1(X10))|(v2_struct_0(X9)|~v2_unialg_1(X9)|~v3_unialg_1(X9)|~v4_unialg_1(X9)|~l1_unialg_1(X9)))&(k1_funct_1(X12,k6_eqrel_1(u1_struct_0(X9),u1_struct_0(X9),k6_alg_1(X9,X10,X11),esk4_4(X9,X10,X11,X12)))!=k3_funct_2(u1_struct_0(X9),u1_struct_0(X10),X11,esk4_4(X9,X10,X11,X12))|X12=k7_alg_1(X9,X10,X11)|(~v1_funct_1(X12)|~v1_funct_2(X12,u1_struct_0(k4_alg_1(X9,k6_alg_1(X9,X10,X11))),u1_struct_0(X10))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k4_alg_1(X9,k6_alg_1(X9,X10,X11))),u1_struct_0(X10)))))|~r1_alg_1(X9,X10,X11)|(~v1_funct_1(X11)|~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X10))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X10)))))|(v2_struct_0(X10)|~v2_unialg_1(X10)|~v3_unialg_1(X10)|~v4_unialg_1(X10)|~l1_unialg_1(X10))|(v2_struct_0(X9)|~v2_unialg_1(X9)|~v3_unialg_1(X9)|~v4_unialg_1(X9)|~l1_unialg_1(X9))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_alg_1])])])])])])).
cnf(c_0_106, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_107, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_93])).
fof(c_0_108, plain, ![X96, X97]:(~v1_relat_1(X97)|~v5_relat_1(X97,X96)|m1_subset_1(k2_relset_1(X96,X97),k1_zfmisc_1(X96))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_relset_1])])).
cnf(c_0_109, plain, (u1_struct_0(X1)=k1_xboole_0|r2_hidden(X2,X3)|v2_struct_0(X1)|v2_struct_0(X4)|X2!=k1_funct_1(k7_alg_1(X4,X1,X5),X6)|X3!=k10_xtuple_0(k7_alg_1(X4,X1,X5))|~r2_hidden(X6,u1_struct_0(k4_alg_1(X4,k6_alg_1(X4,X1,X5))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(X5)|~l1_unialg_1(X1)|~l1_unialg_1(X4)|~v4_unialg_1(X1)|~v4_unialg_1(X4)|~v3_unialg_1(X1)|~v3_unialg_1(X4)|~v2_unialg_1(X1)|~v2_unialg_1(X4)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_95]), c_0_96]), c_0_84])).
fof(c_0_110, plain, ![X67]:(~v1_xboole_0(X67)|X67=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
cnf(c_0_111, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_112, plain, (l1_struct_0(X1)|~l1_unialg_1(X1)), inference(split_conjunct,[status(thm)],[c_0_98])).
cnf(c_0_113, plain, (r2_hidden(X1,k7_eqrel_1(u1_struct_0(X2),k6_alg_1(X2,X3,X4)))|v2_struct_0(X3)|v2_struct_0(X2)|k7_eqrel_1(u1_struct_0(X2),k6_alg_1(X2,X3,X4))!=k7_eqrel_1(u1_struct_0(X2),X5)|X1!=k6_eqrel_1(u1_struct_0(X2),u1_struct_0(X2),X5,X6)|~r2_hidden(X6,u1_struct_0(X2))|~v1_partfun1(k6_alg_1(X2,X3,X4),u1_struct_0(X2))|~v1_partfun1(X5,u1_struct_0(X2))|~v8_relat_2(k6_alg_1(X2,X3,X4))|~v8_relat_2(X5)|~v3_relat_2(k6_alg_1(X2,X3,X4))|~v3_relat_2(X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X4)|~l1_unialg_1(X2)|~l1_unialg_1(X3)|~v4_unialg_1(X2)|~v4_unialg_1(X3)|~v3_unialg_1(X2)|~v3_unialg_1(X3)|~v2_unialg_1(X2)|~v2_unialg_1(X3)), inference(spm,[status(thm)],[c_0_99, c_0_100])).
cnf(c_0_114, plain, (m1_subset_1(k6_eqrel_1(X2,X3,X1,X4),k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_101])).
cnf(c_0_115, plain, (k8_eqrel_1(u1_struct_0(X1),X2)=u1_struct_0(k4_alg_1(X1,X2))|v2_struct_0(X1)|~m1_alg_1(X2,X1)|~l1_unialg_1(X1)|~v4_unialg_1(X1)|~v3_unialg_1(X1)|~v2_unialg_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_102]), c_0_103]), c_0_104])).
cnf(c_0_116, plain, (v3_relat_2(X1)|v2_struct_0(X2)|~m1_alg_1(X1,X2)|~v2_unialg_1(X2)|~v3_unialg_1(X2)|~v4_unialg_1(X2)|~l1_unialg_1(X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_117, plain, (v8_relat_2(X1)|v2_struct_0(X2)|~m1_alg_1(X1,X2)|~v2_unialg_1(X2)|~v3_unialg_1(X2)|~v4_unialg_1(X2)|~l1_unialg_1(X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_118, plain, (v1_partfun1(X1,u1_struct_0(X2))|v2_struct_0(X2)|~m1_alg_1(X1,X2)|~v2_unialg_1(X2)|~v3_unialg_1(X2)|~v4_unialg_1(X2)|~l1_unialg_1(X2)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_119, plain, (k1_funct_1(X1,k6_eqrel_1(u1_struct_0(X2),u1_struct_0(X2),k6_alg_1(X2,X3,X4),X5))=k3_funct_2(u1_struct_0(X2),u1_struct_0(X3),X4,X5)|v2_struct_0(X3)|v2_struct_0(X2)|X1!=k7_alg_1(X2,X3,X4)|~m1_subset_1(X5,u1_struct_0(X2))|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(k4_alg_1(X2,k6_alg_1(X2,X3,X4))),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k4_alg_1(X2,k6_alg_1(X2,X3,X4))),u1_struct_0(X3))))|~r1_alg_1(X2,X3,X4)|~v1_funct_1(X4)|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v2_unialg_1(X3)|~v3_unialg_1(X3)|~v4_unialg_1(X3)|~l1_unialg_1(X3)|~v2_unialg_1(X2)|~v3_unialg_1(X2)|~v4_unialg_1(X2)|~l1_unialg_1(X2)), inference(split_conjunct,[status(thm)],[c_0_105])).
fof(c_0_120, plain, ![X90, X91, X92, X93]:(~m1_subset_1(X92,k1_zfmisc_1(k2_zfmisc_1(X90,X91)))|k6_eqrel_1(X90,X91,X92,X93)=k9_relat_1(X92,X93)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k6_eqrel_1])])).
cnf(c_0_121, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(spm,[status(thm)],[c_0_106, c_0_107])).
cnf(c_0_122, plain, (m1_subset_1(k2_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_108])).
fof(c_0_123, plain, ![X73, X74]:(~v1_relat_1(X74)|~v5_relat_1(X74,X73)|k2_relset_1(X73,X74)=k10_xtuple_0(X74)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k2_relset_1])])).
cnf(c_0_124, plain, (u1_struct_0(X1)=k1_xboole_0|r2_hidden(X2,X3)|v2_struct_0(X4)|v2_struct_0(X1)|X2!=k1_funct_1(k7_alg_1(X4,X1,X5),X6)|X3!=k10_xtuple_0(k7_alg_1(X4,X1,X5))|~r2_hidden(X6,k9_xtuple_0(k7_alg_1(X4,X1,X5)))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(X5)|~l1_unialg_1(X1)|~l1_unialg_1(X4)|~v4_unialg_1(X1)|~v4_unialg_1(X4)|~v3_unialg_1(X1)|~v3_unialg_1(X4)|~v2_unialg_1(X1)|~v2_unialg_1(X4)), inference(spm,[status(thm)],[c_0_109, c_0_95])).
cnf(c_0_125, plain, (r2_hidden(esk9_3(X1,X2,X3),k9_xtuple_0(X1))|~r2_hidden(X3,X2)|X2!=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v1_funct_1(X1)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_126, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_110])).
cnf(c_0_127, plain, (v1_xboole_0(u1_struct_0(X1))|~l1_unialg_1(X1)|~v2_struct_0(X1)), inference(spm,[status(thm)],[c_0_111, c_0_112])).
cnf(c_0_128, plain, (r2_hidden(k6_eqrel_1(u1_struct_0(X1),u1_struct_0(X1),X2,X3),k7_eqrel_1(u1_struct_0(X1),k6_alg_1(X1,X4,X5)))|v2_struct_0(X1)|v2_struct_0(X4)|k7_eqrel_1(u1_struct_0(X1),k6_alg_1(X1,X4,X5))!=k7_eqrel_1(u1_struct_0(X1),X2)|~r2_hidden(X3,u1_struct_0(X1))|~v1_partfun1(k6_alg_1(X1,X4,X5),u1_struct_0(X1))|~v1_partfun1(X2,u1_struct_0(X1))|~v8_relat_2(k6_alg_1(X1,X4,X5))|~v8_relat_2(X2)|~v3_relat_2(k6_alg_1(X1,X4,X5))|~v3_relat_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X4))))|~v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X4))|~v1_funct_1(X5)|~l1_unialg_1(X1)|~l1_unialg_1(X4)|~v4_unialg_1(X1)|~v4_unialg_1(X4)|~v3_unialg_1(X1)|~v3_unialg_1(X4)|~v2_unialg_1(X1)|~v2_unialg_1(X4)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_113]), c_0_114])).
cnf(c_0_129, plain, (k7_eqrel_1(u1_struct_0(X1),k6_alg_1(X1,X2,X3))=u1_struct_0(k4_alg_1(X1,k6_alg_1(X1,X2,X3)))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(X1)|~l1_unialg_1(X2)|~v4_unialg_1(X1)|~v4_unialg_1(X2)|~v3_unialg_1(X1)|~v3_unialg_1(X2)|~v2_unialg_1(X1)|~v2_unialg_1(X2)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_77, c_0_115]), c_0_116]), c_0_117]), c_0_118]), c_0_53])).
cnf(c_0_130, plain, (v3_relat_2(k6_alg_1(X1,X2,X3))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(X1)|~l1_unialg_1(X2)|~v4_unialg_1(X1)|~v4_unialg_1(X2)|~v3_unialg_1(X1)|~v3_unialg_1(X2)|~v2_unialg_1(X1)|~v2_unialg_1(X2)), inference(spm,[status(thm)],[c_0_116, c_0_53])).
cnf(c_0_131, plain, (v8_relat_2(k6_alg_1(X1,X2,X3))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(X1)|~l1_unialg_1(X2)|~v4_unialg_1(X1)|~v4_unialg_1(X2)|~v3_unialg_1(X1)|~v3_unialg_1(X2)|~v2_unialg_1(X1)|~v2_unialg_1(X2)), inference(spm,[status(thm)],[c_0_117, c_0_53])).
cnf(c_0_132, plain, (v1_partfun1(k6_alg_1(X1,X2,X3),u1_struct_0(X1))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(X1)|~l1_unialg_1(X2)|~v4_unialg_1(X1)|~v4_unialg_1(X2)|~v3_unialg_1(X1)|~v3_unialg_1(X2)|~v2_unialg_1(X1)|~v2_unialg_1(X2)), inference(spm,[status(thm)],[c_0_118, c_0_53])).
cnf(c_0_133, plain, (k1_funct_1(k7_alg_1(X1,X2,X3),k6_eqrel_1(u1_struct_0(X1),u1_struct_0(X1),k6_alg_1(X1,X2,X3),X4))=k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,X4)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_alg_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(X2)|~l1_unialg_1(X1)|~v4_unialg_1(X2)|~v4_unialg_1(X1)|~v3_unialg_1(X2)|~v3_unialg_1(X1)|~v2_unialg_1(X2)|~v2_unialg_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_70]), c_0_96]), c_0_71])).
cnf(c_0_134, plain, (k6_eqrel_1(X2,X3,X1,X4)=k9_relat_1(X1,X4)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_135, plain, (r2_hidden(X1,X2)|~v5_relat_1(X3,X2)|~v1_relat_1(X3)|~r2_hidden(X1,k2_relset_1(X2,X3))), inference(spm,[status(thm)],[c_0_121, c_0_122])).
cnf(c_0_136, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_123])).
cnf(c_0_137, plain, (u1_struct_0(X1)=k1_xboole_0|r2_hidden(X2,X3)|v2_struct_0(X1)|v2_struct_0(X4)|X2!=k1_funct_1(k7_alg_1(X4,X1,X5),esk9_3(k7_alg_1(X4,X1,X5),X6,X7))|X3!=k10_xtuple_0(k7_alg_1(X4,X1,X5))|X6!=k10_xtuple_0(k7_alg_1(X4,X1,X5))|~r2_hidden(X7,X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(X5)|~l1_unialg_1(X1)|~l1_unialg_1(X4)|~v4_unialg_1(X1)|~v4_unialg_1(X4)|~v3_unialg_1(X1)|~v3_unialg_1(X4)|~v2_unialg_1(X1)|~v2_unialg_1(X4)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125]), c_0_96]), c_0_84])).
cnf(c_0_138, plain, (u1_struct_0(X1)=k1_xboole_0|~l1_unialg_1(X1)|~v2_struct_0(X1)), inference(spm,[status(thm)],[c_0_126, c_0_127])).
cnf(c_0_139, plain, (r2_hidden(k6_eqrel_1(u1_struct_0(X1),u1_struct_0(X1),X2,X3),u1_struct_0(k4_alg_1(X1,k6_alg_1(X1,X4,X5))))|v2_struct_0(X4)|v2_struct_0(X1)|u1_struct_0(k4_alg_1(X1,k6_alg_1(X1,X4,X5)))!=k7_eqrel_1(u1_struct_0(X1),X2)|~r2_hidden(X3,u1_struct_0(X1))|~v1_partfun1(X2,u1_struct_0(X1))|~v8_relat_2(X2)|~v3_relat_2(X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X1))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X4))))|~v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X4))|~v1_funct_1(X5)|~l1_unialg_1(X1)|~l1_unialg_1(X4)|~v4_unialg_1(X1)|~v4_unialg_1(X4)|~v3_unialg_1(X1)|~v3_unialg_1(X4)|~v2_unialg_1(X1)|~v2_unialg_1(X4)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_129]), c_0_130]), c_0_131]), c_0_132])).
cnf(c_0_140, plain, (k1_funct_1(k7_alg_1(X1,X2,X3),k9_relat_1(k6_alg_1(X1,X2,X3),X4))=k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,X4)|v2_struct_0(X1)|v2_struct_0(X2)|~r1_alg_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(X4,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(X2)|~l1_unialg_1(X1)|~v4_unialg_1(X2)|~v4_unialg_1(X1)|~v3_unialg_1(X2)|~v3_unialg_1(X1)|~v2_unialg_1(X2)|~v2_unialg_1(X1)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_134]), c_0_64])).
fof(c_0_141, plain, ![X105, X106]:(~r2_hidden(X105,X106)|m1_subset_1(X105,X106)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_142, plain, (r2_hidden(X1,X2)|~v5_relat_1(X3,X2)|~v1_relat_1(X3)|~r2_hidden(X1,k10_xtuple_0(X3))), inference(spm,[status(thm)],[c_0_135, c_0_136])).
cnf(c_0_143, plain, (u1_struct_0(X1)=k1_xboole_0|r2_hidden(k1_funct_1(k7_alg_1(X2,X1,X3),esk9_3(k7_alg_1(X2,X1,X3),X4,X5)),X6)|v2_struct_0(X2)|X6!=k10_xtuple_0(k7_alg_1(X2,X1,X3))|X4!=k10_xtuple_0(k7_alg_1(X2,X1,X3))|~r2_hidden(X5,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_unialg_1(X1)|~l1_unialg_1(X2)|~v4_unialg_1(X1)|~v4_unialg_1(X2)|~v3_unialg_1(X1)|~v3_unialg_1(X2)|~v2_unialg_1(X1)|~v2_unialg_1(X2)), inference(csr,[status(thm)],[inference(er,[status(thm)],[c_0_137]), c_0_138])).
cnf(c_0_144, plain, (u1_struct_0(X1)=k1_xboole_0|r2_hidden(X2,X3)|v2_struct_0(X4)|X2!=k1_funct_1(k7_alg_1(X4,X1,X5),k6_eqrel_1(u1_struct_0(X4),u1_struct_0(X4),X6,X7))|u1_struct_0(k4_alg_1(X4,k6_alg_1(X4,X1,X5)))!=k7_eqrel_1(u1_struct_0(X4),X6)|X3!=k10_xtuple_0(k7_alg_1(X4,X1,X5))|~r2_hidden(X7,u1_struct_0(X4))|~v1_partfun1(X6,u1_struct_0(X4))|~v8_relat_2(X6)|~v3_relat_2(X6)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X4))))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(X5)|~l1_unialg_1(X1)|~l1_unialg_1(X4)|~v4_unialg_1(X1)|~v4_unialg_1(X4)|~v3_unialg_1(X1)|~v3_unialg_1(X4)|~v2_unialg_1(X1)|~v2_unialg_1(X4)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_139]), c_0_138])).
cnf(c_0_145, plain, (k1_funct_1(k7_alg_1(X1,X2,X3),k6_eqrel_1(X4,X5,k6_alg_1(X1,X2,X3),X6))=k3_funct_2(u1_struct_0(X1),u1_struct_0(X2),X3,X6)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_alg_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~m1_subset_1(k6_alg_1(X1,X2,X3),k1_zfmisc_1(k2_zfmisc_1(X4,X5)))|~m1_subset_1(X6,u1_struct_0(X1))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(X2)|~l1_unialg_1(X1)|~v4_unialg_1(X2)|~v4_unialg_1(X1)|~v3_unialg_1(X2)|~v3_unialg_1(X1)|~v2_unialg_1(X2)|~v2_unialg_1(X1)), inference(spm,[status(thm)],[c_0_140, c_0_134])).
cnf(c_0_146, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_141])).
fof(c_0_147, negated_conjecture, ~(![X1]:(((((~(v2_struct_0(X1))&v2_unialg_1(X1))&v3_unialg_1(X1))&v4_unialg_1(X1))&l1_unialg_1(X1))=>![X2]:(((((~(v2_struct_0(X2))&v2_unialg_1(X2))&v3_unialg_1(X2))&v4_unialg_1(X2))&l1_unialg_1(X2))=>![X3]:(((v1_funct_1(X3)&v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2)))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2)))))=>(r3_alg_1(X1,X2,X3)=>r4_alg_1(k4_alg_1(X1,k6_alg_1(X1,X2,X3)),X2,k7_alg_1(X1,X2,X3))))))), inference(assume_negation,[status(cth)],[t20_alg_1])).
fof(c_0_148, plain, ![X56, X57, X58]:(((r1_alg_1(X56,X57,X58)|~r2_alg_1(X56,X57,X58)|(~v1_funct_1(X58)|~v1_funct_2(X58,u1_struct_0(X56),u1_struct_0(X57))|~m1_subset_1(X58,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X56),u1_struct_0(X57)))))|(v2_struct_0(X57)|~v2_unialg_1(X57)|~v3_unialg_1(X57)|~v4_unialg_1(X57)|~l1_unialg_1(X57))|(v2_struct_0(X56)|~v2_unialg_1(X56)|~v3_unialg_1(X56)|~v4_unialg_1(X56)|~l1_unialg_1(X56)))&(v2_funct_1(X58)|~r2_alg_1(X56,X57,X58)|(~v1_funct_1(X58)|~v1_funct_2(X58,u1_struct_0(X56),u1_struct_0(X57))|~m1_subset_1(X58,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X56),u1_struct_0(X57)))))|(v2_struct_0(X57)|~v2_unialg_1(X57)|~v3_unialg_1(X57)|~v4_unialg_1(X57)|~l1_unialg_1(X57))|(v2_struct_0(X56)|~v2_unialg_1(X56)|~v3_unialg_1(X56)|~v4_unialg_1(X56)|~l1_unialg_1(X56))))&(~r1_alg_1(X56,X57,X58)|~v2_funct_1(X58)|r2_alg_1(X56,X57,X58)|(~v1_funct_1(X58)|~v1_funct_2(X58,u1_struct_0(X56),u1_struct_0(X57))|~m1_subset_1(X58,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X56),u1_struct_0(X57)))))|(v2_struct_0(X57)|~v2_unialg_1(X57)|~v3_unialg_1(X57)|~v4_unialg_1(X57)|~l1_unialg_1(X57))|(v2_struct_0(X56)|~v2_unialg_1(X56)|~v3_unialg_1(X56)|~v4_unialg_1(X56)|~l1_unialg_1(X56)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d2_alg_1])])])])])).
fof(c_0_149, plain, ![X62, X63, X64]:((r1_alg_1(k4_alg_1(X62,k6_alg_1(X62,X63,X64)),X63,k7_alg_1(X62,X63,X64))|~r1_alg_1(X62,X63,X64)|(~v1_funct_1(X64)|~v1_funct_2(X64,u1_struct_0(X62),u1_struct_0(X63))|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X62),u1_struct_0(X63)))))|(v2_struct_0(X63)|~v2_unialg_1(X63)|~v3_unialg_1(X63)|~v4_unialg_1(X63)|~l1_unialg_1(X63))|(v2_struct_0(X62)|~v2_unialg_1(X62)|~v3_unialg_1(X62)|~v4_unialg_1(X62)|~l1_unialg_1(X62)))&(r2_alg_1(k4_alg_1(X62,k6_alg_1(X62,X63,X64)),X63,k7_alg_1(X62,X63,X64))|~r1_alg_1(X62,X63,X64)|(~v1_funct_1(X64)|~v1_funct_2(X64,u1_struct_0(X62),u1_struct_0(X63))|~m1_subset_1(X64,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X62),u1_struct_0(X63)))))|(v2_struct_0(X63)|~v2_unialg_1(X63)|~v3_unialg_1(X63)|~v4_unialg_1(X63)|~l1_unialg_1(X63))|(v2_struct_0(X62)|~v2_unialg_1(X62)|~v3_unialg_1(X62)|~v4_unialg_1(X62)|~l1_unialg_1(X62)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t19_alg_1])])])])])).
cnf(c_0_150, plain, (u1_struct_0(X1)=k1_xboole_0|r2_hidden(k1_funct_1(k7_alg_1(X2,X1,X3),esk9_3(k7_alg_1(X2,X1,X3),X4,X5)),X6)|v2_struct_0(X2)|k10_xtuple_0(X7)!=k10_xtuple_0(k7_alg_1(X2,X1,X3))|X4!=k10_xtuple_0(k7_alg_1(X2,X1,X3))|~v5_relat_1(X7,X6)|~v1_relat_1(X7)|~r2_hidden(X5,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_unialg_1(X1)|~l1_unialg_1(X2)|~v4_unialg_1(X1)|~v4_unialg_1(X2)|~v3_unialg_1(X1)|~v3_unialg_1(X2)|~v2_unialg_1(X1)|~v2_unialg_1(X2)), inference(spm,[status(thm)],[c_0_142, c_0_143])).
cnf(c_0_151, plain, (u1_struct_0(X1)=k1_xboole_0|r2_hidden(X2,X3)|v2_struct_0(X4)|X2!=k3_funct_2(u1_struct_0(X4),u1_struct_0(X1),X5,X6)|X3!=k10_xtuple_0(k7_alg_1(X4,X1,X5))|~r2_hidden(X6,u1_struct_0(X4))|~r1_alg_1(X4,X1,X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(X5)|~l1_unialg_1(X1)|~l1_unialg_1(X4)|~v4_unialg_1(X1)|~v4_unialg_1(X4)|~v3_unialg_1(X1)|~v3_unialg_1(X4)|~v2_unialg_1(X1)|~v2_unialg_1(X4)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_144, c_0_145]), c_0_146]), c_0_64]), c_0_130]), c_0_131]), c_0_132]), c_0_129]), c_0_138])).
fof(c_0_152, plain, ![X94, X95]:(~v1_relat_1(X95)|~v4_relat_1(X95,X94)|m1_subset_1(k1_relset_1(X94,X95),k1_zfmisc_1(X94))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_relset_1])])).
fof(c_0_153, negated_conjecture, (((((~v2_struct_0(esk1_0)&v2_unialg_1(esk1_0))&v3_unialg_1(esk1_0))&v4_unialg_1(esk1_0))&l1_unialg_1(esk1_0))&(((((~v2_struct_0(esk2_0)&v2_unialg_1(esk2_0))&v3_unialg_1(esk2_0))&v4_unialg_1(esk2_0))&l1_unialg_1(esk2_0))&(((v1_funct_1(esk3_0)&v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0)))&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0)))))&(r3_alg_1(esk1_0,esk2_0,esk3_0)&~r4_alg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)),esk2_0,k7_alg_1(esk1_0,esk2_0,esk3_0)))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_147])])])])).
fof(c_0_154, plain, ![X15, X16, X17]:((((r1_alg_1(X15,X16,X17)|~r4_alg_1(X15,X16,X17)|(~v1_funct_1(X17)|~v1_funct_2(X17,u1_struct_0(X15),u1_struct_0(X16))|~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X15),u1_struct_0(X16)))))|(v2_struct_0(X16)|~v2_unialg_1(X16)|~v3_unialg_1(X16)|~v4_unialg_1(X16)|~l1_unialg_1(X16))|(v2_struct_0(X15)|~v2_unialg_1(X15)|~v3_unialg_1(X15)|~v4_unialg_1(X15)|~l1_unialg_1(X15)))&(k2_relset_1(u1_struct_0(X16),X17)=u1_struct_0(X16)|~r4_alg_1(X15,X16,X17)|(~v1_funct_1(X17)|~v1_funct_2(X17,u1_struct_0(X15),u1_struct_0(X16))|~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X15),u1_struct_0(X16)))))|(v2_struct_0(X16)|~v2_unialg_1(X16)|~v3_unialg_1(X16)|~v4_unialg_1(X16)|~l1_unialg_1(X16))|(v2_struct_0(X15)|~v2_unialg_1(X15)|~v3_unialg_1(X15)|~v4_unialg_1(X15)|~l1_unialg_1(X15))))&(v2_funct_1(X17)|~r4_alg_1(X15,X16,X17)|(~v1_funct_1(X17)|~v1_funct_2(X17,u1_struct_0(X15),u1_struct_0(X16))|~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X15),u1_struct_0(X16)))))|(v2_struct_0(X16)|~v2_unialg_1(X16)|~v3_unialg_1(X16)|~v4_unialg_1(X16)|~l1_unialg_1(X16))|(v2_struct_0(X15)|~v2_unialg_1(X15)|~v3_unialg_1(X15)|~v4_unialg_1(X15)|~l1_unialg_1(X15))))&(~r1_alg_1(X15,X16,X17)|k2_relset_1(u1_struct_0(X16),X17)!=u1_struct_0(X16)|~v2_funct_1(X17)|r4_alg_1(X15,X16,X17)|(~v1_funct_1(X17)|~v1_funct_2(X17,u1_struct_0(X15),u1_struct_0(X16))|~m1_subset_1(X17,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X15),u1_struct_0(X16)))))|(v2_struct_0(X16)|~v2_unialg_1(X16)|~v3_unialg_1(X16)|~v4_unialg_1(X16)|~l1_unialg_1(X16))|(v2_struct_0(X15)|~v2_unialg_1(X15)|~v3_unialg_1(X15)|~v4_unialg_1(X15)|~l1_unialg_1(X15)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t7_alg_1])])])])])).
fof(c_0_155, plain, ![X18, X19, X20]:(((r1_alg_1(X18,X19,X20)|~r3_alg_1(X18,X19,X20)|(~v1_funct_1(X20)|~v1_funct_2(X20,u1_struct_0(X18),u1_struct_0(X19))|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X18),u1_struct_0(X19)))))|(v2_struct_0(X19)|~v2_unialg_1(X19)|~v3_unialg_1(X19)|~v4_unialg_1(X19)|~l1_unialg_1(X19))|(v2_struct_0(X18)|~v2_unialg_1(X18)|~v3_unialg_1(X18)|~v4_unialg_1(X18)|~l1_unialg_1(X18)))&(k2_relset_1(u1_struct_0(X19),X20)=u1_struct_0(X19)|~r3_alg_1(X18,X19,X20)|(~v1_funct_1(X20)|~v1_funct_2(X20,u1_struct_0(X18),u1_struct_0(X19))|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X18),u1_struct_0(X19)))))|(v2_struct_0(X19)|~v2_unialg_1(X19)|~v3_unialg_1(X19)|~v4_unialg_1(X19)|~l1_unialg_1(X19))|(v2_struct_0(X18)|~v2_unialg_1(X18)|~v3_unialg_1(X18)|~v4_unialg_1(X18)|~l1_unialg_1(X18))))&(~r1_alg_1(X18,X19,X20)|k2_relset_1(u1_struct_0(X19),X20)!=u1_struct_0(X19)|r3_alg_1(X18,X19,X20)|(~v1_funct_1(X20)|~v1_funct_2(X20,u1_struct_0(X18),u1_struct_0(X19))|~m1_subset_1(X20,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X18),u1_struct_0(X19)))))|(v2_struct_0(X19)|~v2_unialg_1(X19)|~v3_unialg_1(X19)|~v4_unialg_1(X19)|~l1_unialg_1(X19))|(v2_struct_0(X18)|~v2_unialg_1(X18)|~v3_unialg_1(X18)|~v4_unialg_1(X18)|~l1_unialg_1(X18)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d3_alg_1])])])])])).
cnf(c_0_156, plain, (v2_funct_1(X1)|v2_struct_0(X3)|v2_struct_0(X2)|~r2_alg_1(X2,X3,X1)|~v1_funct_1(X1)|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v2_unialg_1(X3)|~v3_unialg_1(X3)|~v4_unialg_1(X3)|~l1_unialg_1(X3)|~v2_unialg_1(X2)|~v3_unialg_1(X2)|~v4_unialg_1(X2)|~l1_unialg_1(X2)), inference(split_conjunct,[status(thm)],[c_0_148])).
cnf(c_0_157, plain, (r2_alg_1(k4_alg_1(X1,k6_alg_1(X1,X2,X3)),X2,k7_alg_1(X1,X2,X3))|v2_struct_0(X2)|v2_struct_0(X1)|~r1_alg_1(X1,X2,X3)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v2_unialg_1(X2)|~v3_unialg_1(X2)|~v4_unialg_1(X2)|~l1_unialg_1(X2)|~v2_unialg_1(X1)|~v3_unialg_1(X1)|~v4_unialg_1(X1)|~l1_unialg_1(X1)), inference(split_conjunct,[status(thm)],[c_0_149])).
cnf(c_0_158, plain, (u1_struct_0(X1)=k1_xboole_0|r2_hidden(k1_funct_1(k7_alg_1(X2,X1,X3),esk9_3(k7_alg_1(X2,X1,X3),X4,X5)),X6)|v2_struct_0(X2)|X4!=k10_xtuple_0(k7_alg_1(X2,X1,X3))|~v5_relat_1(k7_alg_1(X2,X1,X3),X6)|~v1_relat_1(k7_alg_1(X2,X1,X3))|~r2_hidden(X5,X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_unialg_1(X1)|~l1_unialg_1(X2)|~v4_unialg_1(X1)|~v4_unialg_1(X2)|~v3_unialg_1(X1)|~v3_unialg_1(X2)|~v2_unialg_1(X1)|~v2_unialg_1(X2)), inference(er,[status(thm)],[c_0_150])).
cnf(c_0_159, plain, (X1=k1_funct_1(X2,esk9_3(X2,X3,X1))|~r2_hidden(X1,X3)|X3!=k10_xtuple_0(X2)|~v1_relat_1(X2)|~v1_funct_1(X2)), inference(split_conjunct,[status(thm)],[c_0_81])).
cnf(c_0_160, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2)))), inference(split_conjunct,[status(thm)],[c_0_61])).
cnf(c_0_161, plain, (u1_struct_0(X1)=k1_xboole_0|r2_hidden(k3_funct_2(u1_struct_0(X2),u1_struct_0(X1),X3,X4),X5)|v2_struct_0(X2)|X5!=k10_xtuple_0(k7_alg_1(X2,X1,X3))|~r2_hidden(X4,u1_struct_0(X2))|~r1_alg_1(X2,X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_unialg_1(X1)|~l1_unialg_1(X2)|~v4_unialg_1(X1)|~v4_unialg_1(X2)|~v3_unialg_1(X1)|~v3_unialg_1(X2)|~v2_unialg_1(X1)|~v2_unialg_1(X2)), inference(er,[status(thm)],[c_0_151])).
fof(c_0_162, plain, ![X81, X82, X83, X84]:(v1_xboole_0(X81)|~v1_funct_1(X83)|~v1_funct_2(X83,X81,X82)|~m1_subset_1(X83,k1_zfmisc_1(k2_zfmisc_1(X81,X82)))|~m1_subset_1(X84,X81)|k3_funct_2(X81,X82,X83,X84)=k1_funct_1(X83,X84)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k3_funct_2])])])).
fof(c_0_163, plain, ![X109, X110]:(~r2_hidden(X109,X110)|~v1_xboole_0(X110)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
cnf(c_0_164, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_152])).
fof(c_0_165, plain, ![X222, X223]:((((m1_subset_1(esk13_2(X222,X223),k1_zfmisc_1(k2_zfmisc_1(X222,X223)))&v1_xboole_0(esk13_2(X222,X223)))&v1_relat_1(esk13_2(X222,X223)))&v4_relat_1(esk13_2(X222,X223),X222))&v5_relat_1(esk13_2(X222,X223),X223)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc1_relset_1])])).
cnf(c_0_166, negated_conjecture, (~r4_alg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)),esk2_0,k7_alg_1(esk1_0,esk2_0,esk3_0))), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_167, plain, (r4_alg_1(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r1_alg_1(X1,X2,X3)|k2_relset_1(u1_struct_0(X2),X3)!=u1_struct_0(X2)|~v2_funct_1(X3)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v2_unialg_1(X2)|~v3_unialg_1(X2)|~v4_unialg_1(X2)|~l1_unialg_1(X2)|~v2_unialg_1(X1)|~v3_unialg_1(X1)|~v4_unialg_1(X1)|~l1_unialg_1(X1)), inference(split_conjunct,[status(thm)],[c_0_154])).
cnf(c_0_168, negated_conjecture, (l1_unialg_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_169, negated_conjecture, (v4_unialg_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_170, negated_conjecture, (v3_unialg_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_171, negated_conjecture, (v2_unialg_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_172, negated_conjecture, (~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_173, plain, (r1_alg_1(X1,X2,X3)|v2_struct_0(X2)|v2_struct_0(X1)|~r3_alg_1(X1,X2,X3)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v2_unialg_1(X2)|~v3_unialg_1(X2)|~v4_unialg_1(X2)|~l1_unialg_1(X2)|~v2_unialg_1(X1)|~v3_unialg_1(X1)|~v4_unialg_1(X1)|~l1_unialg_1(X1)), inference(split_conjunct,[status(thm)],[c_0_155])).
cnf(c_0_174, negated_conjecture, (r3_alg_1(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_175, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk1_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_176, negated_conjecture, (v1_funct_2(esk3_0,u1_struct_0(esk1_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_177, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_178, negated_conjecture, (l1_unialg_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_179, negated_conjecture, (v4_unialg_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_180, negated_conjecture, (v3_unialg_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_181, negated_conjecture, (v2_unialg_1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_182, negated_conjecture, (~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_183, plain, (v2_funct_1(k7_alg_1(X1,X2,X3))|v2_struct_0(k4_alg_1(X1,k6_alg_1(X1,X2,X3)))|v2_struct_0(X1)|v2_struct_0(X2)|~r1_alg_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(k4_alg_1(X1,k6_alg_1(X1,X2,X3)))|~l1_unialg_1(X2)|~l1_unialg_1(X1)|~v4_unialg_1(k4_alg_1(X1,k6_alg_1(X1,X2,X3)))|~v4_unialg_1(X2)|~v4_unialg_1(X1)|~v3_unialg_1(k4_alg_1(X1,k6_alg_1(X1,X2,X3)))|~v3_unialg_1(X2)|~v3_unialg_1(X1)|~v2_unialg_1(k4_alg_1(X1,k6_alg_1(X1,X2,X3)))|~v2_unialg_1(X2)|~v2_unialg_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_157]), c_0_96]), c_0_70]), c_0_71])).
cnf(c_0_184, plain, (v3_unialg_1(k4_alg_1(X1,X2))|v2_struct_0(X1)|~v2_unialg_1(X1)|~v3_unialg_1(X1)|~v4_unialg_1(X1)|~l1_unialg_1(X1)|~m1_alg_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_185, plain, (v2_unialg_1(k4_alg_1(X1,X2))|v2_struct_0(X1)|~v2_unialg_1(X1)|~v3_unialg_1(X1)|~v4_unialg_1(X1)|~l1_unialg_1(X1)|~m1_alg_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_186, plain, (v4_unialg_1(k4_alg_1(X1,X2))|v2_struct_0(X1)|~v2_unialg_1(X1)|~v3_unialg_1(X1)|~v4_unialg_1(X1)|~l1_unialg_1(X1)|~m1_alg_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_187, plain, (u1_struct_0(X1)=k1_xboole_0|r2_hidden(X2,X3)|v2_struct_0(X4)|X5!=k10_xtuple_0(k7_alg_1(X4,X1,X6))|~v5_relat_1(k7_alg_1(X4,X1,X6),X3)|~v1_relat_1(k7_alg_1(X4,X1,X6))|~r2_hidden(X2,X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(k7_alg_1(X4,X1,X6))|~v1_funct_1(X6)|~l1_unialg_1(X1)|~l1_unialg_1(X4)|~v4_unialg_1(X1)|~v4_unialg_1(X4)|~v3_unialg_1(X1)|~v3_unialg_1(X4)|~v2_unialg_1(X1)|~v2_unialg_1(X4)), inference(spm,[status(thm)],[c_0_158, c_0_159])).
cnf(c_0_188, plain, (v5_relat_1(k7_alg_1(X1,X2,X3),u1_struct_0(X2))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(X2)|~l1_unialg_1(X1)|~v4_unialg_1(X2)|~v4_unialg_1(X1)|~v3_unialg_1(X2)|~v3_unialg_1(X1)|~v2_unialg_1(X2)|~v2_unialg_1(X1)), inference(spm,[status(thm)],[c_0_160, c_0_71])).
cnf(c_0_189, plain, (u1_struct_0(X1)=k1_xboole_0|r2_hidden(k3_funct_2(u1_struct_0(X2),u1_struct_0(X1),X3,X4),k10_xtuple_0(k7_alg_1(X2,X1,X3)))|v2_struct_0(X2)|~r2_hidden(X4,u1_struct_0(X2))|~r1_alg_1(X2,X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_unialg_1(X1)|~l1_unialg_1(X2)|~v4_unialg_1(X1)|~v4_unialg_1(X2)|~v3_unialg_1(X1)|~v3_unialg_1(X2)|~v2_unialg_1(X1)|~v2_unialg_1(X2)), inference(er,[status(thm)],[c_0_161])).
cnf(c_0_190, plain, (v1_xboole_0(X1)|k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|~v1_funct_1(X2)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3)))|~m1_subset_1(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_162])).
cnf(c_0_191, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_163])).
fof(c_0_192, plain, ![X166, X167, X168]:(~m1_subset_1(X168,k1_zfmisc_1(k2_zfmisc_1(X166,X167)))|v1_relat_1(X168)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_relset_1])])).
cnf(c_0_193, plain, (r2_hidden(X1,X2)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~r2_hidden(X1,k1_relset_1(X2,X3))), inference(spm,[status(thm)],[c_0_121, c_0_164])).
cnf(c_0_194, plain, (v1_xboole_0(esk13_2(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_165])).
fof(c_0_195, plain, ![X218]:(((((((m1_subset_1(esk12_1(X218),k1_zfmisc_1(k2_zfmisc_1(X218,X218)))&v1_relat_1(esk12_1(X218)))&v4_relat_1(esk12_1(X218),X218))&v5_relat_1(esk12_1(X218),X218))&v1_funct_1(esk12_1(X218)))&v1_partfun1(esk12_1(X218),X218))&v1_funct_2(esk12_1(X218),X218,X218))&v3_funct_2(esk12_1(X218),X218,X218)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc2_funct_2])])).
cnf(c_0_196, negated_conjecture, (v2_struct_0(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|k2_relset_1(u1_struct_0(esk2_0),k7_alg_1(esk1_0,esk2_0,esk3_0))!=u1_struct_0(esk2_0)|~v2_funct_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~r1_alg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)),esk2_0,k7_alg_1(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k7_alg_1(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0))),u1_struct_0(esk2_0))))|~v1_funct_2(k7_alg_1(esk1_0,esk2_0,esk3_0),u1_struct_0(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0))),u1_struct_0(esk2_0))|~v1_funct_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~l1_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v4_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v3_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v2_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_167]), c_0_168]), c_0_169]), c_0_170]), c_0_171])]), c_0_172])).
cnf(c_0_197, plain, (r1_alg_1(k4_alg_1(X1,k6_alg_1(X1,X2,X3)),X2,k7_alg_1(X1,X2,X3))|v2_struct_0(X2)|v2_struct_0(X1)|~r1_alg_1(X1,X2,X3)|~v1_funct_1(X3)|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v2_unialg_1(X2)|~v3_unialg_1(X2)|~v4_unialg_1(X2)|~l1_unialg_1(X2)|~v2_unialg_1(X1)|~v3_unialg_1(X1)|~v4_unialg_1(X1)|~l1_unialg_1(X1)), inference(split_conjunct,[status(thm)],[c_0_149])).
cnf(c_0_198, negated_conjecture, (r1_alg_1(esk1_0,esk2_0,esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_174]), c_0_175]), c_0_176]), c_0_177]), c_0_168]), c_0_178]), c_0_169]), c_0_179]), c_0_170]), c_0_180]), c_0_171]), c_0_181])]), c_0_182]), c_0_172])).
cnf(c_0_199, plain, (v2_funct_1(k7_alg_1(X1,X2,X3))|v2_struct_0(k4_alg_1(X1,k6_alg_1(X1,X2,X3)))|v2_struct_0(X2)|v2_struct_0(X1)|~r1_alg_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(X2)|~l1_unialg_1(X1)|~v4_unialg_1(X2)|~v4_unialg_1(X1)|~v3_unialg_1(X2)|~v3_unialg_1(X1)|~v2_unialg_1(X2)|~v2_unialg_1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_184]), c_0_185]), c_0_186]), c_0_103]), c_0_53])).
cnf(c_0_200, plain, (u1_struct_0(X1)=k1_xboole_0|r2_hidden(X2,u1_struct_0(X1))|v2_struct_0(X3)|X4!=k10_xtuple_0(k7_alg_1(X3,X1,X5))|~r2_hidden(X2,X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X5)|~l1_unialg_1(X1)|~l1_unialg_1(X3)|~v4_unialg_1(X1)|~v4_unialg_1(X3)|~v3_unialg_1(X1)|~v3_unialg_1(X3)|~v2_unialg_1(X1)|~v2_unialg_1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_188]), c_0_96]), c_0_84]), c_0_138])).
cnf(c_0_201, plain, (u1_struct_0(X1)=k1_xboole_0|r2_hidden(k1_funct_1(X2,X3),k10_xtuple_0(k7_alg_1(X4,X1,X2)))|~r2_hidden(X3,u1_struct_0(X4))|~r1_alg_1(X4,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~v1_funct_2(X2,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(X2)|~l1_unialg_1(X1)|~l1_unialg_1(X4)|~v4_unialg_1(X1)|~v4_unialg_1(X4)|~v3_unialg_1(X1)|~v3_unialg_1(X4)|~v2_unialg_1(X1)|~v2_unialg_1(X4)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_190]), c_0_146]), c_0_127]), c_0_191])).
cnf(c_0_202, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_192])).
cnf(c_0_203, plain, (r2_hidden(X1,X2)|~v4_relat_1(X3,X2)|~v1_relat_1(X3)|~r2_hidden(X1,k9_xtuple_0(X3))), inference(spm,[status(thm)],[c_0_193, c_0_82])).
fof(c_0_204, plain, ![X123, X124, X125]:(v1_xboole_0(X123)|~v1_xboole_0(X124)|(~m1_subset_1(X125,k1_zfmisc_1(k2_zfmisc_1(X123,X124)))|~v1_partfun1(X125,X123))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_partfun1])])])])).
cnf(c_0_205, plain, (m1_subset_1(esk13_2(X1,X2),k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(split_conjunct,[status(thm)],[c_0_165])).
cnf(c_0_206, plain, (esk13_2(X1,X2)=k1_xboole_0), inference(spm,[status(thm)],[c_0_126, c_0_194])).
cnf(c_0_207, plain, (X1=k1_xboole_0|~v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3)))), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_208, plain, (v1_funct_2(esk12_1(X1),X1,X1)), inference(split_conjunct,[status(thm)],[c_0_195])).
cnf(c_0_209, plain, (m1_subset_1(esk12_1(X1),k1_zfmisc_1(k2_zfmisc_1(X1,X1)))), inference(split_conjunct,[status(thm)],[c_0_195])).
cnf(c_0_210, plain, (k2_relset_1(u1_struct_0(X1),X2)=u1_struct_0(X1)|v2_struct_0(X1)|v2_struct_0(X3)|~r3_alg_1(X3,X1,X2)|~v1_funct_1(X2)|~v1_funct_2(X2,u1_struct_0(X3),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v2_unialg_1(X1)|~v3_unialg_1(X1)|~v4_unialg_1(X1)|~l1_unialg_1(X1)|~v2_unialg_1(X3)|~v3_unialg_1(X3)|~v4_unialg_1(X3)|~l1_unialg_1(X3)), inference(split_conjunct,[status(thm)],[c_0_155])).
cnf(c_0_211, negated_conjecture, (v2_struct_0(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|k2_relset_1(u1_struct_0(esk2_0),k7_alg_1(esk1_0,esk2_0,esk3_0))!=u1_struct_0(esk2_0)|~v2_funct_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k7_alg_1(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0))),u1_struct_0(esk2_0))))|~v1_funct_2(k7_alg_1(esk1_0,esk2_0,esk3_0),u1_struct_0(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0))),u1_struct_0(esk2_0))|~v1_funct_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~l1_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v4_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v3_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v2_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_197]), c_0_198]), c_0_175]), c_0_176]), c_0_177]), c_0_168]), c_0_178]), c_0_169]), c_0_179]), c_0_170]), c_0_180]), c_0_171]), c_0_181])]), c_0_182]), c_0_172])).
cnf(c_0_212, negated_conjecture, (v2_funct_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|v2_struct_0(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_175]), c_0_198]), c_0_176]), c_0_177]), c_0_168]), c_0_178]), c_0_169]), c_0_179]), c_0_170]), c_0_180]), c_0_171]), c_0_181])]), c_0_172]), c_0_182])).
cnf(c_0_213, plain, (u1_struct_0(X1)=k1_xboole_0|r2_hidden(X2,u1_struct_0(X1))|v2_struct_0(X3)|~r2_hidden(X2,k10_xtuple_0(k7_alg_1(X3,X1,X4)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X4)|~l1_unialg_1(X1)|~l1_unialg_1(X3)|~v4_unialg_1(X1)|~v4_unialg_1(X3)|~v3_unialg_1(X1)|~v3_unialg_1(X3)|~v2_unialg_1(X1)|~v2_unialg_1(X3)), inference(er,[status(thm)],[c_0_200])).
cnf(c_0_214, plain, (r2_hidden(esk8_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_215, plain, (u1_struct_0(X1)=k1_xboole_0|r2_hidden(X2,k10_xtuple_0(k7_alg_1(X3,X1,X4)))|X5!=k10_xtuple_0(X4)|~r2_hidden(esk9_3(X4,X5,X2),u1_struct_0(X3))|~r2_hidden(X2,X5)|~r1_alg_1(X3,X1,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X4)|~l1_unialg_1(X1)|~l1_unialg_1(X3)|~v4_unialg_1(X1)|~v4_unialg_1(X3)|~v3_unialg_1(X1)|~v3_unialg_1(X3)|~v2_unialg_1(X1)|~v2_unialg_1(X3)), inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_159]), c_0_202])).
cnf(c_0_216, plain, (r2_hidden(esk9_3(X1,X2,X3),X4)|X2!=k10_xtuple_0(X1)|~v4_relat_1(X1,X4)|~v1_relat_1(X1)|~r2_hidden(X3,X2)|~v1_funct_1(X1)), inference(spm,[status(thm)],[c_0_203, c_0_125])).
cnf(c_0_217, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_partfun1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_204])).
cnf(c_0_218, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))), inference(rw,[status(thm)],[c_0_205, c_0_206])).
cnf(c_0_219, plain, (v1_partfun1(esk12_1(X1),X1)), inference(split_conjunct,[status(thm)],[c_0_195])).
cnf(c_0_220, plain, (esk12_1(X1)=k1_xboole_0|X1!=k1_xboole_0), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_208]), c_0_209])])).
fof(c_0_221, plain, ![X66]:(v2_struct_0(X66)|~l1_struct_0(X66)|~v1_xboole_0(u1_struct_0(X66))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_222, plain, ![X135]:(~v1_xboole_0(X135)|v1_xboole_0(k10_xtuple_0(X135))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc11_relat_1])])).
cnf(c_0_223, negated_conjecture, (k2_relset_1(u1_struct_0(esk2_0),esk3_0)=u1_struct_0(esk2_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_174]), c_0_175]), c_0_176]), c_0_177]), c_0_178]), c_0_168]), c_0_179]), c_0_169]), c_0_180]), c_0_170]), c_0_181]), c_0_171])]), c_0_172]), c_0_182])).
cnf(c_0_224, negated_conjecture, (v1_relat_1(esk3_0)), inference(spm,[status(thm)],[c_0_202, c_0_175])).
cnf(c_0_225, negated_conjecture, (v5_relat_1(esk3_0,u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_160, c_0_175])).
cnf(c_0_226, negated_conjecture, (v2_struct_0(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|k2_relset_1(u1_struct_0(esk2_0),k7_alg_1(esk1_0,esk2_0,esk3_0))!=u1_struct_0(esk2_0)|~m1_subset_1(k7_alg_1(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0))),u1_struct_0(esk2_0))))|~v1_funct_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~l1_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v4_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v3_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v2_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))), inference(csr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_70]), c_0_175]), c_0_176]), c_0_177]), c_0_168]), c_0_178]), c_0_169]), c_0_179]), c_0_170]), c_0_180]), c_0_171]), c_0_181])]), c_0_182]), c_0_172]), c_0_212])).
fof(c_0_227, plain, ![X30, X31]:(((r1_tarski(X30,X31)|X30!=X31)&(r1_tarski(X31,X30)|X30!=X31))&(~r1_tarski(X30,X31)|~r1_tarski(X31,X30)|X30=X31)), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d10_xboole_0])])])).
cnf(c_0_228, plain, (r1_tarski(X1,X2)|~r2_hidden(esk8_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_92])).
cnf(c_0_229, plain, (u1_struct_0(X1)=k1_xboole_0|r1_tarski(k10_xtuple_0(k7_alg_1(X2,X1,X3)),X4)|r2_hidden(esk8_2(k10_xtuple_0(k7_alg_1(X2,X1,X3)),X4),u1_struct_0(X1))|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_unialg_1(X1)|~l1_unialg_1(X2)|~v4_unialg_1(X1)|~v4_unialg_1(X2)|~v3_unialg_1(X1)|~v3_unialg_1(X2)|~v2_unialg_1(X1)|~v2_unialg_1(X2)), inference(spm,[status(thm)],[c_0_213, c_0_214])).
cnf(c_0_230, plain, (u1_struct_0(X1)=k1_xboole_0|r2_hidden(X2,k10_xtuple_0(k7_alg_1(X3,X1,X4)))|X5!=k10_xtuple_0(X4)|~r2_hidden(X2,X5)|~r1_alg_1(X3,X1,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X4)|~l1_unialg_1(X1)|~l1_unialg_1(X3)|~v4_unialg_1(X1)|~v4_unialg_1(X3)|~v3_unialg_1(X1)|~v3_unialg_1(X3)|~v2_unialg_1(X1)|~v2_unialg_1(X3)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_216]), c_0_202]), c_0_74])).
fof(c_0_231, plain, ![X148, X149, X150]:(~v1_xboole_0(X148)|(~m1_subset_1(X150,k1_zfmisc_1(k2_zfmisc_1(X149,X148)))|v1_xboole_0(X150))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_relset_1])])])).
cnf(c_0_232, plain, (v1_xboole_0(X1)|~v1_xboole_0(X2)|~v1_partfun1(k1_xboole_0,X1)), inference(spm,[status(thm)],[c_0_217, c_0_218])).
cnf(c_0_233, plain, (v1_partfun1(k1_xboole_0,X1)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_219, c_0_220])).
cnf(c_0_234, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_221])).
cnf(c_0_235, plain, (v1_xboole_0(k10_xtuple_0(X1))|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_222])).
cnf(c_0_236, negated_conjecture, (k10_xtuple_0(esk3_0)=u1_struct_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_223]), c_0_224])]), c_0_225])])).
cnf(c_0_237, negated_conjecture, (v2_struct_0(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|k2_relset_1(u1_struct_0(esk2_0),k7_alg_1(esk1_0,esk2_0,esk3_0))!=u1_struct_0(esk2_0)|~v1_funct_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~l1_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v4_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v3_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v2_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_71]), c_0_175]), c_0_176]), c_0_177]), c_0_168]), c_0_178]), c_0_169]), c_0_179]), c_0_170]), c_0_180]), c_0_171]), c_0_181])]), c_0_182]), c_0_172])).
cnf(c_0_238, plain, (X1=X2|~r1_tarski(X1,X2)|~r1_tarski(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_227])).
cnf(c_0_239, plain, (u1_struct_0(X1)=k1_xboole_0|r1_tarski(k10_xtuple_0(k7_alg_1(X2,X1,X3)),u1_struct_0(X1))|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X1))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(X1))|~v1_funct_1(X3)|~l1_unialg_1(X1)|~l1_unialg_1(X2)|~v4_unialg_1(X1)|~v4_unialg_1(X2)|~v3_unialg_1(X1)|~v3_unialg_1(X2)|~v2_unialg_1(X1)|~v2_unialg_1(X2)), inference(spm,[status(thm)],[c_0_228, c_0_229])).
cnf(c_0_240, plain, (u1_struct_0(X1)=k1_xboole_0|r1_tarski(X2,X3)|r2_hidden(esk8_2(X2,X3),k10_xtuple_0(k7_alg_1(X4,X1,X5)))|X2!=k10_xtuple_0(X5)|~r1_alg_1(X4,X1,X5)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X1))))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X1))|~v1_funct_1(X5)|~l1_unialg_1(X1)|~l1_unialg_1(X4)|~v4_unialg_1(X1)|~v4_unialg_1(X4)|~v3_unialg_1(X1)|~v3_unialg_1(X4)|~v2_unialg_1(X1)|~v2_unialg_1(X4)), inference(spm,[status(thm)],[c_0_230, c_0_214])).
cnf(c_0_241, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X1)))), inference(split_conjunct,[status(thm)],[c_0_231])).
cnf(c_0_242, plain, (v1_xboole_0(X1)|X1!=k1_xboole_0|~v1_xboole_0(X2)), inference(spm,[status(thm)],[c_0_232, c_0_233])).
cnf(c_0_243, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_244, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_unialg_1(X1)), inference(spm,[status(thm)],[c_0_234, c_0_112])).
cnf(c_0_245, negated_conjecture, (v1_xboole_0(u1_struct_0(esk2_0))|~v1_xboole_0(esk3_0)), inference(spm,[status(thm)],[c_0_235, c_0_236])).
cnf(c_0_246, negated_conjecture, (v2_struct_0(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|k10_xtuple_0(k7_alg_1(esk1_0,esk2_0,esk3_0))!=u1_struct_0(esk2_0)|~v5_relat_1(k7_alg_1(esk1_0,esk2_0,esk3_0),u1_struct_0(esk2_0))|~v1_relat_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~v1_funct_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~l1_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v4_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v3_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v2_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))), inference(spm,[status(thm)],[c_0_237, c_0_136])).
cnf(c_0_247, plain, (k10_xtuple_0(k7_alg_1(X1,X2,X3))=u1_struct_0(X2)|u1_struct_0(X2)=k1_xboole_0|v2_struct_0(X1)|~r1_tarski(u1_struct_0(X2),k10_xtuple_0(k7_alg_1(X1,X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(X2)|~l1_unialg_1(X1)|~v4_unialg_1(X2)|~v4_unialg_1(X1)|~v3_unialg_1(X2)|~v3_unialg_1(X1)|~v2_unialg_1(X2)|~v2_unialg_1(X1)), inference(spm,[status(thm)],[c_0_238, c_0_239])).
cnf(c_0_248, plain, (u1_struct_0(X1)=k1_xboole_0|r1_tarski(X2,k10_xtuple_0(k7_alg_1(X3,X1,X4)))|X2!=k10_xtuple_0(X4)|~r1_alg_1(X3,X1,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X1))))|~v1_funct_2(X4,u1_struct_0(X3),u1_struct_0(X1))|~v1_funct_1(X4)|~l1_unialg_1(X1)|~l1_unialg_1(X3)|~v4_unialg_1(X1)|~v4_unialg_1(X3)|~v3_unialg_1(X1)|~v3_unialg_1(X3)|~v2_unialg_1(X1)|~v2_unialg_1(X3)), inference(spm,[status(thm)],[c_0_228, c_0_240])).
cnf(c_0_249, negated_conjecture, (v1_xboole_0(esk3_0)|~v1_xboole_0(u1_struct_0(esk2_0))), inference(spm,[status(thm)],[c_0_241, c_0_175])).
cnf(c_0_250, plain, (v1_xboole_0(X1)|X1!=k1_xboole_0), inference(spm,[status(thm)],[c_0_242, c_0_243])).
cnf(c_0_251, negated_conjecture, (~v1_xboole_0(esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_245]), c_0_168])]), c_0_172])).
cnf(c_0_252, negated_conjecture, (v2_struct_0(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|k10_xtuple_0(k7_alg_1(esk1_0,esk2_0,esk3_0))!=u1_struct_0(esk2_0)|~v1_relat_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~v1_funct_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~l1_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v4_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v3_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v2_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_188]), c_0_175]), c_0_176]), c_0_177]), c_0_168]), c_0_178]), c_0_169]), c_0_179]), c_0_170]), c_0_180]), c_0_171]), c_0_181])]), c_0_172]), c_0_182])).
cnf(c_0_253, plain, (k10_xtuple_0(k7_alg_1(X1,X2,X3))=u1_struct_0(X2)|u1_struct_0(X2)=k1_xboole_0|v2_struct_0(X1)|u1_struct_0(X2)!=k10_xtuple_0(X3)|~r1_alg_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X3,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l1_unialg_1(X2)|~l1_unialg_1(X1)|~v4_unialg_1(X2)|~v4_unialg_1(X1)|~v3_unialg_1(X2)|~v3_unialg_1(X1)|~v2_unialg_1(X2)|~v2_unialg_1(X1)), inference(spm,[status(thm)],[c_0_247, c_0_248])).
cnf(c_0_254, negated_conjecture, (u1_struct_0(esk2_0)!=k1_xboole_0), inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_250]), c_0_251])).
cnf(c_0_255, negated_conjecture, (v2_struct_0(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v1_relat_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~v1_funct_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~l1_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v4_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v3_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v2_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_253]), c_0_236]), c_0_198]), c_0_175]), c_0_176]), c_0_177]), c_0_168]), c_0_178]), c_0_169]), c_0_179]), c_0_170]), c_0_180]), c_0_171]), c_0_181])]), c_0_254]), c_0_182])).
cnf(c_0_256, negated_conjecture, (v2_struct_0(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~m1_alg_1(k6_alg_1(esk1_0,esk2_0,esk3_0),esk1_0)|~v1_relat_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~v1_funct_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~l1_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v4_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v2_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_184]), c_0_178]), c_0_179]), c_0_180]), c_0_181])]), c_0_182])).
cnf(c_0_257, negated_conjecture, (v2_struct_0(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~m1_alg_1(k6_alg_1(esk1_0,esk2_0,esk3_0),esk1_0)|~v1_relat_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~v1_funct_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~l1_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~v4_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_185]), c_0_178]), c_0_179]), c_0_180]), c_0_181])]), c_0_182])).
cnf(c_0_258, plain, (v2_struct_0(X1)|~v2_struct_0(k4_alg_1(X1,X2))|~v2_unialg_1(X1)|~v3_unialg_1(X1)|~v4_unialg_1(X1)|~l1_unialg_1(X1)|~m1_alg_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_91])).
cnf(c_0_259, negated_conjecture, (v2_struct_0(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))|~m1_alg_1(k6_alg_1(esk1_0,esk2_0,esk3_0),esk1_0)|~v1_relat_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~v1_funct_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~l1_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_186]), c_0_178]), c_0_179]), c_0_180]), c_0_181])]), c_0_182])).
cnf(c_0_260, negated_conjecture, (~m1_alg_1(k6_alg_1(esk1_0,esk2_0,esk3_0),esk1_0)|~v1_relat_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~v1_funct_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~l1_unialg_1(k4_alg_1(esk1_0,k6_alg_1(esk1_0,esk2_0,esk3_0)))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_259]), c_0_178]), c_0_179]), c_0_180]), c_0_181])]), c_0_182])).
cnf(c_0_261, negated_conjecture, (~m1_alg_1(k6_alg_1(esk1_0,esk2_0,esk3_0),esk1_0)|~v1_relat_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~v1_funct_1(k7_alg_1(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_103]), c_0_178]), c_0_179]), c_0_180]), c_0_181])]), c_0_182])).
cnf(c_0_262, negated_conjecture, (~v1_relat_1(k7_alg_1(esk1_0,esk2_0,esk3_0))|~v1_funct_1(k7_alg_1(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_53]), c_0_175]), c_0_176]), c_0_177]), c_0_168]), c_0_178]), c_0_169]), c_0_179]), c_0_170]), c_0_180]), c_0_171]), c_0_181])]), c_0_182]), c_0_172])).
cnf(c_0_263, negated_conjecture, (~v1_funct_1(k7_alg_1(esk1_0,esk2_0,esk3_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_84]), c_0_175]), c_0_176]), c_0_177]), c_0_168]), c_0_178]), c_0_169]), c_0_179]), c_0_170]), c_0_180]), c_0_171]), c_0_181])]), c_0_172]), c_0_182])).
cnf(c_0_264, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_96]), c_0_175]), c_0_176]), c_0_177]), c_0_168]), c_0_178]), c_0_169]), c_0_179]), c_0_170]), c_0_180]), c_0_171]), c_0_181])]), c_0_182]), c_0_172]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 265
# Proof object clause steps            : 171
# Proof object formula steps           : 94
# Proof object conjectures             : 43
# Proof object clause conjectures      : 40
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 82
# Proof object initial formulas used   : 47
# Proof object generating inferences   : 88
# Proof object simplifying inferences  : 253
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 105
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 199
# Removed in clause preprocessing      : 10
# Initial clauses in saturation        : 189
# Processed clauses                    : 1387
# ...of these trivial                  : 46
# ...subsumed                          : 360
# ...remaining for further processing  : 981
# Other redundant clauses eliminated   : 23
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 53
# Backward-rewritten                   : 39
# Generated clauses                    : 7955
# ...of the previous two non-trivial   : 7515
# Contextual simplify-reflections      : 252
# Paramodulations                      : 7834
# Factorizations                       : 0
# Equation resolutions                 : 119
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 885
#    Positive orientable unit clauses  : 77
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 28
#    Non-unit-clauses                  : 780
# Current number of unprocessed clauses: 6200
# ...number of literals in the above   : 87852
# Current number of archived formulas  : 0
# Current number of archived clauses   : 94
# Clause-clause subsumption calls (NU) : 1359961
# Rec. Clause-clause subsumption calls : 27323
# Non-unit clause-clause subsumptions  : 579
# Unit Clause-clause subsumption calls : 5600
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 84
# BW rewrite match successes           : 17
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 352275

# -------------------------------------------------
# User time                : 41.830 s
# System time              : 0.680 s
# Total time               : 42.510 s
# Maximum resident set size: 418616 pages
41.83user 0.71system 0:42.55elapsed 99%CPU (0avgtext+0avgdata 418616maxresident)k
0inputs+216outputs (0major+337770minor)pagefaults 0swaps
