2020-06-01 06:16:57.904028: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_8_bb1/premsel_enigma_01_2020_T30_loop02_epoch_389
2020-06-01 06:16:57.963644: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-06-01 06:16:58.091634: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-06-01 06:16:58.556177: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-06-01 06:17:01.192187: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 3288138 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_8_bb1/premsel_enigma_01_2020_T30_loop02_epoch_389' loaded (query=256; context=768; weigths=real; len_mult=0.000000).
# Preprocessing time       : 4.178 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(s2_finset_1__e5_20__topreal6, axiom, ![X1, X2]:((((v2_pre_topc(X1)&l1_pre_topc(X1))&v1_finset_1(X2))&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))=>(((v1_finset_1(X2)&?[X3]:((m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))&X3=k3_tarski(k1_xboole_0))&v2_compts_1(X3,X1)))&![X4, X5]:(((r2_hidden(X4,X2)&r1_tarski(X5,X2))&?[X6]:((m1_subset_1(X6,k1_zfmisc_1(u1_struct_0(X1)))&X6=k3_tarski(X5))&v2_compts_1(X6,X1)))=>?[X7]:((m1_subset_1(X7,k1_zfmisc_1(u1_struct_0(X1)))&X7=k3_tarski(k2_xboole_0(X5,k1_tarski(X4))))&v2_compts_1(X7,X1))))=>?[X8]:((m1_subset_1(X8,k1_zfmisc_1(u1_struct_0(X1)))&X8=k3_tarski(X2))&v2_compts_1(X8,X1)))), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', s2_finset_1__e5_20__topreal6)).
fof(cc1_compts_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>(v1_xboole_0(X2)=>v2_compts_1(X2,X1)))), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', cc1_compts_1)).
fof(dt_k1_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>m1_subset_1(k1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', dt_k1_struct_0)).
fof(t2_zfmisc_1, axiom, k3_tarski(k1_xboole_0)=k1_xboole_0, file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', t2_zfmisc_1)).
fof(fc3_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>v1_xboole_0(k1_struct_0(X1))), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', fc3_struct_0)).
fof(dt_l1_pre_topc, axiom, ![X1]:(l1_pre_topc(X1)=>l1_struct_0(X1)), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', dt_l1_pre_topc)).
fof(t78_zfmisc_1, axiom, ![X1, X2]:k3_tarski(k2_xboole_0(X1,X2))=k2_xboole_0(k3_tarski(X1),k3_tarski(X2)), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', t78_zfmisc_1)).
fof(t25_zfmisc_1, axiom, ![X1]:k3_tarski(k1_tarski(X1))=X1, file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', t25_zfmisc_1)).
fof(d2_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>k1_struct_0(X1)=k1_xboole_0), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', d2_struct_0)).
fof(commutativity_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', commutativity_k2_xboole_0)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', t4_subset)).
fof(dt_k4_subset_1, axiom, ![X1, X2, X3]:((m1_subset_1(X2,k1_zfmisc_1(X1))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>m1_subset_1(k4_subset_1(X1,X2,X3),k1_zfmisc_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', dt_k4_subset_1)).
fof(t6_boole, axiom, ![X1]:(v1_xboole_0(X1)=>X1=k1_xboole_0), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', t6_boole)).
fof(rc2_subset_1, axiom, ![X1]:?[X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))&v1_xboole_0(X2)), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', rc2_subset_1)).
fof(d3_tarski, axiom, ![X1, X2]:(r1_tarski(X1,X2)<=>![X3]:(r2_hidden(X3,X1)=>r2_hidden(X3,X2))), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', d3_tarski)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', t2_subset)).
fof(redefinition_k4_subset_1, axiom, ![X1, X2, X3]:((m1_subset_1(X2,k1_zfmisc_1(X1))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>k4_subset_1(X1,X2,X3)=k2_xboole_0(X2,X3)), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', redefinition_k4_subset_1)).
fof(commutativity_k4_subset_1, axiom, ![X1, X2, X3]:((m1_subset_1(X2,k1_zfmisc_1(X1))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>k4_subset_1(X1,X2,X3)=k4_subset_1(X1,X3,X2)), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', commutativity_k4_subset_1)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', fc1_xboole_0)).
fof(t10_compts_1, axiom, ![X1]:(l1_pre_topc(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>((v2_compts_1(X2,X1)&v2_compts_1(X3,X1))=>v2_compts_1(k4_subset_1(u1_struct_0(X1),X2,X3),X1))))), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', t10_compts_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', t1_subset)).
fof(idempotence_k4_subset_1, axiom, ![X1, X2, X3]:((m1_subset_1(X2,k1_zfmisc_1(X1))&m1_subset_1(X3,k1_zfmisc_1(X1)))=>k4_subset_1(X1,X2,X2)=X2), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', idempotence_k4_subset_1)).
fof(rc3_subset_1, axiom, ![X1]:?[X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))&~(v1_subset_1(X2,X1))), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', rc3_subset_1)).
fof(t20_topreal6, conjecture, ![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((v1_finset_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))=>(![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X3,X2)=>v2_compts_1(X3,X1)))=>v2_compts_1(k5_setfam_1(u1_struct_0(X1),X2),X1)))), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', t20_topreal6)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', t3_subset)).
fof(fc1_subset_1, axiom, ![X1]:~(v1_xboole_0(k1_zfmisc_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', fc1_subset_1)).
fof(t1_boole, axiom, ![X1]:k2_xboole_0(X1,k1_xboole_0)=X1, file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', t1_boole)).
fof(dt_k5_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>m1_subset_1(k5_setfam_1(X1,X2),k1_zfmisc_1(X1))), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', dt_k5_setfam_1)).
fof(redefinition_k5_setfam_1, axiom, ![X1, X2]:(m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))=>k5_setfam_1(X1,X2)=k3_tarski(X2)), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', redefinition_k5_setfam_1)).
fof(reflexivity_r1_tarski, axiom, ![X1, X2]:r1_tarski(X1,X1), file('/scratch/mptp/allpred/preds__0.005/t20_topreal6', reflexivity_r1_tarski)).
fof(c_0_30, plain, ![X12, X13, X14, X18]:((((((m1_subset_1(esk6_2(X12,X13),k1_zfmisc_1(u1_struct_0(X12)))|(r2_hidden(esk3_2(X12,X13),X13)|(~v1_finset_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|X14!=k3_tarski(k1_xboole_0)|~v2_compts_1(X14,X12))))|(~v2_pre_topc(X12)|~l1_pre_topc(X12)|~v1_finset_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X12))))))&(esk6_2(X12,X13)=k3_tarski(X13)|(r2_hidden(esk3_2(X12,X13),X13)|(~v1_finset_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|X14!=k3_tarski(k1_xboole_0)|~v2_compts_1(X14,X12))))|(~v2_pre_topc(X12)|~l1_pre_topc(X12)|~v1_finset_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X12)))))))&(v2_compts_1(esk6_2(X12,X13),X12)|(r2_hidden(esk3_2(X12,X13),X13)|(~v1_finset_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|X14!=k3_tarski(k1_xboole_0)|~v2_compts_1(X14,X12))))|(~v2_pre_topc(X12)|~l1_pre_topc(X12)|~v1_finset_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X12)))))))&(((m1_subset_1(esk6_2(X12,X13),k1_zfmisc_1(u1_struct_0(X12)))|(r1_tarski(esk4_2(X12,X13),X13)|(~v1_finset_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|X14!=k3_tarski(k1_xboole_0)|~v2_compts_1(X14,X12))))|(~v2_pre_topc(X12)|~l1_pre_topc(X12)|~v1_finset_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X12))))))&(esk6_2(X12,X13)=k3_tarski(X13)|(r1_tarski(esk4_2(X12,X13),X13)|(~v1_finset_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|X14!=k3_tarski(k1_xboole_0)|~v2_compts_1(X14,X12))))|(~v2_pre_topc(X12)|~l1_pre_topc(X12)|~v1_finset_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X12)))))))&(v2_compts_1(esk6_2(X12,X13),X12)|(r1_tarski(esk4_2(X12,X13),X13)|(~v1_finset_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|X14!=k3_tarski(k1_xboole_0)|~v2_compts_1(X14,X12))))|(~v2_pre_topc(X12)|~l1_pre_topc(X12)|~v1_finset_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X12))))))))&(((((m1_subset_1(esk6_2(X12,X13),k1_zfmisc_1(u1_struct_0(X12)))|(m1_subset_1(esk5_2(X12,X13),k1_zfmisc_1(u1_struct_0(X12)))|(~v1_finset_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|X14!=k3_tarski(k1_xboole_0)|~v2_compts_1(X14,X12))))|(~v2_pre_topc(X12)|~l1_pre_topc(X12)|~v1_finset_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X12))))))&(esk6_2(X12,X13)=k3_tarski(X13)|(m1_subset_1(esk5_2(X12,X13),k1_zfmisc_1(u1_struct_0(X12)))|(~v1_finset_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|X14!=k3_tarski(k1_xboole_0)|~v2_compts_1(X14,X12))))|(~v2_pre_topc(X12)|~l1_pre_topc(X12)|~v1_finset_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X12)))))))&(v2_compts_1(esk6_2(X12,X13),X12)|(m1_subset_1(esk5_2(X12,X13),k1_zfmisc_1(u1_struct_0(X12)))|(~v1_finset_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|X14!=k3_tarski(k1_xboole_0)|~v2_compts_1(X14,X12))))|(~v2_pre_topc(X12)|~l1_pre_topc(X12)|~v1_finset_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X12)))))))&(((m1_subset_1(esk6_2(X12,X13),k1_zfmisc_1(u1_struct_0(X12)))|(esk5_2(X12,X13)=k3_tarski(esk4_2(X12,X13))|(~v1_finset_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|X14!=k3_tarski(k1_xboole_0)|~v2_compts_1(X14,X12))))|(~v2_pre_topc(X12)|~l1_pre_topc(X12)|~v1_finset_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X12))))))&(esk6_2(X12,X13)=k3_tarski(X13)|(esk5_2(X12,X13)=k3_tarski(esk4_2(X12,X13))|(~v1_finset_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|X14!=k3_tarski(k1_xboole_0)|~v2_compts_1(X14,X12))))|(~v2_pre_topc(X12)|~l1_pre_topc(X12)|~v1_finset_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X12)))))))&(v2_compts_1(esk6_2(X12,X13),X12)|(esk5_2(X12,X13)=k3_tarski(esk4_2(X12,X13))|(~v1_finset_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|X14!=k3_tarski(k1_xboole_0)|~v2_compts_1(X14,X12))))|(~v2_pre_topc(X12)|~l1_pre_topc(X12)|~v1_finset_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X12))))))))&(((m1_subset_1(esk6_2(X12,X13),k1_zfmisc_1(u1_struct_0(X12)))|(v2_compts_1(esk5_2(X12,X13),X12)|(~v1_finset_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|X14!=k3_tarski(k1_xboole_0)|~v2_compts_1(X14,X12))))|(~v2_pre_topc(X12)|~l1_pre_topc(X12)|~v1_finset_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X12))))))&(esk6_2(X12,X13)=k3_tarski(X13)|(v2_compts_1(esk5_2(X12,X13),X12)|(~v1_finset_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|X14!=k3_tarski(k1_xboole_0)|~v2_compts_1(X14,X12))))|(~v2_pre_topc(X12)|~l1_pre_topc(X12)|~v1_finset_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X12)))))))&(v2_compts_1(esk6_2(X12,X13),X12)|(v2_compts_1(esk5_2(X12,X13),X12)|(~v1_finset_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|X14!=k3_tarski(k1_xboole_0)|~v2_compts_1(X14,X12))))|(~v2_pre_topc(X12)|~l1_pre_topc(X12)|~v1_finset_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X12)))))))))&(((m1_subset_1(esk6_2(X12,X13),k1_zfmisc_1(u1_struct_0(X12)))|(~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X12)))|X18!=k3_tarski(k2_xboole_0(esk4_2(X12,X13),k1_tarski(esk3_2(X12,X13))))|~v2_compts_1(X18,X12)|(~v1_finset_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|X14!=k3_tarski(k1_xboole_0)|~v2_compts_1(X14,X12))))|(~v2_pre_topc(X12)|~l1_pre_topc(X12)|~v1_finset_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X12))))))&(esk6_2(X12,X13)=k3_tarski(X13)|(~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X12)))|X18!=k3_tarski(k2_xboole_0(esk4_2(X12,X13),k1_tarski(esk3_2(X12,X13))))|~v2_compts_1(X18,X12)|(~v1_finset_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|X14!=k3_tarski(k1_xboole_0)|~v2_compts_1(X14,X12))))|(~v2_pre_topc(X12)|~l1_pre_topc(X12)|~v1_finset_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X12)))))))&(v2_compts_1(esk6_2(X12,X13),X12)|(~m1_subset_1(X18,k1_zfmisc_1(u1_struct_0(X12)))|X18!=k3_tarski(k2_xboole_0(esk4_2(X12,X13),k1_tarski(esk3_2(X12,X13))))|~v2_compts_1(X18,X12)|(~v1_finset_1(X13)|(~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(X12)))|X14!=k3_tarski(k1_xboole_0)|~v2_compts_1(X14,X12))))|(~v2_pre_topc(X12)|~l1_pre_topc(X12)|~v1_finset_1(X13)|~m1_subset_1(X13,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X12)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[s2_finset_1__e5_20__topreal6])])])])])).
cnf(c_0_31, plain, (esk6_2(X1,X2)=k3_tarski(X2)|esk5_2(X1,X2)=k3_tarski(esk4_2(X1,X2))|~v1_finset_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k1_xboole_0)|~v2_compts_1(X3,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_32, plain, ![X91, X92]:(~l1_pre_topc(X91)|(~m1_subset_1(X92,k1_zfmisc_1(u1_struct_0(X91)))|(~v1_xboole_0(X92)|v2_compts_1(X92,X91)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_compts_1])])])).
fof(c_0_33, plain, ![X60]:(~l1_struct_0(X60)|m1_subset_1(k1_struct_0(X60),k1_zfmisc_1(u1_struct_0(X60)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k1_struct_0])])).
cnf(c_0_34, plain, (esk6_2(X1,X2)=k3_tarski(X2)|esk5_2(X1,X2)=k3_tarski(esk4_2(X1,X2))|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(cn,[status(thm)],[c_0_31])).
cnf(c_0_35, plain, (k3_tarski(k1_xboole_0)=k1_xboole_0), inference(split_conjunct,[status(thm)],[t2_zfmisc_1])).
cnf(c_0_36, plain, (v2_compts_1(X2,X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_37, plain, (m1_subset_1(k1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_38, plain, ![X73]:(~l1_struct_0(X73)|v1_xboole_0(k1_struct_0(X73))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc3_struct_0])])).
fof(c_0_39, plain, ![X90]:(~l1_pre_topc(X90)|l1_struct_0(X90)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_pre_topc])])).
cnf(c_0_40, plain, (esk6_2(X1,X2)=k3_tarski(X2)|k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_34])).
cnf(c_0_41, plain, (k3_tarski(k1_xboole_0)=k1_xboole_0), inference(evalgc,[status(thm)],[c_0_35])).
cnf(c_0_42, plain, (v2_compts_1(X2,X1)|~l1_pre_topc(X1)|~v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[c_0_36])).
cnf(c_0_43, plain, (m1_subset_1(k1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_37])).
cnf(c_0_44, plain, (v1_xboole_0(k1_struct_0(X1))|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_38])).
cnf(c_0_45, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_46, plain, (esk6_2(X1,X2)=k3_tarski(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k2_xboole_0(esk4_2(X1,X2),k1_tarski(esk3_2(X1,X2))))|~v2_compts_1(X3,X1)|~v1_finset_1(X2)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|X4!=k3_tarski(k1_xboole_0)|~v2_compts_1(X4,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_47, plain, ![X27, X28]:k3_tarski(k2_xboole_0(X27,X28))=k2_xboole_0(k3_tarski(X27),k3_tarski(X28)), inference(variable_rename,[status(thm)],[t78_zfmisc_1])).
fof(c_0_48, plain, ![X31]:k3_tarski(k1_tarski(X31))=X31, inference(variable_rename,[status(thm)],[t25_zfmisc_1])).
cnf(c_0_49, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|esk6_2(X1,X2)=k3_tarski(X2)|X3!=k1_xboole_0|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[c_0_40, c_0_41])).
cnf(c_0_50, plain, (v2_compts_1(k1_struct_0(X1),X1)|~v1_xboole_0(k1_struct_0(X1))|~l1_pre_topc(X1)|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_43])])).
cnf(c_0_51, plain, (v1_xboole_0(k1_struct_0(X1))|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_44])).
cnf(c_0_52, plain, (l1_struct_0(X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[c_0_45])).
fof(c_0_53, plain, ![X61]:(~l1_struct_0(X61)|k1_struct_0(X61)=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_struct_0])])).
cnf(c_0_54, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X4!=k3_tarski(k1_xboole_0)|X3!=k3_tarski(k2_xboole_0(esk4_2(X1,X2),k1_tarski(esk3_2(X1,X2))))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(cn,[status(thm)],[c_0_46])).
cnf(c_0_55, plain, (k3_tarski(k2_xboole_0(X1,X2))=k2_xboole_0(k3_tarski(X1),k3_tarski(X2))), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_56, plain, (k3_tarski(k1_tarski(X1))=X1), inference(split_conjunct,[status(thm)],[c_0_48])).
fof(c_0_57, plain, ![X48, X49]:k2_xboole_0(X48,X49)=k2_xboole_0(X49,X48), inference(variable_rename,[status(thm)],[commutativity_k2_xboole_0])).
cnf(c_0_58, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|esk6_2(X1,X2)=k3_tarski(X2)|k1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(k1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_43])])).
cnf(c_0_59, plain, (v2_compts_1(k1_struct_0(X1),X1)|~l1_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_50, c_0_51]), c_0_52])).
cnf(c_0_60, plain, (k1_struct_0(X1)=k1_xboole_0|~l1_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_53])).
cnf(c_0_61, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X4!=k3_tarski(k1_xboole_0)|X3!=k3_tarski(k2_xboole_0(esk4_2(X1,X2),k1_tarski(esk3_2(X1,X2))))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_54])).
cnf(c_0_62, plain, (k2_xboole_0(k3_tarski(X1),k3_tarski(X2))=k3_tarski(k2_xboole_0(X1,X2))), inference(evalgc,[status(thm)],[c_0_55])).
cnf(c_0_63, plain, (k3_tarski(k1_tarski(X1))=X1), inference(evalgc,[status(thm)],[c_0_56])).
cnf(c_0_64, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_57])).
cnf(c_0_65, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|esk6_2(X1,X2)=k3_tarski(X2)|k1_struct_0(X1)!=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_58, c_0_59]), c_0_52])).
cnf(c_0_66, plain, (k1_struct_0(X1)=k1_xboole_0|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_60])).
fof(c_0_67, plain, ![X37, X38, X39]:(~r2_hidden(X37,X38)|~m1_subset_1(X38,k1_zfmisc_1(X39))|m1_subset_1(X37,X39)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_68, plain, ![X51, X52, X53]:(~m1_subset_1(X52,k1_zfmisc_1(X51))|~m1_subset_1(X53,k1_zfmisc_1(X51))|m1_subset_1(k4_subset_1(X51,X52,X53),k1_zfmisc_1(X51))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k4_subset_1])])).
fof(c_0_69, plain, ![X78]:(~v1_xboole_0(X78)|X78=k1_xboole_0), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_boole])])).
fof(c_0_70, plain, ![X66]:(m1_subset_1(esk8_1(X66),k1_zfmisc_1(X66))&v1_xboole_0(esk8_1(X66))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[rc2_subset_1])])).
cnf(c_0_71, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=k3_tarski(k2_xboole_0(esk4_2(X1,X2),k1_tarski(esk3_2(X1,X2))))|X4!=k1_xboole_0|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[c_0_61, c_0_41])).
cnf(c_0_72, plain, (k3_tarski(k2_xboole_0(X1,k1_tarski(X2)))=k2_xboole_0(k3_tarski(X1),X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_62, c_0_63])])).
cnf(c_0_73, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1)), inference(evalgc,[status(thm)],[c_0_64])).
cnf(c_0_74, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|esk6_2(X1,X2)=k3_tarski(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66])])).
fof(c_0_75, plain, ![X20, X21, X22, X23, X24]:((~r1_tarski(X20,X21)|(~r2_hidden(X22,X20)|r2_hidden(X22,X21)))&((r2_hidden(esk7_2(X23,X24),X23)|r1_tarski(X23,X24))&(~r2_hidden(esk7_2(X23,X24),X24)|r1_tarski(X23,X24)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_tarski])])])])])])).
fof(c_0_76, plain, ![X58, X59]:(~m1_subset_1(X58,X59)|(v1_xboole_0(X59)|r2_hidden(X58,X59))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_77, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_67])).
cnf(c_0_78, plain, (m1_subset_1(k4_subset_1(X2,X1,X3),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_79, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(esk3_2(X1,X2),X2)|~v1_finset_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k1_xboole_0)|~v2_compts_1(X3,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_80, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_81, plain, (v1_xboole_0(esk8_1(X1))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_82, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=k2_xboole_0(esk3_2(X1,X2),k3_tarski(esk4_2(X1,X2)))|X4!=k1_xboole_0|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_71, c_0_72]), c_0_73])])).
cnf(c_0_83, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|esk6_2(X1,X2)=k3_tarski(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_74, c_0_52])).
fof(c_0_84, plain, ![X32, X33, X34]:(~m1_subset_1(X33,k1_zfmisc_1(X32))|~m1_subset_1(X34,k1_zfmisc_1(X32))|k4_subset_1(X32,X33,X34)=k2_xboole_0(X33,X34)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k4_subset_1])])).
cnf(c_0_85, plain, (r1_tarski(X1,X2)|~r2_hidden(esk7_2(X1,X2),X2)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_86, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_76])).
cnf(c_0_87, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_77])).
cnf(c_0_88, plain, (m1_subset_1(k4_subset_1(X2,X1,X3),k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_78])).
cnf(c_0_89, plain, (r2_hidden(esk7_2(X1,X2),X1)|r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_90, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(esk3_2(X1,X2),X2)|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(cn,[status(thm)],[c_0_79])).
cnf(c_0_91, plain, (m1_subset_1(esk8_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_70])).
cnf(c_0_92, plain, (X1=k1_xboole_0|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[c_0_80])).
cnf(c_0_93, plain, (v1_xboole_0(esk8_1(X1))), inference(evalgc,[status(thm)],[c_0_81])).
cnf(c_0_94, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=k2_xboole_0(esk3_2(X1,X2),esk5_2(X1,X2))|X4!=k1_xboole_0|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_82, c_0_83])])).
cnf(c_0_95, plain, (k4_subset_1(X2,X1,X3)=k2_xboole_0(X1,X3)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_84])).
cnf(c_0_96, plain, (esk6_2(X1,X2)=k3_tarski(X2)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k1_xboole_0)|~v2_compts_1(X3,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_97, plain, ![X45, X46, X47]:(~m1_subset_1(X46,k1_zfmisc_1(X45))|~m1_subset_1(X47,k1_zfmisc_1(X45))|k4_subset_1(X45,X46,X47)=k4_subset_1(X45,X47,X46)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k4_subset_1])])).
cnf(c_0_98, plain, (r1_tarski(X1,X2)|~r2_hidden(esk7_2(X1,X2),X2)), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_99, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_100, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k4_subset_1(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_88])])).
cnf(c_0_101, plain, (r1_tarski(X1,X2)|r2_hidden(esk7_2(X1,X2),X1)), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_102, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(esk3_2(X1,X2),X2)|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_103, plain, (m1_subset_1(esk8_1(X1),k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_91])).
cnf(c_0_104, plain, (esk8_1(X1)=k1_xboole_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_93])])).
cnf(c_0_105, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_106, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=k2_xboole_0(esk5_2(X1,X2),esk3_2(X1,X2))|X4!=k1_xboole_0|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[c_0_94, c_0_73])).
cnf(c_0_107, plain, (k4_subset_1(X2,X1,X3)=k2_xboole_0(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_95])).
cnf(c_0_108, plain, (esk6_2(X1,X2)=k3_tarski(X2)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(cn,[status(thm)],[c_0_96])).
fof(c_0_109, plain, ![X42, X43, X44]:(~l1_pre_topc(X42)|(~m1_subset_1(X43,k1_zfmisc_1(u1_struct_0(X42)))|(~m1_subset_1(X44,k1_zfmisc_1(u1_struct_0(X42)))|(~v2_compts_1(X43,X42)|~v2_compts_1(X44,X42)|v2_compts_1(k4_subset_1(u1_struct_0(X42),X43,X44),X42))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t10_compts_1])])])).
cnf(c_0_110, plain, (k4_subset_1(X2,X1,X3)=k4_subset_1(X2,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_111, plain, (r2_hidden(X3,X2)|~r1_tarski(X1,X2)|~r2_hidden(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_75])).
cnf(c_0_112, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(esk7_2(X1,X2),X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99])])).
cnf(c_0_113, plain, (m1_subset_1(esk7_2(k4_subset_1(X1,X2,X3),X4),X1)|r1_tarski(k4_subset_1(X1,X2,X3),X4)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_101])])).
cnf(c_0_114, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(esk3_2(X1,X2),X2)|X3!=k1_xboole_0|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[c_0_102, c_0_41])).
cnf(c_0_115, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[c_0_103, c_0_104])])).
cnf(c_0_116, plain, (v1_xboole_0(k1_xboole_0)), inference(evalgc,[status(thm)],[c_0_105])).
cnf(c_0_117, plain, (v2_compts_1(esk6_2(X1,X2),X1)|esk5_2(X1,X2)=k3_tarski(esk4_2(X1,X2))|~v1_finset_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k1_xboole_0)|~v2_compts_1(X3,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_118, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=k1_xboole_0|~v2_compts_1(X3,X1)|~v2_compts_1(k2_xboole_0(esk5_2(X1,X2),esk3_2(X1,X2)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k2_xboole_0(esk5_2(X1,X2),esk3_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_106])])).
cnf(c_0_119, plain, (m1_subset_1(k2_xboole_0(X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_88, c_0_107])])).
cnf(c_0_120, plain, (esk6_2(X1,X2)=k3_tarski(X2)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_108])).
cnf(c_0_121, plain, (v2_compts_1(k4_subset_1(u1_struct_0(X1),X2,X3),X1)|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v2_compts_1(X2,X1)|~v2_compts_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_109])).
cnf(c_0_122, plain, (k4_subset_1(X2,X1,X3)=k4_subset_1(X2,X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_110])).
cnf(c_0_123, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk5_2(X1,X2),X1)|~v1_finset_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k1_xboole_0)|~v2_compts_1(X3,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(split_conjunct,[status(thm)],[c_0_30])).
fof(c_0_124, plain, ![X54, X55]:(~r2_hidden(X54,X55)|m1_subset_1(X54,X55)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_125, plain, (r2_hidden(X3,X2)|~r2_hidden(X3,X1)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_111])).
cnf(c_0_126, plain, (v1_xboole_0(X1)|r1_tarski(k4_subset_1(X1,X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113])])).
cnf(c_0_127, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(esk3_2(X1,X2),X2)|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_115])])).
cnf(c_0_128, plain, (v2_compts_1(k1_xboole_0,X1)|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_42, c_0_115]), c_0_116])])])).
fof(c_0_129, plain, ![X75, X76, X77]:(~m1_subset_1(X76,k1_zfmisc_1(X75))|~m1_subset_1(X77,k1_zfmisc_1(X75))|k4_subset_1(X75,X76,X76)=X76), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[idempotence_k4_subset_1])])).
fof(c_0_130, plain, ![X93]:(m1_subset_1(esk10_1(X93),k1_zfmisc_1(X93))&~v1_subset_1(esk10_1(X93),X93)), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[rc3_subset_1])])])).
cnf(c_0_131, plain, (esk5_2(X1,X2)=k3_tarski(esk4_2(X1,X2))|v2_compts_1(esk6_2(X1,X2),X1)|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(cn,[status(thm)],[c_0_117])).
cnf(c_0_132, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=k1_xboole_0|~v2_compts_1(k2_xboole_0(esk5_2(X1,X2),esk3_2(X1,X2)),X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_119])])).
cnf(c_0_133, plain, (esk6_2(X1,X2)=k3_tarski(X2)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|X3!=k1_xboole_0|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[c_0_120, c_0_41])).
cnf(c_0_134, plain, (v2_compts_1(k4_subset_1(u1_struct_0(X1),X2,X3),X1)|~l1_pre_topc(X1)|~v2_compts_1(X3,X1)|~v2_compts_1(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[c_0_121])).
cnf(c_0_135, plain, (k4_subset_1(X1,X3,X2)=k2_xboole_0(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_122])])).
cnf(c_0_136, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk5_2(X1,X2),X1)|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(cn,[status(thm)],[c_0_123])).
cnf(c_0_137, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_124])).
cnf(c_0_138, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,k4_subset_1(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_126])])).
cnf(c_0_139, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(esk3_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_127, c_0_128])).
cnf(c_0_140, plain, (k4_subset_1(X2,X1,X1)=X1|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_129])).
cnf(c_0_141, plain, (m1_subset_1(esk10_1(X1),k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_130])).
fof(c_0_142, negated_conjecture, ~(![X1]:((v2_pre_topc(X1)&l1_pre_topc(X1))=>![X2]:((v1_finset_1(X2)&m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))=>(![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))=>(r2_hidden(X3,X2)=>v2_compts_1(X3,X1)))=>v2_compts_1(k5_setfam_1(u1_struct_0(X1),X2),X1))))), inference(assume_negation,[status(cth)],[t20_topreal6])).
fof(c_0_143, plain, ![X40, X41]:((~m1_subset_1(X40,k1_zfmisc_1(X41))|r1_tarski(X40,X41))&(~r1_tarski(X40,X41)|m1_subset_1(X40,k1_zfmisc_1(X41)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_144, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|v2_compts_1(esk6_2(X1,X2),X1)|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_131])).
cnf(c_0_145, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=k1_xboole_0|~v2_compts_1(k2_xboole_0(esk5_2(X1,X2),esk3_2(X1,X2)),X1)|~v2_compts_1(X3,X1)|~m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_132, c_0_133])).
cnf(c_0_146, plain, (v2_compts_1(k2_xboole_0(X3,X2),X1)|~v2_compts_1(X3,X1)|~v2_compts_1(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_135])])).
cnf(c_0_147, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk5_2(X1,X2),X1)|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_136])).
cnf(c_0_148, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_137])).
cnf(c_0_149, plain, (v1_xboole_0(X1)|r2_hidden(esk3_2(X2,k4_subset_1(X1,X3,X4)),X1)|esk6_2(X2,k4_subset_1(X1,X3,X4))=k3_tarski(k4_subset_1(X1,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(k4_subset_1(X1,X3,X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_finset_1(k4_subset_1(X1,X3,X4))|~l1_pre_topc(X2)|~v2_pre_topc(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139])])).
cnf(c_0_150, plain, (k4_subset_1(X2,X1,X1)=X1|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_140])).
cnf(c_0_151, plain, (m1_subset_1(esk10_1(X1),k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_141])).
fof(c_0_152, plain, ![X62]:~v1_xboole_0(k1_zfmisc_1(X62)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc1_subset_1])])).
fof(c_0_153, negated_conjecture, ![X11]:((v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0))&((v1_finset_1(esk2_0)&m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))&((~m1_subset_1(X11,k1_zfmisc_1(u1_struct_0(esk1_0)))|(~r2_hidden(X11,esk2_0)|v2_compts_1(X11,esk1_0)))&~v2_compts_1(k5_setfam_1(u1_struct_0(esk1_0),esk2_0),esk1_0)))), inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_142])])])])).
cnf(c_0_154, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_143])).
cnf(c_0_155, plain, (v2_compts_1(esk6_2(X1,X2),X1)|r2_hidden(esk3_2(X1,X2),X2)|~v1_finset_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k1_xboole_0)|~v2_compts_1(X3,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_156, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|v2_compts_1(esk6_2(X1,X2),X1)|X3!=k1_xboole_0|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[c_0_144, c_0_41])).
cnf(c_0_157, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=k1_xboole_0|~v2_compts_1(X3,X1)|~m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v2_compts_1(esk5_2(X1,X2),X1)|~v2_compts_1(esk3_2(X1,X2),X1)|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_146])])).
cnf(c_0_158, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk5_2(X1,X2),X1)|X3!=k1_xboole_0|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[c_0_147, c_0_41])).
cnf(c_0_159, plain, (m1_subset_1(esk3_2(X1,k4_subset_1(X2,X3,X4)),X2)|esk6_2(X1,k4_subset_1(X2,X3,X4))=k3_tarski(k4_subset_1(X2,X3,X4))|v1_xboole_0(X2)|~m1_subset_1(k4_subset_1(X2,X3,X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v1_finset_1(k4_subset_1(X2,X3,X4))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_149])])).
cnf(c_0_160, plain, (k4_subset_1(X1,X2,X2)=X2|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_151])])).
cnf(c_0_161, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_152])).
cnf(c_0_162, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_163, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_154])).
fof(c_0_164, plain, ![X74]:k2_xboole_0(X74,k1_xboole_0)=X74, inference(variable_rename,[status(thm)],[t1_boole])).
cnf(c_0_165, plain, (r2_hidden(esk3_2(X1,X2),X2)|v2_compts_1(esk6_2(X1,X2),X1)|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(cn,[status(thm)],[c_0_155])).
fof(c_0_166, plain, ![X35, X36]:(~m1_subset_1(X36,k1_zfmisc_1(k1_zfmisc_1(X35)))|m1_subset_1(k5_setfam_1(X35,X36),k1_zfmisc_1(X35))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k5_setfam_1])])).
cnf(c_0_167, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|v2_compts_1(esk6_2(X1,X2),X1)|k1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(k1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_43])])).
cnf(c_0_168, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=k1_xboole_0|~v2_compts_1(esk3_2(X1,X2),X1)|~v2_compts_1(X3,X1)|~m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_157, c_0_133]), c_0_158])).
cnf(c_0_169, plain, (esk6_2(X1,X3)=k3_tarski(X3)|v1_xboole_0(X2)|m1_subset_1(esk3_2(X1,X3),X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v1_finset_1(X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_160])])).
cnf(c_0_170, plain, (~v1_xboole_0(k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_161])).
cnf(c_0_171, negated_conjecture, (v2_compts_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_172, negated_conjecture, (v1_finset_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_173, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))), inference(evalgc,[status(thm)],[c_0_162])).
cnf(c_0_174, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_163])])).
cnf(c_0_175, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(split_conjunct,[status(thm)],[c_0_164])).
cnf(c_0_176, plain, (r2_hidden(esk3_2(X1,X2),X2)|v2_compts_1(esk6_2(X1,X2),X1)|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_165])).
cnf(c_0_177, plain, (m1_subset_1(k5_setfam_1(X2,X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_166])).
fof(c_0_178, plain, ![X29, X30]:(~m1_subset_1(X30,k1_zfmisc_1(k1_zfmisc_1(X29)))|k5_setfam_1(X29,X30)=k3_tarski(X30)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k5_setfam_1])])).
cnf(c_0_179, plain, (v2_compts_1(esk6_2(X1,X2),X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k2_xboole_0(esk4_2(X1,X2),k1_tarski(esk3_2(X1,X2))))|~v2_compts_1(X3,X1)|~v1_finset_1(X2)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|X4!=k3_tarski(k1_xboole_0)|~v2_compts_1(X4,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_180, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|v2_compts_1(esk6_2(X1,X2),X1)|k1_struct_0(X1)!=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_167, c_0_59]), c_0_52])).
cnf(c_0_181, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=k1_xboole_0|~v2_compts_1(esk3_2(X1,X2),X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_169]), c_0_170])])).
cnf(c_0_182, negated_conjecture, (v2_compts_1(X1,esk1_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[c_0_171])).
cnf(c_0_183, negated_conjecture, (v1_finset_1(esk2_0)), inference(evalgc,[status(thm)],[c_0_172])).
cnf(c_0_184, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X1,esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_173])])).
cnf(c_0_185, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k2_xboole_0(X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_119])])).
cnf(c_0_186, plain, (k2_xboole_0(X1,k1_xboole_0)=X1), inference(evalgc,[status(thm)],[c_0_175])).
cnf(c_0_187, plain, (v2_compts_1(esk6_2(X1,X2),X1)|r2_hidden(esk3_2(X1,X2),X2)|X3!=k1_xboole_0|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[c_0_176, c_0_41])).
cnf(c_0_188, plain, (m1_subset_1(k5_setfam_1(X2,X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(evalgc,[status(thm)],[c_0_177])).
cnf(c_0_189, plain, (k5_setfam_1(X2,X1)=k3_tarski(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(split_conjunct,[status(thm)],[c_0_178])).
cnf(c_0_190, plain, (v2_compts_1(esk6_2(X1,X2),X1)|X4!=k3_tarski(k1_xboole_0)|X3!=k3_tarski(k2_xboole_0(esk4_2(X1,X2),k1_tarski(esk3_2(X1,X2))))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(cn,[status(thm)],[c_0_179])).
cnf(c_0_191, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|v2_compts_1(esk6_2(X1,X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_66])])).
cnf(c_0_192, negated_conjecture, (l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_193, negated_conjecture, (v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_194, plain, (esk6_2(X1,X2)=k3_tarski(X2)|~v2_compts_1(esk3_2(X1,X2),X1)|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_115])])).
cnf(c_0_195, negated_conjecture, (v2_compts_1(esk3_2(X1,esk2_0),esk1_0)|esk6_2(X1,esk2_0)=k3_tarski(esk2_0)|~m1_subset_1(esk3_2(X1,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_139]), c_0_183])])])).
cnf(c_0_196, negated_conjecture, (m1_subset_1(esk3_2(X1,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|esk6_2(X1,esk2_0)=k3_tarski(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_139]), c_0_183])])])).
cnf(c_0_197, negated_conjecture, (~v2_compts_1(k5_setfam_1(u1_struct_0(esk1_0),esk2_0),esk1_0)), inference(split_conjunct,[status(thm)],[c_0_153])).
cnf(c_0_198, negated_conjecture, (r2_hidden(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X1,esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_173])])).
cnf(c_0_199, plain, (r2_hidden(esk7_2(k2_xboole_0(X1,X2),X3),X4)|r1_tarski(k2_xboole_0(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(X4))|~m1_subset_1(X1,k1_zfmisc_1(X4))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_101])])).
cnf(c_0_200, plain, (k2_xboole_0(k1_xboole_0,X1)=X1), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_73])])).
cnf(c_0_201, plain, (v2_compts_1(esk6_2(X1,X2),X1)|r2_hidden(esk3_2(X1,X2),X2)|k5_setfam_1(u1_struct_0(X1),X3)!=k1_xboole_0|~v2_compts_1(k5_setfam_1(u1_struct_0(X1),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_188])])).
cnf(c_0_202, plain, (k5_setfam_1(X2,X1)=k3_tarski(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2)))), inference(evalgc,[status(thm)],[c_0_189])).
cnf(c_0_203, plain, (v2_compts_1(esk6_2(X1,X2),X1)|X4!=k3_tarski(k1_xboole_0)|X3!=k3_tarski(k2_xboole_0(esk4_2(X1,X2),k1_tarski(esk3_2(X1,X2))))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_190])).
cnf(c_0_204, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|v2_compts_1(esk6_2(X1,X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_191, c_0_52])).
cnf(c_0_205, negated_conjecture, (l1_pre_topc(esk1_0)), inference(evalgc,[status(thm)],[c_0_192])).
cnf(c_0_206, negated_conjecture, (v2_pre_topc(esk1_0)), inference(evalgc,[status(thm)],[c_0_193])).
cnf(c_0_207, plain, (esk6_2(X1,X2)=k3_tarski(X2)|~v2_compts_1(esk3_2(X1,X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_194, c_0_128])).
cnf(c_0_208, negated_conjecture, (esk6_2(X1,esk2_0)=k3_tarski(esk2_0)|v2_compts_1(esk3_2(X1,esk2_0),esk1_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_196])])).
cnf(c_0_209, negated_conjecture, (~v2_compts_1(k5_setfam_1(u1_struct_0(esk1_0),esk2_0),esk1_0)), inference(evalgc,[status(thm)],[c_0_197])).
cnf(c_0_210, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_143])).
cnf(c_0_211, negated_conjecture, (r1_tarski(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(esk7_2(X1,k1_zfmisc_1(u1_struct_0(esk1_0))),esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_198])])).
cnf(c_0_212, plain, (r1_tarski(X1,X2)|r2_hidden(esk7_2(X1,X2),X3)|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_200]), c_0_115])])])).
cnf(c_0_213, plain, (v2_compts_1(esk6_2(X1,X2),X1)|r2_hidden(esk3_2(X1,X2),X2)|k3_tarski(X3)!=k1_xboole_0|~v2_compts_1(k3_tarski(X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_202])])).
fof(c_0_214, plain, ![X50]:r1_tarski(X50,X50), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r1_tarski])])).
cnf(c_0_215, plain, (v2_compts_1(esk6_2(X1,X2),X1)|X3!=k3_tarski(k2_xboole_0(esk4_2(X1,X2),k1_tarski(esk3_2(X1,X2))))|X4!=k1_xboole_0|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[c_0_203, c_0_41])).
cnf(c_0_216, negated_conjecture, (k3_tarski(esk4_2(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)|v2_compts_1(esk6_2(esk1_0,esk2_0),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_173]), c_0_183]), c_0_205]), c_0_206])])])).
cnf(c_0_217, negated_conjecture, (esk6_2(esk1_0,esk2_0)=k3_tarski(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_208]), c_0_173]), c_0_183]), c_0_205]), c_0_206])])])).
cnf(c_0_218, negated_conjecture, (~v2_compts_1(k3_tarski(esk2_0),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_202]), c_0_173])])])).
cnf(c_0_219, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_210])).
cnf(c_0_220, negated_conjecture, (r1_tarski(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_212])])).
cnf(c_0_221, plain, (v2_compts_1(esk6_2(X1,X2),X1)|r2_hidden(esk3_2(X1,X2),X2)|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_41]), c_0_115])])])).
cnf(c_0_222, plain, (r1_tarski(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_214])).
cnf(c_0_223, plain, (v2_compts_1(esk6_2(X1,X2),X1)|X3!=k2_xboole_0(esk3_2(X1,X2),k3_tarski(esk4_2(X1,X2)))|X4!=k1_xboole_0|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_215, c_0_72]), c_0_73])])).
cnf(c_0_224, negated_conjecture, (k3_tarski(esk4_2(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)), inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_216, c_0_217]), c_0_218])).
cnf(c_0_225, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_220])])).
cnf(c_0_226, plain, (v2_compts_1(esk6_2(X1,X2),X1)|r2_hidden(esk3_2(X1,X2),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_221, c_0_128])).
cnf(c_0_227, plain, (r1_tarski(X1,X1)), inference(evalgc,[status(thm)],[c_0_222])).
cnf(c_0_228, plain, (v2_compts_1(esk6_2(X1,X2),X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k1_xboole_0)|~v2_compts_1(X3,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_229, negated_conjecture, (X1!=k2_xboole_0(esk3_2(esk1_0,esk2_0),esk5_2(esk1_0,esk2_0))|X2!=k1_xboole_0|~v2_compts_1(X2,esk1_0)|~v2_compts_1(X1,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_224]), c_0_217]), c_0_173]), c_0_183]), c_0_205]), c_0_206])]), c_0_218])])).
cnf(c_0_230, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_225])])).
cnf(c_0_231, negated_conjecture, (v2_compts_1(esk6_2(esk1_0,esk2_0),esk1_0)|r2_hidden(esk3_2(esk1_0,esk2_0),esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_173]), c_0_183]), c_0_205]), c_0_206])])])).
cnf(c_0_232, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_227])])).
cnf(c_0_233, plain, (v2_compts_1(esk6_2(X1,X2),X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(cn,[status(thm)],[c_0_228])).
cnf(c_0_234, plain, (v2_compts_1(esk6_2(X1,X2),X1)|v2_compts_1(esk5_2(X1,X2),X1)|~v1_finset_1(X2)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k1_xboole_0)|~v2_compts_1(X3,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_235, negated_conjecture, (X1!=k2_xboole_0(esk5_2(esk1_0,esk2_0),esk3_2(esk1_0,esk2_0))|X2!=k1_xboole_0|~v2_compts_1(X2,esk1_0)|~v2_compts_1(X1,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(rw,[status(thm)],[c_0_229, c_0_73])).
cnf(c_0_236, negated_conjecture, (m1_subset_1(esk3_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|v2_compts_1(esk6_2(esk1_0,esk2_0),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_231]), c_0_232])])])).
cnf(c_0_237, negated_conjecture, (v2_compts_1(esk3_2(esk1_0,esk2_0),esk1_0)|v2_compts_1(esk6_2(esk1_0,esk2_0),esk1_0)|~m1_subset_1(esk3_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_231])])).
cnf(c_0_238, plain, (v2_compts_1(esk6_2(X1,X2),X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_233])).
cnf(c_0_239, plain, (v2_compts_1(esk5_2(X1,X2),X1)|v2_compts_1(esk6_2(X1,X2),X1)|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(cn,[status(thm)],[c_0_234])).
cnf(c_0_240, negated_conjecture, (X1!=k1_xboole_0|~v2_compts_1(X1,esk1_0)|~v2_compts_1(k2_xboole_0(esk5_2(esk1_0,esk2_0),esk3_2(esk1_0,esk2_0)),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_xboole_0(esk5_2(esk1_0,esk2_0),esk3_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_235])])).
cnf(c_0_241, negated_conjecture, (m1_subset_1(esk3_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_236, c_0_217]), c_0_218])])).
cnf(c_0_242, negated_conjecture, (v2_compts_1(esk6_2(esk1_0,esk2_0),esk1_0)|v2_compts_1(esk3_2(esk1_0,esk2_0),esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_236])])).
cnf(c_0_243, plain, (v2_compts_1(esk6_2(X1,X2),X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|X3!=k1_xboole_0|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[c_0_238, c_0_41])).
cnf(c_0_244, plain, (v2_compts_1(esk5_2(X1,X2),X1)|v2_compts_1(esk6_2(X1,X2),X1)|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_239])).
cnf(c_0_245, negated_conjecture, (X1!=k1_xboole_0|~v2_compts_1(k2_xboole_0(esk5_2(esk1_0,esk2_0),esk3_2(esk1_0,esk2_0)),esk1_0)|~v2_compts_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_119]), c_0_241])])])).
cnf(c_0_246, negated_conjecture, (v2_compts_1(esk3_2(esk1_0,esk2_0),esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(rw,[status(thm)],[c_0_242, c_0_217]), c_0_218])])).
cnf(c_0_247, plain, (v2_compts_1(esk6_2(X1,X2),X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_115])])).
cnf(c_0_248, plain, (v2_compts_1(esk6_2(X1,X2),X1)|v2_compts_1(esk5_2(X1,X2),X1)|X3!=k1_xboole_0|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(rw,[status(thm)],[c_0_244, c_0_41])).
cnf(c_0_249, negated_conjecture, (X1!=k1_xboole_0|~v2_compts_1(X1,esk1_0)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_compts_1(esk5_2(esk1_0,esk2_0),esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_146]), c_0_246]), c_0_241]), c_0_205])])])).
cnf(c_0_250, plain, (v2_compts_1(esk6_2(X1,X2),X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_247, c_0_128])).
cnf(c_0_251, plain, (v2_compts_1(esk6_2(X1,X2),X1)|v2_compts_1(esk5_2(X1,X2),X1)|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_248, c_0_115])])).
cnf(c_0_252, negated_conjecture, (X1!=k1_xboole_0|~v2_compts_1(esk5_2(esk1_0,esk2_0),esk1_0)|~v2_compts_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_250]), c_0_217]), c_0_173]), c_0_183]), c_0_205]), c_0_206])]), c_0_218])])).
cnf(c_0_253, plain, (v2_compts_1(esk5_2(X1,X2),X1)|v2_compts_1(esk6_2(X1,X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_251, c_0_128])).
cnf(c_0_254, negated_conjecture, (X1!=k1_xboole_0|~v2_compts_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_253]), c_0_217]), c_0_173]), c_0_183]), c_0_205]), c_0_206])]), c_0_218])])).
cnf(c_0_255, negated_conjecture, (v2_compts_1(k1_xboole_0,esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_205])])).
cnf(c_0_256, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_115]), c_0_255])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 257
# Proof object clause steps            : 198
# Proof object formula steps           : 59
# Proof object conjectures             : 43
# Proof object clause conjectures      : 40
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 48
# Proof object initial formulas used   : 30
# Proof object generating inferences   : 62
# Proof object simplifying inferences  : 106
# Training examples: 114 positive, 1026 negative
# Training: Positive examples begin
cnf(c_0_41, plain, (k3_tarski(k1_xboole_0)=k1_xboole_0)).# trainpos
cnf(c_0_206, negated_conjecture, (v2_pre_topc(esk1_0))).# trainpos
cnf(c_0_205, negated_conjecture, (l1_pre_topc(esk1_0))).# trainpos
cnf(c_0_183, negated_conjecture, (v1_finset_1(esk2_0))).# trainpos
cnf(c_0_116, plain, (v1_xboole_0(k1_xboole_0))).# trainpos
cnf(c_0_209, negated_conjecture, (~v2_compts_1(k5_setfam_1(u1_struct_0(esk1_0),esk2_0),esk1_0))).# trainpos
cnf(c_0_173, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).# trainpos
cnf(c_0_170, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_52, plain, (l1_struct_0(X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_92, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).# trainpos
cnf(c_0_93, plain, (v1_xboole_0(esk8_1(X1)))).# trainpos
cnf(c_0_63, plain, (k3_tarski(k1_tarski(X1))=X1)).# trainpos
cnf(c_0_186, plain, (k2_xboole_0(X1,k1_xboole_0)=X1)).# trainpos
cnf(c_0_227, plain, (r1_tarski(X1,X1))).# trainpos
cnf(c_0_66, plain, (k1_struct_0(X1)=k1_xboole_0|~l1_struct_0(X1))).# trainpos
cnf(c_0_148, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_182, negated_conjecture, (v2_compts_1(X1,esk1_0)|~r2_hidden(X1,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_51, plain, (v1_xboole_0(k1_struct_0(X1))|~l1_struct_0(X1))).# trainpos
cnf(c_0_103, plain, (m1_subset_1(esk8_1(X1),k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_151, plain, (m1_subset_1(esk10_1(X1),k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_73, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1))).# trainpos
cnf(c_0_219, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_163, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_62, plain, (k2_xboole_0(k3_tarski(X1),k3_tarski(X2))=k3_tarski(k2_xboole_0(X1,X2)))).# trainpos
cnf(c_0_87, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_61, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X4!=k3_tarski(k1_xboole_0)|X3!=k3_tarski(k2_xboole_0(esk4_2(X1,X2),k1_tarski(esk3_2(X1,X2))))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).# trainpos
cnf(c_0_99, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_43, plain, (m1_subset_1(k1_struct_0(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).# trainpos
cnf(c_0_101, plain, (r1_tarski(X1,X2)|r2_hidden(esk7_2(X1,X2),X1))).# trainpos
cnf(c_0_98, plain, (r1_tarski(X1,X2)|~r2_hidden(esk7_2(X1,X2),X2))).# trainpos
cnf(c_0_202, plain, (k5_setfam_1(X2,X1)=k3_tarski(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).# trainpos
cnf(c_0_125, plain, (r2_hidden(X3,X2)|~r2_hidden(X3,X1)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_188, plain, (m1_subset_1(k5_setfam_1(X2,X1),k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).# trainpos
cnf(c_0_150, plain, (k4_subset_1(X2,X1,X1)=X1|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_42, plain, (v2_compts_1(X2,X1)|~l1_pre_topc(X1)|~v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_88, plain, (m1_subset_1(k4_subset_1(X2,X1,X3),k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_203, plain, (v2_compts_1(esk6_2(X1,X2),X1)|X4!=k3_tarski(k1_xboole_0)|X3!=k3_tarski(k2_xboole_0(esk4_2(X1,X2),k1_tarski(esk3_2(X1,X2))))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).# trainpos
cnf(c_0_107, plain, (k4_subset_1(X2,X1,X3)=k2_xboole_0(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_134, plain, (v2_compts_1(k4_subset_1(u1_struct_0(X1),X2,X3),X1)|~l1_pre_topc(X1)|~v2_compts_1(X3,X1)|~v2_compts_1(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_122, plain, (k4_subset_1(X2,X1,X3)=k4_subset_1(X2,X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_102, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(esk3_2(X1,X2),X2)|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).# trainpos
cnf(c_0_147, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk5_2(X1,X2),X1)|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).# trainpos
cnf(c_0_40, plain, (esk6_2(X1,X2)=k3_tarski(X2)|k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).# trainpos
cnf(c_0_176, plain, (r2_hidden(esk3_2(X1,X2),X2)|v2_compts_1(esk6_2(X1,X2),X1)|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).# trainpos
cnf(c_0_244, plain, (v2_compts_1(esk5_2(X1,X2),X1)|v2_compts_1(esk6_2(X1,X2),X1)|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).# trainpos
cnf(c_0_120, plain, (esk6_2(X1,X2)=k3_tarski(X2)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).# trainpos
cnf(c_0_144, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|v2_compts_1(esk6_2(X1,X2),X1)|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).# trainpos
cnf(c_0_238, plain, (v2_compts_1(esk6_2(X1,X2),X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).# trainpos
cnf(c_0_218, negated_conjecture, (~v2_compts_1(k3_tarski(esk2_0),esk1_0))).# trainpos
cnf(c_0_104, plain, (esk8_1(X1)=k1_xboole_0)).# trainpos
cnf(c_0_184, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X1,esk2_0))).# trainpos
cnf(c_0_174, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_72, plain, (k3_tarski(k2_xboole_0(X1,k1_tarski(X2)))=k2_xboole_0(k3_tarski(X1),X2))).# trainpos
cnf(c_0_232, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_112, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(esk7_2(X1,X2),X2))).# trainpos
cnf(c_0_160, plain, (k4_subset_1(X1,X2,X2)=X2|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_100, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k4_subset_1(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_135, plain, (k4_subset_1(X1,X3,X2)=k2_xboole_0(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_200, plain, (k2_xboole_0(k1_xboole_0,X1)=X1)).# trainpos
cnf(c_0_119, plain, (m1_subset_1(k2_xboole_0(X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_82, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=k2_xboole_0(esk3_2(X1,X2),k3_tarski(esk4_2(X1,X2)))|X4!=k1_xboole_0|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_146, plain, (v2_compts_1(k2_xboole_0(X3,X2),X1)|~v2_compts_1(X3,X1)|~v2_compts_1(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).# trainpos
cnf(c_0_198, negated_conjecture, (r2_hidden(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X1,esk2_0))).# trainpos
cnf(c_0_115, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_223, plain, (v2_compts_1(esk6_2(X1,X2),X1)|X3!=k2_xboole_0(esk3_2(X1,X2),k3_tarski(esk4_2(X1,X2)))|X4!=k1_xboole_0|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_113, plain, (m1_subset_1(esk7_2(k4_subset_1(X1,X2,X3),X4),X1)|r1_tarski(k4_subset_1(X1,X2,X3),X4)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_211, negated_conjecture, (r1_tarski(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(esk7_2(X1,k1_zfmisc_1(u1_struct_0(esk1_0))),esk2_0))).# trainpos
cnf(c_0_126, plain, (v1_xboole_0(X1)|r1_tarski(k4_subset_1(X1,X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_128, plain, (v2_compts_1(k1_xboole_0,X1)|~l1_pre_topc(X1))).# trainpos
cnf(c_0_185, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k2_xboole_0(X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_127, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(esk3_2(X1,X2),X2)|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_255, negated_conjecture, (v2_compts_1(k1_xboole_0,esk1_0))).# trainpos
cnf(c_0_195, negated_conjecture, (v2_compts_1(esk3_2(X1,esk2_0),esk1_0)|esk6_2(X1,esk2_0)=k3_tarski(esk2_0)|~m1_subset_1(esk3_2(X1,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_196, negated_conjecture, (m1_subset_1(esk3_2(X1,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|esk6_2(X1,esk2_0)=k3_tarski(esk2_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_138, plain, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,k4_subset_1(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_199, plain, (r2_hidden(esk7_2(k2_xboole_0(X1,X2),X3),X4)|r1_tarski(k2_xboole_0(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(X4))|~m1_subset_1(X1,k1_zfmisc_1(X4)))).# trainpos
cnf(c_0_208, negated_conjecture, (esk6_2(X1,esk2_0)=k3_tarski(esk2_0)|v2_compts_1(esk3_2(X1,esk2_0),esk1_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_212, plain, (r1_tarski(X1,X2)|r2_hidden(esk7_2(X1,X2),X3)|~m1_subset_1(X1,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_220, negated_conjecture, (r1_tarski(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).# trainpos
cnf(c_0_149, plain, (v1_xboole_0(X1)|r2_hidden(esk3_2(X2,k4_subset_1(X1,X3,X4)),X1)|esk6_2(X2,k4_subset_1(X1,X3,X4))=k3_tarski(k4_subset_1(X1,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(k4_subset_1(X1,X3,X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_finset_1(k4_subset_1(X1,X3,X4))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).# trainpos
cnf(c_0_225, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).# trainpos
cnf(c_0_50, plain, (v2_compts_1(k1_struct_0(X1),X1)|~v1_xboole_0(k1_struct_0(X1))|~l1_pre_topc(X1)|~l1_struct_0(X1))).# trainpos
cnf(c_0_159, plain, (m1_subset_1(esk3_2(X1,k4_subset_1(X2,X3,X4)),X2)|esk6_2(X1,k4_subset_1(X2,X3,X4))=k3_tarski(k4_subset_1(X2,X3,X4))|v1_xboole_0(X2)|~m1_subset_1(k4_subset_1(X2,X3,X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v1_finset_1(k4_subset_1(X2,X3,X4))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_169, plain, (esk6_2(X1,X3)=k3_tarski(X3)|v1_xboole_0(X2)|m1_subset_1(esk3_2(X1,X3),X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v1_finset_1(X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_230, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(esk2_0)))).# trainpos
cnf(c_0_58, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|esk6_2(X1,X2)=k3_tarski(X2)|k1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(k1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).# trainpos
cnf(c_0_74, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|esk6_2(X1,X2)=k3_tarski(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).# trainpos
cnf(c_0_201, plain, (v2_compts_1(esk6_2(X1,X2),X1)|r2_hidden(esk3_2(X1,X2),X2)|k5_setfam_1(u1_struct_0(X1),X3)!=k1_xboole_0|~v2_compts_1(k5_setfam_1(u1_struct_0(X1),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).# trainpos
cnf(c_0_94, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=k2_xboole_0(esk3_2(X1,X2),esk5_2(X1,X2))|X4!=k1_xboole_0|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_213, plain, (v2_compts_1(esk6_2(X1,X2),X1)|r2_hidden(esk3_2(X1,X2),X2)|k3_tarski(X3)!=k1_xboole_0|~v2_compts_1(k3_tarski(X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_221, plain, (v2_compts_1(esk6_2(X1,X2),X1)|r2_hidden(esk3_2(X1,X2),X2)|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_118, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=k1_xboole_0|~v2_compts_1(X3,X1)|~v2_compts_1(k2_xboole_0(esk5_2(X1,X2),esk3_2(X1,X2)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k2_xboole_0(esk5_2(X1,X2),esk3_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_231, negated_conjecture, (v2_compts_1(esk6_2(esk1_0,esk2_0),esk1_0)|r2_hidden(esk3_2(esk1_0,esk2_0),esk2_0))).# trainpos
cnf(c_0_132, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=k1_xboole_0|~v2_compts_1(k2_xboole_0(esk5_2(X1,X2),esk3_2(X1,X2)),X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_237, negated_conjecture, (v2_compts_1(esk3_2(esk1_0,esk2_0),esk1_0)|v2_compts_1(esk6_2(esk1_0,esk2_0),esk1_0)|~m1_subset_1(esk3_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_236, negated_conjecture, (m1_subset_1(esk3_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|v2_compts_1(esk6_2(esk1_0,esk2_0),esk1_0))).# trainpos
cnf(c_0_157, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=k1_xboole_0|~v2_compts_1(X3,X1)|~m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v2_compts_1(esk5_2(X1,X2),X1)|~v2_compts_1(esk3_2(X1,X2),X1)|~m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_167, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|v2_compts_1(esk6_2(X1,X2),X1)|k1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(k1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).# trainpos
cnf(c_0_242, negated_conjecture, (v2_compts_1(esk6_2(esk1_0,esk2_0),esk1_0)|v2_compts_1(esk3_2(esk1_0,esk2_0),esk1_0))).# trainpos
cnf(c_0_191, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|v2_compts_1(esk6_2(X1,X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).# trainpos
cnf(c_0_181, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=k1_xboole_0|~v2_compts_1(esk3_2(X1,X2),X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_194, plain, (esk6_2(X1,X2)=k3_tarski(X2)|~v2_compts_1(esk3_2(X1,X2),X1)|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_217, negated_conjecture, (esk6_2(esk1_0,esk2_0)=k3_tarski(esk2_0))).# trainpos
cnf(c_0_246, negated_conjecture, (v2_compts_1(esk3_2(esk1_0,esk2_0),esk1_0))).# trainpos
cnf(c_0_241, negated_conjecture, (m1_subset_1(esk3_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_216, negated_conjecture, (k3_tarski(esk4_2(esk1_0,esk2_0))=esk5_2(esk1_0,esk2_0)|v2_compts_1(esk6_2(esk1_0,esk2_0),esk1_0))).# trainpos
cnf(c_0_229, negated_conjecture, (X1!=k2_xboole_0(esk3_2(esk1_0,esk2_0),esk5_2(esk1_0,esk2_0))|X2!=k1_xboole_0|~v2_compts_1(X2,esk1_0)|~v2_compts_1(X1,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_240, negated_conjecture, (X1!=k1_xboole_0|~v2_compts_1(X1,esk1_0)|~v2_compts_1(k2_xboole_0(esk5_2(esk1_0,esk2_0),esk3_2(esk1_0,esk2_0)),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_xboole_0(esk5_2(esk1_0,esk2_0),esk3_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_245, negated_conjecture, (X1!=k1_xboole_0|~v2_compts_1(k2_xboole_0(esk5_2(esk1_0,esk2_0),esk3_2(esk1_0,esk2_0)),esk1_0)|~v2_compts_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_249, negated_conjecture, (X1!=k1_xboole_0|~v2_compts_1(X1,esk1_0)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v2_compts_1(esk5_2(esk1_0,esk2_0),esk1_0))).# trainpos
cnf(c_0_251, plain, (v2_compts_1(esk6_2(X1,X2),X1)|v2_compts_1(esk5_2(X1,X2),X1)|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_247, plain, (v2_compts_1(esk6_2(X1,X2),X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).# trainpos
cnf(c_0_252, negated_conjecture, (X1!=k1_xboole_0|~v2_compts_1(esk5_2(esk1_0,esk2_0),esk1_0)|~v2_compts_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_254, negated_conjecture, (X1!=k1_xboole_0|~v2_compts_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_55, plain, (o_0_0_xboole_0=k1_xboole_0)).#trainneg
cnf(i_0_121, plain, (l1_pre_topc(esk16_0))).#trainneg
cnf(i_0_74, plain, (v1_xboole_0(esk11_0))).#trainneg
cnf(i_0_88, plain, (v1_xboole_0(esk14_0))).#trainneg
cnf(i_0_108, plain, (v1_xboole_0(esk15_0))).#trainneg
cnf(i_0_105, plain, (v1_xreal_0(esk15_0))).#trainneg
cnf(i_0_122, plain, (v1_xreal_0(esk17_0))).#trainneg
cnf(i_0_44, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_87, plain, (v1_xxreal_0(esk14_0))).#trainneg
cnf(i_0_106, plain, (v1_xxreal_0(esk15_0))).#trainneg
cnf(i_0_107, plain, (v1_xcmplx_0(esk15_0))).#trainneg
cnf(i_0_76, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_78, plain, (v1_finset_1(k1_tarski(X1)))).#trainneg
cnf(i_0_61, plain, (k2_xboole_0(X1,X1)=X1)).#trainneg
cnf(i_0_50, plain, (~v1_xboole_0(k1_tarski(X1)))).#trainneg
cnf(i_0_67, plain, (m1_subset_1(esk9_1(X1),X1))).#trainneg
cnf(i_0_71, plain, (~v1_subset_1(esk10_1(X1),X1))).#trainneg
cnf(i_0_65, plain, (v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_100, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_81, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_77, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_75, plain, (v2_compts_1(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_80, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_118, plain, (v2_membered(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_54, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_115, plain, (v5_xxreal_2(X1)|~v1_finset_1(X1)|~v3_membered(X1))).#trainneg
cnf(i_0_109, plain, (v5_finset_1(k1_zfmisc_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_79, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_117, plain, (v3_ordinal1(k3_tarski(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_68, plain, (m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_83, plain, (v2_tops_1(esk12_1(X1),X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_63, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_xboole_0(X2,X1)))).#trainneg
cnf(i_0_64, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_xboole_0(X1,X2)))).#trainneg
cnf(i_0_119, plain, (v1_finset_1(X2)|~v5_finset_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_102, plain, (v1_xreal_0(X2)|~v3_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_53, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_111, plain, (v1_xboole_0(X1)|v1_xxreal_2(X1)|~v1_finset_1(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_110, plain, (v1_xboole_0(X1)|v2_xxreal_2(X1)|~v1_finset_1(X1)|~v2_membered(X1))).#trainneg
cnf(i_0_120, plain, (v1_xxreal_0(X2)|~v2_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_82, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_89, plain, (v2_tops_1(X2,X1)|~l1_pre_topc(X1)|~v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_101, plain, (v2_connsp_1(X2,X1)|~l1_pre_topc(X1)|~v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_91, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_85, plain, (v3_tops_1(esk13_1(X1),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_84, plain, (m1_subset_1(esk12_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_49, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_104, plain, (v3_membered(X2)|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_66, plain, (v1_finset_1(k2_xboole_0(X1,X2))|~v1_finset_1(X2)|~v1_finset_1(X1))).#trainneg
cnf(i_0_86, plain, (m1_subset_1(esk13_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~v2_pre_topc(X1)|~l1_pre_topc(X1))).#trainneg
cnf(i_0_73, plain, (v3_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_90, plain, (v4_pre_topc(X2,X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_103, plain, (v3_pre_topc(k2_xboole_0(X2,X3),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v3_pre_topc(X3,X1)|~v3_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_114, plain, (v4_pre_topc(k2_xboole_0(X2,X3),X1)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v4_pre_topc(X3,X1)|~v4_pre_topc(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_20, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(esk4_2(X1,X2),X2)|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_19, plain, (v2_compts_1(esk6_2(X1,X2),X1)|r1_tarski(esk4_2(X1,X2),X2)|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_24, plain, (r2_hidden(esk3_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_12, plain, (v2_compts_1(esk5_2(X1,X2),X1)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_21, plain, (r1_tarski(esk4_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_15, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_9, plain, (m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|X4!=k3_tarski(k1_xboole_0)|X3!=k3_tarski(k2_xboole_0(esk4_2(X1,X2),k1_tarski(esk3_2(X1,X2))))|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_18, plain, (m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|X3!=k3_tarski(k1_xboole_0)|~v2_pre_topc(X1)|~l1_pre_topc(X1)|~v1_finset_1(X2)|~v2_compts_1(X3,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_123, plain, (esk15_0=k1_xboole_0)).#trainneg
cnf(i_0_124, plain, (esk14_0=k1_xboole_0)).#trainneg
cnf(i_0_125, plain, (esk11_0=k1_xboole_0)).#trainneg
cnf(i_0_131, plain, (v1_finset_1(k1_xboole_0))).#trainneg
cnf(i_0_133, plain, (v1_xreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_134, plain, (v1_xxreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_132, plain, (v2_compts_1(esk9_1(k1_zfmisc_1(u1_struct_0(X1))),X1)|~v1_finset_1(esk9_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_306, plain, (v2_compts_1(k4_subset_1(u1_struct_0(X1),X2,X3),X1)|~v1_finset_1(k4_subset_1(u1_struct_0(X1),X2,X3))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_277, plain, (v2_compts_1(k5_setfam_1(u1_struct_0(X1),X2),X1)|~v1_finset_1(k5_setfam_1(u1_struct_0(X1),X2))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_234, negated_conjecture, (v2_compts_1(esk7_2(esk2_0,X1),esk1_0)|r1_tarski(esk2_0,X1)|~m1_subset_1(esk7_2(esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_191, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk10_1(X2)))).#trainneg
cnf(i_0_250, negated_conjecture, (v3_membered(esk2_0)|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_137, plain, (v1_finset_1(k1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_184, negated_conjecture, (v1_xreal_0(esk2_0)|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_203, negated_conjecture, (v1_xxreal_0(esk2_0)|~v2_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_279, plain, (v1_finset_1(k5_setfam_1(X1,X2))|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_144, plain, (X1=k1_struct_0(X2)|~v1_xboole_0(X1)|~l1_struct_0(X2))).#trainneg
cnf(i_0_283, plain, (m1_subset_1(k3_tarski(X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_224, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_386, plain, (r1_tarski(esk4_2(X1,X2),X2)|v2_compts_1(esk6_2(X1,X2),X1)|k4_subset_1(u1_struct_0(X1),X3,X4)!=k1_xboole_0|~v2_compts_1(k4_subset_1(u1_struct_0(X1),X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_354, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(esk4_2(X1,X2),X2)|k4_subset_1(u1_struct_0(X1),X3,X4)!=k1_xboole_0|~v2_compts_1(k4_subset_1(u1_struct_0(X1),X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_387, plain, (r1_tarski(esk4_2(X1,X2),X2)|v2_compts_1(esk6_2(X1,X2),X1)|k5_setfam_1(u1_struct_0(X1),X3)!=k1_xboole_0|~v2_compts_1(k5_setfam_1(u1_struct_0(X1),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_197, negated_conjecture, (v2_compts_1(X1,esk1_0)|v1_xboole_0(esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_338, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(esk3_2(X1,X2),X2)|k4_subset_1(u1_struct_0(X1),X3,X4)!=k1_xboole_0|~v2_compts_1(k4_subset_1(u1_struct_0(X1),X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_282, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k5_setfam_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_192, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk9_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_186, plain, (v1_finset_1(esk10_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_208, negated_conjecture, (v7_ordinal1(esk2_0)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_151, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk10_1(X1)))).#trainneg
cnf(i_0_146, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_tarski(X1)))).#trainneg
cnf(i_0_213, negated_conjecture, (v3_ordinal1(esk2_0)|~v3_ordinal1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_135, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,esk9_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_280, plain, (v1_xboole_0(k5_setfam_1(X1,X2))|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_281, plain, (v3_membered(k5_setfam_1(X1,X2))|~v3_membered(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_150, plain, (v2_compts_1(esk10_1(u1_struct_0(X1)),X1)|~v1_finset_1(esk10_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_355, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(esk4_2(X1,X2),X2)|k5_setfam_1(u1_struct_0(X1),X3)!=k1_xboole_0|~v2_compts_1(k5_setfam_1(u1_struct_0(X1),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_235, plain, (m1_subset_1(esk7_2(X1,X2),X1)|r1_tarski(X1,X2))).#trainneg
cnf(i_0_362, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|esk6_2(X1,X2)=k3_tarski(X2)|k4_subset_1(u1_struct_0(X1),X3,X4)!=k1_xboole_0|~v2_compts_1(k4_subset_1(u1_struct_0(X1),X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_370, plain, (v2_compts_1(esk6_2(X1,X2),X1)|r2_hidden(esk3_2(X1,X2),X2)|k4_subset_1(u1_struct_0(X1),X3,X4)!=k1_xboole_0|~v2_compts_1(k4_subset_1(u1_struct_0(X1),X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_339, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(esk3_2(X1,X2),X2)|k5_setfam_1(u1_struct_0(X1),X3)!=k1_xboole_0|~v2_compts_1(k5_setfam_1(u1_struct_0(X1),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_461, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,esk10_1(X2)))).#trainneg
cnf(i_0_478, plain, (v1_finset_1(X1)|~v5_finset_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_106, plain, (v1_xxreal_0(k1_xboole_0))).#trainneg
cnf(i_0_107, plain, (v1_xcmplx_0(k1_xboole_0))).#trainneg
cnf(i_0_105, plain, (v1_xreal_0(k1_xboole_0))).#trainneg
cnf(i_0_479, plain, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_480, plain, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_458, plain, (v2_compts_1(k3_tarski(X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(k3_tarski(X2))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_456, plain, (v2_compts_1(k2_xboole_0(X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(k2_xboole_0(X2,X3))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_499, plain, (v2_compts_1(u1_struct_0(X1),X1)|~v1_finset_1(u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_462, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,esk9_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_541, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_567, negated_conjecture, (v1_xboole_0(esk2_0)|v2_compts_1(esk13_1(esk1_0),esk1_0)|~m1_subset_1(esk13_1(esk1_0),esk2_0))).#trainneg
cnf(i_0_481, plain, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_477, negated_conjecture, (v1_xxreal_0(esk2_0)|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_568, negated_conjecture, (v1_xboole_0(esk2_0)|v2_compts_1(esk12_1(esk1_0),esk1_0)|~m1_subset_1(esk12_1(esk1_0),esk2_0))).#trainneg
cnf(i_0_482, plain, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_575, negated_conjecture, (v1_xboole_0(esk2_0)|v2_compts_1(u1_struct_0(esk1_0),esk1_0)|~m1_subset_1(u1_struct_0(esk1_0),esk2_0))).#trainneg
cnf(i_0_168, plain, (v1_xboole_0(k3_tarski(X1))|~v1_xboole_0(k3_tarski(k2_xboole_0(X2,X1))))).#trainneg
cnf(i_0_170, plain, (v1_xboole_0(k3_tarski(X1))|~v1_xboole_0(k3_tarski(k2_xboole_0(X1,X2))))).#trainneg
cnf(i_0_453, negated_conjecture, (m1_subset_1(esk7_2(esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_tarski(esk2_0,X1))).#trainneg
cnf(i_0_471, plain, (k3_tarski(k2_xboole_0(k1_tarski(X2),X1))=k2_xboole_0(k3_tarski(X1),X2))).#trainneg
cnf(i_0_689, plain, (m1_subset_1(k3_tarski(k2_xboole_0(X1,X2)),k1_zfmisc_1(X3))|~m1_subset_1(k3_tarski(X2),k1_zfmisc_1(X3))|~m1_subset_1(k3_tarski(X1),k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_538, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_tarski(X3))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_187, plain, (v1_finset_1(esk9_1(k1_zfmisc_1(X1)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_571, negated_conjecture, (v1_xboole_0(esk2_0)|v2_compts_1(k1_struct_0(esk1_0),esk1_0)|~m1_subset_1(k1_struct_0(esk1_0),esk2_0)|~l1_struct_0(esk1_0))).#trainneg
cnf(i_0_9, plain, (m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|X3!=k2_xboole_0(esk3_2(X1,X2),k3_tarski(esk4_2(X1,X2)))|X4!=k1_xboole_0|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_569, negated_conjecture, (v1_xboole_0(esk2_0)|v2_compts_1(esk10_1(u1_struct_0(esk1_0)),esk1_0)|~m1_subset_1(esk10_1(u1_struct_0(esk1_0)),esk2_0))).#trainneg
cnf(i_0_570, negated_conjecture, (v1_xboole_0(esk2_0)|v2_compts_1(esk9_1(k1_zfmisc_1(u1_struct_0(esk1_0))),esk1_0)|~m1_subset_1(esk9_1(k1_zfmisc_1(u1_struct_0(esk1_0))),esk2_0))).#trainneg
cnf(i_0_454, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|v1_xboole_0(esk2_0)|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_498, plain, (v2_compts_1(u1_struct_0(X1),X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_176, plain, (v1_finset_1(esk10_1(X1))|~v5_finset_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_466, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k5_setfam_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_562, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k2_xboole_0(X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_539, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_tarski(X3))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_469, plain, (k3_tarski(k2_xboole_0(k2_xboole_0(X1,k1_tarski(X2)),X3))=k2_xboole_0(k2_xboole_0(k3_tarski(X1),X2),k3_tarski(X3)))).#trainneg
cnf(i_0_634, plain, (v2_compts_1(esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X2),X1)|r1_tarski(k1_zfmisc_1(u1_struct_0(X1)),X2)|~v1_finset_1(esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X2))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_658, plain, (v2_compts_1(esk6_2(X1,X2),X1)|r2_hidden(esk3_2(X1,X2),X2)|k4_subset_1(u1_struct_0(X1),X3,X4)!=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1))).#trainneg
cnf(i_0_681, plain, (v2_compts_1(k2_xboole_0(X1,X2),X3)|~v1_xboole_0(k2_xboole_0(X1,X2))|~l1_pre_topc(X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X3))))).#trainneg
cnf(i_0_507, plain, (v2_compts_1(k5_setfam_1(u1_struct_0(X1),X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_finset_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_177, plain, (v1_finset_1(esk9_1(X1))|~v5_finset_1(X1))).#trainneg
cnf(i_0_533, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(X3))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_181, plain, (v1_xreal_0(esk10_1(X1))|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_182, plain, (v1_xreal_0(esk9_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_220, plain, (~v1_xboole_0(u1_struct_0(X1))|~r2_hidden(X2,k1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_230, plain, (~v1_xboole_0(u1_struct_0(X1))|~r2_hidden(X2,esk12_1(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_610, plain, (v1_finset_1(k5_setfam_1(X1,X2))|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_272, plain, (v1_finset_1(k5_setfam_1(X1,X2))|~v5_finset_1(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_200, plain, (v1_xxreal_0(esk10_1(X1))|~v2_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_201, plain, (v1_xxreal_0(esk9_1(X1))|~v2_membered(X1))).#trainneg
cnf(i_0_950, negated_conjecture, (~v1_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_700, plain, (v2_compts_1(k3_tarski(k2_xboole_0(X1,X2)),X3)|~v2_compts_1(k3_tarski(X1),X3)|~v2_compts_1(k3_tarski(X2),X3)|~m1_subset_1(k3_tarski(X1),k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(k3_tarski(X2),k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3))).#trainneg
cnf(i_0_747, plain, (v2_compts_1(k2_xboole_0(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~v1_finset_1(X2)|~v1_finset_1(X1))).#trainneg
cnf(i_0_739, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_740, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_717, negated_conjecture, (v1_xboole_0(esk2_0)|v2_compts_1(k1_xboole_0,esk1_0)|~m1_subset_1(k1_xboole_0,esk2_0))).#trainneg
cnf(i_0_205, plain, (v7_ordinal1(esk10_1(X1))|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_829, negated_conjecture, (r1_tarski(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_895, negated_conjecture, (v1_xboole_0(esk2_0)|m1_subset_1(esk9_1(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_699, plain, (v1_xxreal_0(X1)|~v3_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_273, plain, (v1_xreal_0(k5_setfam_1(X1,X2))|~v3_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_206, plain, (v7_ordinal1(esk9_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_210, plain, (v3_ordinal1(esk10_1(X1))|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_799, plain, (r1_tarski(X2,X3)|m1_subset_1(esk7_2(X2,X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_870, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_tarski(k2_xboole_0(X3,X4)))|~m1_subset_1(k3_tarski(X4),k1_zfmisc_1(X2))|~m1_subset_1(k3_tarski(X3),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_830, negated_conjecture, (r1_tarski(esk2_0,X1)|v2_compts_1(esk7_2(esk2_0,X1),esk1_0))).#trainneg
cnf(i_0_735, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k1_xboole_0))).#trainneg
cnf(i_0_211, plain, (v3_ordinal1(esk9_1(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_738, plain, (v3_membered(k1_xboole_0)|~v3_membered(X1))).#trainneg
cnf(i_0_960, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))|~r2_hidden(X1,k3_tarski(esk2_0)))).#trainneg
cnf(i_0_715, plain, (v7_ordinal1(k1_xboole_0)|~v6_membered(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_274, plain, (v1_xxreal_0(k5_setfam_1(X1,X2))|~v2_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_275, plain, (v7_ordinal1(k5_setfam_1(X1,X2))|~v6_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_216, plain, (v1_xreal_0(k1_struct_0(X1))|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_741, plain, (v2_compts_1(k3_tarski(k2_xboole_0(X1,X2)),X3)|~m1_subset_1(k3_tarski(X2),k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(k3_tarski(X1),k1_zfmisc_1(u1_struct_0(X3)))|~v1_finset_1(k3_tarski(k2_xboole_0(X1,X2)))|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_921, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_tarski(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_709, plain, (v2_compts_1(k2_xboole_0(k3_tarski(X1),X2),X3)|~m1_subset_1(k2_xboole_0(X1,k1_tarski(X2)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~v1_finset_1(k2_xboole_0(k3_tarski(X1),X2))|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_867, plain, (v1_finset_1(k3_tarski(k2_xboole_0(X1,X2)))|~v1_finset_1(X3)|~m1_subset_1(k3_tarski(X2),k1_zfmisc_1(X3))|~m1_subset_1(k3_tarski(X1),k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_891, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_tarski(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_880, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_tarski(esk10_1(k1_zfmisc_1(X2)))))).#trainneg
cnf(i_0_217, plain, (v1_xxreal_0(k1_struct_0(X1))|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_716, plain, (v3_ordinal1(k1_xboole_0)|~v3_ordinal1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_967, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_276, plain, (v3_ordinal1(k5_setfam_1(X1,X2))|~v3_ordinal1(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_612, plain, (X1=k5_setfam_1(X2,X3)|~v1_xboole_0(X1)|~v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_218, plain, (v7_ordinal1(k1_struct_0(X1))|~v6_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_219, plain, (v3_ordinal1(k1_struct_0(X1))|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_936, plain, (k3_tarski(k2_xboole_0(X3,k2_xboole_0(X1,k1_tarski(X2))))=k2_xboole_0(k2_xboole_0(k3_tarski(X1),X2),k3_tarski(X3)))).#trainneg
cnf(i_0_908, plain, (m1_subset_1(esk7_2(k2_xboole_0(X1,X2),X3),X4)|r1_tarski(k2_xboole_0(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(X4))|~m1_subset_1(X1,k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_871, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_tarski(k2_xboole_0(X3,X4)))|~m1_subset_1(k3_tarski(X4),k1_zfmisc_1(X2))|~m1_subset_1(k3_tarski(X3),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_580, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(esk3_2(X1,X2),X2)|k4_subset_1(u1_struct_0(X1),X3,X4)!=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1))).#trainneg
cnf(i_0_910, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_tarski(esk10_1(k1_zfmisc_1(X2)))))).#trainneg
cnf(i_0_1130, plain, (v2_compts_1(k1_xboole_0,esk16_0))).#trainneg
cnf(i_0_998, negated_conjecture, (esk2_0=k1_xboole_0|v2_compts_1(k1_xboole_0,esk1_0)|~m1_subset_1(k1_xboole_0,esk2_0))).#trainneg
cnf(i_0_222, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,k1_struct_0(X2))|~l1_struct_0(X2))).#trainneg
cnf(i_0_223, plain, (v2_compts_1(k1_struct_0(X1),X1)|~v1_finset_1(k1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_1255, plain, (m1_subset_1(esk3_2(X1,X2),X2)|esk6_2(X1,X2)=k3_tarski(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1057, plain, (v2_compts_1(esk7_2(X1,X2),X3)|r1_tarski(X1,X2)|~v1_finset_1(esk7_2(X1,X2))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3)))))).#trainneg
cnf(i_0_881, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_tarski(esk9_1(k1_zfmisc_1(k1_zfmisc_1(X2))))))).#trainneg
cnf(i_0_911, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_tarski(esk9_1(k1_zfmisc_1(k1_zfmisc_1(X2))))))).#trainneg
cnf(i_0_1015, negated_conjecture, (v1_xboole_0(esk2_0)|v2_compts_1(esk9_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_964, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(k3_tarski(X3)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_225, plain, (v1_finset_1(esk12_1(X1))|~v5_finset_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1011, negated_conjecture, (v1_xboole_0(esk9_1(esk2_0))|v1_xboole_0(esk2_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1012, negated_conjecture, (v3_membered(esk9_1(esk2_0))|v1_xboole_0(esk2_0)|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1097, plain, (v1_xreal_0(k5_setfam_1(X1,X2))|~v6_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1098, plain, (v1_xxreal_0(k5_setfam_1(X1,X2))|~v6_membered(k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_226, plain, (v1_xreal_0(esk12_1(X1))|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_227, plain, (v1_xxreal_0(esk12_1(X1))|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1004, negated_conjecture, (v1_finset_1(esk9_1(esk2_0))|v1_xboole_0(esk2_0)|~v5_finset_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_266, plain, (~v1_xboole_0(u1_struct_0(X1))|~r2_hidden(X2,esk13_1(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_228, plain, (v7_ordinal1(esk12_1(X1))|~v6_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1005, negated_conjecture, (v1_xreal_0(esk9_1(esk2_0))|v1_xboole_0(esk2_0)|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1006, negated_conjecture, (v1_xxreal_0(esk9_1(esk2_0))|v1_xboole_0(esk2_0)|~v2_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_973, plain, (v2_compts_1(k5_setfam_1(u1_struct_0(X1),X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_975, plain, (v2_compts_1(k5_setfam_1(u1_struct_0(X1),X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v5_finset_1(k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_229, plain, (v3_ordinal1(esk12_1(X1))|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_231, plain, (v1_finset_1(esk12_1(X1))|~v1_finset_1(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1112, plain, (v2_compts_1(X1,X2)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_finset_1(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1256, plain, (esk6_2(X2,X1)=k3_tarski(X1)|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_finset_1(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1007, negated_conjecture, (v7_ordinal1(esk9_1(esk2_0))|v1_xboole_0(esk2_0)|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_956, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(esk10_1(k1_zfmisc_1(X1)))))).#trainneg
cnf(i_0_232, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,esk12_1(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_1008, negated_conjecture, (v3_ordinal1(esk9_1(esk2_0))|v1_xboole_0(esk2_0)|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1009, negated_conjecture, (v1_xboole_0(esk2_0)|~v1_xboole_0(u1_struct_0(esk1_0))|~r2_hidden(X1,esk9_1(esk2_0)))).#trainneg
cnf(i_0_574, negated_conjecture, (v1_xboole_0(esk2_0)|v2_compts_1(k5_setfam_1(u1_struct_0(esk1_0),X1),esk1_0)|~m1_subset_1(k5_setfam_1(u1_struct_0(esk1_0),X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_583, plain, (m1_subset_1(X1,X2)|v1_xboole_0(k5_setfam_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k5_setfam_1(X2,X3)))).#trainneg
cnf(i_0_233, plain, (v2_compts_1(esk12_1(X1),X1)|~v1_finset_1(esk12_1(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_239, plain, (v1_xboole_0(esk12_1(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1186, plain, (v2_compts_1(esk7_2(k2_xboole_0(X1,X2),X3),X4)|r1_tarski(k2_xboole_0(X1,X2),X3)|~v1_finset_1(esk7_2(k2_xboole_0(X1,X2),X3))|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4))))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4)))))).#trainneg
cnf(i_0_1087, plain, (v1_xboole_0(X1)|r1_tarski(k2_xboole_0(X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_986, plain, (v2_compts_1(k3_tarski(k2_xboole_0(X1,X2)),X3)|~m1_subset_1(k3_tarski(X2),k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(k3_tarski(X1),k1_zfmisc_1(u1_struct_0(X3)))|~v1_finset_1(k3_tarski(X2))|~v1_finset_1(k3_tarski(X1))|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_888, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_tarski(k3_tarski(X3)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X2)))))).#trainneg
cnf(i_0_1094, negated_conjecture, (r1_tarski(k3_tarski(esk2_0),X1)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_957, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(esk9_1(k1_zfmisc_1(k1_zfmisc_1(X1))))))).#trainneg
cnf(i_0_240, plain, (v1_xboole_0(esk10_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_884, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k3_tarski(esk2_0)))).#trainneg
cnf(i_0_1095, negated_conjecture, (v1_xboole_0(k3_tarski(esk2_0))|~v1_xboole_0(u1_struct_0(esk1_0))|~m1_subset_1(X1,k3_tarski(esk2_0)))).#trainneg
cnf(i_0_901, plain, (r2_hidden(X1,X2)|v1_xboole_0(k5_setfam_1(X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k5_setfam_1(X2,X3)))).#trainneg
cnf(i_0_241, plain, (v1_xboole_0(esk9_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1352, negated_conjecture, (v2_compts_1(esk3_2(X1,esk2_0),esk1_0)|esk6_2(X1,esk2_0)=k3_tarski(esk2_0)|~v1_finset_1(esk3_2(X1,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1409, plain, (r1_tarski(k2_xboole_0(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_1341, negated_conjecture, (v2_compts_1(esk6_2(esk1_0,X1),esk1_0)|r2_hidden(esk3_2(esk1_0,X1),X1)|esk6_2(X2,esk2_0)=k3_tarski(esk2_0)|esk3_2(X2,esk2_0)!=k1_xboole_0|~v2_compts_1(esk3_2(X2,esk2_0),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_finset_1(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1485, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_tarski(k3_tarski(k1_zfmisc_1(k1_zfmisc_1(X2))))))).#trainneg
cnf(i_0_1319, negated_conjecture, (esk9_1(esk2_0)=k1_xboole_0|v1_xboole_0(esk2_0)|~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_246, plain, (v3_membered(esk12_1(X1))|~v3_membered(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1359, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1363, negated_conjecture, (~v5_finset_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1138, plain, (k5_setfam_1(X1,X2)=k5_setfam_1(X3,X4)|~v1_xboole_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3)))|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_299, plain, (v2_compts_1(k5_setfam_1(u1_struct_0(X1),X2),X1)|~v1_xboole_0(k5_setfam_1(u1_struct_0(X1),X2))|~l1_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_247, plain, (v3_membered(esk10_1(X1))|~v3_membered(X1))).#trainneg
cnf(i_0_248, plain, (v3_membered(esk9_1(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_1338, negated_conjecture, (esk6_2(esk1_0,X1)=k3_tarski(X1)|r1_tarski(esk4_2(esk1_0,X1),X1)|esk6_2(X2,esk2_0)=k3_tarski(esk2_0)|esk3_2(X2,esk2_0)!=k1_xboole_0|~v2_compts_1(esk3_2(X2,esk2_0),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_finset_1(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1372, plain, (v1_xboole_0(X1)|r2_hidden(esk7_2(k4_subset_1(X1,X2,X3),X4),X1)|r1_tarski(k4_subset_1(X1,X2,X3),X4)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1472, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_tarski(k3_tarski(esk10_1(k1_zfmisc_1(k1_zfmisc_1(X2)))))))).#trainneg
cnf(i_0_1473, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_tarski(k3_tarski(esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X2))))))))).#trainneg
cnf(i_0_1313, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(k3_tarski(k3_tarski(X3))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1))))))).#trainneg
cnf(i_0_251, plain, (v3_membered(k1_struct_0(X1))|~v3_membered(u1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_1432, negated_conjecture, (v1_xboole_0(esk2_0)|~m1_subset_1(k5_setfam_1(u1_struct_0(esk1_0),esk2_0),esk2_0))).#trainneg
cnf(i_0_1496, plain, (X1=esk10_1(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_253, plain, (v1_finset_1(k3_tarski(k2_xboole_0(X1,X2)))|~v1_finset_1(k3_tarski(X2))|~v1_finset_1(k3_tarski(X1)))).#trainneg
cnf(i_0_261, plain, (v1_finset_1(esk13_1(X1))|~v5_finset_1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1292, plain, (v1_finset_1(esk3_2(X1,X2))|esk6_2(X1,X2)=k3_tarski(X2)|~v5_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1345, negated_conjecture, (r2_hidden(esk3_2(esk1_0,X1),X1)|m1_subset_1(esk6_2(esk1_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|esk6_2(X2,esk2_0)=k3_tarski(esk2_0)|esk3_2(X2,esk2_0)!=k1_xboole_0|~v2_compts_1(esk3_2(X2,esk2_0),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_finset_1(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1343, negated_conjecture, (r1_tarski(esk4_2(esk1_0,X1),X1)|v2_compts_1(esk6_2(esk1_0,X1),esk1_0)|esk6_2(X2,esk2_0)=k3_tarski(esk2_0)|esk3_2(X2,esk2_0)!=k1_xboole_0|~v2_compts_1(esk3_2(X2,esk2_0),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_finset_1(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_782, plain, (v2_compts_1(esk7_2(k4_subset_1(k1_zfmisc_1(u1_struct_0(X1)),X2,X3),X4),X1)|r1_tarski(k4_subset_1(k1_zfmisc_1(u1_struct_0(X1)),X2,X3),X4)|~v1_finset_1(esk7_2(k4_subset_1(k1_zfmisc_1(u1_struct_0(X1)),X2,X3),X4))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_554, plain, (r1_tarski(esk4_2(X1,X2),X2)|v2_compts_1(esk6_2(X1,X2),X1)|k4_subset_1(u1_struct_0(X1),X3,X4)!=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1))).#trainneg
cnf(i_0_582, plain, (m1_subset_1(esk7_2(k5_setfam_1(X1,X2),X3),X1)|r1_tarski(k5_setfam_1(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_918, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_tarski(k3_tarski(X3)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X2)))))).#trainneg
cnf(i_0_1724, negated_conjecture, (v1_xboole_0(esk2_0)|~m1_subset_1(k3_tarski(esk2_0),esk2_0))).#trainneg
cnf(i_0_262, plain, (v1_xreal_0(esk13_1(X1))|~v3_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_263, plain, (v1_xxreal_0(esk13_1(X1))|~v2_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_264, plain, (v7_ordinal1(esk13_1(X1))|~v6_membered(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1794, plain, (v1_xboole_0(X1)|r1_tarski(k5_setfam_1(X1,X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1734, plain, (v2_compts_1(k3_tarski(k2_xboole_0(X1,X2)),X3)|~m1_subset_1(k2_xboole_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~v1_finset_1(k3_tarski(X2))|~v1_finset_1(k3_tarski(X1)))).#trainneg
cnf(i_0_1647, plain, (v2_compts_1(k5_setfam_1(u1_struct_0(X1),X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_1493, plain, (v1_finset_1(esk10_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1416, negated_conjecture, (v1_xreal_0(esk9_1(esk2_0))|v1_xboole_0(esk2_0)|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_265, plain, (v3_ordinal1(esk13_1(X1))|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_267, plain, (v1_finset_1(esk13_1(X1))|~v1_finset_1(u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_268, plain, (v1_xboole_0(esk13_1(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1561, plain, (r1_tarski(k3_tarski(k2_xboole_0(X1,X2)),X3)|~m1_subset_1(k3_tarski(X2),k1_zfmisc_1(X3))|~m1_subset_1(k3_tarski(X1),k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_1535, plain, (m1_subset_1(esk7_2(X1,X2),X3)|r1_tarski(X1,X2)|~m1_subset_1(X5,k1_zfmisc_1(X3))|~m1_subset_1(X4,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_xboole_0(X4,X5))))).#trainneg
cnf(i_0_1530, negated_conjecture, (v2_compts_1(esk7_2(X1,X2),esk1_0)|r1_tarski(X1,X2)|~m1_subset_1(esk7_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_1116, plain, (v2_compts_1(k2_xboole_0(k3_tarski(X1),X2),X3)|~v1_finset_1(k2_xboole_0(k3_tarski(X1),X2))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3)))))).#trainneg
cnf(i_0_535, plain, (v1_finset_1(k3_tarski(X1))|~v1_finset_1(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_1034, plain, (v1_xreal_0(esk9_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_1417, negated_conjecture, (v1_xxreal_0(esk9_1(esk2_0))|v1_xboole_0(esk2_0)|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_269, plain, (v3_membered(esk13_1(X1))|~v3_membered(u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_270, plain, (m1_subset_1(X1,u1_struct_0(X2))|~r2_hidden(X1,esk13_1(X2))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_271, plain, (v2_compts_1(esk13_1(X1),X1)|~v1_finset_1(esk13_1(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1661, negated_conjecture, (v2_compts_1(esk7_2(k4_subset_1(esk2_0,X1,X2),X3),esk1_0)|v1_xboole_0(esk2_0)|r1_tarski(k4_subset_1(esk2_0,X1,X2),X3)|~m1_subset_1(esk7_2(k4_subset_1(esk2_0,X1,X2),X3),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_1529, negated_conjecture, (m1_subset_1(esk7_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_914, negated_conjecture, (r2_hidden(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k3_tarski(esk2_0)))).#trainneg
cnf(i_0_1795, plain, (r1_tarski(k3_tarski(X2),X3)|m1_subset_1(esk7_2(k3_tarski(X2),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1410, plain, (r1_tarski(k3_tarski(k2_xboole_0(X1,X2)),X3)|r2_hidden(esk7_2(k3_tarski(k2_xboole_0(X1,X2)),X3),X4)|~m1_subset_1(k3_tarski(X2),k1_zfmisc_1(X4))|~m1_subset_1(k3_tarski(X1),k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_1970, plain, (esk6_2(X1,X3)=k3_tarski(X3)|v1_xboole_0(X2)|r2_hidden(esk3_2(X1,X3),X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~v1_finset_1(X3)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1815, negated_conjecture, (esk2_0=k1_xboole_0|~m1_subset_1(k3_tarski(esk2_0),esk2_0))).#trainneg
cnf(i_0_293, plain, (v2_compts_1(esk13_1(X1),X1)|~v1_xboole_0(esk13_1(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1820, negated_conjecture, (r2_hidden(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_294, plain, (v2_compts_1(esk12_1(X1),X1)|~v1_xboole_0(esk12_1(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_295, plain, (v2_compts_1(esk10_1(u1_struct_0(X1)),X1)|~v1_xboole_0(esk10_1(u1_struct_0(X1)))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1865, plain, (r1_tarski(X1,X2)|m1_subset_1(esk7_2(X1,X2),X3)|~m1_subset_1(X1,k1_zfmisc_1(X4))|~m1_subset_1(X4,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_1812, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_tarski(k3_tarski(k1_zfmisc_1(k1_zfmisc_1(X2))))))).#trainneg
cnf(i_0_1798, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_tarski(k3_tarski(esk10_1(k1_zfmisc_1(k1_zfmisc_1(X2)))))))).#trainneg
cnf(i_0_296, plain, (v2_compts_1(esk9_1(k1_zfmisc_1(u1_struct_0(X1))),X1)|~v1_xboole_0(esk9_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_1316, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(k3_tarski(k1_zfmisc_1(k1_zfmisc_1(X1))))))).#trainneg
cnf(i_0_300, plain, (v1_finset_1(k4_subset_1(X1,X2,X3))|~v5_finset_1(k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2004, negated_conjecture, (r2_hidden(esk7_2(k3_tarski(esk2_0),X1),u1_struct_0(esk1_0))|r1_tarski(k3_tarski(esk2_0),X1))).#trainneg
cnf(i_0_1825, plain, (v1_xboole_0(X1)|r1_tarski(k3_tarski(X2),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1889, negated_conjecture, (r1_tarski(X1,X2)|v2_compts_1(esk7_2(X1,X2),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_1799, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_tarski(k3_tarski(esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X2))))))))).#trainneg
cnf(i_0_1906, plain, (v1_finset_1(k3_tarski(k1_zfmisc_1(X1)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_1718, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(k3_tarski(k3_tarski(k3_tarski(X3)))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1)))))))).#trainneg
cnf(i_0_301, plain, (v1_xreal_0(k4_subset_1(X1,X2,X3))|~v3_membered(k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_1494, plain, (esk10_1(X1)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1303, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(k3_tarski(esk10_1(k1_zfmisc_1(k1_zfmisc_1(X1)))))))).#trainneg
cnf(i_0_1304, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(k3_tarski(esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1))))))))).#trainneg
cnf(i_0_1599, plain, (m1_subset_1(k5_setfam_1(X3,X4),k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1)|~v1_xboole_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3))))).#trainneg
cnf(i_0_302, plain, (v1_xxreal_0(k4_subset_1(X1,X2,X3))|~v2_membered(k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_303, plain, (v7_ordinal1(k4_subset_1(X1,X2,X3))|~v6_membered(k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2253, negated_conjecture, (r1_tarski(k3_tarski(esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2159, negated_conjecture, (v2_compts_1(X1,esk1_0)|~v1_finset_1(X1)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_2168, negated_conjecture, (v1_finset_1(X1)|~v1_finset_1(k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_304, plain, (v3_ordinal1(k4_subset_1(X1,X2,X3))|~v3_ordinal1(k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2310, plain, (~v1_subset_1(k1_xboole_0,X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_305, plain, (v2_compts_1(k4_subset_1(u1_struct_0(X1),X2,X3),X1)|~v1_xboole_0(k4_subset_1(u1_struct_0(X1),X2,X3))|~l1_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_309, plain, (v1_finset_1(k4_subset_1(X1,X2,X3))|~v1_finset_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2207, plain, (r1_tarski(k4_subset_1(X1,X2,X3),X4)|m1_subset_1(esk7_2(k4_subset_1(X1,X2,X3),X4),X5)|~m1_subset_1(X1,k1_zfmisc_1(X5))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2056, plain, (r1_tarski(X1,X2)|m1_subset_1(esk7_2(X1,X2),X3)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(k1_zfmisc_1(X3))))).#trainneg
cnf(i_0_2125, negated_conjecture, (v2_compts_1(esk3_2(X1,X2),esk1_0)|esk6_2(X1,X2)=k3_tarski(X2)|v1_xboole_0(esk2_0)|~m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_2250, negated_conjecture, (m1_subset_1(esk7_2(k3_tarski(esk2_0),X1),u1_struct_0(esk1_0))|r1_tarski(k3_tarski(esk2_0),X1))).#trainneg
cnf(i_0_2208, plain, (r1_tarski(k2_xboole_0(X1,X2),X3)|m1_subset_1(esk7_2(k2_xboole_0(X1,X2),X3),X4)|~m1_subset_1(X5,k1_zfmisc_1(X4))|~m1_subset_1(X2,k1_zfmisc_1(X5))|~m1_subset_1(X1,k1_zfmisc_1(X5)))).#trainneg
cnf(i_0_1250, plain, (esk6_2(X2,k1_xboole_0)=k1_xboole_0|~v1_xboole_0(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_994, plain, (r2_hidden(esk7_2(k1_xboole_0,X1),X2)|r1_tarski(k1_xboole_0,X1))).#trainneg
cnf(i_0_2457, negated_conjecture, (m1_subset_1(k3_tarski(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2158, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_2682, negated_conjecture, (r1_tarski(k1_xboole_0,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_310, plain, (v1_xboole_0(k4_subset_1(X1,X2,X3))|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2683, plain, (r1_tarski(k1_xboole_0,X1))).#trainneg
cnf(i_0_2187, negated_conjecture, (r2_hidden(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_311, plain, (v3_membered(k4_subset_1(X1,X2,X3))|~v3_membered(X1)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_315, plain, (v3_pre_topc(k3_tarski(k2_xboole_0(X1,X2)),X3)|~v3_pre_topc(k3_tarski(X2),X3)|~v3_pre_topc(k3_tarski(X1),X3)|~m1_subset_1(k3_tarski(X2),k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(k3_tarski(X1),k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_2480, plain, (v2_compts_1(k4_subset_1(u1_struct_0(X1),X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_finset_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_2484, plain, (v1_finset_1(k2_xboole_0(X3,X2))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_2444, plain, (v2_compts_1(esk3_2(X1,X2),X3)|esk6_2(X1,X2)=k3_tarski(X2)|~v1_finset_1(esk3_2(X1,X2))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_1721, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(k3_tarski(k3_tarski(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1))))))))).#trainneg
cnf(i_0_2474, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_320, plain, (v4_pre_topc(k3_tarski(k2_xboole_0(X1,X2)),X3)|~v4_pre_topc(k3_tarski(X2),X3)|~v4_pre_topc(k3_tarski(X1),X3)|~m1_subset_1(k3_tarski(X2),k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(k3_tarski(X1),k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_2629, plain, (esk6_2(X1,k1_xboole_0)=k1_xboole_0|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_2612, negated_conjecture, (r1_tarski(k2_xboole_0(X1,X2),X3)|m1_subset_1(esk7_2(k2_xboole_0(X1,X2),X3),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_2568, plain, (v2_compts_1(esk7_2(X1,X2),X3)|r1_tarski(X1,X2)|~v1_finset_1(esk7_2(X1,X2))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))))).#trainneg
cnf(i_0_2531, plain, (v1_xboole_0(X1)|r1_tarski(k4_subset_1(X2,X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_2582, plain, (v1_xboole_0(X1)|r1_tarski(X2,X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_835, plain, (k3_tarski(k2_xboole_0(X1,k2_xboole_0(k1_tarski(X2),X3)))=k2_xboole_0(k3_tarski(X1),k2_xboole_0(k3_tarski(X3),X2)))).#trainneg
cnf(i_0_2162, negated_conjecture, (v1_finset_1(X1)|~v5_finset_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_2170, negated_conjecture, (v3_membered(X1)|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_1708, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(k3_tarski(k3_tarski(esk10_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1)))))))))).#trainneg
cnf(i_0_2712, negated_conjecture, (~v1_xboole_0(k3_tarski(esk2_0)))).#trainneg
cnf(i_0_2713, negated_conjecture, (~v1_finset_1(k3_tarski(esk2_0)))).#trainneg
cnf(i_0_2719, plain, (k4_subset_1(X1,X2,X3)=k1_xboole_0|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2726, plain, (v1_xboole_0(k2_xboole_0(X2,X3))|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2818, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk3_2(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~v1_finset_1(X2)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v2_pre_topc(X3)|~v2_pre_topc(X1)|~v5_finset_1(X2))).#trainneg
cnf(i_0_2910, plain, (m1_subset_1(k4_subset_1(X1,X2,X3),k1_zfmisc_1(X4))|v1_xboole_0(X4)|~m1_subset_1(X1,k1_zfmisc_1(X4))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2817, plain, (v1_finset_1(k2_xboole_0(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_2284, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(k3_tarski(k3_tarski(k3_tarski(k3_tarski(X3))))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1))))))))).#trainneg
cnf(i_0_340, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk5_2(X1,X2),X1)|esk13_1(X1)!=k1_xboole_0|~v2_compts_1(esk13_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_2692, negated_conjecture, (v3_membered(k3_tarski(esk2_0))|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_341, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk5_2(X1,X2),X1)|esk12_1(X1)!=k1_xboole_0|~v2_compts_1(esk12_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_342, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk5_2(X1,X2),X1)|esk10_1(u1_struct_0(X1))!=k1_xboole_0|~v2_compts_1(esk10_1(u1_struct_0(X1)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_2780, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|esk6_2(X1,X2)=k3_tarski(X2)|k4_subset_1(u1_struct_0(X1),X3,X4)!=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_finset_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_2718, plain, (v1_finset_1(k4_subset_1(X1,X2,X3))|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2787, plain, (v2_compts_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_2846, negated_conjecture, (v1_xboole_0(X1)|m1_subset_1(esk3_2(X2,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|esk6_2(X2,X1)=k3_tarski(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_finset_1(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_1891, plain, (v1_finset_1(k3_tarski(esk10_1(k1_zfmisc_1(X1))))|~v1_finset_1(X1))).#trainneg
cnf(i_0_2688, negated_conjecture, (v3_ordinal1(k3_tarski(esk2_0))|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1709, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(k3_tarski(k3_tarski(esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1))))))))))).#trainneg
cnf(i_0_343, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk5_2(X1,X2),X1)|esk9_1(k1_zfmisc_1(u1_struct_0(X1)))!=k1_xboole_0|~v2_compts_1(esk9_1(k1_zfmisc_1(u1_struct_0(X1))),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_2685, negated_conjecture, (v1_xreal_0(k3_tarski(esk2_0))|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2686, negated_conjecture, (v1_xxreal_0(k3_tarski(esk2_0))|~v2_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_345, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk5_2(X1,X2),X1)|k1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(k1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_346, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk5_2(X1,X2),X1)|k4_subset_1(u1_struct_0(X1),X3,X4)!=k1_xboole_0|~v2_compts_1(k4_subset_1(u1_struct_0(X1),X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_2916, plain, (v1_xboole_0(X1)|r1_tarski(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3188, negated_conjecture, (v2_compts_1(esk3_2(X1,X2),esk1_0)|esk6_2(X1,X2)=k3_tarski(X2)|~v1_finset_1(esk3_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3078, plain, (v1_xboole_0(X1)|m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3124, plain, (v1_finset_1(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3081, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_1892, plain, (v1_finset_1(k3_tarski(esk9_1(k1_zfmisc_1(k1_zfmisc_1(X1)))))|~v1_finset_1(X1))).#trainneg
cnf(i_0_347, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk5_2(X1,X2),X1)|k5_setfam_1(u1_struct_0(X1),X3)!=k1_xboole_0|~v2_compts_1(k5_setfam_1(u1_struct_0(X1),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_2687, negated_conjecture, (v7_ordinal1(k3_tarski(esk2_0))|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3203, negated_conjecture, (v1_xxreal_0(k3_tarski(esk2_0))|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_348, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(esk4_2(X1,X2),X2)|esk13_1(X1)!=k1_xboole_0|~v2_compts_1(esk13_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_349, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(esk4_2(X1,X2),X2)|esk12_1(X1)!=k1_xboole_0|~v2_compts_1(esk12_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3238, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k3_tarski(esk2_0))))).#trainneg
cnf(i_0_2976, plain, (k2_xboole_0(X2,X3)=k1_xboole_0|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3230, plain, (v1_xboole_0(X1)|r1_tarski(X2,X1)|esk6_2(X3,X4)=k3_tarski(X4)|~m1_subset_1(X2,k1_zfmisc_1(esk3_2(X3,X4)))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v1_finset_1(X4)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_3074, plain, (v1_xboole_0(X1)|m1_subset_1(k2_xboole_0(X3,X4),k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_1035, plain, (v1_xxreal_0(esk9_1(X1))|~v6_membered(X1))).#trainneg
cnf(i_0_2752, negated_conjecture, (r2_hidden(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))),k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_3103, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(k3_tarski(k3_tarski(k3_tarski(k3_tarski(k3_tarski(X3)))))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1)))))))))).#trainneg
cnf(i_0_350, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(esk4_2(X1,X2),X2)|esk10_1(u1_struct_0(X1))!=k1_xboole_0|~v2_compts_1(esk10_1(u1_struct_0(X1)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_2848, negated_conjecture, (m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))),k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_2756, negated_conjecture, (r2_hidden(k3_tarski(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_2334, plain, (m1_subset_1(k5_setfam_1(X1,X2),k1_zfmisc_1(X3))|~v1_xboole_0(X3)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_351, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(esk4_2(X1,X2),X2)|esk9_1(k1_zfmisc_1(u1_struct_0(X1)))!=k1_xboole_0|~v2_compts_1(esk9_1(k1_zfmisc_1(u1_struct_0(X1))),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_353, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(esk4_2(X1,X2),X2)|k1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(k1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_3272, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(k3_tarski(esk2_0))))).#trainneg
cnf(i_0_3245, negated_conjecture, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk3_2(X1,X2),esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v5_finset_1(X2))).#trainneg
cnf(i_0_3325, negated_conjecture, (~m1_subset_1(k3_tarski(esk2_0),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_3321, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk3_2(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~v1_finset_1(X2)|~l1_pre_topc(X3)|~l1_pre_topc(X1)|~v2_pre_topc(X3)|~v2_pre_topc(X1)|~m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_3277, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X2,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_3322, negated_conjecture, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk3_2(X1,X2),esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_3330, negated_conjecture, (~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_3329, negated_conjecture, (v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_356, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|esk6_2(X1,X2)=k3_tarski(X2)|esk13_1(X1)!=k1_xboole_0|~v2_compts_1(esk13_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_357, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|esk6_2(X1,X2)=k3_tarski(X2)|esk12_1(X1)!=k1_xboole_0|~v2_compts_1(esk12_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_358, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|esk6_2(X1,X2)=k3_tarski(X2)|esk10_1(u1_struct_0(X1))!=k1_xboole_0|~v2_compts_1(esk10_1(u1_struct_0(X1)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3266, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|esk6_2(X2,X3)=k3_tarski(X3)|~m1_subset_1(X1,k1_zfmisc_1(esk3_2(X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(esk2_0))|~v1_finset_1(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_3419, negated_conjecture, (m1_subset_1(k2_xboole_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_3367, plain, (k2_xboole_0(X1,k4_subset_1(X2,X3,X4))=k1_xboole_0|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3275, plain, (v1_xboole_0(X1)|m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_3345, plain, (k2_xboole_0(X1,esk12_1(X2))=k1_xboole_0|~v1_xboole_0(u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_3373, plain, (X1=k1_xboole_0|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3374, plain, (k2_xboole_0(X1,k5_setfam_1(X2,X3))=k1_xboole_0|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_3271, plain, (v1_xboole_0(X1)|m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_xboole_0(X3,X4)))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3544, negated_conjecture, (~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_359, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|esk6_2(X1,X2)=k3_tarski(X2)|esk9_1(k1_zfmisc_1(u1_struct_0(X1)))!=k1_xboole_0|~v2_compts_1(esk9_1(k1_zfmisc_1(u1_struct_0(X1))),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_363, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|esk6_2(X1,X2)=k3_tarski(X2)|k5_setfam_1(u1_struct_0(X1),X3)!=k1_xboole_0|~v2_compts_1(k5_setfam_1(u1_struct_0(X1),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_3764, plain, (k3_tarski(k2_xboole_0(X1,X2))=k1_xboole_0|~v1_xboole_0(X3)|~m1_subset_1(k3_tarski(X2),k1_zfmisc_1(X3))|~m1_subset_1(k3_tarski(X1),k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_3778, plain, (v1_xboole_0(k3_tarski(X1))|~v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3766, plain, (k5_setfam_1(X1,X2)=k1_xboole_0|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_1515, plain, (v1_finset_1(esk9_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3522, negated_conjecture, (m1_subset_1(esk10_1(k3_tarski(esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3453, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(k3_tarski(k3_tarski(k3_tarski(k3_tarski(k3_tarski(k3_tarski(X3))))))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1))))))))))).#trainneg
cnf(i_0_364, plain, (v2_compts_1(esk6_2(X1,X2),X1)|r2_hidden(esk3_2(X1,X2),X2)|esk13_1(X1)!=k1_xboole_0|~v2_compts_1(esk13_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3523, negated_conjecture, (m1_subset_1(esk9_1(k1_zfmisc_1(k3_tarski(esk2_0))),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3508, plain, (v3_membered(k5_setfam_1(X1,X2))|~v3_membered(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_3693, plain, (v1_xboole_0(X1)|m1_subset_1(k5_setfam_1(X2,X3),k1_zfmisc_1(X1))|~v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_365, plain, (v2_compts_1(esk6_2(X1,X2),X1)|r2_hidden(esk3_2(X1,X2),X2)|esk12_1(X1)!=k1_xboole_0|~v2_compts_1(esk12_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_366, plain, (v2_compts_1(esk6_2(X1,X2),X1)|r2_hidden(esk3_2(X1,X2),X2)|esk10_1(u1_struct_0(X1))!=k1_xboole_0|~v2_compts_1(esk10_1(u1_struct_0(X1)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3634, plain, (v1_xboole_0(k3_tarski(X1))|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3866, plain, (v1_xboole_0(k3_tarski(X1))|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3903, plain, (k3_tarski(X2)=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3865, plain, (k3_tarski(k2_xboole_0(X1,X2))=k1_xboole_0|~v1_xboole_0(X3)|~m1_subset_1(k3_tarski(X1),k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3))))).#trainneg
cnf(i_0_3631, plain, (k2_xboole_0(k4_subset_1(X1,X2,X3),X4)=k1_xboole_0|~v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_367, plain, (v2_compts_1(esk6_2(X1,X2),X1)|r2_hidden(esk3_2(X1,X2),X2)|esk9_1(k1_zfmisc_1(u1_struct_0(X1)))!=k1_xboole_0|~v2_compts_1(esk9_1(k1_zfmisc_1(u1_struct_0(X1))),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_369, plain, (v2_compts_1(esk6_2(X1,X2),X1)|r2_hidden(esk3_2(X1,X2),X2)|k1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(k1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_3811, plain, (k2_xboole_0(X1,k3_tarski(X3))=k1_xboole_0|~v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_4097, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k5_setfam_1(X3,X4))|~v1_xboole_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3))))).#trainneg
cnf(i_0_3942, negated_conjecture, (v3_membered(esk10_1(k3_tarski(esk2_0)))|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_372, plain, (v2_compts_1(esk6_2(X1,X2),X1)|v2_compts_1(esk5_2(X1,X2),X1)|esk13_1(X1)!=k1_xboole_0|~v2_compts_1(esk13_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4095, plain, (v3_membered(k5_setfam_1(X1,X2))|~v3_membered(X3)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_373, plain, (v2_compts_1(esk6_2(X1,X2),X1)|v2_compts_1(esk5_2(X1,X2),X1)|esk12_1(X1)!=k1_xboole_0|~v2_compts_1(esk12_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_374, plain, (v2_compts_1(esk6_2(X1,X2),X1)|v2_compts_1(esk5_2(X1,X2),X1)|esk10_1(u1_struct_0(X1))!=k1_xboole_0|~v2_compts_1(esk10_1(u1_struct_0(X1)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3397, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v1_xboole_0(X3)|r1_tarski(esk3_2(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3521, plain, (r2_hidden(X1,X2)|esk6_2(X3,X2)=k3_tarski(X2)|~r2_hidden(X1,esk4_2(X3,X2))|k1_struct_0(X3)!=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~v1_finset_1(X2)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_3366, plain, (k2_xboole_0(X1,k1_struct_0(X2))=k1_xboole_0|~v1_xboole_0(u1_struct_0(X2))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~l1_struct_0(X2))).#trainneg
cnf(i_0_3232, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|esk6_2(X2,X3)=k3_tarski(X3)|~m1_subset_1(X1,k1_zfmisc_1(esk3_2(X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(esk2_0))|~v1_finset_1(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_4195, plain, (v1_xboole_0(k3_tarski(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1442, plain, (v2_compts_1(esk12_1(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_finset_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_375, plain, (v2_compts_1(esk6_2(X1,X2),X1)|v2_compts_1(esk5_2(X1,X2),X1)|esk9_1(k1_zfmisc_1(u1_struct_0(X1)))!=k1_xboole_0|~v2_compts_1(esk9_1(k1_zfmisc_1(u1_struct_0(X1))),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_377, plain, (v2_compts_1(esk6_2(X1,X2),X1)|v2_compts_1(esk5_2(X1,X2),X1)|k1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(k1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_378, plain, (v2_compts_1(esk6_2(X1,X2),X1)|v2_compts_1(esk5_2(X1,X2),X1)|k4_subset_1(u1_struct_0(X1),X3,X4)!=k1_xboole_0|~v2_compts_1(k4_subset_1(u1_struct_0(X1),X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_4316, plain, (v2_compts_1(esk6_2(X1,X2),X1)|r2_hidden(esk3_2(X1,X2),X2)|k5_setfam_1(u1_struct_0(X1),X3)!=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_finset_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_4374, plain, (m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(X3))|esk6_2(X1,X2)=k3_tarski(X2)|v1_xboole_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4382, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(esk3_2(X3,esk4_2(X1,X2)),X2)|esk6_2(X3,esk4_2(X1,X2))=k3_tarski(esk4_2(X1,X2))|k1_struct_0(X1)!=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~v1_finset_1(esk4_2(X1,X2))|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_1917, plain, (v2_compts_1(esk13_1(X1),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_finset_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_3986, plain, (~v1_xboole_0(X1)|~r2_hidden(X2,k3_tarski(k3_tarski(k3_tarski(k3_tarski(k3_tarski(k3_tarski(k3_tarski(k3_tarski(X3)))))))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X1)))))))))))).#trainneg
cnf(i_0_379, plain, (v2_compts_1(esk6_2(X1,X2),X1)|v2_compts_1(esk5_2(X1,X2),X1)|k5_setfam_1(u1_struct_0(X1),X3)!=k1_xboole_0|~v2_compts_1(k5_setfam_1(u1_struct_0(X1),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_4219, plain, (k3_tarski(k1_zfmisc_1(X1))=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_3930, negated_conjecture, (v1_xreal_0(esk10_1(k3_tarski(esk2_0)))|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3931, negated_conjecture, (v1_xxreal_0(esk10_1(k3_tarski(esk2_0)))|~v2_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_380, plain, (r1_tarski(esk4_2(X1,X2),X2)|v2_compts_1(esk6_2(X1,X2),X1)|esk13_1(X1)!=k1_xboole_0|~v2_compts_1(esk13_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_381, plain, (r1_tarski(esk4_2(X1,X2),X2)|v2_compts_1(esk6_2(X1,X2),X1)|esk12_1(X1)!=k1_xboole_0|~v2_compts_1(esk12_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4151, plain, (v2_compts_1(esk5_2(X1,X2),X3)|esk6_2(X1,X2)=k3_tarski(X2)|~m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~v1_finset_1(esk5_2(X1,X2))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4227, plain, (k2_xboole_0(k3_tarski(X2),X1)=k1_xboole_0|~v1_xboole_0(X3)|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3))))).#trainneg
cnf(i_0_4187, plain, (v1_xboole_0(k3_tarski(k2_xboole_0(X1,X2)))|~v1_xboole_0(X3)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_4457, negated_conjecture, (~v1_xboole_0(esk2_0))).#trainneg
cnf(i_0_4658, plain, (k3_tarski(k1_zfmisc_1(k1_struct_0(X1)))=k1_xboole_0|~l1_struct_0(X1))).#trainneg
cnf(i_0_4453, plain, (k3_tarski(X1)=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4458, plain, (v1_xboole_0(k3_tarski(X1))|~v1_xboole_0(k2_xboole_0(X2,X1)))).#trainneg
cnf(i_0_4493, plain, (v2_compts_1(esk6_2(X1,X2),X1)|r2_hidden(esk3_2(X1,X2),X2)|k3_tarski(k2_xboole_0(X3,X4))!=k1_xboole_0|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(k2_xboole_0(X3,X4),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v2_compts_1(k3_tarski(X3),X1)|~v2_compts_1(k3_tarski(X4),X1)|~m1_subset_1(k3_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k3_tarski(X4),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_4452, plain, (v1_finset_1(k3_tarski(X1))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_4661, plain, (k3_tarski(k1_zfmisc_1(k1_xboole_0))=k1_xboole_0)).#trainneg
cnf(i_0_4657, negated_conjecture, (k3_tarski(k1_zfmisc_1(esk2_0))=k1_xboole_0|~m1_subset_1(k3_tarski(esk2_0),esk2_0))).#trainneg
cnf(i_0_382, plain, (r1_tarski(esk4_2(X1,X2),X2)|v2_compts_1(esk6_2(X1,X2),X1)|esk10_1(u1_struct_0(X1))!=k1_xboole_0|~v2_compts_1(esk10_1(u1_struct_0(X1)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3932, negated_conjecture, (v7_ordinal1(esk10_1(k3_tarski(esk2_0)))|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3933, negated_conjecture, (v3_ordinal1(esk10_1(k3_tarski(esk2_0)))|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_383, plain, (r1_tarski(esk4_2(X1,X2),X2)|v2_compts_1(esk6_2(X1,X2),X1)|esk9_1(k1_zfmisc_1(u1_struct_0(X1)))!=k1_xboole_0|~v2_compts_1(esk9_1(k1_zfmisc_1(u1_struct_0(X1))),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_385, plain, (r1_tarski(esk4_2(X1,X2),X2)|v2_compts_1(esk6_2(X1,X2),X1)|k1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(k1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_4526, negated_conjecture, (v1_xboole_0(esk2_0)|v2_compts_1(esk3_2(X1,X2),esk1_0)|esk6_2(X1,X2)=k3_tarski(X2)|~m1_subset_1(esk3_2(X1,X2),esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4459, plain, (v1_xboole_0(k3_tarski(X1))|~v1_xboole_0(k2_xboole_0(X1,X2)))).#trainneg
cnf(i_0_4280, plain, (k2_xboole_0(k2_xboole_0(X2,X3),X4)=k1_xboole_0|~v1_xboole_0(X1)|~m1_subset_1(X4,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_4342, plain, (k2_xboole_0(X1,k3_tarski(k1_zfmisc_1(X2)))=k1_xboole_0|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_4560, negated_conjecture, (esk6_2(X1,X2)=k3_tarski(X2)|v1_xboole_0(esk2_0)|v2_compts_1(esk3_2(X1,X2),esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_4509, negated_conjecture, (m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|esk6_2(X1,X2)=k3_tarski(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k3_tarski(esk2_0))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4098, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k5_setfam_1(X3,X4))|~v1_xboole_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3))))).#trainneg
cnf(i_0_1015, negated_conjecture, (v2_compts_1(esk9_1(esk2_0),esk1_0))).#trainneg
cnf(i_0_388, plain, (esk6_2(X1,X2)=k3_tarski(X2)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk13_1(X1)!=k1_xboole_0|~v2_compts_1(esk13_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4734, negated_conjecture, (~m1_subset_1(k3_tarski(esk2_0),esk2_0))).#trainneg
cnf(i_0_4045, plain, (v1_xreal_0(k5_setfam_1(X1,X2))|~v3_membered(k1_zfmisc_1(X3))|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_389, plain, (esk6_2(X1,X2)=k3_tarski(X2)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk12_1(X1)!=k1_xboole_0|~v2_compts_1(esk12_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_390, plain, (esk6_2(X1,X2)=k3_tarski(X2)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk10_1(u1_struct_0(X1))!=k1_xboole_0|~v2_compts_1(esk10_1(u1_struct_0(X1)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4836, plain, (k3_tarski(k2_xboole_0(X1,k1_zfmisc_1(k1_xboole_0)))=k3_tarski(X1))).#trainneg
cnf(i_0_4831, plain, (k3_tarski(k2_xboole_0(k1_zfmisc_1(k1_xboole_0),X1))=k3_tarski(X1))).#trainneg
cnf(i_0_4842, plain, (~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_4970, plain, (X1=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_4743, plain, (k3_tarski(k2_xboole_0(X1,k1_zfmisc_1(k1_struct_0(X2))))=k3_tarski(X1)|~l1_struct_0(X2))).#trainneg
cnf(i_0_4758, plain, (k3_tarski(k1_struct_0(X1))=k1_xboole_0|~l1_struct_0(X1))).#trainneg
cnf(i_0_895, negated_conjecture, (m1_subset_1(esk9_1(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_391, plain, (esk6_2(X1,X2)=k3_tarski(X2)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk9_1(k1_zfmisc_1(u1_struct_0(X1)))!=k1_xboole_0|~v2_compts_1(esk9_1(k1_zfmisc_1(u1_struct_0(X1))),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4046, plain, (v1_xxreal_0(k5_setfam_1(X1,X2))|~v2_membered(k1_zfmisc_1(X3))|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_4047, plain, (v7_ordinal1(k5_setfam_1(X1,X2))|~v6_membered(k1_zfmisc_1(X3))|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_393, plain, (esk6_2(X1,X2)=k3_tarski(X2)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(k1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_394, plain, (esk6_2(X1,X2)=k3_tarski(X2)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k4_subset_1(u1_struct_0(X1),X3,X4)!=k1_xboole_0|~v2_compts_1(k4_subset_1(u1_struct_0(X1),X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_4505, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v1_xboole_0(X3)|r1_tarski(esk3_2(X4,X5),X3)|esk6_2(X4,X5)=k3_tarski(X5)|v1_xboole_0(esk3_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X3)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X5,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4))))|~m1_subset_1(X5,k1_zfmisc_1(k1_zfmisc_1(esk3_2(X1,X2))))|~v1_finset_1(X5)|~l1_pre_topc(X4)|~v2_pre_topc(X4))).#trainneg
cnf(i_0_4319, plain, (k2_xboole_0(X1,k3_tarski(esk10_1(k1_zfmisc_1(X2))))=k1_xboole_0|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_5131, plain, (k3_tarski(X1)=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_5225, negated_conjecture, (m1_subset_1(esk3_2(esk1_0,esk2_0),esk2_0)|v2_compts_1(esk6_2(esk1_0,esk2_0),esk1_0))).#trainneg
cnf(i_0_5113, plain, (esk10_1(k1_xboole_0)=k1_xboole_0)).#trainneg
cnf(i_0_395, plain, (esk6_2(X1,X2)=k3_tarski(X2)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k5_setfam_1(u1_struct_0(X1),X3)!=k1_xboole_0|~v2_compts_1(k5_setfam_1(u1_struct_0(X1),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_5171, plain, (v3_ordinal1(k1_xboole_0)|~v3_ordinal1(k1_struct_0(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_4048, plain, (v3_ordinal1(k5_setfam_1(X1,X2))|~v3_ordinal1(k1_zfmisc_1(X3))|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_5231, plain, (v1_xxreal_0(k5_setfam_1(X1,X2))|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v3_membered(k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_396, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|v2_compts_1(esk6_2(X1,X2),X1)|esk13_1(X1)!=k1_xboole_0|~v2_compts_1(esk13_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_397, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|v2_compts_1(esk6_2(X1,X2),X1)|esk12_1(X1)!=k1_xboole_0|~v2_compts_1(esk12_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_5132, plain, (k5_setfam_1(k1_xboole_0,X1)=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_5125, plain, (esk3_2(X1,X2)=k1_xboole_0|esk6_2(X1,X2)=k3_tarski(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_5114, plain, (esk9_1(k1_zfmisc_1(k1_xboole_0))=k1_xboole_0)).#trainneg
cnf(i_0_398, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|v2_compts_1(esk6_2(X1,X2),X1)|esk10_1(u1_struct_0(X1))!=k1_xboole_0|~v2_compts_1(esk10_1(u1_struct_0(X1)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4456, plain, (X1=k3_tarski(X2)|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_5198, negated_conjecture, (v3_membered(esk9_1(esk2_0))|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4086, plain, (v2_compts_1(k5_setfam_1(X1,X2),X3)|~v1_finset_1(u1_struct_0(X3))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_399, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|v2_compts_1(esk6_2(X1,X2),X1)|esk9_1(k1_zfmisc_1(u1_struct_0(X1)))!=k1_xboole_0|~v2_compts_1(esk9_1(k1_zfmisc_1(u1_struct_0(X1))),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_5122, plain, (esk7_2(X1,X2)=k1_xboole_0|r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_5177, plain, (k3_tarski(k2_xboole_0(X1,k1_struct_0(X2)))=k3_tarski(X1)|~l1_struct_0(X2))).#trainneg
cnf(i_0_5070, plain, (k3_tarski(k2_xboole_0(k2_xboole_0(k1_zfmisc_1(k1_xboole_0),X1),X2))=k3_tarski(k2_xboole_0(X1,X2)))).#trainneg
cnf(i_0_5691, plain, (k3_tarski(k2_xboole_0(X1,k2_xboole_0(k1_zfmisc_1(k1_xboole_0),X1)))=k3_tarski(X1))).#trainneg
cnf(i_0_5359, plain, (k3_tarski(esk10_1(k1_zfmisc_1(k1_xboole_0)))=k1_xboole_0)).#trainneg
cnf(i_0_5490, negated_conjecture, (esk6_2(esk1_0,esk2_0)=k3_tarski(esk2_0)|X1!=k1_xboole_0|~v2_compts_1(esk3_2(esk1_0,esk2_0),esk1_0)|~v2_compts_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_404, plain, (r2_hidden(esk3_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk13_1(X1)!=k1_xboole_0|~v2_compts_1(esk13_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_5442, plain, (~v1_subset_1(k1_xboole_0,k1_xboole_0))).#trainneg
cnf(i_0_405, plain, (r2_hidden(esk3_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk12_1(X1)!=k1_xboole_0|~v2_compts_1(esk12_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_406, plain, (r2_hidden(esk3_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk10_1(u1_struct_0(X1))!=k1_xboole_0|~v2_compts_1(esk10_1(u1_struct_0(X1)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_5530, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=k3_tarski(esk4_2(X1,X2))|X4!=k1_xboole_0|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_5680, plain, (k3_tarski(k2_xboole_0(k2_xboole_0(X1,k1_zfmisc_1(k1_xboole_0)),X2))=k3_tarski(k2_xboole_0(X1,X2)))).#trainneg
cnf(i_0_407, plain, (r2_hidden(esk3_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk9_1(k1_zfmisc_1(u1_struct_0(X1)))!=k1_xboole_0|~v2_compts_1(esk9_1(k1_zfmisc_1(u1_struct_0(X1))),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4069, negated_conjecture, (v1_xboole_0(esk2_0)|v2_compts_1(k5_setfam_1(X1,X2),esk1_0)|~m1_subset_1(k5_setfam_1(X1,X2),esk2_0)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_409, plain, (r2_hidden(esk3_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(k1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_410, plain, (r2_hidden(esk3_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k4_subset_1(u1_struct_0(X1),X3,X4)!=k1_xboole_0|~v2_compts_1(k4_subset_1(u1_struct_0(X1),X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_5075, plain, (k3_tarski(k2_xboole_0(X1,k2_xboole_0(k1_zfmisc_1(k1_xboole_0),X2)))=k3_tarski(k2_xboole_0(X1,X2)))).#trainneg
cnf(i_0_5745, plain, (k3_tarski(k2_xboole_0(X1,esk10_1(k1_zfmisc_1(k1_xboole_0))))=k3_tarski(X1))).#trainneg
cnf(i_0_5740, plain, (k3_tarski(k2_xboole_0(esk10_1(k1_zfmisc_1(k1_xboole_0)),X1))=k3_tarski(X1))).#trainneg
cnf(i_0_5912, plain, (k3_tarski(k2_xboole_0(X2,k2_xboole_0(X1,k1_zfmisc_1(k1_xboole_0))))=k3_tarski(k2_xboole_0(X1,X2)))).#trainneg
cnf(i_0_1444, plain, (v1_finset_1(esk12_1(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_5184, negated_conjecture, (v1_xreal_0(esk9_1(esk2_0))|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5836, negated_conjecture, (esk6_2(esk1_0,X1)=k3_tarski(X1)|X2!=k3_tarski(esk4_2(esk1_0,X1))|esk9_1(esk2_0)!=k1_xboole_0|~v2_compts_1(X2,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_411, plain, (r2_hidden(esk3_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k5_setfam_1(u1_struct_0(X1),X3)!=k1_xboole_0|~v2_compts_1(k5_setfam_1(u1_struct_0(X1),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_2163, negated_conjecture, (v1_xreal_0(X1)|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_2164, negated_conjecture, (v1_xxreal_0(X1)|~v2_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_412, plain, (v2_compts_1(esk5_2(X1,X2),X1)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk13_1(X1)!=k1_xboole_0|~v2_compts_1(esk13_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_413, plain, (v2_compts_1(esk5_2(X1,X2),X1)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk12_1(X1)!=k1_xboole_0|~v2_compts_1(esk12_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_5117, plain, (esk7_2(k2_xboole_0(X1,X2),X3)=k1_xboole_0|r1_tarski(k2_xboole_0(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_5227, negated_conjecture, (r2_hidden(esk3_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|v2_compts_1(esk6_2(esk1_0,esk2_0),esk1_0))).#trainneg
cnf(i_0_5377, plain, (k3_tarski(k3_tarski(X1))=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_5073, plain, (v2_compts_1(k3_tarski(X1),X2)|~m1_subset_1(k2_xboole_0(k1_zfmisc_1(k1_xboole_0),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_finset_1(k3_tarski(X1))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_5185, negated_conjecture, (v1_xxreal_0(esk9_1(esk2_0))|~v2_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5783, negated_conjecture, (esk6_2(esk1_0,esk2_0)=k3_tarski(esk2_0)|X1!=k1_xboole_0|~v2_compts_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk3_2(esk1_0,esk2_0),esk2_0))).#trainneg
cnf(i_0_414, plain, (v2_compts_1(esk5_2(X1,X2),X1)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk10_1(u1_struct_0(X1))!=k1_xboole_0|~v2_compts_1(esk10_1(u1_struct_0(X1)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_5739, plain, (v3_ordinal1(k1_xboole_0)|~v3_ordinal1(esk10_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_5186, negated_conjecture, (v7_ordinal1(esk9_1(esk2_0))|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2165, negated_conjecture, (v7_ordinal1(X1)|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_2166, negated_conjecture, (v3_ordinal1(X1)|~v3_ordinal1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_415, plain, (v2_compts_1(esk5_2(X1,X2),X1)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk9_1(k1_zfmisc_1(u1_struct_0(X1)))!=k1_xboole_0|~v2_compts_1(esk9_1(k1_zfmisc_1(u1_struct_0(X1))),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_417, plain, (v2_compts_1(esk5_2(X1,X2),X1)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(k1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_5865, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=k3_tarski(esk4_2(X1,X2))|~v2_compts_1(k1_xboole_0,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_5512, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=esk5_2(X1,X2)|X4!=k1_xboole_0|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_6284, plain, (esk6_2(X1,X2)=k3_tarski(X2)|esk6_2(X3,X4)=k3_tarski(X4)|esk3_2(X3,X4)!=k3_tarski(esk4_2(X1,X2))|~v2_compts_1(esk3_2(X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X4)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_5187, negated_conjecture, (v3_ordinal1(esk9_1(esk2_0))|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_6306, negated_conjecture, (esk6_2(esk1_0,X1)=k3_tarski(X1)|X2!=esk5_2(esk1_0,X1)|esk9_1(esk2_0)!=k1_xboole_0|~v2_compts_1(X2,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_418, plain, (v2_compts_1(esk5_2(X1,X2),X1)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k4_subset_1(u1_struct_0(X1),X3,X4)!=k1_xboole_0|~v2_compts_1(k4_subset_1(u1_struct_0(X1),X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_6259, negated_conjecture, (v1_xreal_0(esk9_1(esk2_0))|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_6148, negated_conjecture, (v1_xxreal_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk2_0))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_4085, plain, (v2_compts_1(k5_setfam_1(X1,X2),X3)|~v1_finset_1(k5_setfam_1(X1,X2))|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_419, plain, (v2_compts_1(esk5_2(X1,X2),X1)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k5_setfam_1(u1_struct_0(X1),X3)!=k1_xboole_0|~v2_compts_1(k5_setfam_1(u1_struct_0(X1),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_420, plain, (v2_compts_1(esk6_2(X1,X2),X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk13_1(X1)!=k1_xboole_0|~v2_compts_1(esk13_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_6091, plain, (esk6_2(X1,X2)=k3_tarski(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v5_finset_1(X2))).#trainneg
cnf(i_0_6296, plain, (esk6_2(X1,X2)=k3_tarski(X2)|k3_tarski(X3)!=k3_tarski(esk4_2(X1,X2))|~v2_compts_1(k3_tarski(X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_6216, plain, (k3_tarski(esk5_2(X1,X2))=k1_xboole_0|esk6_2(X1,X2)=k3_tarski(X2)|~m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_6293, plain, (esk6_2(X1,X2)=k3_tarski(X2)|k2_xboole_0(X3,X4)!=k3_tarski(esk4_2(X1,X2))|~v2_compts_1(k2_xboole_0(X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_6197, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(X2)))=k3_tarski(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_6260, negated_conjecture, (v1_xxreal_0(esk9_1(esk2_0))|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_6268, negated_conjecture, (esk6_2(esk1_0,X1)=k3_tarski(X1)|esk9_1(esk2_0)!=k3_tarski(esk4_2(esk1_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_421, plain, (v2_compts_1(esk6_2(X1,X2),X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk12_1(X1)!=k1_xboole_0|~v2_compts_1(esk12_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4845, negated_conjecture, (v1_xreal_0(esk10_1(k3_tarski(esk2_0)))|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4084, plain, (v2_compts_1(k5_setfam_1(X1,X2),X3)|~v1_xboole_0(k5_setfam_1(X1,X2))|~l1_pre_topc(X3)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_422, plain, (v2_compts_1(esk6_2(X1,X2),X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk10_1(u1_struct_0(X1))!=k1_xboole_0|~v2_compts_1(esk10_1(u1_struct_0(X1)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_423, plain, (v2_compts_1(esk6_2(X1,X2),X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk9_1(k1_zfmisc_1(u1_struct_0(X1)))!=k1_xboole_0|~v2_compts_1(esk9_1(k1_zfmisc_1(u1_struct_0(X1))),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_6286, negated_conjecture, (esk6_2(esk1_0,X1)=k3_tarski(X1)|esk6_2(X2,esk2_0)=k3_tarski(esk2_0)|esk3_2(X2,esk2_0)!=k3_tarski(esk4_2(esk1_0,X1))|~v2_compts_1(esk3_2(X2,esk2_0),esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~v1_finset_1(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_6089, negated_conjecture, (esk6_2(esk1_0,X1)=k3_tarski(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_6291, plain, (esk6_2(X1,X2)=k3_tarski(X2)|k4_subset_1(u1_struct_0(X1),X3,X4)!=k3_tarski(esk4_2(X1,X2))|~v2_compts_1(k4_subset_1(u1_struct_0(X1),X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_6298, plain, (esk6_2(X1,X2)=k3_tarski(X2)|k5_setfam_1(u1_struct_0(X1),X3)!=k3_tarski(esk4_2(X1,X2))|~v2_compts_1(k5_setfam_1(u1_struct_0(X1),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_6335, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=esk5_2(X1,X2)|~v2_compts_1(k1_xboole_0,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_6085, plain, (esk6_2(X1,X2)=k3_tarski(X2)|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_6090, plain, (esk6_2(X1,X2)=k3_tarski(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk3_2(X1,X2),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_6156, plain, (r1_tarski(k2_xboole_0(X1,X2),X3)|~r2_hidden(k1_xboole_0,X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_4846, negated_conjecture, (v1_xxreal_0(esk10_1(k3_tarski(esk2_0)))|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_425, plain, (v2_compts_1(esk6_2(X1,X2),X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(k1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_4864, negated_conjecture, (r1_tarski(esk4_2(esk1_0,esk2_0),esk2_0)|v2_compts_1(esk6_2(esk1_0,esk2_0),esk1_0)|k1_struct_0(esk1_0)!=k1_xboole_0)).#trainneg
cnf(i_0_426, plain, (v2_compts_1(esk6_2(X1,X2),X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k4_subset_1(u1_struct_0(X1),X3,X4)!=k1_xboole_0|~v2_compts_1(k4_subset_1(u1_struct_0(X1),X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_427, plain, (v2_compts_1(esk6_2(X1,X2),X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k5_setfam_1(u1_struct_0(X1),X3)!=k1_xboole_0|~v2_compts_1(k5_setfam_1(u1_struct_0(X1),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_4802, negated_conjecture, (r2_hidden(esk3_2(esk1_0,esk2_0),esk2_0))).#trainneg
cnf(i_0_5225, negated_conjecture, (m1_subset_1(esk3_2(esk1_0,esk2_0),esk2_0))).#trainneg
cnf(i_0_5227, negated_conjecture, (r2_hidden(esk3_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_6458, plain, (k3_tarski(k2_xboole_0(X1,esk5_2(X2,X3)))=k3_tarski(X1)|esk6_2(X2,X3)=k3_tarski(X3)|~m1_subset_1(esk4_2(X2,X3),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_finset_1(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_6218, plain, (k3_tarski(k3_tarski(X1))=k1_xboole_0|~m1_subset_1(k2_xboole_0(esk10_1(k1_zfmisc_1(k1_xboole_0)),X1),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_6224, plain, (k3_tarski(k3_tarski(X1))=k1_xboole_0|~m1_subset_1(k2_xboole_0(X1,esk10_1(k1_zfmisc_1(k1_xboole_0))),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_5360, plain, (k3_tarski(esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))=k1_xboole_0)).#trainneg
cnf(i_0_6488, negated_conjecture, (esk6_2(esk1_0,X1)=k3_tarski(X1)|esk9_1(esk2_0)!=esk5_2(esk1_0,X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~v1_finset_1(X1))).#trainneg
cnf(i_0_428, plain, (r1_tarski(esk4_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk13_1(X1)!=k1_xboole_0|~v2_compts_1(esk13_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_429, plain, (r1_tarski(esk4_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk12_1(X1)!=k1_xboole_0|~v2_compts_1(esk12_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_430, plain, (r1_tarski(esk4_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk10_1(u1_struct_0(X1))!=k1_xboole_0|~v2_compts_1(esk10_1(u1_struct_0(X1)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_6047, plain, (m1_subset_1(k3_tarski(X1),k1_zfmisc_1(X2))|~m1_subset_1(k2_xboole_0(esk10_1(k1_zfmisc_1(k1_xboole_0)),X1),k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_6575, plain, (r1_tarski(X1,X2)|~r2_hidden(k1_xboole_0,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_6460, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(X2)))=k3_tarski(X1)|~m1_subset_1(k2_xboole_0(esk10_1(k1_zfmisc_1(k1_xboole_0)),X2),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_6088, negated_conjecture, (esk6_2(esk1_0,X1)=k3_tarski(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~v1_finset_1(X1)|~m1_subset_1(esk3_2(esk1_0,X1),esk2_0))).#trainneg
cnf(i_0_6461, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(k2_xboole_0(X2,X3))))=k3_tarski(X1)|~m1_subset_1(k2_xboole_0(k2_xboole_0(X2,k1_zfmisc_1(k1_xboole_0)),X3),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_431, plain, (r1_tarski(esk4_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk9_1(k1_zfmisc_1(u1_struct_0(X1)))!=k1_xboole_0|~v2_compts_1(esk9_1(k1_zfmisc_1(u1_struct_0(X1))),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_4007, negated_conjecture, (v3_membered(esk9_1(k1_zfmisc_1(k3_tarski(esk2_0))))|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3242, plain, (v1_xboole_0(X1)|r1_tarski(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k5_setfam_1(X1,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_3276, plain, (v1_xboole_0(X1)|m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k5_setfam_1(X1,X3)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_433, plain, (r1_tarski(esk4_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(k1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_434, plain, (r1_tarski(esk4_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k4_subset_1(u1_struct_0(X1),X3,X4)!=k1_xboole_0|~v2_compts_1(k4_subset_1(u1_struct_0(X1),X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_6775, plain, (k3_tarski(k2_xboole_0(X1,esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))=k3_tarski(X1))).#trainneg
cnf(i_0_6803, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r2_hidden(k1_xboole_0,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_6869, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(X2)))=k3_tarski(X1)|~m1_subset_1(k2_xboole_0(k2_xboole_0(esk10_1(k1_zfmisc_1(k1_xboole_0)),k1_zfmisc_1(k1_xboole_0)),X2),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_435, plain, (r1_tarski(esk4_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k5_setfam_1(u1_struct_0(X1),X3)!=k1_xboole_0|~v2_compts_1(k5_setfam_1(u1_struct_0(X1),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_6920, plain, (v1_xboole_0(X1)|r1_tarski(k5_setfam_1(X2,X3),X1)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_436, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk13_1(X1)!=k1_xboole_0|~v2_compts_1(esk13_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_437, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk12_1(X1)!=k1_xboole_0|~v2_compts_1(esk12_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_6769, plain, (k3_tarski(k2_xboole_0(esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))),X1))=k3_tarski(X1))).#trainneg
cnf(i_0_6858, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(X2)))=k3_tarski(X1)|~m1_subset_1(k2_xboole_0(k1_zfmisc_1(k1_xboole_0),X2),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_6804, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,X3)|~r2_hidden(k1_xboole_0,X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_6655, negated_conjecture, (v1_finset_1(esk3_2(esk1_0,esk2_0))|~v5_finset_1(esk2_0))).#trainneg
cnf(i_0_6628, negated_conjecture, (m1_subset_1(esk4_2(esk1_0,esk2_0),k1_zfmisc_1(esk2_0))|k1_struct_0(esk1_0)!=k1_xboole_0)).#trainneg
cnf(i_0_438, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk10_1(u1_struct_0(X1))!=k1_xboole_0|~v2_compts_1(esk10_1(u1_struct_0(X1)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_6768, plain, (v3_ordinal1(k1_xboole_0)|~v3_ordinal1(esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_6656, negated_conjecture, (v1_xreal_0(esk3_2(esk1_0,esk2_0))|~v3_membered(esk2_0))).#trainneg
cnf(i_0_6919, plain, (v1_xboole_0(X1)|r1_tarski(k5_setfam_1(k5_setfam_1(X1,X2),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X1,X2)))))).#trainneg
cnf(i_0_439, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk9_1(k1_zfmisc_1(u1_struct_0(X1)))!=k1_xboole_0|~v2_compts_1(esk9_1(k1_zfmisc_1(u1_struct_0(X1))),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_441, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(k1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_6868, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(X2)))=k3_tarski(X1)|~m1_subset_1(k2_xboole_0(k2_xboole_0(X2,k1_zfmisc_1(k1_xboole_0)),X2),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_5637, plain, (k3_tarski(k2_xboole_0(k1_struct_0(X2),X1))=k3_tarski(X1)|~l1_struct_0(X2))).#trainneg
cnf(i_0_6657, negated_conjecture, (v1_xxreal_0(esk3_2(esk1_0,esk2_0))|~v2_membered(esk2_0))).#trainneg
cnf(i_0_6509, negated_conjecture, (esk6_2(X1,esk2_0)=k3_tarski(esk2_0)|esk6_2(esk1_0,X2)=k3_tarski(X2)|esk3_2(X1,esk2_0)!=esk5_2(esk1_0,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_442, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k4_subset_1(u1_struct_0(X1),X3,X4)!=k1_xboole_0|~v2_compts_1(k4_subset_1(u1_struct_0(X1),X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_6658, negated_conjecture, (v7_ordinal1(esk3_2(esk1_0,esk2_0))|~v6_membered(esk2_0))).#trainneg
cnf(i_0_6659, negated_conjecture, (v3_ordinal1(esk3_2(esk1_0,esk2_0))|~v3_ordinal1(esk2_0))).#trainneg
cnf(i_0_7067, plain, (v1_xboole_0(X1)|r1_tarski(k5_setfam_1(X2,X3),X1)|~v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_6947, plain, (v1_xboole_0(X1)|m1_subset_1(k5_setfam_1(k5_setfam_1(X1,X2),X3),k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X1,X2)))))).#trainneg
cnf(i_0_443, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k5_setfam_1(u1_struct_0(X1),X3)!=k1_xboole_0|~v2_compts_1(k5_setfam_1(u1_struct_0(X1),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_445, plain, (m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk13_1(X1)!=k1_xboole_0|~v2_compts_1(esk13_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_7105, plain, (v2_compts_1(k3_tarski(X1),X2)|~m1_subset_1(k2_xboole_0(esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))),X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_finset_1(k3_tarski(X1))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_7102, plain, (k3_tarski(k2_xboole_0(k2_xboole_0(esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))),X1),X2))=k3_tarski(k2_xboole_0(X1,X2)))).#trainneg
cnf(i_0_7041, plain, (m1_subset_1(k3_tarski(X1),k1_zfmisc_1(X2))|~r2_hidden(k1_xboole_0,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_446, plain, (m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk12_1(X1)!=k1_xboole_0|~v2_compts_1(esk12_1(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_7045, plain, (m1_subset_1(k1_zfmisc_1(k1_xboole_0),k1_zfmisc_1(X1))|~r2_hidden(k1_xboole_0,X1))).#trainneg
cnf(i_0_6679, negated_conjecture, (v3_membered(esk3_2(esk1_0,esk2_0))|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_447, plain, (m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk10_1(u1_struct_0(X1))!=k1_xboole_0|~v2_compts_1(esk10_1(u1_struct_0(X1)),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_448, plain, (m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|esk9_1(k1_zfmisc_1(u1_struct_0(X1)))!=k1_xboole_0|~v2_compts_1(esk9_1(k1_zfmisc_1(u1_struct_0(X1))),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_5071, plain, (m1_subset_1(k3_tarski(X1),k1_zfmisc_1(X2))|~m1_subset_1(k2_xboole_0(k1_zfmisc_1(k1_xboole_0),X1),k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_6989, plain, (k3_tarski(k2_xboole_0(k2_xboole_0(X1,esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))),X2))=k3_tarski(k2_xboole_0(X1,X2)))).#trainneg
cnf(i_0_6466, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(X2)))=k3_tarski(X1)|~m1_subset_1(k2_xboole_0(X2,esk10_1(k1_zfmisc_1(k1_xboole_0))),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_5311, plain, (k3_tarski(esk10_1(k1_zfmisc_1(X1)))=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_6743, plain, (k3_tarski(k2_xboole_0(esk5_2(X2,X3),X1))=k3_tarski(X1)|esk6_2(X2,X3)=k3_tarski(X3)|~m1_subset_1(esk4_2(X2,X3),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_finset_1(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_7310, negated_conjecture, (v1_xreal_0(esk3_2(esk1_0,esk2_0))|~v6_membered(esk2_0))).#trainneg
cnf(i_0_450, plain, (m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(k1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_struct_0(X1))).#trainneg
cnf(i_0_7311, negated_conjecture, (v1_xxreal_0(esk3_2(esk1_0,esk2_0))|~v6_membered(esk2_0))).#trainneg
cnf(i_0_451, plain, (m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k4_subset_1(u1_struct_0(X1),X3,X4)!=k1_xboole_0|~v2_compts_1(k4_subset_1(u1_struct_0(X1),X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_452, plain, (m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k5_setfam_1(u1_struct_0(X1),X3)!=k1_xboole_0|~v2_compts_1(k5_setfam_1(u1_struct_0(X1),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_7492, plain, (k3_tarski(k2_xboole_0(X1,k2_xboole_0(esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))),X1)))=k3_tarski(X1))).#trainneg
cnf(i_0_1847, plain, (v1_finset_1(esk13_1(X1))|~v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_459, plain, (r2_hidden(X1,u1_struct_0(X2))|~r2_hidden(X1,esk13_1(X2))|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_6665, negated_conjecture, (v1_xreal_0(esk3_2(esk1_0,esk2_0))|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7384, plain, (v1_finset_1(k5_setfam_1(k5_setfam_1(X1,X2),X3))|v1_xboole_0(X1)|~v1_finset_1(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X1,X2))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_7386, plain, (v3_membered(k5_setfam_1(k5_setfam_1(X1,X2),X3))|v1_xboole_0(X1)|~v3_membered(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X1,X2))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_460, plain, (r2_hidden(X1,u1_struct_0(X2))|~r2_hidden(X1,esk12_1(X2))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_464, plain, (r2_hidden(X1,u1_struct_0(X2))|~r2_hidden(X1,k1_struct_0(X2))|~l1_struct_0(X2))).#trainneg
cnf(i_0_7572, plain, (k3_tarski(k2_xboole_0(X2,k2_xboole_0(X1,esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))))=k3_tarski(k2_xboole_0(X1,X2)))).#trainneg
cnf(i_0_7617, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(X4)))=k3_tarski(X1)|esk6_2(X2,X3)=k3_tarski(X3)|~m1_subset_1(k2_xboole_0(esk5_2(X2,X3),X4),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(esk4_2(X2,X3),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_finset_1(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_7648, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(X4)))=k3_tarski(X1)|esk6_2(X2,X3)=k3_tarski(X3)|~m1_subset_1(k2_xboole_0(k2_xboole_0(esk5_2(X2,X3),k1_zfmisc_1(k1_xboole_0)),X4),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(esk4_2(X2,X3),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_finset_1(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_7447, plain, (k3_tarski(k2_xboole_0(X2,k2_xboole_0(esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))),X1)))=k3_tarski(k2_xboole_0(X2,X1)))).#trainneg
cnf(i_0_6666, negated_conjecture, (v1_xxreal_0(esk3_2(esk1_0,esk2_0))|~v2_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_465, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k4_subset_1(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_6667, negated_conjecture, (v7_ordinal1(esk3_2(esk1_0,esk2_0))|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_6668, negated_conjecture, (v3_ordinal1(esk3_2(esk1_0,esk2_0))|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7321, plain, (v1_finset_1(k5_setfam_1(k5_setfam_1(X1,X2),X3))|v1_xboole_0(X1)|~v5_finset_1(k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X1,X2))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_468, plain, (v3_ordinal1(k2_xboole_0(k3_tarski(X1),X2))|~v3_ordinal1(k2_xboole_0(X1,k1_tarski(X2))))).#trainneg
cnf(i_0_7630, plain, (k2_xboole_0(k3_tarski(esk5_2(X1,X2)),X3)=X3|esk6_2(X1,X2)=k3_tarski(X2)|~m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_7494, plain, (m1_subset_1(k3_tarski(X1),k1_zfmisc_1(X2))|v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k1_xboole_0,X2))).#trainneg
cnf(i_0_7618, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(X4)))=k3_tarski(X1)|esk6_2(X2,X3)=k3_tarski(X3)|~m1_subset_1(k2_xboole_0(k1_zfmisc_1(k1_xboole_0),k2_xboole_0(esk5_2(X2,X3),X4)),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(esk4_2(X2,X3),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_finset_1(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_4162, plain, (v1_xboole_0(k3_tarski(esk10_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7973, negated_conjecture, (v1_xreal_0(esk3_2(esk1_0,esk2_0))|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_474, plain, (m1_subset_1(esk7_2(esk10_1(X1),X2),X1)|r1_tarski(esk10_1(X1),X2))).#trainneg
cnf(i_0_7974, negated_conjecture, (v1_xxreal_0(esk3_2(esk1_0,esk2_0))|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7322, plain, (v1_xreal_0(k5_setfam_1(k5_setfam_1(X1,X2),X3))|v1_xboole_0(X1)|~v3_membered(k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X1,X2))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_7323, plain, (v1_xxreal_0(k5_setfam_1(k5_setfam_1(X1,X2),X3))|v1_xboole_0(X1)|~v2_membered(k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X1,X2))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_475, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk10_1(X2))|~m1_subset_1(X1,esk10_1(X2)))).#trainneg
cnf(i_0_484, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk5_2(X1,X2),X1)|u1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(u1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_8109, negated_conjecture, (v2_compts_1(k3_tarski(X1),esk1_0)|~m1_subset_1(k3_tarski(X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8152, plain, (v1_xboole_0(X1)|m1_subset_1(k3_tarski(X2),k1_zfmisc_1(X1))|~m1_subset_1(k2_xboole_0(esk10_1(k1_zfmisc_1(k1_xboole_0)),X2),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_8028, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(X4)))=k3_tarski(X1)|esk6_2(X2,X3)=k3_tarski(X3)|~m1_subset_1(k2_xboole_0(k2_xboole_0(k3_tarski(esk5_2(X2,X3)),k1_zfmisc_1(k1_xboole_0)),X4),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(esk4_2(X2,X3),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_finset_1(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_8162, plain, (v1_xboole_0(X1)|m1_subset_1(k3_tarski(X2),k1_zfmisc_1(X1))|~m1_subset_1(k2_xboole_0(k1_zfmisc_1(k1_xboole_0),X2),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_8163, plain, (v1_xboole_0(X1)|m1_subset_1(k3_tarski(X2),k1_zfmisc_1(X1))|~m1_subset_1(k2_xboole_0(X2,esk10_1(k1_zfmisc_1(k1_xboole_0))),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_8003, plain, (k3_tarski(k2_xboole_0(X3,k3_tarski(esk5_2(X1,X2))))=k3_tarski(X3)|esk6_2(X1,X2)=k3_tarski(X2)|~m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_8196, plain, (v1_finset_1(k3_tarski(esk10_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_7497, negated_conjecture, (m1_subset_1(k1_zfmisc_1(k1_xboole_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~r2_hidden(k1_xboole_0,esk2_0))).#trainneg
cnf(i_0_485, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(esk4_2(X1,X2),X2)|u1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(u1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_7324, plain, (v7_ordinal1(k5_setfam_1(k5_setfam_1(X1,X2),X3))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X1,X2))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_7325, plain, (v3_ordinal1(k5_setfam_1(k5_setfam_1(X1,X2),X3))|v1_xboole_0(X1)|~v3_ordinal1(k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X1,X2))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_487, plain, (esk6_2(X1,X2)=k3_tarski(X2)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|u1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(u1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_489, plain, (v2_compts_1(esk6_2(X1,X2),X1)|v2_compts_1(esk5_2(X1,X2),X1)|u1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(u1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_7527, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(k2_xboole_0(X2,X3))))=k3_tarski(X1)|~m1_subset_1(k2_xboole_0(k2_xboole_0(X2,esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))),X3),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_8176, plain, (v1_xboole_0(X1)|m1_subset_1(k3_tarski(X2),k1_zfmisc_1(X1))|~m1_subset_1(k2_xboole_0(X2,k3_tarski(X3)),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k1_xboole_0,X1)|~m1_subset_1(k2_xboole_0(k1_zfmisc_1(k1_xboole_0),X3),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_490, plain, (r1_tarski(esk4_2(X1,X2),X2)|v2_compts_1(esk6_2(X1,X2),X1)|u1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(u1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3965, negated_conjecture, (v2_compts_1(esk10_1(k3_tarski(esk2_0)),esk1_0)|~v1_xboole_0(esk10_1(k3_tarski(esk2_0))))).#trainneg
cnf(i_0_8461, plain, (v1_xreal_0(k5_setfam_1(k5_setfam_1(X1,X2),X3))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X1,X2))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_8462, plain, (v1_xxreal_0(k5_setfam_1(k5_setfam_1(X1,X2),X3))|v1_xboole_0(X1)|~v6_membered(k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X1,X2))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_492, plain, (r2_hidden(esk3_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|u1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(u1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_493, plain, (v2_compts_1(esk5_2(X1,X2),X1)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|u1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(u1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_8342, negated_conjecture, (v2_compts_1(k3_tarski(X1),esk1_0)|~m1_subset_1(k2_xboole_0(X1,esk10_1(k1_zfmisc_1(k1_xboole_0))),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0))|~m1_subset_1(k3_tarski(X1),esk2_0))).#trainneg
cnf(i_0_8337, negated_conjecture, (v2_compts_1(k3_tarski(k2_xboole_0(X1,X2)),esk1_0)|~m1_subset_1(k2_xboole_0(k2_xboole_0(X1,esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))),X2),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0))|~m1_subset_1(k3_tarski(k2_xboole_0(X1,X2)),esk2_0))).#trainneg
cnf(i_0_8357, negated_conjecture, (v2_compts_1(k3_tarski(X1),esk1_0)|~m1_subset_1(k2_xboole_0(X1,k1_zfmisc_1(k1_xboole_0)),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0))|~m1_subset_1(k3_tarski(X1),esk2_0))).#trainneg
cnf(i_0_8331, negated_conjecture, (v2_compts_1(k3_tarski(X1),esk1_0)|~m1_subset_1(k2_xboole_0(esk10_1(k1_zfmisc_1(k1_xboole_0)),X1),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0))|~m1_subset_1(k3_tarski(X1),esk2_0))).#trainneg
cnf(i_0_8341, negated_conjecture, (v2_compts_1(k3_tarski(X1),esk1_0)|~m1_subset_1(k2_xboole_0(k1_zfmisc_1(k1_xboole_0),X1),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0))|~m1_subset_1(k3_tarski(X1),esk2_0))).#trainneg
cnf(i_0_7469, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(k2_xboole_0(X2,X3))))=k3_tarski(X1)|~m1_subset_1(k2_xboole_0(k2_xboole_0(k2_xboole_0(esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))),X2),k1_zfmisc_1(k1_xboole_0)),X3),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_7885, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(X2)))=k3_tarski(X1)|esk6_2(X3,X4)=k3_tarski(X4)|~m1_subset_1(k2_xboole_0(X2,esk5_2(X3,X4)),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(esk4_2(X3,X4),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~v1_finset_1(X4)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_7624, plain, (k3_tarski(esk5_2(X3,X4))=k3_tarski(esk5_2(X1,X2))|esk6_2(X3,X4)=k3_tarski(X4)|esk6_2(X1,X2)=k3_tarski(X2)|~m1_subset_1(esk4_2(X3,X4),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~v1_finset_1(X4)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_7610, plain, (k3_tarski(k3_tarski(X3))=k1_xboole_0|esk6_2(X1,X2)=k3_tarski(X2)|~m1_subset_1(k2_xboole_0(esk5_2(X1,X2),X3),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3307, plain, (v1_finset_1(k3_tarski(X1))|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_3027, plain, (v1_xboole_0(k2_xboole_0(X1,X2))|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_494, plain, (v2_compts_1(esk6_2(X1,X2),X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|u1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(u1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_3241, plain, (v1_xboole_0(X1)|r1_tarski(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_3966, negated_conjecture, (v2_compts_1(esk10_1(k3_tarski(esk2_0)),esk1_0)|~v1_finset_1(esk10_1(k3_tarski(esk2_0))))).#trainneg
cnf(i_0_495, plain, (r1_tarski(esk4_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|u1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(u1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_496, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|u1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(u1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_8725, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(X2)))=k3_tarski(X1)|~m1_subset_1(k2_xboole_0(k2_xboole_0(k1_zfmisc_1(k1_xboole_0),esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))),X2),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_8735, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(X2)))=k3_tarski(X1)|~m1_subset_1(k2_xboole_0(k1_zfmisc_1(k1_xboole_0),k2_xboole_0(esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))),X2)),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_8822, negated_conjecture, (k2_xboole_0(esk5_2(esk1_0,esk2_0),k3_tarski(X1))=k3_tarski(k2_xboole_0(esk4_2(esk1_0,esk2_0),X1)))).#trainneg
cnf(i_0_8819, negated_conjecture, (k3_tarski(esk5_2(esk1_0,esk2_0))=k1_xboole_0|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_7594, plain, (k3_tarski(k2_xboole_0(X1,esk10_1(k1_zfmisc_1(X2))))=k3_tarski(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_8353, negated_conjecture, (v2_compts_1(k3_tarski(X1),esk1_0)|~m1_subset_1(k2_xboole_0(X1,k3_tarski(k2_xboole_0(X2,X3))),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0))|~m1_subset_1(k3_tarski(X1),esk2_0)|~m1_subset_1(k2_xboole_0(k2_xboole_0(X2,k1_zfmisc_1(k1_xboole_0)),X3),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_8811, plain, (v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_8786, plain, (v1_finset_1(k3_tarski(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_497, plain, (m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|u1_struct_0(X1)!=k1_xboole_0|~v2_compts_1(u1_struct_0(X1),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_9169, negated_conjecture, (~m1_subset_1(esk2_0,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_509, plain, (k1_struct_0(X1)=k1_struct_0(X2)|~l1_struct_0(X2)|~l1_struct_0(X1))).#trainneg
cnf(i_0_8113, plain, (v2_compts_1(esk6_2(X1,X2),X1)|v2_compts_1(esk5_2(X1,X2),X1)|v1_xboole_0(u1_struct_0(X1))|k3_tarski(X3)!=k1_xboole_0|~v2_compts_1(k3_tarski(X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k1_xboole_0,u1_struct_0(X1)))).#trainneg
cnf(i_0_8335, negated_conjecture, (v2_compts_1(k3_tarski(k2_xboole_0(X1,X2)),esk1_0)|~m1_subset_1(k2_xboole_0(k2_xboole_0(X1,k1_zfmisc_1(k1_xboole_0)),X2),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0))|~m1_subset_1(k3_tarski(k2_xboole_0(X1,X2)),esk2_0))).#trainneg
cnf(i_0_7557, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(k2_xboole_0(X2,X3))))=k3_tarski(X1)|~m1_subset_1(k2_xboole_0(k2_xboole_0(k2_xboole_0(X2,esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))),k1_zfmisc_1(k1_xboole_0)),X3),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_8178, plain, (v1_xboole_0(X1)|m1_subset_1(k3_tarski(X2),k1_zfmisc_1(X1))|~m1_subset_1(k2_xboole_0(X2,k1_zfmisc_1(k1_xboole_0)),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k1_xboole_0,X1))).#trainneg
cnf(i_0_8789, plain, (v1_finset_1(k2_xboole_0(X1,X2))|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_511, plain, (v1_finset_1(k3_tarski(X1))|~v5_finset_1(k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_9192, plain, (X1=X2|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_8817, negated_conjecture, (v1_xboole_0(esk5_2(esk1_0,esk2_0))|~v1_xboole_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_8818, negated_conjecture, (v1_finset_1(esk5_2(esk1_0,esk2_0))|~v1_xboole_0(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_512, plain, (v1_xreal_0(k3_tarski(X1))|~v3_membered(k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_513, plain, (v1_xxreal_0(k3_tarski(X1))|~v2_membered(k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_7493, negated_conjecture, (m1_subset_1(k3_tarski(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~r2_hidden(k1_xboole_0,esk2_0))).#trainneg
cnf(i_0_8759, plain, (v1_finset_1(k3_tarski(esk10_1(k1_zfmisc_1(X1))))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_514, plain, (v7_ordinal1(k3_tarski(X1))|~v6_membered(k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_515, plain, (v3_ordinal1(k3_tarski(X1))|~v3_ordinal1(k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_517, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk5_2(X1,X2),X1)|k3_tarski(X3)!=k1_xboole_0|~v2_compts_1(k3_tarski(X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_9334, negated_conjecture, (v2_compts_1(k3_tarski(k2_xboole_0(X1,X2)),esk1_0)|~m1_subset_1(k2_xboole_0(k2_xboole_0(k1_zfmisc_1(k1_xboole_0),X1),X2),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k3_tarski(k2_xboole_0(X1,X2)),esk2_0)|~m1_subset_1(k1_xboole_0,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9385, plain, (v1_finset_1(k2_xboole_0(X2,X1))|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_9443, negated_conjecture, (v1_xboole_0(esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_518, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(esk4_2(X1,X2),X2)|k3_tarski(X3)!=k1_xboole_0|~v2_compts_1(k3_tarski(X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_520, plain, (esk6_2(X1,X2)=k3_tarski(X2)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k3_tarski(X3)!=k1_xboole_0|~v2_compts_1(k3_tarski(X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_522, plain, (v2_compts_1(esk6_2(X1,X2),X1)|v2_compts_1(esk5_2(X1,X2),X1)|k3_tarski(X3)!=k1_xboole_0|~v2_compts_1(k3_tarski(X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_9364, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(k2_xboole_0(X2,X3))))=k3_tarski(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k2_xboole_0(k1_zfmisc_1(k1_xboole_0),k2_xboole_0(X2,esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_9503, negated_conjecture, (m1_subset_1(k3_tarski(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k1_xboole_0,esk2_0))).#trainneg
cnf(i_0_9349, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(k2_xboole_0(k3_tarski(esk5_2(X2,X3)),X4))))=k3_tarski(X1)|esk6_2(X2,X3)=k3_tarski(X3)|~m1_subset_1(k2_xboole_0(k2_xboole_0(k1_zfmisc_1(k1_xboole_0),esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))),X4),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(esk4_2(X2,X3),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~v1_finset_1(X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2))).#trainneg
cnf(i_0_9657, negated_conjecture, (X1!=k1_xboole_0|~v2_compts_1(X1,esk1_0)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_finset_1(esk3_2(esk1_0,esk2_0))|~v1_finset_1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_523, plain, (r1_tarski(esk4_2(X1,X2),X2)|v2_compts_1(esk6_2(X1,X2),X1)|k3_tarski(X3)!=k1_xboole_0|~v2_compts_1(k3_tarski(X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_8821, negated_conjecture, (v3_ordinal1(esk5_2(esk1_0,esk2_0))|~v3_ordinal1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_9444, negated_conjecture, (v1_finset_1(esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_4088, plain, (v1_xboole_0(X1)|r1_tarski(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k5_setfam_1(X3,X4)))|~v1_xboole_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3))))).#trainneg
cnf(i_0_4089, plain, (v1_xboole_0(X1)|m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(k5_setfam_1(X3,X4)))|~v1_xboole_0(X3)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(X3))))).#trainneg
cnf(i_0_525, plain, (r2_hidden(esk3_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k3_tarski(X3)!=k1_xboole_0|~v2_compts_1(k3_tarski(X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_526, plain, (v2_compts_1(esk5_2(X1,X2),X1)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k3_tarski(X3)!=k1_xboole_0|~v2_compts_1(k3_tarski(X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_9359, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(X2)))=k3_tarski(X1)|~m1_subset_1(k2_xboole_0(k1_zfmisc_1(k1_xboole_0),k2_xboole_0(X2,esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_9696, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(esk4_2(X1,X2),X2)|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_7937, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(k2_xboole_0(X2,X3))))=k3_tarski(X1)|~m1_subset_1(k2_xboole_0(k2_xboole_0(X2,k1_zfmisc_1(k1_xboole_0)),k2_xboole_0(esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))),X3)),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_9055, negated_conjecture, (k3_tarski(k2_xboole_0(esk4_2(esk1_0,esk2_0),k3_tarski(X1)))=esk5_2(esk1_0,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_10110, plain, (r2_hidden(X1,X2)|esk6_2(X3,X2)=k3_tarski(X2)|~r2_hidden(X1,esk4_2(X3,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~v1_finset_1(X2)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_10109, plain, (m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(X2))|esk6_2(X1,X2)=k3_tarski(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_9945, plain, (r1_tarski(esk4_2(X1,X2),X2)|v2_compts_1(esk6_2(X1,X2),X1)|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_8813, plain, (v1_xboole_0(k2_xboole_0(X2,X1))|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_9894, negated_conjecture, (X1!=k1_xboole_0|~v2_compts_1(esk5_2(esk1_0,esk2_0),esk1_0)|~v2_compts_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|k1_struct_0(esk1_0)!=k1_xboole_0)).#trainneg
cnf(i_0_527, plain, (v2_compts_1(esk6_2(X1,X2),X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k3_tarski(X3)!=k1_xboole_0|~v2_compts_1(k3_tarski(X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_3995, negated_conjecture, (v1_xreal_0(esk9_1(k1_zfmisc_1(k3_tarski(esk2_0))))|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3996, negated_conjecture, (v1_xxreal_0(esk9_1(k1_zfmisc_1(k3_tarski(esk2_0))))|~v2_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_528, plain, (r1_tarski(esk4_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k3_tarski(X3)!=k1_xboole_0|~v2_compts_1(k3_tarski(X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_529, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k3_tarski(X3)!=k1_xboole_0|~v2_compts_1(k3_tarski(X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_10310, plain, (v1_finset_1(esk4_2(X1,X2))|esk6_2(X1,X2)=k3_tarski(X2)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_10226, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(esk7_2(esk4_2(X1,X2),X3),X2)|r1_tarski(esk4_2(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_10340, negated_conjecture, (r1_tarski(esk4_2(esk1_0,esk2_0),esk2_0))).#trainneg
cnf(i_0_10230, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(esk3_2(X3,esk4_2(X1,X2)),X2)|esk6_2(X3,esk4_2(X1,X2))=k3_tarski(esk4_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~v1_finset_1(esk4_2(X1,X2))|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_10231, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(esk3_2(X3,X4),X2)|esk6_2(X3,X4)=k3_tarski(X4)|v1_xboole_0(esk4_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(esk4_2(X1,X2)))|~v1_finset_1(X4)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_10185, negated_conjecture, (k3_tarski(k2_xboole_0(esk4_2(esk1_0,esk2_0),k3_tarski(X1)))=esk5_2(esk1_0,esk2_0)|~m1_subset_1(k2_xboole_0(esk10_1(k1_zfmisc_1(k1_xboole_0)),X1),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_10303, plain, (v1_xboole_0(X1)|r1_tarski(X2,X1)|esk6_2(X3,X1)=k3_tarski(X1)|~m1_subset_1(X2,k1_zfmisc_1(esk4_2(X3,X1)))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~v1_finset_1(X1)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_10360, plain, (k2_xboole_0(X1,X2)=k1_xboole_0|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_10589, negated_conjecture, (m1_subset_1(esk4_2(esk1_0,esk2_0),k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_10386, negated_conjecture, (k1_struct_0(esk1_0)!=k1_xboole_0|X1!=k1_xboole_0|~v2_compts_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_530, plain, (m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k3_tarski(X3)!=k1_xboole_0|~v2_compts_1(k3_tarski(X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))).#trainneg
cnf(i_0_3997, negated_conjecture, (v7_ordinal1(esk9_1(k1_zfmisc_1(k3_tarski(esk2_0))))|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3998, negated_conjecture, (v3_ordinal1(esk9_1(k1_zfmisc_1(k3_tarski(esk2_0))))|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9980, plain, (v1_xboole_0(X1)|r1_tarski(k5_setfam_1(k5_setfam_1(X2,X3),X4),X1)|~v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X2,X3)))))).#trainneg
cnf(i_0_531, plain, (v2_compts_1(k3_tarski(X1),X2)|~v1_xboole_0(k3_tarski(X1))|~l1_pre_topc(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2)))))).#trainneg
cnf(i_0_536, plain, (v1_xboole_0(k3_tarski(X1))|~v1_xboole_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_10705, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(esk3_2(X3,X4),X2)|esk6_2(X3,X4)=k3_tarski(X4)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(esk4_2(X1,X2))))|~v1_finset_1(X4)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_10590, negated_conjecture, (r2_hidden(X1,esk2_0)|~r2_hidden(X1,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_10707, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(k4_subset_1(esk4_2(X1,X2),X3,X4),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(esk4_2(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(esk4_2(X1,X2))))).#trainneg
cnf(i_0_10715, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(k5_setfam_1(esk4_2(X1,X2),X3),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(esk4_2(X1,X2)))))).#trainneg
cnf(i_0_10582, plain, (m1_subset_1(esk7_2(esk4_2(X1,X2),X3),X2)|esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(esk4_2(X1,X2),X3)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_11006, plain, (v2_compts_1(X1,X2)|~v1_xboole_0(X1)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2))).#trainneg
cnf(i_0_10897, negated_conjecture, (v1_finset_1(esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_9891, negated_conjecture, (X1!=k1_xboole_0|~v2_compts_1(esk5_2(esk1_0,esk2_0),esk1_0)|~v2_compts_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|u1_struct_0(esk1_0)!=k1_xboole_0|~v2_compts_1(u1_struct_0(esk1_0),esk1_0))).#trainneg
cnf(i_0_537, plain, (v3_membered(k3_tarski(X1))|~v3_membered(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_10899, negated_conjecture, (v3_membered(esk4_2(esk1_0,esk2_0))|~v3_membered(esk2_0))).#trainneg
cnf(i_0_10011, plain, (v1_xboole_0(X1)|m1_subset_1(k5_setfam_1(k5_setfam_1(X2,X3),X4),k1_zfmisc_1(X1))|~v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(X2)))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X2,X3)))))).#trainneg
cnf(i_0_540, plain, (m1_subset_1(k2_xboole_0(k3_tarski(X1),X2),k1_zfmisc_1(X3))|~m1_subset_1(k2_xboole_0(X1,k1_tarski(X2)),k1_zfmisc_1(k1_zfmisc_1(X3))))).#trainneg
cnf(i_0_566, plain, (m1_subset_1(X1,X2)|v1_xboole_0(k4_subset_1(X2,X3,X4))|~m1_subset_1(X4,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,k4_subset_1(X2,X3,X4)))).#trainneg
cnf(i_0_11084, plain, (r2_hidden(X1,X2)|esk6_2(X3,X2)=k3_tarski(X2)|~r2_hidden(X1,k4_subset_1(esk4_2(X3,X2),X4,X5))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~m1_subset_1(X5,k1_zfmisc_1(esk4_2(X3,X2)))|~m1_subset_1(X4,k1_zfmisc_1(esk4_2(X3,X2)))|~v1_finset_1(X2)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_10901, negated_conjecture, (m1_subset_1(X1,esk2_0)|~r2_hidden(X1,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_11076, negated_conjecture, (r2_hidden(esk7_2(esk4_2(esk1_0,esk2_0),X1),esk2_0)|r1_tarski(esk4_2(esk1_0,esk2_0),X1))).#trainneg
cnf(i_0_11092, plain, (r2_hidden(X1,X2)|esk6_2(X3,X2)=k3_tarski(X2)|~r2_hidden(X1,k5_setfam_1(esk4_2(X3,X2),X4))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(esk4_2(X3,X2))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~v1_finset_1(X2)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_10907, negated_conjecture, (r1_tarski(esk4_2(esk1_0,esk2_0),X1)|m1_subset_1(esk7_2(esk4_2(esk1_0,esk2_0),X1),X2)|~m1_subset_1(esk2_0,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_572, negated_conjecture, (v1_xboole_0(esk2_0)|v2_compts_1(k4_subset_1(u1_struct_0(esk1_0),X1,X2),esk1_0)|~m1_subset_1(k4_subset_1(u1_struct_0(esk1_0),X1,X2),esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_11350, plain, (v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X3,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_11317, plain, (v3_membered(k5_setfam_1(k5_setfam_1(X1,X2),X3))|v1_xboole_0(X4)|~v3_membered(X4)|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X1,X2))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_573, negated_conjecture, (v1_xboole_0(esk2_0)|v2_compts_1(k3_tarski(X1),esk1_0)|~m1_subset_1(k3_tarski(X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_595, plain, (m1_subset_1(esk7_2(esk9_1(k1_zfmisc_1(X1)),X2),X1)|r1_tarski(esk9_1(k1_zfmisc_1(X1)),X2))).#trainneg
cnf(i_0_11370, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(X3,X2)|~r2_hidden(X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(esk4_2(X1,X2)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_11530, negated_conjecture, (v2_compts_1(X1,esk1_0)|~m1_subset_1(X1,esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_11568, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_11614, negated_conjecture, (v2_compts_1(esk5_2(esk1_0,esk2_0),esk1_0)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(esk5_2(esk1_0,esk2_0),esk2_0))).#trainneg
cnf(i_0_11770, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(X3,X2)|esk6_2(X4,X5)=k3_tarski(X5)|~r2_hidden(X3,esk3_2(X4,X5))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X5,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4))))|~m1_subset_1(X5,k1_zfmisc_1(k1_zfmisc_1(esk4_2(X1,X2))))|~v1_finset_1(X5)|~l1_pre_topc(X4)|~v2_pre_topc(X4))).#trainneg
cnf(i_0_11081, negated_conjecture, (r2_hidden(esk3_2(X1,X2),esk2_0)|esk6_2(X1,X2)=k3_tarski(X2)|v1_xboole_0(esk4_2(esk1_0,esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(esk4_2(esk1_0,esk2_0)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_10702, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(esk4_2(X3,esk4_2(X1,X2)),X2)|esk6_2(X3,esk4_2(X1,X2))=k3_tarski(esk4_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~v1_finset_1(esk4_2(X1,X2))|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_11194, negated_conjecture, (v2_compts_1(X1,esk1_0)|~v1_xboole_0(X1)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_11043, plain, (v1_xboole_0(k3_tarski(k1_zfmisc_1(X1)))|~v1_xboole_0(X1))).#trainneg
cnf(i_0_10885, negated_conjecture, (v1_xreal_0(esk4_2(esk1_0,esk2_0))|~v3_membered(k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_11199, negated_conjecture, (u1_struct_0(esk1_0)!=k1_xboole_0|X1!=k1_xboole_0|~v2_compts_1(u1_struct_0(esk1_0),esk1_0)|~v2_compts_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_596, plain, (m1_subset_1(X1,X2)|v1_xboole_0(esk9_1(k1_zfmisc_1(X2)))|~m1_subset_1(X1,esk9_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_10886, negated_conjecture, (v1_xxreal_0(esk4_2(esk1_0,esk2_0))|~v2_membered(k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_10887, negated_conjecture, (v7_ordinal1(esk4_2(esk1_0,esk2_0))|~v6_membered(k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_609, plain, (k5_setfam_1(X1,X2)=k1_struct_0(X3)|~l1_struct_0(X3)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_10709, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(k2_xboole_0(X3,X4),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(esk4_2(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(esk4_2(X1,X2))))).#trainneg
cnf(i_0_11870, negated_conjecture, (v2_compts_1(X1,esk1_0)|~m1_subset_1(X1,esk2_0))).#trainneg
cnf(i_0_11839, negated_conjecture, (X1!=k1_xboole_0|~v2_compts_1(k2_xboole_0(esk5_2(esk1_0,esk2_0),esk3_2(esk1_0,esk2_0)),esk1_0)|~v2_compts_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_xboole_0(esk5_2(esk1_0,esk2_0),esk3_2(esk1_0,esk2_0)),esk2_0))).#trainneg
cnf(i_0_10891, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(esk4_2(esk1_0,esk2_0))))).#trainneg
cnf(i_0_10540, negated_conjecture, (v2_compts_1(esk7_2(esk4_2(X1,esk2_0),X2),esk1_0)|esk6_2(X1,esk2_0)=k3_tarski(esk2_0)|r1_tarski(esk4_2(X1,esk2_0),X2)|~m1_subset_1(esk7_2(esk4_2(X1,esk2_0),X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_890, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_tarski(k5_setfam_1(k1_zfmisc_1(X2),X3)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X2)))))).#trainneg
cnf(i_0_10888, negated_conjecture, (v3_ordinal1(esk4_2(esk1_0,esk2_0))|~v3_ordinal1(k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_11837, negated_conjecture, (X1!=k1_xboole_0|~v2_compts_1(esk5_2(esk1_0,esk2_0),esk1_0)|~v2_compts_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk5_2(esk1_0,esk2_0),esk2_0))).#trainneg
cnf(i_0_11953, negated_conjecture, (v1_xreal_0(esk4_2(esk1_0,esk2_0))|~v6_membered(k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_11954, negated_conjecture, (v1_xxreal_0(esk4_2(esk1_0,esk2_0))|~v6_membered(k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_615, plain, (v2_compts_1(esk10_1(u1_struct_0(X1)),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~v1_finset_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_11386, negated_conjecture, (m1_subset_1(esk7_2(esk4_2(esk1_0,esk2_0),X1),esk2_0)|r1_tarski(esk4_2(esk1_0,esk2_0),X1))).#trainneg
cnf(i_0_10908, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(esk4_2(esk1_0,esk2_0))))).#trainneg
cnf(i_0_11514, negated_conjecture, (v1_xboole_0(X1)|r1_tarski(esk4_2(esk1_0,esk2_0),X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_10355, plain, (k2_xboole_0(X1,X2)=k1_struct_0(X3)|~l1_struct_0(X3)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_12263, negated_conjecture, (m1_subset_1(esk4_2(esk1_0,esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_12264, negated_conjecture, (m1_subset_1(esk4_2(esk1_0,esk2_0),k1_zfmisc_1(X1))|v1_xboole_0(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_12157, negated_conjecture, (~m1_subset_1(k5_setfam_1(u1_struct_0(esk1_0),esk2_0),esk2_0))).#trainneg
cnf(i_0_12213, negated_conjecture, (v3_ordinal1(esk5_2(esk1_0,esk2_0))|~v3_ordinal1(k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_9906, negated_conjecture, (X1!=k1_xboole_0|~v2_compts_1(X1,esk1_0)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_finset_1(esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk3_2(esk1_0,esk2_0),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_619, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk5_2(X1,X2),X1)|r1_tarski(k1_zfmisc_1(u1_struct_0(X1)),X3)|esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X3)!=k1_xboole_0|~v2_compts_1(esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_8883, negated_conjecture, (v2_compts_1(esk10_1(k3_tarski(esk2_0)),esk1_0)|~v5_finset_1(k1_zfmisc_1(k3_tarski(esk2_0))))).#trainneg
cnf(i_0_10957, negated_conjecture, (v1_xreal_0(esk9_1(k1_zfmisc_1(k3_tarski(esk2_0))))|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_622, plain, (esk6_2(X1,X2)=k3_tarski(X2)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|r1_tarski(k1_zfmisc_1(u1_struct_0(X1)),X3)|esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X3)!=k1_xboole_0|~v2_compts_1(esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_624, plain, (v2_compts_1(esk6_2(X1,X2),X1)|v2_compts_1(esk5_2(X1,X2),X1)|r1_tarski(k1_zfmisc_1(u1_struct_0(X1)),X3)|esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X3)!=k1_xboole_0|~v2_compts_1(esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_12194, negated_conjecture, (esk6_2(X1,esk2_0)=k3_tarski(esk2_0)|r1_tarski(esk4_2(X1,esk2_0),X2)|v2_compts_1(esk7_2(esk4_2(X1,esk2_0),X2),esk1_0)|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk4_2(X1,esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_12265, negated_conjecture, (r2_hidden(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_12058, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(k2_xboole_0(X3,esk10_1(esk4_2(X1,X2))),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(esk4_2(X1,X2)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_11852, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X2,esk2_0))).#trainneg
cnf(i_0_10359, plain, (k2_xboole_0(X1,X2)=k3_tarski(X3)|~v1_xboole_0(X3)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_920, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_tarski(k5_setfam_1(k1_zfmisc_1(X2),X3)))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X2)))))).#trainneg
cnf(i_0_11231, plain, (v3_membered(k3_tarski(k1_zfmisc_1(X1)))|~v3_membered(X1))).#trainneg
cnf(i_0_10958, negated_conjecture, (v1_xxreal_0(esk9_1(k1_zfmisc_1(k3_tarski(esk2_0))))|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9907, negated_conjecture, (X1!=k1_xboole_0|~v2_compts_1(X1,esk1_0)|~m1_subset_1(esk5_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_finset_1(esk5_2(esk1_0,esk2_0))|~v5_finset_1(esk2_0))).#trainneg
cnf(i_0_627, plain, (r2_hidden(esk3_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|r1_tarski(k1_zfmisc_1(u1_struct_0(X1)),X3)|esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X3)!=k1_xboole_0|~v2_compts_1(esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_628, plain, (v2_compts_1(esk5_2(X1,X2),X1)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|r1_tarski(k1_zfmisc_1(u1_struct_0(X1)),X3)|esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X3)!=k1_xboole_0|~v2_compts_1(esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_629, plain, (v2_compts_1(esk6_2(X1,X2),X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|r1_tarski(k1_zfmisc_1(u1_struct_0(X1)),X3)|esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X3)!=k1_xboole_0|~v2_compts_1(esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_12594, negated_conjecture, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,esk4_2(esk1_0,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_12474, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X1,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_12580, negated_conjecture, (v2_compts_1(esk4_2(esk1_0,esk2_0),X1)|v1_xboole_0(u1_struct_0(X1))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_12633, negated_conjecture, (v1_xboole_0(X1)|r2_hidden(esk7_2(esk4_2(esk1_0,esk2_0),X2),X1)|r1_tarski(esk4_2(esk1_0,esk2_0),X2)|~m1_subset_1(esk2_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_12641, plain, (r2_hidden(X1,X2)|esk6_2(X4,X2)=k3_tarski(X2)|~r2_hidden(X1,k2_xboole_0(X3,esk10_1(esk4_2(X4,X2))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4))))|~m1_subset_1(X3,k1_zfmisc_1(esk4_2(X4,X2)))|~v1_finset_1(X2)|~l1_pre_topc(X4)|~v2_pre_topc(X4))).#trainneg
cnf(i_0_12682, negated_conjecture, (m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|esk6_2(X1,X2)=k3_tarski(X2)|~m1_subset_1(X2,esk2_0)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_12584, negated_conjecture, (v1_xboole_0(X1)|m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(esk4_2(esk1_0,esk2_0)))|~m1_subset_1(esk2_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_12559, negated_conjecture, (m1_subset_1(X1,X2)|~r2_hidden(X1,esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(X2))))).#trainneg
cnf(i_0_12579, negated_conjecture, (v2_compts_1(esk4_2(esk1_0,esk2_0),X1)|v1_xboole_0(u1_struct_0(X1))|~v1_xboole_0(esk4_2(esk1_0,esk2_0))|~l1_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_630, plain, (r1_tarski(esk4_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|r1_tarski(k1_zfmisc_1(u1_struct_0(X1)),X3)|esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X3)!=k1_xboole_0|~v2_compts_1(esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_631, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|r1_tarski(k1_zfmisc_1(u1_struct_0(X1)),X3)|esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X3)!=k1_xboole_0|~v2_compts_1(esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_632, plain, (m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|r1_tarski(k1_zfmisc_1(u1_struct_0(X1)),X3)|esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X3)!=k1_xboole_0|~v2_compts_1(esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X3),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_10232, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(X3,X2)|v1_xboole_0(esk4_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,esk4_2(X1,X2)))).#trainneg
cnf(i_0_8951, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(X2)))=k3_tarski(X1)|~m1_subset_1(k2_xboole_0(X2,k2_xboole_0(k1_zfmisc_1(k1_xboole_0),esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_10588, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(esk4_2(X1,X2),X3)|r2_hidden(k1_xboole_0,X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))).#trainneg
cnf(i_0_12605, negated_conjecture, (k3_tarski(esk5_2(esk1_0,esk2_0))=k1_xboole_0|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_12806, negated_conjecture, (v2_compts_1(esk7_2(esk4_2(esk1_0,esk2_0),X1),esk1_0)|r1_tarski(esk4_2(esk1_0,esk2_0),X1)|~m1_subset_1(esk7_2(esk4_2(esk1_0,esk2_0),X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1482, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_tarski(k3_tarski(k3_tarski(X3))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X2))))))).#trainneg
cnf(i_0_633, plain, (v2_compts_1(esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X2),X1)|r1_tarski(k1_zfmisc_1(u1_struct_0(X1)),X2)|~v1_xboole_0(esk7_2(k1_zfmisc_1(u1_struct_0(X1)),X2))|~l1_pre_topc(X1))).#trainneg
cnf(i_0_13010, negated_conjecture, (v3_ordinal1(k1_xboole_0)|~v3_ordinal1(esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_636, plain, (v1_finset_1(esk7_2(k1_zfmisc_1(X1),X2))|r1_tarski(k1_zfmisc_1(X1),X2)|~v1_finset_1(X1))).#trainneg
cnf(i_0_637, plain, (v1_xboole_0(esk7_2(k1_zfmisc_1(X1),X2))|r1_tarski(k1_zfmisc_1(X1),X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_12856, negated_conjecture, (r1_tarski(esk4_2(esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_12779, negated_conjecture, (m1_subset_1(esk7_2(esk4_2(esk1_0,esk2_0),X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_tarski(esk4_2(esk1_0,esk2_0),X1))).#trainneg
cnf(i_0_1809, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_tarski(k3_tarski(k3_tarski(X3))))|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X2))))))).#trainneg
cnf(i_0_12486, negated_conjecture, (v3_membered(esk5_2(esk1_0,esk2_0))|~v3_membered(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_638, plain, (v3_membered(esk7_2(k1_zfmisc_1(X1),X2))|r1_tarski(k1_zfmisc_1(X1),X2)|~v3_membered(X1))).#trainneg
cnf(i_0_639, plain, (m1_subset_1(X1,X2)|r1_tarski(k1_zfmisc_1(X2),X3)|~r2_hidden(X1,esk7_2(k1_zfmisc_1(X2),X3)))).#trainneg
cnf(i_0_640, plain, (r2_hidden(X1,X2)|r1_tarski(k1_zfmisc_1(X2),X3)|~r2_hidden(X1,esk7_2(k1_zfmisc_1(X2),X3)))).#trainneg
cnf(i_0_10604, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(esk3_2(X3,esk4_2(X4,esk4_2(X1,X2))),X2)|esk6_2(X3,esk4_2(X4,esk4_2(X1,X2)))=k3_tarski(esk4_2(X4,esk4_2(X1,X2)))|esk6_2(X4,esk4_2(X1,X2))=k3_tarski(esk4_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk4_2(X4,esk4_2(X1,X2)),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4))))|~v1_finset_1(esk4_2(X1,X2))|~l1_pre_topc(X4)|~l1_pre_topc(X3)|~v2_pre_topc(X4)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_7615, plain, (v2_compts_1(k3_tarski(X3),X4)|esk6_2(X1,X2)=k3_tarski(X2)|~m1_subset_1(k2_xboole_0(esk5_2(X1,X2),X3),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X4))))|~v1_finset_1(k3_tarski(X3))|~l1_pre_topc(X4)|~v2_pre_topc(X4)|~m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_10195, negated_conjecture, (k3_tarski(k2_xboole_0(esk4_2(esk1_0,esk2_0),k3_tarski(X1)))=esk5_2(esk1_0,esk2_0)|~m1_subset_1(k2_xboole_0(k1_zfmisc_1(k1_xboole_0),X1),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_13040, negated_conjecture, (r1_tarski(esk4_2(esk1_0,esk2_0),X1)|v2_compts_1(esk7_2(esk4_2(esk1_0,esk2_0),X1),esk1_0))).#trainneg
cnf(i_0_13095, negated_conjecture, (r2_hidden(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_hidden(X1,esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_10712, plain, (esk6_2(X1,X2)=k3_tarski(X2)|r1_tarski(k3_tarski(X3),X2)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(esk4_2(X1,X2)))))).#trainneg
cnf(i_0_9324, plain, (v1_xboole_0(u1_struct_0(X1))|v2_compts_1(esk5_2(X1,X2),X1)|v2_compts_1(esk6_2(X1,X2),X1)|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(k1_xboole_0,u1_struct_0(X1))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_641, plain, (v1_finset_1(esk7_2(X1,X2))|r1_tarski(X1,X2)|~v5_finset_1(X1))).#trainneg
cnf(i_0_12472, negated_conjecture, (v3_membered(esk4_2(esk1_0,esk2_0))|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_642, plain, (v1_xreal_0(esk7_2(X1,X2))|r1_tarski(X1,X2)|~v3_membered(X1))).#trainneg
cnf(i_0_643, plain, (v1_xxreal_0(esk7_2(X1,X2))|r1_tarski(X1,X2)|~v2_membered(X1))).#trainneg
cnf(i_0_12918, negated_conjecture, (v2_compts_1(esk4_2(X1,X2),esk1_0)|esk6_2(X1,X2)=k3_tarski(X2)|~v1_finset_1(esk4_2(X1,X2))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,esk2_0)|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_12492, negated_conjecture, (v2_compts_1(esk6_2(esk1_0,esk4_2(esk1_0,esk2_0)),esk1_0)|r2_hidden(esk3_2(esk1_0,esk4_2(esk1_0,esk2_0)),esk4_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_10051, plain, (r2_hidden(esk3_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_7966, plain, (r2_hidden(esk7_2(k4_subset_1(X1,X2,X3),X4),X1)|r1_tarski(k4_subset_1(X1,X2,X3),X4)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_7896, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(X2)))=k3_tarski(X1)|esk6_2(X3,X4)=k3_tarski(X4)|~m1_subset_1(k2_xboole_0(X2,k2_xboole_0(k1_zfmisc_1(k1_xboole_0),esk5_2(X3,X4))),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(esk4_2(X3,X4),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~v1_finset_1(X4)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_11835, negated_conjecture, (r1_tarski(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk7_2(X1,k1_zfmisc_1(u1_struct_0(esk1_0))),esk2_0))).#trainneg
cnf(i_0_13074, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_tarski(k3_tarski(k3_tarski(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X2))))))))).#trainneg
cnf(i_0_13171, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_tarski(k3_tarski(k3_tarski(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X2))))))))).#trainneg
cnf(i_0_12487, negated_conjecture, (v1_xreal_0(esk5_2(esk1_0,esk2_0))|~v3_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_644, plain, (v7_ordinal1(esk7_2(X1,X2))|r1_tarski(X1,X2)|~v6_membered(X1))).#trainneg
cnf(i_0_12488, negated_conjecture, (v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v2_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_12489, negated_conjecture, (v7_ordinal1(esk5_2(esk1_0,esk2_0))|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_11237, plain, (v1_finset_1(k5_setfam_1(k5_setfam_1(X1,X2),X3))|v1_xboole_0(X4)|~v5_finset_1(k1_zfmisc_1(X4))|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X1,X2))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_645, plain, (v3_ordinal1(esk7_2(X1,X2))|r1_tarski(X1,X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_660, plain, (v1_finset_1(k2_xboole_0(X1,X2))|~v5_finset_1(k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_13578, plain, (r1_tarski(k4_subset_1(X1,X2,X3),X1)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_13605, negated_conjecture, (r1_tarski(k2_xboole_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_13577, negated_conjecture, (r1_tarski(k4_subset_1(esk2_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_13043, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_tarski(k3_tarski(k3_tarski(esk10_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X2)))))))))).#trainneg
cnf(i_0_12490, negated_conjecture, (v3_ordinal1(esk5_2(esk1_0,esk2_0))|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_661, plain, (v1_xreal_0(k2_xboole_0(X1,X2))|~v3_membered(k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_11238, plain, (v1_xreal_0(k5_setfam_1(k5_setfam_1(X1,X2),X3))|v1_xboole_0(X4)|~v3_membered(k1_zfmisc_1(X4))|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X1,X2))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_11239, plain, (v1_xxreal_0(k5_setfam_1(k5_setfam_1(X1,X2),X3))|v1_xboole_0(X4)|~v2_membered(k1_zfmisc_1(X4))|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X1,X2))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_662, plain, (v1_xxreal_0(k2_xboole_0(X1,X2))|~v2_membered(k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_663, plain, (v7_ordinal1(k2_xboole_0(X1,X2))|~v6_membered(k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_13526, negated_conjecture, (v2_compts_1(esk7_2(k4_subset_1(esk2_0,X1,X2),X3),esk1_0)|r1_tarski(k4_subset_1(esk2_0,X1,X2),X3)|~m1_subset_1(esk7_2(k4_subset_1(esk2_0,X1,X2),X3),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_12601, negated_conjecture, (r1_tarski(esk4_2(esk1_0,esk2_0),X1)|m1_subset_1(esk7_2(esk4_2(esk1_0,esk2_0),X1),X2)|v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(esk2_0,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_13140, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_tarski(k3_tarski(k3_tarski(esk10_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X2)))))))))).#trainneg
cnf(i_0_13646, negated_conjecture, (v1_xreal_0(esk5_2(esk1_0,esk2_0))|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_10734, negated_conjecture, (X1!=k1_xboole_0|~v2_compts_1(X1,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_xboole_0(esk5_2(esk1_0,esk2_0))|~m1_subset_1(esk3_2(esk1_0,esk2_0),k1_zfmisc_1(esk5_2(esk1_0,esk2_0))))).#trainneg
cnf(i_0_664, plain, (v3_ordinal1(k2_xboole_0(X1,X2))|~v3_ordinal1(k1_zfmisc_1(X3))|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_13647, negated_conjecture, (v1_xxreal_0(esk5_2(esk1_0,esk2_0))|~v6_membered(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_667, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk5_2(X1,X2),X1)|k2_xboole_0(X3,X4)!=k1_xboole_0|~v2_compts_1(k2_xboole_0(X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_670, plain, (esk6_2(X1,X2)=k3_tarski(X2)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k2_xboole_0(X3,X4)!=k1_xboole_0|~v2_compts_1(k2_xboole_0(X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_13044, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k3_tarski(k3_tarski(k3_tarski(esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X2))))))))))).#trainneg
cnf(i_0_12458, negated_conjecture, (v1_xreal_0(esk4_2(esk1_0,esk2_0))|~v3_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_672, plain, (v2_compts_1(esk6_2(X1,X2),X1)|v2_compts_1(esk5_2(X1,X2),X1)|k2_xboole_0(X3,X4)!=k1_xboole_0|~v2_compts_1(k2_xboole_0(X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_9187, plain, (k3_tarski(X1)=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_676, plain, (v2_compts_1(esk5_2(X1,X2),X1)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k2_xboole_0(X3,X4)!=k1_xboole_0|~v2_compts_1(k2_xboole_0(X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_677, plain, (v2_compts_1(esk6_2(X1,X2),X1)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k2_xboole_0(X3,X4)!=k1_xboole_0|~v2_compts_1(k2_xboole_0(X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_14355, plain, (k3_tarski(k2_xboole_0(X1,X2))=k3_tarski(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_14346, plain, (k3_tarski(k2_xboole_0(X1,X2))=k3_tarski(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_14209, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=esk3_2(X1,X2)|X4!=k1_xboole_0|~v2_compts_1(X4,X1)|~v2_compts_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk4_2(X1,X2),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_12514, negated_conjecture, (esk6_2(X1,X2)=k3_tarski(X2)|r2_hidden(X3,X2)|v1_xboole_0(esk4_2(X1,X2))|~r2_hidden(X3,esk4_2(esk1_0,esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(esk2_0,k1_zfmisc_1(esk4_2(X1,X2))))).#trainneg
cnf(i_0_10606, negated_conjecture, (v2_compts_1(esk3_2(X1,esk4_2(X2,esk2_0)),esk1_0)|esk6_2(X1,esk4_2(X2,esk2_0))=k3_tarski(esk4_2(X2,esk2_0))|esk6_2(X2,esk2_0)=k3_tarski(esk2_0)|~m1_subset_1(esk3_2(X1,esk4_2(X2,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk4_2(X2,esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(esk2_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X2))))|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_13141, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k3_tarski(k3_tarski(k3_tarski(esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(X2))))))))))).#trainneg
cnf(i_0_14205, negated_conjecture, (esk6_2(esk1_0,X1)=k3_tarski(X1)|esk9_1(esk2_0)!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))|~v1_finset_1(X1)|~m1_subset_1(esk4_2(esk1_0,X1),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_678, plain, (r1_tarski(esk4_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k2_xboole_0(X3,X4)!=k1_xboole_0|~v2_compts_1(k2_xboole_0(X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_14345, plain, (v3_ordinal1(k1_xboole_0)|~v3_ordinal1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_11240, plain, (v7_ordinal1(k5_setfam_1(k5_setfam_1(X1,X2),X3))|v1_xboole_0(X4)|~v6_membered(k1_zfmisc_1(X4))|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X1,X2))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_679, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k2_xboole_0(X3,X4)!=k1_xboole_0|~v2_compts_1(k2_xboole_0(X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_680, plain, (m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|k2_xboole_0(X3,X4)!=k1_xboole_0|~v2_compts_1(k2_xboole_0(X3,X4),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~m1_subset_1(X4,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_12459, negated_conjecture, (v1_xxreal_0(esk4_2(esk1_0,esk2_0))|~v2_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_11241, plain, (v3_ordinal1(k5_setfam_1(k5_setfam_1(X1,X2),X3))|v1_xboole_0(X4)|~v3_ordinal1(k1_zfmisc_1(X4))|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X1,X2))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1))))).#trainneg
cnf(i_0_13826, plain, (v1_xxreal_0(k5_setfam_1(k5_setfam_1(X1,X2),X3))|v1_xboole_0(X4)|~v1_xboole_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k5_setfam_1(X1,X2))))|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(X1)))|~v3_membered(k1_zfmisc_1(X4)))).#trainneg
cnf(i_0_686, plain, (v3_membered(k2_xboole_0(X1,X2))|~v3_membered(X3)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_14449, plain, (k2_xboole_0(k3_tarski(X2),X1)=X1|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_14556, plain, (k3_tarski(X2)=k3_tarski(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_14423, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(X2)))=k3_tarski(X1)|~m1_subset_1(k2_xboole_0(X2,X3),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(X3,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_14540, plain, (k3_tarski(k3_tarski(X2))=k3_tarski(X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_12460, negated_conjecture, (v7_ordinal1(esk4_2(esk1_0,esk2_0))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_695, plain, (r2_hidden(esk7_2(esk10_1(X1),X2),X1)|r1_tarski(esk10_1(X1),X2))).#trainneg
cnf(i_0_696, plain, (r2_hidden(X1,X2)|v1_xboole_0(esk10_1(X2))|~m1_subset_1(X1,esk10_1(X2)))).#trainneg
cnf(i_0_698, plain, (esk6_2(X1,X2)=k3_tarski(X2)|X3!=k1_xboole_0|~v2_compts_1(X3,X1)|~v2_compts_1(k2_xboole_0(esk3_2(X1,X2),k3_tarski(esk4_2(X1,X2))),X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k2_xboole_0(esk3_2(X1,X2),k3_tarski(esk4_2(X1,X2))),k1_zfmisc_1(u1_struct_0(X1)))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_14780, plain, (k3_tarski(k2_xboole_0(X2,k3_tarski(X1)))=k3_tarski(X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_12461, negated_conjecture, (v3_ordinal1(esk4_2(esk1_0,esk2_0))|~v3_ordinal1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_14651, plain, (v3_ordinal1(k1_xboole_0)|~m1_subset_1(k3_tarski(X1),k1_zfmisc_1(k1_xboole_0))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_719, plain, (esk6_2(X1,X2)=k3_tarski(X2)|v2_compts_1(esk5_2(X1,X2),X1)|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_722, plain, (esk6_2(X1,X2)=k3_tarski(X2)|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_15020, plain, (v2_compts_1(k3_tarski(X4),X3)|~v2_compts_1(k3_tarski(X1),X3)|~v2_compts_1(k3_tarski(X2),X3)|~m1_subset_1(k3_tarski(X1),k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(k3_tarski(X2),k1_zfmisc_1(u1_struct_0(X3)))|~l1_pre_topc(X3)|~m1_subset_1(X4,k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(k2_xboole_0(X1,X2),k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_14504, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(X3)))=k3_tarski(X1)|~m1_subset_1(k2_xboole_0(X2,X3),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(X2,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_8825, negated_conjecture, (v2_compts_1(esk5_2(esk1_0,esk2_0),X1)|~m1_subset_1(esk4_2(esk1_0,esk2_0),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(esk5_2(esk1_0,esk2_0))|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_7438, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(k2_xboole_0(X2,X3))))=k3_tarski(X1)|~m1_subset_1(k2_xboole_0(k2_xboole_0(esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))),X2),X3),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_13739, negated_conjecture, (m1_subset_1(k4_subset_1(esk2_0,X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_8740, plain, (k3_tarski(k2_xboole_0(X1,k3_tarski(k2_xboole_0(X2,X3))))=k3_tarski(X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k2_xboole_0(k1_zfmisc_1(k1_xboole_0),k2_xboole_0(esk9_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))),X2)),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
cnf(i_0_13950, negated_conjecture, (r1_tarski(k2_xboole_0(X1,X2),X3)|v2_compts_1(esk7_2(k2_xboole_0(X1,X2),X3),esk1_0)|~m1_subset_1(esk7_2(k2_xboole_0(X1,X2),X3),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,k1_zfmisc_1(esk2_0))|~m1_subset_1(X1,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_15629, negated_conjecture, (v1_xreal_0(esk4_2(esk1_0,esk2_0))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_14646, negated_conjecture, (v3_ordinal1(k1_xboole_0)|~m1_subset_1(esk3_2(esk1_0,esk2_0),k1_zfmisc_1(k1_xboole_0))|~v3_ordinal1(esk2_0))).#trainneg
cnf(i_0_728, plain, (v2_compts_1(esk5_2(X1,X2),X1)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_12903, negated_conjecture, (m1_subset_1(k2_xboole_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|esk6_2(X3,X4)=k3_tarski(X4)|~m1_subset_1(X2,k1_zfmisc_1(esk4_2(X3,X4)))|~m1_subset_1(X1,k1_zfmisc_1(esk4_2(X3,X4)))|~m1_subset_1(X4,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X3))))|~m1_subset_1(X4,esk2_0)|~v1_finset_1(X4)|~l1_pre_topc(X3)|~v2_pre_topc(X3))).#trainneg
cnf(i_0_16055, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X1,k1_zfmisc_1(k4_subset_1(esk2_0,X2,X3)))|~m1_subset_1(X3,k1_zfmisc_1(esk2_0))|~m1_subset_1(X2,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_16018, negated_conjecture, (v2_compts_1(esk5_2(esk1_0,esk2_0),esk1_0)|~v1_finset_1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_15630, negated_conjecture, (v1_xxreal_0(esk4_2(esk1_0,esk2_0))|~v6_membered(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))))).#trainneg
cnf(i_0_16188, negated_conjecture, (X1!=k1_xboole_0|~v2_compts_1(X1,esk1_0)|~m1_subset_1(esk3_2(esk1_0,esk2_0),k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_finset_1(esk5_2(esk1_0,esk2_0)))).#trainneg
cnf(i_0_730, plain, (r1_tarski(esk4_2(X1,X2),X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_15893, plain, (v3_ordinal1(k1_xboole_0)|~m1_subset_1(k3_tarski(k3_tarski(X1)),k1_zfmisc_1(k1_xboole_0))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_731, plain, (k3_tarski(esk4_2(X1,X2))=esk5_2(X1,X2)|m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_732, plain, (m1_subset_1(esk6_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|m1_subset_1(esk5_2(X1,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v2_compts_1(k1_xboole_0,X1)|~m1_subset_1(X2,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~v1_finset_1(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1))).#trainneg
cnf(i_0_16075, negated_conjecture, (m1_subset_1(k2_xboole_0(X1,X2),k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0))))|~m1_subset_1(X2,k1_zfmisc_1(k4_subset_1(esk2_0,X3,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k4_subset_1(esk2_0,X3,X4)))|~m1_subset_1(X4,k1_zfmisc_1(esk2_0))|~m1_subset_1(X3,k1_zfmisc_1(esk2_0)))).#trainneg
cnf(i_0_8175, plain, (v1_xboole_0(X1)|m1_subset_1(k3_tarski(X2),k1_zfmisc_1(X1))|~m1_subset_1(k2_xboole_0(X2,k3_tarski(X3)),k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0))))|~m1_subset_1(k1_xboole_0,X1)|~m1_subset_1(X3,k1_zfmisc_1(k1_zfmisc_1(k1_zfmisc_1(k1_xboole_0)))))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 85
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 122
# Removed in clause preprocessing      : 12
# Initial clauses in saturation        : 110
# Processed clauses                    : 2262
# ...of these trivial                  : 38
# ...subsumed                          : 1084
# ...remaining for further processing  : 1140
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 255
# Backward-rewritten                   : 27
# Generated clauses                    : 16438
# ...of the previous two non-trivial   : 15026
# Contextual simplify-reflections      : 76
# Paramodulations                      : 16428
# Factorizations                       : 0
# Equation resolutions                 : 8
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 856
#    Positive orientable unit clauses  : 75
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 19
#    Non-unit-clauses                  : 761
# Current number of unprocessed clauses: 12527
# ...number of literals in the above   : 79583
# Current number of archived formulas  : 0
# Current number of archived clauses   : 1425
# Proof object given clauses           : 114
# Proof search given clauses           : 1140
# Clause-clause subsumption calls (NU) : 440749
# Rec. Clause-clause subsumption calls : 63525
# Non-unit clause-clause subsumptions  : 1168
# Unit Clause-clause subsumption calls : 1912
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 85
# BW rewrite match successes           : 19
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 397873

# -------------------------------------------------
# User time                : 39.156 s
# System time              : 1.572 s
# Total time               : 40.728 s
# Maximum resident set size: 355240 pages
39.15user 1.65system 0:41.57elapsed 98%CPU (0avgtext+0avgdata 355240maxresident)k
0inputs+576outputs (0major+407479minor)pagefaults 0swaps
