# Preprocessing time       : 0.025 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t21_jordan9, conjecture, ![X1]:(m1_subset_1(X1,k5_numbers)=>![X2]:(m1_subset_1(X2,k5_numbers)=>![X3]:(((((((~(v3_relat_1(X3))&v1_matrix_1(X3))&v2_goboard1(X3))&v3_goboard1(X3))&v4_goboard1(X3))&v5_goboard1(X3))&m2_finseq_1(X3,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))=>![X4]:(m1_subset_1(X4,u1_struct_0(k15_euclid(np__2)))=>((((((r1_xxreal_0(np__1,X1)&r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(X3)))&r1_xxreal_0(np__1,X2))&r1_xxreal_0(k2_nat_1(X2,np__1),k1_matrix_1(X3)))&r2_hidden(X4,k16_matrix_1(X3)))&r2_hidden(X4,k3_goboard5(X3,X1,X2)))=>r1_sppol_1(np__2,X4,k3_goboard5(X3,X1,X2))))))), file('/scratch/mptp/minsubpreds3/t21_jordan9', t21_jordan9)).
fof(redefinition_m2_finseq_1, axiom, ![X1, X2]:(m2_finseq_1(X2,X1)<=>m1_finseq_1(X2,X1)), file('/scratch/mptp/minsubpreds3/t21_jordan9', redefinition_m2_finseq_1)).
fof(cc8_ordinal1, axiom, ![X1]:(m1_subset_1(X1,k4_ordinal1)=>v7_ordinal1(X1)), file('/scratch/mptp/minsubpreds3/t21_jordan9', cc8_ordinal1)).
fof(dt_k3_goboard5, axiom, ![X1, X2, X3]:((((v1_matrix_1(X1)&m1_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))&v7_ordinal1(X2))&v7_ordinal1(X3))=>m1_subset_1(k3_goboard5(X1,X2,X3),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), file('/scratch/mptp/minsubpreds3/t21_jordan9', dt_k3_goboard5)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/minsubpreds3/t21_jordan9', redefinition_k5_numbers)).
fof(d1_sppol_1, axiom, ![X1]:(m2_subset_1(X1,k1_numbers,k5_numbers)=>![X2]:(m1_subset_1(X2,u1_struct_0(k15_euclid(X1)))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X1))))=>(r1_sppol_1(X1,X2,X3)<=>(r2_hidden(X2,X3)&![X4]:(m1_subset_1(X4,u1_struct_0(k15_euclid(X1)))=>![X5]:(m1_subset_1(X5,u1_struct_0(k15_euclid(X1)))=>~((((r2_hidden(X2,k1_rltopsp1(k15_euclid(X1),X4,X5))&r1_tarski(k1_rltopsp1(k15_euclid(X1),X4,X5),X3))&X2!=X4)&X2!=X5))))))))), file('/scratch/mptp/minsubpreds3/t21_jordan9', d1_sppol_1)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/scratch/mptp/minsubpreds3/t21_jordan9', t4_subset)).
fof(dt_k15_euclid, axiom, ![X1]:(v7_ordinal1(X1)=>(v5_rltopsp1(k15_euclid(X1))&l1_rltopsp1(k15_euclid(X1)))), file('/scratch/mptp/minsubpreds3/t21_jordan9', dt_k15_euclid)).
fof(spc2_numerals, axiom, (((v2_xxreal_0(np__2)&m2_subset_1(np__2,k1_numbers,k5_numbers))&m1_subset_1(np__2,k5_numbers))&m1_subset_1(np__2,k1_numbers)), file('/scratch/mptp/minsubpreds3/t21_jordan9', spc2_numerals)).
fof(dt_l1_rltopsp1, axiom, ![X1]:(l1_rltopsp1(X1)=>(l1_rlvect_1(X1)&l1_pre_topc(X1))), file('/scratch/mptp/minsubpreds3/t21_jordan9', dt_l1_rltopsp1)).
fof(fc6_euclid, axiom, ![X1]:(v7_ordinal1(X1)=>(((((((((v2_pre_topc(k15_euclid(X1))&v13_algstr_0(k15_euclid(X1)))&v2_rlvect_1(k15_euclid(X1)))&v3_rlvect_1(k15_euclid(X1)))&v4_rlvect_1(k15_euclid(X1)))&v5_rlvect_1(k15_euclid(X1)))&v6_rlvect_1(k15_euclid(X1)))&v7_rlvect_1(k15_euclid(X1)))&v8_rlvect_1(k15_euclid(X1)))&v5_rltopsp1(k15_euclid(X1)))), file('/scratch/mptp/minsubpreds3/t21_jordan9', fc6_euclid)).
fof(t68_rltopsp1, axiom, ![X1]:((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v5_rlvect_1(X1))&v6_rlvect_1(X1))&v7_rlvect_1(X1))&v8_rlvect_1(X1))&l1_rlvect_1(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>r2_hidden(X2,k1_rltopsp1(X1,X2,X3))))), file('/scratch/mptp/minsubpreds3/t21_jordan9', t68_rltopsp1)).
fof(commutativity_k1_rltopsp1, axiom, ![X1, X2, X3]:((((((((((((~(v2_struct_0(X1))&v13_algstr_0(X1))&v2_rlvect_1(X1))&v3_rlvect_1(X1))&v4_rlvect_1(X1))&v5_rlvect_1(X1))&v6_rlvect_1(X1))&v7_rlvect_1(X1))&v8_rlvect_1(X1))&l1_rlvect_1(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>k1_rltopsp1(X1,X2,X3)=k1_rltopsp1(X1,X3,X2)), file('/scratch/mptp/minsubpreds3/t21_jordan9', commutativity_k1_rltopsp1)).
fof(dt_k17_euclid, axiom, ![X1]:(m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))=>m1_subset_1(k17_euclid(X1),k1_numbers)), file('/scratch/mptp/minsubpreds3/t21_jordan9', dt_k17_euclid)).
fof(t17_jordan9, axiom, ![X1]:(((((((~(v3_relat_1(X1))&v1_matrix_1(X1))&v2_goboard1(X1))&v3_goboard1(X1))&v4_goboard1(X1))&v5_goboard1(X1))&m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))=>![X2]:(m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))=>![X3]:(m1_subset_1(X3,k5_numbers)=>![X4]:(m1_subset_1(X4,k5_numbers)=>((((r1_xxreal_0(np__1,X3)&r1_xxreal_0(k2_nat_1(X3,np__1),k3_finseq_1(X1)))&r1_xxreal_0(np__1,X4))&r1_xxreal_0(k2_nat_1(X4,np__1),k1_matrix_1(X1)))=>(r2_hidden(X2,k3_goboard5(X1,X3,X4))<=>(((r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X3,X4)),k17_euclid(X2))&r1_xxreal_0(k17_euclid(X2),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X3,np__1),X4))))&r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X3,X4)),k18_euclid(X2)))&r1_xxreal_0(k18_euclid(X2),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X3,k2_nat_1(X4,np__1))))))))))), file('/scratch/mptp/minsubpreds3/t21_jordan9', t17_jordan9)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/scratch/mptp/minsubpreds3/t21_jordan9', t3_subset)).
fof(t19_jordan9, axiom, ![X1]:(m1_subset_1(X1,k5_numbers)=>![X2]:(m1_subset_1(X2,k5_numbers)=>![X3]:(((((((~(v3_relat_1(X3))&v1_matrix_1(X3))&v2_goboard1(X3))&v3_goboard1(X3))&v4_goboard1(X3))&v5_goboard1(X3))&m2_finseq_1(X3,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))=>![X4]:(m1_subset_1(X4,u1_struct_0(k15_euclid(np__2)))=>~((((((((((r1_xxreal_0(np__1,X1)&r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(X3)))&r1_xxreal_0(np__1,X2))&r1_xxreal_0(k2_nat_1(X2,np__1),k1_matrix_1(X3)))&r2_hidden(X4,k16_matrix_1(X3)))&r2_hidden(X4,k3_goboard5(X3,X1,X2)))&X4!=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,X1,X2))&X4!=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,X1,k2_nat_1(X2,np__1)))&X4!=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,k2_nat_1(X1,np__1),k2_nat_1(X2,np__1)))&X4!=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,k2_nat_1(X1,np__1),X2))))))), file('/scratch/mptp/minsubpreds3/t21_jordan9', t19_jordan9)).
fof(cc1_xreal_0, axiom, ![X1]:(m1_subset_1(X1,k1_numbers)=>v1_xreal_0(X1)), file('/scratch/mptp/minsubpreds3/t21_jordan9', cc1_xreal_0)).
fof(cc4_xreal_0, axiom, ![X1]:(v1_xreal_0(X1)=>v1_xxreal_0(X1)), file('/scratch/mptp/minsubpreds3/t21_jordan9', cc4_xreal_0)).
fof(t7_boole, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2))), file('/scratch/mptp/minsubpreds3/t21_jordan9', t7_boole)).
fof(t16_jordan9, axiom, ![X1]:(m1_subset_1(X1,k5_numbers)=>![X2]:(m1_subset_1(X2,k5_numbers)=>![X3]:(((((((~(v3_relat_1(X3))&v1_matrix_1(X3))&v2_goboard1(X3))&v3_goboard1(X3))&v4_goboard1(X3))&v5_goboard1(X3))&m2_finseq_1(X3,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))=>((((r1_xxreal_0(np__1,X1)&r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(X3)))&r1_xxreal_0(np__1,X2))&r1_xxreal_0(k2_nat_1(X2,np__1),k1_matrix_1(X3)))=>(((k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,X1,X2))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,X1,k2_nat_1(X2,np__1)))&k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,X1,X2))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,k2_nat_1(X1,np__1),X2)))&k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,k2_nat_1(X1,np__1),k2_nat_1(X2,np__1)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,k2_nat_1(X1,np__1),X2)))&k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,k2_nat_1(X1,np__1),k2_nat_1(X2,np__1)))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,X1,k2_nat_1(X2,np__1)))))))), file('/scratch/mptp/minsubpreds3/t21_jordan9', t16_jordan9)).
fof(connectedness_r1_xxreal_0, axiom, ![X1, X2]:((v1_xxreal_0(X1)&v1_xxreal_0(X2))=>(r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1))), file('/scratch/mptp/minsubpreds3/t21_jordan9', connectedness_r1_xxreal_0)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('/scratch/mptp/minsubpreds3/t21_jordan9', t2_subset)).
fof(t3_topreal1, axiom, ![X1]:(m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))=>![X2]:(m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))=>((r1_xxreal_0(k17_euclid(X1),k17_euclid(X2))&r2_hidden(X3,k1_rltopsp1(k15_euclid(np__2),X1,X2)))=>(r1_xxreal_0(k17_euclid(X1),k17_euclid(X3))&r1_xxreal_0(k17_euclid(X3),k17_euclid(X2))))))), file('/scratch/mptp/minsubpreds3/t21_jordan9', t3_topreal1)).
fof(t1_xxreal_0, axiom, ![X1]:(v1_xxreal_0(X1)=>![X2]:(v1_xxreal_0(X2)=>((r1_xxreal_0(X1,X2)&r1_xxreal_0(X2,X1))=>X1=X2))), file('/scratch/mptp/minsubpreds3/t21_jordan9', t1_xxreal_0)).
fof(dt_k18_euclid, axiom, ![X1]:(m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))=>m1_subset_1(k18_euclid(X1),k1_numbers)), file('/scratch/mptp/minsubpreds3/t21_jordan9', dt_k18_euclid)).
fof(t18_sppol_1, axiom, ![X1]:(m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))=>![X2]:(m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))=>![X4]:(m1_subset_1(X4,u1_struct_0(k15_euclid(np__2)))=>(((r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),X2,X3))&r2_hidden(X4,k1_rltopsp1(k15_euclid(np__2),X2,X3)))&k17_euclid(X1)=k17_euclid(X4))=>(k18_euclid(X1)=k18_euclid(X4)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X2,X3)))))))), file('/scratch/mptp/minsubpreds3/t21_jordan9', t18_sppol_1)).
fof(t6_topreal3, axiom, ![X1]:(m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))=>![X2]:(m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))=>((k17_euclid(X1)=k17_euclid(X2)&k18_euclid(X1)=k18_euclid(X2))=>X1=X2))), file('/scratch/mptp/minsubpreds3/t21_jordan9', t6_topreal3)).
fof(t4_topreal1, axiom, ![X1]:(m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))=>![X2]:(m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))=>![X3]:(m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))=>((r1_xxreal_0(k18_euclid(X1),k18_euclid(X2))&r2_hidden(X3,k1_rltopsp1(k15_euclid(np__2),X1,X2)))=>(r1_xxreal_0(k18_euclid(X1),k18_euclid(X3))&r1_xxreal_0(k18_euclid(X3),k18_euclid(X2))))))), file('/scratch/mptp/minsubpreds3/t21_jordan9', t4_topreal1)).
fof(t16_sppol_1, axiom, ![X1]:(m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))=>![X2]:(m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))=>(k17_euclid(X1)=k17_euclid(X2)<=>v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,X2))))), file('/scratch/mptp/minsubpreds3/t21_jordan9', t16_sppol_1)).
fof(fc5_euclid, axiom, ![X1]:(v7_ordinal1(X1)=>(~(v2_struct_0(k15_euclid(X1)))&v5_rltopsp1(k15_euclid(X1)))), file('/scratch/mptp/minsubpreds3/t21_jordan9', fc5_euclid)).
fof(c_0_31, negated_conjecture, ~(![X1]:(m1_subset_1(X1,k5_numbers)=>![X2]:(m1_subset_1(X2,k5_numbers)=>![X3]:(((((((~(v3_relat_1(X3))&v1_matrix_1(X3))&v2_goboard1(X3))&v3_goboard1(X3))&v4_goboard1(X3))&v5_goboard1(X3))&m2_finseq_1(X3,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))=>![X4]:(m1_subset_1(X4,u1_struct_0(k15_euclid(np__2)))=>((((((r1_xxreal_0(np__1,X1)&r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(X3)))&r1_xxreal_0(np__1,X2))&r1_xxreal_0(k2_nat_1(X2,np__1),k1_matrix_1(X3)))&r2_hidden(X4,k16_matrix_1(X3)))&r2_hidden(X4,k3_goboard5(X3,X1,X2)))=>r1_sppol_1(np__2,X4,k3_goboard5(X3,X1,X2)))))))), inference(assume_negation,[status(cth)],[t21_jordan9])).
fof(c_0_32, plain, ![X35, X36]:((~m2_finseq_1(X36,X35)|m1_finseq_1(X36,X35))&(~m1_finseq_1(X36,X35)|m2_finseq_1(X36,X35))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_m2_finseq_1])])).
fof(c_0_33, negated_conjecture, (m1_subset_1(esk1_0,k5_numbers)&(m1_subset_1(esk2_0,k5_numbers)&(((((((~v3_relat_1(esk3_0)&v1_matrix_1(esk3_0))&v2_goboard1(esk3_0))&v3_goboard1(esk3_0))&v4_goboard1(esk3_0))&v5_goboard1(esk3_0))&m2_finseq_1(esk3_0,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))&(m1_subset_1(esk4_0,u1_struct_0(k15_euclid(np__2)))&((((((r1_xxreal_0(np__1,esk1_0)&r1_xxreal_0(k2_nat_1(esk1_0,np__1),k3_finseq_1(esk3_0)))&r1_xxreal_0(np__1,esk2_0))&r1_xxreal_0(k2_nat_1(esk2_0,np__1),k1_matrix_1(esk3_0)))&r2_hidden(esk4_0,k16_matrix_1(esk3_0)))&r2_hidden(esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))&~r1_sppol_1(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_31])])])])).
fof(c_0_34, plain, ![X12]:(~m1_subset_1(X12,k4_ordinal1)|v7_ordinal1(X12)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_ordinal1])])).
fof(c_0_35, plain, ![X28, X29, X30]:(~v1_matrix_1(X28)|~m1_finseq_1(X28,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))|~v7_ordinal1(X29)|~v7_ordinal1(X30)|m1_subset_1(k3_goboard5(X28,X29,X30),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k3_goboard5])])).
cnf(c_0_36, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_32])).
cnf(c_0_37, negated_conjecture, (m2_finseq_1(esk3_0,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_38, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_39, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
cnf(c_0_40, plain, (m1_subset_1(k3_goboard5(X1,X2,X3),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_matrix_1(X1)|~m1_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))|~v7_ordinal1(X2)|~v7_ordinal1(X3)), inference(split_conjunct,[status(thm)],[c_0_35])).
cnf(c_0_41, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_36])).
cnf(c_0_42, negated_conjecture, (m2_finseq_1(esk3_0,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[c_0_37])).
cnf(c_0_43, negated_conjecture, (v1_matrix_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_44, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(evalgc,[status(thm)],[c_0_38])).
cnf(c_0_45, plain, (k4_ordinal1=k5_numbers), inference(evalgc,[status(thm)],[c_0_39])).
cnf(c_0_46, negated_conjecture, (m1_subset_1(esk2_0,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_47, plain, ![X18, X19, X20, X21, X22]:(((r2_hidden(X19,X20)|~r1_sppol_1(X18,X19,X20)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(k15_euclid(X18))))|~m1_subset_1(X19,u1_struct_0(k15_euclid(X18)))|~m2_subset_1(X18,k1_numbers,k5_numbers))&(~m1_subset_1(X21,u1_struct_0(k15_euclid(X18)))|(~m1_subset_1(X22,u1_struct_0(k15_euclid(X18)))|(~r2_hidden(X19,k1_rltopsp1(k15_euclid(X18),X21,X22))|~r1_tarski(k1_rltopsp1(k15_euclid(X18),X21,X22),X20)|X19=X21|X19=X22))|~r1_sppol_1(X18,X19,X20)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(k15_euclid(X18))))|~m1_subset_1(X19,u1_struct_0(k15_euclid(X18)))|~m2_subset_1(X18,k1_numbers,k5_numbers)))&((m1_subset_1(esk5_3(X18,X19,X20),u1_struct_0(k15_euclid(X18)))|~r2_hidden(X19,X20)|r1_sppol_1(X18,X19,X20)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(k15_euclid(X18))))|~m1_subset_1(X19,u1_struct_0(k15_euclid(X18)))|~m2_subset_1(X18,k1_numbers,k5_numbers))&((m1_subset_1(esk6_3(X18,X19,X20),u1_struct_0(k15_euclid(X18)))|~r2_hidden(X19,X20)|r1_sppol_1(X18,X19,X20)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(k15_euclid(X18))))|~m1_subset_1(X19,u1_struct_0(k15_euclid(X18)))|~m2_subset_1(X18,k1_numbers,k5_numbers))&((((r2_hidden(X19,k1_rltopsp1(k15_euclid(X18),esk5_3(X18,X19,X20),esk6_3(X18,X19,X20)))|~r2_hidden(X19,X20)|r1_sppol_1(X18,X19,X20)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(k15_euclid(X18))))|~m1_subset_1(X19,u1_struct_0(k15_euclid(X18)))|~m2_subset_1(X18,k1_numbers,k5_numbers))&(r1_tarski(k1_rltopsp1(k15_euclid(X18),esk5_3(X18,X19,X20),esk6_3(X18,X19,X20)),X20)|~r2_hidden(X19,X20)|r1_sppol_1(X18,X19,X20)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(k15_euclid(X18))))|~m1_subset_1(X19,u1_struct_0(k15_euclid(X18)))|~m2_subset_1(X18,k1_numbers,k5_numbers)))&(X19!=esk5_3(X18,X19,X20)|~r2_hidden(X19,X20)|r1_sppol_1(X18,X19,X20)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(k15_euclid(X18))))|~m1_subset_1(X19,u1_struct_0(k15_euclid(X18)))|~m2_subset_1(X18,k1_numbers,k5_numbers)))&(X19!=esk6_3(X18,X19,X20)|~r2_hidden(X19,X20)|r1_sppol_1(X18,X19,X20)|~m1_subset_1(X20,k1_zfmisc_1(u1_struct_0(k15_euclid(X18))))|~m1_subset_1(X19,u1_struct_0(k15_euclid(X18)))|~m2_subset_1(X18,k1_numbers,k5_numbers)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_sppol_1])])])])])).
fof(c_0_48, plain, ![X65, X66, X67]:(~r2_hidden(X65,X66)|~m1_subset_1(X66,k1_zfmisc_1(X67))|m1_subset_1(X65,X67)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_49, plain, (m1_subset_1(k3_goboard5(X1,X2,X3),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_matrix_1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[c_0_40])).
cnf(c_0_50, negated_conjecture, (m1_finseq_1(esk3_0,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42])])).
cnf(c_0_51, negated_conjecture, (v1_matrix_1(esk3_0)), inference(evalgc,[status(thm)],[c_0_43])).
cnf(c_0_52, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_44, c_0_45])).
cnf(c_0_53, negated_conjecture, (m1_subset_1(esk2_0,k5_numbers)), inference(evalgc,[status(thm)],[c_0_46])).
cnf(c_0_54, negated_conjecture, (m1_subset_1(esk1_0,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_55, plain, (m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(k15_euclid(X1)))|r1_sppol_1(X1,X2,X3)|~r2_hidden(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X1))))|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_56, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_57, negated_conjecture, (m1_subset_1(k3_goboard5(esk3_0,X1,X2),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_50]), c_0_51])])])).
cnf(c_0_58, negated_conjecture, (v7_ordinal1(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53])])).
cnf(c_0_59, negated_conjecture, (m1_subset_1(esk1_0,k5_numbers)), inference(evalgc,[status(thm)],[c_0_54])).
fof(c_0_60, plain, ![X25]:((v5_rltopsp1(k15_euclid(X25))|~v7_ordinal1(X25))&(l1_rltopsp1(k15_euclid(X25))|~v7_ordinal1(X25))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k15_euclid])])])).
cnf(c_0_61, plain, (m1_subset_1(np__2,k5_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
cnf(c_0_62, plain, (r1_tarski(k1_rltopsp1(k15_euclid(X1),esk5_3(X1,X2,X3),esk6_3(X1,X2,X3)),X3)|r1_sppol_1(X1,X2,X3)|~r2_hidden(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X1))))|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_63, plain, (r1_sppol_1(X1,X2,X3)|m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(k15_euclid(X1)))|~r2_hidden(X2,X3)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X1))))), inference(evalgc,[status(thm)],[c_0_55])).
cnf(c_0_64, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_56])).
cnf(c_0_65, negated_conjecture, (m1_subset_1(k3_goboard5(esk3_0,X1,esk2_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_58])])).
cnf(c_0_66, negated_conjecture, (v7_ordinal1(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_59])])).
cnf(c_0_67, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
fof(c_0_68, plain, ![X31]:((l1_rlvect_1(X31)|~l1_rltopsp1(X31))&(l1_pre_topc(X31)|~l1_rltopsp1(X31))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_rltopsp1])])])).
cnf(c_0_69, plain, (l1_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_70, plain, (m1_subset_1(np__2,k5_numbers)), inference(evalgc,[status(thm)],[c_0_61])).
fof(c_0_71, plain, ![X33]:((((((((((v2_pre_topc(k15_euclid(X33))|~v7_ordinal1(X33))&(v13_algstr_0(k15_euclid(X33))|~v7_ordinal1(X33)))&(v2_rlvect_1(k15_euclid(X33))|~v7_ordinal1(X33)))&(v3_rlvect_1(k15_euclid(X33))|~v7_ordinal1(X33)))&(v4_rlvect_1(k15_euclid(X33))|~v7_ordinal1(X33)))&(v5_rlvect_1(k15_euclid(X33))|~v7_ordinal1(X33)))&(v6_rlvect_1(k15_euclid(X33))|~v7_ordinal1(X33)))&(v7_rlvect_1(k15_euclid(X33))|~v7_ordinal1(X33)))&(v8_rlvect_1(k15_euclid(X33))|~v7_ordinal1(X33)))&(v5_rltopsp1(k15_euclid(X33))|~v7_ordinal1(X33))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc6_euclid])])])).
cnf(c_0_72, plain, (m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(k15_euclid(X1)))|r1_sppol_1(X1,X2,X3)|~r2_hidden(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X1))))|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_73, plain, (r1_sppol_1(X1,X2,X3)|r1_tarski(k1_rltopsp1(k15_euclid(X1),esk5_3(X1,X2,X3),esk6_3(X1,X2,X3)),X3)|~r2_hidden(X2,X3)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X1))))), inference(evalgc,[status(thm)],[c_0_62])).
fof(c_0_74, plain, ![X74, X75, X76]:(v2_struct_0(X74)|~v13_algstr_0(X74)|~v2_rlvect_1(X74)|~v3_rlvect_1(X74)|~v4_rlvect_1(X74)|~v5_rlvect_1(X74)|~v6_rlvect_1(X74)|~v7_rlvect_1(X74)|~v8_rlvect_1(X74)|~l1_rlvect_1(X74)|(~m1_subset_1(X75,u1_struct_0(X74))|(~m1_subset_1(X76,u1_struct_0(X74))|r2_hidden(X75,k1_rltopsp1(X74,X75,X76))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t68_rltopsp1])])])])).
cnf(c_0_75, plain, (r1_sppol_1(X1,X2,X3)|m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(k15_euclid(X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r2_hidden(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X1))))), inference(csr,[status(thm)],[c_0_63, c_0_64])).
cnf(c_0_76, negated_conjecture, (m1_subset_1(k3_goboard5(esk3_0,esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66])])).
cnf(c_0_77, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers)), inference(evalgc,[status(thm)],[c_0_67])).
cnf(c_0_78, negated_conjecture, (r2_hidden(esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_79, negated_conjecture, (~r1_sppol_1(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_80, plain, (l1_rlvect_1(X1)|~l1_rltopsp1(X1)), inference(split_conjunct,[status(thm)],[c_0_68])).
cnf(c_0_81, plain, (l1_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_69])).
cnf(c_0_82, plain, (v7_ordinal1(np__2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_70])])).
cnf(c_0_83, plain, (v8_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_84, plain, (v7_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_85, plain, (v6_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_86, plain, (v5_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_87, plain, (v4_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_88, plain, (v3_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_89, plain, (v2_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_90, plain, (v13_algstr_0(k15_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_71])).
cnf(c_0_91, plain, (r1_sppol_1(X1,X2,X3)|m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(k15_euclid(X1)))|~r2_hidden(X2,X3)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X1))))), inference(evalgc,[status(thm)],[c_0_72])).
fof(c_0_92, plain, ![X13, X14, X15]:(v2_struct_0(X13)|~v13_algstr_0(X13)|~v2_rlvect_1(X13)|~v3_rlvect_1(X13)|~v4_rlvect_1(X13)|~v5_rlvect_1(X13)|~v6_rlvect_1(X13)|~v7_rlvect_1(X13)|~v8_rlvect_1(X13)|~l1_rlvect_1(X13)|~m1_subset_1(X14,u1_struct_0(X13))|~m1_subset_1(X15,u1_struct_0(X13))|k1_rltopsp1(X13,X14,X15)=k1_rltopsp1(X13,X15,X14)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[commutativity_k1_rltopsp1])])])).
fof(c_0_93, plain, ![X26]:(~m1_subset_1(X26,u1_struct_0(k15_euclid(np__2)))|m1_subset_1(k17_euclid(X26),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k17_euclid])])).
fof(c_0_94, plain, ![X42, X43, X44, X45]:(((((r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X42,X44,X45)),k17_euclid(X43))|~r2_hidden(X43,k3_goboard5(X42,X44,X45))|(~r1_xxreal_0(np__1,X44)|~r1_xxreal_0(k2_nat_1(X44,np__1),k3_finseq_1(X42))|~r1_xxreal_0(np__1,X45)|~r1_xxreal_0(k2_nat_1(X45,np__1),k1_matrix_1(X42)))|~m1_subset_1(X45,k5_numbers)|~m1_subset_1(X44,k5_numbers)|~m1_subset_1(X43,u1_struct_0(k15_euclid(np__2)))|(v3_relat_1(X42)|~v1_matrix_1(X42)|~v2_goboard1(X42)|~v3_goboard1(X42)|~v4_goboard1(X42)|~v5_goboard1(X42)|~m2_finseq_1(X42,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))))&(r1_xxreal_0(k17_euclid(X43),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X42,k2_nat_1(X44,np__1),X45)))|~r2_hidden(X43,k3_goboard5(X42,X44,X45))|(~r1_xxreal_0(np__1,X44)|~r1_xxreal_0(k2_nat_1(X44,np__1),k3_finseq_1(X42))|~r1_xxreal_0(np__1,X45)|~r1_xxreal_0(k2_nat_1(X45,np__1),k1_matrix_1(X42)))|~m1_subset_1(X45,k5_numbers)|~m1_subset_1(X44,k5_numbers)|~m1_subset_1(X43,u1_struct_0(k15_euclid(np__2)))|(v3_relat_1(X42)|~v1_matrix_1(X42)|~v2_goboard1(X42)|~v3_goboard1(X42)|~v4_goboard1(X42)|~v5_goboard1(X42)|~m2_finseq_1(X42,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))))&(r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X42,X44,X45)),k18_euclid(X43))|~r2_hidden(X43,k3_goboard5(X42,X44,X45))|(~r1_xxreal_0(np__1,X44)|~r1_xxreal_0(k2_nat_1(X44,np__1),k3_finseq_1(X42))|~r1_xxreal_0(np__1,X45)|~r1_xxreal_0(k2_nat_1(X45,np__1),k1_matrix_1(X42)))|~m1_subset_1(X45,k5_numbers)|~m1_subset_1(X44,k5_numbers)|~m1_subset_1(X43,u1_struct_0(k15_euclid(np__2)))|(v3_relat_1(X42)|~v1_matrix_1(X42)|~v2_goboard1(X42)|~v3_goboard1(X42)|~v4_goboard1(X42)|~v5_goboard1(X42)|~m2_finseq_1(X42,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))))&(r1_xxreal_0(k18_euclid(X43),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X42,X44,k2_nat_1(X45,np__1))))|~r2_hidden(X43,k3_goboard5(X42,X44,X45))|(~r1_xxreal_0(np__1,X44)|~r1_xxreal_0(k2_nat_1(X44,np__1),k3_finseq_1(X42))|~r1_xxreal_0(np__1,X45)|~r1_xxreal_0(k2_nat_1(X45,np__1),k1_matrix_1(X42)))|~m1_subset_1(X45,k5_numbers)|~m1_subset_1(X44,k5_numbers)|~m1_subset_1(X43,u1_struct_0(k15_euclid(np__2)))|(v3_relat_1(X42)|~v1_matrix_1(X42)|~v2_goboard1(X42)|~v3_goboard1(X42)|~v4_goboard1(X42)|~v5_goboard1(X42)|~m2_finseq_1(X42,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))))&(~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X42,X44,X45)),k17_euclid(X43))|~r1_xxreal_0(k17_euclid(X43),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X42,k2_nat_1(X44,np__1),X45)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X42,X44,X45)),k18_euclid(X43))|~r1_xxreal_0(k18_euclid(X43),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X42,X44,k2_nat_1(X45,np__1))))|r2_hidden(X43,k3_goboard5(X42,X44,X45))|(~r1_xxreal_0(np__1,X44)|~r1_xxreal_0(k2_nat_1(X44,np__1),k3_finseq_1(X42))|~r1_xxreal_0(np__1,X45)|~r1_xxreal_0(k2_nat_1(X45,np__1),k1_matrix_1(X42)))|~m1_subset_1(X45,k5_numbers)|~m1_subset_1(X44,k5_numbers)|~m1_subset_1(X43,u1_struct_0(k15_euclid(np__2)))|(v3_relat_1(X42)|~v1_matrix_1(X42)|~v2_goboard1(X42)|~v3_goboard1(X42)|~v4_goboard1(X42)|~v5_goboard1(X42)|~m2_finseq_1(X42,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t17_jordan9])])])])])).
fof(c_0_95, plain, ![X60, X61]:((~m1_subset_1(X60,k1_zfmisc_1(X61))|r1_tarski(X60,X61))&(~r1_tarski(X60,X61)|m1_subset_1(X60,k1_zfmisc_1(X61)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_96, plain, (r1_tarski(k1_rltopsp1(k15_euclid(X1),esk5_3(X1,X2,X3),esk6_3(X1,X2,X3)),X3)|r1_sppol_1(X1,X2,X3)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r2_hidden(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X1))))), inference(csr,[status(thm)],[c_0_73, c_0_64])).
cnf(c_0_97, plain, (v2_struct_0(X1)|r2_hidden(X2,k1_rltopsp1(X1,X2,X3))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~l1_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_74])).
cnf(c_0_98, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|m1_subset_1(esk5_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76]), c_0_77])])])).
cnf(c_0_99, negated_conjecture, (r2_hidden(esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))), inference(evalgc,[status(thm)],[c_0_78])).
cnf(c_0_100, negated_conjecture, (~r1_sppol_1(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))), inference(evalgc,[status(thm)],[c_0_79])).
cnf(c_0_101, plain, (l1_rlvect_1(X1)|~l1_rltopsp1(X1)), inference(evalgc,[status(thm)],[c_0_80])).
cnf(c_0_102, plain, (l1_rltopsp1(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82])])).
cnf(c_0_103, plain, (v8_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_83])).
cnf(c_0_104, plain, (v7_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_84])).
cnf(c_0_105, plain, (v6_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_106, plain, (v5_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_107, plain, (v4_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_87])).
cnf(c_0_108, plain, (v3_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_88])).
cnf(c_0_109, plain, (v2_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_110, plain, (v13_algstr_0(k15_euclid(X1))|~v7_ordinal1(X1)), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_111, plain, (r1_sppol_1(X1,X2,X3)|m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(k15_euclid(X1)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r2_hidden(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X1))))), inference(csr,[status(thm)],[c_0_91, c_0_64])).
cnf(c_0_112, plain, (v2_struct_0(X1)|k1_rltopsp1(X1,X2,X3)=k1_rltopsp1(X1,X3,X2)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~l1_rlvect_1(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_92])).
fof(c_0_113, plain, ![X50, X51, X52, X53]:(~m1_subset_1(X50,k5_numbers)|(~m1_subset_1(X51,k5_numbers)|(v3_relat_1(X52)|~v1_matrix_1(X52)|~v2_goboard1(X52)|~v3_goboard1(X52)|~v4_goboard1(X52)|~v5_goboard1(X52)|~m2_finseq_1(X52,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))|(~m1_subset_1(X53,u1_struct_0(k15_euclid(np__2)))|(~r1_xxreal_0(np__1,X50)|~r1_xxreal_0(k2_nat_1(X50,np__1),k3_finseq_1(X52))|~r1_xxreal_0(np__1,X51)|~r1_xxreal_0(k2_nat_1(X51,np__1),k1_matrix_1(X52))|~r2_hidden(X53,k16_matrix_1(X52))|~r2_hidden(X53,k3_goboard5(X52,X50,X51))|X53=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X52,X50,X51)|X53=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X52,X50,k2_nat_1(X51,np__1))|X53=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X52,k2_nat_1(X50,np__1),k2_nat_1(X51,np__1))|X53=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X52,k2_nat_1(X50,np__1),X51)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t19_jordan9])])])])).
fof(c_0_114, plain, ![X10]:(~m1_subset_1(X10,k1_numbers)|v1_xreal_0(X10)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_xreal_0])])).
cnf(c_0_115, plain, (m1_subset_1(k17_euclid(X1),k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(split_conjunct,[status(thm)],[c_0_93])).
cnf(c_0_116, plain, (r1_xxreal_0(k17_euclid(X1),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X2,k2_nat_1(X3,np__1),X4)))|v3_relat_1(X2)|~r2_hidden(X1,k3_goboard5(X2,X3,X4))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(k2_nat_1(X3,np__1),k3_finseq_1(X2))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(k2_nat_1(X4,np__1),k1_matrix_1(X2))|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_matrix_1(X2)|~v2_goboard1(X2)|~v3_goboard1(X2)|~v4_goboard1(X2)|~v5_goboard1(X2)|~m2_finseq_1(X2,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_117, negated_conjecture, (v5_goboard1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_118, negated_conjecture, (v4_goboard1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_119, negated_conjecture, (v3_goboard1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_120, negated_conjecture, (v2_goboard1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_121, negated_conjecture, (~v3_relat_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_122, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_95])).
cnf(c_0_123, negated_conjecture, (r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_76]), c_0_77])])])).
cnf(c_0_124, plain, (v2_struct_0(X1)|r2_hidden(X2,k1_rltopsp1(X1,X2,X3))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~l1_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_97])).
cnf(c_0_125, negated_conjecture, (m1_subset_1(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_98, c_0_99]), c_0_100])])).
cnf(c_0_126, plain, (l1_rlvect_1(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_102])])).
cnf(c_0_127, plain, (v8_rlvect_1(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_82])])).
cnf(c_0_128, plain, (v7_rlvect_1(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_82])])).
cnf(c_0_129, plain, (v6_rlvect_1(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_82])])).
cnf(c_0_130, plain, (v5_rlvect_1(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_82])])).
cnf(c_0_131, plain, (v4_rlvect_1(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_82])])).
cnf(c_0_132, plain, (v3_rlvect_1(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_82])])).
cnf(c_0_133, plain, (v2_rlvect_1(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_82])])).
cnf(c_0_134, plain, (v13_algstr_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_82])])).
cnf(c_0_135, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|m1_subset_1(esk6_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_76]), c_0_77])])])).
cnf(c_0_136, plain, (k1_rltopsp1(X1,X2,X3)=k1_rltopsp1(X1,X3,X2)|v2_struct_0(X1)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~l1_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_112])).
cnf(c_0_137, plain, (v3_relat_1(X3)|X4=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,X1,X2)|X4=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,X1,k2_nat_1(X2,np__1))|X4=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,k2_nat_1(X1,np__1),k2_nat_1(X2,np__1))|X4=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,k2_nat_1(X1,np__1),X2)|~m1_subset_1(X1,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~v1_matrix_1(X3)|~v2_goboard1(X3)|~v3_goboard1(X3)|~v4_goboard1(X3)|~v5_goboard1(X3)|~m2_finseq_1(X3,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X4,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(np__1,X1)|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(X3))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k2_nat_1(X2,np__1),k1_matrix_1(X3))|~r2_hidden(X4,k16_matrix_1(X3))|~r2_hidden(X4,k3_goboard5(X3,X1,X2))), inference(split_conjunct,[status(thm)],[c_0_113])).
fof(c_0_138, plain, ![X11]:(~v1_xreal_0(X11)|v1_xxreal_0(X11)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc4_xreal_0])])).
cnf(c_0_139, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(split_conjunct,[status(thm)],[c_0_114])).
cnf(c_0_140, plain, (m1_subset_1(k17_euclid(X1),k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[c_0_115])).
cnf(c_0_141, plain, (v3_relat_1(X2)|r1_xxreal_0(k17_euclid(X1),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X2,k2_nat_1(X3,np__1),X4)))|~v1_matrix_1(X2)|~v2_goboard1(X2)|~v3_goboard1(X2)|~v4_goboard1(X2)|~v5_goboard1(X2)|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X3,k5_numbers)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r2_hidden(X1,k3_goboard5(X2,X3,X4))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k2_nat_1(X4,np__1),k1_matrix_1(X2))|~r1_xxreal_0(k2_nat_1(X3,np__1),k3_finseq_1(X2))|~m2_finseq_1(X2,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[c_0_116])).
cnf(c_0_142, negated_conjecture, (v5_goboard1(esk3_0)), inference(evalgc,[status(thm)],[c_0_117])).
cnf(c_0_143, negated_conjecture, (v4_goboard1(esk3_0)), inference(evalgc,[status(thm)],[c_0_118])).
cnf(c_0_144, negated_conjecture, (v3_goboard1(esk3_0)), inference(evalgc,[status(thm)],[c_0_119])).
cnf(c_0_145, negated_conjecture, (v2_goboard1(esk3_0)), inference(evalgc,[status(thm)],[c_0_120])).
cnf(c_0_146, negated_conjecture, (~v3_relat_1(esk3_0)), inference(evalgc,[status(thm)],[c_0_121])).
cnf(c_0_147, negated_conjecture, (r1_xxreal_0(k2_nat_1(esk2_0,np__1),k1_matrix_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_148, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_149, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_122])).
cnf(c_0_150, negated_conjecture, (r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k3_goboard5(esk3_0,esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_99]), c_0_100])])).
cnf(c_0_151, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),X1,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_125]), c_0_126]), c_0_127]), c_0_128]), c_0_129]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134])])])).
cnf(c_0_152, negated_conjecture, (m1_subset_1(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_99]), c_0_100])])).
cnf(c_0_153, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),X1,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),X1)|v2_struct_0(k15_euclid(np__2))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_125]), c_0_126]), c_0_127]), c_0_128]), c_0_129]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134])])])).
fof(c_0_154, plain, ![X80, X81]:(~r2_hidden(X80,X81)|~v1_xboole_0(X81)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
fof(c_0_155, plain, ![X37, X38, X39]:((((k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X39,X37,X38))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X39,X37,k2_nat_1(X38,np__1)))|(~r1_xxreal_0(np__1,X37)|~r1_xxreal_0(k2_nat_1(X37,np__1),k3_finseq_1(X39))|~r1_xxreal_0(np__1,X38)|~r1_xxreal_0(k2_nat_1(X38,np__1),k1_matrix_1(X39)))|(v3_relat_1(X39)|~v1_matrix_1(X39)|~v2_goboard1(X39)|~v3_goboard1(X39)|~v4_goboard1(X39)|~v5_goboard1(X39)|~m2_finseq_1(X39,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))|~m1_subset_1(X38,k5_numbers)|~m1_subset_1(X37,k5_numbers))&(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X39,X37,X38))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X39,k2_nat_1(X37,np__1),X38))|(~r1_xxreal_0(np__1,X37)|~r1_xxreal_0(k2_nat_1(X37,np__1),k3_finseq_1(X39))|~r1_xxreal_0(np__1,X38)|~r1_xxreal_0(k2_nat_1(X38,np__1),k1_matrix_1(X39)))|(v3_relat_1(X39)|~v1_matrix_1(X39)|~v2_goboard1(X39)|~v3_goboard1(X39)|~v4_goboard1(X39)|~v5_goboard1(X39)|~m2_finseq_1(X39,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))|~m1_subset_1(X38,k5_numbers)|~m1_subset_1(X37,k5_numbers)))&(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X39,k2_nat_1(X37,np__1),k2_nat_1(X38,np__1)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X39,k2_nat_1(X37,np__1),X38))|(~r1_xxreal_0(np__1,X37)|~r1_xxreal_0(k2_nat_1(X37,np__1),k3_finseq_1(X39))|~r1_xxreal_0(np__1,X38)|~r1_xxreal_0(k2_nat_1(X38,np__1),k1_matrix_1(X39)))|(v3_relat_1(X39)|~v1_matrix_1(X39)|~v2_goboard1(X39)|~v3_goboard1(X39)|~v4_goboard1(X39)|~v5_goboard1(X39)|~m2_finseq_1(X39,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))|~m1_subset_1(X38,k5_numbers)|~m1_subset_1(X37,k5_numbers)))&(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X39,k2_nat_1(X37,np__1),k2_nat_1(X38,np__1)))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X39,X37,k2_nat_1(X38,np__1)))|(~r1_xxreal_0(np__1,X37)|~r1_xxreal_0(k2_nat_1(X37,np__1),k3_finseq_1(X39))|~r1_xxreal_0(np__1,X38)|~r1_xxreal_0(k2_nat_1(X38,np__1),k1_matrix_1(X39)))|(v3_relat_1(X39)|~v1_matrix_1(X39)|~v2_goboard1(X39)|~v3_goboard1(X39)|~v4_goboard1(X39)|~v5_goboard1(X39)|~m2_finseq_1(X39,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))|~m1_subset_1(X38,k5_numbers)|~m1_subset_1(X37,k5_numbers))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t16_jordan9])])])])])).
cnf(c_0_156, plain, (X4=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,X1,X2)|X4=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,X1,k2_nat_1(X2,np__1))|X4=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,k2_nat_1(X1,np__1),X2)|X4=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,k2_nat_1(X1,np__1),k2_nat_1(X2,np__1))|v3_relat_1(X3)|~v1_matrix_1(X3)|~v2_goboard1(X3)|~v3_goboard1(X3)|~v4_goboard1(X3)|~v5_goboard1(X3)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r2_hidden(X4,k16_matrix_1(X3))|~r2_hidden(X4,k3_goboard5(X3,X1,X2))|~m1_subset_1(X4,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k2_nat_1(X2,np__1),k1_matrix_1(X3))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(X3))|~m2_finseq_1(X3,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[c_0_137])).
fof(c_0_157, plain, ![X16, X17]:(~v1_xxreal_0(X16)|~v1_xxreal_0(X17)|(r1_xxreal_0(X16,X17)|r1_xxreal_0(X17,X16))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[connectedness_r1_xxreal_0])])).
cnf(c_0_158, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(split_conjunct,[status(thm)],[c_0_138])).
cnf(c_0_159, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers)), inference(evalgc,[status(thm)],[c_0_139])).
cnf(c_0_160, negated_conjecture, (m1_subset_1(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_125])])).
cnf(c_0_161, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X2,np__1),X3)))|~r2_hidden(X1,k3_goboard5(esk3_0,X2,X3))|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_42]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_51])]), c_0_146])])).
cnf(c_0_162, negated_conjecture, (r1_xxreal_0(k2_nat_1(esk2_0,np__1),k1_matrix_1(esk3_0))), inference(evalgc,[status(thm)],[c_0_147])).
cnf(c_0_163, negated_conjecture, (r1_xxreal_0(np__1,esk2_0)), inference(evalgc,[status(thm)],[c_0_148])).
cnf(c_0_164, negated_conjecture, (r1_xxreal_0(k2_nat_1(esk1_0,np__1),k3_finseq_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_165, negated_conjecture, (r1_xxreal_0(np__1,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_33])).
fof(c_0_166, plain, ![X58, X59]:(~m1_subset_1(X58,X59)|(v1_xboole_0(X59)|r2_hidden(X58,X59))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_167, negated_conjecture, (m1_subset_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_zfmisc_1(k3_goboard5(esk3_0,esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_150])])).
cnf(c_0_168, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152])])).
cnf(c_0_169, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_152])])).
cnf(c_0_170, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_154])).
cnf(c_0_171, plain, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1),k2_nat_1(X3,np__1)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1),X3))|v3_relat_1(X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~v1_matrix_1(X1)|~v2_goboard1(X1)|~v3_goboard1(X1)|~v4_goboard1(X1)|~v5_goboard1(X1)|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_155])).
cnf(c_0_172, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X2,np__1),k2_nat_1(X3,np__1))|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X2,np__1),X3)|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X2,k2_nat_1(X3,np__1))|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X2,X3)|~r2_hidden(X1,k3_goboard5(esk3_0,X2,X3))|~r2_hidden(X1,k16_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_42]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_51])]), c_0_146])])).
cnf(c_0_173, plain, (r2_hidden(X1,k1_rltopsp1(k15_euclid(X2),esk5_3(X2,X1,X3),esk6_3(X2,X1,X3)))|r1_sppol_1(X2,X1,X3)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X2))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_47])).
fof(c_0_174, plain, ![X62, X63, X64]:((r1_xxreal_0(k17_euclid(X62),k17_euclid(X64))|(~r1_xxreal_0(k17_euclid(X62),k17_euclid(X63))|~r2_hidden(X64,k1_rltopsp1(k15_euclid(np__2),X62,X63)))|~m1_subset_1(X64,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X63,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X62,u1_struct_0(k15_euclid(np__2))))&(r1_xxreal_0(k17_euclid(X64),k17_euclid(X63))|(~r1_xxreal_0(k17_euclid(X62),k17_euclid(X63))|~r2_hidden(X64,k1_rltopsp1(k15_euclid(np__2),X62,X63)))|~m1_subset_1(X64,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X63,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X62,u1_struct_0(k15_euclid(np__2))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_topreal1])])])])).
cnf(c_0_175, plain, (r1_xxreal_0(X1,X2)|r1_xxreal_0(X2,X1)|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)), inference(split_conjunct,[status(thm)],[c_0_157])).
cnf(c_0_176, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1)), inference(evalgc,[status(thm)],[c_0_158])).
cnf(c_0_177, negated_conjecture, (v1_xreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_160])])).
cnf(c_0_178, negated_conjecture, (m1_subset_1(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_152])])).
cnf(c_0_179, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X2,np__1),esk2_0)))|~r2_hidden(X1,k3_goboard5(esk3_0,X2,esk2_0))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_162]), c_0_163]), c_0_53])])])).
cnf(c_0_180, negated_conjecture, (r1_xxreal_0(k2_nat_1(esk1_0,np__1),k3_finseq_1(esk3_0))), inference(evalgc,[status(thm)],[c_0_164])).
cnf(c_0_181, negated_conjecture, (r1_xxreal_0(np__1,esk1_0)), inference(evalgc,[status(thm)],[c_0_165])).
cnf(c_0_182, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_166])).
cnf(c_0_183, negated_conjecture, (m1_subset_1(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_167])])).
cnf(c_0_184, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_169])])).
cnf(c_0_185, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_170])).
cnf(c_0_186, plain, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1),k2_nat_1(X3,np__1)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1),X3))|v3_relat_1(X1)|~v1_matrix_1(X1)|~v2_goboard1(X1)|~v3_goboard1(X1)|~v4_goboard1(X1)|~v5_goboard1(X1)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[c_0_171])).
cnf(c_0_187, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X2,np__1),k2_nat_1(esk2_0,np__1))|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X2,k2_nat_1(esk2_0,np__1))|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X2,np__1),esk2_0)|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X2,esk2_0)|~r2_hidden(X1,k3_goboard5(esk3_0,X2,esk2_0))|~r2_hidden(X1,k16_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_162]), c_0_163]), c_0_53])])])).
cnf(c_0_188, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(k15_euclid(np__2)))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_189, plain, (r1_sppol_1(X2,X1,X3)|r2_hidden(X1,k1_rltopsp1(k15_euclid(X2),esk5_3(X2,X1,X3),esk6_3(X2,X1,X3)))|~r2_hidden(X1,X3)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X2))))), inference(evalgc,[status(thm)],[c_0_173])).
cnf(c_0_190, plain, (r1_xxreal_0(k17_euclid(X1),k17_euclid(X2))|~r1_xxreal_0(k17_euclid(X3),k17_euclid(X2))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),X3,X2))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_191, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[c_0_175])).
cnf(c_0_192, negated_conjecture, (v1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_177])])).
cnf(c_0_193, negated_conjecture, (v1_xreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_178])])).
fof(c_0_194, plain, ![X56, X57]:(~v1_xxreal_0(X56)|(~v1_xxreal_0(X57)|(~r1_xxreal_0(X56,X57)|~r1_xxreal_0(X57,X56)|X56=X57))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xxreal_0])])])).
cnf(c_0_195, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_180]), c_0_181]), c_0_59])])])).
cnf(c_0_196, negated_conjecture, (m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_76])])).
cnf(c_0_197, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_182])).
cnf(c_0_198, negated_conjecture, (m1_subset_1(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_184])])).
cnf(c_0_199, negated_conjecture, (~v1_xboole_0(k3_goboard5(esk3_0,esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_99])])).
cnf(c_0_200, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X1,np__1),k2_nat_1(X2,np__1)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X1,np__1),X2))|~r1_xxreal_0(k2_nat_1(X2,np__1),k1_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_42]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_51])]), c_0_146])])).
cnf(c_0_201, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k16_matrix_1(esk3_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_180]), c_0_181]), c_0_59])])])).
cnf(c_0_202, negated_conjecture, (r2_hidden(esk4_0,k16_matrix_1(esk3_0))), inference(split_conjunct,[status(thm)],[c_0_33])).
cnf(c_0_203, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[c_0_188])).
cnf(c_0_204, plain, (r1_sppol_1(X1,X2,X3)|r2_hidden(X2,k1_rltopsp1(k15_euclid(X1),esk5_3(X1,X2,X3),esk6_3(X1,X2,X3)))|~m2_subset_1(X1,k1_numbers,k5_numbers)|~r2_hidden(X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X1))))), inference(csr,[status(thm)],[c_0_189, c_0_64])).
cnf(c_0_205, plain, (r1_xxreal_0(k17_euclid(X1),k17_euclid(X2))|~r1_xxreal_0(k17_euclid(X3),k17_euclid(X2))|~m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),X3,X2))), inference(evalgc,[status(thm)],[c_0_190])).
cnf(c_0_206, negated_conjecture, (r1_xxreal_0(X1,k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),X1)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_192])])).
cnf(c_0_207, negated_conjecture, (v1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_193])])).
cnf(c_0_208, plain, (X1=X2|~v1_xxreal_0(X1)|~v1_xxreal_0(X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_194])).
cnf(c_0_209, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))), inference(csr,[status(thm)],[c_0_195, c_0_196])).
cnf(c_0_210, negated_conjecture, (r2_hidden(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_198]), c_0_199])])).
cnf(c_0_211, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X1,np__1),k2_nat_1(esk2_0,np__1)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X1,np__1),esk2_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_162]), c_0_163]), c_0_53])])])).
cnf(c_0_212, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k16_matrix_1(esk3_0))), inference(csr,[status(thm)],[c_0_201, c_0_196])).
cnf(c_0_213, negated_conjecture, (r2_hidden(esk4_0,k16_matrix_1(esk3_0))), inference(evalgc,[status(thm)],[c_0_202])).
cnf(c_0_214, negated_conjecture, (m1_subset_1(k17_euclid(esk4_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_203])])).
cnf(c_0_215, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_76]), c_0_77])])])).
cnf(c_0_216, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_169]), c_0_152]), c_0_125])])])).
cnf(c_0_217, negated_conjecture, (r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_207])])).
cnf(c_0_218, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),X1,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_152]), c_0_126]), c_0_127]), c_0_128]), c_0_129]), c_0_130]), c_0_131]), c_0_132]), c_0_133]), c_0_134])])])).
cnf(c_0_219, plain, (X1=X2|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_208])).
cnf(c_0_220, negated_conjecture, (r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_210])])).
cnf(c_0_221, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_180]), c_0_181]), c_0_59])])])).
cnf(c_0_222, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_99]), c_0_213])])])).
cnf(c_0_223, negated_conjecture, (v1_xreal_0(k17_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_214])])).
cnf(c_0_224, negated_conjecture, (r2_hidden(esk4_0,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_99]), c_0_100])])).
cnf(c_0_225, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_217])])).
cnf(c_0_226, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_125])])).
cnf(c_0_227, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))=k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|~v1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_220]), c_0_207])])])).
cnf(c_0_228, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_222])])).
cnf(c_0_229, negated_conjecture, (v1_xxreal_0(k17_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_223])])).
cnf(c_0_230, negated_conjecture, (r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_224]), c_0_125]), c_0_152]), c_0_203])])])).
cnf(c_0_231, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_224]), c_0_203])])])).
cnf(c_0_232, negated_conjecture, (m1_subset_1(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_226])])).
cnf(c_0_233, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_228]), c_0_229])])])).
cnf(c_0_234, negated_conjecture, (r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_231])])).
cnf(c_0_235, negated_conjecture, (r2_hidden(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_232]), c_0_199])])).
fof(c_0_236, plain, ![X27]:(~m1_subset_1(X27,u1_struct_0(k15_euclid(np__2)))|m1_subset_1(k18_euclid(X27),k1_numbers)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k18_euclid])])).
cnf(c_0_237, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_234])])).
cnf(c_0_238, negated_conjecture, (r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_235])])).
cnf(c_0_239, plain, (m1_subset_1(k18_euclid(X1),k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(split_conjunct,[status(thm)],[c_0_236])).
fof(c_0_240, plain, ![X46, X47, X48, X49]:(~m1_subset_1(X46,u1_struct_0(k15_euclid(np__2)))|(~m1_subset_1(X47,u1_struct_0(k15_euclid(np__2)))|(~m1_subset_1(X48,u1_struct_0(k15_euclid(np__2)))|(~m1_subset_1(X49,u1_struct_0(k15_euclid(np__2)))|(~r2_hidden(X46,k1_rltopsp1(k15_euclid(np__2),X47,X48))|~r2_hidden(X49,k1_rltopsp1(k15_euclid(np__2),X47,X48))|k17_euclid(X46)!=k17_euclid(X49)|(k18_euclid(X46)=k18_euclid(X49)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X47,X48)))))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t18_sppol_1])])])).
cnf(c_0_241, plain, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,k2_nat_1(X3,np__1)))|v3_relat_1(X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~v1_matrix_1(X1)|~v2_goboard1(X1)|~v3_goboard1(X1)|~v4_goboard1(X1)|~v5_goboard1(X1)|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_155])).
cnf(c_0_242, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_237]), c_0_229]), c_0_192])])])).
cnf(c_0_243, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_228])])).
cnf(c_0_244, plain, (r1_xxreal_0(k18_euclid(X1),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X2,X3,k2_nat_1(X4,np__1))))|v3_relat_1(X2)|~r2_hidden(X1,k3_goboard5(X2,X3,X4))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(k2_nat_1(X3,np__1),k3_finseq_1(X2))|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(k2_nat_1(X4,np__1),k1_matrix_1(X2))|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~v1_matrix_1(X2)|~v2_goboard1(X2)|~v3_goboard1(X2)|~v4_goboard1(X2)|~v5_goboard1(X2)|~m2_finseq_1(X2,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_245, plain, (m1_subset_1(k18_euclid(X1),k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[c_0_239])).
cnf(c_0_246, plain, (k18_euclid(X1)=k18_euclid(X4)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X2,X3))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X4,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),X2,X3))|~r2_hidden(X4,k1_rltopsp1(k15_euclid(np__2),X2,X3))|k17_euclid(X1)!=k17_euclid(X4)), inference(split_conjunct,[status(thm)],[c_0_240])).
cnf(c_0_247, plain, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,k2_nat_1(X3,np__1)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3))|v3_relat_1(X1)|~v1_matrix_1(X1)|~v2_goboard1(X1)|~v3_goboard1(X1)|~v4_goboard1(X1)|~v5_goboard1(X1)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[c_0_241])).
cnf(c_0_248, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(csr,[status(thm)],[c_0_242, c_0_243])).
cnf(c_0_249, plain, (v3_relat_1(X2)|r1_xxreal_0(k18_euclid(X1),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X2,X3,k2_nat_1(X4,np__1))))|~v1_matrix_1(X2)|~v2_goboard1(X2)|~v3_goboard1(X2)|~v4_goboard1(X2)|~v5_goboard1(X2)|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X3,k5_numbers)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r2_hidden(X1,k3_goboard5(X2,X3,X4))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k2_nat_1(X4,np__1),k1_matrix_1(X2))|~r1_xxreal_0(k2_nat_1(X3,np__1),k3_finseq_1(X2))|~m2_finseq_1(X2,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[c_0_244])).
cnf(c_0_250, negated_conjecture, (m1_subset_1(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_125])])).
fof(c_0_251, plain, ![X78, X79]:(~m1_subset_1(X78,u1_struct_0(k15_euclid(np__2)))|(~m1_subset_1(X79,u1_struct_0(k15_euclid(np__2)))|(k17_euclid(X78)!=k17_euclid(X79)|k18_euclid(X78)!=k18_euclid(X79)|X78=X79))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t6_topreal3])])])).
cnf(c_0_252, plain, (k18_euclid(X1)=k18_euclid(X4)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X2,X3))|k17_euclid(X1)!=k17_euclid(X4)|~m1_subset_1(X4,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X4,k1_rltopsp1(k15_euclid(np__2),X2,X3))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),X2,X3))), inference(evalgc,[status(thm)],[c_0_246])).
cnf(c_0_253, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,k2_nat_1(X2,np__1)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,X2))|~r1_xxreal_0(k2_nat_1(X2,np__1),k1_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_42]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_51])]), c_0_146])])).
cnf(c_0_254, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_248]), c_0_229])])])).
cnf(c_0_255, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|~v1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_238]), c_0_192])])])).
cnf(c_0_256, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X2,k2_nat_1(X3,np__1))))|~r2_hidden(X1,k3_goboard5(esk3_0,X2,X3))|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_42]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_51])]), c_0_146])])).
cnf(c_0_257, plain, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1),k2_nat_1(X3,np__1)))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,k2_nat_1(X3,np__1)))|v3_relat_1(X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~v1_matrix_1(X1)|~v2_goboard1(X1)|~v3_goboard1(X1)|~v4_goboard1(X1)|~v5_goboard1(X1)|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_155])).
fof(c_0_258, plain, ![X68, X69, X70]:((r1_xxreal_0(k18_euclid(X68),k18_euclid(X70))|(~r1_xxreal_0(k18_euclid(X68),k18_euclid(X69))|~r2_hidden(X70,k1_rltopsp1(k15_euclid(np__2),X68,X69)))|~m1_subset_1(X70,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X69,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X68,u1_struct_0(k15_euclid(np__2))))&(r1_xxreal_0(k18_euclid(X70),k18_euclid(X69))|(~r1_xxreal_0(k18_euclid(X68),k18_euclid(X69))|~r2_hidden(X70,k1_rltopsp1(k15_euclid(np__2),X68,X69)))|~m1_subset_1(X70,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X69,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X68,u1_struct_0(k15_euclid(np__2))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_topreal1])])])])).
cnf(c_0_259, negated_conjecture, (v1_xreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_250])])).
cnf(c_0_260, negated_conjecture, (m1_subset_1(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_152])])).
cnf(c_0_261, plain, (X1=X2|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|k17_euclid(X1)!=k17_euclid(X2)|k18_euclid(X1)!=k18_euclid(X2)), inference(split_conjunct,[status(thm)],[c_0_251])).
cnf(c_0_262, negated_conjecture, (k18_euclid(X1)=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_168]), c_0_152]), c_0_125])])])).
cnf(c_0_263, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,k2_nat_1(esk2_0,np__1)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,esk2_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_162]), c_0_163]), c_0_53])])])).
cnf(c_0_264, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_234])])).
cnf(c_0_265, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_248]), c_0_229])])])).
cnf(c_0_266, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X2,k2_nat_1(esk2_0,np__1))))|~r2_hidden(X1,k3_goboard5(esk3_0,X2,esk2_0))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_162]), c_0_163]), c_0_53])])])).
cnf(c_0_267, plain, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1),k2_nat_1(X3,np__1)))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,k2_nat_1(X3,np__1)))|v3_relat_1(X1)|~v1_matrix_1(X1)|~v2_goboard1(X1)|~v3_goboard1(X1)|~v4_goboard1(X1)|~v5_goboard1(X1)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[c_0_257])).
cnf(c_0_268, plain, (r1_xxreal_0(k18_euclid(X1),k18_euclid(X2))|~r1_xxreal_0(k18_euclid(X3),k18_euclid(X2))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),X3,X2))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))), inference(split_conjunct,[status(thm)],[c_0_258])).
cnf(c_0_269, negated_conjecture, (v1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_259])])).
cnf(c_0_270, negated_conjecture, (v1_xreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_260])])).
cnf(c_0_271, plain, (X1=X2|k17_euclid(X1)!=k17_euclid(X2)|k18_euclid(X1)!=k18_euclid(X2)|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[c_0_261])).
fof(c_0_272, plain, ![X40, X41]:((k17_euclid(X40)!=k17_euclid(X41)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X40,X41))|~m1_subset_1(X41,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X40,u1_struct_0(k15_euclid(np__2))))&(~v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X40,X41))|k17_euclid(X40)=k17_euclid(X41)|~m1_subset_1(X41,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X40,u1_struct_0(k15_euclid(np__2))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t16_sppol_1])])])])).
cnf(c_0_273, negated_conjecture, (k18_euclid(X1)=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_169])])).
cnf(c_0_274, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_180]), c_0_181]), c_0_59])])])).
cnf(c_0_275, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(csr,[status(thm)],[c_0_264, c_0_265])).
cnf(c_0_276, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_180]), c_0_181]), c_0_59])])])).
cnf(c_0_277, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X1,np__1),k2_nat_1(X2,np__1)))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,k2_nat_1(X2,np__1)))|~r1_xxreal_0(k2_nat_1(X2,np__1),k1_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_42]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_51])]), c_0_146])])).
cnf(c_0_278, plain, (r1_xxreal_0(k18_euclid(X1),k18_euclid(X2))|~r1_xxreal_0(k18_euclid(X3),k18_euclid(X2))|~m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),X3,X2))), inference(evalgc,[status(thm)],[c_0_268])).
cnf(c_0_279, negated_conjecture, (r1_xxreal_0(X1,k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),X1)|~v1_xxreal_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_269])])).
cnf(c_0_280, negated_conjecture, (v1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_270])])).
cnf(c_0_281, negated_conjecture, (X1=esk4_0|k17_euclid(X1)!=k17_euclid(esk4_0)|k18_euclid(X1)!=k18_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_203])])).
cnf(c_0_282, plain, (k17_euclid(X1)=k17_euclid(X2)|~v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,X2))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(split_conjunct,[status(thm)],[c_0_272])).
cnf(c_0_283, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_224]), c_0_203])])])).
cnf(c_0_284, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_274, c_0_275])])).
cnf(c_0_285, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))), inference(csr,[status(thm)],[c_0_276, c_0_196])).
cnf(c_0_286, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X1,np__1),k2_nat_1(esk2_0,np__1)))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,k2_nat_1(esk2_0,np__1)))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_277, c_0_162]), c_0_163]), c_0_53])])])).
cnf(c_0_287, negated_conjecture, (m1_subset_1(k18_euclid(esk4_0),k1_numbers)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_203])])).
cnf(c_0_288, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_169]), c_0_152]), c_0_125])])])).
cnf(c_0_289, negated_conjecture, (r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_279, c_0_280])])).
cnf(c_0_290, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_281, c_0_152])])).
cnf(c_0_291, plain, (k17_euclid(X1)=k17_euclid(X2)|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,X2))), inference(evalgc,[status(thm)],[c_0_282])).
cnf(c_0_292, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_284])])).
cnf(c_0_293, plain, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1),X3))|v3_relat_1(X1)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~v1_matrix_1(X1)|~v2_goboard1(X1)|~v3_goboard1(X1)|~v4_goboard1(X1)|~v5_goboard1(X1)|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_155])).
cnf(c_0_294, negated_conjecture, (r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_210])])).
cnf(c_0_295, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_180]), c_0_181]), c_0_59])])])).
cnf(c_0_296, negated_conjecture, (v1_xreal_0(k18_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_287])])).
cnf(c_0_297, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_288, c_0_289])])).
cnf(c_0_298, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_284])])).
cnf(c_0_299, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_292]), c_0_152]), c_0_125])])])).
cnf(c_0_300, plain, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1),X3))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3))|v3_relat_1(X1)|~v1_matrix_1(X1)|~v2_goboard1(X1)|~v3_goboard1(X1)|~v4_goboard1(X1)|~v5_goboard1(X1)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[c_0_293])).
cnf(c_0_301, negated_conjecture, (r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_285, c_0_235])])).
cnf(c_0_302, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|~v1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_294]), c_0_280])])])).
cnf(c_0_303, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_295, c_0_222])])).
cnf(c_0_304, negated_conjecture, (v1_xxreal_0(k18_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_296])])).
cnf(c_0_305, negated_conjecture, (r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_278, c_0_224]), c_0_125]), c_0_152]), c_0_203])])])).
cnf(c_0_306, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_297, c_0_224]), c_0_203])])])).
cnf(c_0_307, negated_conjecture, (k18_euclid(X1)=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(X1)!=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_224]), c_0_203]), c_0_152]), c_0_125])])])).
cnf(c_0_308, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_298, c_0_299])])).
cnf(c_0_309, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X1,np__1),X2))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,X2))|~r1_xxreal_0(k2_nat_1(X2,np__1),k1_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_300, c_0_42]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_51])]), c_0_146])])).
cnf(c_0_310, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|~v1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_301]), c_0_269])])])).
cnf(c_0_311, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_303]), c_0_304])])])).
cnf(c_0_312, negated_conjecture, (r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_305, c_0_306])])).
cnf(c_0_313, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_307, c_0_226]), c_0_125])])])).
cnf(c_0_314, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_284, c_0_308])])).
cnf(c_0_315, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X1,np__1),esk2_0))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,esk2_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_162]), c_0_163]), c_0_53])])])).
cnf(c_0_316, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_303]), c_0_304])])])).
cnf(c_0_317, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_312])])).
cnf(c_0_318, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_281, c_0_125])])).
cnf(c_0_319, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_313, c_0_314])])).
cnf(c_0_320, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_315, c_0_180]), c_0_181]), c_0_59])])])).
cnf(c_0_321, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_316, c_0_317])])).
cnf(c_0_322, plain, (r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k18_euclid(X4))|v3_relat_1(X1)|~r2_hidden(X4,k3_goboard5(X1,X2,X3))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X4,u1_struct_0(k15_euclid(np__2)))|~v1_matrix_1(X1)|~v2_goboard1(X1)|~v3_goboard1(X1)|~v4_goboard1(X1)|~v5_goboard1(X1)|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_323, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_314])])).
cnf(c_0_324, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_319]), c_0_152]), c_0_125])])])).
cnf(c_0_325, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_320, c_0_321])])).
cnf(c_0_326, plain, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,X2))|k17_euclid(X1)!=k17_euclid(X2)|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(split_conjunct,[status(thm)],[c_0_272])).
cnf(c_0_327, plain, (v3_relat_1(X1)|r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k18_euclid(X4))|~v1_matrix_1(X1)|~v2_goboard1(X1)|~v3_goboard1(X1)|~v4_goboard1(X1)|~v5_goboard1(X1)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r2_hidden(X4,k3_goboard5(X1,X2,X3))|~m1_subset_1(X4,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[c_0_322])).
cnf(c_0_328, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_323, c_0_324])])).
cnf(c_0_329, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_302, c_0_325]), c_0_304])])])).
cnf(c_0_330, plain, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,X2))|k17_euclid(X1)!=k17_euclid(X2)|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[c_0_326])).
cnf(c_0_331, negated_conjecture, (r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,X2)),k18_euclid(X3))|~r2_hidden(X3,k3_goboard5(esk3_0,X1,X2))|~r1_xxreal_0(k2_nat_1(X2,np__1),k1_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_327, c_0_42]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_51])]), c_0_146])])).
cnf(c_0_332, plain, (r1_sppol_1(X2,X1,X3)|X1!=esk6_3(X2,X1,X3)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X2))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_333, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_328])])).
cnf(c_0_334, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_329, c_0_312])])).
cnf(c_0_335, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_325]), c_0_304])])])).
cnf(c_0_336, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,esk4_0))|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_203])])).
cnf(c_0_337, negated_conjecture, (r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,esk2_0)),k18_euclid(X2))|~r2_hidden(X2,k3_goboard5(esk3_0,X1,esk2_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_331, c_0_162]), c_0_163]), c_0_53])])])).
cnf(c_0_338, plain, (r1_sppol_1(X2,X1,X3)|X1!=esk5_3(X2,X1,X3)|~r2_hidden(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X2))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(X2)))|~m2_subset_1(X2,k1_numbers,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_339, plain, (r1_sppol_1(X2,X1,X3)|esk6_3(X2,X1,X3)!=X1|~r2_hidden(X1,X3)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X2))))), inference(evalgc,[status(thm)],[c_0_332])).
cnf(c_0_340, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)), inference(csr,[status(thm)],[c_0_333, c_0_314])).
cnf(c_0_341, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(csr,[status(thm)],[c_0_334, c_0_335])).
cnf(c_0_342, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_336, c_0_152])])).
cnf(c_0_343, negated_conjecture, (r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(X1))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_337, c_0_180]), c_0_181]), c_0_59])])])).
cnf(c_0_344, plain, (r1_sppol_1(X2,X1,X3)|esk5_3(X2,X1,X3)!=X1|~r2_hidden(X1,X3)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X2))))), inference(evalgc,[status(thm)],[c_0_338])).
cnf(c_0_345, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|esk6_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_339, c_0_76]), c_0_77])])])).
cnf(c_0_346, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_340, c_0_341])])).
cnf(c_0_347, plain, (r1_xxreal_0(k18_euclid(X1),k18_euclid(X2))|~r1_xxreal_0(k18_euclid(X1),k18_euclid(X3))|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),X1,X3))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(split_conjunct,[status(thm)],[c_0_258])).
cnf(c_0_348, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_342, c_0_328])])).
cnf(c_0_349, negated_conjecture, (r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(X1))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))), inference(csr,[status(thm)],[c_0_343, c_0_196])).
cnf(c_0_350, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|esk5_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_76]), c_0_77])])])).
cnf(c_0_351, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|esk6_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))), inference(csr,[status(thm)],[c_0_345, c_0_196])).
cnf(c_0_352, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(csr,[status(thm)],[c_0_346, c_0_323])).
cnf(c_0_353, plain, (r1_xxreal_0(k18_euclid(X1),k18_euclid(X2))|~r1_xxreal_0(k18_euclid(X1),k18_euclid(X3))|~m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),X1,X3))), inference(evalgc,[status(thm)],[c_0_347])).
cnf(c_0_354, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|v2_struct_0(k15_euclid(np__2))), inference(csr,[status(thm)],[c_0_348, c_0_314])).
cnf(c_0_355, negated_conjecture, (r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_349, c_0_235])])).
cnf(c_0_356, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|esk5_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))), inference(csr,[status(thm)],[c_0_350, c_0_196])).
cnf(c_0_357, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_351, c_0_352]), c_0_99])]), c_0_100])])).
cnf(c_0_358, negated_conjecture, (r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|v2_struct_0(k15_euclid(np__2))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_353, c_0_169]), c_0_125]), c_0_152])])])).
cnf(c_0_359, negated_conjecture, (X1=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(X1)!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_152])])).
cnf(c_0_360, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_354]), c_0_203]), c_0_152])])])).
cnf(c_0_361, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|~v1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_355]), c_0_269])])])).
cnf(c_0_362, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_357]), c_0_99])]), c_0_100])])).
cnf(c_0_363, negated_conjecture, (r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_349, c_0_210])])).
cnf(c_0_364, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_358, c_0_289])])).
cnf(c_0_365, negated_conjecture, (X1=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|k18_euclid(X1)!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_359, c_0_360])])).
cnf(c_0_366, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_362]), c_0_304])])])).
cnf(c_0_367, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_301, c_0_303])])).
cnf(c_0_368, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|~v1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_363]), c_0_280])])])).
cnf(c_0_369, negated_conjecture, (r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_353, c_0_224]), c_0_152]), c_0_203]), c_0_125])])])).
cnf(c_0_370, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_364, c_0_224]), c_0_203])])])).
cnf(c_0_371, plain, (r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k17_euclid(X4))|v3_relat_1(X1)|~r2_hidden(X4,k3_goboard5(X1,X2,X3))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X4,u1_struct_0(k15_euclid(np__2)))|~v1_matrix_1(X1)|~v2_goboard1(X1)|~v3_goboard1(X1)|~v4_goboard1(X1)|~v5_goboard1(X1)|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))), inference(split_conjunct,[status(thm)],[c_0_94])).
cnf(c_0_372, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_365, c_0_320])])).
cnf(c_0_373, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_367])])).
cnf(c_0_374, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_368, c_0_362]), c_0_304])])])).
cnf(c_0_375, negated_conjecture, (r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_369, c_0_370])])).
cnf(c_0_376, plain, (v3_relat_1(X1)|r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k17_euclid(X4))|~v1_matrix_1(X1)|~v2_goboard1(X1)|~v3_goboard1(X1)|~v4_goboard1(X1)|~v5_goboard1(X1)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r2_hidden(X4,k3_goboard5(X1,X2,X3))|~m1_subset_1(X4,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))), inference(evalgc,[status(thm)],[c_0_371])).
cnf(c_0_377, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_372, c_0_373]), c_0_203])])])).
cnf(c_0_378, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_374, c_0_375])])).
cnf(c_0_379, negated_conjecture, (r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,X2)),k17_euclid(X3))|~r2_hidden(X3,k3_goboard5(esk3_0,X1,X2))|~r1_xxreal_0(k2_nat_1(X2,np__1),k1_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_376, c_0_42]), c_0_142]), c_0_143]), c_0_144]), c_0_145]), c_0_51])]), c_0_146])])).
cnf(c_0_380, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))), inference(csr,[status(thm)],[c_0_377, c_0_323])).
cnf(c_0_381, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_366, c_0_378])])).
cnf(c_0_382, negated_conjecture, (r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,esk2_0)),k17_euclid(X2))|~r2_hidden(X2,k3_goboard5(esk3_0,X1,esk2_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,k5_numbers)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_379, c_0_162]), c_0_163]), c_0_53])])])).
cnf(c_0_383, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))!=k18_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_380, c_0_381])])).
cnf(c_0_384, negated_conjecture, (r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(X1))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_382, c_0_180]), c_0_181]), c_0_59])])])).
cnf(c_0_385, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_383, c_0_362]), c_0_323])).
cnf(c_0_386, plain, (r1_xxreal_0(k17_euclid(X1),k17_euclid(X2))|~r1_xxreal_0(k17_euclid(X1),k17_euclid(X3))|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),X1,X3))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(split_conjunct,[status(thm)],[c_0_174])).
cnf(c_0_387, negated_conjecture, (r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(X1))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))), inference(csr,[status(thm)],[c_0_384, c_0_196])).
cnf(c_0_388, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_274, c_0_385])])).
cnf(c_0_389, plain, (r1_xxreal_0(k17_euclid(X1),k17_euclid(X2))|~r1_xxreal_0(k17_euclid(X1),k17_euclid(X3))|~m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),X1,X3))), inference(evalgc,[status(thm)],[c_0_386])).
cnf(c_0_390, negated_conjecture, (r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_210])])).
cnf(c_0_391, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_351, c_0_388]), c_0_99])]), c_0_100])])).
cnf(c_0_392, negated_conjecture, (r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|v2_struct_0(k15_euclid(np__2))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_389, c_0_169]), c_0_125]), c_0_152])])])).
cnf(c_0_393, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|~v1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_390]), c_0_207])])])).
cnf(c_0_394, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_391]), c_0_99])]), c_0_100])])).
cnf(c_0_395, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_392, c_0_217])])).
cnf(c_0_396, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_394]), c_0_229])])])).
cnf(c_0_397, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_228])])).
cnf(c_0_398, negated_conjecture, (r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_387, c_0_235])])).
cnf(c_0_399, negated_conjecture, (r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_389, c_0_224]), c_0_152]), c_0_203]), c_0_125])])])).
cnf(c_0_400, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_395, c_0_224]), c_0_203])])])).
cnf(c_0_401, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_397])])).
cnf(c_0_402, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|~v1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_398]), c_0_192])])])).
cnf(c_0_403, negated_conjecture, (r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_399, c_0_400])])).
cnf(c_0_404, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_401]), c_0_229])])])).
cnf(c_0_405, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_390, c_0_394])])).
cnf(c_0_406, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_394]), c_0_229])])])).
cnf(c_0_407, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_396, c_0_403])])).
cnf(c_0_408, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_359, c_0_274])])).
cnf(c_0_409, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_404, c_0_405])])).
cnf(c_0_410, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_407])])).
cnf(c_0_411, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_408, c_0_409]), c_0_203])])])).
cnf(c_0_412, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_410])])).
cnf(c_0_413, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)), inference(csr,[status(thm)],[c_0_411, c_0_290])).
cnf(c_0_414, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_412]), c_0_152]), c_0_125])])])).
cnf(c_0_415, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_243])])).
cnf(c_0_416, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_413, c_0_414])])).
cnf(c_0_417, negated_conjecture, (X1=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(X1)!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_125])])).
cnf(c_0_418, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_415]), c_0_229])])])).
cnf(c_0_419, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_398, c_0_394])])).
cnf(c_0_420, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))!=k17_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_416, c_0_410])])).
cnf(c_0_421, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_417, c_0_274])])).
cnf(c_0_422, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_418, c_0_419])])).
cnf(c_0_423, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(csr,[status(thm)],[c_0_420, c_0_394])).
cnf(c_0_424, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_421, c_0_422]), c_0_203])])])).
cnf(c_0_425, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_313, c_0_423])])).
cnf(c_0_426, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_320, c_0_401])])).
cnf(c_0_427, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)), inference(csr,[status(thm)],[c_0_424, c_0_318])).
cnf(c_0_428, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_425]), c_0_152]), c_0_125])])])).
cnf(c_0_429, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_408, c_0_426]), c_0_203])])])).
cnf(c_0_430, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_427, c_0_428])])).
cnf(c_0_431, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)), inference(csr,[status(thm)],[c_0_429, c_0_290])).
cnf(c_0_432, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))!=k17_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_423])])).
cnf(c_0_433, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_431, c_0_341])])).
cnf(c_0_434, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(csr,[status(thm)],[c_0_432, c_0_394])).
cnf(c_0_435, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))!=k17_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_433, c_0_434])])).
cnf(c_0_436, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_423])])).
cnf(c_0_437, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_435, c_0_394]), c_0_436])).
cnf(c_0_438, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_351, c_0_437]), c_0_99])]), c_0_100])])).
cnf(c_0_439, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_438]), c_0_99])]), c_0_100])])).
cnf(c_0_440, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_368, c_0_439]), c_0_304])])])).
cnf(c_0_441, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_439]), c_0_304])])])).
cnf(c_0_442, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_440, c_0_375])])).
cnf(c_0_443, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_441, c_0_442])])).
cnf(c_0_444, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_413, c_0_443])])).
cnf(c_0_445, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))!=k17_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_444, c_0_434])])).
cnf(c_0_446, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_445, c_0_394]), c_0_436])).
cnf(c_0_447, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_351, c_0_446]), c_0_99])]), c_0_100])])).
cnf(c_0_448, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_447]), c_0_99])]), c_0_100])])).
cnf(c_0_449, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_448]), c_0_229])])])).
cnf(c_0_450, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_448]), c_0_229])])])).
cnf(c_0_451, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_449, c_0_403])])).
cnf(c_0_452, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_450, c_0_451])])).
cnf(c_0_453, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_452])])).
cnf(c_0_454, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_452])])).
cnf(c_0_455, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_453]), c_0_152]), c_0_125])])])).
cnf(c_0_456, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_454, c_0_455])])).
cnf(c_0_457, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_452, c_0_456])])).
cnf(c_0_458, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_313, c_0_457])])).
cnf(c_0_459, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_318, c_0_457])])).
cnf(c_0_460, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_291, c_0_458]), c_0_152]), c_0_125])])])).
cnf(c_0_461, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_459, c_0_460])])).
cnf(c_0_462, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_368, c_0_448]), c_0_304])])])).
cnf(c_0_463, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_290, c_0_461])])).
cnf(c_0_464, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_361, c_0_448]), c_0_304])])])).
cnf(c_0_465, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_462, c_0_375])])).
cnf(c_0_466, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)), inference(csr,[status(thm)],[c_0_463, c_0_457])).
cnf(c_0_467, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_464, c_0_465])])).
cnf(c_0_468, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_466, c_0_467])])).
fof(c_0_469, plain, ![X32]:((~v2_struct_0(k15_euclid(X32))|~v7_ordinal1(X32))&(v5_rltopsp1(k15_euclid(X32))|~v7_ordinal1(X32))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc5_euclid])])])])).
cnf(c_0_470, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(csr,[status(thm)],[c_0_468, c_0_459])).
cnf(c_0_471, plain, (~v2_struct_0(k15_euclid(X1))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_469])).
cnf(c_0_472, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_351, c_0_470]), c_0_99])]), c_0_100])])).
cnf(c_0_473, plain, (~v7_ordinal1(X1)|~v2_struct_0(k15_euclid(X1))), inference(evalgc,[status(thm)],[c_0_471])).
cnf(c_0_474, negated_conjecture, (v2_struct_0(k15_euclid(np__2))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_472]), c_0_99])]), c_0_100])])).
cnf(c_0_475, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_473, c_0_474]), c_0_82])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 476
# Proof object clause steps            : 415
# Proof object formula steps           : 61
# Proof object conjectures             : 298
# Proof object clause conjectures      : 295
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 69
# Proof object initial formulas used   : 31
# Proof object generating inferences   : 248
# Proof object simplifying inferences  : 366
# Training examples: 316 positive, 1066 negative
# Training: Positive examples begin
cnf(c_0_146, negated_conjecture, (~v3_relat_1(esk3_0))).# trainpos
cnf(c_0_51, negated_conjecture, (v1_matrix_1(esk3_0))).# trainpos
cnf(c_0_145, negated_conjecture, (v2_goboard1(esk3_0))).# trainpos
cnf(c_0_144, negated_conjecture, (v3_goboard1(esk3_0))).# trainpos
cnf(c_0_143, negated_conjecture, (v4_goboard1(esk3_0))).# trainpos
cnf(c_0_142, negated_conjecture, (v5_goboard1(esk3_0))).# trainpos
cnf(c_0_100, negated_conjecture, (~r1_sppol_1(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))).# trainpos
cnf(c_0_59, negated_conjecture, (m1_subset_1(esk1_0,k5_numbers))).# trainpos
cnf(c_0_53, negated_conjecture, (m1_subset_1(esk2_0,k5_numbers))).# trainpos
cnf(c_0_181, negated_conjecture, (r1_xxreal_0(np__1,esk1_0))).# trainpos
cnf(c_0_163, negated_conjecture, (r1_xxreal_0(np__1,esk2_0))).# trainpos
cnf(c_0_473, plain, (~v7_ordinal1(X1)|~v2_struct_0(k15_euclid(X1)))).# trainpos
cnf(c_0_45, plain, (k4_ordinal1=k5_numbers)).# trainpos
cnf(c_0_77, plain, (m2_subset_1(np__2,k1_numbers,k5_numbers))).# trainpos
cnf(c_0_213, negated_conjecture, (r2_hidden(esk4_0,k16_matrix_1(esk3_0)))).# trainpos
cnf(c_0_203, negated_conjecture, (m1_subset_1(esk4_0,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_42, negated_conjecture, (m2_finseq_1(esk3_0,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_180, negated_conjecture, (r1_xxreal_0(k2_nat_1(esk1_0,np__1),k3_finseq_1(esk3_0)))).# trainpos
cnf(c_0_162, negated_conjecture, (r1_xxreal_0(k2_nat_1(esk2_0,np__1),k1_matrix_1(esk3_0)))).# trainpos
cnf(c_0_99, negated_conjecture, (r2_hidden(esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))).# trainpos
cnf(c_0_70, plain, (m1_subset_1(np__2,k5_numbers))).# trainpos
cnf(c_0_176, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).# trainpos
cnf(c_0_185, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_101, plain, (l1_rlvect_1(X1)|~l1_rltopsp1(X1))).# trainpos
cnf(c_0_159, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).# trainpos
cnf(c_0_44, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).# trainpos
cnf(c_0_291, plain, (k17_euclid(X1)=k17_euclid(X2)|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,X2)))).# trainpos
cnf(c_0_110, plain, (v13_algstr_0(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_109, plain, (v2_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_108, plain, (v3_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_219, plain, (X1=X2|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2))).# trainpos
cnf(c_0_205, plain, (r1_xxreal_0(k17_euclid(X1),k17_euclid(X2))|~r1_xxreal_0(k17_euclid(X3),k17_euclid(X2))|~m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),X3,X2)))).# trainpos
cnf(c_0_107, plain, (v4_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_106, plain, (v5_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_105, plain, (v6_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_104, plain, (v7_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_103, plain, (v8_rlvect_1(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_81, plain, (l1_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_271, plain, (X1=X2|k17_euclid(X1)!=k17_euclid(X2)|k18_euclid(X1)!=k18_euclid(X2)|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_197, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_389, plain, (r1_xxreal_0(k17_euclid(X1),k17_euclid(X2))|~r1_xxreal_0(k17_euclid(X1),k17_euclid(X3))|~m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),X1,X3)))).# trainpos
cnf(c_0_41, plain, (m1_finseq_1(X1,X2)|~m2_finseq_1(X1,X2))).# trainpos
cnf(c_0_140, plain, (m1_subset_1(k17_euclid(X1),k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_245, plain, (m1_subset_1(k18_euclid(X1),k1_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_149, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_278, plain, (r1_xxreal_0(k18_euclid(X1),k18_euclid(X2))|~r1_xxreal_0(k18_euclid(X3),k18_euclid(X2))|~m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),X3,X2)))).# trainpos
cnf(c_0_353, plain, (r1_xxreal_0(k18_euclid(X1),k18_euclid(X2))|~r1_xxreal_0(k18_euclid(X1),k18_euclid(X3))|~m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),X1,X3)))).# trainpos
cnf(c_0_330, plain, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,X2))|k17_euclid(X1)!=k17_euclid(X2)|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_344, plain, (r1_sppol_1(X2,X1,X3)|esk5_3(X2,X1,X3)!=X1|~r2_hidden(X1,X3)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X2)))))).# trainpos
cnf(c_0_339, plain, (r1_sppol_1(X2,X1,X3)|esk6_3(X2,X1,X3)!=X1|~r2_hidden(X1,X3)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X2)))))).# trainpos
cnf(c_0_191, plain, (r1_xxreal_0(X2,X1)|r1_xxreal_0(X1,X2)|~v1_xxreal_0(X2)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_64, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_49, plain, (m1_subset_1(k3_goboard5(X1,X2,X3),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v1_matrix_1(X1)|~v7_ordinal1(X3)|~v7_ordinal1(X2)|~m1_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_247, plain, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,k2_nat_1(X3,np__1)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3))|v3_relat_1(X1)|~v1_matrix_1(X1)|~v2_goboard1(X1)|~v3_goboard1(X1)|~v4_goboard1(X1)|~v5_goboard1(X1)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_300, plain, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1),X3))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3))|v3_relat_1(X1)|~v1_matrix_1(X1)|~v2_goboard1(X1)|~v3_goboard1(X1)|~v4_goboard1(X1)|~v5_goboard1(X1)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_267, plain, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1),k2_nat_1(X3,np__1)))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,k2_nat_1(X3,np__1)))|v3_relat_1(X1)|~v1_matrix_1(X1)|~v2_goboard1(X1)|~v3_goboard1(X1)|~v4_goboard1(X1)|~v5_goboard1(X1)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_186, plain, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1),k2_nat_1(X3,np__1)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1),X3))|v3_relat_1(X1)|~v1_matrix_1(X1)|~v2_goboard1(X1)|~v3_goboard1(X1)|~v4_goboard1(X1)|~v5_goboard1(X1)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_124, plain, (v2_struct_0(X1)|r2_hidden(X2,k1_rltopsp1(X1,X2,X3))|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~l1_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_63, plain, (r1_sppol_1(X1,X2,X3)|m1_subset_1(esk5_3(X1,X2,X3),u1_struct_0(k15_euclid(X1)))|~r2_hidden(X2,X3)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))).# trainpos
cnf(c_0_91, plain, (r1_sppol_1(X1,X2,X3)|m1_subset_1(esk6_3(X1,X2,X3),u1_struct_0(k15_euclid(X1)))|~r2_hidden(X2,X3)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))).# trainpos
cnf(c_0_252, plain, (k18_euclid(X1)=k18_euclid(X4)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X2,X3))|k17_euclid(X1)!=k17_euclid(X4)|~m1_subset_1(X4,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X4,k1_rltopsp1(k15_euclid(np__2),X2,X3))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),X2,X3)))).# trainpos
cnf(c_0_376, plain, (v3_relat_1(X1)|r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k17_euclid(X4))|~v1_matrix_1(X1)|~v2_goboard1(X1)|~v3_goboard1(X1)|~v4_goboard1(X1)|~v5_goboard1(X1)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r2_hidden(X4,k3_goboard5(X1,X2,X3))|~m1_subset_1(X4,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_136, plain, (k1_rltopsp1(X1,X2,X3)=k1_rltopsp1(X1,X3,X2)|v2_struct_0(X1)|~v13_algstr_0(X1)|~v2_rlvect_1(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v5_rlvect_1(X1)|~v6_rlvect_1(X1)|~v7_rlvect_1(X1)|~v8_rlvect_1(X1)|~l1_rlvect_1(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_156, plain, (X4=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,X1,X2)|X4=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,X1,k2_nat_1(X2,np__1))|X4=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,k2_nat_1(X1,np__1),X2)|X4=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X3,k2_nat_1(X1,np__1),k2_nat_1(X2,np__1))|v3_relat_1(X3)|~v1_matrix_1(X3)|~v2_goboard1(X3)|~v3_goboard1(X3)|~v4_goboard1(X3)|~v5_goboard1(X3)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~r2_hidden(X4,k16_matrix_1(X3))|~r2_hidden(X4,k3_goboard5(X3,X1,X2))|~m1_subset_1(X4,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k2_nat_1(X2,np__1),k1_matrix_1(X3))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(X3))|~m2_finseq_1(X3,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_327, plain, (v3_relat_1(X1)|r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k18_euclid(X4))|~v1_matrix_1(X1)|~v2_goboard1(X1)|~v3_goboard1(X1)|~v4_goboard1(X1)|~v5_goboard1(X1)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~r2_hidden(X4,k3_goboard5(X1,X2,X3))|~m1_subset_1(X4,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_249, plain, (v3_relat_1(X2)|r1_xxreal_0(k18_euclid(X1),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X2,X3,k2_nat_1(X4,np__1))))|~v1_matrix_1(X2)|~v2_goboard1(X2)|~v3_goboard1(X2)|~v4_goboard1(X2)|~v5_goboard1(X2)|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X3,k5_numbers)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r2_hidden(X1,k3_goboard5(X2,X3,X4))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k2_nat_1(X4,np__1),k1_matrix_1(X2))|~r1_xxreal_0(k2_nat_1(X3,np__1),k3_finseq_1(X2))|~m2_finseq_1(X2,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_141, plain, (v3_relat_1(X2)|r1_xxreal_0(k17_euclid(X1),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X2,k2_nat_1(X3,np__1),X4)))|~v1_matrix_1(X2)|~v2_goboard1(X2)|~v3_goboard1(X2)|~v4_goboard1(X2)|~v5_goboard1(X2)|~m1_subset_1(X4,k5_numbers)|~m1_subset_1(X3,k5_numbers)|~r1_xxreal_0(np__1,X4)|~r1_xxreal_0(np__1,X3)|~r2_hidden(X1,k3_goboard5(X2,X3,X4))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k2_nat_1(X4,np__1),k1_matrix_1(X2))|~r1_xxreal_0(k2_nat_1(X3,np__1),k3_finseq_1(X2))|~m2_finseq_1(X2,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_189, plain, (r1_sppol_1(X2,X1,X3)|r2_hidden(X1,k1_rltopsp1(k15_euclid(X2),esk5_3(X2,X1,X3),esk6_3(X2,X1,X3)))|~r2_hidden(X1,X3)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X2)))))).# trainpos
cnf(c_0_73, plain, (r1_sppol_1(X1,X2,X3)|r1_tarski(k1_rltopsp1(k15_euclid(X1),esk5_3(X1,X2,X3),esk6_3(X1,X2,X3)),X3)|~r2_hidden(X2,X3)|~m2_subset_1(X1,k1_numbers,k5_numbers)|~m1_subset_1(X2,u1_struct_0(k15_euclid(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(k15_euclid(X1)))))).# trainpos
cnf(c_0_199, negated_conjecture, (~v1_xboole_0(k3_goboard5(esk3_0,esk1_0,esk2_0)))).# trainpos
cnf(c_0_253, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,k2_nat_1(X2,np__1)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,X2))|~r1_xxreal_0(k2_nat_1(X2,np__1),k1_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_58, negated_conjecture, (v7_ordinal1(esk2_0))).# trainpos
cnf(c_0_66, negated_conjecture, (v7_ordinal1(esk1_0))).# trainpos
cnf(c_0_82, plain, (v7_ordinal1(np__2))).# trainpos
cnf(c_0_134, plain, (v13_algstr_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_133, plain, (v2_rlvect_1(k15_euclid(np__2)))).# trainpos
cnf(c_0_132, plain, (v3_rlvect_1(k15_euclid(np__2)))).# trainpos
cnf(c_0_131, plain, (v4_rlvect_1(k15_euclid(np__2)))).# trainpos
cnf(c_0_130, plain, (v5_rlvect_1(k15_euclid(np__2)))).# trainpos
cnf(c_0_281, negated_conjecture, (X1=esk4_0|k17_euclid(X1)!=k17_euclid(esk4_0)|k18_euclid(X1)!=k18_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_263, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,k2_nat_1(esk2_0,np__1)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,esk2_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_129, plain, (v6_rlvect_1(k15_euclid(np__2)))).# trainpos
cnf(c_0_128, plain, (v7_rlvect_1(k15_euclid(np__2)))).# trainpos
cnf(c_0_127, plain, (v8_rlvect_1(k15_euclid(np__2)))).# trainpos
cnf(c_0_102, plain, (l1_rltopsp1(k15_euclid(np__2)))).# trainpos
cnf(c_0_126, plain, (l1_rlvect_1(k15_euclid(np__2)))).# trainpos
cnf(c_0_50, negated_conjecture, (m1_finseq_1(esk3_0,k3_finseq_2(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_309, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X1,np__1),X2))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,X2))|~r1_xxreal_0(k2_nat_1(X2,np__1),k1_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_274, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).# trainpos
cnf(c_0_315, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X1,np__1),esk2_0))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,esk2_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_320, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).# trainpos
cnf(c_0_57, negated_conjecture, (m1_subset_1(k3_goboard5(esk3_0,X1,X2),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v7_ordinal1(X2)|~v7_ordinal1(X1))).# trainpos
cnf(c_0_277, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X1,np__1),k2_nat_1(X2,np__1)))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,k2_nat_1(X2,np__1)))|~r1_xxreal_0(k2_nat_1(X2,np__1),k1_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_65, negated_conjecture, (m1_subset_1(k3_goboard5(esk3_0,X1,esk2_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v7_ordinal1(X1))).# trainpos
cnf(c_0_214, negated_conjecture, (m1_subset_1(k17_euclid(esk4_0),k1_numbers))).# trainpos
cnf(c_0_286, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X1,np__1),k2_nat_1(esk2_0,np__1)))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,k2_nat_1(esk2_0,np__1)))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_223, negated_conjecture, (v1_xreal_0(k17_euclid(esk4_0)))).# trainpos
cnf(c_0_229, negated_conjecture, (v1_xxreal_0(k17_euclid(esk4_0)))).# trainpos
cnf(c_0_295, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))).# trainpos
cnf(c_0_336, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,esk4_0))|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_76, negated_conjecture, (m1_subset_1(k3_goboard5(esk3_0,esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).# trainpos
cnf(c_0_196, negated_conjecture, (m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0)))).# trainpos
cnf(c_0_287, negated_conjecture, (m1_subset_1(k18_euclid(esk4_0),k1_numbers))).# trainpos
cnf(c_0_296, negated_conjecture, (v1_xreal_0(k18_euclid(esk4_0)))).# trainpos
cnf(c_0_304, negated_conjecture, (v1_xxreal_0(k18_euclid(esk4_0)))).# trainpos
cnf(c_0_350, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|esk5_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_345, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|esk6_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_98, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|m1_subset_1(esk5_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0)))).# trainpos
cnf(c_0_200, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X1,np__1),k2_nat_1(X2,np__1)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X1,np__1),X2))|~r1_xxreal_0(k2_nat_1(X2,np__1),k1_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_125, negated_conjecture, (m1_subset_1(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_318, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0))).# trainpos
cnf(c_0_417, negated_conjecture, (X1=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(X1)!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_421, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_211, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X1,np__1),k2_nat_1(esk2_0,np__1)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X1,np__1),esk2_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_221, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))).# trainpos
cnf(c_0_135, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|m1_subset_1(esk6_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0)))).# trainpos
cnf(c_0_152, negated_conjecture, (m1_subset_1(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_290, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0))).# trainpos
cnf(c_0_342, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).# trainpos
cnf(c_0_359, negated_conjecture, (X1=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(X1)!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_408, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_151, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),X1,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_218, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),X1,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_153, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),X1,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),X1)|v2_struct_0(k15_euclid(np__2))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_160, negated_conjecture, (m1_subset_1(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_numbers))).# trainpos
cnf(c_0_177, negated_conjecture, (v1_xreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_192, negated_conjecture, (v1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_206, negated_conjecture, (r1_xxreal_0(X1,k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),X1)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_250, negated_conjecture, (m1_subset_1(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_numbers))).# trainpos
cnf(c_0_259, negated_conjecture, (v1_xreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_269, negated_conjecture, (v1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_279, negated_conjecture, (r1_xxreal_0(X1,k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),X1)|~v1_xxreal_0(X1))).# trainpos
cnf(c_0_178, negated_conjecture, (m1_subset_1(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_numbers))).# trainpos
cnf(c_0_193, negated_conjecture, (v1_xreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_207, negated_conjecture, (v1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_260, negated_conjecture, (m1_subset_1(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_numbers))).# trainpos
cnf(c_0_270, negated_conjecture, (v1_xreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_280, negated_conjecture, (v1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_169, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_123, negated_conjecture, (r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0)))).# trainpos
cnf(c_0_215, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0)))).# trainpos
cnf(c_0_392, negated_conjecture, (r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|v2_struct_0(k15_euclid(np__2))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_358, negated_conjecture, (r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|v2_struct_0(k15_euclid(np__2))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_216, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_288, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_224, negated_conjecture, (r2_hidden(esk4_0,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_230, negated_conjecture, (r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_305, negated_conjecture, (r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_399, negated_conjecture, (r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_369, negated_conjecture, (r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_307, negated_conjecture, (k18_euclid(X1)=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(X1)!=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_379, negated_conjecture, (r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,X2)),k17_euclid(X3))|~r2_hidden(X3,k3_goboard5(esk3_0,X1,X2))|~r1_xxreal_0(k2_nat_1(X2,np__1),k1_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_382, negated_conjecture, (r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,esk2_0)),k17_euclid(X2))|~r2_hidden(X2,k3_goboard5(esk3_0,X1,esk2_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_384, negated_conjecture, (r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(X1))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_331, negated_conjecture, (r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,X2)),k18_euclid(X3))|~r2_hidden(X3,k3_goboard5(esk3_0,X1,X2))|~r1_xxreal_0(k2_nat_1(X2,np__1),k1_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X2)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X3,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_337, negated_conjecture, (r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X1,esk2_0)),k18_euclid(X2))|~r2_hidden(X2,k3_goboard5(esk3_0,X1,esk2_0))|~r1_xxreal_0(k2_nat_1(X1,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,k5_numbers))).# trainpos
cnf(c_0_343, negated_conjecture, (r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(X1))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_150, negated_conjecture, (r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k3_goboard5(esk3_0,esk1_0,esk2_0)))).# trainpos
cnf(c_0_256, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X2,k2_nat_1(X3,np__1))))|~r2_hidden(X1,k3_goboard5(esk3_0,X2,X3))|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_266, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X2,k2_nat_1(esk2_0,np__1))))|~r2_hidden(X1,k3_goboard5(esk3_0,X2,esk2_0))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_276, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_161, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X2,np__1),X3)))|~r2_hidden(X1,k3_goboard5(esk3_0,X2,X3))|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_179, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X2,np__1),esk2_0)))|~r2_hidden(X1,k3_goboard5(esk3_0,X2,esk2_0))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_195, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_172, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X2,np__1),k2_nat_1(X3,np__1))|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X2,np__1),X3)|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X2,k2_nat_1(X3,np__1))|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X2,X3)|~r2_hidden(X1,k3_goboard5(esk3_0,X2,X3))|~r2_hidden(X1,k16_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_187, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X2,np__1),k2_nat_1(esk2_0,np__1))|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X2,k2_nat_1(esk2_0,np__1))|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(X2,np__1),esk2_0)|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,X2,esk2_0)|~r2_hidden(X1,k3_goboard5(esk3_0,X2,esk2_0))|~r2_hidden(X1,k16_matrix_1(esk3_0))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,X2)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X2,k5_numbers))).# trainpos
cnf(c_0_201, negated_conjecture, (X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))|X1=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k16_matrix_1(esk3_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_222, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).# trainpos
cnf(c_0_303, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).# trainpos
cnf(c_0_228, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).# trainpos
cnf(c_0_167, negated_conjecture, (m1_subset_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_zfmisc_1(k3_goboard5(esk3_0,esk1_0,esk2_0))))).# trainpos
cnf(c_0_183, negated_conjecture, (m1_subset_1(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_217, negated_conjecture, (r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_395, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_225, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_289, negated_conjecture, (r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_364, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_297, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_168, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_262, negated_conjecture, (k18_euclid(X1)=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_273, negated_conjecture, (k18_euclid(X1)=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_283, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).# trainpos
cnf(c_0_226, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_313, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).# trainpos
cnf(c_0_232, negated_conjecture, (m1_subset_1(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_235, negated_conjecture, (r2_hidden(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_398, negated_conjecture, (r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_402, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|~v1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))))).# trainpos
cnf(c_0_355, negated_conjecture, (r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_361, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|~v1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))))).# trainpos
cnf(c_0_301, negated_conjecture, (r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_310, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|~v1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_316, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_367, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).# trainpos
cnf(c_0_238, negated_conjecture, (r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_255, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|~v1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_243, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).# trainpos
cnf(c_0_184, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_198, negated_conjecture, (m1_subset_1(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_210, negated_conjecture, (r2_hidden(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_390, negated_conjecture, (r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_393, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|~v1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))))).# trainpos
cnf(c_0_363, negated_conjecture, (r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_368, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|~v1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))))).# trainpos
cnf(c_0_294, negated_conjecture, (r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_302, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|~v1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_311, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_220, negated_conjecture, (r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_227, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))=k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|~v1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_233, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_397, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).# trainpos
cnf(c_0_400, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).# trainpos
cnf(c_0_403, negated_conjecture, (r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).# trainpos
cnf(c_0_231, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_234, negated_conjecture, (r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_237, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_242, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).# trainpos
cnf(c_0_265, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_254, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_264, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_284, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_298, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0))).# trainpos
cnf(c_0_370, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).# trainpos
cnf(c_0_375, negated_conjecture, (r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).# trainpos
cnf(c_0_306, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_312, negated_conjecture, (r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_317, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_321, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0))).# trainpos
cnf(c_0_325, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_335, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_329, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_334, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_292, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).# trainpos
cnf(c_0_299, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_308, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).# trainpos
cnf(c_0_314, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0)).# trainpos
cnf(c_0_323, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0))).# trainpos
cnf(c_0_319, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).# trainpos
cnf(c_0_324, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_328, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).# trainpos
cnf(c_0_333, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0))).# trainpos
cnf(c_0_348, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).# trainpos
cnf(c_0_360, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_346, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0))).# trainpos
cnf(c_0_357, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_362, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_366, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).# trainpos
cnf(c_0_374, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).# trainpos
cnf(c_0_373, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0)).# trainpos
cnf(c_0_365, negated_conjecture, (X1=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|k18_euclid(X1)!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_372, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).# trainpos
cnf(c_0_377, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).# trainpos
cnf(c_0_378, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).# trainpos
cnf(c_0_381, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0))).# trainpos
cnf(c_0_383, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))!=k18_euclid(esk4_0))).# trainpos
cnf(c_0_388, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_391, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_394, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_406, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).# trainpos
cnf(c_0_396, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).# trainpos
cnf(c_0_415, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0)).# trainpos
cnf(c_0_401, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0)).# trainpos
cnf(c_0_426, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_429, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).# trainpos
cnf(c_0_433, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).# trainpos
cnf(c_0_418, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_404, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).# trainpos
cnf(c_0_419, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).# trainpos
cnf(c_0_422, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_424, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).# trainpos
cnf(c_0_405, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).# trainpos
cnf(c_0_409, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_411, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).# trainpos
cnf(c_0_407, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).# trainpos
cnf(c_0_410, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0))).# trainpos
cnf(c_0_412, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).# trainpos
cnf(c_0_414, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_416, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).# trainpos
cnf(c_0_420, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))!=k17_euclid(esk4_0))).# trainpos
cnf(c_0_436, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0))).# trainpos
cnf(c_0_425, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0)).# trainpos
cnf(c_0_428, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_430, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).# trainpos
cnf(c_0_432, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))!=k17_euclid(esk4_0))).# trainpos
cnf(c_0_435, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))!=k17_euclid(esk4_0))).# trainpos
cnf(c_0_438, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_439, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_441, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).# trainpos
cnf(c_0_440, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).# trainpos
cnf(c_0_442, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).# trainpos
cnf(c_0_443, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0))).# trainpos
cnf(c_0_444, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).# trainpos
cnf(c_0_445, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))!=k17_euclid(esk4_0))).# trainpos
cnf(c_0_447, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_448, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_450, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).# trainpos
cnf(c_0_464, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).# trainpos
cnf(c_0_449, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).# trainpos
cnf(c_0_462, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).# trainpos
cnf(c_0_451, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).# trainpos
cnf(c_0_452, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0))).# trainpos
cnf(c_0_454, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0))).# trainpos
cnf(c_0_465, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).# trainpos
cnf(c_0_467, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0))).# trainpos
cnf(c_0_453, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0))).# trainpos
cnf(c_0_455, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_456, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).# trainpos
cnf(c_0_457, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0)).# trainpos
cnf(c_0_459, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0))).# trainpos
cnf(c_0_458, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0)).# trainpos
cnf(c_0_460, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_461, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).# trainpos
cnf(c_0_463, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0))).# trainpos
cnf(c_0_468, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0))).# trainpos
cnf(c_0_472, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2)))).# trainpos
cnf(c_0_474, negated_conjecture, (v2_struct_0(k15_euclid(np__2)))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_31, plain, (k1_xboole_0=o_0_0_xboole_0)).#trainneg
cnf(i_0_39, plain, (v1_xboole_0(o_0_0_xboole_0))).#trainneg
cnf(i_0_59, plain, (v2_xxreal_0(np__2))).#trainneg
cnf(i_0_85, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_56, plain, (m1_subset_1(np__2,k1_numbers))).#trainneg
cnf(i_0_73, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_54, plain, (m2_finseq_1(X1,X2)|~m1_finseq_1(X1,X2))).#trainneg
cnf(i_0_77, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_37, plain, (l1_pre_topc(X1)|~l1_rltopsp1(X1))).#trainneg
cnf(i_0_83, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_33, plain, (v5_rltopsp1(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_51, plain, (v2_pre_topc(k15_euclid(X1))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_30, plain, (r2_hidden(X1,X2)|~r1_sppol_1(X3,X1,X2)|~m2_subset_1(X3,k1_numbers,k5_numbers)|~m1_subset_1(X1,u1_struct_0(k15_euclid(X3)))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(X3)))))).#trainneg
cnf(i_0_29, plain, (X3=X4|X1=X4|~r1_sppol_1(X2,X4,X5)|~m2_subset_1(X2,k1_numbers,k5_numbers)|~m1_subset_1(X4,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X3,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(X2)))|~m1_subset_1(X5,k1_zfmisc_1(u1_struct_0(k15_euclid(X2))))|~r2_hidden(X4,k1_rltopsp1(k15_euclid(X2),X1,X3))|~r1_tarski(k1_rltopsp1(k15_euclid(X2),X1,X3),X5))).#trainneg
cnf(i_0_66, plain, (v3_relat_1(X1)|r2_hidden(X4,k3_goboard5(X1,X2,X3))|~v1_matrix_1(X1)|~v2_goboard1(X1)|~v3_goboard1(X1)|~v4_goboard1(X1)|~v5_goboard1(X1)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers)|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~m1_subset_1(X4,u1_struct_0(k15_euclid(np__2)))|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k17_euclid(X4))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k18_euclid(X4))|~r1_xxreal_0(k18_euclid(X4),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,k2_nat_1(X3,np__1))))|~r1_xxreal_0(k17_euclid(X4),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1),X3))))).#trainneg
cnf(i_0_89, negated_conjecture, (~v1_xboole_0(k16_matrix_1(esk3_0)))).#trainneg
cnf(i_0_98, negated_conjecture, (esk2_0=np__1|~v1_xxreal_0(np__1)|~v1_xxreal_0(esk2_0)|~r1_xxreal_0(esk2_0,np__1))).#trainneg
cnf(i_0_99, negated_conjecture, (esk1_0=np__1|~v1_xxreal_0(np__1)|~v1_xxreal_0(esk1_0)|~r1_xxreal_0(esk1_0,np__1))).#trainneg
cnf(i_0_96, negated_conjecture, (k2_nat_1(esk2_0,np__1)=k1_matrix_1(esk3_0)|~v1_xxreal_0(k2_nat_1(esk2_0,np__1))|~v1_xxreal_0(k1_matrix_1(esk3_0))|~r1_xxreal_0(k1_matrix_1(esk3_0),k2_nat_1(esk2_0,np__1)))).#trainneg
cnf(i_0_97, negated_conjecture, (k2_nat_1(esk1_0,np__1)=k3_finseq_1(esk3_0)|~v1_xxreal_0(k2_nat_1(esk1_0,np__1))|~v1_xxreal_0(k3_finseq_1(esk3_0))|~r1_xxreal_0(k3_finseq_1(esk3_0),k2_nat_1(esk1_0,np__1)))).#trainneg
cnf(i_0_104, plain, (v1_xboole_0(k1_numbers)|r2_hidden(np__2,k1_numbers))).#trainneg
cnf(i_0_90, plain, (v1_xreal_0(np__2))).#trainneg
cnf(i_0_124, plain, (v1_xxreal_0(np__2))).#trainneg
cnf(i_0_102, negated_conjecture, (v1_xboole_0(k5_numbers)|r2_hidden(esk2_0,k5_numbers))).#trainneg
cnf(i_0_155, plain, (v5_rltopsp1(k15_euclid(np__2)))).#trainneg
cnf(i_0_157, plain, (v2_pre_topc(k15_euclid(np__2)))).#trainneg
cnf(i_0_163, plain, (l1_pre_topc(k15_euclid(np__2)))).#trainneg
cnf(i_0_95, negated_conjecture, (m1_subset_1(esk4_0,k16_matrix_1(esk3_0)))).#trainneg
cnf(i_0_125, negated_conjecture, (v13_algstr_0(k15_euclid(esk2_0)))).#trainneg
cnf(i_0_103, negated_conjecture, (v1_xboole_0(k5_numbers)|r2_hidden(esk1_0,k5_numbers))).#trainneg
cnf(i_0_94, negated_conjecture, (m1_subset_1(esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_105, plain, (v1_xboole_0(k5_numbers)|r2_hidden(np__2,k5_numbers))).#trainneg
cnf(i_0_126, negated_conjecture, (v2_rlvect_1(k15_euclid(esk2_0)))).#trainneg
cnf(i_0_127, negated_conjecture, (v3_rlvect_1(k15_euclid(esk2_0)))).#trainneg
cnf(i_0_128, negated_conjecture, (v4_rlvect_1(k15_euclid(esk2_0)))).#trainneg
cnf(i_0_129, negated_conjecture, (v5_rlvect_1(k15_euclid(esk2_0)))).#trainneg
cnf(i_0_130, negated_conjecture, (v6_rlvect_1(k15_euclid(esk2_0)))).#trainneg
cnf(i_0_131, negated_conjecture, (v7_rlvect_1(k15_euclid(esk2_0)))).#trainneg
cnf(i_0_132, negated_conjecture, (v8_rlvect_1(k15_euclid(esk2_0)))).#trainneg
cnf(i_0_133, negated_conjecture, (v5_rltopsp1(k15_euclid(esk2_0)))).#trainneg
cnf(i_0_158, plain, (r1_xxreal_0(X1,np__2)|r1_xxreal_0(np__2,X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_173, plain, (r1_xxreal_0(np__2,np__2))).#trainneg
cnf(i_0_134, negated_conjecture, (l1_rltopsp1(k15_euclid(esk2_0)))).#trainneg
cnf(i_0_135, negated_conjecture, (v2_pre_topc(k15_euclid(esk2_0)))).#trainneg
cnf(i_0_136, negated_conjecture, (v13_algstr_0(k15_euclid(esk1_0)))).#trainneg
cnf(i_0_137, negated_conjecture, (v2_rlvect_1(k15_euclid(esk1_0)))).#trainneg
cnf(i_0_138, negated_conjecture, (v3_rlvect_1(k15_euclid(esk1_0)))).#trainneg
cnf(i_0_139, negated_conjecture, (v4_rlvect_1(k15_euclid(esk1_0)))).#trainneg
cnf(i_0_140, negated_conjecture, (v5_rlvect_1(k15_euclid(esk1_0)))).#trainneg
cnf(i_0_141, negated_conjecture, (v6_rlvect_1(k15_euclid(esk1_0)))).#trainneg
cnf(i_0_142, negated_conjecture, (v7_rlvect_1(k15_euclid(esk1_0)))).#trainneg
cnf(i_0_143, negated_conjecture, (v8_rlvect_1(k15_euclid(esk1_0)))).#trainneg
cnf(i_0_144, negated_conjecture, (v5_rltopsp1(k15_euclid(esk1_0)))).#trainneg
cnf(i_0_145, negated_conjecture, (l1_rltopsp1(k15_euclid(esk1_0)))).#trainneg
cnf(i_0_146, negated_conjecture, (v2_pre_topc(k15_euclid(esk1_0)))).#trainneg
cnf(i_0_176, negated_conjecture, (l1_rlvect_1(k15_euclid(esk2_0)))).#trainneg
cnf(i_0_177, negated_conjecture, (l1_pre_topc(k15_euclid(esk2_0)))).#trainneg
cnf(i_0_182, negated_conjecture, (l1_rlvect_1(k15_euclid(esk1_0)))).#trainneg
cnf(i_0_183, negated_conjecture, (l1_pre_topc(k15_euclid(esk1_0)))).#trainneg
cnf(i_0_101, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(np__2)))|r2_hidden(esk4_0,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_179, negated_conjecture, (m1_subset_1(k3_goboard5(esk3_0,X1,esk1_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_180, negated_conjecture, (m1_subset_1(k3_goboard5(esk3_0,X1,np__2),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~v7_ordinal1(X1))).#trainneg
cnf(i_0_203, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk4_0)))).#trainneg
cnf(i_0_206, negated_conjecture, (~v1_xboole_0(u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_205, negated_conjecture, (r1_tarski(k3_goboard5(esk3_0,esk1_0,esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_197, negated_conjecture, (r1_xxreal_0(X1,k17_euclid(esk4_0))|r1_xxreal_0(k17_euclid(esk4_0),X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_199, negated_conjecture, (r2_hidden(X1,k3_goboard5(esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))|~r1_xxreal_0(k18_euclid(X1),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k17_euclid(X1),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(k2_nat_1(esk1_0,np__1),np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(X1))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(X1))|~r1_xxreal_0(k2_nat_1(k2_nat_1(esk1_0,np__1),np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,k2_nat_1(esk1_0,np__1))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(k2_nat_1(esk1_0,np__1),k5_numbers))).#trainneg
cnf(i_0_214, negated_conjecture, (r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_219, negated_conjecture, (r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk4_0)))).#trainneg
cnf(i_0_196, negated_conjecture, (r1_xxreal_0(np__2,k17_euclid(esk4_0))|r1_xxreal_0(k17_euclid(esk4_0),np__2))).#trainneg
cnf(i_0_222, negated_conjecture, (r1_xxreal_0(np__2,k18_euclid(esk4_0))|r1_xxreal_0(k18_euclid(esk4_0),np__2))).#trainneg
cnf(i_0_224, negated_conjecture, (r1_xxreal_0(X1,k18_euclid(esk4_0))|r1_xxreal_0(k18_euclid(esk4_0),X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_227, negated_conjecture, (r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk4_0)))).#trainneg
cnf(i_0_238, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_244, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_240, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(X1)!=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_247, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_246, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_259, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_257, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(X1)!=k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_263, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_262, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_189, negated_conjecture, (v1_xboole_0(k1_numbers)|r2_hidden(k17_euclid(esk4_0),k1_numbers))).#trainneg
cnf(i_0_114, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),X1,esk4_0))|~l1_rlvect_1(k15_euclid(np__2))|~v8_rlvect_1(k15_euclid(np__2))|~v7_rlvect_1(k15_euclid(np__2))|~v6_rlvect_1(k15_euclid(np__2))|~v5_rlvect_1(k15_euclid(np__2))|~v4_rlvect_1(k15_euclid(np__2))|~v3_rlvect_1(k15_euclid(np__2))|~v2_rlvect_1(k15_euclid(np__2))|~v13_algstr_0(k15_euclid(np__2))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_271, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(esk4_0,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk4_0)))).#trainneg
cnf(i_0_276, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|~v1_xboole_0(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk4_0)))).#trainneg
cnf(i_0_275, negated_conjecture, (m1_subset_1(esk4_0,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk4_0))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_186, negated_conjecture, (m1_subset_1(k3_goboard5(esk3_0,esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_285, negated_conjecture, (~v1_xboole_0(u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,esk2_0)))).#trainneg
cnf(i_0_286, negated_conjecture, (m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,esk2_0)))).#trainneg
cnf(i_0_284, negated_conjecture, (r1_tarski(k3_goboard5(esk3_0,esk2_0,esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_293, negated_conjecture, (r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,esk2_0))|~r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_287, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk2_0))|esk5_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk2_0))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_288, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk2_0))|esk6_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk2_0))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_217, negated_conjecture, (v1_xboole_0(k1_numbers)|r2_hidden(k18_euclid(esk4_0),k1_numbers))).#trainneg
cnf(i_0_289, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk2_0))|m1_subset_1(esk5_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,esk2_0)))).#trainneg
cnf(i_0_298, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0)=k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_300, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|~v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_188, negated_conjecture, (m1_subset_1(k3_goboard5(esk3_0,np__2,esk2_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_308, negated_conjecture, (~v1_xboole_0(u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,esk2_0)))).#trainneg
cnf(i_0_309, negated_conjecture, (m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,esk2_0)))).#trainneg
cnf(i_0_307, negated_conjecture, (r1_tarski(k3_goboard5(esk3_0,np__2,esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_316, negated_conjecture, (r2_hidden(X1,k3_goboard5(esk3_0,np__2,esk2_0))|~r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_310, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,esk2_0))|esk5_3(np__2,X1,k3_goboard5(esk3_0,np__2,esk2_0))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_311, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,esk2_0))|esk6_3(np__2,X1,k3_goboard5(esk3_0,np__2,esk2_0))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_191, negated_conjecture, (m1_subset_1(k3_goboard5(esk3_0,esk2_0,esk1_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_320, negated_conjecture, (~v1_xboole_0(u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,esk1_0)))).#trainneg
cnf(i_0_328, negated_conjecture, (r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,esk1_0))|~r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk1_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_303, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|v2_struct_0(k15_euclid(np__2))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_321, negated_conjecture, (m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,esk1_0)))).#trainneg
cnf(i_0_319, negated_conjecture, (r1_tarski(k3_goboard5(esk3_0,esk2_0,esk1_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_322, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk1_0))|esk5_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk1_0))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,esk1_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_323, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk1_0))|esk6_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk1_0))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,esk1_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_192, negated_conjecture, (m1_subset_1(k3_goboard5(esk3_0,esk1_0,esk1_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_332, negated_conjecture, (~v1_xboole_0(u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk1_0)))).#trainneg
cnf(i_0_333, negated_conjecture, (m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk1_0)))).#trainneg
cnf(i_0_331, negated_conjecture, (r1_tarski(k3_goboard5(esk3_0,esk1_0,esk1_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_340, negated_conjecture, (r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk1_0))|~r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk1_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_334, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk1_0))|esk5_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk1_0))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk1_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_335, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk1_0))|esk6_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk1_0))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk1_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_290, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk2_0))|m1_subset_1(esk6_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,esk2_0)))).#trainneg
cnf(i_0_304, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|v2_struct_0(k15_euclid(np__2))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_193, negated_conjecture, (m1_subset_1(k3_goboard5(esk3_0,np__2,esk1_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_344, negated_conjecture, (~v1_xboole_0(u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,esk1_0)))).#trainneg
cnf(i_0_345, negated_conjecture, (m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,esk1_0)))).#trainneg
cnf(i_0_343, negated_conjecture, (r1_tarski(k3_goboard5(esk3_0,np__2,esk1_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_352, negated_conjecture, (r2_hidden(X1,k3_goboard5(esk3_0,np__2,esk1_0))|~r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,esk1_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_346, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,esk1_0))|esk5_3(np__2,X1,k3_goboard5(esk3_0,np__2,esk1_0))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,esk1_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_347, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,esk1_0))|esk6_3(np__2,X1,k3_goboard5(esk3_0,np__2,esk1_0))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,esk1_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_200, negated_conjecture, (m1_subset_1(k3_goboard5(esk3_0,esk2_0,np__2),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_356, negated_conjecture, (~v1_xboole_0(u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,np__2)))).#trainneg
cnf(i_0_357, negated_conjecture, (m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,np__2)))).#trainneg
cnf(i_0_355, negated_conjecture, (r1_tarski(k3_goboard5(esk3_0,esk2_0,np__2),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_364, negated_conjecture, (r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,np__2))|~r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,np__2))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_301, negated_conjecture, (r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|v2_struct_0(k15_euclid(np__2))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_358, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,np__2))|esk5_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,np__2))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,np__2))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_359, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,np__2))|esk6_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,np__2))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,np__2))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_201, negated_conjecture, (m1_subset_1(k3_goboard5(esk3_0,esk1_0,np__2),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_368, negated_conjecture, (~v1_xboole_0(u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,np__2)))).#trainneg
cnf(i_0_369, negated_conjecture, (m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,np__2)))).#trainneg
cnf(i_0_367, negated_conjecture, (r1_tarski(k3_goboard5(esk3_0,esk1_0,np__2),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_376, negated_conjecture, (r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,np__2))|~r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,np__2))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_370, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,np__2))|esk5_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,np__2))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,np__2))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_371, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,np__2))|esk6_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,np__2))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,np__2))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_202, negated_conjecture, (m1_subset_1(k3_goboard5(esk3_0,np__2,np__2),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_380, negated_conjecture, (~v1_xboole_0(u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,np__2)))).#trainneg
cnf(i_0_388, negated_conjecture, (r2_hidden(X1,k3_goboard5(esk3_0,np__2,np__2))|~r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,np__2))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_302, negated_conjecture, (r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|v2_struct_0(k15_euclid(np__2))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_381, negated_conjecture, (m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,np__2)))).#trainneg
cnf(i_0_379, negated_conjecture, (r1_tarski(k3_goboard5(esk3_0,np__2,np__2),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_382, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,np__2))|esk5_3(np__2,X1,k3_goboard5(esk3_0,np__2,np__2))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,np__2))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_383, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,np__2))|esk6_3(np__2,X1,k3_goboard5(esk3_0,np__2,np__2))!=X1|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,np__2))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_204, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|r2_hidden(k3_goboard5(esk3_0,esk1_0,esk2_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_250, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),X1,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),X1)|v2_struct_0(k15_euclid(np__2))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_393, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0)=k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_395, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|~v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_232, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(np__2)))|r2_hidden(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_249, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(np__2)))|r2_hidden(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_312, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,esk2_0))|m1_subset_1(esk5_3(np__2,X1,k3_goboard5(esk3_0,np__2,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,esk2_0)))).#trainneg
cnf(i_0_398, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|v2_struct_0(k15_euclid(np__2))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_223, negated_conjecture, (r1_xxreal_0(k17_euclid(esk4_0),k18_euclid(esk4_0))|r1_xxreal_0(k18_euclid(esk4_0),k17_euclid(esk4_0)))).#trainneg
cnf(i_0_313, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,esk2_0))|m1_subset_1(esk6_3(np__2,X1,k3_goboard5(esk3_0,np__2,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,esk2_0)))).#trainneg
cnf(i_0_324, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk1_0))|m1_subset_1(esk5_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk1_0)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,esk1_0)))).#trainneg
cnf(i_0_325, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk1_0))|m1_subset_1(esk6_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk1_0)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,esk1_0)))).#trainneg
cnf(i_0_336, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk1_0))|m1_subset_1(esk5_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk1_0)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk1_0)))).#trainneg
cnf(i_0_337, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk1_0))|m1_subset_1(esk6_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk1_0)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk1_0)))).#trainneg
cnf(i_0_399, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|v2_struct_0(k15_euclid(np__2))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_348, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,esk1_0))|m1_subset_1(esk5_3(np__2,X1,k3_goboard5(esk3_0,np__2,esk1_0)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,esk1_0)))).#trainneg
cnf(i_0_349, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,esk1_0))|m1_subset_1(esk6_3(np__2,X1,k3_goboard5(esk3_0,np__2,esk1_0)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,esk1_0)))).#trainneg
cnf(i_0_360, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,np__2))|m1_subset_1(esk5_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,np__2)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,np__2)))).#trainneg
cnf(i_0_361, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,np__2))|m1_subset_1(esk6_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,np__2)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,np__2)))).#trainneg
cnf(i_0_372, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,np__2))|m1_subset_1(esk5_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,np__2)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,np__2)))).#trainneg
cnf(i_0_373, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,np__2))|m1_subset_1(esk6_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,np__2)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,np__2)))).#trainneg
cnf(i_0_384, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,np__2))|m1_subset_1(esk5_3(np__2,X1,k3_goboard5(esk3_0,np__2,np__2)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,np__2)))).#trainneg
cnf(i_0_385, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,np__2))|m1_subset_1(esk6_3(np__2,X1,k3_goboard5(esk3_0,np__2,np__2)),u1_struct_0(k15_euclid(np__2)))|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,np__2)))).#trainneg
cnf(i_0_396, negated_conjecture, (r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|v2_struct_0(k15_euclid(np__2))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_438, negated_conjecture, (r1_xxreal_0(X1,k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_283, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|r2_hidden(k3_goboard5(esk3_0,esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_306, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|r2_hidden(k3_goboard5(esk3_0,np__2,esk2_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_318, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|r2_hidden(k3_goboard5(esk3_0,esk2_0,esk1_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_330, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|r2_hidden(k3_goboard5(esk3_0,esk1_0,esk1_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_342, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|r2_hidden(k3_goboard5(esk3_0,np__2,esk1_0),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_354, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|r2_hidden(k3_goboard5(esk3_0,esk2_0,np__2),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_366, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|r2_hidden(k3_goboard5(esk3_0,esk1_0,np__2),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_116, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),X1,esk4_0)=k1_rltopsp1(k15_euclid(np__2),esk4_0,X1)|v2_struct_0(k15_euclid(np__2))|~l1_rlvect_1(k15_euclid(np__2))|~v8_rlvect_1(k15_euclid(np__2))|~v7_rlvect_1(k15_euclid(np__2))|~v6_rlvect_1(k15_euclid(np__2))|~v5_rlvect_1(k15_euclid(np__2))|~v4_rlvect_1(k15_euclid(np__2))|~v3_rlvect_1(k15_euclid(np__2))|~v2_rlvect_1(k15_euclid(np__2))|~v13_algstr_0(k15_euclid(np__2))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_397, negated_conjecture, (r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|v2_struct_0(k15_euclid(np__2))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_470, negated_conjecture, (r1_xxreal_0(X1,k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_378, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|r2_hidden(k3_goboard5(esk3_0,np__2,np__2),k1_zfmisc_1(u1_struct_0(k15_euclid(np__2)))))).#trainneg
cnf(i_0_481, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|~v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_245, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_266, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(esk4_0,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_490, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|~v1_xboole_0(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_274, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(esk4_0,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_497, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|~v1_xboole_0(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_489, negated_conjecture, (m1_subset_1(esk4_0,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_496, negated_conjecture, (m1_subset_1(esk4_0,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_408, negated_conjecture, (r1_xxreal_0(np__2,k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),np__2))).#trainneg
cnf(i_0_420, negated_conjecture, (r1_xxreal_0(np__2,k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),np__2))).#trainneg
cnf(i_0_261, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_291, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk2_0))|r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk2_0)),esk6_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk2_0))))|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,esk2_0)))).#trainneg
cnf(i_0_433, negated_conjecture, (r1_xxreal_0(np__2,k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),np__2))).#trainneg
cnf(i_0_464, negated_conjecture, (r1_xxreal_0(np__2,k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),np__2))).#trainneg
cnf(i_0_495, negated_conjecture, (k18_euclid(X1)=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_502, negated_conjecture, (k18_euclid(X1)=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_292, negated_conjecture, (r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk2_0)),esk6_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk2_0))),k3_goboard5(esk3_0,esk2_0,esk2_0))|r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk2_0))|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,esk2_0)))).#trainneg
cnf(i_0_315, negated_conjecture, (r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,np__2,esk2_0)),esk6_3(np__2,X1,k3_goboard5(esk3_0,np__2,esk2_0))),k3_goboard5(esk3_0,np__2,esk2_0))|r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,esk2_0))|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,esk2_0)))).#trainneg
cnf(i_0_405, negated_conjecture, (v1_xboole_0(k1_numbers)|r2_hidden(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_numbers))).#trainneg
cnf(i_0_314, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,esk2_0))|r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,np__2,esk2_0)),esk6_3(np__2,X1,k3_goboard5(esk3_0,np__2,esk2_0))))|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,esk2_0)))).#trainneg
cnf(i_0_299, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|esk4_0=X1|v2_struct_0(k15_euclid(np__2))|~r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),X2)|~r1_sppol_1(np__2,X1,X2)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_327, negated_conjecture, (r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk1_0)),esk6_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk1_0))),k3_goboard5(esk3_0,esk2_0,esk1_0))|r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk1_0))|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,esk1_0)))).#trainneg
cnf(i_0_339, negated_conjecture, (r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk1_0)),esk6_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk1_0))),k3_goboard5(esk3_0,esk1_0,esk1_0))|r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk1_0))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk1_0)))).#trainneg
cnf(i_0_351, negated_conjecture, (r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,np__2,esk1_0)),esk6_3(np__2,X1,k3_goboard5(esk3_0,np__2,esk1_0))),k3_goboard5(esk3_0,np__2,esk1_0))|r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,esk1_0))|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,esk1_0)))).#trainneg
cnf(i_0_363, negated_conjecture, (r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,np__2)),esk6_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,np__2))),k3_goboard5(esk3_0,esk2_0,np__2))|r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,np__2))|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,np__2)))).#trainneg
cnf(i_0_375, negated_conjecture, (r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,np__2)),esk6_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,np__2))),k3_goboard5(esk3_0,esk1_0,np__2))|r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,np__2))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,np__2)))).#trainneg
cnf(i_0_387, negated_conjecture, (r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,np__2,np__2)),esk6_3(np__2,X1,k3_goboard5(esk3_0,np__2,np__2))),k3_goboard5(esk3_0,np__2,np__2))|r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,np__2))|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,np__2)))).#trainneg
cnf(i_0_326, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk1_0))|r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk1_0)),esk6_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,esk1_0))))|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,esk1_0)))).#trainneg
cnf(i_0_338, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk1_0))|r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk1_0)),esk6_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk1_0))))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk1_0)))).#trainneg
cnf(i_0_350, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,esk1_0))|r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,np__2,esk1_0)),esk6_3(np__2,X1,k3_goboard5(esk3_0,np__2,esk1_0))))|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,esk1_0)))).#trainneg
cnf(i_0_362, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk2_0,np__2))|r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,np__2)),esk6_3(np__2,X1,k3_goboard5(esk3_0,esk2_0,np__2))))|~r2_hidden(X1,k3_goboard5(esk3_0,esk2_0,np__2)))).#trainneg
cnf(i_0_417, negated_conjecture, (v1_xboole_0(k1_numbers)|r2_hidden(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_numbers))).#trainneg
cnf(i_0_374, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,np__2))|r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,np__2)),esk6_3(np__2,X1,k3_goboard5(esk3_0,esk1_0,np__2))))|~r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,np__2)))).#trainneg
cnf(i_0_394, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|esk4_0=X1|v2_struct_0(k15_euclid(np__2))|~r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),X2)|~r1_sppol_1(np__2,X1,X2)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_386, negated_conjecture, (r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,np__2,np__2))|r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,X1,k3_goboard5(esk3_0,np__2,np__2)),esk6_3(np__2,X1,k3_goboard5(esk3_0,np__2,np__2))))|~r2_hidden(X1,k3_goboard5(esk3_0,np__2,np__2)))).#trainneg
cnf(i_0_514, negated_conjecture, (~v1_xboole_0(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_513, negated_conjecture, (m1_subset_1(esk4_0,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_409, negated_conjecture, (r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).#trainneg
cnf(i_0_430, negated_conjecture, (v1_xboole_0(k1_numbers)|r2_hidden(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_numbers))).#trainneg
cnf(i_0_523, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_529, negated_conjecture, (r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(esk4_0)))).#trainneg
cnf(i_0_530, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|~v1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_410, negated_conjecture, (r1_xxreal_0(k18_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).#trainneg
cnf(i_0_421, negated_conjecture, (r1_xxreal_0(k17_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).#trainneg
cnf(i_0_422, negated_conjecture, (r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).#trainneg
cnf(i_0_535, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_459, negated_conjecture, (v1_xboole_0(k1_numbers)|r2_hidden(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_numbers))).#trainneg
cnf(i_0_522, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_542, negated_conjecture, (r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(esk4_0)))).#trainneg
cnf(i_0_543, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|~v1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_434, negated_conjecture, (r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).#trainneg
cnf(i_0_546, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_435, negated_conjecture, (r1_xxreal_0(k18_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).#trainneg
cnf(i_0_534, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_550, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|~r1_sppol_1(np__2,X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_555, negated_conjecture, (r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))))).#trainneg
cnf(i_0_556, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))=k18_euclid(esk4_0)|~v1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk4_0)))).#trainneg
cnf(i_0_465, negated_conjecture, (r1_xxreal_0(k17_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).#trainneg
cnf(i_0_466, negated_conjecture, (r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).#trainneg
cnf(i_0_561, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_413, negated_conjecture, (r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_545, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_566, negated_conjecture, (r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))))).#trainneg
cnf(i_0_567, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))=k17_euclid(esk4_0)|~v1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(esk4_0)))).#trainneg
cnf(i_0_560, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_425, negated_conjecture, (r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_441, negated_conjecture, (r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_480, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|v2_struct_0(k15_euclid(np__2))|~r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),X2)|~r1_sppol_1(np__2,X1,X2)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_471, negated_conjecture, (r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_269, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0)))).#trainneg
cnf(i_0_572, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|~v1_xboole_0(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0)))).#trainneg
cnf(i_0_577, negated_conjecture, (k18_euclid(X1)=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_580, negated_conjecture, (k18_euclid(X1)=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_581, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_270, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0)))).#trainneg
cnf(i_0_583, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|~v1_xboole_0(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0)))).#trainneg
cnf(i_0_588, negated_conjecture, (k18_euclid(X1)=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_597, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_599, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_598, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_600, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_609, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_611, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_604, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_594, negated_conjecture, (r2_hidden(X1,k3_goboard5(esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|~r1_xxreal_0(k17_euclid(X1),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(k2_nat_1(esk1_0,np__1),np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(X1))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(X1))|~r1_xxreal_0(k2_nat_1(k2_nat_1(esk1_0,np__1),np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,k2_nat_1(esk1_0,np__1))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(k2_nat_1(esk1_0,np__1),k5_numbers))).#trainneg
cnf(i_0_606, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_610, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_612, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_605, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_607, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_592, negated_conjecture, (k18_euclid(X1)=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_613, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_571, negated_conjecture, (m1_subset_1(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_578, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_305, negated_conjecture, (k18_euclid(X1)=k18_euclid(X2)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(X2)|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_582, negated_conjecture, (m1_subset_1(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_589, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_615, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|m1_subset_1(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_630, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|m1_subset_1(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_602, negated_conjecture, (r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|~r1_xxreal_0(k17_euclid(X1),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(X1))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(X1))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_400, negated_conjecture, (k18_euclid(X1)=k18_euclid(X2)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(X2)|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_649, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_648, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk4_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_653, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_647, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk4_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_175, negated_conjecture, (r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),k3_goboard5(X1,X2,X3))|v3_relat_1(X1)|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,k2_nat_1(X3,np__1))))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1),X3)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))|~v5_goboard1(X1)|~v4_goboard1(X1)|~v3_goboard1(X1)|~v2_goboard1(X1)|~v1_matrix_1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_654, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_646, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(X1))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(X1))|~r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_657, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk4_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_660, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_656, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(esk4_0))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk4_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_661, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk4_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_662, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_601, negated_conjecture, (r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),k3_goboard5(X1,X2,X3))|v3_relat_1(X1)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,k2_nat_1(X3,np__1))))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1),X3)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k18_euclid(esk4_0))|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))|~v5_goboard1(X1)|~v4_goboard1(X1)|~v3_goboard1(X1)|~v2_goboard1(X1)|~v1_matrix_1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_664, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k3_goboard5(esk3_0,esk1_0,esk2_0)))|r2_hidden(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_zfmisc_1(k3_goboard5(esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_423, negated_conjecture, (r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_486, negated_conjecture, (k18_euclid(X1)=k18_euclid(X2)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(X2)|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_198, negated_conjecture, (r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),k3_goboard5(X1,X2,X3))|v3_relat_1(X1)|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,k2_nat_1(X3,np__1))))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1),X3)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))|~v5_goboard1(X1)|~v4_goboard1(X1)|~v3_goboard1(X1)|~v2_goboard1(X1)|~v1_matrix_1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_437, negated_conjecture, (r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_467, negated_conjecture, (r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_469, negated_conjecture, (r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_688, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|~v1_xboole_0(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_265, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_699, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|~v1_xboole_0(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_272, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_706, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|~v1_xboole_0(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_725, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_718, negated_conjecture, (k18_euclid(X1)=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_740, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)|v2_struct_0(k15_euclid(np__2))|~r2_hidden(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k16_matrix_1(esk3_0)))).#trainneg
cnf(i_0_732, negated_conjecture, (r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_744, negated_conjecture, (esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0))).#trainneg
cnf(i_0_747, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_743, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),X1,esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_742, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),X1,esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))=k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),X1)|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_754, negated_conjecture, (X1=esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(X1)!=k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_749, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0)=k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_765, negated_conjecture, (k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_755, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_771, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_733, negated_conjecture, (r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_775, negated_conjecture, (esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0))).#trainneg
cnf(i_0_778, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_774, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),X1,esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_780, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0)=k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_795, negated_conjecture, (k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_773, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),X1,esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))=k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),X1)|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_768, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_787, negated_conjecture, (X1=esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(X1)!=k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_788, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_806, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_769, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_798, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_764, negated_conjecture, (esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|esk4_0=X1|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),X2)|~r1_sppol_1(np__2,X1,X2)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_799, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_816, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_815, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_824, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_827, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_826, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_741, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(np__2)))|r2_hidden(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_772, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(np__2)))|r2_hidden(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_687, negated_conjecture, (m1_subset_1(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_794, negated_conjecture, (esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|esk4_0=X1|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),X2)|~r1_sppol_1(np__2,X1,X2)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_766, negated_conjecture, (r1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_750, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_839, negated_conjecture, (k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_752, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_846, negated_conjecture, (k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_781, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_853, negated_conjecture, (k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_783, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_860, negated_conjecture, (k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_842, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_843, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_849, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_850, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_856, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_857, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_863, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_864, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_759, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r2_hidden(esk4_0,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_867, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v1_xboole_0(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_900, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)|v2_struct_0(k15_euclid(np__2))|~r2_hidden(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k16_matrix_1(esk3_0)))).#trainneg
cnf(i_0_892, negated_conjecture, (r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_904, negated_conjecture, (esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0))).#trainneg
cnf(i_0_907, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_903, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),X1,esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_902, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),X1,esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))=k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),X1)|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_840, negated_conjecture, (r1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_909, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0)=k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_932, negated_conjecture, (k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_918, negated_conjecture, (X1=esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(X1)!=k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_919, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_939, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_893, negated_conjecture, (r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_943, negated_conjecture, (esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0))).#trainneg
cnf(i_0_946, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_942, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),X1,esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_941, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),X1,esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))=k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),X1)|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_935, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_948, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0)=k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_975, negated_conjecture, (k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_959, negated_conjecture, (X1=esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(X1)!=k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_960, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_982, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_931, negated_conjecture, (esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|esk4_0=X1|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),X2)|~r1_sppol_1(np__2,X1,X2)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_936, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_978, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_979, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_988, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).#trainneg
cnf(i_0_992, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_991, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1003, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1002, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_901, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(np__2)))|r2_hidden(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_698, negated_conjecture, (m1_subset_1(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_974, negated_conjecture, (esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|esk4_0=X1|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),X2)|~r1_sppol_1(np__2,X1,X2)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_841, negated_conjecture, (r1_xxreal_0(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_940, negated_conjecture, (v1_xboole_0(u1_struct_0(k15_euclid(np__2)))|r2_hidden(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_910, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1014, negated_conjecture, (k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_912, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1021, negated_conjecture, (k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_949, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1028, negated_conjecture, (k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1017, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1018, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_951, negated_conjecture, (k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1035, negated_conjecture, (k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1024, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1025, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1031, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1032, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1038, negated_conjecture, (r1_xxreal_0(k17_euclid(X1),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1039, negated_conjecture, (r1_xxreal_0(k18_euclid(X1),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_705, negated_conjecture, (m1_subset_1(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_838, negated_conjecture, (esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),X2)|~r1_sppol_1(np__2,X1,X2)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_847, negated_conjecture, (r1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_872, negated_conjecture, (k18_euclid(X1)=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_793, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r2_hidden(esk4_0,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1043, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v1_xboole_0(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1048, negated_conjecture, (k18_euclid(X1)=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_866, negated_conjecture, (m1_subset_1(esk4_0,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_925, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r2_hidden(esk4_0,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1051, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v1_xboole_0(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1056, negated_conjecture, (k18_euclid(X1)=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_967, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r2_hidden(esk4_0,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1058, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v1_xboole_0(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_712, negated_conjecture, (m1_subset_1(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_845, negated_conjecture, (esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),X2)|~r1_sppol_1(np__2,X1,X2)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_848, negated_conjecture, (r1_xxreal_0(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1063, negated_conjecture, (k18_euclid(X1)=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1042, negated_conjecture, (m1_subset_1(esk4_0,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1050, negated_conjecture, (m1_subset_1(esk4_0,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1057, negated_conjecture, (m1_subset_1(esk4_0,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_854, negated_conjecture, (r1_xxreal_0(k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_855, negated_conjecture, (r1_xxreal_0(k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_861, negated_conjecture, (r1_xxreal_0(k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_862, negated_conjecture, (r1_xxreal_0(k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1015, negated_conjecture, (r1_xxreal_0(k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1016, negated_conjecture, (r1_xxreal_0(k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_837, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|m1_subset_1(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_852, negated_conjecture, (esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),X2)|~r1_sppol_1(np__2,X1,X2)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1022, negated_conjecture, (r1_xxreal_0(k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1023, negated_conjecture, (r1_xxreal_0(k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1029, negated_conjecture, (r1_xxreal_0(k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1030, negated_conjecture, (r1_xxreal_0(k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1036, negated_conjecture, (r1_xxreal_0(k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1037, negated_conjecture, (r1_xxreal_0(k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_767, negated_conjecture, (r1_xxreal_0(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_796, negated_conjecture, (r1_xxreal_0(k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_797, negated_conjecture, (r1_xxreal_0(k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_933, negated_conjecture, (r1_xxreal_0(k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_934, negated_conjecture, (r1_xxreal_0(k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_859, negated_conjecture, (esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),X2)|~r1_sppol_1(np__2,X1,X2)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_976, negated_conjecture, (r1_xxreal_0(k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_977, negated_conjecture, (r1_xxreal_0(k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1069, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1073, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1075, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1074, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1013, negated_conjecture, (esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),X2)|~r1_sppol_1(np__2,X1,X2)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1020, negated_conjecture, (esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),X2)|~r1_sppol_1(np__2,X1,X2)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1080, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1084, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).#trainneg
cnf(i_0_1092, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_1103, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1105, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1104, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1106, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1102, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_1117, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0))).#trainneg
cnf(i_0_1142, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_1109, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1121, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1111, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1140, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1120, negated_conjecture, (X1=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|k18_euclid(X1)!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1173, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1172, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1177, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1122, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(X1))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1141, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1123, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1135, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(X1))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1163, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_1164, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_1166, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_1167, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_1119, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_1136, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1148, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1169, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_1095, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).#trainneg
cnf(i_0_1191, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1190, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1101, negated_conjecture, (r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|~r1_xxreal_0(k17_euclid(X1),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(X1))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(X1))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1097, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).#trainneg
cnf(i_0_1207, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1206, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1215, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1217, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1216, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1221, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1223, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1218, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1255, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_1253, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1234, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1233, negated_conjecture, (X1=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|k18_euclid(X1)!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1286, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1285, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1235, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1288, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1254, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1236, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1248, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1276, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_1277, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_1279, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_1200, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk4_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1280, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_1232, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_1282, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_1249, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1289, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1027, negated_conjecture, (esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),X2)|~r1_sppol_1(np__2,X1,X2)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1199, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk4_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1306, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1292, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).#trainneg
cnf(i_0_1310, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1220, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).#trainneg
cnf(i_0_1318, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1213, negated_conjecture, (r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|~r1_xxreal_0(k17_euclid(X1),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(X1))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(X1))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1317, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1222, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).#trainneg
cnf(i_0_1333, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1332, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1324, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk4_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1323, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk4_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1340, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1343, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1347, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1349, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1261, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1348, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1354, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1358, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).#trainneg
cnf(i_0_1360, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1366, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).#trainneg
cnf(i_0_1371, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1034, negated_conjecture, (esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=X1|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_tarski(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),X2)|~r1_sppol_1(np__2,X1,X2)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(k15_euclid(np__2))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1263, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1426, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1428, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1424, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_1417, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1474, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|X1=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(X1)!=k18_euclid(esk4_0)|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1471, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1421, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1419, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1473, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0))).#trainneg
cnf(i_0_1422, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1441, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1479, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_1442, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1480, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_1423, negated_conjecture, (r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|~r1_xxreal_0(k17_euclid(X1),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(X1))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(X1))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1493, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1481, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_1482, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_1453, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1494, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1454, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1505, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1536, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk4_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1535, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk4_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1559, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1416, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).#trainneg
cnf(i_0_1564, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1563, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1418, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).#trainneg
cnf(i_0_1571, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1570, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1504, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1268, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1528, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1530, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).#trainneg
cnf(i_0_1592, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1270, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1430, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).#trainneg
cnf(i_0_1603, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1602, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1432, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).#trainneg
cnf(i_0_1610, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1609, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1650, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1632, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1620, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1621, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1633, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1665, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_770, negated_conjecture, (k18_euclid(X1)=k18_euclid(X2)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(X2)|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1639, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1667, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).#trainneg
cnf(i_0_1679, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1408, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_800, negated_conjecture, (k18_euclid(X1)=k18_euclid(X2)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(X2)|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1641, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1541, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1545, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).#trainneg
cnf(i_0_1702, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_745, negated_conjecture, (m1_subset_1(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_numbers)|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_937, negated_conjecture, (k18_euclid(X1)=k18_euclid(X2)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(X2)|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1646, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1708, negated_conjecture, (v1_xreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1710, negated_conjecture, (v1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1718, negated_conjecture, (r1_xxreal_0(X1,k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),X1)|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1488, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_746, negated_conjecture, (m1_subset_1(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_numbers)|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_980, negated_conjecture, (k18_euclid(X1)=k18_euclid(X2)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(X2)|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1648, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1733, negated_conjecture, (v1_xreal_0(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1735, negated_conjecture, (v1_xxreal_0(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1744, negated_conjecture, (r1_xxreal_0(X1,k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),X1)|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_776, negated_conjecture, (m1_subset_1(k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_numbers)|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_844, negated_conjecture, (k18_euclid(X1)=k18_euclid(X2)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(X2)|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1390, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1755, negated_conjecture, (v1_xreal_0(k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1756, negated_conjecture, (v1_xxreal_0(k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1766, negated_conjecture, (r1_xxreal_0(X1,k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),X1)|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_777, negated_conjecture, (m1_subset_1(k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_numbers)|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_851, negated_conjecture, (k18_euclid(X1)=k18_euclid(X2)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(X2)|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1392, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1778, negated_conjecture, (v1_xreal_0(k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1779, negated_conjecture, (v1_xxreal_0(k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1790, negated_conjecture, (r1_xxreal_0(X1,k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),X1)|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_905, negated_conjecture, (m1_subset_1(k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k1_numbers)|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_858, negated_conjecture, (k18_euclid(X1)=k18_euclid(X2)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(X2)|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1397, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1803, negated_conjecture, (v1_xreal_0(k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1804, negated_conjecture, (v1_xxreal_0(k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1816, negated_conjecture, (r1_xxreal_0(X1,k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),X1)|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_1863, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1930, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1944, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1964, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_1987, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_1961, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1943, negated_conjecture, (X1=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|k18_euclid(X1)!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2000, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1999, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1866, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|X1=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(X1)!=k18_euclid(esk4_0)|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2001, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_2004, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1963, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2003, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1945, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(X1))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1946, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1958, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1985, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_1986, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_1988, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_1989, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_2005, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_2006, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_2007, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_2008, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_1942, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_1833, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1991, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_1957, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1834, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2013, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_1845, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1973, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2015, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).#trainneg
cnf(i_0_2046, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1846, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1975, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2019, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2025, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_2035, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).#trainneg
cnf(i_0_2064, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1977, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2018, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2031, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1979, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_865, negated_conjecture, (k18_euclid(X1)=k18_euclid(X2)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(X2)|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1972, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2112, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_2123, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_2154, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_2204, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2184, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_2258, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_2117, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|X1=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(X1)!=k18_euclid(esk4_0)|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2151, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_2152, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2726, negated_conjecture, (r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0),k3_goboard5(X1,X2,X3))|v3_relat_1(X1)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,k2_nat_1(X3,np__1))))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1),X3)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k18_euclid(esk4_0))|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))|~v5_goboard1(X1)|~v4_goboard1(X1)|~v3_goboard1(X1)|~v2_goboard1(X1)|~v1_matrix_1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_2751, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_2736, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0)).#trainneg
cnf(i_0_2745, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_2205, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(X1))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2747, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_2758, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_2760, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_2206, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2752, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2765, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2750, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2083, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2264, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_2763, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2793, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2803, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2132, negated_conjecture, (X1=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(X1)!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2811, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2810, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2153, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2265, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_2133, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2822, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2821, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2267, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_2268, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_2728, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).#trainneg
cnf(i_0_2084, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2730, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).#trainneg
cnf(i_0_2183, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2218, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(X1))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2253, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_2254, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_2256, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_2257, negated_conjecture, (k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_2095, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_2134, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k17_euclid(esk4_0)|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3386, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0),k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3402, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_3420, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_3388, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3421, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_3423, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_3433, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_3435, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_3397, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0))).#trainneg
cnf(i_0_3401, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0))).#trainneg
cnf(i_0_3454, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_3456, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_3466, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_3426, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_3428, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_3438, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_3440, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_3468, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_3449, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_3451, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_3461, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_3463, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_1019, negated_conjecture, (k18_euclid(X1)=k18_euclid(X2)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(X2)|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1399, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1026, negated_conjecture, (k18_euclid(X1)=k18_euclid(X2)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(X2)|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1638, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3500, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_3502, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_3495, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_3497, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_3520, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_3515, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_3517, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_3522, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_3492, negated_conjecture, (r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|~r1_xxreal_0(k17_euclid(X1),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(X1))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(X1))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3535, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3550, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0))).#trainneg
cnf(i_0_3596, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_3554, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3597, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_3553, negated_conjecture, (X1=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|k18_euclid(X1)!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3648, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3647, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3534, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk4_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3645, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|X1=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(X1)!=k18_euclid(esk4_0)|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3533, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk4_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3662, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3542, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3552, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_3555, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3556, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3568, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3602, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_3603, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_3605, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_3606, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_3541, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3569, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3656, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_3657, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_3658, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_3659, negated_conjecture, (k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_3499, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).#trainneg
cnf(i_0_3684, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3683, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3512, negated_conjecture, (r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|~r1_xxreal_0(k17_euclid(X1),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(X1))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(X1))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3695, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3501, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).#trainneg
cnf(i_0_3702, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3701, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3519, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).#trainneg
cnf(i_0_3694, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk4_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3711, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3693, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk4_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3719, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3710, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3521, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).#trainneg
cnf(i_0_3726, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3725, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3672, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).#trainneg
cnf(i_0_3581, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3734, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3612, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3583, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3624, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3506, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_3750, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0))).#trainneg
cnf(i_0_3790, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_3794, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|X1=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(X1)!=k18_euclid(esk4_0)|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3795, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_3805, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3796, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_3797, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_3798, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_3799, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_3755, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3817, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3816, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1033, negated_conjecture, (k18_euclid(X1)=k18_euclid(X2)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(X2)|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3588, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3840, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_3528, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_3862, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0))).#trainneg
cnf(i_0_3866, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3390, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(X1))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(X1))|~r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|~r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3912, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3918, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_1040, negated_conjecture, (k18_euclid(X1)=k18_euclid(X2)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(X1)!=k17_euclid(X2)|~r2_hidden(X2,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X2,u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3590, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3842, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_3906, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_3826, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4012, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_4031, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_3992, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4039, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4009, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3991, negated_conjecture, (X1=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|k18_euclid(X1)!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4052, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4051, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4050, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4053, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_3981, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|X1=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(X1)!=k18_euclid(esk4_0)|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4056, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4011, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4055, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3993, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(X1))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3994, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3990, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_3940, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4006, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4037, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4038, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4040, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4041, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4057, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4058, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4059, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4060, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4043, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_4005, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4067, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).#trainneg
cnf(i_0_4088, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3952, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_3953, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4075, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4071, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_4021, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4083, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4023, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4171, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0))).#trainneg
cnf(i_0_4164, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_4196, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_4711, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0)).#trainneg
cnf(i_0_4712, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4175, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4775, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_4730, negated_conjecture, (esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_4786, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))!=k18_euclid(esk4_0))).#trainneg
cnf(i_0_4715, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0)).#trainneg
cnf(i_0_4713, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4771, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|X1=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k18_euclid(X1)!=k18_euclid(esk4_0)|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4717, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4718, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4951, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|X1=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(X1)!=k18_euclid(esk4_0)|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4165, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4157, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_4970, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4783, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4156, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|X1=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(X1)!=k18_euclid(esk4_0)|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4990, negated_conjecture, (k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4991, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4728, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4810, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4174, negated_conjecture, (X1=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|k18_euclid(X1)!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5042, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5043, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))).#trainneg
cnf(i_0_5041, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4155, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4971, negated_conjecture, (X1=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(X1)!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5055, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5054, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4709, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(X1,k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(X1))|~r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|~r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4703, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5073, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0),k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5048, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4702, negated_conjecture, (r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0),k3_goboard5(X1,X2,X3))|v3_relat_1(X1)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,k2_nat_1(X3,np__1))))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1),X3)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k18_euclid(esk4_0))|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))|~v5_goboard1(X1)|~v4_goboard1(X1)|~v3_goboard1(X1)|~v2_goboard1(X1)|~v1_matrix_1(X1)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0),u1_struct_0(k15_euclid(np__2)))|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_4953, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4954, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4955, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4956, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4176, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(X1))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4992, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4972, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5084, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5083, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4705, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).#trainneg
cnf(i_0_4177, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4859, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4860, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4862, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4863, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_5027, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_5028, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_5030, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_5031, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4736, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5067, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(esk4_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5100, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),k3_goboard5(esk3_0,esk1_0,esk2_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(esk4_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5101, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r2_hidden(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),k3_goboard5(esk3_0,esk1_0,esk2_0))|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4818, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4869, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4913, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4707, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0)).#trainneg
cnf(i_0_4121, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4748, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4122, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4830, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4881, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4925, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4189, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(X1))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5018, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4778, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4779, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4781, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4973, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4782, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_4747, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4829, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4882, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4926, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4190, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_4133, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5693, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_5694, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_5697, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_678, negated_conjecture, (r2_hidden(esk4_0,k3_goboard5(X1,X2,X3))|v3_relat_1(X1)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,k2_nat_1(X3,np__1))))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,k2_nat_1(X2,np__1),X3)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),X1,X2,X3)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k2_nat_1(X3,np__1),k1_matrix_1(X1))|~r1_xxreal_0(k2_nat_1(X2,np__1),k3_finseq_1(X1))|~r1_xxreal_0(np__1,X3)|~r1_xxreal_0(np__1,X2)|~m2_finseq_1(X1,k3_finseq_2(u1_struct_0(k15_euclid(np__2))))|~v5_goboard1(X1)|~v4_goboard1(X1)|~v3_goboard1(X1)|~v2_goboard1(X1)|~v1_matrix_1(X1)|~m1_subset_1(X3,k5_numbers)|~m1_subset_1(X2,k5_numbers))).#trainneg
cnf(i_0_5698, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))))).#trainneg
cnf(i_0_5729, negated_conjecture, (X1=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|k18_euclid(X1)!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5711, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|r2_hidden(esk4_0,k3_goboard5(esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(esk4_0))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(k2_nat_1(esk1_0,np__1),np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k2_nat_1(k2_nat_1(esk1_0,np__1),np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,k2_nat_1(esk1_0,np__1))|~m1_subset_1(k2_nat_1(esk1_0,np__1),k5_numbers))).#trainneg
cnf(i_0_5713, negated_conjecture, (k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)=esk4_0|r2_hidden(esk4_0,k3_goboard5(esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(k2_nat_1(esk1_0,np__1),np__1),esk2_0)))|~r1_xxreal_0(k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)),k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)))|~r1_xxreal_0(k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0)),k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1))))|~r1_xxreal_0(k2_nat_1(k2_nat_1(esk1_0,np__1),np__1),k3_finseq_1(esk3_0))|~r1_xxreal_0(np__1,k2_nat_1(esk1_0,np__1))|~m1_subset_1(k2_nat_1(esk1_0,np__1),k5_numbers))).#trainneg
cnf(i_0_5771, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5770, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5730, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5728, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_5719, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5774, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_5731, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5732, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5744, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5775, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_5777, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_5778, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_5780, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_5718, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5745, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5834, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0))).#trainneg
cnf(i_0_5835, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|X1=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k18_euclid(X1)!=k18_euclid(esk4_0)|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5869, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_5801, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5805, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5806, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5870, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_5871, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_5872, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_5817, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5840, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5852, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5800, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5841, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5818, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5787, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).#trainneg
cnf(i_0_5897, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5851, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5756, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5879, negated_conjecture, (k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).#trainneg
cnf(i_0_5907, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5758, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5891, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).#trainneg
cnf(i_0_5916, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5762, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5764, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5823, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5825, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5829, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5831, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k18_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5859, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5961, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_6018, negated_conjecture, (X1=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|k18_euclid(X1)!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_6019, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),X1,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_6061, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0)=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,esk2_0))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_6060, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k18_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,esk1_0,k2_nat_1(esk2_0,np__1)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),esk2_0))!=k17_euclid(esk4_0)|~m1_subset_1(k3_matrix_1(u1_struct_0(k15_euclid(np__2)),esk3_0,k2_nat_1(esk1_0,np__1),k2_nat_1(esk2_0,np__1)),u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_6064, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_5959, negated_conjecture, (k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|X1=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k18_euclid(X1)!=k18_euclid(esk4_0)|k17_euclid(X1)!=k17_euclid(esk4_0)|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_6017, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_6020, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk4_0),k17_euclid(X1))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_6062, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_6021, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_6033, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_6063, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_6065, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_6066, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_6072, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_6073, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_6068, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|v2_struct_0(k15_euclid(np__2)))).#trainneg
cnf(i_0_5926, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_6074, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_6075, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))=esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|k18_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk4_0)|k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
cnf(i_0_6032, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5927, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk4_0))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5938, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(X1))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_6046, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_6082, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0)))).#trainneg
cnf(i_0_6102, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(X1))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_5939, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(X1),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k18_euclid(esk4_0),k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_6048, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk6_3(np__2,esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_6096, negated_conjecture, (k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(esk4_0)))).#trainneg
cnf(i_0_6112, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_xxreal_0(k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))),k18_euclid(X1))|k17_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))=k17_euclid(esk4_0)|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk4_0,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_6050, negated_conjecture, (v2_struct_0(k15_euclid(np__2))|r1_sppol_1(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))|r1_xxreal_0(k17_euclid(X1),k17_euclid(esk4_0))|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|~r2_hidden(X1,k1_rltopsp1(k15_euclid(np__2),esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))))|~r1_xxreal_0(k17_euclid(esk5_3(np__2,esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),k3_goboard5(esk3_0,esk1_0,esk2_0))),k17_euclid(esk4_0))|~m1_subset_1(X1,u1_struct_0(k15_euclid(np__2))))).#trainneg
cnf(i_0_6155, negated_conjecture, (esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k18_euclid(esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k18_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))))).#trainneg
cnf(i_0_6162, negated_conjecture, (v2_sppol_1(k1_rltopsp1(k15_euclid(np__2),esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)),esk4_0))|esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|esk6_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0))=esk4_0|v2_struct_0(k15_euclid(np__2))|k17_euclid(esk5_3(np__2,esk4_0,k3_goboard5(esk3_0,esk1_0,esk2_0)))!=k17_euclid(esk4_0))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 37
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 87
# Removed in clause preprocessing      : 1
# Initial clauses in saturation        : 86
# Processed clauses                    : 2309
# ...of these trivial                  : 0
# ...subsumed                          : 927
# ...remaining for further processing  : 1382
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 598
# Backward-rewritten                   : 369
# Generated clauses                    : 6562
# ...of the previous two non-trivial   : 4855
# Contextual simplify-reflections      : 211
# Paramodulations                      : 6537
# Factorizations                       : 0
# Equation resolutions                 : 25
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 415
#    Positive orientable unit clauses  : 130
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 5
#    Non-unit-clauses                  : 280
# Current number of unprocessed clauses: 896
# ...number of literals in the above   : 5021
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2350
# Proof object given clauses           : 316
# Proof search given clauses           : 1382
# Clause-clause subsumption calls (NU) : 1095887
# Rec. Clause-clause subsumption calls : 27500
# Non-unit clause-clause subsumptions  : 1729
# Unit Clause-clause subsumption calls : 4312
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 57
# BW rewrite match successes           : 1
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 357350

# -------------------------------------------------
# User time                : 0.941 s
# System time              : 0.012 s
# Total time               : 0.953 s
# Maximum resident set size: 9844 pages
