# ENIGMATIC: Connected to the TF server '127.0.0.1:8888'.
# ENIGMATIC: LightGBM model './big03/model.lgb' loaded with 17544 features 'C(l,x[c=3],r,h[b=4096],v,c[b=256],d,a):G:P'
# Preprocessing time       : 0.776 s
# ENIGMATIC: LightGBM parental guidance model #1 './froc_gnn_model2/model.lgb' loaded with 13963 features 'C(l,x[c=3],r,h[b=2048],v,c[b=256],d,a):G:F'
#TF#SERVER: Lgb model filter skipped 53 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=435, query=173, total=622).
#TF#SERVER: Sent 238359 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2436 bytes expected).
#TF#SERVER: Received 2436 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 60 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=539, query=201, total=754).
#TF#SERVER: Sent 397981 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (2964 bytes expected).
#TF#SERVER: Received 2964 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 78 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=614, query=181, total=809).
#TF#SERVER: Sent 430426 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3184 bytes expected).
#TF#SERVER: Received 3184 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 77 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=682, query=181, total=877).
#TF#SERVER: Sent 462460 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3456 bytes expected).
#TF#SERVER: Received 3456 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 95 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=733, query=165, total=912).
#TF#SERVER: Sent 563763 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3596 bytes expected).
#TF#SERVER: Received 3596 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 152 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=116, total=898).
#TF#SERVER: Sent 504642 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3540 bytes expected).
#TF#SERVER: Received 3540 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 166 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=90, total=872).
#TF#SERVER: Sent 459067 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3436 bytes expected).
#TF#SERVER: Received 3436 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 122 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=135, total=917).
#TF#SERVER: Sent 524629 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3616 bytes expected).
#TF#SERVER: Received 3616 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 162 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=96, total=878).
#TF#SERVER: Sent 452048 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3460 bytes expected).
#TF#SERVER: Received 3460 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 129 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=157, total=939).
#TF#SERVER: Sent 549396 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3704 bytes expected).
#TF#SERVER: Received 3704 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 153 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=104, total=886).
#TF#SERVER: Sent 483059 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3492 bytes expected).
#TF#SERVER: Received 3492 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 132 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=127, total=909).
#TF#SERVER: Sent 519299 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3584 bytes expected).
#TF#SERVER: Received 3584 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 149 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=107, total=889).
#TF#SERVER: Sent 480432 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3504 bytes expected).
#TF#SERVER: Received 3504 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 130 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=144, total=926).
#TF#SERVER: Sent 502292 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3652 bytes expected).
#TF#SERVER: Received 3652 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 112 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=147, total=929).
#TF#SERVER: Sent 480373 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3664 bytes expected).
#TF#SERVER: Received 3664 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 97 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=163, total=945).
#TF#SERVER: Sent 437759 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3728 bytes expected).
#TF#SERVER: Received 3728 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 88 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=169, total=951).
#TF#SERVER: Sent 455343 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3752 bytes expected).
#TF#SERVER: Received 3752 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 97 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=168, total=950).
#TF#SERVER: Sent 467079 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3748 bytes expected).
#TF#SERVER: Received 3748 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 108 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=153, total=935).
#TF#SERVER: Sent 467531 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3688 bytes expected).
#TF#SERVER: Received 3688 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 125 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=135, total=917).
#TF#SERVER: Sent 466110 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3616 bytes expected).
#TF#SERVER: Received 3616 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 139 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=125, total=907).
#TF#SERVER: Sent 455172 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3576 bytes expected).
#TF#SERVER: Received 3576 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 114 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=167, total=949).
#TF#SERVER: Sent 541048 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3744 bytes expected).
#TF#SERVER: Received 3744 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 143 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=125, total=907).
#TF#SERVER: Sent 454562 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3576 bytes expected).
#TF#SERVER: Received 3576 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 85 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=177, total=959).
#TF#SERVER: Sent 543578 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3784 bytes expected).
#TF#SERVER: Received 3784 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 138 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=124, total=906).
#TF#SERVER: Sent 494297 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3572 bytes expected).
#TF#SERVER: Received 3572 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 134 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=133, total=915).
#TF#SERVER: Sent 509349 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3608 bytes expected).
#TF#SERVER: Received 3608 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 143 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=115, total=897).
#TF#SERVER: Sent 465523 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3536 bytes expected).
#TF#SERVER: Received 3536 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 185 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=90, total=872).
#TF#SERVER: Sent 439424 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3436 bytes expected).
#TF#SERVER: Received 3436 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 179 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=99, total=881).
#TF#SERVER: Sent 458660 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3472 bytes expected).
#TF#SERVER: Received 3472 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 148 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=112, total=894).
#TF#SERVER: Sent 509139 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3524 bytes expected).
#TF#SERVER: Received 3524 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 159 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=97, total=879).
#TF#SERVER: Sent 457362 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3464 bytes expected).
#TF#SERVER: Received 3464 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 156 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=112, total=894).
#TF#SERVER: Sent 486909 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3524 bytes expected).
#TF#SERVER: Received 3524 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 173 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=86, total=868).
#TF#SERVER: Sent 436459 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3420 bytes expected).
#TF#SERVER: Received 3420 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 201 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=89, total=871).
#TF#SERVER: Sent 433633 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3432 bytes expected).
#TF#SERVER: Received 3432 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 184 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=74, total=856).
#TF#SERVER: Sent 439456 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3372 bytes expected).
#TF#SERVER: Received 3372 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 176 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=155, total=937).
#TF#SERVER: Sent 551117 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3696 bytes expected).
#TF#SERVER: Received 3696 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 183 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=100, total=882).
#TF#SERVER: Sent 486606 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3476 bytes expected).
#TF#SERVER: Received 3476 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 167 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=92, total=874).
#TF#SERVER: Sent 476101 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3444 bytes expected).
#TF#SERVER: Received 3444 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 166 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=102, total=884).
#TF#SERVER: Sent 530340 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3484 bytes expected).
#TF#SERVER: Received 3484 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 233 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=66, total=848).
#TF#SERVER: Sent 441444 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3340 bytes expected).
#TF#SERVER: Received 3340 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 149 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=111, total=893).
#TF#SERVER: Sent 464623 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3520 bytes expected).
#TF#SERVER: Received 3520 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 189 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=90, total=872).
#TF#SERVER: Sent 434925 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3436 bytes expected).
#TF#SERVER: Received 3436 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 173 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=93, total=875).
#TF#SERVER: Sent 477448 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3448 bytes expected).
#TF#SERVER: Received 3448 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 217 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=85, total=867).
#TF#SERVER: Sent 453689 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3416 bytes expected).
#TF#SERVER: Received 3416 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 160 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=97, total=879).
#TF#SERVER: Sent 442498 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3464 bytes expected).
#TF#SERVER: Received 3464 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 166 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=103, total=885).
#TF#SERVER: Sent 464067 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3488 bytes expected).
#TF#SERVER: Received 3488 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 134 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=134, total=916).
#TF#SERVER: Sent 510566 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3612 bytes expected).
#TF#SERVER: Received 3612 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 163 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=109, total=891).
#TF#SERVER: Sent 473690 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3512 bytes expected).
#TF#SERVER: Received 3512 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 145 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=116, total=898).
#TF#SERVER: Sent 478134 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3540 bytes expected).
#TF#SERVER: Received 3540 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 98 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=163, total=945).
#TF#SERVER: Sent 512578 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3728 bytes expected).
#TF#SERVER: Received 3728 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 206 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=60, total=842).
#TF#SERVER: Sent 421832 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3316 bytes expected).
#TF#SERVER: Received 3316 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 124 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=177, total=959).
#TF#SERVER: Sent 586759 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3784 bytes expected).
#TF#SERVER: Received 3784 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 137 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=123, total=905).
#TF#SERVER: Sent 505365 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3568 bytes expected).
#TF#SERVER: Received 3568 bytes from the TF server.
#TF#SERVER: Lgb model filter skipped 143 clauses
#TF#SERVER: Sending data to the TF server (conj=14, context=768, query=132, total=914).
#TF#SERVER: Sent 479301 bytes to the TF server.
#TF#SERVER: Receiving data from the TF server (3604 bytes expected).
#TF#SERVER: Received 3604 bytes from the TF server.

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(d6_waybel_6, axiom, ![X1]:((~(v2_struct_0(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v5_waybel_6(X2,X1)<=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~(((r1_orders_2(X1,k11_lattice3(X1,X3,X4),X2)&~(r1_orders_2(X1,X3,X2)))&~(r1_orders_2(X1,X4,X2))))))))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', d6_waybel_6)).
fof(redefinition_k12_lattice3, axiom, ![X1, X2, X3]:(((((v5_orders_2(X1)&v2_lattice3(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', redefinition_k12_lattice3)).
fof(cc2_lattice3, axiom, ![X1]:(l1_orders_2(X1)=>(v2_lattice3(X1)=>~(v2_struct_0(X1)))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', cc2_lattice3)).
fof(t40_yellow_0, axiom, ![X1]:(((((v3_orders_2(X1)&v4_orders_2(X1))&v5_orders_2(X1))&v2_lattice3(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>![X3]:(m1_subset_1(X3,u1_struct_0(X1))=>k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k12_lattice3(X1,X2,X3)))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', t40_yellow_0)).
fof(redefinition_k7_domain_1, axiom, ![X1, X2, X3]:(((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))&m1_subset_1(X3,X1))=>k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', redefinition_k7_domain_1)).
fof(fc2_struct_0, axiom, ![X1]:((~(v2_struct_0(X1))&l1_struct_0(X1))=>~(v1_xboole_0(u1_struct_0(X1)))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', fc2_struct_0)).
fof(dt_l1_orders_2, axiom, ![X1]:(l1_orders_2(X1)=>l1_struct_0(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', dt_l1_orders_2)).
fof(d2_tarski, axiom, ![X1, X2, X3]:(X3=k2_tarski(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(X4=X1|X4=X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', d2_tarski)).
fof(t22_waybel_6, conjecture, ![X1]:(((((v3_orders_2(X1)&v4_orders_2(X1))&v5_orders_2(X1))&v2_lattice3(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v5_waybel_6(X2,X1)<=>![X3]:(((~(v1_xboole_0(X3))&v1_finset_1(X3))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))=>~((r1_orders_2(X1,k2_yellow_0(X1,X3),X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X3)&r1_orders_2(X1,X4,X2)))))))))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', t22_waybel_6)).
fof(fc2_finset_1, axiom, ![X1, X2]:v1_finset_1(k2_tarski(X1,X2)), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', fc2_finset_1)).
fof(fc3_xboole_0, axiom, ![X1, X2]:~(v1_xboole_0(k2_tarski(X1,X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', fc3_xboole_0)).
fof(t4_yellow_2, axiom, ![X1]:((((v4_orders_2(X1)&v5_orders_2(X1))&v2_lattice3(X1))&l1_orders_2(X1))=>![X2, X3]:((r2_yellow_0(X1,X2)&r2_yellow_0(X1,X3))=>(r2_yellow_0(X1,k2_xboole_0(X2,X3))&k2_yellow_0(X1,k2_xboole_0(X2,X3))=k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', t4_yellow_2)).
fof(t39_yellow_0, axiom, ![X1]:((((~(v2_struct_0(X1))&v3_orders_2(X1))&v5_orders_2(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2&k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', t39_yellow_0)).
fof(commutativity_k12_lattice3, axiom, ![X1, X2, X3]:(((((v5_orders_2(X1)&v2_lattice3(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))&m1_subset_1(X3,u1_struct_0(X1)))=>k12_lattice3(X1,X2,X3)=k12_lattice3(X1,X3,X2)), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', commutativity_k12_lattice3)).
fof(redefinition_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>k6_domain_1(X1,X2)=k1_tarski(X2)), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', redefinition_k6_domain_1)).
fof(dt_k2_yellow_0, axiom, ![X1, X2]:(l1_orders_2(X1)=>m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', dt_k2_yellow_0)).
fof(t55_yellow_0, axiom, ![X1]:(((((~(v2_struct_0(X1))&v3_orders_2(X1))&v4_orders_2(X1))&v5_orders_2(X1))&l1_orders_2(X1))=>(v2_lattice3(X1)<=>![X2]:(((~(v1_xboole_0(X2))&v1_finset_1(X2))&m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))=>r2_yellow_0(X1,X2)))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', t55_yellow_0)).
fof(dt_k7_domain_1, axiom, ![X1, X2, X3]:(((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))&m1_subset_1(X3,X1))=>m1_subset_1(k7_domain_1(X1,X2,X3),k1_zfmisc_1(X1))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', dt_k7_domain_1)).
fof(fc1_finset_1, axiom, ![X1]:v1_finset_1(k1_tarski(X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', fc1_finset_1)).
fof(fc2_xboole_0, axiom, ![X1]:~(v1_xboole_0(k1_tarski(X1))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', fc2_xboole_0)).
fof(s2_finset_1__e5_41_1__waybel_6, axiom, ![X1, X2, X3]:(((((((((v3_orders_2(X1)&v4_orders_2(X1))&v5_orders_2(X1))&v2_lattice3(X1))&l1_orders_2(X1))&m1_subset_1(X2,u1_struct_0(X1)))&~(v1_xboole_0(X3)))&v1_finset_1(X3))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))=>(((v1_finset_1(X3)&~(((~(v1_xboole_0(k1_xboole_0))&r1_orders_2(X1,k2_yellow_0(X1,k1_xboole_0),X2))&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,k1_xboole_0)&r1_orders_2(X1,X4,X2)))))))&![X5, X6]:(((r2_hidden(X5,X3)&r1_tarski(X6,X3))&~(((~(v1_xboole_0(X6))&r1_orders_2(X1,k2_yellow_0(X1,X6),X2))&![X7]:(m1_subset_1(X7,u1_struct_0(X1))=>~((r2_hidden(X7,X6)&r1_orders_2(X1,X7,X2)))))))=>~(((~(v1_xboole_0(k2_xboole_0(X6,k1_tarski(X5))))&r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X6,k1_tarski(X5))),X2))&![X8]:(m1_subset_1(X8,u1_struct_0(X1))=>~((r2_hidden(X8,k2_xboole_0(X6,k1_tarski(X5)))&r1_orders_2(X1,X8,X2))))))))=>~(((~(v1_xboole_0(X3))&r1_orders_2(X1,k2_yellow_0(X1,X3),X2))&![X9]:(m1_subset_1(X9,u1_struct_0(X1))=>~((r2_hidden(X9,X3)&r1_orders_2(X1,X9,X2)))))))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', s2_finset_1__e5_41_1__waybel_6)).
fof(d3_xboole_0, axiom, ![X1, X2, X3]:(X3=k2_xboole_0(X1,X2)<=>![X4]:(r2_hidden(X4,X3)<=>(r2_hidden(X4,X1)|r2_hidden(X4,X2)))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', d3_xboole_0)).
fof(d1_tarski, axiom, ![X1, X2]:(X2=k1_tarski(X1)<=>![X3]:(r2_hidden(X3,X2)<=>X3=X1)), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', d1_tarski)).
fof(fc1_xboole_0, axiom, v1_xboole_0(k1_xboole_0), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', fc1_xboole_0)).
fof(idempotence_k2_xboole_0, axiom, ![X1, X2]:k2_xboole_0(X1,X1)=X1, file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', idempotence_k2_xboole_0)).
fof(dt_k6_domain_1, axiom, ![X1, X2]:((~(v1_xboole_0(X1))&m1_subset_1(X2,X1))=>m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', dt_k6_domain_1)).
fof(t1_xboole_1, axiom, ![X1, X2, X3]:((r1_tarski(X1,X2)&r1_tarski(X2,X3))=>r1_tarski(X1,X3)), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', t1_xboole_1)).
fof(t3_subset, axiom, ![X1, X2]:(m1_subset_1(X1,k1_zfmisc_1(X2))<=>r1_tarski(X1,X2)), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', t3_subset)).
fof(t7_boole, axiom, ![X1, X2]:~((r2_hidden(X1,X2)&v1_xboole_0(X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', t7_boole)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', t4_subset)).
fof(reflexivity_r1_tarski, axiom, ![X1, X2]:r1_tarski(X1,X1), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', reflexivity_r1_tarski)).
fof(t2_subset, axiom, ![X1, X2]:(m1_subset_1(X1,X2)=>(v1_xboole_0(X2)|r2_hidden(X1,X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', t2_subset)).
fof(cc2_finset_1, axiom, ![X1]:(v1_finset_1(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_finset_1(X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', cc2_finset_1)).
fof(dt_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)=>m1_subset_1(X3,X1))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', dt_m2_subset_1)).
fof(cc1_subset_1, axiom, ![X1]:(v1_xboole_0(X1)=>![X2]:(m1_subset_1(X2,k1_zfmisc_1(X1))=>v1_xboole_0(X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', cc1_subset_1)).
fof(redefinition_m2_subset_1, axiom, ![X1, X2]:(((~(v1_xboole_0(X1))&~(v1_xboole_0(X2)))&m1_subset_1(X2,k1_zfmisc_1(X1)))=>![X3]:(m2_subset_1(X3,X1,X2)<=>m1_subset_1(X3,X2))), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', redefinition_m2_subset_1)).
fof(t1_subset, axiom, ![X1, X2]:(r2_hidden(X1,X2)=>m1_subset_1(X1,X2)), file('/home/mptp/big1/bushy_np/problems_small_np/t22_waybel_6', t1_subset)).
fof(c_0_37, plain, ![X182, X183, X184, X185]:((~v5_waybel_6(X183,X182)|(~m1_subset_1(X184,u1_struct_0(X182))|(~m1_subset_1(X185,u1_struct_0(X182))|(~r1_orders_2(X182,k11_lattice3(X182,X184,X185),X183)|r1_orders_2(X182,X184,X183)|r1_orders_2(X182,X185,X183))))|~m1_subset_1(X183,u1_struct_0(X182))|(v2_struct_0(X182)|~l1_orders_2(X182)))&((m1_subset_1(esk8_2(X182,X183),u1_struct_0(X182))|v5_waybel_6(X183,X182)|~m1_subset_1(X183,u1_struct_0(X182))|(v2_struct_0(X182)|~l1_orders_2(X182)))&((m1_subset_1(esk9_2(X182,X183),u1_struct_0(X182))|v5_waybel_6(X183,X182)|~m1_subset_1(X183,u1_struct_0(X182))|(v2_struct_0(X182)|~l1_orders_2(X182)))&(((r1_orders_2(X182,k11_lattice3(X182,esk8_2(X182,X183),esk9_2(X182,X183)),X183)|v5_waybel_6(X183,X182)|~m1_subset_1(X183,u1_struct_0(X182))|(v2_struct_0(X182)|~l1_orders_2(X182)))&(~r1_orders_2(X182,esk8_2(X182,X183),X183)|v5_waybel_6(X183,X182)|~m1_subset_1(X183,u1_struct_0(X182))|(v2_struct_0(X182)|~l1_orders_2(X182))))&(~r1_orders_2(X182,esk9_2(X182,X183),X183)|v5_waybel_6(X183,X182)|~m1_subset_1(X183,u1_struct_0(X182))|(v2_struct_0(X182)|~l1_orders_2(X182))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_waybel_6])])])])])])).
fof(c_0_38, plain, ![X325, X326, X327]:(~v5_orders_2(X325)|~v2_lattice3(X325)|~l1_orders_2(X325)|~m1_subset_1(X326,u1_struct_0(X325))|~m1_subset_1(X327,u1_struct_0(X325))|k12_lattice3(X325,X326,X327)=k11_lattice3(X325,X326,X327)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k12_lattice3])])).
cnf(c_0_39, plain, (r1_orders_2(X1,k11_lattice3(X1,esk8_2(X1,X2),esk9_2(X1,X2)),X2)|v5_waybel_6(X2,X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_40, plain, (k12_lattice3(X1,X2,X3)=k11_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_38])).
fof(c_0_41, plain, ![X56]:(~l1_orders_2(X56)|(~v2_lattice3(X56)|~v2_struct_0(X56))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_lattice3])])])).
fof(c_0_42, plain, ![X364, X365, X366]:(~v3_orders_2(X364)|~v4_orders_2(X364)|~v5_orders_2(X364)|~v2_lattice3(X364)|~l1_orders_2(X364)|(~m1_subset_1(X365,u1_struct_0(X364))|(~m1_subset_1(X366,u1_struct_0(X364))|k2_yellow_0(X364,k7_domain_1(u1_struct_0(X364),X365,X366))=k12_lattice3(X364,X365,X366)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t40_yellow_0])])])).
fof(c_0_43, plain, ![X330, X331, X332]:(v1_xboole_0(X330)|~m1_subset_1(X331,X330)|~m1_subset_1(X332,X330)|k7_domain_1(X330,X331,X332)=k2_tarski(X331,X332)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k7_domain_1])])])).
fof(c_0_44, plain, ![X226]:(v2_struct_0(X226)|~l1_struct_0(X226)|~v1_xboole_0(u1_struct_0(X226))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_struct_0])])])).
fof(c_0_45, plain, ![X203]:(~l1_orders_2(X203)|l1_struct_0(X203)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_orders_2])])).
fof(c_0_46, plain, ![X164, X165, X166, X167, X168, X169, X170, X171]:(((~r2_hidden(X167,X166)|(X167=X164|X167=X165)|X166!=k2_tarski(X164,X165))&((X168!=X164|r2_hidden(X168,X166)|X166!=k2_tarski(X164,X165))&(X168!=X165|r2_hidden(X168,X166)|X166!=k2_tarski(X164,X165))))&(((esk6_3(X169,X170,X171)!=X169|~r2_hidden(esk6_3(X169,X170,X171),X171)|X171=k2_tarski(X169,X170))&(esk6_3(X169,X170,X171)!=X170|~r2_hidden(esk6_3(X169,X170,X171),X171)|X171=k2_tarski(X169,X170)))&(r2_hidden(esk6_3(X169,X170,X171),X171)|(esk6_3(X169,X170,X171)=X169|esk6_3(X169,X170,X171)=X170)|X171=k2_tarski(X169,X170)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d2_tarski])])])])])])).
fof(c_0_47, negated_conjecture, ~(![X1]:(((((v3_orders_2(X1)&v4_orders_2(X1))&v5_orders_2(X1))&v2_lattice3(X1))&l1_orders_2(X1))=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(v5_waybel_6(X2,X1)<=>![X3]:(((~(v1_xboole_0(X3))&v1_finset_1(X3))&m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))=>~((r1_orders_2(X1,k2_yellow_0(X1,X3),X2)&![X4]:(m1_subset_1(X4,u1_struct_0(X1))=>~((r2_hidden(X4,X3)&r1_orders_2(X1,X4,X2))))))))))), inference(assume_negation,[status(cth)],[t22_waybel_6])).
cnf(c_0_48, plain, (v2_struct_0(X1)|v5_waybel_6(X2,X1)|r1_orders_2(X1,k11_lattice3(X1,esk8_2(X1,X2),esk9_2(X1,X2)),X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_39])).
cnf(c_0_49, plain, (k11_lattice3(X1,X2,X3)=k12_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_40])).
cnf(c_0_50, plain, (m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|v5_waybel_6(X2,X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_51, plain, (m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|v5_waybel_6(X2,X1)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_52, plain, (~l1_orders_2(X1)|~v2_lattice3(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_41])).
cnf(c_0_53, plain, (k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k12_lattice3(X1,X2,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_42])).
cnf(c_0_54, plain, (v1_xboole_0(X1)|k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_43])).
cnf(c_0_55, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_44])).
cnf(c_0_56, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_45])).
cnf(c_0_57, plain, (X1=X3|X1=X4|~r2_hidden(X1,X2)|X2!=k2_tarski(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_46])).
fof(c_0_58, negated_conjecture, ![X13, X14]:(((((v3_orders_2(esk1_0)&v4_orders_2(esk1_0))&v5_orders_2(esk1_0))&v2_lattice3(esk1_0))&l1_orders_2(esk1_0))&(m1_subset_1(esk2_0,u1_struct_0(esk1_0))&(((((~v1_xboole_0(esk3_0)|~v5_waybel_6(esk2_0,esk1_0))&(v1_finset_1(esk3_0)|~v5_waybel_6(esk2_0,esk1_0)))&(m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v5_waybel_6(esk2_0,esk1_0)))&((r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk2_0)|~v5_waybel_6(esk2_0,esk1_0))&(~m1_subset_1(X13,u1_struct_0(esk1_0))|(~r2_hidden(X13,esk3_0)|~r1_orders_2(esk1_0,X13,esk2_0))|~v5_waybel_6(esk2_0,esk1_0))))&((m1_subset_1(esk4_1(X14),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X14),esk2_0)|(v1_xboole_0(X14)|~v1_finset_1(X14)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(esk1_0))))|v5_waybel_6(esk2_0,esk1_0))&((r2_hidden(esk4_1(X14),X14)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X14),esk2_0)|(v1_xboole_0(X14)|~v1_finset_1(X14)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(esk1_0))))|v5_waybel_6(esk2_0,esk1_0))&(r1_orders_2(esk1_0,esk4_1(X14),esk2_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X14),esk2_0)|(v1_xboole_0(X14)|~v1_finset_1(X14)|~m1_subset_1(X14,k1_zfmisc_1(u1_struct_0(esk1_0))))|v5_waybel_6(esk2_0,esk1_0))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_47])])])])])])).
fof(c_0_59, plain, ![X224, X225]:v1_finset_1(k2_tarski(X224,X225)), inference(variable_rename,[status(thm)],[fc2_finset_1])).
fof(c_0_60, plain, ![X236, X237]:~v1_xboole_0(k2_tarski(X236,X237)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc3_xboole_0])])).
cnf(c_0_61, plain, (v2_struct_0(X1)|r1_orders_2(X1,k12_lattice3(X1,esk8_2(X1,X2),esk9_2(X1,X2)),X2)|v5_waybel_6(X2,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_49])])).
cnf(c_0_62, plain, (v2_struct_0(X1)|v5_waybel_6(X2,X1)|m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_50])).
cnf(c_0_63, plain, (v2_struct_0(X1)|v5_waybel_6(X2,X1)|m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_51])).
cnf(c_0_64, plain, (~v2_lattice3(X1)|~l1_orders_2(X1)|~v2_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_52])).
cnf(c_0_65, plain, (k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k12_lattice3(X1,X2,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_53])).
cnf(c_0_66, plain, (k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)|v1_xboole_0(X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_54])).
cnf(c_0_67, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_55])).
cnf(c_0_68, plain, (l1_struct_0(X1)|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_56])).
fof(c_0_69, plain, ![X370, X371, X372]:((r2_yellow_0(X370,k2_xboole_0(X371,X372))|(~r2_yellow_0(X370,X371)|~r2_yellow_0(X370,X372))|(~v4_orders_2(X370)|~v5_orders_2(X370)|~v2_lattice3(X370)|~l1_orders_2(X370)))&(k2_yellow_0(X370,k2_xboole_0(X371,X372))=k12_lattice3(X370,k2_yellow_0(X370,X371),k2_yellow_0(X370,X372))|(~r2_yellow_0(X370,X371)|~r2_yellow_0(X370,X372))|(~v4_orders_2(X370)|~v5_orders_2(X370)|~v2_lattice3(X370)|~l1_orders_2(X370)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_yellow_2])])])])).
fof(c_0_70, plain, ![X360, X361]:((k1_yellow_0(X360,k6_domain_1(u1_struct_0(X360),X361))=X361|~m1_subset_1(X361,u1_struct_0(X360))|(v2_struct_0(X360)|~v3_orders_2(X360)|~v5_orders_2(X360)|~l1_orders_2(X360)))&(k2_yellow_0(X360,k6_domain_1(u1_struct_0(X360),X361))=X361|~m1_subset_1(X361,u1_struct_0(X360))|(v2_struct_0(X360)|~v3_orders_2(X360)|~v5_orders_2(X360)|~l1_orders_2(X360)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t39_yellow_0])])])])])).
cnf(c_0_71, plain, (X1=X4|X1=X3|X2!=k2_tarski(X3,X4)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_57])).
cnf(c_0_72, negated_conjecture, (r2_hidden(esk4_1(X1),X1)|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_73, plain, (v1_finset_1(k2_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_59])).
cnf(c_0_74, plain, (~v1_xboole_0(k2_tarski(X1,X2))), inference(split_conjunct,[status(thm)],[c_0_60])).
cnf(c_0_75, plain, (r1_orders_2(X1,k12_lattice3(X1,esk8_2(X1,X2),esk9_2(X1,X2)),X2)|v5_waybel_6(X2,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_61, c_0_62]), c_0_63]), c_0_64])).
cnf(c_0_76, plain, (k12_lattice3(X1,X2,X3)=k2_yellow_0(X1,k2_tarski(X2,X3))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66])])).
cnf(c_0_77, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_67, c_0_68])])).
cnf(c_0_78, plain, (k2_yellow_0(X1,k2_xboole_0(X2,X3))=k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,X3)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_69])).
cnf(c_0_79, plain, (k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_70])).
fof(c_0_80, plain, ![X147, X148, X149]:(~v5_orders_2(X147)|~v2_lattice3(X147)|~l1_orders_2(X147)|~m1_subset_1(X148,u1_struct_0(X147))|~m1_subset_1(X149,u1_struct_0(X147))|k12_lattice3(X147,X148,X149)=k12_lattice3(X147,X149,X148)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k12_lattice3])])).
cnf(c_0_81, plain, (X1=X2|X1=X3|~r2_hidden(X1,k2_tarski(X2,X3))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_71])])).
cnf(c_0_82, negated_conjecture, (v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|r2_hidden(esk4_1(X1),X1)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)), inference(evalgc,[status(thm)],[c_0_72])).
cnf(c_0_83, plain, (v1_finset_1(k2_tarski(X1,X2))), inference(evalgc,[status(thm)],[c_0_73])).
cnf(c_0_84, plain, (~v1_xboole_0(k2_tarski(X1,X2))), inference(evalgc,[status(thm)],[c_0_74])).
cnf(c_0_85, plain, (r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|v5_waybel_6(X2,X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_75, c_0_76])])).
cnf(c_0_86, plain, (~l1_orders_2(X1)|~v2_lattice3(X1)|~v1_xboole_0(u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_77])])).
cnf(c_0_87, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_88, negated_conjecture, (l1_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_89, negated_conjecture, (v2_lattice3(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_90, negated_conjecture, (v5_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_91, negated_conjecture, (v4_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_92, negated_conjecture, (v3_orders_2(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_93, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(X2,X3))|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2)), inference(evalgc,[status(thm)],[c_0_78])).
cnf(c_0_94, plain, (k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2|v2_struct_0(X1)|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_79])).
fof(c_0_95, plain, ![X328, X329]:(v1_xboole_0(X328)|~m1_subset_1(X329,X328)|k6_domain_1(X328,X329)=k1_tarski(X329)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_k6_domain_1])])])).
cnf(c_0_96, plain, (k12_lattice3(X1,X2,X3)=k12_lattice3(X1,X3,X2)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))), inference(split_conjunct,[status(thm)],[c_0_80])).
fof(c_0_97, plain, ![X196, X197]:(~l1_orders_2(X196)|m1_subset_1(k2_yellow_0(X196,X197),u1_struct_0(X196))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_yellow_0])])).
cnf(c_0_98, negated_conjecture, (r1_orders_2(esk1_0,esk4_1(X1),esk2_0)|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_99, negated_conjecture, (esk4_1(k2_tarski(X1,X2))=X2|esk4_1(k2_tarski(X1,X2))=X1|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk2_0)|~m1_subset_1(k2_tarski(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_81, c_0_82]), c_0_83])]), c_0_84])])).
cnf(c_0_100, plain, (r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|v5_waybel_6(X2,X1)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_85, c_0_86])).
cnf(c_0_101, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[c_0_87])).
cnf(c_0_102, negated_conjecture, (l1_orders_2(esk1_0)), inference(evalgc,[status(thm)],[c_0_88])).
cnf(c_0_103, negated_conjecture, (v2_lattice3(esk1_0)), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_104, negated_conjecture, (v5_orders_2(esk1_0)), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_105, negated_conjecture, (v4_orders_2(esk1_0)), inference(evalgc,[status(thm)],[c_0_91])).
cnf(c_0_106, negated_conjecture, (v3_orders_2(esk1_0)), inference(evalgc,[status(thm)],[c_0_92])).
cnf(c_0_107, plain, (k2_yellow_0(X1,k2_xboole_0(X2,k6_domain_1(u1_struct_0(X1),X3)))=k12_lattice3(X1,k2_yellow_0(X1,X2),X3)|v2_struct_0(X1)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X3))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_94])])).
cnf(c_0_108, plain, (v1_xboole_0(X1)|k6_domain_1(X1,X2)=k1_tarski(X2)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_95])).
fof(c_0_109, plain, ![X373, X374]:((~v2_lattice3(X373)|(v1_xboole_0(X374)|~v1_finset_1(X374)|~m1_subset_1(X374,k1_zfmisc_1(u1_struct_0(X373)))|r2_yellow_0(X373,X374))|(v2_struct_0(X373)|~v3_orders_2(X373)|~v4_orders_2(X373)|~v5_orders_2(X373)|~l1_orders_2(X373)))&((((~v1_xboole_0(esk73_1(X373))|v2_lattice3(X373)|(v2_struct_0(X373)|~v3_orders_2(X373)|~v4_orders_2(X373)|~v5_orders_2(X373)|~l1_orders_2(X373)))&(v1_finset_1(esk73_1(X373))|v2_lattice3(X373)|(v2_struct_0(X373)|~v3_orders_2(X373)|~v4_orders_2(X373)|~v5_orders_2(X373)|~l1_orders_2(X373))))&(m1_subset_1(esk73_1(X373),k1_zfmisc_1(u1_struct_0(X373)))|v2_lattice3(X373)|(v2_struct_0(X373)|~v3_orders_2(X373)|~v4_orders_2(X373)|~v5_orders_2(X373)|~l1_orders_2(X373))))&(~r2_yellow_0(X373,esk73_1(X373))|v2_lattice3(X373)|(v2_struct_0(X373)|~v3_orders_2(X373)|~v4_orders_2(X373)|~v5_orders_2(X373)|~l1_orders_2(X373))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t55_yellow_0])])])])])])).
cnf(c_0_110, plain, (k12_lattice3(X1,X2,X3)=k12_lattice3(X1,X3,X2)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_96])).
cnf(c_0_111, plain, (m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_97])).
cnf(c_0_112, negated_conjecture, (v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|r1_orders_2(esk1_0,esk4_1(X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)), inference(evalgc,[status(thm)],[c_0_98])).
cnf(c_0_113, negated_conjecture, (esk4_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)))=esk8_2(esk1_0,esk2_0)|esk4_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)))=esk9_2(esk1_0,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])])).
fof(c_0_114, plain, ![X200, X201, X202]:(v1_xboole_0(X200)|~m1_subset_1(X201,X200)|~m1_subset_1(X202,X200)|m1_subset_1(k7_domain_1(X200,X201,X202),k1_zfmisc_1(X200))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k7_domain_1])])])).
cnf(c_0_115, plain, (r1_orders_2(X2,X3,X1)|r1_orders_2(X2,X4,X1)|v2_struct_0(X2)|~v5_waybel_6(X1,X2)|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~r1_orders_2(X2,k11_lattice3(X2,X3,X4),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~l1_orders_2(X2)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_116, plain, (k2_yellow_0(X1,k2_xboole_0(X2,k6_domain_1(u1_struct_0(X1),X3)))=k12_lattice3(X1,k2_yellow_0(X1,X2),X3)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X3))|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_107, c_0_64])).
cnf(c_0_117, plain, (k6_domain_1(X1,X2)=k1_tarski(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_108])).
cnf(c_0_118, plain, (v1_xboole_0(X2)|r2_yellow_0(X1,X2)|v2_struct_0(X1)|~v2_lattice3(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_109])).
fof(c_0_119, plain, ![X221]:v1_finset_1(k1_tarski(X221)), inference(variable_rename,[status(thm)],[fc1_finset_1])).
fof(c_0_120, plain, ![X227]:~v1_xboole_0(k1_tarski(X227)), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[fc2_xboole_0])])).
fof(c_0_121, plain, ![X343, X344, X345, X346, X350]:(((((((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r2_hidden(esk69_3(X343,X344,X345),X345)|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r2_hidden(esk69_3(X343,X344,X345),X345)|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r2_hidden(esk69_3(X343,X344,X345),X345)|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))))&((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_tarski(esk70_3(X343,X344,X345),X345)|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_tarski(esk70_3(X343,X344,X345),X345)|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_tarski(esk70_3(X343,X344,X345),X345)|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343))))))))&(((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(m1_subset_1(esk71_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(m1_subset_1(esk71_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(m1_subset_1(esk71_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))))&(((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r2_hidden(esk71_3(X343,X344,X345),esk70_3(X343,X344,X345))|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r2_hidden(esk71_3(X343,X344,X345),esk70_3(X343,X344,X345))|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r2_hidden(esk71_3(X343,X344,X345),esk70_3(X343,X344,X345))|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))))&((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_orders_2(X343,esk71_3(X343,X344,X345),X344)|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_orders_2(X343,esk71_3(X343,X344,X345),X344)|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_orders_2(X343,esk71_3(X343,X344,X345),X344)|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343))))))))))&((((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(~v1_xboole_0(k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345))))|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(~v1_xboole_0(k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345))))|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(~v1_xboole_0(k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345))))|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))))&((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_orders_2(X343,k2_yellow_0(X343,k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345)))),X344)|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_orders_2(X343,k2_yellow_0(X343,k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345)))),X344)|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_orders_2(X343,k2_yellow_0(X343,k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345)))),X344)|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343))))))))&((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(~m1_subset_1(X350,u1_struct_0(X343))|(~r2_hidden(X350,k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345))))|~r1_orders_2(X343,X350,X344))|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(~m1_subset_1(X350,u1_struct_0(X343))|(~r2_hidden(X350,k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345))))|~r1_orders_2(X343,X350,X344))|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(~m1_subset_1(X350,u1_struct_0(X343))|(~r2_hidden(X350,k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345))))|~r1_orders_2(X343,X350,X344))|(~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))))))&(((((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r2_hidden(esk69_3(X343,X344,X345),X345)|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r2_hidden(esk69_3(X343,X344,X345),X345)|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r2_hidden(esk69_3(X343,X344,X345),X345)|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))))&((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_tarski(esk70_3(X343,X344,X345),X345)|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_tarski(esk70_3(X343,X344,X345),X345)|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_tarski(esk70_3(X343,X344,X345),X345)|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343))))))))&(((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(m1_subset_1(esk71_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(m1_subset_1(esk71_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(m1_subset_1(esk71_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))))&(((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r2_hidden(esk71_3(X343,X344,X345),esk70_3(X343,X344,X345))|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r2_hidden(esk71_3(X343,X344,X345),esk70_3(X343,X344,X345))|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r2_hidden(esk71_3(X343,X344,X345),esk70_3(X343,X344,X345))|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))))&((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_orders_2(X343,esk71_3(X343,X344,X345),X344)|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_orders_2(X343,esk71_3(X343,X344,X345),X344)|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_orders_2(X343,esk71_3(X343,X344,X345),X344)|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343))))))))))&((((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(~v1_xboole_0(k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345))))|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(~v1_xboole_0(k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345))))|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(~v1_xboole_0(k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345))))|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))))&((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_orders_2(X343,k2_yellow_0(X343,k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345)))),X344)|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_orders_2(X343,k2_yellow_0(X343,k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345)))),X344)|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_orders_2(X343,k2_yellow_0(X343,k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345)))),X344)|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343))))))))&((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(~m1_subset_1(X350,u1_struct_0(X343))|(~r2_hidden(X350,k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345))))|~r1_orders_2(X343,X350,X344))|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(~m1_subset_1(X350,u1_struct_0(X343))|(~r2_hidden(X350,k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345))))|~r1_orders_2(X343,X350,X344))|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(~m1_subset_1(X350,u1_struct_0(X343))|(~r2_hidden(X350,k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345))))|~r1_orders_2(X343,X350,X344))|(r1_orders_2(X343,k2_yellow_0(X343,k1_xboole_0),X344)|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343))))))))))&(((((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r2_hidden(esk69_3(X343,X344,X345),X345)|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r2_hidden(esk69_3(X343,X344,X345),X345)|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r2_hidden(esk69_3(X343,X344,X345),X345)|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))))&((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_tarski(esk70_3(X343,X344,X345),X345)|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_tarski(esk70_3(X343,X344,X345),X345)|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_tarski(esk70_3(X343,X344,X345),X345)|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343))))))))&(((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(m1_subset_1(esk71_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(m1_subset_1(esk71_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(m1_subset_1(esk71_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))))&(((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r2_hidden(esk71_3(X343,X344,X345),esk70_3(X343,X344,X345))|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r2_hidden(esk71_3(X343,X344,X345),esk70_3(X343,X344,X345))|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r2_hidden(esk71_3(X343,X344,X345),esk70_3(X343,X344,X345))|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))))&((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_orders_2(X343,esk71_3(X343,X344,X345),X344)|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_orders_2(X343,esk71_3(X343,X344,X345),X344)|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_orders_2(X343,esk71_3(X343,X344,X345),X344)|(v1_xboole_0(esk70_3(X343,X344,X345))|~r1_orders_2(X343,k2_yellow_0(X343,esk70_3(X343,X344,X345)),X344))|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343))))))))))&((((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(~v1_xboole_0(k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345))))|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(~v1_xboole_0(k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345))))|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(~v1_xboole_0(k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345))))|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))))&((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_orders_2(X343,k2_yellow_0(X343,k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345)))),X344)|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_orders_2(X343,k2_yellow_0(X343,k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345)))),X344)|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(r1_orders_2(X343,k2_yellow_0(X343,k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345)))),X344)|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343))))))))&((m1_subset_1(esk72_3(X343,X344,X345),u1_struct_0(X343))|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(~m1_subset_1(X350,u1_struct_0(X343))|(~r2_hidden(X350,k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345))))|~r1_orders_2(X343,X350,X344))|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&((r2_hidden(esk72_3(X343,X344,X345),X345)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(~m1_subset_1(X350,u1_struct_0(X343))|(~r2_hidden(X350,k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345))))|~r1_orders_2(X343,X350,X344))|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))&(r1_orders_2(X343,esk72_3(X343,X344,X345),X344)|(v1_xboole_0(X345)|~r1_orders_2(X343,k2_yellow_0(X343,X345),X344))|(~m1_subset_1(X350,u1_struct_0(X343))|(~r2_hidden(X350,k2_xboole_0(esk70_3(X343,X344,X345),k1_tarski(esk69_3(X343,X344,X345))))|~r1_orders_2(X343,X350,X344))|(~m1_subset_1(X346,u1_struct_0(X343))|(~r2_hidden(X346,k1_xboole_0)|~r1_orders_2(X343,X346,X344))|~v1_finset_1(X345)))|(~v3_orders_2(X343)|~v4_orders_2(X343)|~v5_orders_2(X343)|~v2_lattice3(X343)|~l1_orders_2(X343)|~m1_subset_1(X344,u1_struct_0(X343))|v1_xboole_0(X345)|~v1_finset_1(X345)|~m1_subset_1(X345,k1_zfmisc_1(u1_struct_0(X343)))))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[s2_finset_1__e5_41_1__waybel_6])])])])])])).
cnf(c_0_122, plain, (k12_lattice3(X1,k2_yellow_0(X1,X3),k2_yellow_0(X1,X2))=k2_yellow_0(X1,k2_xboole_0(X2,X3))|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2)|~v4_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_93])])).
cnf(c_0_123, plain, (m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)), inference(evalgc,[status(thm)],[c_0_111])).
cnf(c_0_124, negated_conjecture, (r1_orders_2(esk1_0,esk9_2(esk1_0,esk2_0),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|esk4_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)))=esk8_2(esk1_0,esk2_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0))),esk2_0)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_113]), c_0_83])]), c_0_84])])).
cnf(c_0_125, plain, (v1_xboole_0(X1)|m1_subset_1(k7_domain_1(X1,X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)), inference(split_conjunct,[status(thm)],[c_0_114])).
cnf(c_0_126, plain, (v2_struct_0(X2)|r1_orders_2(X2,X4,X1)|r1_orders_2(X2,X3,X1)|~l1_orders_2(X2)|~v5_waybel_6(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,k11_lattice3(X2,X3,X4),X1)), inference(evalgc,[status(thm)],[c_0_115])).
cnf(c_0_127, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),X3)=k2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(X3)))|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,k1_tarski(X3))|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_117])])).
cnf(c_0_128, plain, (v1_xboole_0(X2)|v2_struct_0(X1)|r2_yellow_0(X1,X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[c_0_118])).
cnf(c_0_129, plain, (v1_finset_1(k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_119])).
cnf(c_0_130, plain, (~v1_xboole_0(k1_tarski(X1))), inference(split_conjunct,[status(thm)],[c_0_120])).
cnf(c_0_131, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~r1_orders_2(X1,X4,X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
fof(c_0_132, plain, ![X173, X174, X175, X176, X177, X178, X179, X180]:(((~r2_hidden(X176,X175)|(r2_hidden(X176,X173)|r2_hidden(X176,X174))|X175!=k2_xboole_0(X173,X174))&((~r2_hidden(X177,X173)|r2_hidden(X177,X175)|X175!=k2_xboole_0(X173,X174))&(~r2_hidden(X177,X174)|r2_hidden(X177,X175)|X175!=k2_xboole_0(X173,X174))))&(((~r2_hidden(esk7_3(X178,X179,X180),X178)|~r2_hidden(esk7_3(X178,X179,X180),X180)|X180=k2_xboole_0(X178,X179))&(~r2_hidden(esk7_3(X178,X179,X180),X179)|~r2_hidden(esk7_3(X178,X179,X180),X180)|X180=k2_xboole_0(X178,X179)))&(r2_hidden(esk7_3(X178,X179,X180),X180)|(r2_hidden(esk7_3(X178,X179,X180),X178)|r2_hidden(esk7_3(X178,X179,X180),X179))|X180=k2_xboole_0(X178,X179)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d3_xboole_0])])])])])])).
fof(c_0_133, plain, ![X157, X158, X159, X160, X161, X162]:(((~r2_hidden(X159,X158)|X159=X157|X158!=k1_tarski(X157))&(X160!=X157|r2_hidden(X160,X158)|X158!=k1_tarski(X157)))&((~r2_hidden(esk5_2(X161,X162),X162)|esk5_2(X161,X162)!=X161|X162=k1_tarski(X161))&(r2_hidden(esk5_2(X161,X162),X162)|esk5_2(X161,X162)=X161|X162=k1_tarski(X161)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(fof_nnf,[status(thm)],[d1_tarski])])])])])])).
cnf(c_0_134, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(X3,X2))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_122, c_0_123]), c_0_123])).
cnf(c_0_135, negated_conjecture, (r1_orders_2(esk1_0,esk8_2(esk1_0,esk2_0),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|r1_orders_2(esk1_0,esk9_2(esk1_0,esk2_0),esk2_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0))),esk2_0)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_124]), c_0_83])]), c_0_84])])).
cnf(c_0_136, plain, (v1_xboole_0(X1)|m1_subset_1(k7_domain_1(X1,X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_125])).
cnf(c_0_137, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X3)|r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,k12_lattice3(X1,X2,X4),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_49])])).
cnf(c_0_138, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),X3)=k2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(X3)))|~r2_yellow_0(X1,k1_tarski(X3))|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_127, c_0_86])).
cnf(c_0_139, plain, (r2_yellow_0(X1,X2)|v1_xboole_0(X2)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_128, c_0_64])).
cnf(c_0_140, plain, (v1_finset_1(k1_tarski(X1))), inference(evalgc,[status(thm)],[c_0_129])).
cnf(c_0_141, plain, (~v1_xboole_0(k1_tarski(X1))), inference(evalgc,[status(thm)],[c_0_130])).
cnf(c_0_142, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_143, plain, (v1_xboole_0(X3)|r2_hidden(esk72_3(X1,X2,X3),X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r2_hidden(X4,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))), inference(cn,[status(thm)],[c_0_131])).
cnf(c_0_144, plain, (v1_xboole_0(k1_xboole_0)), inference(split_conjunct,[status(thm)],[fc1_xboole_0])).
cnf(c_0_145, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X4,X2)), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_146, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(split_conjunct,[status(thm)],[c_0_133])).
cnf(c_0_147, plain, (k2_yellow_0(X1,k1_tarski(X2))=X2|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_117])])).
cnf(c_0_148, plain, (k2_yellow_0(X1,k2_xboole_0(X3,k6_domain_1(u1_struct_0(X1),X2)))=k12_lattice3(X1,X2,k2_yellow_0(X1,X3))|v2_struct_0(X1)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_94])])).
fof(c_0_149, plain, ![X249]:k2_xboole_0(X249,X249)=X249, inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[idempotence_k2_xboole_0])])).
cnf(c_0_150, plain, (v5_waybel_6(X2,X1)|v2_struct_0(X1)|~r1_orders_2(X1,esk9_2(X1,X2),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_151, negated_conjecture, (r1_orders_2(esk1_0,esk9_2(esk1_0,esk2_0),esk2_0)|r1_orders_2(esk1_0,esk8_2(esk1_0,esk2_0),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_135, c_0_100]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])])).
cnf(c_0_152, plain, (v1_xboole_0(X1)|m1_subset_1(k2_tarski(X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_136, c_0_66])])).
cnf(c_0_153, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_86, c_0_103]), c_0_102])])])).
cnf(c_0_154, plain, (r1_orders_2(X1,X2,X3)|r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,k12_lattice3(X1,X4,X2),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)), inference(csr,[status(thm)],[c_0_137, c_0_64])).
cnf(c_0_155, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),X3)=k2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(X3)))|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_139]), c_0_140])]), c_0_141])])).
cnf(c_0_156, plain, (v1_xboole_0(X3)|r2_hidden(esk72_3(X1,X2,X3),X3)|r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)), inference(cn,[status(thm)],[c_0_142])).
cnf(c_0_157, plain, (v1_xboole_0(X3)|r2_hidden(esk72_3(X1,X2,X3),X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r2_hidden(X4,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))), inference(evalgc,[status(thm)],[c_0_143])).
cnf(c_0_158, plain, (v1_xboole_0(k1_xboole_0)), inference(evalgc,[status(thm)],[c_0_144])).
cnf(c_0_159, plain, (r2_hidden(X1,X3)|X3!=k2_xboole_0(X4,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_145])).
cnf(c_0_160, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2)), inference(evalgc,[status(thm)],[c_0_146])).
cnf(c_0_161, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(X3,k1_tarski(X2)))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,k1_tarski(X2))|~r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_147])])).
cnf(c_0_162, plain, (k2_yellow_0(X1,k2_xboole_0(X2,k6_domain_1(u1_struct_0(X1),X3)))=k12_lattice3(X1,X3,k2_yellow_0(X1,X2))|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X3))|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_148, c_0_64])).
cnf(c_0_163, plain, (k2_xboole_0(X1,X1)=X1), inference(split_conjunct,[status(thm)],[c_0_149])).
cnf(c_0_164, plain, (v5_waybel_6(X2,X1)|v2_struct_0(X1)|~r1_orders_2(X1,esk8_2(X1,X2),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)), inference(split_conjunct,[status(thm)],[c_0_37])).
cnf(c_0_165, plain, (v2_struct_0(X1)|v5_waybel_6(X2,X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,esk9_2(X1,X2),X2)), inference(evalgc,[status(thm)],[c_0_150])).
cnf(c_0_166, negated_conjecture, (r1_orders_2(esk1_0,esk8_2(esk1_0,esk2_0),esk2_0)|r1_orders_2(esk1_0,esk9_2(esk1_0,esk2_0),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_152]), c_0_153])])).
cnf(c_0_167, plain, (r1_orders_2(X1,X2,X3)|r1_orders_2(X1,k2_yellow_0(X1,X4),X3)|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X4,k1_tarski(X2))),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~r2_yellow_0(X1,X4)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_155])])).
cnf(c_0_168, plain, (v1_xboole_0(X3)|r2_hidden(esk72_3(X1,X2,X3),X3)|r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)), inference(evalgc,[status(thm)],[c_0_156])).
cnf(c_0_169, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r2_hidden(X4,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,X4,X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_157, c_0_158])])).
cnf(c_0_170, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X3)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_159])])).
cnf(c_0_171, plain, (r2_hidden(X1,X2)|X2!=k1_tarski(X1)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_160])])).
cnf(c_0_172, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_173, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(X3,k1_tarski(X2)))|~r2_yellow_0(X1,k1_tarski(X2))|~r2_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_161, c_0_86]), c_0_64])).
cnf(c_0_174, plain, (k2_yellow_0(X1,k2_xboole_0(X2,k6_domain_1(u1_struct_0(X1),X3)))=k12_lattice3(X1,X3,k2_yellow_0(X1,X2))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_139])])).
cnf(c_0_175, plain, (k2_xboole_0(X1,X1)=X1), inference(evalgc,[status(thm)],[c_0_163])).
fof(c_0_176, plain, ![X198, X199]:(v1_xboole_0(X198)|~m1_subset_1(X199,X198)|m1_subset_1(k6_domain_1(X198,X199),k1_zfmisc_1(X198))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k6_domain_1])])])).
cnf(c_0_177, plain, (v2_struct_0(X1)|v5_waybel_6(X2,X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,esk8_2(X1,X2),X2)), inference(evalgc,[status(thm)],[c_0_164])).
cnf(c_0_178, negated_conjecture, (v2_struct_0(esk1_0)|v5_waybel_6(esk2_0,esk1_0)|r1_orders_2(esk1_0,esk8_2(esk1_0,esk2_0),esk2_0)|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_166]), c_0_101]), c_0_102])])])).
cnf(c_0_179, plain, (r1_orders_2(X1,k2_yellow_0(X1,X2),X3)|r1_orders_2(X1,X4,X3)|~r2_yellow_0(X1,X2)|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(X4))),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(k1_tarski(X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_167, c_0_123])).
cnf(c_0_180, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_168, c_0_158])])).
cnf(c_0_181, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,X4,X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_hidden(X4,k1_tarski(esk69_3(X1,X2,X3)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_170])])).
cnf(c_0_182, plain, (r2_hidden(X1,k1_tarski(X1))), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_171])])).
cnf(c_0_183, plain, (v1_xboole_0(X3)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)), inference(cn,[status(thm)],[c_0_172])).
cnf(c_0_184, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(X3,k1_tarski(X2)))|~r2_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_139]), c_0_140])]), c_0_141])])).
cnf(c_0_185, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2)))=k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X2))|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_175])])).
cnf(c_0_186, plain, (v1_xboole_0(X1)|m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_176])).
cnf(c_0_187, negated_conjecture, (v2_struct_0(esk1_0)|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_178]), c_0_101]), c_0_102])])])).
cnf(c_0_188, plain, (r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|r1_orders_2(X1,esk69_3(X1,X2,X3),X2)|r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r2_yellow_0(X1,esk70_3(X1,X2,X3))|~v5_waybel_6(X2,X1)|~m1_subset_1(k1_tarski(esk69_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_180])])).
cnf(c_0_189, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,esk69_3(X1,X2,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_182])])).
cnf(c_0_190, plain, (v1_xboole_0(X3)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)), inference(evalgc,[status(thm)],[c_0_183])).
cnf(c_0_191, plain, (k2_yellow_0(X1,k2_xboole_0(k6_domain_1(u1_struct_0(X1),X3),k1_tarski(X2)))=k12_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_94])])).
cnf(c_0_192, plain, (k12_lattice3(X1,X3,X2)=k2_yellow_0(X1,k2_tarski(X2,X3))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_76])])).
cnf(c_0_193, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2)))=k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X2))|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_185, c_0_139])).
cnf(c_0_194, plain, (v1_xboole_0(X1)|m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[c_0_186])).
cnf(c_0_195, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|r1_tarski(esk70_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_196, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v5_waybel_6(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_197, negated_conjecture, (v2_struct_0(esk1_0)|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_63]), c_0_101]), c_0_102])])])).
cnf(c_0_198, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|v1_xboole_0(X3)|~r2_yellow_0(X1,esk70_3(X1,X2,X3))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(k1_tarski(esk69_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_188, c_0_189])).
cnf(c_0_199, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~r1_orders_2(X1,X4,X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_200, plain, (r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_190, c_0_158])])).
cnf(c_0_201, plain, (k2_yellow_0(X1,k2_xboole_0(k6_domain_1(u1_struct_0(X1),X2),k1_tarski(X3)))=k12_lattice3(X1,X3,X2)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_191, c_0_64])).
cnf(c_0_202, plain, (k12_lattice3(X1,X2,X3)=k2_yellow_0(X1,k2_tarski(X3,X2))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_192, c_0_86])).
cnf(c_0_203, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,k1_tarski(X2)))=k2_yellow_0(X1,k1_tarski(X2))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_117]), c_0_140])]), c_0_141])])).
cnf(c_0_204, plain, (v1_xboole_0(X1)|m1_subset_1(k1_tarski(X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_117])])).
fof(c_0_205, plain, ![X355, X356, X357]:(~r1_tarski(X355,X356)|~r1_tarski(X356,X357)|r1_tarski(X355,X357)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_xboole_1])])).
fof(c_0_206, plain, ![X362, X363]:((~m1_subset_1(X362,k1_zfmisc_1(X363))|r1_tarski(X362,X363))&(~r1_tarski(X362,X363)|m1_subset_1(X362,k1_zfmisc_1(X363)))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t3_subset])])).
cnf(c_0_207, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|r1_tarski(esk70_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_208, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk2_0)|~v5_waybel_6(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_209, negated_conjecture, (v1_finset_1(esk3_0)|~v5_waybel_6(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_210, negated_conjecture, (~v1_xboole_0(esk3_0)|~v5_waybel_6(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_211, plain, (v1_xboole_0(X3)|r1_tarski(esk70_3(X1,X2,X3),X3)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)), inference(cn,[status(thm)],[c_0_195])).
cnf(c_0_212, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v5_waybel_6(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_196])).
cnf(c_0_213, negated_conjecture, (v2_struct_0(esk1_0)|v5_waybel_6(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_62]), c_0_101]), c_0_102])])])).
cnf(c_0_214, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(k1_tarski(esk69_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(esk70_3(X1,X2,X3))|~m1_subset_1(esk70_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_139])])).
cnf(c_0_215, plain, (v1_xboole_0(X3)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r2_hidden(X4,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))), inference(cn,[status(thm)],[c_0_199])).
cnf(c_0_216, plain, (r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|r1_orders_2(X1,esk69_3(X1,X2,X3),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r2_yellow_0(X1,esk70_3(X1,X2,X3))|~v5_waybel_6(X2,X1)|~m1_subset_1(k1_tarski(esk69_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_200])])).
cnf(c_0_217, plain, (k2_yellow_0(X1,k2_xboole_0(k1_tarski(X2),k1_tarski(X3)))=k12_lattice3(X1,X3,X2)|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_117])])).
cnf(c_0_218, plain, (k12_lattice3(X1,X3,X2)=k2_yellow_0(X1,k2_tarski(X3,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_202])])).
cnf(c_0_219, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,k1_tarski(X2)))=k2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_203, c_0_204]), c_0_86])).
cnf(c_0_220, plain, (r1_tarski(X1,X3)|~r1_tarski(X1,X2)|~r1_tarski(X2,X3)), inference(split_conjunct,[status(thm)],[c_0_205])).
cnf(c_0_221, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(split_conjunct,[status(thm)],[c_0_206])).
cnf(c_0_222, plain, (v1_xboole_0(X3)|r2_hidden(esk72_3(X1,X2,X3),X3)|r1_tarski(esk70_3(X1,X2,X3),X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)), inference(cn,[status(thm)],[c_0_207])).
cnf(c_0_223, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk2_0)|~v5_waybel_6(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_208])).
cnf(c_0_224, negated_conjecture, (v1_finset_1(esk3_0)|~v5_waybel_6(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_209])).
cnf(c_0_225, negated_conjecture, (~v1_xboole_0(esk3_0)|~v5_waybel_6(esk2_0,esk1_0)), inference(evalgc,[status(thm)],[c_0_210])).
cnf(c_0_226, plain, (m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|v1_xboole_0(X3)|r1_tarski(esk70_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_227, plain, (v1_xboole_0(X3)|r1_tarski(esk70_3(X1,X2,X3),X3)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)), inference(evalgc,[status(thm)],[c_0_211])).
cnf(c_0_228, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_213])])).
cnf(c_0_229, negated_conjecture, (~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk3_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~v5_waybel_6(esk2_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_58])).
cnf(c_0_230, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(esk70_3(X1,X2,X3))|~v1_finset_1(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(esk70_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_204])])).
cnf(c_0_231, plain, (v1_xboole_0(X3)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r2_hidden(X4,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))), inference(evalgc,[status(thm)],[c_0_215])).
cnf(c_0_232, plain, (r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|r1_orders_2(X1,esk69_3(X1,X2,X3),X2)|v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(k1_tarski(esk69_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(esk70_3(X1,X2,X3))|~m1_subset_1(esk70_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_139])])).
cnf(c_0_233, plain, (r2_hidden(esk7_3(X1,X2,X3),X3)|r2_hidden(esk7_3(X1,X2,X3),X1)|r2_hidden(esk7_3(X1,X2,X3),X2)|X3=k2_xboole_0(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_234, plain, (k2_yellow_0(X1,k2_xboole_0(k1_tarski(X2),k1_tarski(X3)))=k12_lattice3(X1,X3,X2)|~r2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_217, c_0_204]), c_0_86])).
cnf(c_0_235, plain, (k2_yellow_0(X1,k2_tarski(X2,k2_yellow_0(X1,k1_tarski(X2))))=k2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,k1_tarski(X2)),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_219])])).
cnf(c_0_236, plain, (r1_tarski(X1,X3)|~r1_tarski(X2,X3)|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_220])).
cnf(c_0_237, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[c_0_221])).
cnf(c_0_238, plain, (v1_xboole_0(X3)|r2_hidden(esk72_3(X1,X2,X3),X3)|r1_tarski(esk70_3(X1,X2,X3),X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)), inference(evalgc,[status(thm)],[c_0_222])).
cnf(c_0_239, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk2_0)|v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_213])])).
cnf(c_0_240, negated_conjecture, (v1_finset_1(esk3_0)|v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_213])])).
cnf(c_0_241, negated_conjecture, (v2_struct_0(esk1_0)|~v1_xboole_0(esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_213])])).
cnf(c_0_242, plain, (v1_xboole_0(X3)|r1_tarski(esk70_3(X1,X2,X3),X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)), inference(cn,[status(thm)],[c_0_226])).
cnf(c_0_243, plain, (r1_tarski(esk70_3(X1,X2,X3),X3)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_227, c_0_158])])).
cnf(c_0_244, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_228]), c_0_102]), c_0_103])])])).
cnf(c_0_245, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|r2_hidden(esk71_3(X1,X2,X3),esk70_3(X1,X2,X3))|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_246, negated_conjecture, (~r2_hidden(X1,esk3_0)|~v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,X1,esk2_0)), inference(evalgc,[status(thm)],[c_0_229])).
cnf(c_0_247, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(esk70_3(X1,X2,X3))|~v1_finset_1(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(esk70_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_230, c_0_86])).
cnf(c_0_248, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|r2_hidden(esk71_3(X1,X2,X3),esk70_3(X1,X2,X3))|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_249, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r2_hidden(X4,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,X4,X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_231, c_0_158])])).
cnf(c_0_250, plain, (r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|r1_orders_2(X1,esk69_3(X1,X2,X3),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(esk70_3(X1,X2,X3))|~v1_finset_1(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(esk70_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_204])])).
cnf(c_0_251, plain, (X3=k2_xboole_0(X1,X2)|~r2_hidden(esk7_3(X1,X2,X3),X2)|~r2_hidden(esk7_3(X1,X2,X3),X3)), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_252, plain, (X3=k2_xboole_0(X1,X2)|r2_hidden(esk7_3(X1,X2,X3),X3)|r2_hidden(esk7_3(X1,X2,X3),X2)|r2_hidden(esk7_3(X1,X2,X3),X1)), inference(evalgc,[status(thm)],[c_0_233])).
cnf(c_0_253, plain, (k12_lattice3(X1,X2,X2)=k2_yellow_0(X1,k1_tarski(X2))|~r2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_175])])).
cnf(c_0_254, plain, (k2_yellow_0(X1,k2_tarski(X2,k2_yellow_0(X1,k1_tarski(X2))))=k2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_235, c_0_123])).
cnf(c_0_255, plain, (m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|v1_xboole_0(X3)|r2_hidden(esk69_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_256, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_206])).
cnf(c_0_257, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_237])])).
cnf(c_0_258, plain, (r1_tarski(esk70_3(X1,X2,X3),X3)|r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_238, c_0_158])])).
cnf(c_0_259, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_239]), c_0_102]), c_0_103])])])).
cnf(c_0_260, negated_conjecture, (v1_finset_1(esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_240]), c_0_102]), c_0_103])])])).
cnf(c_0_261, negated_conjecture, (~v1_xboole_0(esk3_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_241]), c_0_102]), c_0_103])])])).
cnf(c_0_262, plain, (v1_xboole_0(X3)|r1_tarski(esk70_3(X1,X2,X3),X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)), inference(evalgc,[status(thm)],[c_0_242])).
cnf(c_0_263, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,esk3_0),esk3_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,esk3_0),X1)|v1_xboole_0(esk3_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),X1)|~v1_finset_1(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_244]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])])).
cnf(c_0_264, plain, (r2_hidden(X1,X3)|~r2_hidden(X1,X2)|X3!=k2_xboole_0(X2,X4)), inference(split_conjunct,[status(thm)],[c_0_132])).
cnf(c_0_265, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r2_hidden(esk72_3(X1,X2,X3),X3)|r2_hidden(esk71_3(X1,X2,X3),esk70_3(X1,X2,X3))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)), inference(cn,[status(thm)],[c_0_245])).
cnf(c_0_266, negated_conjecture, (v2_struct_0(esk1_0)|~r2_hidden(X1,esk3_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_213])])).
cnf(c_0_267, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,X1),X1)|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,X1)),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~v1_finset_1(X1)|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_213]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])])).
cnf(c_0_268, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|r1_orders_2(X1,esk71_3(X1,X2,X3),X2)|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_269, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|r2_hidden(esk69_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_270, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|r2_hidden(esk71_3(X1,X2,X3),esk70_3(X1,X2,X3))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)), inference(cn,[status(thm)],[c_0_248])).
cnf(c_0_271, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,X4,X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_hidden(X4,k1_tarski(esk69_3(X1,X2,X3)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_170])])).
cnf(c_0_272, plain, (r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|r1_orders_2(X1,esk69_3(X1,X2,X3),X2)|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(esk70_3(X1,X2,X3))|~v1_finset_1(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(esk70_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_250, c_0_86])).
cnf(c_0_273, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|r1_orders_2(X1,esk71_3(X1,X2,X3),X2)|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
fof(c_0_274, plain, ![X380, X381]:(~r2_hidden(X380,X381)|~v1_xboole_0(X381)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t7_boole])])).
cnf(c_0_275, plain, (X3=k2_xboole_0(X1,X2)|~r2_hidden(esk7_3(X1,X2,X3),X3)|~r2_hidden(esk7_3(X1,X2,X3),X2)), inference(evalgc,[status(thm)],[c_0_251])).
cnf(c_0_276, plain, (k2_xboole_0(X11,X10)=X10|r2_hidden(esk7_3(X11,X10,X10),X11)|r2_hidden(esk7_3(X11,X10,X10),X10)), inference(evalgc,[status(thm)],[inference(ef,[status(thm)],[c_0_252])])).
cnf(c_0_277, plain, (k12_lattice3(X1,X2,X2)=k2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_139]), c_0_140])]), c_0_141])])).
cnf(c_0_278, plain, (k2_yellow_0(X1,k2_tarski(X2,X2))=X2|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_254, c_0_147])])).
cnf(c_0_279, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|r2_hidden(esk69_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_280, plain, (m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|v1_xboole_0(X3)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~m1_subset_1(X4,u1_struct_0(X1))|~r2_hidden(X4,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~r1_orders_2(X1,X4,X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
fof(c_0_281, plain, ![X367, X368, X369]:(~r2_hidden(X367,X368)|~m1_subset_1(X368,k1_zfmisc_1(X369))|m1_subset_1(X367,X369)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
cnf(c_0_282, plain, (v1_xboole_0(X3)|r2_hidden(esk69_3(X1,X2,X3),X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)), inference(cn,[status(thm)],[c_0_255])).
cnf(c_0_283, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2)), inference(evalgc,[status(thm)],[c_0_256])).
cnf(c_0_284, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~r1_tarski(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_244])])).
cnf(c_0_285, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk3_0),esk3_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_259]), c_0_260]), c_0_244]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])]), c_0_261])])).
cnf(c_0_286, plain, (r1_tarski(esk70_3(X1,X2,X3),X3)|v1_xboole_0(X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_262, c_0_158])])).
cnf(c_0_287, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,esk3_0),esk3_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,esk3_0),X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_263, c_0_260])]), c_0_261])).
cnf(c_0_288, plain, (r2_hidden(X1,X3)|X3!=k2_xboole_0(X2,X4)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_264])).
cnf(c_0_289, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r2_hidden(esk72_3(X1,X2,X3),X3)|r2_hidden(esk71_3(X1,X2,X3),esk70_3(X1,X2,X3))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)), inference(evalgc,[status(thm)],[c_0_265])).
cnf(c_0_290, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk3_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_finset_1(esk3_0)|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_267])])).
cnf(c_0_291, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r2_hidden(esk72_3(X1,X2,X3),X3)|r1_orders_2(X1,esk71_3(X1,X2,X3),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)), inference(cn,[status(thm)],[c_0_268])).
cnf(c_0_292, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|m1_subset_1(esk71_3(X1,X2,X3),u1_struct_0(X1))|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_293, plain, (v1_xboole_0(X3)|r2_hidden(esk69_3(X1,X2,X3),X3)|r2_hidden(esk72_3(X1,X2,X3),X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)), inference(cn,[status(thm)],[c_0_269])).
cnf(c_0_294, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|m1_subset_1(esk71_3(X1,X2,X3),u1_struct_0(X1))|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_295, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|r2_hidden(esk71_3(X1,X2,X3),esk70_3(X1,X2,X3))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)), inference(evalgc,[status(thm)],[c_0_270])).
cnf(c_0_296, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,esk69_3(X1,X2,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_271, c_0_182])])).
cnf(c_0_297, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,X1)),esk2_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~v1_finset_1(X1)|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_272, c_0_213]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])])).
cnf(c_0_298, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r1_orders_2(X1,esk71_3(X1,X2,X3),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)), inference(cn,[status(thm)],[c_0_273])).
cnf(c_0_299, plain, (~r2_hidden(X1,X2)|~v1_xboole_0(X2)), inference(split_conjunct,[status(thm)],[c_0_274])).
cnf(c_0_300, plain, (k2_xboole_0(X2,X1)=X1|r2_hidden(esk7_3(X2,X1,X1),X2)|~r2_hidden(esk7_3(X2,X1,X1),X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_276])])).
cnf(c_0_301, plain, (k2_yellow_0(X1,k2_tarski(X2,X2))=k2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_277])])).
cnf(c_0_302, plain, (k2_yellow_0(X1,k2_tarski(X2,X2))=X2|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_278, c_0_86]), c_0_64])).
cnf(c_0_303, plain, (v1_xboole_0(X3)|r2_hidden(esk69_3(X1,X2,X3),X3)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)), inference(cn,[status(thm)],[c_0_279])).
fof(c_0_304, plain, ![X339]:r1_tarski(X339,X339), inference(variable_rename,[status(thm)],[inference(fof_simplification,[status(thm)],[reflexivity_r1_tarski])])).
cnf(c_0_305, plain, (v1_xboole_0(X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r2_hidden(X4,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))), inference(cn,[status(thm)],[c_0_280])).
fof(c_0_306, plain, ![X358, X359]:(~m1_subset_1(X358,X359)|(v1_xboole_0(X359)|r2_hidden(X358,X359))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t2_subset])])).
cnf(c_0_307, plain, (m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|v1_xboole_0(X3)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_308, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_281])).
cnf(c_0_309, plain, (v1_xboole_0(X3)|r2_hidden(esk69_3(X1,X2,X3),X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)), inference(evalgc,[status(thm)],[c_0_282])).
cnf(c_0_310, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_tarski(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_284])])).
fof(c_0_311, plain, ![X52, X53]:(~v1_finset_1(X52)|(~m1_subset_1(X53,k1_zfmisc_1(X52))|v1_finset_1(X53))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc2_finset_1])])])).
cnf(c_0_312, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_285])])).
cnf(c_0_313, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk3_0),esk3_0)|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_286, c_0_259]), c_0_260]), c_0_244]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])]), c_0_261])])).
cnf(c_0_314, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk3_0),esk3_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_287, c_0_259]), c_0_101])])])).
fof(c_0_315, plain, ![X204, X205, X206]:(v1_xboole_0(X204)|v1_xboole_0(X205)|~m1_subset_1(X205,k1_zfmisc_1(X204))|(~m2_subset_1(X206,X204,X205)|m1_subset_1(X206,X204))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_m2_subset_1])])])])).
fof(c_0_316, plain, ![X39, X40]:(~v1_xboole_0(X39)|(~m1_subset_1(X40,k1_zfmisc_1(X39))|v1_xboole_0(X40))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc1_subset_1])])])).
fof(c_0_317, plain, ![X333, X334, X335]:((~m2_subset_1(X335,X333,X334)|m1_subset_1(X335,X334)|(v1_xboole_0(X333)|v1_xboole_0(X334)|~m1_subset_1(X334,k1_zfmisc_1(X333))))&(~m1_subset_1(X335,X334)|m2_subset_1(X335,X333,X334)|(v1_xboole_0(X333)|v1_xboole_0(X334)|~m1_subset_1(X334,k1_zfmisc_1(X333))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[redefinition_m2_subset_1])])])])])).
cnf(c_0_318, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[inference(er,[status(thm)],[c_0_288])])).
cnf(c_0_319, plain, (r2_hidden(esk71_3(X1,X2,X3),esk70_3(X1,X2,X3))|r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_289, c_0_158])])).
cnf(c_0_320, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_290, c_0_259]), c_0_260]), c_0_244])]), c_0_261])).
cnf(c_0_321, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r2_hidden(esk72_3(X1,X2,X3),X3)|r1_orders_2(X1,esk71_3(X1,X2,X3),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)), inference(evalgc,[status(thm)],[c_0_291])).
cnf(c_0_322, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r2_hidden(esk72_3(X1,X2,X3),X3)|m1_subset_1(esk71_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)), inference(cn,[status(thm)],[c_0_292])).
cnf(c_0_323, plain, (v1_xboole_0(X3)|r2_hidden(esk69_3(X1,X2,X3),X3)|r2_hidden(esk72_3(X1,X2,X3),X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)), inference(evalgc,[status(thm)],[c_0_293])).
cnf(c_0_324, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|m1_subset_1(esk71_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)), inference(cn,[status(thm)],[c_0_294])).
cnf(c_0_325, plain, (r2_hidden(esk71_3(X1,X2,X3),esk70_3(X1,X2,X3))|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_295, c_0_158])])).
cnf(c_0_326, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(X1)|v2_struct_0(esk1_0)|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,X1)),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_296, c_0_297]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])])).
cnf(c_0_327, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r1_orders_2(X1,esk71_3(X1,X2,X3),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)), inference(evalgc,[status(thm)],[c_0_298])).
cnf(c_0_328, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_299])).
cnf(c_0_329, plain, (k2_xboole_0(X1,X2)=X2|r2_hidden(esk7_3(X1,X2,X2),X1)), inference(csr,[status(thm)],[c_0_300, c_0_276])).
cnf(c_0_330, plain, (k2_yellow_0(X1,k1_tarski(X2))=X2|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_301, c_0_302])])).
cnf(c_0_331, plain, (v1_xboole_0(X3)|r2_hidden(esk69_3(X1,X2,X3),X3)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)), inference(evalgc,[status(thm)],[c_0_303])).
cnf(c_0_332, plain, (r1_tarski(X1,X1)), inference(split_conjunct,[status(thm)],[c_0_304])).
cnf(c_0_333, plain, (v1_xboole_0(X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r2_hidden(X4,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))), inference(evalgc,[status(thm)],[c_0_305])).
cnf(c_0_334, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_306])).
cnf(c_0_335, plain, (v1_xboole_0(X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))), inference(cn,[status(thm)],[c_0_307])).
fof(c_0_336, plain, ![X353, X354]:(~r2_hidden(X353,X354)|m1_subset_1(X353,X354)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_subset])])).
cnf(c_0_337, plain, (m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|v1_xboole_0(X3)|r2_hidden(esk71_3(X1,X2,X3),esk70_3(X1,X2,X3))|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_338, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_308])).
cnf(c_0_339, plain, (r2_hidden(esk69_3(X1,X2,X3),X3)|v1_xboole_0(X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_309, c_0_158])])).
cnf(c_0_340, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_285])])).
cnf(c_0_341, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_311])).
cnf(c_0_342, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_312])])).
cnf(c_0_343, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_313])])).
cnf(c_0_344, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_314])])).
cnf(c_0_345, plain, (v1_xboole_0(X1)|v1_xboole_0(X2)|m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m2_subset_1(X3,X1,X2)), inference(split_conjunct,[status(thm)],[c_0_315])).
cnf(c_0_346, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(split_conjunct,[status(thm)],[c_0_316])).
cnf(c_0_347, plain, (m2_subset_1(X1,X3,X2)|v1_xboole_0(X3)|v1_xboole_0(X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_317])).
cnf(c_0_348, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,X4,X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_hidden(X4,esk70_3(X1,X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_318])])).
cnf(c_0_349, negated_conjecture, (r2_hidden(esk71_3(esk1_0,esk2_0,esk3_0),esk70_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_319, c_0_320]), c_0_259]), c_0_260]), c_0_244]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])]), c_0_261])])).
cnf(c_0_350, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|r1_orders_2(X1,esk71_3(X1,X2,X3),X2)|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_321, c_0_158])])).
cnf(c_0_351, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r2_hidden(esk72_3(X1,X2,X3),X3)|m1_subset_1(esk71_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)), inference(evalgc,[status(thm)],[c_0_322])).
cnf(c_0_352, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|r2_hidden(esk69_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_323, c_0_158])])).
cnf(c_0_353, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|m1_subset_1(esk71_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)), inference(evalgc,[status(thm)],[c_0_324])).
cnf(c_0_354, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,X4,X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_hidden(X4,esk70_3(X1,X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_318])])).
cnf(c_0_355, negated_conjecture, (r2_hidden(esk71_3(esk1_0,esk2_0,X1),esk70_3(esk1_0,esk2_0,X1))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_325, c_0_326]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])])).
cnf(c_0_356, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|r1_orders_2(X1,esk71_3(X1,X2,X3),X2)|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_327, c_0_158])])).
cnf(c_0_357, plain, (k2_xboole_0(X1,X2)=X2|~v1_xboole_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_328, c_0_329])])).
cnf(c_0_358, plain, (k2_yellow_0(X1,k1_tarski(X2))=X2|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_330, c_0_204])])).
cnf(c_0_359, plain, (r2_hidden(esk69_3(X1,X2,X3),X3)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_331, c_0_158])])).
cnf(c_0_360, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_237])])).
cnf(c_0_361, plain, (r1_tarski(X1,X1)), inference(evalgc,[status(thm)],[c_0_332])).
cnf(c_0_362, plain, (v1_xboole_0(X1)|m1_subset_1(esk72_3(X2,X3,X1),u1_struct_0(X2))|~r2_hidden(X4,k2_xboole_0(esk70_3(X2,X3,X1),k1_tarski(esk69_3(X2,X3,X1))))|~r1_orders_2(X2,k2_yellow_0(X2,X1),X3)|~r1_orders_2(X2,X4,X3)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_333, c_0_158])])).
cnf(c_0_363, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2)), inference(evalgc,[status(thm)],[c_0_334])).
cnf(c_0_364, plain, (v1_xboole_0(X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))), inference(evalgc,[status(thm)],[c_0_335])).
cnf(c_0_365, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(split_conjunct,[status(thm)],[c_0_336])).
cnf(c_0_366, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|r2_hidden(esk71_3(X1,X2,X3),esk70_3(X1,X2,X3))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)), inference(cn,[status(thm)],[c_0_337])).
cnf(c_0_367, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_338, c_0_244])])).
cnf(c_0_368, negated_conjecture, (r2_hidden(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_339, c_0_259]), c_0_260]), c_0_244]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])]), c_0_261])])).
cnf(c_0_369, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_340])])).
cnf(c_0_370, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_313])])).
cnf(c_0_371, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_310, c_0_314])])).
cnf(c_0_372, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_341])).
cnf(c_0_373, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk3_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_342, c_0_343]), c_0_344])).
cnf(c_0_374, plain, (m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|v1_xboole_0(X3)|r1_orders_2(X1,esk71_3(X1,X2,X3),X2)|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_375, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_345])).
cnf(c_0_376, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[c_0_346])).
cnf(c_0_377, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(evalgc,[status(thm)],[c_0_347])).
cnf(c_0_378, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~r2_hidden(X1,esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_348, c_0_259]), c_0_260]), c_0_244]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])]), c_0_261])])).
cnf(c_0_379, negated_conjecture, (v2_struct_0(esk1_0)|r2_hidden(esk71_3(esk1_0,esk2_0,esk3_0),esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[c_0_349, c_0_266])).
cnf(c_0_380, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|r1_orders_2(esk1_0,esk71_3(esk1_0,esk2_0,esk3_0),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_350, c_0_320]), c_0_259]), c_0_260]), c_0_244]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])]), c_0_261])])).
cnf(c_0_381, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|m1_subset_1(esk71_3(X1,X2,X3),u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_351, c_0_158])])).
cnf(c_0_382, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|r2_hidden(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_352, c_0_259]), c_0_260]), c_0_244]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])]), c_0_261])])).
cnf(c_0_383, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|m1_subset_1(esk71_3(X1,X2,X3),u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_353, c_0_158])])).
cnf(c_0_384, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(X1)|v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~r1_orders_2(esk1_0,esk71_3(esk1_0,esk2_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk71_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_354, c_0_355]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])])).
cnf(c_0_385, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|r1_orders_2(esk1_0,esk71_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_356, c_0_326]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])])).
cnf(c_0_386, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3))),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_xboole_0(esk70_3(X1,X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_357])])).
cnf(c_0_387, plain, (k2_yellow_0(X1,k1_tarski(X2))=X2|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_358, c_0_86])).
cnf(c_0_388, negated_conjecture, (r2_hidden(esk69_3(esk1_0,X1,X2),X2)|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,X2),X1)|v1_xboole_0(X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X2),X1)|~v1_finset_1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_359, c_0_360]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])])).
cnf(c_0_389, plain, (m1_subset_1(X1,k1_zfmisc_1(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_283, c_0_361])])).
cnf(c_0_390, plain, (m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|v1_xboole_0(X3)|m1_subset_1(esk71_3(X1,X2,X3),u1_struct_0(X1))|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|~v1_xboole_0(k1_xboole_0)|~v1_finset_1(X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))), inference(split_conjunct,[status(thm)],[c_0_121])).
cnf(c_0_391, plain, (v1_xboole_0(X1)|m1_subset_1(esk72_3(X2,X3,X1),u1_struct_0(X2))|v1_xboole_0(k2_xboole_0(esk70_3(X2,X3,X1),k1_tarski(esk69_3(X2,X3,X1))))|~r1_orders_2(X2,k2_yellow_0(X2,X1),X3)|~r1_orders_2(X2,X4,X3)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X4,k2_xboole_0(esk70_3(X2,X3,X1),k1_tarski(esk69_3(X2,X3,X1))))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_362, c_0_363])])).
cnf(c_0_392, plain, (v1_xboole_0(X1)|m1_subset_1(esk72_3(X2,X3,X1),u1_struct_0(X2))|~r1_orders_2(X2,k2_yellow_0(X2,X1),X3)|~v1_finset_1(X1)|~v1_xboole_0(k2_xboole_0(esk70_3(X2,X3,X1),k1_tarski(esk69_3(X2,X3,X1))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_364, c_0_158])])).
cnf(c_0_393, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[c_0_365])).
cnf(c_0_394, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|r2_hidden(esk71_3(X1,X2,X3),esk70_3(X1,X2,X3))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)), inference(evalgc,[status(thm)],[c_0_366])).
cnf(c_0_395, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_267]), c_0_259]), c_0_260]), c_0_244])]), c_0_261])])).
cnf(c_0_396, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_368])])).
cnf(c_0_397, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_369, c_0_370]), c_0_371])).
cnf(c_0_398, negated_conjecture, (v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_372, c_0_373]), c_0_260])])])).
cnf(c_0_399, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r1_orders_2(X1,esk71_3(X1,X2,X3),X2)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)), inference(cn,[status(thm)],[c_0_374])).
cnf(c_0_400, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~m2_subset_1(X2,X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X3))), inference(csr,[status(thm)],[c_0_375, c_0_376])).
cnf(c_0_401, plain, (m2_subset_1(X1,X2,X3)|v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(X2))|~m1_subset_1(X1,X3)), inference(csr,[status(thm)],[c_0_377, c_0_376])).
cnf(c_0_402, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk71_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk71_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_378, c_0_379])])).
cnf(c_0_403, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk71_3(esk1_0,esk2_0,esk3_0),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[c_0_380, c_0_266])).
cnf(c_0_404, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk71_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_381, c_0_320]), c_0_259]), c_0_260]), c_0_244]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])]), c_0_261])])).
cnf(c_0_405, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_367, c_0_382])])).
cnf(c_0_406, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|m1_subset_1(esk71_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_383, c_0_326]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])])])).
cnf(c_0_407, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~v1_finset_1(X1)|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk71_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[c_0_384, c_0_385])).
cnf(c_0_408, plain, (r1_orders_2(X1,esk69_3(X1,X2,X3),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~v1_xboole_0(esk70_3(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_386, c_0_387])])).
cnf(c_0_409, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(esk3_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_338, c_0_360])])).
cnf(c_0_410, negated_conjecture, (r2_hidden(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_388, c_0_259]), c_0_260]), c_0_101]), c_0_389])]), c_0_261])])).
cnf(c_0_411, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|m1_subset_1(esk71_3(X1,X2,X3),u1_struct_0(X1))|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)), inference(cn,[status(thm)],[c_0_390])).
cnf(c_0_412, plain, (v1_xboole_0(X1)|m1_subset_1(esk72_3(X2,X3,X1),u1_struct_0(X2))|~r1_orders_2(X2,k2_yellow_0(X2,X1),X3)|~r1_orders_2(X2,X4,X3)|~v1_finset_1(X1)|~m1_subset_1(X4,k2_xboole_0(esk70_3(X2,X3,X1),k1_tarski(esk69_3(X2,X3,X1))))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)), inference(csr,[status(thm)],[c_0_391, c_0_392])).
cnf(c_0_413, plain, (m1_subset_1(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_318])])).
cnf(c_0_414, plain, (r2_hidden(esk71_3(X1,X2,X3),esk70_3(X1,X2,X3))|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_394, c_0_158])])).
cnf(c_0_415, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_395, c_0_396]), c_0_397]), c_0_398])).
cnf(c_0_416, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r1_orders_2(X1,esk71_3(X1,X2,X3),X2)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)), inference(evalgc,[status(thm)],[c_0_399])).
cnf(c_0_417, negated_conjecture, (v2_struct_0(esk1_0)|v1_xboole_0(esk3_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_363])])).
cnf(c_0_418, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3))),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_xboole_0(esk70_3(X1,X2,X3))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_357])])).
cnf(c_0_419, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m1_subset_1(X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_400, c_0_401])])).
cnf(c_0_420, negated_conjecture, (v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk71_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_402, c_0_403]), c_0_266])).
cnf(c_0_421, negated_conjecture, (v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk71_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_404, c_0_405]), c_0_397]), c_0_398]), c_0_266])).
cnf(c_0_422, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_405])])).
cnf(c_0_423, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~v1_finset_1(X1)|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))), inference(csr,[status(thm)],[c_0_406, c_0_407])).
cnf(c_0_424, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~v1_xboole_0(esk70_3(X1,X2,X3))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(csr,[status(thm)],[c_0_408, c_0_296])).
cnf(c_0_425, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_409, c_0_410]), c_0_389])])])).
cnf(c_0_426, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|m1_subset_1(esk71_3(X1,X2,X3),u1_struct_0(X1))|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)), inference(evalgc,[status(thm)],[c_0_411])).
cnf(c_0_427, plain, (v1_xboole_0(X1)|m1_subset_1(esk72_3(X2,X3,X1),u1_struct_0(X2))|~r1_orders_2(X2,k2_yellow_0(X2,X1),X3)|~r1_orders_2(X2,X4,X3)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_hidden(X4,esk70_3(X2,X3,X1))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_412, c_0_413])])).
cnf(c_0_428, negated_conjecture, (r2_hidden(esk71_3(esk1_0,esk2_0,esk3_0),esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_414, c_0_415]), c_0_259]), c_0_260]), c_0_244]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])]), c_0_261])])).
cnf(c_0_429, plain, (r1_orders_2(X1,esk71_3(X1,X2,X3),X2)|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_416, c_0_158])])).
cnf(c_0_430, negated_conjecture, (v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,esk3_0)), inference(sr,[status(thm)],[c_0_417, c_0_261])).
cnf(c_0_431, negated_conjecture, (r2_hidden(esk72_3(esk1_0,X1,esk3_0),esk3_0)|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk69_3(esk1_0,X1,esk3_0))),X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),X1)|~v1_xboole_0(esk70_3(esk1_0,X1,esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_418, c_0_244]), c_0_260]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])]), c_0_261])])).
cnf(c_0_432, negated_conjecture, (v1_xboole_0(esk3_0)|m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_419, c_0_244])])).
cnf(c_0_433, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_382])])).
cnf(c_0_434, negated_conjecture, (v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_420, c_0_421])])).
cnf(c_0_435, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)), inference(csr,[status(thm)],[c_0_422, c_0_396])).
cnf(c_0_436, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_423, c_0_371]), c_0_259]), c_0_260]), c_0_244])]), c_0_261])])).
cnf(c_0_437, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_424, c_0_425]), c_0_259]), c_0_260]), c_0_244]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])]), c_0_261])])).
cnf(c_0_438, plain, (v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|m1_subset_1(esk71_3(X1,X2,X3),u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_426, c_0_158])])).
cnf(c_0_439, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk71_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk71_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_427, c_0_428]), c_0_259]), c_0_260]), c_0_244]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])]), c_0_261])])).
cnf(c_0_440, negated_conjecture, (r1_orders_2(esk1_0,esk71_3(esk1_0,esk2_0,esk3_0),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_429, c_0_415]), c_0_259]), c_0_260]), c_0_244]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])]), c_0_261])])).
cnf(c_0_441, negated_conjecture, (v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(esk69_3(esk1_0,esk2_0,esk3_0))),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(esk69_3(esk1_0,esk2_0,esk3_0))),esk3_0)|~v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_430, c_0_431]), c_0_259]), c_0_101])])])).
cnf(c_0_442, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,esk3_0)), inference(sr,[status(thm)],[c_0_432, c_0_261])).
cnf(c_0_443, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_433])])).
cnf(c_0_444, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_368])])).
cnf(c_0_445, negated_conjecture, (v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_434, c_0_435]), c_0_397]), c_0_398])).
cnf(c_0_446, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_436, c_0_425]), c_0_398]), c_0_437])).
cnf(c_0_447, negated_conjecture, (v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|m1_subset_1(esk71_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_438, c_0_415]), c_0_259]), c_0_260]), c_0_244]), c_0_101]), c_0_102]), c_0_103]), c_0_104]), c_0_105]), c_0_106])]), c_0_261])])).
cnf(c_0_448, negated_conjecture, (v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk71_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[c_0_439, c_0_440])).
cnf(c_0_449, negated_conjecture, (v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(esk69_3(esk1_0,esk2_0,esk3_0))),esk3_0)), inference(csr,[status(thm)],[c_0_441, c_0_442])).
cnf(c_0_450, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)), inference(csr,[status(thm)],[c_0_443, c_0_444])).
cnf(c_0_451, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_393, c_0_410])])).
cnf(c_0_452, negated_conjecture, (v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_445, c_0_446])])).
cnf(c_0_453, negated_conjecture, (v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[c_0_447, c_0_448])).
cnf(c_0_454, negated_conjecture, (v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_449, c_0_147]), c_0_102]), c_0_104]), c_0_106])]), c_0_153])])).
cnf(c_0_455, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_450, c_0_451])])).
cnf(c_0_456, negated_conjecture, (v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))), inference(csr,[status(thm)],[c_0_452, c_0_453])).
cnf(c_0_457, negated_conjecture, (v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_454, c_0_455]), c_0_405]), c_0_456])).
cnf(c_0_458, negated_conjecture, (v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_457])])).
cnf(c_0_459, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_409, c_0_457]), c_0_389])])])).
cnf(c_0_460, negated_conjecture, (v2_struct_0(esk1_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))), inference(csr,[status(thm)],[c_0_458, c_0_446])).
cnf(c_0_461, negated_conjecture, (v2_struct_0(esk1_0)), inference(csr,[status(thm)],[c_0_459, c_0_460])).
cnf(c_0_462, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_64, c_0_461]), c_0_102]), c_0_103])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 463
# Proof object clause steps            : 389
# Proof object formula steps           : 74
# Proof object conjectures             : 134
# Proof object clause conjectures      : 131
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 78
# Proof object initial formulas used   : 37
# Proof object generating inferences   : 140
# Proof object simplifying inferences  : 444
# Training examples: 218 positive, 2148 negative
# Training: Positive examples begin
cnf(c_0_106, negated_conjecture, (v3_orders_2(esk1_0))).# trainpos
cnf(c_0_105, negated_conjecture, (v4_orders_2(esk1_0))).# trainpos
cnf(c_0_104, negated_conjecture, (v5_orders_2(esk1_0))).# trainpos
cnf(c_0_224, negated_conjecture, (v1_finset_1(esk3_0)|~v5_waybel_6(esk2_0,esk1_0))).# trainpos
cnf(c_0_103, negated_conjecture, (v2_lattice3(esk1_0))).# trainpos
cnf(c_0_102, negated_conjecture, (l1_orders_2(esk1_0))).# trainpos
cnf(c_0_101, negated_conjecture, (m1_subset_1(esk2_0,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_393, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_158, plain, (v1_xboole_0(k1_xboole_0))).# trainpos
cnf(c_0_225, negated_conjecture, (~v1_xboole_0(esk3_0)|~v5_waybel_6(esk2_0,esk1_0))).# trainpos
cnf(c_0_212, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v5_waybel_6(esk2_0,esk1_0))).# trainpos
cnf(c_0_141, plain, (~v1_xboole_0(k1_tarski(X1)))).# trainpos
cnf(c_0_223, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk2_0)|~v5_waybel_6(esk2_0,esk1_0))).# trainpos
cnf(c_0_82, negated_conjecture, (v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|r2_hidden(esk4_1(X1),X1)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0))).# trainpos
cnf(c_0_140, plain, (v1_finset_1(k1_tarski(X1)))).# trainpos
cnf(c_0_84, plain, (~v1_xboole_0(k2_tarski(X1,X2)))).# trainpos
cnf(c_0_361, plain, (r1_tarski(X1,X1))).# trainpos
cnf(c_0_175, plain, (k2_xboole_0(X1,X1)=X1)).# trainpos
cnf(c_0_83, plain, (v1_finset_1(k2_tarski(X1,X2)))).# trainpos
cnf(c_0_68, plain, (l1_struct_0(X1)|~l1_orders_2(X1))).# trainpos
cnf(c_0_112, negated_conjecture, (v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|r1_orders_2(esk1_0,esk4_1(X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0))).# trainpos
cnf(c_0_328, plain, (~v1_xboole_0(X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_64, plain, (~v2_lattice3(X1)|~l1_orders_2(X1)|~v2_struct_0(X1))).# trainpos
cnf(c_0_363, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).# trainpos
cnf(c_0_246, negated_conjecture, (~r2_hidden(X1,esk3_0)|~v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,X1,esk2_0))).# trainpos
cnf(c_0_160, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2))).# trainpos
cnf(c_0_67, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).# trainpos
cnf(c_0_283, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_237, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_376, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_372, plain, (v1_finset_1(X2)|~v1_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_123, plain, (m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1))).# trainpos
cnf(c_0_117, plain, (k6_domain_1(X1,X2)=k1_tarski(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_71, plain, (X1=X4|X1=X3|X2!=k2_tarski(X3,X4)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_236, plain, (r1_tarski(X1,X3)|~r1_tarski(X2,X3)|~r1_tarski(X1,X2))).# trainpos
cnf(c_0_159, plain, (r2_hidden(X1,X3)|X3!=k2_xboole_0(X4,X2)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_288, plain, (r2_hidden(X1,X3)|X3!=k2_xboole_0(X2,X4)|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_338, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_194, plain, (v1_xboole_0(X1)|m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_66, plain, (k7_domain_1(X1,X2,X3)=k2_tarski(X2,X3)|v1_xboole_0(X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_275, plain, (X3=k2_xboole_0(X1,X2)|~r2_hidden(esk7_3(X1,X2,X3),X3)|~r2_hidden(esk7_3(X1,X2,X3),X2))).# trainpos
cnf(c_0_375, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m1_subset_1(X3,X1)|~m2_subset_1(X3,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_136, plain, (v1_xboole_0(X1)|m1_subset_1(k7_domain_1(X1,X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_377, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).# trainpos
cnf(c_0_62, plain, (v2_struct_0(X1)|v5_waybel_6(X2,X1)|m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_63, plain, (v2_struct_0(X1)|v5_waybel_6(X2,X1)|m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_177, plain, (v2_struct_0(X1)|v5_waybel_6(X2,X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,esk8_2(X1,X2),X2))).# trainpos
cnf(c_0_165, plain, (v2_struct_0(X1)|v5_waybel_6(X2,X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~r1_orders_2(X1,esk9_2(X1,X2),X2))).# trainpos
cnf(c_0_94, plain, (k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2|v2_struct_0(X1)|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_128, plain, (v1_xboole_0(X2)|v2_struct_0(X1)|r2_yellow_0(X1,X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_252, plain, (X3=k2_xboole_0(X1,X2)|r2_hidden(esk7_3(X1,X2,X3),X3)|r2_hidden(esk7_3(X1,X2,X3),X2)|r2_hidden(esk7_3(X1,X2,X3),X1))).# trainpos
cnf(c_0_110, plain, (k12_lattice3(X1,X2,X3)=k12_lattice3(X1,X3,X2)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_49, plain, (k11_lattice3(X1,X2,X3)=k12_lattice3(X1,X2,X3)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_48, plain, (v2_struct_0(X1)|v5_waybel_6(X2,X1)|r1_orders_2(X1,k11_lattice3(X1,esk8_2(X1,X2),esk9_2(X1,X2)),X2)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_93, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(X2,X3))|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2))).# trainpos
cnf(c_0_65, plain, (k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))=k12_lattice3(X1,X2,X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_126, plain, (v2_struct_0(X2)|r1_orders_2(X2,X4,X1)|r1_orders_2(X2,X3,X1)|~l1_orders_2(X2)|~v5_waybel_6(X1,X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~m1_subset_1(X1,u1_struct_0(X2))|~r1_orders_2(X2,k11_lattice3(X2,X3,X4),X1))).# trainpos
cnf(c_0_323, plain, (v1_xboole_0(X3)|r2_hidden(esk69_3(X1,X2,X3),X3)|r2_hidden(esk72_3(X1,X2,X3),X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2))).# trainpos
cnf(c_0_238, plain, (v1_xboole_0(X3)|r2_hidden(esk72_3(X1,X2,X3),X3)|r1_tarski(esk70_3(X1,X2,X3),X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2))).# trainpos
cnf(c_0_309, plain, (v1_xboole_0(X3)|r2_hidden(esk69_3(X1,X2,X3),X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2))).# trainpos
cnf(c_0_262, plain, (v1_xboole_0(X3)|r1_tarski(esk70_3(X1,X2,X3),X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2))).# trainpos
cnf(c_0_331, plain, (v1_xboole_0(X3)|r2_hidden(esk69_3(X1,X2,X3),X3)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2))).# trainpos
cnf(c_0_364, plain, (v1_xboole_0(X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))))).# trainpos
cnf(c_0_227, plain, (v1_xboole_0(X3)|r1_tarski(esk70_3(X1,X2,X3),X3)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2))).# trainpos
cnf(c_0_157, plain, (v1_xboole_0(X3)|r2_hidden(esk72_3(X1,X2,X3),X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r2_hidden(X4,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))))).# trainpos
cnf(c_0_333, plain, (v1_xboole_0(X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r2_hidden(X4,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))))).# trainpos
cnf(c_0_231, plain, (v1_xboole_0(X3)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~r1_orders_2(X1,X4,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r2_hidden(X4,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))))).# trainpos
cnf(c_0_168, plain, (v1_xboole_0(X3)|r2_hidden(esk72_3(X1,X2,X3),X3)|r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2))).# trainpos
cnf(c_0_190, plain, (v1_xboole_0(X3)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2))).# trainpos
cnf(c_0_351, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r2_hidden(esk72_3(X1,X2,X3),X3)|m1_subset_1(esk71_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2))).# trainpos
cnf(c_0_321, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r2_hidden(esk72_3(X1,X2,X3),X3)|r1_orders_2(X1,esk71_3(X1,X2,X3),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2))).# trainpos
cnf(c_0_426, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|m1_subset_1(esk71_3(X1,X2,X3),u1_struct_0(X1))|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2))).# trainpos
cnf(c_0_416, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r1_orders_2(X1,esk71_3(X1,X2,X3),X2)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2))).# trainpos
cnf(c_0_353, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|m1_subset_1(esk71_3(X1,X2,X3),u1_struct_0(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2))).# trainpos
cnf(c_0_327, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r1_orders_2(X1,esk71_3(X1,X2,X3),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2))).# trainpos
cnf(c_0_289, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r2_hidden(esk72_3(X1,X2,X3),X3)|r2_hidden(esk71_3(X1,X2,X3),esk70_3(X1,X2,X3))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2))).# trainpos
cnf(c_0_394, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|r2_hidden(esk71_3(X1,X2,X3),esk70_3(X1,X2,X3))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2))).# trainpos
cnf(c_0_295, plain, (v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|r2_hidden(esk71_3(X1,X2,X3),esk70_3(X1,X2,X3))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2))).# trainpos
cnf(c_0_171, plain, (r2_hidden(X1,X2)|X2!=k1_tarski(X1))).# trainpos
cnf(c_0_147, plain, (k2_yellow_0(X1,k1_tarski(X2))=X2|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_76, plain, (k12_lattice3(X1,X2,X3)=k2_yellow_0(X1,k2_tarski(X2,X3))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_122, plain, (k12_lattice3(X1,k2_yellow_0(X1,X3),k2_yellow_0(X1,X2))=k2_yellow_0(X1,k2_xboole_0(X2,X3))|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2)|~v4_orders_2(X1))).# trainpos
cnf(c_0_61, plain, (v2_struct_0(X1)|r1_orders_2(X1,k12_lattice3(X1,esk8_2(X1,X2),esk9_2(X1,X2)),X2)|v5_waybel_6(X2,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~v2_lattice3(X1)|~v5_orders_2(X1))).# trainpos
cnf(c_0_137, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X3)|r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,k12_lattice3(X1,X2,X4),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1))).# trainpos
cnf(c_0_107, plain, (k2_yellow_0(X1,k2_xboole_0(X2,k6_domain_1(u1_struct_0(X1),X3)))=k12_lattice3(X1,k2_yellow_0(X1,X2),X3)|v2_struct_0(X1)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X3))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v3_orders_2(X1))).# trainpos
cnf(c_0_419, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(X3))|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_77, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1))).# trainpos
cnf(c_0_318, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_170, plain, (r2_hidden(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X3))).# trainpos
cnf(c_0_389, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).# trainpos
cnf(c_0_204, plain, (v1_xboole_0(X1)|m1_subset_1(k1_tarski(X2),k1_zfmisc_1(X1))|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_152, plain, (v1_xboole_0(X1)|m1_subset_1(k2_tarski(X2,X3),k1_zfmisc_1(X1))|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).# trainpos
cnf(c_0_257, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).# trainpos
cnf(c_0_81, plain, (X1=X2|X1=X3|~r2_hidden(X1,k2_tarski(X2,X3)))).# trainpos
cnf(c_0_276, plain, (k2_xboole_0(X11,X10)=X10|r2_hidden(esk7_3(X11,X10,X10),X11)|r2_hidden(esk7_3(X11,X10,X10),X10))).# trainpos
cnf(c_0_99, negated_conjecture, (esk4_1(k2_tarski(X1,X2))=X2|esk4_1(k2_tarski(X1,X2))=X1|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk2_0)|~m1_subset_1(k2_tarski(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_86, plain, (~l1_orders_2(X1)|~v2_lattice3(X1)|~v1_xboole_0(u1_struct_0(X1)))).# trainpos
cnf(c_0_182, plain, (r2_hidden(X1,k1_tarski(X1)))).# trainpos
cnf(c_0_85, plain, (r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|v5_waybel_6(X2,X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_348, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,X4,X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_hidden(X4,esk70_3(X1,X2,X3)))).# trainpos
cnf(c_0_127, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),X3)=k2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(X3)))|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,k1_tarski(X3))|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_354, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,X4,X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_hidden(X4,esk70_3(X1,X2,X3)))).# trainpos
cnf(c_0_192, plain, (k12_lattice3(X1,X3,X2)=k2_yellow_0(X1,k2_tarski(X2,X3))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_181, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,X4,X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_hidden(X4,k1_tarski(esk69_3(X1,X2,X3))))).# trainpos
cnf(c_0_161, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(X3,k1_tarski(X2)))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,k1_tarski(X2))|~r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1))).# trainpos
cnf(c_0_153, negated_conjecture, (~v1_xboole_0(u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_218, plain, (k12_lattice3(X1,X3,X2)=k2_yellow_0(X1,k2_tarski(X3,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_113, negated_conjecture, (esk4_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)))=esk8_2(esk1_0,esk2_0)|esk4_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)))=esk9_2(esk1_0,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_189, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,esk69_3(X1,X2,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_155, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),X3)=k2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(X3)))|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_184, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(X3,k1_tarski(X2)))|~r2_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_124, negated_conjecture, (r1_orders_2(esk1_0,esk9_2(esk1_0,esk2_0),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|esk4_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)))=esk8_2(esk1_0,esk2_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0))),esk2_0)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_167, plain, (r1_orders_2(X1,X2,X3)|r1_orders_2(X1,k2_yellow_0(X1,X4),X3)|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X4,k1_tarski(X2))),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~r2_yellow_0(X1,X4)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_216, plain, (r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|r1_orders_2(X1,esk69_3(X1,X2,X3),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r2_yellow_0(X1,esk70_3(X1,X2,X3))|~v5_waybel_6(X2,X1)|~m1_subset_1(k1_tarski(esk69_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_188, plain, (r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|r1_orders_2(X1,esk69_3(X1,X2,X3),X2)|r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r2_yellow_0(X1,esk70_3(X1,X2,X3))|~v5_waybel_6(X2,X1)|~m1_subset_1(k1_tarski(esk69_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_135, negated_conjecture, (r1_orders_2(esk1_0,esk8_2(esk1_0,esk2_0),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|r1_orders_2(esk1_0,esk9_2(esk1_0,esk2_0),esk2_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0))),esk2_0)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_151, negated_conjecture, (r1_orders_2(esk1_0,esk9_2(esk1_0,esk2_0),esk2_0)|r1_orders_2(esk1_0,esk8_2(esk1_0,esk2_0),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_214, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(k1_tarski(esk69_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(esk70_3(X1,X2,X3))|~m1_subset_1(esk70_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_232, plain, (r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|r1_orders_2(X1,esk69_3(X1,X2,X3),X2)|v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(k1_tarski(esk69_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(esk70_3(X1,X2,X3))|~m1_subset_1(esk70_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_166, negated_conjecture, (r1_orders_2(esk1_0,esk8_2(esk1_0,esk2_0),esk2_0)|r1_orders_2(esk1_0,esk9_2(esk1_0,esk2_0),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_230, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(esk70_3(X1,X2,X3))|~v1_finset_1(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(esk70_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_250, plain, (r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|r1_orders_2(X1,esk69_3(X1,X2,X3),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(esk70_3(X1,X2,X3))|~v1_finset_1(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(esk70_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_148, plain, (k2_yellow_0(X1,k2_xboole_0(X3,k6_domain_1(u1_struct_0(X1),X2)))=k12_lattice3(X1,X2,k2_yellow_0(X1,X3))|v2_struct_0(X1)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1))).# trainpos
cnf(c_0_191, plain, (k2_yellow_0(X1,k2_xboole_0(k6_domain_1(u1_struct_0(X1),X3),k1_tarski(X2)))=k12_lattice3(X1,X2,X3)|v2_struct_0(X1)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).# trainpos
cnf(c_0_271, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,X4,X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_hidden(X4,k1_tarski(esk69_3(X1,X2,X3))))).# trainpos
cnf(c_0_178, negated_conjecture, (v2_struct_0(esk1_0)|v5_waybel_6(esk2_0,esk1_0)|r1_orders_2(esk1_0,esk8_2(esk1_0,esk2_0),esk2_0)|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_296, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,esk69_3(X1,X2,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_413, plain, (m1_subset_1(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X2))).# trainpos
cnf(c_0_187, negated_conjecture, (v2_struct_0(esk1_0)|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_217, plain, (k2_yellow_0(X1,k2_xboole_0(k1_tarski(X2),k1_tarski(X3)))=k12_lattice3(X1,X3,X2)|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_197, negated_conjecture, (v2_struct_0(esk1_0)|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_253, plain, (k12_lattice3(X1,X2,X2)=k2_yellow_0(X1,k1_tarski(X2))|~r2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_391, plain, (v1_xboole_0(X1)|m1_subset_1(esk72_3(X2,X3,X1),u1_struct_0(X2))|v1_xboole_0(k2_xboole_0(esk70_3(X2,X3,X1),k1_tarski(esk69_3(X2,X3,X1))))|~r1_orders_2(X2,k2_yellow_0(X2,X1),X3)|~r1_orders_2(X2,X4,X3)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~m1_subset_1(X4,k2_xboole_0(esk70_3(X2,X3,X1),k1_tarski(esk69_3(X2,X3,X1)))))).# trainpos
cnf(c_0_213, negated_conjecture, (v2_struct_0(esk1_0)|v5_waybel_6(esk2_0,esk1_0))).# trainpos
cnf(c_0_277, plain, (k12_lattice3(X1,X2,X2)=k2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_427, plain, (v1_xboole_0(X1)|m1_subset_1(esk72_3(X2,X3,X1),u1_struct_0(X2))|~r1_orders_2(X2,k2_yellow_0(X2,X1),X3)|~r1_orders_2(X2,X4,X3)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_hidden(X4,esk70_3(X2,X3,X1)))).# trainpos
cnf(c_0_174, plain, (k2_yellow_0(X1,k2_xboole_0(X2,k6_domain_1(u1_struct_0(X1),X3)))=k12_lattice3(X1,X3,k2_yellow_0(X1,X2))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_239, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk2_0)|v2_struct_0(esk1_0))).# trainpos
cnf(c_0_240, negated_conjecture, (v1_finset_1(esk3_0)|v2_struct_0(esk1_0))).# trainpos
cnf(c_0_228, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0)))|v2_struct_0(esk1_0))).# trainpos
cnf(c_0_241, negated_conjecture, (v2_struct_0(esk1_0)|~v1_xboole_0(esk3_0))).# trainpos
cnf(c_0_266, negated_conjecture, (v2_struct_0(esk1_0)|~r2_hidden(X1,esk3_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_301, plain, (k2_yellow_0(X1,k2_tarski(X2,X2))=k2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1))))).# trainpos
cnf(c_0_267, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,X1),X1)|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,X1)),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~v1_finset_1(X1)|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_297, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,X1)),esk2_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~v1_finset_1(X1)|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_244, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_260, negated_conjecture, (v1_finset_1(esk3_0))).# trainpos
cnf(c_0_261, negated_conjecture, (~v1_xboole_0(esk3_0))).# trainpos
cnf(c_0_259, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk2_0))).# trainpos
cnf(c_0_417, negated_conjecture, (v2_struct_0(esk1_0)|v1_xboole_0(esk3_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,esk3_0))).# trainpos
cnf(c_0_290, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk3_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_finset_1(esk3_0)|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_326, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(X1)|v2_struct_0(esk1_0)|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,X1)),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).# trainpos
cnf(c_0_185, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2)))=k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X2))|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_382, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|r2_hidden(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0))).# trainpos
cnf(c_0_285, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk3_0),esk3_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0))).# trainpos
cnf(c_0_367, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk3_0))).# trainpos
cnf(c_0_432, negated_conjecture, (v1_xboole_0(esk3_0)|m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,esk3_0))).# trainpos
cnf(c_0_313, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk3_0),esk3_0)|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_368, negated_conjecture, (r2_hidden(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_378, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~r2_hidden(X1,esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_284, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~r1_tarski(X1,esk3_0))).# trainpos
cnf(c_0_349, negated_conjecture, (r2_hidden(esk71_3(esk1_0,esk2_0,esk3_0),esk70_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_263, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,esk3_0),esk3_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,esk3_0),X1)|v1_xboole_0(esk3_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),X1)|~v1_finset_1(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_380, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|r1_orders_2(esk1_0,esk71_3(esk1_0,esk2_0,esk3_0),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_310, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_tarski(X1,esk3_0))).# trainpos
cnf(c_0_402, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk71_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk71_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_314, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk3_0),esk3_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0))).# trainpos
cnf(c_0_433, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0))).# trainpos
cnf(c_0_312, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0))).# trainpos
cnf(c_0_405, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0))).# trainpos
cnf(c_0_396, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_444, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_343, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_371, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0))).# trainpos
cnf(c_0_340, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0))).# trainpos
cnf(c_0_422, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_370, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_360, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).# trainpos
cnf(c_0_443, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_344, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0))).# trainpos
cnf(c_0_342, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_369, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_398, negated_conjecture, (v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0))).# trainpos
cnf(c_0_409, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(esk3_0)))).# trainpos
cnf(c_0_388, negated_conjecture, (r2_hidden(esk69_3(esk1_0,X1,X2),X2)|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,X2),X1)|v1_xboole_0(X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X2),X1)|~v1_finset_1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(esk3_0)))).# trainpos
cnf(c_0_395, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_410, negated_conjecture, (r2_hidden(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0))).# trainpos
cnf(c_0_428, negated_conjecture, (r2_hidden(esk71_3(esk1_0,esk2_0,esk3_0),esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0))).# trainpos
cnf(c_0_425, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0))).# trainpos
cnf(c_0_451, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0))).# trainpos
cnf(c_0_440, negated_conjecture, (r1_orders_2(esk1_0,esk71_3(esk1_0,esk2_0,esk3_0),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0))).# trainpos
cnf(c_0_203, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,k1_tarski(X2)))=k2_yellow_0(X1,k1_tarski(X2))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_355, negated_conjecture, (r2_hidden(esk71_3(esk1_0,esk2_0,X1),esk70_3(esk1_0,esk2_0,X1))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_439, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk71_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk71_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_385, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|r1_orders_2(esk1_0,esk71_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_455, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0))).# trainpos
cnf(c_0_384, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(X1)|v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~r1_orders_2(esk1_0,esk71_3(esk1_0,esk2_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk71_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_235, plain, (k2_yellow_0(X1,k2_tarski(X2,k2_yellow_0(X1,k1_tarski(X2))))=k2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,k1_tarski(X2)),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_278, plain, (k2_yellow_0(X1,k2_tarski(X2,X2))=X2|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_404, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk71_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_434, negated_conjecture, (v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_330, plain, (k2_yellow_0(X1,k1_tarski(X2))=X2|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_358, plain, (k2_yellow_0(X1,k1_tarski(X2))=X2|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).# trainpos
cnf(c_0_447, negated_conjecture, (v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|m1_subset_1(esk71_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0))).# trainpos
cnf(c_0_406, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|m1_subset_1(esk71_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_300, plain, (k2_xboole_0(X2,X1)=X1|r2_hidden(esk7_3(X2,X1,X1),X2)|~r2_hidden(esk7_3(X2,X1,X1),X1))).# trainpos
cnf(c_0_357, plain, (k2_xboole_0(X1,X2)=X2|~v1_xboole_0(X1))).# trainpos
cnf(c_0_418, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3))),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_xboole_0(esk70_3(X1,X2,X3)))).# trainpos
cnf(c_0_386, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3))),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_xboole_0(esk70_3(X1,X2,X3)))).# trainpos
cnf(c_0_408, plain, (r1_orders_2(X1,esk69_3(X1,X2,X3),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~v1_xboole_0(esk70_3(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1)))).# trainpos
cnf(c_0_431, negated_conjecture, (r2_hidden(esk72_3(esk1_0,X1,esk3_0),esk3_0)|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk69_3(esk1_0,X1,esk3_0))),X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),X1)|~v1_xboole_0(esk70_3(esk1_0,X1,esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_437, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_441, negated_conjecture, (v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(esk69_3(esk1_0,esk2_0,esk3_0))),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(esk69_3(esk1_0,esk2_0,esk3_0))),esk3_0)|~v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0)))).# trainpos
cnf(c_0_436, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_452, negated_conjecture, (v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_454, negated_conjecture, (v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_458, negated_conjecture, (v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).# trainpos
cnf(c_0_459, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_277, plain, (o_0_0_xboole_0=k1_xboole_0)).#trainneg
cnf(i_0_504, plain, (k5_numbers=k4_ordinal1)).#trainneg
cnf(i_0_349, plain, (v3_orders_2(esk14_0))).#trainneg
cnf(i_0_373, plain, (v3_orders_2(esk23_0))).#trainneg
cnf(i_0_382, plain, (v3_orders_2(esk24_0))).#trainneg
cnf(i_0_434, plain, (v3_orders_2(esk42_0))).#trainneg
cnf(i_0_348, plain, (v4_orders_2(esk14_0))).#trainneg
cnf(i_0_372, plain, (v4_orders_2(esk23_0))).#trainneg
cnf(i_0_381, plain, (v4_orders_2(esk24_0))).#trainneg
cnf(i_0_433, plain, (v4_orders_2(esk42_0))).#trainneg
cnf(i_0_585, plain, (~v1_xboole_0(np__1))).#trainneg
cnf(i_0_337, plain, (~v1_xboole_0(k4_ordinal1))).#trainneg
cnf(i_0_355, plain, (~v1_xboole_0(esk16_0))).#trainneg
cnf(i_0_363, plain, (~v1_xboole_0(esk20_0))).#trainneg
cnf(i_0_404, plain, (~v1_xboole_0(esk33_0))).#trainneg
cnf(i_0_414, plain, (~v1_xboole_0(esk35_0))).#trainneg
cnf(i_0_424, plain, (~v1_xboole_0(esk39_0))).#trainneg
cnf(i_0_442, plain, (~v1_xboole_0(esk44_0))).#trainneg
cnf(i_0_347, plain, (v5_orders_2(esk14_0))).#trainneg
cnf(i_0_371, plain, (v5_orders_2(esk23_0))).#trainneg
cnf(i_0_380, plain, (v5_orders_2(esk24_0))).#trainneg
cnf(i_0_432, plain, (v5_orders_2(esk42_0))).#trainneg
cnf(i_0_378, plain, (v2_lattice3(esk24_0))).#trainneg
cnf(i_0_310, plain, (l1_orders_2(esk10_0))).#trainneg
cnf(i_0_352, plain, (l1_orders_2(esk14_0))).#trainneg
cnf(i_0_376, plain, (l1_orders_2(esk23_0))).#trainneg
cnf(i_0_385, plain, (l1_orders_2(esk24_0))).#trainneg
cnf(i_0_437, plain, (l1_orders_2(esk42_0))).#trainneg
cnf(i_0_461, plain, (~v1_xboole_0(esk51_0))).#trainneg
cnf(i_0_478, plain, (~v1_xboole_0(esk59_0))).#trainneg
cnf(i_0_495, plain, (~v1_xboole_0(esk66_0))).#trainneg
cnf(i_0_498, plain, (~v1_xboole_0(esk67_0))).#trainneg
cnf(i_0_502, plain, (~v1_xboole_0(esk68_0))).#trainneg
cnf(i_0_339, plain, (~v1_finset_1(k4_ordinal1))).#trainneg
cnf(i_0_438, plain, (~v1_finset_1(esk43_0))).#trainneg
cnf(i_0_351, plain, (~v2_struct_0(esk14_0))).#trainneg
cnf(i_0_386, plain, (v1_xboole_0(esk25_0))).#trainneg
cnf(i_0_470, plain, (v1_xboole_0(esk55_0))).#trainneg
cnf(i_0_472, plain, (v1_xboole_0(esk56_0))).#trainneg
cnf(i_0_354, plain, (v1_finset_1(esk16_0))).#trainneg
cnf(i_0_397, plain, (v1_finset_1(esk31_0))).#trainneg
cnf(i_0_439, plain, (v1_finset_1(esk44_0))).#trainneg
cnf(i_0_497, plain, (v1_finset_1(esk67_0))).#trainneg
cnf(i_0_311, plain, (l1_struct_0(esk11_0))).#trainneg
cnf(i_0_379, plain, (v1_lattice3(esk24_0))).#trainneg
cnf(i_0_48, plain, (v3_ordinal1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_49, plain, (v1_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_370, plain, (v3_lattice3(esk23_0))).#trainneg
cnf(i_0_375, plain, (~v2_struct_0(esk23_0))).#trainneg
cnf(i_0_384, plain, (~v2_struct_0(esk24_0))).#trainneg
cnf(i_0_436, plain, (~v2_struct_0(esk42_0))).#trainneg
cnf(i_0_474, plain, (~v3_funct_1(esk58_0))).#trainneg
cnf(i_0_431, plain, (v3_lattice3(esk42_0))).#trainneg
cnf(i_0_332, plain, (v1_card_1(k4_ordinal1))).#trainneg
cnf(i_0_353, plain, (v1_card_1(esk15_0))).#trainneg
cnf(i_0_396, plain, (v1_card_1(esk31_0))).#trainneg
cnf(i_0_336, plain, (v3_ordinal1(k4_ordinal1))).#trainneg
cnf(i_0_361, plain, (v3_ordinal1(esk19_0))).#trainneg
cnf(i_0_398, plain, (v3_ordinal1(esk31_0))).#trainneg
cnf(i_0_411, plain, (v3_ordinal1(esk35_0))).#trainneg
cnf(i_0_357, plain, (v1_funct_1(esk17_0))).#trainneg
cnf(i_0_359, plain, (v1_funct_1(esk18_0))).#trainneg
cnf(i_0_406, plain, (v1_funct_1(esk33_0))).#trainneg
cnf(i_0_409, plain, (v1_funct_1(esk34_0))).#trainneg
cnf(i_0_440, plain, (v1_funct_1(esk44_0))).#trainneg
cnf(i_0_443, plain, (v1_funct_1(esk45_0))).#trainneg
cnf(i_0_54, plain, (v1_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_458, plain, (v1_funct_1(esk51_0))).#trainneg
cnf(i_0_321, plain, (~v1_xboole_0(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_475, plain, (v1_funct_1(esk58_0))).#trainneg
cnf(i_0_500, plain, (v1_funct_1(esk68_0))).#trainneg
cnf(i_0_350, plain, (v2_orders_2(esk14_0))).#trainneg
cnf(i_0_374, plain, (v2_orders_2(esk23_0))).#trainneg
cnf(i_0_383, plain, (v2_orders_2(esk24_0))).#trainneg
cnf(i_0_435, plain, (v2_orders_2(esk42_0))).#trainneg
cnf(i_0_400, plain, (v1_ordinal1(esk31_0))).#trainneg
cnf(i_0_413, plain, (v1_ordinal1(esk35_0))).#trainneg
cnf(i_0_399, plain, (v2_ordinal1(esk31_0))).#trainneg
cnf(i_0_412, plain, (v2_ordinal1(esk35_0))).#trainneg
cnf(i_0_358, plain, (v1_relat_1(esk17_0))).#trainneg
cnf(i_0_360, plain, (v1_relat_1(esk18_0))).#trainneg
cnf(i_0_362, plain, (v1_relat_1(esk20_0))).#trainneg
cnf(i_0_407, plain, (v1_relat_1(esk33_0))).#trainneg
cnf(i_0_410, plain, (v1_relat_1(esk34_0))).#trainneg
cnf(i_0_416, plain, (v1_relat_1(esk36_0))).#trainneg
cnf(i_0_441, plain, (v1_relat_1(esk44_0))).#trainneg
cnf(i_0_53, plain, (v2_ordinal1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_445, plain, (v1_relat_1(esk45_0))).#trainneg
cnf(i_0_460, plain, (v1_relat_1(esk51_0))).#trainneg
cnf(i_0_476, plain, (v1_relat_1(esk58_0))).#trainneg
cnf(i_0_501, plain, (v1_relat_1(esk68_0))).#trainneg
cnf(i_0_346, plain, (v16_waybel_0(esk14_0))).#trainneg
cnf(i_0_369, plain, (v16_waybel_0(esk23_0))).#trainneg
cnf(i_0_387, plain, (v1_xreal_0(esk26_0))).#trainneg
cnf(i_0_425, plain, (v1_xreal_0(esk40_0))).#trainneg
cnf(i_0_450, plain, (v1_xreal_0(esk48_0))).#trainneg
cnf(i_0_467, plain, (v1_xreal_0(esk55_0))).#trainneg
cnf(i_0_377, plain, (v1_waybel_5(esk24_0))).#trainneg
cnf(i_0_356, plain, (v1_funcop_1(esk17_0))).#trainneg
cnf(i_0_408, plain, (v2_funct_1(esk34_0))).#trainneg
cnf(i_0_462, plain, (v7_ordinal1(esk52_0))).#trainneg
cnf(i_0_477, plain, (v7_ordinal1(esk59_0))).#trainneg
cnf(i_0_388, plain, (v1_xxreal_0(esk27_0))).#trainneg
cnf(i_0_427, plain, (v1_xxreal_0(esk40_0))).#trainneg
cnf(i_0_430, plain, (v1_xxreal_0(esk41_0))).#trainneg
cnf(i_0_452, plain, (v1_xxreal_0(esk48_0))).#trainneg
cnf(i_0_455, plain, (v1_xxreal_0(esk49_0))).#trainneg
cnf(i_0_103, plain, (v1_xreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_468, plain, (v1_xxreal_0(esk55_0))).#trainneg
cnf(i_0_471, plain, (v1_xxreal_0(esk56_0))).#trainneg
cnf(i_0_428, plain, (v1_xcmplx_0(esk40_0))).#trainneg
cnf(i_0_453, plain, (v1_xcmplx_0(esk48_0))).#trainneg
cnf(i_0_469, plain, (v1_xcmplx_0(esk55_0))).#trainneg
cnf(i_0_589, plain, (v2_xxreal_0(np__1))).#trainneg
cnf(i_0_426, plain, (v2_xxreal_0(esk40_0))).#trainneg
cnf(i_0_429, plain, (v2_xxreal_0(esk41_0))).#trainneg
cnf(i_0_451, plain, (v3_xxreal_0(esk48_0))).#trainneg
cnf(i_0_454, plain, (v3_xxreal_0(esk49_0))).#trainneg
cnf(i_0_405, plain, (v3_funct_1(esk33_0))).#trainneg
cnf(i_0_415, plain, (v2_relat_1(esk36_0))).#trainneg
cnf(i_0_444, plain, (v2_relat_1(esk45_0))).#trainneg
cnf(i_0_459, plain, (v2_relat_1(esk51_0))).#trainneg
cnf(i_0_496, plain, (v5_finset_1(esk67_0))).#trainneg
cnf(i_0_494, plain, (v4_funct_1(esk66_0))).#trainneg
cnf(i_0_499, plain, (v2_finset_1(esk68_0))).#trainneg
cnf(i_0_334, plain, (v2_card_1(k4_ordinal1))).#trainneg
cnf(i_0_417, plain, (v1_xboole_0(esk37_1(X1)))).#trainneg
cnf(i_0_132, plain, (v1_xcmplx_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_3, negated_conjecture, (v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|m1_subset_1(esk4_1(X1),u1_struct_0(esk1_0))|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0))).#trainneg
cnf(i_0_609, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_448, plain, (~v1_subset_1(esk47_1(X1),X1))).#trainneg
cnf(i_0_139, plain, (v1_finset_1(X1)|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_586, plain, (m1_subset_1(np__1,k1_numbers))).#trainneg
cnf(i_0_587, plain, (m1_subset_1(np__1,k5_numbers))).#trainneg
cnf(i_0_591, plain, (m1_subset_1(k1_xboole_0,k4_ordinal1))).#trainneg
cnf(i_0_301, plain, (m1_subset_1(k5_numbers,k1_zfmisc_1(k1_numbers)))).#trainneg
cnf(i_0_588, plain, (m2_subset_1(np__1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_590, plain, (k2_xboole_0(X1,k1_xboole_0)=X1)).#trainneg
cnf(i_0_314, plain, (v3_card_1(k1_tarski(X1),np__1))).#trainneg
cnf(i_0_312, plain, (m1_subset_1(esk12_1(X1),X1))).#trainneg
cnf(i_0_168, plain, (v1_finset_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_259, plain, (v6_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_74, plain, (v1_card_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_115, plain, (v3_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_110, plain, (v1_card_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_201, plain, (v3_ordinal1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_163, plain, (v1_xxreal_0(X1)|~v1_xreal_0(X1))).#trainneg
cnf(i_0_255, plain, (~v1_xboole_0(X1)|~v3_card_1(X1,np__1))).#trainneg
cnf(i_0_51, plain, (~l1_orders_2(X1)|~v2_struct_0(X1)|~v1_lattice3(X1))).#trainneg
cnf(i_0_135, plain, (~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v2_xxreal_0(X1))).#trainneg
cnf(i_0_50, plain, (v1_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_55, plain, (v1_relat_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_225, plain, (v7_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_104, plain, (v1_xxreal_0(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_143, plain, (v5_ordinal1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_197, plain, (v5_finset_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_224, plain, (v4_funct_1(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_418, plain, (m1_subset_1(esk37_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_449, plain, (m1_subset_1(esk47_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_264, plain, (k2_tarski(X1,X2)=k2_tarski(X2,X1))).#trainneg
cnf(i_0_265, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X2,X1))).#trainneg
cnf(i_0_327, plain, (v5_finset_1(k1_zfmisc_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_138, plain, (v1_finset_1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_401, plain, (v1_xboole_0(X1)|v1_finset_1(esk32_1(X1)))).#trainneg
cnf(i_0_342, plain, (v1_finset_1(X1)|~v1_finset_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_364, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk21_1(X1)))).#trainneg
cnf(i_0_402, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk32_1(X1)))).#trainneg
cnf(i_0_194, plain, (~v1_xboole_0(X1)|~v1_xxreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_133, plain, (~v1_xxreal_0(X1)|~v2_xxreal_0(X1)|~v3_xxreal_0(X1))).#trainneg
cnf(i_0_25, plain, (~l1_struct_0(X1)|~v2_struct_0(X1)|~v13_struct_0(X1,np__1))).#trainneg
cnf(i_0_479, plain, (v1_xboole_0(X1)|v1_zfmisc_1(esk60_1(X1)))).#trainneg
cnf(i_0_491, plain, (v1_zfmisc_1(X1)|v1_finset_1(esk65_1(X1)))).#trainneg
cnf(i_0_317, plain, (v1_finset_1(k1_zfmisc_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_395, plain, (l1_struct_0(esk30_1(X1))|~v1_card_1(X1))).#trainneg
cnf(i_0_485, plain, (v1_funct_1(esk62_1(X1))|~v1_card_1(X1))).#trainneg
cnf(i_0_486, plain, (v1_relat_1(esk62_1(X1))|~v1_card_1(X1))).#trainneg
cnf(i_0_326, plain, (v5_finset_1(k1_tarski(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_237, plain, (v3_card_1(X1,k1_xboole_0)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_480, plain, (v1_xboole_0(X1)|~v1_xboole_0(esk60_1(X1)))).#trainneg
cnf(i_0_456, plain, (v1_finset_1(X1)|~v1_finset_1(esk50_1(X1)))).#trainneg
cnf(i_0_487, plain, (v1_zfmisc_1(X1)|~v1_zfmisc_1(esk63_1(X1)))).#trainneg
cnf(i_0_47, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_611, plain, (X1=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_222, plain, (v1_xboole_0(X1)|~v3_card_1(X1,k1_xboole_0))).#trainneg
cnf(i_0_68, plain, (v1_xreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_241, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_15, plain, (~r2_hidden(X2,X1)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_608, plain, (~v1_xboole_0(X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_148, plain, (~v1_xboole_0(X1)|~v1_subset_1(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_492, plain, (v1_zfmisc_1(X1)|~v1_zfmisc_1(esk65_1(X1)))).#trainneg
cnf(i_0_254, plain, (v1_zfmisc_1(X1)|~v3_card_1(X1,np__1))).#trainneg
cnf(i_0_489, plain, (v1_xboole_0(X1)|v3_card_1(esk64_1(X1),np__1))).#trainneg
cnf(i_0_465, plain, (v1_xboole_0(X1)|v1_subset_1(esk54_1(X1),X1))).#trainneg
cnf(i_0_473, plain, (v3_card_1(esk57_1(X1),X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_484, plain, (v3_card_1(esk62_1(X1),X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_394, plain, (v13_struct_0(esk30_1(X1),X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_367, plain, (v1_waybel_0(esk22_1(X1),X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_366, plain, (v2_waybel_0(esk22_1(X1),X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_446, plain, (v6_orders_2(esk46_1(X1),X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_335, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_xboole_0(X2,X1)))).#trainneg
cnf(i_0_333, plain, (v1_xboole_0(X1)|~v1_xboole_0(k2_xboole_0(X1,X2)))).#trainneg
cnf(i_0_107, plain, (v4_orders_2(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v13_struct_0(X1,np__1))).#trainneg
cnf(i_0_166, plain, (v2_yellow_0(X1)|~l1_orders_2(X1)|~v3_yellow_0(X1))).#trainneg
cnf(i_0_52, plain, (v2_orders_2(X1)|~v3_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_84, plain, (v3_ordinal1(X1)|~v1_ordinal1(X1)|~v2_ordinal1(X1))).#trainneg
cnf(i_0_167, plain, (v1_yellow_0(X1)|~l1_orders_2(X1)|~v3_yellow_0(X1))).#trainneg
cnf(i_0_196, plain, (v7_ordinal1(X1)|~v1_finset_1(X1)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_116, plain, (v3_relat_1(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_144, plain, (v2_relat_1(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_146, plain, (v8_struct_0(X1)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_202, plain, (v8_struct_0(X1)|~l1_struct_0(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_256, plain, (v2_finset_1(X1)|~v1_xboole_0(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_365, plain, (v1_xboole_0(X1)|m1_subset_1(esk21_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_403, plain, (v1_xboole_0(X1)|m1_subset_1(esk32_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_466, plain, (v1_xboole_0(X1)|m1_subset_1(esk54_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_481, plain, (v1_xboole_0(X1)|m1_subset_1(esk60_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_490, plain, (v1_xboole_0(X1)|m1_subset_1(esk64_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_457, plain, (v1_finset_1(X1)|m1_subset_1(esk50_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_338, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v1_zfmisc_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_106, plain, (v5_orders_2(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v13_struct_0(X1,np__1))).#trainneg
cnf(i_0_164, plain, (v1_xboole_0(X1)|v2_xxreal_0(X1)|v3_xxreal_0(X1)|~v1_xxreal_0(X1))).#trainneg
cnf(i_0_320, plain, (v1_xboole_0(u1_struct_0(X1))|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_341, plain, (v1_finset_1(u1_struct_0(X1))|~l1_struct_0(X1)|~v8_struct_0(X1))).#trainneg
cnf(i_0_390, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(esk28_1(X1)))).#trainneg
cnf(i_0_463, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(esk53_1(X1)))).#trainneg
cnf(i_0_392, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v1_zfmisc_1(esk29_1(X1)))).#trainneg
cnf(i_0_344, plain, (v8_struct_0(X1)|~l1_struct_0(X1)|~v1_finset_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_260, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v13_struct_0(X1,k1_xboole_0))).#trainneg
cnf(i_0_488, plain, (v1_zfmisc_1(X1)|m1_subset_1(esk63_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_493, plain, (v1_zfmisc_1(X1)|m1_subset_1(esk65_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_270, plain, (X1=X3|X2!=k1_tarski(X3)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_389, plain, (v2_struct_0(X1)|v1_zfmisc_1(esk28_1(X1))|~l1_struct_0(X1))).#trainneg
cnf(i_0_16, plain, (v1_xboole_0(X1)|v3_card_1(X1,np__1)|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_340, plain, (v1_zfmisc_1(u1_struct_0(X1))|~l1_struct_0(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_315, plain, (v4_funct_1(k1_tarski(X1))|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_242, plain, (v13_struct_0(X1,k1_xboole_0)|~l1_struct_0(X1)|~v2_struct_0(X1))).#trainneg
cnf(i_0_24, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v13_struct_0(X1,np__1))).#trainneg
cnf(i_0_195, plain, (v3_yellow_0(X1)|~l1_orders_2(X1)|~v2_yellow_0(X1)|~v1_yellow_0(X1))).#trainneg
cnf(i_0_111, plain, (v1_funcop_1(X1)|~v1_xboole_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_223, plain, (v1_finset_1(X2)|~v5_finset_1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_173, plain, (v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_239, plain, (v1_funct_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_240, plain, (v1_relat_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_76, plain, (v2_funcop_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_funcop_1(X1))).#trainneg
cnf(i_0_274, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X4,X2))).#trainneg
cnf(i_0_275, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k2_tarski(X2,X4))).#trainneg
cnf(i_0_368, plain, (m1_subset_1(esk22_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_447, plain, (m1_subset_1(esk46_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1))).#trainneg
cnf(i_0_295, plain, (m1_subset_1(k1_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_176, plain, (v1_zfmisc_1(X2)|~v1_zfmisc_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_17, plain, (v6_ordinal1(X2)|~v6_ordinal1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_85, plain, (v1_relat_1(X2)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_238, plain, (v5_finset_1(X2)|~v5_finset_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_258, plain, (v4_funct_1(X2)|~v4_funct_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_172, plain, (v3_funct_1(X1)|~v1_zfmisc_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_422, plain, (v2_struct_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v1_xboole_0(esk38_1(X1)))).#trainneg
cnf(i_0_250, plain, (v3_orders_2(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_79, plain, (v2_funct_1(X1)|~v1_xboole_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_140, plain, (v3_funct_1(X1)|~v1_xboole_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_181, plain, (v3_orders_2(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v9_waybel_1(X1))).#trainneg
cnf(i_0_105, plain, (v3_lattice3(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v13_struct_0(X1,np__1))).#trainneg
cnf(i_0_89, plain, (v2_waybel_1(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v13_struct_0(X1,np__1))).#trainneg
cnf(i_0_88, plain, (v10_waybel_1(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v13_struct_0(X1,np__1))).#trainneg
cnf(i_0_119, plain, (v1_waybel_2(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v13_struct_0(X1,np__1))).#trainneg
cnf(i_0_343, plain, (v1_finset_1(k2_xboole_0(X1,X2))|~v1_finset_1(X2)|~v1_finset_1(X1))).#trainneg
cnf(i_0_330, plain, (v1_relat_1(k2_xboole_0(X1,X2))|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_328, plain, (v5_finset_1(k2_tarski(X1,X2))|~v1_finset_1(X2)|~v1_finset_1(X1))).#trainneg
cnf(i_0_329, plain, (v5_finset_1(k2_xboole_0(X1,X2))|~v5_finset_1(X2)|~v5_finset_1(X1))).#trainneg
cnf(i_0_249, plain, (v4_orders_2(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_180, plain, (v4_orders_2(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v9_waybel_1(X1))).#trainneg
cnf(i_0_248, plain, (v5_orders_2(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_179, plain, (v5_orders_2(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v9_waybel_1(X1))).#trainneg
cnf(i_0_246, plain, (v2_lattice3(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_177, plain, (v2_lattice3(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v9_waybel_1(X1))).#trainneg
cnf(i_0_69, plain, (v2_lattice3(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_70, plain, (v2_struct_0(X1)|v1_lattice3(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_60, plain, (v2_lattice3(X1)|v2_struct_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v16_waybel_0(X1))).#trainneg
cnf(i_0_247, plain, (v2_struct_0(X1)|v1_lattice3(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_178, plain, (v2_struct_0(X1)|v1_lattice3(X1)|~l1_orders_2(X1)|~v9_waybel_1(X1))).#trainneg
cnf(i_0_22, plain, (v2_struct_0(X1)|v9_waybel_1(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_72, plain, (v2_struct_0(X1)|v3_lattice3(X1)|~l1_orders_2(X1)|~v1_waybel_5(X1))).#trainneg
cnf(i_0_122, plain, (v2_waybel_3(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v13_struct_0(X1,np__1))).#trainneg
cnf(i_0_19, plain, (v2_yellow_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v24_waybel_0(X1))).#trainneg
cnf(i_0_114, plain, (v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_98, plain, (v1_waybel_5(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~v13_struct_0(X1,np__1))).#trainneg
cnf(i_0_227, plain, (v2_struct_0(X1)|v2_yellow_0(X1)|~l1_orders_2(X1)|~v9_waybel_1(X1))).#trainneg
cnf(i_0_244, plain, (v2_struct_0(X1)|v2_waybel_1(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_206, plain, (v2_struct_0(X1)|v2_waybel_1(X1)|~l1_orders_2(X1)|~v9_waybel_1(X1))).#trainneg
cnf(i_0_243, plain, (v2_struct_0(X1)|v10_waybel_1(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_245, plain, (v2_struct_0(X1)|v3_yellow_0(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_136, plain, (v2_struct_0(X1)|v3_yellow_0(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_421, plain, (v2_struct_0(X1)|v1_finset_1(esk38_1(X1))|~v3_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_27, plain, (v2_struct_0(X1)|v24_waybel_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_38, plain, (v2_lattice3(X1)|v2_struct_0(X1)|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~v25_waybel_0(X1))).#trainneg
cnf(i_0_61, plain, (v2_struct_0(X1)|v1_lattice3(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v16_waybel_0(X1))).#trainneg
cnf(i_0_18, plain, (v2_struct_0(X1)|v13_struct_0(X1,np__1)|~l1_struct_0(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_154, plain, (v2_struct_0(X1)|v24_waybel_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v3_waybel_3(X1))).#trainneg
cnf(i_0_150, plain, (v2_struct_0(X1)|v24_waybel_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v2_waybel_2(X1))).#trainneg
cnf(i_0_157, plain, (v9_waybel_1(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_waybel_5(X1))).#trainneg
cnf(i_0_126, plain, (v3_lattice3(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_waybel_5(X1))).#trainneg
cnf(i_0_125, plain, (v3_waybel_3(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v1_waybel_5(X1))).#trainneg
cnf(i_0_26, plain, (v2_struct_0(X1)|v25_waybel_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_30, plain, (v2_struct_0(X1)|v1_yellow_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v25_waybel_0(X1))).#trainneg
cnf(i_0_391, plain, (v2_struct_0(X1)|m1_subset_1(esk28_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_464, plain, (v2_struct_0(X1)|m1_subset_1(esk53_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_393, plain, (v7_struct_0(X1)|m1_subset_1(esk29_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~l1_struct_0(X1))).#trainneg
cnf(i_0_267, plain, (X2=k1_tarski(X1)|esk5_2(X1,X2)=X1|r2_hidden(esk5_2(X1,X2),X2))).#trainneg
cnf(i_0_268, plain, (X2=k1_tarski(X1)|esk5_2(X1,X2)!=X1|~r2_hidden(esk5_2(X1,X2),X2))).#trainneg
cnf(i_0_420, plain, (v2_struct_0(X1)|v1_waybel_0(esk38_1(X1),X1)|~v3_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_419, plain, (v2_struct_0(X1)|v2_waybel_0(esk38_1(X1),X1)|~v3_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_203, plain, (v2_struct_0(X1)|v16_waybel_0(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_149, plain, (v2_struct_0(X1)|v1_waybel_2(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v2_waybel_2(X1))).#trainneg
cnf(i_0_318, plain, (v3_card_1(u1_struct_0(X2),X1)|~l1_struct_0(X2)|~v1_card_1(X1)|~v13_struct_0(X2,X1))).#trainneg
cnf(i_0_606, plain, (v2_lattice3(X1)|v2_struct_0(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~v1_xboole_0(esk73_1(X1)))).#trainneg
cnf(i_0_153, plain, (v2_struct_0(X1)|v2_waybel_3(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v3_waybel_3(X1))).#trainneg
cnf(i_0_59, plain, (v1_waybel_0(X1,X2)|~l1_orders_2(X2)|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_58, plain, (v2_waybel_0(X1,X2)|~l1_orders_2(X2)|~v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_83, plain, (v6_orders_2(X2,X1)|~l1_orders_2(X1)|~v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_186, plain, (v3_waybel_3(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v24_waybel_0(X1)|~v1_yellow_0(X1)|~v2_waybel_3(X1))).#trainneg
cnf(i_0_183, plain, (v2_struct_0(X1)|v2_waybel_2(X1)|~v3_orders_2(X1)|~l1_orders_2(X1)|~v24_waybel_0(X1)|~v1_waybel_2(X1))).#trainneg
cnf(i_0_316, plain, (v4_funct_1(k2_tarski(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_43, plain, (v2_struct_0(X1)|v1_lattice3(X1)|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~v2_yellow_0(X1)|~v25_waybel_0(X1))).#trainneg
cnf(i_0_87, plain, (v1_xboole_0(X1)|v1_subset_1(X2,X1)|~v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_423, plain, (v2_struct_0(X1)|m1_subset_1(esk38_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_283, plain, (r2_hidden(X1,X4)|r2_hidden(X1,X3)|X2!=k2_xboole_0(X3,X4)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_272, plain, (X3=k2_tarski(X1,X2)|esk6_3(X1,X2,X3)!=X2|~r2_hidden(esk6_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_273, plain, (X3=k2_tarski(X1,X2)|esk6_3(X1,X2,X3)!=X1|~r2_hidden(esk6_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_482, plain, (v2_struct_0(X1)|v5_waybel_6(esk61_1(X1),X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~v2_yellow_0(X1))).#trainneg
cnf(i_0_483, plain, (v2_struct_0(X1)|m1_subset_1(esk61_1(X1),u1_struct_0(X1))|~v5_orders_2(X1)|~l1_orders_2(X1)|~v2_yellow_0(X1))).#trainneg
cnf(i_0_313, plain, (v1_xboole_0(X2)|v1_xboole_0(X1)|m2_subset_1(esk13_2(X1,X2),X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_64, plain, (v2_struct_0(X2)|v1_waybel_0(X1,X2)|~l1_orders_2(X2)|~v16_waybel_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_603, plain, (v2_lattice3(X1)|v2_struct_0(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~r2_yellow_0(X1,esk73_1(X1)))).#trainneg
cnf(i_0_605, plain, (v2_lattice3(X1)|v2_struct_0(X1)|v1_finset_1(esk73_1(X1))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_63, plain, (v2_struct_0(X2)|v2_waybel_0(X1,X2)|~l1_orders_2(X2)|~v16_waybel_0(X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_508, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_214, plain, (v2_struct_0(X1)|v2_waybel_3(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~v3_lattice3(X1)|~v16_waybel_0(X1))).#trainneg
cnf(i_0_208, plain, (v2_struct_0(X1)|v2_waybel_2(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~v9_waybel_1(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_229, plain, (v2_struct_0(X1)|v9_waybel_1(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~v3_lattice3(X1)|~v2_waybel_1(X1)|~v2_waybel_2(X1))).#trainneg
cnf(i_0_33, plain, (v2_struct_0(X1)|v3_lattice3(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v24_waybel_0(X1)|~v1_yellow_0(X1))).#trainneg
cnf(i_0_93, plain, (v2_struct_0(X1)|v3_lattice3(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~v24_waybel_0(X1)|~v1_yellow_0(X1)|~v16_waybel_0(X1))).#trainneg
cnf(i_0_280, plain, (X3=k2_xboole_0(X1,X2)|~r2_hidden(esk7_3(X1,X2,X3),X3)|~r2_hidden(esk7_3(X1,X2,X3),X1))).#trainneg
cnf(i_0_266, plain, (k7_domain_1(X1,X2,X3)=k7_domain_1(X1,X3,X2)|v1_xboole_0(X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_290, plain, (m1_subset_1(k11_lattice3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_271, plain, (X3=k2_tarski(X1,X2)|esk6_3(X1,X2,X3)=X2|esk6_3(X1,X2,X3)=X1|r2_hidden(esk6_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_261, plain, (v2_struct_0(X1)|v11_waybel_1(X1)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_lattice3(X1)|~v2_waybel_1(X1)|~v10_waybel_1(X1)|~v3_yellow_0(X1))).#trainneg
cnf(i_0_510, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~r3_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_604, plain, (v2_lattice3(X1)|v2_struct_0(X1)|m1_subset_1(esk73_1(X1),k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_66, plain, (v2_struct_0(X2)|v1_waybel_0(X1,X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~v6_orders_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_65, plain, (v2_struct_0(X2)|v2_waybel_0(X1,X2)|~v3_orders_2(X2)|~v4_orders_2(X2)|~v5_orders_2(X2)|~l1_orders_2(X2)|~v6_orders_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2))))).#trainneg
cnf(i_0_512, plain, (v2_struct_0(X1)|r3_orders_2(X1,X2,X2)|~v3_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_602, plain, (r2_yellow_0(X1,k2_xboole_0(X2,X3))|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2))).#trainneg
cnf(i_0_291, plain, (m1_subset_1(k12_lattice3(X1,X2,X3),u1_struct_0(X1))|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_509, plain, (v2_struct_0(X1)|r3_orders_2(X1,X2,X3)|~v3_orders_2(X1)|~l1_orders_2(X1)|~r1_orders_2(X1,X2,X3)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_596, plain, (k1_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))=X2|v2_struct_0(X1)|~v3_orders_2(X1)|~v5_orders_2(X1)|~l1_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_568, plain, (v1_xboole_0(X3)|r2_hidden(esk72_3(X1,X2,X3),X3)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))))).#trainneg
cnf(i_0_567, plain, (v1_xboole_0(X3)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))))).#trainneg
cnf(i_0_566, plain, (v1_xboole_0(X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),X2)|~v3_orders_2(X1)|~v4_orders_2(X1)|~v5_orders_2(X1)|~v2_lattice3(X1)|~l1_orders_2(X1)|~v1_finset_1(X3)|~v1_xboole_0(k1_xboole_0)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2))).#trainneg
cnf(i_0_870, plain, (~v1_xboole_0(esk43_0))).#trainneg
cnf(i_0_1149, plain, (~v1_xboole_0(esk41_0))).#trainneg
cnf(i_0_1150, plain, (~v1_xboole_0(esk40_0))).#trainneg
cnf(i_0_1170, plain, (~v1_xboole_0(esk49_0))).#trainneg
cnf(i_0_1171, plain, (~v1_xboole_0(esk48_0))).#trainneg
cnf(i_0_1172, plain, (~v2_xxreal_0(esk49_0))).#trainneg
cnf(i_0_1246, negated_conjecture, (esk4_1(X1)=X2|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|X1!=k1_tarski(X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1709, negated_conjecture, (r3_orders_2(esk1_0,X1,X1)|v2_struct_0(esk1_0)|v1_xboole_0(X2)|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X2),esk2_0)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_274, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X3,X1))).#trainneg
cnf(i_0_1311, negated_conjecture, (v3_ordinal1(esk2_0)|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_275, plain, (r2_hidden(X1,X2)|X2!=k2_tarski(X1,X3))).#trainneg
cnf(i_0_1268, plain, (~v1_xboole_0(k1_numbers))).#trainneg
cnf(i_0_1280, plain, (~v1_finset_1(k1_numbers))).#trainneg
cnf(i_0_1445, plain, (~v1_xboole_0(esk58_0))).#trainneg
cnf(i_0_1173, plain, (~v2_xxreal_0(esk48_0))).#trainneg
cnf(i_0_1444, plain, (~v1_zfmisc_1(esk58_0))).#trainneg
cnf(i_0_1166, plain, (~m1_subset_1(esk43_0,k4_ordinal1))).#trainneg
cnf(i_0_1167, plain, (~m1_subset_1(k4_ordinal1,k4_ordinal1))).#trainneg
cnf(i_0_1710, negated_conjecture, (r3_orders_2(esk1_0,X1,X1)|v2_struct_0(esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1755, plain, (k2_yellow_0(X1,k2_xboole_0(k6_domain_1(u1_struct_0(X1),X2),X3))=k12_lattice3(X1,X2,k2_yellow_0(X1,X3))|v2_struct_0(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_1706, plain, (r3_orders_2(X1,X2,X2)|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1802, plain, (r1_tarski(esk70_3(X1,X2,esk38_1(X1)),esk38_1(X1))|r1_orders_2(X1,esk72_3(X1,X2,esk38_1(X1)),X2)|v1_xboole_0(esk38_1(X1))|v2_struct_0(X1)|~r1_orders_2(X1,k2_yellow_0(X1,esk38_1(X1)),X2)|~v1_finset_1(esk38_1(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1760, plain, (k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X2))=k12_lattice3(X1,X2,X3)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1328, negated_conjecture, (v1_funct_1(esk2_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1130, negated_conjecture, (m1_subset_1(esk4_1(X1),X1)|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1131, plain, (esk56_0=k1_xboole_0)).#trainneg
cnf(i_0_1132, plain, (esk55_0=k1_xboole_0)).#trainneg
cnf(i_0_1134, plain, (esk25_0=k1_xboole_0)).#trainneg
cnf(i_0_1135, plain, (v1_finset_1(esk59_0))).#trainneg
cnf(i_0_1136, plain, (v1_finset_1(esk52_0))).#trainneg
cnf(i_0_1137, plain, (v1_card_1(esk59_0))).#trainneg
cnf(i_0_1138, plain, (v1_card_1(esk52_0))).#trainneg
cnf(i_0_1140, plain, (v3_ordinal1(esk59_0))).#trainneg
cnf(i_0_1764, plain, (r2_hidden(esk72_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),k6_domain_1(u1_struct_0(X1),X3))|r2_hidden(esk69_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),k6_domain_1(u1_struct_0(X1),X3))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|v2_struct_0(X1)|~r1_orders_2(X1,X3,X2)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1782, plain, (r2_hidden(esk69_3(X1,X2,esk38_1(X1)),esk38_1(X1))|r1_orders_2(X1,esk72_3(X1,X2,esk38_1(X1)),X2)|v1_xboole_0(esk38_1(X1))|v2_struct_0(X1)|~r1_orders_2(X1,k2_yellow_0(X1,esk38_1(X1)),X2)|~v1_finset_1(esk38_1(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1818, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,X4,X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))))).#trainneg
cnf(i_0_1644, plain, (k7_domain_1(X1,X3,X2)=k2_tarski(X2,X3)|v1_xboole_0(X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1621, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|~r2_hidden(X1,k2_xboole_0(X2,X3)))).#trainneg
cnf(i_0_1345, negated_conjecture, (v1_relat_1(esk2_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1490, plain, (r2_hidden(X1,X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_1133, plain, (esk37_1(X1)=k1_xboole_0)).#trainneg
cnf(i_0_1461, plain, (v2_lattice3(esk42_0))).#trainneg
cnf(i_0_1462, plain, (v2_lattice3(esk23_0))).#trainneg
cnf(i_0_1141, plain, (v3_ordinal1(esk52_0))).#trainneg
cnf(i_0_1145, plain, (v1_xxreal_0(esk26_0))).#trainneg
cnf(i_0_1155, plain, (v1_xxreal_0(esk59_0))).#trainneg
cnf(i_0_1156, plain, (v1_xxreal_0(esk52_0))).#trainneg
cnf(i_0_1312, plain, (v3_ordinal1(k1_xboole_0))).#trainneg
cnf(i_0_1726, plain, (k1_yellow_0(X1,k1_tarski(X2))=X2|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1766, plain, (r1_tarski(esk70_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),k6_domain_1(u1_struct_0(X1),X3))|r2_hidden(esk72_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),k6_domain_1(u1_struct_0(X1),X3))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|v2_struct_0(X1)|~r1_orders_2(X1,X3,X2)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1790, plain, (r2_hidden(esk69_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),k6_domain_1(u1_struct_0(X1),X3))|r1_orders_2(X1,esk72_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),X2)|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X3)),X2)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1765, plain, (r2_hidden(esk72_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k7_domain_1(u1_struct_0(X1),X3,X4))|r2_hidden(esk69_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k7_domain_1(u1_struct_0(X1),X3,X4))|v1_xboole_0(k7_domain_1(u1_struct_0(X1),X3,X4))|~r1_orders_2(X1,k12_lattice3(X1,X3,X4),X2)|~v1_finset_1(k7_domain_1(u1_struct_0(X1),X3,X4))|~m1_subset_1(k7_domain_1(u1_struct_0(X1),X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1810, plain, (r1_tarski(esk70_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),k6_domain_1(u1_struct_0(X1),X3))|r1_orders_2(X1,esk72_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),X2)|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X3)),X2)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1767, plain, (r1_tarski(esk70_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k7_domain_1(u1_struct_0(X1),X3,X4))|r2_hidden(esk72_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k7_domain_1(u1_struct_0(X1),X3,X4))|v1_xboole_0(k7_domain_1(u1_struct_0(X1),X3,X4))|~r1_orders_2(X1,k12_lattice3(X1,X3,X4),X2)|~v1_finset_1(k7_domain_1(u1_struct_0(X1),X3,X4))|~m1_subset_1(k7_domain_1(u1_struct_0(X1),X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1294, negated_conjecture, (v1_finset_1(esk2_0)|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1574, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k4_ordinal1))).#trainneg
cnf(i_0_1465, plain, (v2_lattice3(esk14_0))).#trainneg
cnf(i_0_1314, plain, (v3_ordinal1(np__1))).#trainneg
cnf(i_0_1466, plain, (v3_lattice3(esk24_0))).#trainneg
cnf(i_0_1503, plain, (v1_lattice3(esk23_0))).#trainneg
cnf(i_0_1504, plain, (v1_lattice3(esk14_0))).#trainneg
cnf(i_0_1508, plain, (v9_waybel_1(esk24_0))).#trainneg
cnf(i_0_1756, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k12_lattice3(X1,X3,X4))=k2_yellow_0(X1,k2_xboole_0(X2,k7_domain_1(u1_struct_0(X1),X3,X4)))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_1757, plain, (k12_lattice3(X1,k12_lattice3(X1,X2,X3),k2_yellow_0(X1,X4))=k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X2,X3),X4))|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_1387, plain, (v1_zfmisc_1(k4_ordinal1)|~v1_zfmisc_1(k1_numbers))).#trainneg
cnf(i_0_1401, plain, (v6_ordinal1(k4_ordinal1)|~v6_ordinal1(k1_numbers))).#trainneg
cnf(i_0_1415, plain, (v1_relat_1(k4_ordinal1)|~v1_relat_1(k1_numbers))).#trainneg
cnf(i_0_1429, plain, (v5_finset_1(k4_ordinal1)|~v5_finset_1(k1_numbers))).#trainneg
cnf(i_0_1443, plain, (v4_funct_1(k4_ordinal1)|~v4_funct_1(k1_numbers))).#trainneg
cnf(i_0_1239, plain, (v1_xboole_0(u1_struct_0(X1))|~v2_struct_0(X1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_1247, plain, (X1=X2|v1_xboole_0(X3)|X3!=k1_tarski(X2)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_1590, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,k6_domain_1(X2,X3))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_1162, plain, (k2_xboole_0(k1_xboole_0,X1)=X1)).#trainneg
cnf(i_0_1157, plain, (v1_xxreal_0(X1)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_1221, plain, (v1_xxreal_0(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_1186, plain, (v1_xxreal_0(X1)|~m1_subset_1(X1,k1_numbers))).#trainneg
cnf(i_0_1197, plain, (v1_zfmisc_1(k1_tarski(X1)))).#trainneg
cnf(i_0_1235, plain, (v2_struct_0(esk11_0)|~v1_xboole_0(u1_struct_0(esk11_0)))).#trainneg
cnf(i_0_1193, plain, (~r2_hidden(X1,esk47_1(X2))|~v1_xboole_0(X2))).#trainneg
cnf(i_0_1188, plain, (v1_xxreal_0(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_1189, plain, (v1_card_1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_1190, plain, (v3_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_1222, plain, (v1_card_1(X1)|~v3_ordinal1(X1)|~v1_finset_1(X1))).#trainneg
cnf(i_0_1256, plain, (v1_zfmisc_1(u1_struct_0(X1))|~l1_struct_0(X1)|~v13_struct_0(X1,np__1))).#trainneg
cnf(i_0_1770, plain, (r1_tarski(esk70_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),k6_domain_1(u1_struct_0(X1),X3))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|m1_subset_1(esk72_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),u1_struct_0(X1))|v2_struct_0(X1)|~r1_orders_2(X1,X3,X2)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1768, plain, (r2_hidden(esk69_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),k6_domain_1(u1_struct_0(X1),X3))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|m1_subset_1(esk72_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),u1_struct_0(X1))|v2_struct_0(X1)|~r1_orders_2(X1,X3,X2)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1659, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,k7_domain_1(X2,X3,X4))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_1482, plain, (v2_yellow_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v11_waybel_1(X1))).#trainneg
cnf(i_0_471, plain, (v1_xxreal_0(k1_xboole_0))).#trainneg
cnf(i_0_1890, plain, (~m1_subset_1(k1_numbers,k4_ordinal1))).#trainneg
cnf(i_0_1950, plain, (~v1_xboole_0(u1_struct_0(esk42_0)))).#trainneg
cnf(i_0_1951, plain, (~v1_xboole_0(u1_struct_0(esk24_0)))).#trainneg
cnf(i_0_1952, plain, (~v1_xboole_0(u1_struct_0(esk23_0)))).#trainneg
cnf(i_0_1953, plain, (~v1_xboole_0(u1_struct_0(esk14_0)))).#trainneg
cnf(i_0_1978, negated_conjecture, (r2_hidden(esk4_1(k2_xboole_0(X1,X2)),X2)|r2_hidden(esk4_1(k2_xboole_0(X1,X2)),X1)|v1_xboole_0(k2_xboole_0(X1,X2))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X1,X2)),esk2_0)|~v1_finset_1(k2_xboole_0(X1,X2))|~m1_subset_1(k2_xboole_0(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1892, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1887, plain, (r2_hidden(X1,k2_tarski(X1,X2)))).#trainneg
cnf(i_0_1875, plain, (r2_hidden(X1,k2_tarski(X2,X1)))).#trainneg
cnf(i_0_1483, plain, (v2_yellow_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_469, plain, (v1_xcmplx_0(k1_xboole_0))).#trainneg
cnf(i_0_467, plain, (v1_xreal_0(k1_xboole_0))).#trainneg
cnf(i_0_2257, plain, (v1_xxreal_0(esk31_0))).#trainneg
cnf(i_0_2177, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~r1_tarski(X1,k6_domain_1(X2,X3))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_2236, plain, (X1=X2|X1=X3|~m1_subset_1(X1,k2_tarski(X3,X2)))).#trainneg
cnf(i_0_1897, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1860, plain, (k2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3)))=k2_yellow_0(X1,k2_xboole_0(X2,X3))|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_418, plain, (m1_subset_1(k1_xboole_0,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_2256, plain, (v1_xxreal_0(esk35_0)|~v1_finset_1(esk35_0))).#trainneg
cnf(i_0_2258, plain, (v1_xxreal_0(esk19_0)|~v1_finset_1(esk19_0))).#trainneg
cnf(i_0_2261, plain, (v1_xxreal_0(np__1)|~v1_finset_1(np__1))).#trainneg
cnf(i_0_2278, plain, (v1_card_1(esk35_0)|~v1_finset_1(esk35_0))).#trainneg
cnf(i_0_2280, plain, (v1_card_1(esk19_0)|~v1_finset_1(esk19_0))).#trainneg
cnf(i_0_2281, plain, (v1_card_1(k1_xboole_0)|~v1_finset_1(k1_xboole_0))).#trainneg
cnf(i_0_1866, plain, (k2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3)))=k2_yellow_0(X1,k2_xboole_0(X3,X2))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,X3))).#trainneg
cnf(i_0_2045, plain, (r1_tarski(esk70_3(X1,X2,k2_tarski(X3,X4)),k2_tarski(X3,X4))|r2_hidden(esk72_3(X1,X2,k2_tarski(X3,X4)),k2_tarski(X3,X4))|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k12_lattice3(X1,X3,X4),X2)|~m1_subset_1(k2_tarski(X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2008, plain, (r1_tarski(esk70_3(X1,X2,k1_tarski(X3)),k1_tarski(X3))|r2_hidden(esk72_3(X1,X2,k1_tarski(X3)),k1_tarski(X3))|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,X3,X2)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2179, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~r1_tarski(X1,k7_domain_1(X2,X3,X4))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_2283, plain, (v1_card_1(np__1)|~v1_finset_1(np__1))).#trainneg
cnf(i_0_2218, plain, (X1=X2|k1_zfmisc_1(X1)!=k1_tarski(X2))).#trainneg
cnf(i_0_2263, plain, (v1_xxreal_0(X1)|~v1_finset_1(X1)|~v1_card_1(X1))).#trainneg
cnf(i_0_1985, plain, (~v1_xboole_0(k2_xboole_0(X2,X3))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_2029, plain, (~v1_xboole_0(k2_xboole_0(X2,X3))|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_1949, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_1895, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(k1_tarski(X2),X3))|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,k1_tarski(X2))|~r2_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2016, plain, (r2_hidden(esk69_3(X1,X2,k2_tarski(X3,X4)),k2_tarski(X3,X4))|r2_hidden(esk72_3(X1,X2,k2_tarski(X3,X4)),k2_tarski(X3,X4))|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k12_lattice3(X1,X3,X4),X2)|~m1_subset_1(k2_tarski(X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2226, negated_conjecture, (v1_xboole_0(X1)|m1_subset_1(esk4_1(k6_domain_1(X1,X2)),X1)|v1_xboole_0(k6_domain_1(X1,X2))|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X2,X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k6_domain_1(X1,X2)),esk2_0)|~v1_finset_1(k6_domain_1(X1,X2))|~m1_subset_1(k6_domain_1(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1514, plain, (v1_yellow_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~v3_lattice3(X1))).#trainneg
cnf(i_0_2101, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,k1_tarski(X3))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_1512, plain, (v24_waybel_0(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~v1_waybel_5(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1))).#trainneg
cnf(i_0_2205, negated_conjecture, (esk2_0=X1|v1_xboole_0(u1_struct_0(esk1_0))|k1_tarski(X1)!=u1_struct_0(esk1_0))).#trainneg
cnf(i_0_2181, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k4_ordinal1))).#trainneg
cnf(i_0_1922, negated_conjecture, (v5_waybel_6(esk2_0,esk1_0)|~r2_hidden(X1,esk4_1(k1_zfmisc_1(X2)))|~v1_xboole_0(X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(X2)),esk2_0)|~v1_finset_1(k1_zfmisc_1(X2))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2272, negated_conjecture, (v1_xboole_0(esk47_1(X1))|v5_waybel_6(esk2_0,esk1_0)|~v1_xboole_0(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk47_1(X1)),esk2_0)|~v1_finset_1(esk47_1(X1))|~m1_subset_1(esk47_1(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1846, negated_conjecture, (esk4_1(k1_tarski(X1))=X1|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X1)),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1592, plain, (v2_struct_0(X1)|v2_lattice3(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_1879, plain, (r1_orders_2(X1,k2_yellow_0(X1,X2),X3)|r1_orders_2(X1,k2_yellow_0(X1,X4),X3)|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X4,X2)),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,X4)|~v4_orders_2(X1))).#trainneg
cnf(i_0_1959, plain, (r2_hidden(esk69_3(X1,X2,k1_tarski(X3)),k1_tarski(X3))|r2_hidden(esk72_3(X1,X2,k1_tarski(X3)),k1_tarski(X3))|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,X3,X2)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1870, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),X3)=k2_yellow_0(X1,k2_xboole_0(k1_tarski(X3),X2))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,k1_tarski(X3))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_1593, plain, (v1_lattice3(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~v7_struct_0(X1))).#trainneg
cnf(i_0_1555, plain, (m1_subset_1(esk5_2(X1,X2),X2)|esk5_2(X1,X2)=X1|X2=k1_tarski(X1))).#trainneg
cnf(i_0_2200, plain, (esk12_1(X1)=X2|v1_xboole_0(X1)|X1!=k1_tarski(X2))).#trainneg
cnf(i_0_1988, plain, (X1=X2|k2_xboole_0(X3,X4)!=k1_tarski(X2)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_2032, plain, (X1=X2|k2_xboole_0(X3,X4)!=k1_tarski(X2)|~r2_hidden(X1,X4))).#trainneg
cnf(i_0_1613, plain, (v1_zfmisc_1(esk38_1(X1))|v2_struct_0(X1)|~v1_zfmisc_1(u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1614, plain, (v6_ordinal1(esk38_1(X1))|v2_struct_0(X1)|~v6_ordinal1(u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1615, plain, (v1_relat_1(esk38_1(X1))|v2_struct_0(X1)|~v1_relat_1(u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1616, plain, (v5_finset_1(esk38_1(X1))|v2_struct_0(X1)|~v5_finset_1(u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1617, plain, (v4_funct_1(esk38_1(X1))|v2_struct_0(X1)|~v4_funct_1(u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1882, plain, (r1_orders_2(X1,X2,X3)|r1_orders_2(X1,X4,X3)|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(X4,X2)),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1894, plain, (k2_yellow_0(X1,k2_xboole_0(X2,k6_domain_1(u1_struct_0(X1),X3)))=k12_lattice3(X1,k2_yellow_0(X1,X2),X3)|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1869, plain, (k2_yellow_0(X1,k2_xboole_0(k6_domain_1(u1_struct_0(X1),X3),X2))=k12_lattice3(X1,k2_yellow_0(X1,X2),X3)|v2_struct_0(X1)|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X3))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_2180, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~r1_tarski(X1,k2_tarski(X3,X4))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_2364, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk2_0,esk2_0))).#trainneg
cnf(i_0_2451, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_1557, plain, (esk5_2(X1,X2)=X3|esk5_2(X1,X2)=X1|X2=k1_tarski(X1)|X2!=k1_tarski(X3))).#trainneg
cnf(i_0_2582, plain, (v24_waybel_0(esk24_0))).#trainneg
cnf(i_0_2363, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk4_1(X1),esk4_1(X1))|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2402, plain, (v2_struct_0(X1)|r1_orders_2(X1,esk9_2(X1,X2),esk9_2(X1,X2))|v5_waybel_6(X2,X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2403, plain, (v2_struct_0(X1)|r1_orders_2(X1,esk8_2(X1,X2),esk8_2(X1,X2))|v5_waybel_6(X2,X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2624, negated_conjecture, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,esk4_1(esk70_3(X1,X2,X3)),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk4_1(esk70_3(X1,X2,X3)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(X1,X2,X3)),esk2_0)|~v1_finset_1(esk70_3(X1,X2,X3))|~m1_subset_1(esk70_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2597, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X1)),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2443, plain, (v1_finset_1(k1_xboole_0)|~v1_finset_1(X1))).#trainneg
cnf(i_0_1575, plain, (X1=k1_tarski(X2)|v1_xboole_0(X1)|esk5_2(X2,X1)!=X2|~m1_subset_1(esk5_2(X2,X1),X1))).#trainneg
cnf(i_0_2505, plain, (v1_xxreal_0(esk15_0)|~v1_finset_1(esk15_0))).#trainneg
cnf(i_0_2358, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk61_1(esk1_0),esk61_1(esk1_0))|~v2_yellow_0(esk1_0))).#trainneg
cnf(i_0_2359, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk12_1(u1_struct_0(esk1_0)),esk12_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2362, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk4_1(u1_struct_0(esk1_0)),esk4_1(u1_struct_0(esk1_0)))|v1_xboole_0(u1_struct_0(esk1_0))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,u1_struct_0(esk1_0)),esk2_0)|~v1_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2410, plain, (v2_struct_0(X1)|r1_orders_2(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X2))|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2584, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~r1_tarski(X1,k4_ordinal1))).#trainneg
cnf(i_0_1597, plain, (v2_waybel_3(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~v1_waybel_5(X1)|~v1_lattice3(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1))).#trainneg
cnf(i_0_2383, plain, (~r2_hidden(k1_tarski(X1),X1))).#trainneg
cnf(i_0_2450, plain, (~r2_hidden(X1,k1_xboole_0)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_2444, plain, (v1_zfmisc_1(k1_xboole_0)|~v1_zfmisc_1(X1))).#trainneg
cnf(i_0_2445, plain, (v6_ordinal1(k1_xboole_0)|~v6_ordinal1(X1))).#trainneg
cnf(i_0_2446, plain, (v1_relat_1(k1_xboole_0)|~v1_relat_1(X1))).#trainneg
cnf(i_0_2447, plain, (v5_finset_1(k1_xboole_0)|~v5_finset_1(X1))).#trainneg
cnf(i_0_2448, plain, (v4_funct_1(k1_xboole_0)|~v4_funct_1(X1))).#trainneg
cnf(i_0_2360, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk9_2(esk1_0,X1),esk9_2(esk1_0,X1))|v5_waybel_6(X1,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2609, plain, (r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3))),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3)))|~r2_yellow_0(X1,esk70_3(X1,X2,X3))|~v5_waybel_6(X2,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1))).#trainneg
cnf(i_0_2454, negated_conjecture, (esk4_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)))=k2_yellow_0(esk1_0,X1)|esk4_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)))=k2_yellow_0(esk1_0,X2)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X2,X1)),esk2_0)|~m1_subset_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_yellow_0(esk1_0,X1)|~r2_yellow_0(esk1_0,X2))).#trainneg
cnf(i_0_2411, negated_conjecture, (esk4_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)))=k2_yellow_0(esk1_0,X1)|esk4_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)))=k2_yellow_0(esk1_0,X2)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X1,X2)),esk2_0)|~m1_subset_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,X1))).#trainneg
cnf(i_0_2368, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X1)))).#trainneg
cnf(i_0_2596, negated_conjecture, (v5_waybel_6(esk2_0,esk1_0)|m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X1)),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2371, plain, (~r2_hidden(k2_tarski(X1,X2),X1))).#trainneg
cnf(i_0_2377, plain, (~r2_hidden(k2_tarski(X2,X1),X1))).#trainneg
cnf(i_0_2382, plain, (m1_subset_1(X1,k1_tarski(X1)))).#trainneg
cnf(i_0_2666, plain, (esk12_1(k1_tarski(X1))=X1)).#trainneg
cnf(i_0_2365, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X1),k1_yellow_0(esk1_0,X1)))).#trainneg
cnf(i_0_2452, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,k1_xboole_0))).#trainneg
cnf(i_0_2388, plain, (r1_tarski(esk70_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|r2_hidden(esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|v5_waybel_6(X2,X1)|~m1_subset_1(k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_2361, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk8_2(esk1_0,X1),esk8_2(esk1_0,X1))|v5_waybel_6(X1,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2394, plain, (r1_tarski(k6_domain_1(X1,X2),X1)|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_2535, plain, (r2_hidden(esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|v5_waybel_6(X2,X1)|~r2_hidden(X3,esk70_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))))|~r1_orders_2(X1,X3,X2)|~m1_subset_1(k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_2387, plain, (r2_hidden(esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|r2_hidden(esk69_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|v5_waybel_6(X2,X1)|~m1_subset_1(k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_2408, plain, (v2_struct_0(X1)|r1_orders_2(X1,k11_lattice3(X1,X2,X3),k11_lattice3(X1,X2,X3))|~l1_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2544, plain, (r2_hidden(esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|r2_hidden(esk69_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|v5_waybel_6(X2,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2713, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,X4,X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,k1_tarski(esk69_3(X1,X2,X3))))).#trainneg
cnf(i_0_2133, plain, (m1_subset_1(X1,X2)|v1_xboole_0(X2)|~r2_hidden(X1,k2_tarski(X3,X4))|~m1_subset_1(X4,X2)|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_1628, plain, (X1=k2_tarski(X2,X3)|v1_xboole_0(X1)|esk6_3(X2,X3,X1)!=X3|~m1_subset_1(esk6_3(X2,X3,X1),X1))).#trainneg
cnf(i_0_2384, plain, (X1=X2|k1_tarski(X1)!=k1_tarski(X2))).#trainneg
cnf(i_0_2372, plain, (X1=X2|k2_tarski(X1,X3)!=k1_tarski(X2))).#trainneg
cnf(i_0_2378, plain, (X1=X2|k2_tarski(X3,X1)!=k1_tarski(X2))).#trainneg
cnf(i_0_2593, negated_conjecture, (v5_waybel_6(esk2_0,esk1_0)|v1_xboole_0(esk4_1(k1_zfmisc_1(X1)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(X1)),esk2_0)|~v1_finset_1(k1_zfmisc_1(X1))|~v1_xboole_0(X1)|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk4_1(k1_zfmisc_1(X1))),esk2_0)|~v1_finset_1(esk4_1(k1_zfmisc_1(X1)))|~m1_subset_1(esk4_1(k1_zfmisc_1(X1)),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2679, plain, (X1=X2|k2_xboole_0(k1_tarski(X1),X3)!=k1_tarski(X2))).#trainneg
cnf(i_0_2697, plain, (X1=X2|k2_xboole_0(X3,k1_tarski(X1))!=k1_tarski(X2))).#trainneg
cnf(i_0_2682, plain, (X1=X2|k2_xboole_0(k2_tarski(X3,X1),X4)!=k1_tarski(X2))).#trainneg
cnf(i_0_2490, plain, (r1_tarski(esk70_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|r2_hidden(esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|v5_waybel_6(X2,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2607, plain, (r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3))),X2)|r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3)))|~r2_yellow_0(X1,esk70_3(X1,X2,X3))|~v5_waybel_6(X2,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1))).#trainneg
cnf(i_0_2409, plain, (v2_struct_0(X1)|r1_orders_2(X1,k12_lattice3(X1,X2,X3),k12_lattice3(X1,X2,X3))|~l1_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_lattice3(X1)|~v5_orders_2(X1))).#trainneg
cnf(i_0_2536, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(k1_tarski(X2),X3))|~r2_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_2616, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),X3)=k2_yellow_0(X1,k2_xboole_0(k1_tarski(X3),X2))|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1923, negated_conjecture, (m1_subset_1(X1,X2)|v5_waybel_6(esk2_0,esk1_0)|~r2_hidden(X1,esk4_1(k1_zfmisc_1(X2)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(X2)),esk2_0)|~v1_finset_1(k1_zfmisc_1(X2))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1629, plain, (X1=k2_tarski(X2,X3)|v1_xboole_0(X1)|esk6_3(X2,X3,X1)!=X2|~m1_subset_1(esk6_3(X2,X3,X1),X1))).#trainneg
cnf(i_0_2370, plain, (m1_subset_1(X1,k2_tarski(X1,X2)))).#trainneg
cnf(i_0_2776, plain, (v1_finset_1(k1_xboole_0))).#trainneg
cnf(i_0_2846, plain, (v1_zfmisc_1(k1_xboole_0))).#trainneg
cnf(i_0_2750, negated_conjecture, (r1_orders_2(esk1_0,esk2_0,esk2_0))).#trainneg
cnf(i_0_2798, negated_conjecture, (esk4_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)))=esk8_2(esk1_0,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|esk9_2(esk1_0,esk2_0)!=esk8_2(esk1_0,esk2_0)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2911, negated_conjecture, (r2_hidden(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk38_1(esk1_0)),esk38_1(esk1_0)),esk38_1(esk1_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk38_1(esk1_0)),esk38_1(esk1_0)),k2_yellow_0(esk1_0,esk38_1(esk1_0)))|v2_struct_0(esk1_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk38_1(esk1_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2912, negated_conjecture, (r1_tarski(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk38_1(esk1_0)),esk38_1(esk1_0)),esk38_1(esk1_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk38_1(esk1_0)),esk38_1(esk1_0)),k2_yellow_0(esk1_0,esk38_1(esk1_0)))|v2_struct_0(esk1_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk38_1(esk1_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2764, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,X1)),esk2_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~v1_finset_1(X1)|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk4_1(esk70_3(esk1_0,esk2_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2762, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,X1,X1)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X1)),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2833, plain, (v1_xboole_0(X1)|~m1_subset_1(k1_tarski(X1),X1))).#trainneg
cnf(i_0_1634, plain, (X1=k2_xboole_0(X2,X3)|v1_xboole_0(X1)|~r2_hidden(esk7_3(X2,X3,X1),X3)|~m1_subset_1(esk7_3(X2,X3,X1),X1))).#trainneg
cnf(i_0_2957, plain, (r1_tarski(k1_xboole_0,X1))).#trainneg
cnf(i_0_2802, negated_conjecture, (v5_waybel_6(esk2_0,esk1_0)|m1_subset_1(esk9_2(esk1_0,esk2_0),k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)))|esk4_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)))=esk8_2(esk1_0,esk2_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0))),esk2_0)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2831, plain, (~r2_hidden(k1_tarski(k2_xboole_0(X1,X2)),X2))).#trainneg
cnf(i_0_2810, plain, (r2_hidden(esk72_3(X1,k2_yellow_0(X1,X2),X2),X2)|r2_hidden(esk69_3(X1,k2_yellow_0(X1,X2),X2),X2)|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2641, plain, (v1_lattice3(X1)|v2_struct_0(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~l1_struct_0(X1)|~v1_zfmisc_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_3160, plain, (v1_finset_1(X1)|~v5_finset_1(k2_tarski(X1,X2)))).#trainneg
cnf(i_0_1635, plain, (v1_xboole_0(X1)|m1_subset_1(esk13_2(X2,X1),X1)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_2832, plain, (~r2_hidden(k1_tarski(k2_xboole_0(X1,X2)),X1))).#trainneg
cnf(i_0_2830, plain, (m1_subset_1(X1,k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_2966, plain, (m1_subset_1(k6_domain_1(k4_ordinal1,X1),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_3161, plain, (v3_ordinal1(X1)|~v3_ordinal1(k2_tarski(X1,X2)))).#trainneg
cnf(i_0_3162, plain, (v1_funct_1(X1)|~v4_funct_1(k2_tarski(X1,X2)))).#trainneg
cnf(i_0_2625, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,X4,X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,esk70_3(X1,X2,X3)))).#trainneg
cnf(i_0_2748, plain, (r1_tarski(k2_tarski(X1,X2),X3)|v1_xboole_0(X3)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_2600, plain, (v2_struct_0(X1)|v2_lattice3(X1)|~l1_orders_2(X1)|~v3_orders_2(X1)|~l1_struct_0(X1)|~v1_zfmisc_1(u1_struct_0(X1)))).#trainneg
cnf(i_0_2941, plain, (v1_xboole_0(X1)|~m1_subset_1(k2_tarski(X1,X2),X1))).#trainneg
cnf(i_0_1637, plain, (v1_xboole_0(X1)|m1_subset_1(esk13_2(X2,X1),X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_3163, plain, (v1_relat_1(X1)|~v4_funct_1(k2_tarski(X1,X2)))).#trainneg
cnf(i_0_2946, plain, (v1_xboole_0(X2)|~m1_subset_1(k2_tarski(X1,X2),X2))).#trainneg
cnf(i_0_2958, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_2539, plain, (r2_hidden(esk72_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))),k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)))|r2_hidden(esk69_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))),k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)))|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X4,X3)),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X4))).#trainneg
cnf(i_0_2538, plain, (r2_hidden(esk72_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))),k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)))|r2_hidden(esk69_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))),k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)))|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X3,X4)),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,X3))).#trainneg
cnf(i_0_2917, negated_conjecture, (r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|r2_hidden(esk69_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|v1_xboole_0(X1)|v2_struct_0(esk1_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2811, plain, (r1_tarski(esk70_3(X1,k2_yellow_0(X1,X2),X2),X2)|r2_hidden(esk72_3(X1,k2_yellow_0(X1,X2),X2),X2)|v1_xboole_0(X2)|v2_struct_0(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2176, negated_conjecture, (r1_tarski(X1,X2)|v5_waybel_6(esk2_0,esk1_0)|~r1_tarski(X1,esk4_1(k1_zfmisc_1(X2)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(X2)),esk2_0)|~v1_finset_1(k1_zfmisc_1(X2))|~m1_subset_1(k1_zfmisc_1(X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2678, negated_conjecture, (esk4_1(X1)=X2|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|k2_xboole_0(X1,X3)!=k1_tarski(X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1643, plain, (X1=k2_xboole_0(X2,X3)|v1_xboole_0(X1)|~r2_hidden(esk7_3(X2,X3,X1),X2)|~m1_subset_1(esk7_3(X2,X3,X1),X1))).#trainneg
cnf(i_0_2808, plain, (r1_orders_2(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,k2_xboole_0(X2,X3)))|r1_orders_2(X1,k2_yellow_0(X1,X3),k2_yellow_0(X1,k2_xboole_0(X2,X3)))|v2_struct_0(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2)|~v5_waybel_6(k2_yellow_0(X1,k2_xboole_0(X2,X3)),X1)|~m1_subset_1(k2_yellow_0(X1,k2_xboole_0(X2,X3)),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2530, plain, (r2_hidden(esk72_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))),k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)))|~r2_hidden(X5,esk70_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))))|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X3,X4)),X2)|~r1_orders_2(X1,X5,X2)|~m1_subset_1(k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,X3))).#trainneg
cnf(i_0_2366, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,k11_lattice3(esk1_0,X1,X2),k11_lattice3(esk1_0,X1,X2))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2876, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3))),X2)|r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r2_yellow_0(X1,esk70_3(X1,X2,X3))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(esk69_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_2415, plain, (r1_tarski(esk70_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))),k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)))|r2_hidden(esk72_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))),k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)))|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X3,X4)),X2)|~m1_subset_1(k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,X3))).#trainneg
cnf(i_0_2485, plain, (r1_tarski(esk70_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))),k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)))|r2_hidden(esk72_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))),k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)))|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X4,X3)),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X4))).#trainneg
cnf(i_0_2491, plain, (m1_subset_1(esk70_3(X1,X2,k1_tarski(X3)),k1_zfmisc_1(k1_tarski(X3)))|r2_hidden(esk72_3(X1,X2,k1_tarski(X3)),k1_tarski(X3))|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2696, negated_conjecture, (esk4_1(X1)=X2|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|k2_xboole_0(X3,X1)!=k1_tarski(X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1672, plain, (esk6_3(X1,X2,X3)=X4|esk6_3(X1,X2,X3)=X1|esk6_3(X1,X2,X3)=X2|X3=k2_tarski(X1,X2)|X3!=k1_tarski(X4))).#trainneg
cnf(i_0_2376, plain, (m1_subset_1(X1,k2_tarski(X2,X1)))).#trainneg
cnf(i_0_2683, plain, (X1=X2|k2_xboole_0(k2_tarski(X1,X3),X4)!=k1_tarski(X2))).#trainneg
cnf(i_0_2700, plain, (X1=X2|k2_xboole_0(X3,k2_tarski(X4,X1))!=k1_tarski(X2))).#trainneg
cnf(i_0_2281, plain, (v1_card_1(k1_xboole_0))).#trainneg
cnf(i_0_3170, negated_conjecture, (r1_orders_2(esk1_0,esk8_2(esk1_0,esk2_0),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0))),esk2_0)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|esk9_2(esk1_0,esk2_0)!=esk8_2(esk1_0,esk2_0)|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3178, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk38_1(esk1_0)),esk38_1(esk1_0)),k2_yellow_0(esk1_0,esk38_1(esk1_0)))|~v1_xboole_0(esk38_1(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk38_1(esk1_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3570, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3))),X2)|r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(k1_tarski(esk69_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(esk70_3(X1,X2,X3))|~m1_subset_1(esk70_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_3521, negated_conjecture, (r1_tarski(esk4_1(k1_zfmisc_1(X1)),X1)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(X1)),esk2_0)|~v1_finset_1(k1_zfmisc_1(X1))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3389, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X1,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_3179, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk38_1(esk1_0)),esk38_1(esk1_0)),esk38_1(esk1_0))|v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk38_1(esk1_0)),esk38_1(esk1_0)),k2_yellow_0(esk1_0,esk38_1(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,esk38_1(esk1_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3184, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk38_1(esk1_0)),esk38_1(esk1_0)),k1_zfmisc_1(esk38_1(esk1_0)))|v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk38_1(esk1_0)),esk38_1(esk1_0)),k2_yellow_0(esk1_0,esk38_1(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,esk38_1(esk1_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3538, plain, (r1_orders_2(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X2))|~r2_yellow_0(X1,X2)|~v5_waybel_6(k2_yellow_0(X1,X2),X1)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3497, negated_conjecture, (v2_struct_0(esk1_0)|r2_hidden(esk69_3(esk1_0,X1,k1_tarski(X1)),k1_tarski(X1))|r2_hidden(esk72_3(esk1_0,X1,k1_tarski(X1)),k1_tarski(X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3227, plain, (esk69_3(X1,k2_yellow_0(X1,k2_tarski(X2,X3)),k2_tarski(X2,X3))=X3|esk69_3(X1,k2_yellow_0(X1,k2_tarski(X2,X3)),k2_tarski(X2,X3))=X2|r2_hidden(esk72_3(X1,k2_yellow_0(X1,k2_tarski(X2,X3)),k2_tarski(X2,X3)),k2_tarski(X2,X3))|~m1_subset_1(k2_tarski(X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3632, plain, (r1_tarski(esk70_3(X1,X2,k2_tarski(k2_yellow_0(X1,X4),k2_yellow_0(X1,X3))),k2_tarski(k2_yellow_0(X1,X4),k2_yellow_0(X1,X3)))|r2_hidden(esk72_3(X1,X2,k2_tarski(k2_yellow_0(X1,X4),k2_yellow_0(X1,X3))),k2_tarski(k2_yellow_0(X1,X4),k2_yellow_0(X1,X3)))|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X4)|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X4,X3)),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3349, plain, (v1_xboole_0(X1)|~m1_subset_1(X2,X1)|~m1_subset_1(k1_zfmisc_1(X1),X1))).#trainneg
cnf(i_0_3294, plain, (v1_zfmisc_1(X1)|~v1_zfmisc_1(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_3295, plain, (v6_ordinal1(X1)|~v6_ordinal1(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_3296, plain, (v1_relat_1(X1)|~v1_relat_1(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_3297, plain, (v5_finset_1(X1)|~v5_finset_1(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_3572, plain, (r1_orders_2(X1,k2_yellow_0(X1,X2),X3)|r1_orders_2(X1,X4,X3)|~r2_yellow_0(X1,X2)|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(k1_tarski(X4),X2)),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(k1_tarski(X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3568, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,X2),k12_lattice3(esk1_0,X1,X2))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3301, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_3664, plain, (v1_finset_1(X1)|~v5_finset_1(k2_tarski(X2,X1)))).#trainneg
cnf(i_0_3322, plain, (v3_ordinal1(X1)|~v3_ordinal1(k2_tarski(X2,X1)))).#trainneg
cnf(i_0_3298, plain, (v4_funct_1(X1)|~v4_funct_1(k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_3311, plain, (v1_zfmisc_1(k6_domain_1(k4_ordinal1,X1))|~v1_zfmisc_1(k1_numbers)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_3312, plain, (v6_ordinal1(k6_domain_1(k4_ordinal1,X1))|~v6_ordinal1(k1_numbers)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_3236, plain, (esk69_3(X1,k2_yellow_0(X1,X2),X2)=X3|r2_hidden(esk72_3(X1,k2_yellow_0(X1,X2),X2),X2)|v1_xboole_0(X2)|X2!=k1_tarski(X3)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3035, plain, (k2_yellow_0(X1,k2_tarski(X3,k2_yellow_0(X1,X2)))=k2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(X3)))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X2)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_2902, negated_conjecture, (esk4_1(k2_tarski(X1,k2_yellow_0(esk1_0,X2)))=k2_yellow_0(esk1_0,X2)|esk4_1(k2_tarski(X1,k2_yellow_0(esk1_0,X2)))=X1|v5_waybel_6(esk2_0,esk1_0)|v2_struct_0(esk1_0)|~r2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X1),X2)),esk2_0)|~m1_subset_1(k2_tarski(X1,k2_yellow_0(esk1_0,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2526, plain, (r2_hidden(esk72_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),k6_domain_1(u1_struct_0(X1),X3))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|v2_struct_0(X1)|~r2_hidden(X4,esk70_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)))|~r1_orders_2(X1,X3,X2)|~r1_orders_2(X1,X4,X2)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_2910, negated_conjecture, (esk4_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)))=k2_yellow_0(esk1_0,X2)|esk4_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)))=k2_yellow_0(esk1_0,X1)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X1,X2)),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2531, plain, (r2_hidden(esk72_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))),k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)))|~r2_hidden(X5,esk70_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))))|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X4,X3)),X2)|~r1_orders_2(X1,X5,X2)|~m1_subset_1(k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X4))).#trainneg
cnf(i_0_3493, negated_conjecture, (esk69_3(esk1_0,k2_yellow_0(esk1_0,X1),X1)=X2|v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|v1_xboole_0(X1)|X1!=k1_tarski(X2)|~v1_finset_1(X1)|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2498, plain, (r1_tarski(k7_domain_1(X1,X2,X3),X1)|v1_xboole_0(X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_1694, plain, (v6_ordinal1(esk73_1(X1))|v2_struct_0(X1)|v2_lattice3(X1)|~v6_ordinal1(u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3666, plain, (v1_funct_1(X1)|~v4_funct_1(k2_tarski(X2,X1)))).#trainneg
cnf(i_0_3313, plain, (v1_relat_1(k6_domain_1(k4_ordinal1,X1))|~v1_relat_1(k1_numbers)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_3314, plain, (v5_finset_1(k6_domain_1(k4_ordinal1,X1))|~v5_finset_1(k1_numbers)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_3315, plain, (v4_funct_1(k6_domain_1(k4_ordinal1,X1))|~v4_funct_1(k1_numbers)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_3321, plain, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_3667, plain, (v1_relat_1(X1)|~v4_funct_1(k2_tarski(X2,X1)))).#trainneg
cnf(i_0_2809, plain, (r1_orders_2(X1,X2,k2_yellow_0(X1,k2_tarski(X2,X3)))|r1_orders_2(X1,X3,k2_yellow_0(X1,k2_tarski(X2,X3)))|v2_struct_0(X1)|~v5_waybel_6(k2_yellow_0(X1,k2_tarski(X2,X3)),X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,k2_tarski(X2,X3)),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2893, negated_conjecture, (esk4_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)))=k2_yellow_0(esk1_0,X2)|esk4_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)))=k2_yellow_0(esk1_0,X1)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,X1)|~r2_yellow_0(esk1_0,X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X2,X1)),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3112, plain, (k2_yellow_0(X1,k2_tarski(X2,k2_yellow_0(X1,X3)))=k2_yellow_0(X1,k2_xboole_0(k1_tarski(X2),X3))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X3)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_2390, plain, (r1_tarski(esk70_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|m1_subset_1(esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),u1_struct_0(X1))|v5_waybel_6(X2,X1)|~m1_subset_1(k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_3509, plain, (r1_tarski(X1,X2)|r2_hidden(esk72_3(X3,k2_yellow_0(X3,X2),X2),X2)|v1_xboole_0(X2)|~r1_tarski(X1,esk70_3(X3,k2_yellow_0(X3,X2),X2))|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X3)))|~l1_orders_2(X3)|~v2_lattice3(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3))).#trainneg
cnf(i_0_3245, plain, (r2_hidden(esk69_3(X1,k2_yellow_0(X1,k2_xboole_0(X3,X2)),k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))),k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3)))|r2_hidden(esk72_3(X1,k2_yellow_0(X1,k2_xboole_0(X3,X2)),k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))),k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3)))|~m1_subset_1(k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,X3))).#trainneg
cnf(i_0_3703, negated_conjecture, (m1_subset_1(esk4_1(k1_zfmisc_1(X1)),k1_zfmisc_1(X1))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(X1)),esk2_0)|~v1_finset_1(k1_zfmisc_1(X1))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1695, plain, (v1_relat_1(esk73_1(X1))|v2_struct_0(X1)|v2_lattice3(X1)|~v1_relat_1(u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3702, negated_conjecture, (m1_subset_1(esk4_1(k1_zfmisc_1(k4_ordinal1)),k1_zfmisc_1(k1_numbers))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(k4_ordinal1)),esk2_0)|~v1_finset_1(k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(k1_zfmisc_1(k4_ordinal1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3684, negated_conjecture, (r1_orders_2(esk1_0,esk8_2(esk1_0,esk2_0),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|esk9_2(esk1_0,esk2_0)!=esk8_2(esk1_0,esk2_0)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3716, plain, (m1_subset_1(esk12_1(k1_zfmisc_1(k1_xboole_0)),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_3719, negated_conjecture, (m1_subset_1(esk4_1(k1_zfmisc_1(k1_xboole_0)),k1_zfmisc_1(X1))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(k1_xboole_0)),esk2_0)|~v1_finset_1(k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(k1_zfmisc_1(k1_xboole_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3747, plain, (esk72_3(X1,k2_yellow_0(X1,k2_tarski(X2,X3)),k2_tarski(X2,X3))=X3|esk72_3(X1,k2_yellow_0(X1,k2_tarski(X2,X3)),k2_tarski(X2,X3))=X2|esk69_3(X1,k2_yellow_0(X1,k2_tarski(X2,X3)),k2_tarski(X2,X3))=X2|esk69_3(X1,k2_yellow_0(X1,k2_tarski(X2,X3)),k2_tarski(X2,X3))=X3|~m1_subset_1(k2_tarski(X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3949, negated_conjecture, (esk4_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X1)))=k2_yellow_0(esk1_0,X1)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3872, plain, (esk69_3(X1,k2_yellow_0(X1,k1_tarski(X2)),k1_tarski(X2))=X2|r2_hidden(esk72_3(X1,k2_yellow_0(X1,k1_tarski(X2)),k1_tarski(X2)),k1_tarski(X2))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1696, plain, (v5_finset_1(esk73_1(X1))|v2_struct_0(X1)|v2_lattice3(X1)|~v5_finset_1(u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_3855, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(k2_tarski(X2,X1),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_3856, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(k2_tarski(X1,X2),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_3926, plain, (k2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X3),X2))=k2_yellow_0(X1,k2_xboole_0(X3,k1_tarski(X2)))|~r2_yellow_0(X1,X3)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_4102, plain, (k2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X3),X2))=k2_yellow_0(X1,k2_xboole_0(k1_tarski(X2),X3))|~r2_yellow_0(X1,X3)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1697, plain, (v4_funct_1(esk73_1(X1))|v2_struct_0(X1)|v2_lattice3(X1)|~v4_funct_1(u1_struct_0(X1))|~l1_orders_2(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2701, plain, (X1=X2|k2_xboole_0(X3,k2_tarski(X1,X4))!=k1_tarski(X2))).#trainneg
cnf(i_0_3302, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_4021, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k1_tarski(X2))|~m1_subset_1(X2,k4_ordinal1))).#trainneg
cnf(i_0_4022, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k1_tarski(X2))|~m1_subset_1(X2,k4_ordinal1))).#trainneg
cnf(i_0_3340, plain, (m1_subset_1(k2_tarski(X1,X2),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k4_ordinal1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_3989, plain, (m1_subset_1(k7_domain_1(k4_ordinal1,X1,X2),k1_zfmisc_1(k1_numbers))|~m1_subset_1(X2,k4_ordinal1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_2918, negated_conjecture, (r1_tarski(esk70_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|v1_xboole_0(X1)|v2_struct_0(esk1_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3079, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3))),X2)|r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|v1_xboole_0(X3)|~r2_yellow_0(X1,esk70_3(X1,X2,X3))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(esk69_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_3244, plain, (r2_hidden(esk69_3(X1,k2_yellow_0(X1,k2_xboole_0(X2,X3)),k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))),k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3)))|r2_hidden(esk72_3(X1,k2_yellow_0(X1,k2_xboole_0(X2,X3)),k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))),k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3)))|~m1_subset_1(k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2))).#trainneg
cnf(i_0_3851, negated_conjecture, (m1_subset_1(esk4_1(X1),k1_numbers)|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(k4_ordinal1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_3390, plain, (r1_tarski(X1,X2)|~r1_tarski(X1,X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_xboole_0)))).#trainneg
cnf(i_0_2916, negated_conjecture, (r1_orders_2(esk1_0,X1,k2_yellow_0(esk1_0,k2_tarski(X1,X2)))|r1_orders_2(esk1_0,X2,k2_yellow_0(esk1_0,k2_tarski(X1,X2)))|v2_struct_0(esk1_0)|~v5_waybel_6(k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk1_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(X1,X2)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_3082, plain, (r1_orders_2(X1,X2,k12_lattice3(X1,X3,X2))|r1_orders_2(X1,X3,k12_lattice3(X1,X3,X2))|~v5_waybel_6(k12_lattice3(X1,X3,X2),X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k12_lattice3(X1,X3,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2407, plain, (v2_struct_0(X1)|r1_orders_2(X1,k1_yellow_0(X1,X2),k1_yellow_0(X1,X2))|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_4562, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3))),X2)|r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|v1_xboole_0(X3)|v1_xboole_0(esk70_3(X1,X2,X3))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(k1_tarski(esk69_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(esk70_3(X1,X2,X3))|~m1_subset_1(esk70_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1700, plain, (m1_subset_1(X1,u1_struct_0(X2))|v2_struct_0(X2)|v2_lattice3(X2)|~r2_hidden(X1,esk73_1(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_4515, plain, (~m1_subset_1(X1,k4_ordinal1)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k4_ordinal1))).#trainneg
cnf(i_0_4220, negated_conjecture, (r1_orders_2(esk1_0,esk8_2(esk1_0,esk2_0),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|esk9_2(esk1_0,esk2_0)!=esk8_2(esk1_0,esk2_0)|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4264, plain, (esk69_3(X1,k2_yellow_0(X1,k2_tarski(X2,X3)),k2_tarski(X2,X3))=X3|esk69_3(X1,k2_yellow_0(X1,k2_tarski(X2,X3)),k2_tarski(X2,X3))=X2|esk72_3(X1,k2_yellow_0(X1,k2_tarski(X2,X3)),k2_tarski(X2,X3))=X2|esk72_3(X1,k2_yellow_0(X1,k2_tarski(X2,X3)),k2_tarski(X2,X3))=X3|v1_xboole_0(u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_4552, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k1_tarski(X1)),k1_tarski(X1))|v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,X1,k1_tarski(X1)),k1_tarski(X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1727, plain, (X4=X5|r2_hidden(esk7_3(X5,X5,X4),X5)|r2_hidden(esk7_3(X5,X5,X4),X4))).#trainneg
cnf(i_0_4469, plain, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X1,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_4495, plain, (r1_tarski(k1_tarski(X1),k1_numbers)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_4624, plain, (r2_hidden(esk72_3(X1,k2_yellow_0(X1,X2),k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X2))),k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X2)))|r2_hidden(esk69_3(X1,k2_yellow_0(X1,X2),k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X2))),k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X2)))|~r2_yellow_0(X1,X2)|~m1_subset_1(k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2613, plain, (esk69_3(X1,X2,k1_tarski(X3))=X4|r2_hidden(esk72_3(X1,X2,k1_tarski(X3)),k1_tarski(X3))|k1_tarski(X3)!=k1_tarski(X4)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2795, plain, (r2_hidden(esk72_3(X1,X2,k2_tarski(X3,X4)),k2_tarski(X3,X4))|r2_hidden(esk69_3(X1,X2,k2_tarski(X3,X4)),k2_tarski(X3,X4))|~r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(X3,X4)),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_4630, negated_conjecture, (esk4_1(X1)=X2|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|k1_tarski(X2)!=k1_numbers|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_1728, plain, (k2_xboole_0(X6,X7)=X6|r2_hidden(esk7_3(X6,X7,X6),X7)|r2_hidden(esk7_3(X6,X7,X6),X6))).#trainneg
cnf(i_0_4486, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X2,k4_ordinal1)|~m1_subset_1(X1,k1_tarski(X2)))).#trainneg
cnf(i_0_4470, plain, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X2,k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_4496, plain, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X2,k4_ordinal1)|~m1_subset_1(X1,k1_zfmisc_1(k1_tarski(X2))))).#trainneg
cnf(i_0_4631, negated_conjecture, (v5_waybel_6(esk2_0,esk1_0)|m1_subset_1(X1,k1_numbers)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X1)),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_4485, negated_conjecture, (m1_subset_1(esk4_1(k1_tarski(X1)),k1_numbers)|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X1,k4_ordinal1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X1)),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2398, plain, (esk12_1(k2_tarski(X1,X2))=X1|esk12_1(k2_tarski(X1,X2))=X2)).#trainneg
cnf(i_0_3088, plain, (r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(X3,X2)),k2_yellow_0(X1,k2_tarski(X3,X2)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v3_orders_2(X1)|~v4_orders_2(X1))).#trainneg
cnf(i_0_2877, negated_conjecture, (esk4_1(k2_tarski(k2_yellow_0(esk1_0,X1),X2))=X2|esk4_1(k2_tarski(k2_yellow_0(esk1_0,X1),X2))=k2_yellow_0(esk1_0,X1)|v5_waybel_6(esk2_0,esk1_0)|v2_struct_0(esk1_0)|~r2_yellow_0(esk1_0,X1)|~r2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X2))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k6_domain_1(u1_struct_0(esk1_0),X2),X1)),esk2_0)|~m1_subset_1(k2_tarski(k2_yellow_0(esk1_0,X1),X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2425, plain, (k2_yellow_0(X1,k2_xboole_0(X2,k6_domain_1(u1_struct_0(X1),X3)))=k2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X2),X3))|v2_struct_0(X1)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X3))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_1946, negated_conjecture, (v1_xboole_0(esk4_1(k1_zfmisc_1(X1)))|m1_subset_1(X2,X1)|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X2,esk4_1(k1_zfmisc_1(X1)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(X1)),esk2_0)|~v1_finset_1(k1_zfmisc_1(X1))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4512, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k2_tarski(X2,X3))|~m1_subset_1(X3,k4_ordinal1)|~m1_subset_1(X2,k4_ordinal1))).#trainneg
cnf(i_0_4513, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k2_tarski(X2,X3))|~m1_subset_1(X3,k4_ordinal1)|~m1_subset_1(X2,k4_ordinal1))).#trainneg
cnf(i_0_4533, plain, (m1_subset_1(X1,k1_numbers)|~r2_hidden(X1,k7_domain_1(k4_ordinal1,X2,X3))|~m1_subset_1(X3,k4_ordinal1)|~m1_subset_1(X2,k4_ordinal1))).#trainneg
cnf(i_0_4534, plain, (r1_tarski(X1,k1_numbers)|~r1_tarski(X1,k7_domain_1(k4_ordinal1,X2,X3))|~m1_subset_1(X3,k4_ordinal1)|~m1_subset_1(X2,k4_ordinal1))).#trainneg
cnf(i_0_4505, plain, (v1_zfmisc_1(k2_tarski(X1,X2))|~v1_zfmisc_1(k1_numbers)|~m1_subset_1(X2,k4_ordinal1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_4506, plain, (v6_ordinal1(k2_tarski(X1,X2))|~v6_ordinal1(k1_numbers)|~m1_subset_1(X2,k4_ordinal1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_3511, plain, (r1_tarski(esk70_3(X1,X2,k1_tarski(X2)),k1_tarski(X2))|r2_hidden(esk72_3(X1,X2,k1_tarski(X2)),k1_tarski(X2))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_4299, plain, (esk69_3(X1,X2,k1_tarski(X2))=X2|r2_hidden(esk72_3(X1,X2,k1_tarski(X2)),k1_tarski(X2))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2541, plain, (r2_hidden(esk72_3(X1,X2,k2_tarski(X3,X4)),k2_tarski(X3,X4))|r2_hidden(esk69_3(X1,X2,k2_tarski(X3,X4)),k2_tarski(X3,X4))|~r1_orders_2(X1,k12_lattice3(X1,X4,X3),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2404, negated_conjecture, (v2_struct_0(X1)|r1_orders_2(X1,esk4_1(u1_struct_0(X1)),esk4_1(u1_struct_0(X1)))|v1_xboole_0(u1_struct_0(X1))|v5_waybel_6(esk2_0,esk1_0)|~l1_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,u1_struct_0(X1)),esk2_0)|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(u1_struct_0(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_4507, plain, (v1_relat_1(k2_tarski(X1,X2))|~v1_relat_1(k1_numbers)|~m1_subset_1(X2,k4_ordinal1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_4508, plain, (v5_finset_1(k2_tarski(X1,X2))|~v5_finset_1(k1_numbers)|~m1_subset_1(X2,k4_ordinal1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_4509, plain, (v4_funct_1(k2_tarski(X1,X2))|~v4_funct_1(k1_numbers)|~m1_subset_1(X2,k4_ordinal1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_2995, plain, (esk69_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk9_2(X1,X2)|esk69_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|r2_hidden(esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|v5_waybel_6(X2,X1)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2471, plain, (k2_yellow_0(X1,k2_xboole_0(k6_domain_1(u1_struct_0(X1),X3),X2))=k2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X2),X3))|v2_struct_0(X1)|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X3))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_3085, plain, (r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X3,X2)),k2_yellow_0(X1,k2_xboole_0(X3,X2)))|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,X3)|~v4_orders_2(X1))).#trainneg
cnf(i_0_2986, plain, (v2_struct_0(X1)|r1_orders_2(X1,X2,k11_lattice3(X1,X2,X3))|r1_orders_2(X1,X3,k11_lattice3(X1,X2,X3))|~v5_waybel_6(k11_lattice3(X1,X2,X3),X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k11_lattice3(X1,X2,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1871, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k12_lattice3(X1,X3,X4))=k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X3,X4),X2))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_3697, negated_conjecture, (r1_tarski(esk4_1(k1_zfmisc_1(esk4_1(k1_zfmisc_1(X1)))),X1)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(X1)),esk2_0)|~v1_finset_1(k1_zfmisc_1(X1))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(esk4_1(k1_zfmisc_1(X1)))),esk2_0)|~v1_finset_1(k1_zfmisc_1(esk4_1(k1_zfmisc_1(X1))))|~m1_subset_1(k1_zfmisc_1(esk4_1(k1_zfmisc_1(X1))),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2580, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X1)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,k1_tarski(X3)))).#trainneg
cnf(i_0_4654, plain, (r1_tarski(X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(X1,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_4626, negated_conjecture, (v1_finset_1(esk4_1(X1))|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|~v5_finset_1(k1_numbers)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_2430, plain, (k2_yellow_0(X1,k2_tarski(X2,k2_yellow_0(X1,X3)))=k2_yellow_0(X1,k2_xboole_0(k1_tarski(X2),X3))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,k1_tarski(X2))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_4695, negated_conjecture, (v2_struct_0(esk1_0)|v5_waybel_6(esk2_0,esk1_0)|esk9_2(esk1_0,esk2_0)!=esk8_2(esk1_0,esk2_0)|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5198, plain, (r1_orders_2(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X2))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1734, plain, (m1_subset_1(esk7_3(X1,X2,X3),X3)|X3=k2_xboole_0(X1,X2)|r2_hidden(esk7_3(X1,X2,X3),X1)|r2_hidden(esk7_3(X1,X2,X3),X2))).#trainneg
cnf(i_0_5005, plain, (esk12_1(k2_tarski(X1,X1))=X1)).#trainneg
cnf(i_0_5294, negated_conjecture, (v1_finset_1(esk8_2(esk1_0,esk2_0))|v5_waybel_6(esk2_0,esk1_0)|~v5_finset_1(k1_numbers)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0))),esk2_0)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(k4_ordinal1))|esk9_2(esk1_0,esk2_0)!=esk8_2(esk1_0,esk2_0)|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5296, negated_conjecture, (v1_finset_1(esk8_2(esk1_0,esk2_0))|v5_waybel_6(esk2_0,esk1_0)|r1_orders_2(esk1_0,esk9_2(esk1_0,esk2_0),esk2_0)|~v5_finset_1(k1_numbers)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0))),esk2_0)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5295, negated_conjecture, (v1_finset_1(esk9_2(esk1_0,esk2_0))|v5_waybel_6(esk2_0,esk1_0)|esk4_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)))=esk8_2(esk1_0,esk2_0)|~v5_finset_1(k1_numbers)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0))),esk2_0)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4706, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,esk2_0)),k2_tarski(X1,esk2_0))=esk2_0|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,esk2_0)),k2_tarski(X1,esk2_0))=X1|esk69_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,esk2_0)),k2_tarski(X1,esk2_0))=X1|esk69_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,esk2_0)),k2_tarski(X1,esk2_0))=esk2_0|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5182, plain, (esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk9_2(X1,X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|esk69_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|esk69_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk9_2(X1,X2)|v5_waybel_6(X2,X1)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5015, plain, (r2_hidden(esk72_3(X1,k2_yellow_0(X1,k2_tarski(X2,X3)),k2_tarski(X2,X3)),k2_tarski(X2,X3))|r2_hidden(esk69_3(X1,k2_yellow_0(X1,k2_tarski(X2,X3)),k2_tarski(X2,X3)),k2_tarski(X2,X3))|~m1_subset_1(k2_tarski(X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k2_yellow_0(X1,k2_tarski(X2,X3)),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_5016, plain, (r1_tarski(esk70_3(X1,k2_yellow_0(X1,k2_tarski(X2,X3)),k2_tarski(X2,X3)),k2_tarski(X2,X3))|r2_hidden(esk72_3(X1,k2_yellow_0(X1,k2_tarski(X2,X3)),k2_tarski(X2,X3)),k2_tarski(X2,X3))|~m1_subset_1(k2_tarski(X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k2_yellow_0(X1,k2_tarski(X2,X3)),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_5385, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,X4,X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,k1_tarski(esk69_3(X1,X2,X3))))).#trainneg
cnf(i_0_3157, negated_conjecture, (v5_waybel_6(esk2_0,esk1_0)|m1_subset_1(esk4_1(esk4_1(k1_zfmisc_1(X1))),X1)|v1_xboole_0(esk4_1(k1_zfmisc_1(X1)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(X1)),esk2_0)|~v1_finset_1(k1_zfmisc_1(X1))|~m1_subset_1(k1_zfmisc_1(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk4_1(k1_zfmisc_1(X1))),esk2_0)|~v1_finset_1(esk4_1(k1_zfmisc_1(X1)))|~m1_subset_1(esk4_1(k1_zfmisc_1(X1)),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1736, plain, (esk7_3(X1,X2,X3)=X4|X3=k2_xboole_0(X1,X2)|r2_hidden(esk7_3(X1,X2,X3),X1)|r2_hidden(esk7_3(X1,X2,X3),X2)|X3!=k1_tarski(X4))).#trainneg
cnf(i_0_4992, plain, (r1_tarski(k1_tarski(X1),k1_numbers)|~m1_subset_1(X2,k4_ordinal1)|~m1_subset_1(X1,k1_tarski(X2)))).#trainneg
cnf(i_0_5293, negated_conjecture, (v1_finset_1(X1)|v5_waybel_6(esk2_0,esk1_0)|~v5_finset_1(k1_numbers)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X1)),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_5041, plain, (k2_yellow_0(X1,k2_xboole_0(X2,k6_domain_1(u1_struct_0(X1),X3)))=k2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X2),X3))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_2126, plain, (v1_zfmisc_1(k2_tarski(X1,X2))|v1_xboole_0(X3)|~v1_zfmisc_1(X3)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_1740, plain, (m1_subset_1(esk7_3(X1,X2,X3),X2)|X3=k2_xboole_0(X1,X2)|r2_hidden(esk7_3(X1,X2,X3),X1)|r2_hidden(esk7_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_2030, plain, (m1_subset_1(X1,k2_xboole_0(X2,X3))|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_2684, plain, (X1=X2|v1_xboole_0(X3)|k2_xboole_0(X3,X4)!=k1_tarski(X2)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_4997, negated_conjecture, (v5_waybel_6(esk2_0,esk1_0)|m1_subset_1(X1,k1_numbers)|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4988, negated_conjecture, (r1_tarski(esk4_1(k1_zfmisc_1(k1_tarski(X1))),k1_numbers)|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X1,k4_ordinal1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(k1_tarski(X1))),esk2_0)|~v1_finset_1(k1_zfmisc_1(k1_tarski(X1)))|~m1_subset_1(k1_zfmisc_1(k1_tarski(X1)),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5071, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(X3,k4_ordinal1)|~m1_subset_1(X2,k4_ordinal1)|~m1_subset_1(X1,k2_tarski(X2,X3)))).#trainneg
cnf(i_0_5138, plain, (esk72_3(X1,X2,k1_tarski(X2))=X3|esk69_3(X1,X2,k1_tarski(X2))=X2|k1_tarski(X2)!=k1_tarski(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5185, plain, (k2_yellow_0(X1,k2_xboole_0(k6_domain_1(u1_struct_0(X1),X2),X3))=k2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X3),X2))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X2))|~r2_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_4860, plain, (esk69_3(X1,X2,k1_tarski(X3))=X3|r2_hidden(esk72_3(X1,X2,k1_tarski(X3)),k1_tarski(X3))|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5233, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,k2_tarski(X3,X4)))=k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X3,X4),X2))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))|~r2_yellow_0(X1,X2)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2127, plain, (v6_ordinal1(k2_tarski(X1,X2))|v1_xboole_0(X3)|~v6_ordinal1(X3)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_1742, plain, (esk7_3(X1,X2,X3)=X4|X3=k2_xboole_0(X1,X2)|r2_hidden(esk7_3(X1,X2,X3),X1)|r2_hidden(esk7_3(X1,X2,X3),X3)|X2!=k1_tarski(X4))).#trainneg
cnf(i_0_2702, plain, (X1=X2|v1_xboole_0(X4)|k2_xboole_0(X3,X4)!=k1_tarski(X2)|~m1_subset_1(X1,X4))).#trainneg
cnf(i_0_2969, plain, (r1_tarski(k1_tarski(X2),X1)|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_5088, plain, (r1_tarski(X1,k1_numbers)|~m1_subset_1(X3,k4_ordinal1)|~m1_subset_1(X2,k4_ordinal1)|~m1_subset_1(X1,k1_zfmisc_1(k2_tarski(X2,X3))))).#trainneg
cnf(i_0_5184, plain, (k2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X3),X2))=k2_yellow_0(X1,k2_xboole_0(k1_tarski(X2),X3))|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,k1_tarski(X2))|~r2_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5363, plain, (k2_yellow_0(X1,k2_xboole_0(k1_tarski(X3),k6_domain_1(u1_struct_0(X1),X2)))=k12_lattice3(X1,X3,X2)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5991, plain, (esk69_3(X1,X2,k1_tarski(X2))=X2|esk72_3(X1,X2,k1_tarski(X2))=X2|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5390, plain, (r1_orders_2(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X2))|v1_xboole_0(X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1744, plain, (m1_subset_1(esk7_3(X1,X2,X3),X1)|X3=k2_xboole_0(X1,X2)|r2_hidden(esk7_3(X1,X2,X3),X2)|r2_hidden(esk7_3(X1,X2,X3),X3))).#trainneg
cnf(i_0_5499, negated_conjecture, (v1_finset_1(esk8_2(esk1_0,esk2_0))|v5_waybel_6(esk2_0,esk1_0)|esk9_2(esk1_0,esk2_0)!=esk8_2(esk1_0,esk2_0)|~v5_finset_1(k1_numbers)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5500, negated_conjecture, (r1_orders_2(esk1_0,esk9_2(esk1_0,esk2_0),esk2_0)|v1_finset_1(esk8_2(esk1_0,esk2_0))|v5_waybel_6(esk2_0,esk1_0)|~v5_finset_1(k1_numbers)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5506, negated_conjecture, (v1_finset_1(esk8_2(esk1_0,esk2_0))|v5_waybel_6(esk2_0,esk1_0)|v1_finset_1(esk9_2(esk1_0,esk2_0))|~v5_finset_1(k1_numbers)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0))),esk2_0)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5503, negated_conjecture, (r1_orders_2(esk1_0,esk8_2(esk1_0,esk2_0),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|v1_finset_1(esk9_2(esk1_0,esk2_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0))),esk2_0)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v5_finset_1(k1_numbers)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5539, plain, (r2_hidden(esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|esk69_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk9_2(X1,X2)|v5_waybel_6(X2,X1)|~r1_orders_2(X1,esk9_2(X1,X2),X2)|~r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_6349, plain, (k2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,k1_tarski(X2)),X3))=k12_lattice3(X1,X2,X3)|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|~r2_yellow_0(X1,k1_tarski(X2))|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2128, plain, (v1_relat_1(k2_tarski(X1,X2))|v1_xboole_0(X3)|~v1_relat_1(X3)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_5881, plain, (v1_finset_1(X1)|~v5_finset_1(k2_xboole_0(X2,X3))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_1746, plain, (esk7_3(X1,X2,X3)=X4|X3=k2_xboole_0(X1,X2)|r2_hidden(esk7_3(X1,X2,X3),X2)|r2_hidden(esk7_3(X1,X2,X3),X3)|X1!=k1_tarski(X4))).#trainneg
cnf(i_0_5947, negated_conjecture, (esk2_0=X1|k2_xboole_0(u1_struct_0(esk1_0),X2)!=k1_tarski(X1))).#trainneg
cnf(i_0_6214, negated_conjecture, (esk2_0=X1|k2_xboole_0(X2,u1_struct_0(esk1_0))!=k1_tarski(X1))).#trainneg
cnf(i_0_5960, plain, (k1_xboole_0=X1|k2_xboole_0(k1_zfmisc_1(X2),X3)!=k1_tarski(X1))).#trainneg
cnf(i_0_5518, negated_conjecture, (esk69_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0)).#trainneg
cnf(i_0_5508, negated_conjecture, (esk69_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk61_1(esk1_0),esk2_0)),k2_tarski(esk61_1(esk1_0),esk2_0))=esk2_0|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk61_1(esk1_0),esk2_0)),k2_tarski(esk61_1(esk1_0),esk2_0))=esk2_0|esk69_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk61_1(esk1_0),esk2_0)),k2_tarski(esk61_1(esk1_0),esk2_0))=esk61_1(esk1_0)|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk61_1(esk1_0),esk2_0)),k2_tarski(esk61_1(esk1_0),esk2_0))=esk61_1(esk1_0)|v2_struct_0(esk1_0)|~v2_yellow_0(esk1_0))).#trainneg
cnf(i_0_5516, negated_conjecture, (esk69_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk4_1(u1_struct_0(esk1_0)),esk2_0)),k2_tarski(esk4_1(u1_struct_0(esk1_0)),esk2_0))=esk2_0|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk4_1(u1_struct_0(esk1_0)),esk2_0)),k2_tarski(esk4_1(u1_struct_0(esk1_0)),esk2_0))=esk2_0|esk69_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk4_1(u1_struct_0(esk1_0)),esk2_0)),k2_tarski(esk4_1(u1_struct_0(esk1_0)),esk2_0))=esk4_1(u1_struct_0(esk1_0))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk4_1(u1_struct_0(esk1_0)),esk2_0)),k2_tarski(esk4_1(u1_struct_0(esk1_0)),esk2_0))=esk4_1(u1_struct_0(esk1_0))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,u1_struct_0(esk1_0)),esk2_0)|~v1_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6350, plain, (r1_orders_2(X1,k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2)),X3)|r1_orders_2(X1,X4,X3)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~r1_orders_2(X1,k12_lattice3(X1,X4,X2),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(k1_tarski(X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2129, plain, (v5_finset_1(k2_tarski(X1,X2))|v1_xboole_0(X3)|~v5_finset_1(X3)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_1769, plain, (r2_hidden(esk69_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k7_domain_1(u1_struct_0(X1),X3,X4))|v1_xboole_0(k7_domain_1(u1_struct_0(X1),X3,X4))|m1_subset_1(esk72_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),u1_struct_0(X1))|~r1_orders_2(X1,k12_lattice3(X1,X3,X4),X2)|~v1_finset_1(k7_domain_1(u1_struct_0(X1),X3,X4))|~m1_subset_1(k7_domain_1(u1_struct_0(X1),X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_5961, plain, (k4_ordinal1=X1|k2_xboole_0(k1_zfmisc_1(k1_numbers),X2)!=k1_tarski(X1))).#trainneg
cnf(i_0_5972, plain, (X1=X2|k2_xboole_0(k1_zfmisc_1(X1),X3)!=k1_tarski(X2))).#trainneg
cnf(i_0_6227, plain, (k1_xboole_0=X1|k2_xboole_0(X3,k1_zfmisc_1(X2))!=k1_tarski(X1))).#trainneg
cnf(i_0_5743, negated_conjecture, (v1_zfmisc_1(k2_tarski(X1,esk2_0))|~v1_zfmisc_1(u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6152, negated_conjecture, (v6_ordinal1(k2_tarski(X1,esk2_0))|~v6_ordinal1(u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6228, plain, (k4_ordinal1=X1|k2_xboole_0(X2,k1_zfmisc_1(k1_numbers))!=k1_tarski(X1))).#trainneg
cnf(i_0_6347, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X3)|r1_orders_2(X1,k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X4)),X3)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X4))|~r2_yellow_0(X1,k1_tarski(X2))|~r1_orders_2(X1,k12_lattice3(X1,X2,X4),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_6348, plain, (k12_lattice3(X1,k2_yellow_0(X1,k1_tarski(X2)),X3)=k12_lattice3(X1,X2,X3)|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|~r2_yellow_0(X1,k1_tarski(X2))|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2130, plain, (v4_funct_1(k2_tarski(X1,X2))|v1_xboole_0(X3)|~v4_funct_1(X3)|~m1_subset_1(X2,X3)|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_5900, plain, (v1_finset_1(X1)|~v5_finset_1(k2_xboole_0(X2,X3))|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_1771, plain, (r1_tarski(esk70_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k7_domain_1(u1_struct_0(X1),X3,X4))|v1_xboole_0(k7_domain_1(u1_struct_0(X1),X3,X4))|m1_subset_1(esk72_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),u1_struct_0(X1))|~r1_orders_2(X1,k12_lattice3(X1,X3,X4),X2)|~v1_finset_1(k7_domain_1(u1_struct_0(X1),X3,X4))|~m1_subset_1(k7_domain_1(u1_struct_0(X1),X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_6239, plain, (X1=X2|k2_xboole_0(X3,k1_zfmisc_1(X1))!=k1_tarski(X2))).#trainneg
cnf(i_0_5973, negated_conjecture, (v2_struct_0(esk1_0)|v5_waybel_6(esk2_0,esk1_0)|m1_subset_1(X1,k1_numbers)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_5882, plain, (v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X2,X3))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_5883, plain, (v1_funct_1(X1)|~v4_funct_1(k2_xboole_0(X2,X3))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_5040, plain, (k2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X2),X3))=k2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(X3)))|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,k1_tarski(X3))|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_5536, plain, (r2_hidden(esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|esk69_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk9_2(X1,X2)|v5_waybel_6(X2,X1)|~r2_hidden(X3,k1_tarski(esk9_2(X1,X2)))|~r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_5192, plain, (r1_orders_2(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(X3))))|r1_orders_2(X1,X3,k2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(X3))))|~r2_yellow_0(X1,X2)|~v5_waybel_6(k2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(X3))),X1)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(X3))),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k1_tarski(X3)))).#trainneg
cnf(i_0_7006, plain, (v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_6696, plain, (v1_finset_1(X1)|~v5_finset_1(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_1791, plain, (r2_hidden(esk69_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k7_domain_1(u1_struct_0(X1),X3,X4))|r1_orders_2(X1,esk72_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),X2)|v1_xboole_0(k7_domain_1(u1_struct_0(X1),X3,X4))|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4)),X2)|~v1_finset_1(k7_domain_1(u1_struct_0(X1),X3,X4))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_6703, negated_conjecture, (esk2_0=X1|k1_tarski(X1)!=u1_struct_0(esk1_0))).#trainneg
cnf(i_0_6533, negated_conjecture, (v1_finset_1(esk8_2(esk1_0,esk2_0))|v5_waybel_6(esk2_0,esk1_0)|esk9_2(esk1_0,esk2_0)!=esk8_2(esk1_0,esk2_0)|~v5_finset_1(k1_numbers)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6534, negated_conjecture, (r1_orders_2(esk1_0,esk9_2(esk1_0,esk2_0),esk2_0)|v1_finset_1(esk8_2(esk1_0,esk2_0))|v5_waybel_6(esk2_0,esk1_0)|~v5_finset_1(k1_numbers)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6535, negated_conjecture, (v1_finset_1(esk9_2(esk1_0,esk2_0))|v1_finset_1(esk8_2(esk1_0,esk2_0))|v5_waybel_6(esk2_0,esk1_0)|~v5_finset_1(k1_numbers)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6536, negated_conjecture, (r1_orders_2(esk1_0,esk8_2(esk1_0,esk2_0),esk2_0)|v1_finset_1(esk9_2(esk1_0,esk2_0))|v5_waybel_6(esk2_0,esk1_0)|~v5_finset_1(k1_numbers)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6388, plain, (r2_hidden(esk69_3(X1,k2_yellow_0(X1,X2),X2),X2)|v1_xboole_0(X2)|m1_subset_1(esk72_3(X1,k2_yellow_0(X1,X2),X2),u1_struct_0(X1))|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6373, plain, (r1_orders_2(X1,esk72_3(X1,X2,k1_tarski(X2)),X2)|esk72_3(X1,X2,k1_tarski(X2))=X2|~r1_orders_2(X1,X2,X2)|~r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6732, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r1_orders_2(esk1_0,esk2_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6389, plain, (r1_tarski(esk70_3(X1,k2_yellow_0(X1,X2),X2),X2)|v1_xboole_0(X2)|m1_subset_1(esk72_3(X1,k2_yellow_0(X1,X2),X2),u1_struct_0(X1))|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1811, plain, (r1_tarski(esk70_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k7_domain_1(u1_struct_0(X1),X3,X4))|r1_orders_2(X1,esk72_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),X2)|v1_xboole_0(k7_domain_1(u1_struct_0(X1),X3,X4))|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4)),X2)|~v1_finset_1(k7_domain_1(u1_struct_0(X1),X3,X4))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_6870, plain, (k4_ordinal1=X1|k1_tarski(X1)!=k1_zfmisc_1(k1_numbers))).#trainneg
cnf(i_0_6730, negated_conjecture, (r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r1_orders_2(esk1_0,esk2_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6734, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~v1_xboole_0(k2_xboole_0(esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k1_tarski(esk2_0)))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6719, negated_conjecture, (r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~v1_xboole_0(k2_xboole_0(esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k1_tarski(esk2_0)))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6720, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~v1_xboole_0(k2_xboole_0(esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k1_tarski(esk2_0)))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6742, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk61_1(esk1_0)|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk2_0|esk69_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk2_0|v2_struct_0(esk1_0)|esk61_1(esk1_0)!=esk2_0|~v2_yellow_0(esk1_0))).#trainneg
cnf(i_0_6371, plain, (r2_hidden(esk72_3(X1,X2,k1_tarski(X2)),k1_tarski(X2))|esk72_3(X1,X2,k1_tarski(X2))=X2|~r1_orders_2(X1,X2,X2)|~r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6715, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r2_hidden(X1,k2_xboole_0(esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k1_tarski(esk2_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~r1_orders_2(esk1_0,X1,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6356, plain, (r1_orders_2(X1,esk72_3(X1,X2,k1_tarski(X2)),X2)|esk72_3(X1,X2,k1_tarski(X2))=X2|~r2_hidden(X3,k2_xboole_0(esk70_3(X1,X2,k1_tarski(X2)),k1_tarski(X2)))|~r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X2)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6721, negated_conjecture, (r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r2_hidden(X1,k2_xboole_0(esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k1_tarski(esk2_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~r1_orders_2(esk1_0,X1,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6359, plain, (r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(esk70_3(X1,X2,k1_tarski(X2)),k1_tarski(X2))),X2)|r1_orders_2(X1,esk72_3(X1,X2,k1_tarski(X2)),X2)|esk72_3(X1,X2,k1_tarski(X2))=X2|~r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X2)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6367, plain, (r1_orders_2(X1,esk72_3(X1,X2,k1_tarski(X2)),X2)|esk72_3(X1,X2,k1_tarski(X2))=X2|~r2_hidden(X3,k1_tarski(X2))|~r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X2)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6552, plain, (r2_hidden(esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk9_2(X1,X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|v5_waybel_6(X2,X1)|~r1_orders_2(X1,esk8_2(X1,X2),X2)|~r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,esk9_2(X1,X2),X2)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_7014, plain, (v1_funct_1(X1)|~v4_funct_1(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_1817, negated_conjecture, (r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,esk4_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk4_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),esk2_0)|~v1_finset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~m1_subset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_6711, plain, (k1_xboole_0=X1|k1_zfmisc_1(X2)!=k1_tarski(X1))).#trainneg
cnf(i_0_6717, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k1_tarski(esk2_0))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6716, negated_conjecture, (r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k1_tarski(esk2_0))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6766, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0))))=esk4_1(u1_struct_0(esk1_0))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0))))=esk2_0|esk69_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0))))=esk2_0|v5_waybel_6(esk2_0,esk1_0)|esk4_1(u1_struct_0(esk1_0))!=esk2_0|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,u1_struct_0(esk1_0)),esk2_0)|~v1_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6718, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k1_tarski(esk2_0))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6722, negated_conjecture, (r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r2_yellow_0(esk1_0,k1_tarski(esk2_0))|~r2_yellow_0(esk1_0,esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~v5_waybel_6(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),esk1_0)|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7118, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X3)|r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X4)),X3)|~r2_yellow_0(X1,k1_tarski(X4))|~r2_yellow_0(X1,k1_tarski(X2))|~r1_orders_2(X1,k12_lattice3(X1,X4,X2),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_6537, plain, (r1_orders_2(X1,esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk9_2(X1,X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|v5_waybel_6(X2,X1)|~r2_hidden(X3,k2_xboole_0(esk70_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k1_tarski(esk8_2(X1,X2))))|~r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,esk9_2(X1,X2),X2)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_6548, plain, (r1_orders_2(X1,esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk9_2(X1,X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|v5_waybel_6(X2,X1)|~r2_hidden(X3,k1_tarski(esk8_2(X1,X2)))|~r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,esk9_2(X1,X2),X2)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_6790, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X3)|r1_orders_2(X1,X4,X3)|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,k1_tarski(X2))|~r1_orders_2(X1,k12_lattice3(X1,X4,X2),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(k1_tarski(X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6390, plain, (r2_hidden(esk72_3(X1,k2_yellow_0(X1,X2),X2),X2)|v1_xboole_0(X2)|~r2_hidden(X3,esk70_3(X1,k2_yellow_0(X1,X2),X2))|~r1_orders_2(X1,X3,k2_yellow_0(X1,X2))|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1824, negated_conjecture, (v1_xboole_0(X1)|m1_subset_1(esk72_3(X2,X3,X1),u1_struct_0(X2))|v1_xboole_0(k2_xboole_0(esk70_3(X2,X3,X1),k1_tarski(esk69_3(X2,X3,X1))))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(X2,k2_yellow_0(X2,X1),X3)|~r1_orders_2(X2,esk4_1(k2_xboole_0(esk70_3(X2,X3,X1),k1_tarski(esk69_3(X2,X3,X1)))),X3)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(esk4_1(k2_xboole_0(esk70_3(X2,X3,X1),k1_tarski(esk69_3(X2,X3,X1)))),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(esk70_3(X2,X3,X1),k1_tarski(esk69_3(X2,X3,X1)))),esk2_0)|~v1_finset_1(k2_xboole_0(esk70_3(X2,X3,X1),k1_tarski(esk69_3(X2,X3,X1))))|~m1_subset_1(k2_xboole_0(esk70_3(X2,X3,X1),k1_tarski(esk69_3(X2,X3,X1))),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_6723, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r2_yellow_0(esk1_0,k1_tarski(esk2_0))|~r2_yellow_0(esk1_0,esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~v5_waybel_6(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),esk1_0)|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6758, negated_conjecture, (r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk61_1(esk1_0)|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk2_0|esk69_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk2_0|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk61_1(esk1_0),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))))|~m1_subset_1(esk61_1(esk1_0),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk61_1(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),u1_struct_0(esk1_0))|~v2_yellow_0(esk1_0))).#trainneg
cnf(i_0_6771, negated_conjecture, (r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0))))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0))))=esk4_1(u1_struct_0(esk1_0))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0))))=esk2_0|esk69_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0))))=esk2_0|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))))|~v1_xboole_0(k2_xboole_0(esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k1_tarski(esk4_1(u1_struct_0(esk1_0)))))|~m1_subset_1(k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,u1_struct_0(esk1_0)),esk2_0)|~v1_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6745, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0))),k1_tarski(esk61_1(esk1_0)))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))))|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0))),u1_struct_0(esk1_0))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk61_1(esk1_0)|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk2_0|esk69_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk2_0|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))))|~m1_subset_1(k2_tarski(esk2_0,esk61_1(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),u1_struct_0(esk1_0))|~v2_yellow_0(esk1_0))).#trainneg
cnf(i_0_6553, plain, (r1_orders_2(X1,esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk9_2(X1,X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|v5_waybel_6(X2,X1)|~r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,k1_tarski(esk8_2(X1,X2)))|~m1_subset_1(k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,esk9_2(X1,X2),X2)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_6546, plain, (r2_hidden(esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk9_2(X1,X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|v5_waybel_6(X2,X1)|~r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,k2_xboole_0(esk70_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k1_tarski(esk8_2(X1,X2))))|~m1_subset_1(k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,esk9_2(X1,X2),X2)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_6724, negated_conjecture, (r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~r1_orders_2(esk1_0,X1,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(X1,k2_xboole_0(esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k1_tarski(esk2_0)))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7121, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X3)|r1_orders_2(X1,X4,X3)|~r2_yellow_0(X1,k1_tarski(X2))|~r1_orders_2(X1,k12_lattice3(X1,X2,X4),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(k1_tarski(X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k1_tarski(X4))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_6543, plain, (r2_hidden(esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk9_2(X1,X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|v5_waybel_6(X2,X1)|~r2_hidden(X3,k2_xboole_0(esk70_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k1_tarski(esk8_2(X1,X2))))|~r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,esk9_2(X1,X2),X2)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_6549, plain, (r2_hidden(esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk9_2(X1,X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|v5_waybel_6(X2,X1)|~r2_hidden(X3,k1_tarski(esk8_2(X1,X2)))|~r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,esk9_2(X1,X2),X2)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_6357, plain, (r2_hidden(esk72_3(X1,X2,k1_tarski(X2)),k1_tarski(X2))|r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(esk70_3(X1,X2,k1_tarski(X2)),k1_tarski(X2))),X2)|esk72_3(X1,X2,k1_tarski(X2))=X2|~r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X2)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6695, plain, (v1_finset_1(X1)|~v5_finset_1(X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_6744, negated_conjecture, (r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0))),k1_tarski(esk61_1(esk1_0)))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk61_1(esk1_0)|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk2_0|esk69_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk2_0|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))))|~m1_subset_1(k2_tarski(esk2_0,esk61_1(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),u1_struct_0(esk1_0))|~v2_yellow_0(esk1_0))).#trainneg
cnf(i_0_6769, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k1_tarski(esk4_1(u1_struct_0(esk1_0))))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))))|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),u1_struct_0(esk1_0))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0))))=esk4_1(u1_struct_0(esk1_0))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0))))=esk2_0|esk69_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0))))=esk2_0|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))))|~m1_subset_1(k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,u1_struct_0(esk1_0)),esk2_0)|~v1_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6768, negated_conjecture, (r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0))))|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k1_tarski(esk4_1(u1_struct_0(esk1_0))))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0))))=esk4_1(u1_struct_0(esk1_0))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0))))=esk2_0|esk69_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0))))=esk2_0|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))))|~m1_subset_1(k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk4_1(u1_struct_0(esk1_0)))),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,u1_struct_0(esk1_0)),esk2_0)|~v1_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6727, negated_conjecture, (r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r2_hidden(X1,k1_tarski(esk2_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~r1_orders_2(esk1_0,X1,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4870, plain, (r2_hidden(esk69_3(X1,X2,k2_tarski(X3,k2_yellow_0(X1,X4))),k2_tarski(X3,k2_yellow_0(X1,X4)))|r2_hidden(esk72_3(X1,X2,k2_tarski(X3,k2_yellow_0(X1,X4))),k2_tarski(X3,k2_yellow_0(X1,X4)))|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(k1_tarski(X3),X4)),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X4)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_7767, negated_conjecture, (r2_hidden(esk72_3(X1,k2_yellow_0(X1,X2),X2),X2)|v1_xboole_0(X2)|v1_xboole_0(esk70_3(X1,k2_yellow_0(X1,X2),X2))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(X1,esk4_1(esk70_3(X1,k2_yellow_0(X1,X2),X2)),k2_yellow_0(X1,X2))|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk4_1(esk70_3(X1,k2_yellow_0(X1,X2),X2)),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(X1,k2_yellow_0(X1,X2),X2)),esk2_0)|~v1_finset_1(esk70_3(X1,k2_yellow_0(X1,X2),X2))|~m1_subset_1(esk70_3(X1,k2_yellow_0(X1,X2),X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7327, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~r1_orders_2(esk1_0,esk2_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7732, plain, (r2_hidden(esk72_3(X1,X2,k1_tarski(X2)),k1_tarski(X2))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~r2_hidden(X3,esk70_3(X1,X2,k1_tarski(X2)))|~r1_orders_2(X1,X3,X2)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_7715, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2))),X3)|r1_orders_2(X1,k2_yellow_0(X1,X4),X3)|~r2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2)))|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X4,X2)),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(k2_yellow_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,X4))).#trainneg
cnf(i_0_8024, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|~r2_hidden(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k1_tarski(esk2_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7412, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~r1_orders_2(esk1_0,X1,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))))).#trainneg
cnf(i_0_1831, negated_conjecture, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,esk4_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk4_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),esk2_0)|~v1_finset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~m1_subset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7230, negated_conjecture, (v1_finset_1(esk8_2(esk1_0,esk2_0))|v1_finset_1(esk9_2(esk1_0,esk2_0))|v5_waybel_6(esk2_0,esk1_0)|~v5_finset_1(k1_numbers)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7231, negated_conjecture, (r1_orders_2(esk1_0,esk8_2(esk1_0,esk2_0),esk2_0)|v1_finset_1(esk9_2(esk1_0,esk2_0))|v5_waybel_6(esk2_0,esk1_0)|~v5_finset_1(k1_numbers)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7349, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~v1_xboole_0(k2_xboole_0(k1_tarski(esk2_0),esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7716, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2))),X3)|r1_orders_2(X1,k2_yellow_0(X1,X4),X3)|~r2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2)))|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X2,X4)),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(k2_yellow_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,X2))).#trainneg
cnf(i_0_7852, plain, (esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk9_2(X1,X2)|v5_waybel_6(X2,X1)|~r1_orders_2(X1,esk9_2(X1,X2),X2)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_hidden(X3,esk70_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))))).#trainneg
cnf(i_0_7936, plain, (esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk9_2(X1,X2)|v5_waybel_6(X2,X1)|~r1_orders_2(X1,esk9_2(X1,X2),X2)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,k1_tarski(esk8_2(X1,X2))))).#trainneg
cnf(i_0_7722, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X3)|r1_orders_2(X1,X4,X3)|~r2_yellow_0(X1,k1_tarski(X2))|~r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(X2,X4)),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_7219, plain, (r2_hidden(esk69_3(X1,X2,k2_tarski(X3,X4)),k2_tarski(X3,X4))|r1_orders_2(X1,esk72_3(X1,X2,k2_tarski(X3,X4)),X2)|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(X3,X4)),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1832, plain, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,X4,X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))))).#trainneg
cnf(i_0_7859, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~r1_orders_2(esk1_0,esk12_1(k2_xboole_0(k1_tarski(esk2_0),esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk12_1(k2_xboole_0(k1_tarski(esk2_0),esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8016, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|v5_waybel_6(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),esk1_0)|~r2_hidden(k12_lattice3(esk1_0,esk8_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))),esk9_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))),k1_tarski(esk2_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k12_lattice3(esk1_0,esk8_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))),esk9_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8015, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|v2_struct_0(esk1_0)|v5_waybel_6(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),esk1_0)|~r2_hidden(k11_lattice3(esk1_0,esk8_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))),esk9_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))),k1_tarski(esk2_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k11_lattice3(esk1_0,esk8_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))),esk9_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8023, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|v5_waybel_6(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),esk1_0)|~r2_hidden(k2_yellow_0(esk1_0,k2_tarski(esk8_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))),esk9_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))))),k1_tarski(esk2_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk8_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))),esk9_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))))),u1_struct_0(esk1_0))|~m1_subset_1(esk9_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8029, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~r2_hidden(k2_yellow_0(esk1_0,esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))),k1_tarski(esk2_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))),u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)))|~r2_yellow_0(esk1_0,k1_tarski(esk2_0))|~v5_waybel_6(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),esk1_0))).#trainneg
cnf(i_0_5584, plain, (r1_orders_2(X1,esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|esk69_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk9_2(X1,X2)|v5_waybel_6(X2,X1)|~r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,k1_tarski(esk9_2(X1,X2)))|~m1_subset_1(k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_6893, plain, (r1_orders_2(X1,k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2)),X3)|r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X4)),X3)|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X2))|~r2_yellow_0(X1,k1_tarski(X4))|~r1_orders_2(X1,k12_lattice3(X1,X4,X2),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(k1_tarski(X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_7723, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X3)|r1_orders_2(X1,X4,X3)|~r2_yellow_0(X1,k1_tarski(X2))|~r1_orders_2(X1,k12_lattice3(X1,X2,X4),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_7082, plain, (k12_lattice3(X1,X2,k12_lattice3(X1,X4,X3))=k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(k1_tarski(X3),k1_tarski(X4)),k1_tarski(X2)))|~r2_yellow_0(X1,k2_xboole_0(k1_tarski(X3),k1_tarski(X4)))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k1_tarski(X3))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_7990, plain, (v1_xboole_0(X1)|m1_subset_1(esk72_3(X2,X3,X1),u1_struct_0(X2))|~r1_orders_2(X2,k2_yellow_0(X2,X1),X3)|~r1_orders_2(X2,X4,X3)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r2_hidden(X4,k1_tarski(esk69_3(X2,X3,X1))))).#trainneg
cnf(i_0_1833, plain, (r2_hidden(esk72_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|r2_hidden(esk69_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~v1_finset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~m1_subset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_7005, plain, (v3_ordinal1(X1)|~v3_ordinal1(X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_7013, plain, (v1_funct_1(X1)|~v4_funct_1(X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_7818, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0))),u1_struct_0(esk1_0))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk61_1(esk1_0)|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk2_0|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))))|~v1_xboole_0(k2_xboole_0(esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0))),k1_tarski(esk2_0)))|~m1_subset_1(k2_tarski(esk2_0,esk61_1(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),u1_struct_0(esk1_0))|~v2_yellow_0(esk1_0)|~r1_orders_2(esk1_0,esk61_1(esk1_0),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))))|~m1_subset_1(esk61_1(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7814, negated_conjecture, (r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk61_1(esk1_0)|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk2_0|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk2_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))))|~m1_subset_1(k2_tarski(esk2_0,esk61_1(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),u1_struct_0(esk1_0))|~v2_yellow_0(esk1_0)|~r1_orders_2(esk1_0,esk61_1(esk1_0),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))))|~m1_subset_1(esk61_1(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7802, negated_conjecture, (r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk61_1(esk1_0)|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk2_0|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))))|~v1_xboole_0(k2_xboole_0(esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0))),k1_tarski(esk2_0)))|~m1_subset_1(k2_tarski(esk2_0,esk61_1(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),u1_struct_0(esk1_0))|~v2_yellow_0(esk1_0)|~r1_orders_2(esk1_0,esk61_1(esk1_0),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))))|~m1_subset_1(esk61_1(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8014, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k1_tarski(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))))=k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))|~r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k1_tarski(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))),k1_tarski(esk2_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k1_tarski(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7474, plain, (r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,k1_tarski(X2))),X2)|r1_orders_2(X1,X2,X2)|esk72_3(X1,X2,k1_tarski(X2))=X2|r1_orders_2(X1,esk72_3(X1,X2,k1_tarski(X2)),X2)|~r2_yellow_0(X1,esk70_3(X1,X2,k1_tarski(X2)))|~v5_waybel_6(X2,X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X2))).#trainneg
cnf(i_0_7308, plain, (r1_tarski(esk70_3(X1,X2,k2_tarski(X3,X4)),k2_tarski(X3,X4))|r1_orders_2(X1,esk72_3(X1,X2,k2_tarski(X3,X4)),X2)|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(X3,X4)),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_7653, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X3)|r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X4)),X3)|~r2_yellow_0(X1,k1_tarski(X2))|~r2_yellow_0(X1,k1_tarski(X4))|~r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(X4,X2)),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_7435, plain, (esk72_3(X1,X2,k1_tarski(X2))=X2|r1_orders_2(X1,esk72_3(X1,X2,k1_tarski(X2)),X2)|~r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X2)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_hidden(X3,esk70_3(X1,X2,k1_tarski(X2))))).#trainneg
cnf(i_0_7166, negated_conjecture, (v3_ordinal1(esk4_1(X1))|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|~v3_ordinal1(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1834, plain, (r1_tarski(esk70_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|r2_hidden(esk72_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~v1_finset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~m1_subset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_7460, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|v1_xboole_0(k2_xboole_0(k1_tarski(esk2_0),esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~r1_orders_2(esk1_0,esk4_1(k2_xboole_0(k1_tarski(esk2_0),esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk4_1(k2_xboole_0(k1_tarski(esk2_0),esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)))),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(esk2_0),esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)))),esk2_0)|~v1_finset_1(k2_xboole_0(k1_tarski(esk2_0),esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))))|~m1_subset_1(k2_xboole_0(k1_tarski(esk2_0),esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7410, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|v1_xboole_0(k2_xboole_0(k1_tarski(esk2_0),esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~r1_orders_2(esk1_0,esk4_1(k2_xboole_0(k1_tarski(esk2_0),esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk4_1(k2_xboole_0(k1_tarski(esk2_0),esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)))),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(esk2_0),esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)))),esk2_0)|~v1_finset_1(k2_xboole_0(k1_tarski(esk2_0),esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))))|~m1_subset_1(k2_xboole_0(k1_tarski(esk2_0),esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_7800, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0))),k1_tarski(esk2_0))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))))|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0))),u1_struct_0(esk1_0))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk61_1(esk1_0)|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk2_0|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))))|~m1_subset_1(k2_tarski(esk2_0,esk61_1(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),u1_struct_0(esk1_0))|~v2_yellow_0(esk1_0)|~r1_orders_2(esk1_0,esk61_1(esk1_0),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))))|~m1_subset_1(esk61_1(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7768, plain, (r2_hidden(esk72_3(X1,k2_yellow_0(X1,X2),X2),X2)|v1_xboole_0(X2)|v1_xboole_0(esk70_3(X1,k2_yellow_0(X1,X2),X2))|~r1_orders_2(X1,X3,k2_yellow_0(X1,X2))|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,esk70_3(X1,k2_yellow_0(X1,X2),X2)))).#trainneg
cnf(i_0_1835, plain, (r2_hidden(esk69_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|m1_subset_1(esk72_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),u1_struct_0(X1))|r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~v1_finset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~m1_subset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_8669, plain, (v3_ordinal1(X1)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_8681, plain, (v1_funct_1(X1)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_8215, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r2_hidden(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k1_tarski(esk2_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_8134, plain, (k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2)))=k2_yellow_0(X1,X2)|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(k1_tarski(k2_yellow_0(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_8137, plain, (k2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(k2_yellow_0(X1,X2))))=k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2)))|~r2_yellow_0(X1,X2)|~m1_subset_1(k1_tarski(k2_yellow_0(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_8198, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3))))),X2)|r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3)))))|~r2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3)))|~r2_yellow_0(X1,esk70_3(X1,X2,X3))|~v5_waybel_6(X2,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1836, plain, (r1_tarski(esk70_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|m1_subset_1(esk72_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),u1_struct_0(X1))|r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~v1_finset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~m1_subset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_8283, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,esk70_3(X1,X2,X3)))),X2)|r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3))),X2)|r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|~r2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,esk70_3(X1,X2,X3))))|~r2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3)))|~r2_yellow_0(X1,esk70_3(X1,X2,X3))|~v5_waybel_6(X2,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_8340, negated_conjecture, (v1_xboole_0(X1)|m1_subset_1(esk72_3(X2,X3,X1),u1_struct_0(X2))|v1_xboole_0(esk70_3(X2,X3,X1))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(X2,k2_yellow_0(X2,X1),X3)|~r1_orders_2(X2,esk4_1(esk70_3(X2,X3,X1)),X3)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(esk4_1(esk70_3(X2,X3,X1)),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(X2,X3,X1)),esk2_0)|~v1_finset_1(esk70_3(X2,X3,X1))|~m1_subset_1(esk70_3(X2,X3,X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_8155, plain, (r1_orders_2(X1,X2,X3)|~r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1837, plain, (r2_hidden(esk72_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|r2_hidden(esk69_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|v1_xboole_0(X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v1_finset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~m1_subset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_8688, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),k2_tarski(esk2_0,esk61_1(esk1_0)))=esk2_0|v2_struct_0(esk1_0)|esk61_1(esk1_0)!=esk2_0|~v2_yellow_0(esk1_0)|~r1_orders_2(esk1_0,esk61_1(esk1_0),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))))|~r1_orders_2(esk1_0,esk2_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk61_1(esk1_0))),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk61_1(esk1_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk61_1(esk1_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8426, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~r2_yellow_0(esk1_0,esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)))|~r2_yellow_0(esk1_0,k1_tarski(esk2_0))|~r2_hidden(k2_yellow_0(esk1_0,esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))),k1_tarski(esk2_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~v5_waybel_6(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),esk1_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_8409, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|v5_waybel_6(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),esk1_0)|~m1_subset_1(k11_lattice3(esk1_0,esk8_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))),esk9_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k11_lattice3(esk1_0,esk8_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))),esk9_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))),k1_tarski(esk2_0)))).#trainneg
cnf(i_0_8403, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|v5_waybel_6(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(k12_lattice3(esk1_0,esk8_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))),esk9_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k12_lattice3(esk1_0,esk8_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0))),esk9_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))),k1_tarski(esk2_0)))).#trainneg
cnf(i_0_9005, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|esk69_3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~m1_subset_1(k1_tarski(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_8996, plain, (k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2)))=k2_yellow_0(X1,X2)|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_8914, negated_conjecture, (v2_struct_0(esk1_0)|r2_hidden(esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk3_0)|v1_xboole_0(esk3_0)|~r1_orders_2(esk1_0,esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1838, plain, (r1_tarski(esk70_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|r2_hidden(esk72_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|v1_xboole_0(X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v1_finset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~m1_subset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_8970, plain, (v1_xxreal_0(X1)|~v1_finset_1(X1)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_8971, plain, (v1_card_1(X1)|~v1_finset_1(X1)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_8913, negated_conjecture, (v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|v1_xboole_0(esk3_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk3_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9015, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_yellow_0(esk1_0,k1_tarski(esk2_0)))|~r1_orders_2(esk1_0,esk2_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9026, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r2_hidden(k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k1_tarski(esk2_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_yellow_0(esk1_0,k1_tarski(esk2_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9298, plain, (r1_orders_2(X1,X2,k2_yellow_0(X1,k1_tarski(X2)))|v2_struct_0(X1)|~v5_waybel_6(k2_yellow_0(X1,k1_tarski(X2)),X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,k1_tarski(X2)),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_9289, plain, (r2_hidden(esk72_3(X1,X2,X3),X3)|r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,esk70_3(X1,X2,X3)))),X2)|r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3))),X2)|v1_xboole_0(X3)|~r2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3)))|~r2_yellow_0(X1,esk70_3(X1,X2,X3))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(k2_yellow_0(X1,esk70_3(X1,X2,X3))),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_8200, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3))))),X2)|r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3)))))|~r2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3)))|~r2_yellow_0(X1,esk70_3(X1,X2,X3))|~v5_waybel_6(X2,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1839, plain, (r2_hidden(esk69_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|m1_subset_1(esk72_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),u1_struct_0(X1))|v1_xboole_0(X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v1_finset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~m1_subset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_9016, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_yellow_0(esk1_0,k1_tarski(esk2_0)))|~v1_xboole_0(k2_xboole_0(k1_tarski(esk2_0),esk70_3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_tarski(esk2_0,esk2_0))))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9020, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|~r1_orders_2(esk1_0,esk12_1(k2_xboole_0(k1_tarski(esk2_0),esk70_3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_tarski(esk2_0,esk2_0)))),k2_yellow_0(esk1_0,k1_tarski(esk2_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_yellow_0(esk1_0,k1_tarski(esk2_0)))|~m1_subset_1(esk12_1(k2_xboole_0(k1_tarski(esk2_0),esk70_3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_tarski(esk2_0,esk2_0)))),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9225, plain, (k12_lattice3(X1,k2_yellow_0(X1,k1_tarski(X2)),X2)=k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,k1_tarski(X2))))|~r2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(k2_yellow_0(X1,k1_tarski(X2)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(k2_yellow_0(X1,k1_tarski(X2))),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_8285, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,esk70_3(X1,X2,X3)))),X2)|r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3))),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,esk70_3(X1,X2,X3))))|~r2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3)))|~r2_yellow_0(X1,esk70_3(X1,X2,X3))|~v5_waybel_6(X2,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_8370, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~r1_orders_2(esk1_0,X1,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(X1,k2_xboole_0(k1_tarski(esk2_0),esk70_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk2_0,esk2_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_8353, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),k2_yellow_0(X1,k2_tarski(X2,X3)))|r1_orders_2(X1,X3,k2_yellow_0(X1,k2_tarski(X2,X3)))|~r2_yellow_0(X1,k1_tarski(X2))|~v5_waybel_6(k2_yellow_0(X1,k2_tarski(X2,X3)),X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,k2_tarski(X2,X3)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k2_tarski(X2,X3),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1840, plain, (r1_tarski(esk70_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|m1_subset_1(esk72_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),u1_struct_0(X1))|v1_xboole_0(X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~v1_finset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~m1_subset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_9022, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|v5_waybel_6(k2_yellow_0(esk1_0,k1_tarski(esk2_0)),esk1_0)|~r2_hidden(k11_lattice3(esk1_0,esk8_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0))),esk9_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)))),k1_tarski(esk2_0))|~m1_subset_1(k11_lattice3(esk1_0,esk8_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0))),esk9_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)))),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9023, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|v5_waybel_6(k2_yellow_0(esk1_0,k1_tarski(esk2_0)),esk1_0)|~r2_hidden(k12_lattice3(esk1_0,esk8_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0))),esk9_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)))),k1_tarski(esk2_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_yellow_0(esk1_0,k1_tarski(esk2_0)))|~m1_subset_1(k12_lattice3(esk1_0,esk8_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0))),esk9_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)))),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9342, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|v5_waybel_6(k2_yellow_0(esk1_0,k1_tarski(esk2_0)),esk1_0)|~m1_subset_1(k11_lattice3(esk1_0,esk8_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0))),esk9_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)))),u1_struct_0(esk1_0))|~m1_subset_1(k11_lattice3(esk1_0,esk8_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0))),esk9_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)))),k1_tarski(esk2_0))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9349, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_tarski(esk2_0,esk2_0))=esk2_0|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|v5_waybel_6(k2_yellow_0(esk1_0,k1_tarski(esk2_0)),esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),k2_yellow_0(esk1_0,k1_tarski(esk2_0)))|~m1_subset_1(k12_lattice3(esk1_0,esk8_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0))),esk9_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)))),u1_struct_0(esk1_0))|~m1_subset_1(k12_lattice3(esk1_0,esk8_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0))),esk9_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)))),k1_tarski(esk2_0))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_8708, plain, (r1_tarski(esk70_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|r1_orders_2(X1,esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|v5_waybel_6(X2,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_9153, negated_conjecture, (esk72_3(esk1_0,esk2_0,X1)=X2|v2_struct_0(esk1_0)|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,X1)),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|X1!=k1_tarski(X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~v1_finset_1(X1)|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_8636, plain, (v1_xboole_0(X1)|m1_subset_1(esk72_3(X2,X3,X1),u1_struct_0(X2))|~r1_orders_2(X2,k2_yellow_0(X2,X1),X3)|~r1_orders_2(X2,esk69_3(X2,X3,X1),X3)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(esk69_3(X2,X3,X1),u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_8367, plain, (r2_hidden(esk69_3(X1,k2_yellow_0(X1,k2_tarski(X2,X3)),k2_tarski(X2,X3)),k2_tarski(X2,X3))|r1_orders_2(X1,esk72_3(X1,k2_yellow_0(X1,k2_tarski(X2,X3)),k2_tarski(X2,X3)),k2_yellow_0(X1,k2_tarski(X2,X3)))|~m1_subset_1(k2_yellow_0(X1,k2_tarski(X2,X3)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k2_tarski(X2,X3),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_8707, plain, (r1_tarski(esk70_3(X1,k2_yellow_0(X1,k2_tarski(X2,X3)),k2_tarski(X2,X3)),k2_tarski(X2,X3))|r1_orders_2(X1,esk72_3(X1,k2_yellow_0(X1,k2_tarski(X2,X3)),k2_tarski(X2,X3)),k2_yellow_0(X1,k2_tarski(X2,X3)))|~m1_subset_1(k2_yellow_0(X1,k2_tarski(X2,X3)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k2_tarski(X2,X3),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_7122, plain, (k2_yellow_0(X1,k2_xboole_0(k1_tarski(X3),k1_tarski(X2)))=k12_lattice3(X1,X3,X2)|~r2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_9450, negated_conjecture, (v2_struct_0(esk1_0)|~m1_subset_1(esk2_0,esk3_0))).#trainneg
cnf(i_0_1841, plain, (r2_hidden(esk72_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|r2_hidden(esk69_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~v1_finset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~m1_subset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_9369, negated_conjecture, (v1_zfmisc_1(esk3_0)|~v1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9370, negated_conjecture, (v6_ordinal1(esk3_0)|~v6_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9371, negated_conjecture, (v1_relat_1(esk3_0)|~v1_relat_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9372, negated_conjecture, (v5_finset_1(esk3_0)|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9373, negated_conjecture, (v4_funct_1(esk3_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9487, negated_conjecture, (v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9435, negated_conjecture, (v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_9451, negated_conjecture, (v2_struct_0(esk1_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),esk3_0))).#trainneg
cnf(i_0_9424, plain, (k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2)))=k2_yellow_0(X1,X2)|v1_xboole_0(X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_9356, negated_conjecture, (v3_ordinal1(esk3_0)|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9359, negated_conjecture, (esk3_0=X1|k1_tarski(X1)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1842, plain, (r1_tarski(esk70_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|r2_hidden(esk72_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~v1_finset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~m1_subset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_9357, negated_conjecture, (v1_funct_1(esk3_0)|~v4_funct_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9358, negated_conjecture, (v1_relat_1(esk3_0)|~v4_funct_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9374, negated_conjecture, (v1_funct_1(esk3_0)|~v1_relat_1(u1_struct_0(esk1_0))|~v1_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9555, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk69_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9556, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9593, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,k1_tarski(X2)))=k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,k1_tarski(X2))))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,k1_tarski(X2)),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~r2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(k1_tarski(k2_yellow_0(X1,k1_tarski(X2))),k1_zfmisc_1(u1_struct_0(X1)))|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_9490, negated_conjecture, (esk72_3(esk1_0,esk2_0,k2_tarski(esk2_0,esk2_0))=esk2_0|v2_struct_0(esk1_0)|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1843, plain, (r2_hidden(esk69_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|m1_subset_1(esk72_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),u1_struct_0(X1))|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~v1_finset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~m1_subset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_9360, negated_conjecture, (esk3_0=X1|k2_xboole_0(k1_zfmisc_1(u1_struct_0(esk1_0)),X2)!=k1_tarski(X1))).#trainneg
cnf(i_0_9361, negated_conjecture, (esk3_0=X1|k2_xboole_0(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))!=k1_tarski(X1))).#trainneg
cnf(i_0_9436, negated_conjecture, (v2_struct_0(esk1_0)|esk72_3(esk1_0,esk2_0,k1_tarski(esk2_0))=esk2_0|~m1_subset_1(esk72_3(esk1_0,esk2_0,k1_tarski(esk2_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk72_3(esk1_0,esk2_0,k1_tarski(esk2_0)),esk3_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(esk2_0)),esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9478, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9569, negated_conjecture, (v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk3_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9422, negated_conjecture, (esk69_3(esk1_0,esk2_0,k2_tarski(esk2_0,esk2_0))=esk2_0|esk72_3(esk1_0,esk2_0,k2_tarski(esk2_0,esk2_0))=esk2_0|v2_struct_0(esk1_0)|~m1_subset_1(k1_tarski(esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9559, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~r2_hidden(X1,esk70_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)))|~r1_orders_2(esk1_0,X1,esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9381, negated_conjecture, (r2_hidden(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|v1_xboole_0(esk3_0)|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk3_0))).#trainneg
cnf(i_0_9824, negated_conjecture, (v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9959, negated_conjecture, (~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),esk3_0))).#trainneg
cnf(i_0_9852, negated_conjecture, (m1_subset_1(esk4_1(esk3_0),u1_struct_0(esk1_0))|v5_waybel_6(esk2_0,esk1_0))).#trainneg
cnf(i_0_1844, plain, (r1_tarski(esk70_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|v1_xboole_0(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|m1_subset_1(esk72_3(X1,X2,k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3)))),u1_struct_0(X1))|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~v1_finset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))))|~m1_subset_1(k2_xboole_0(esk70_3(X1,X2,X3),k1_tarski(esk69_3(X1,X2,X3))),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_9854, negated_conjecture, (~m1_subset_1(esk2_0,esk3_0))).#trainneg
cnf(i_0_9818, negated_conjecture, (v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v5_finset_1(esk3_0))).#trainneg
cnf(i_0_9819, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v3_ordinal1(esk3_0))).#trainneg
cnf(i_0_9820, negated_conjecture, (v1_funct_1(esk69_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v4_funct_1(esk3_0))).#trainneg
cnf(i_0_9828, negated_conjecture, (r1_tarski(X1,esk3_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~r1_tarski(X1,esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_9931, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10087, negated_conjecture, (m1_subset_1(esk71_3(esk1_0,esk2_0,esk3_0),esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9958, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~r1_tarski(X1,X2)|~r1_tarski(X2,esk3_0))).#trainneg
cnf(i_0_1856, plain, (r2_hidden(esk69_3(X1,X2,k1_tarski(X3)),k1_tarski(X3))|m1_subset_1(esk72_3(X1,X2,k1_tarski(X3)),u1_struct_0(X1))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,X3,X2)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_9860, negated_conjecture, (~m1_subset_1(k2_tarski(u1_struct_0(esk1_0),X1),esk3_0))).#trainneg
cnf(i_0_9817, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),k1_numbers)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_9826, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk3_0),k1_numbers)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_9970, negated_conjecture, (v5_waybel_6(esk2_0,esk1_0)|m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|v1_xboole_0(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,X1)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9953, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,esk4_1(esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~m1_subset_1(esk4_1(esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9954, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_9833, negated_conjecture, (m1_subset_1(esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),u1_struct_0(esk1_0))|r2_hidden(esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_9920, negated_conjecture, (r1_tarski(X1,esk3_0)|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~r1_tarski(X1,esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_9875, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_9867, negated_conjecture, (X1=X2|k1_tarski(X2)!=u1_struct_0(esk1_0)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_1857, plain, (r1_tarski(esk70_3(X1,X2,k1_tarski(X3)),k1_tarski(X3))|m1_subset_1(esk72_3(X1,X2,k1_tarski(X3)),u1_struct_0(X1))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,X3,X2)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_9861, negated_conjecture, (~m1_subset_1(k2_tarski(X1,u1_struct_0(esk1_0)),esk3_0))).#trainneg
cnf(i_0_9918, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk3_0),k1_numbers)|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_9924, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),k1_numbers)|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10269, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9832, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10148, negated_conjecture, (r1_tarski(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),k2_yellow_0(esk1_0,esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10093, negated_conjecture, (esk71_3(esk1_0,esk2_0,esk3_0)=X1|v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|k1_tarski(X1)!=esk70_3(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1874, plain, (k12_lattice3(X1,k12_lattice3(X1,X2,X3),k2_yellow_0(X1,X4))=k2_yellow_0(X1,k2_xboole_0(X4,k7_domain_1(u1_struct_0(X1),X2,X3)))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,X4)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_5884, plain, (v1_relat_1(X1)|~v4_funct_1(k2_xboole_0(X2,X3))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_10260, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10262, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),k1_numbers)|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10270, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10281, negated_conjecture, (v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10505, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1896, plain, (k2_yellow_0(X1,k2_xboole_0(k6_domain_1(u1_struct_0(X1),X2),X3))=k12_lattice3(X1,X2,k2_yellow_0(X1,X3))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X2))|~r2_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_10560, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk12_1(esk3_0),esk12_1(esk3_0)))).#trainneg
cnf(i_0_10276, negated_conjecture, (v3_ordinal1(esk72_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v3_ordinal1(esk3_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10272, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10491, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r2_hidden(X1,esk3_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_10342, negated_conjecture, (v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk71_3(esk1_0,esk2_0,esk3_0),esk3_0))).#trainneg
cnf(i_0_10626, negated_conjecture, (v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10725, plain, (v1_relat_1(X1)|~v4_funct_1(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_1898, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k12_lattice3(X1,X4,X3))=k2_yellow_0(X1,k2_xboole_0(X2,k7_domain_1(u1_struct_0(X1),X3,X4)))|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_10568, negated_conjecture, (esk12_1(esk3_0)=X1|k1_tarski(X1)!=u1_struct_0(esk1_0))).#trainneg
cnf(i_0_10724, plain, (v1_relat_1(X1)|~v4_funct_1(X2)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_10416, negated_conjecture, (v1_finset_1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v5_finset_1(esk3_0))).#trainneg
cnf(i_0_10417, negated_conjecture, (v3_ordinal1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v3_ordinal1(esk3_0))).#trainneg
cnf(i_0_10337, negated_conjecture, (m1_subset_1(k2_tarski(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,esk3_0)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_10453, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10334, negated_conjecture, (m1_subset_1(k6_domain_1(esk3_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_10531, negated_conjecture, (m1_subset_1(esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),u1_struct_0(esk1_0))|m1_subset_1(esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_10329, negated_conjecture, (m1_subset_1(esk70_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|r2_hidden(esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_10387, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_1899, plain, (r2_hidden(esk72_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k7_domain_1(u1_struct_0(X1),X3,X4))|r2_hidden(esk69_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k7_domain_1(u1_struct_0(X1),X3,X4))|v1_xboole_0(k7_domain_1(u1_struct_0(X1),X3,X4))|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k12_lattice3(X1,X4,X3),X2)|~v1_finset_1(k7_domain_1(u1_struct_0(X1),X3,X4))|~m1_subset_1(k7_domain_1(u1_struct_0(X1),X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_10287, negated_conjecture, (esk4_1(esk3_0)=X1|v5_waybel_6(esk2_0,esk1_0)|k1_tarski(X1)!=u1_struct_0(esk1_0))).#trainneg
cnf(i_0_10418, negated_conjecture, (v1_funct_1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v4_funct_1(esk3_0))).#trainneg
cnf(i_0_10434, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10413, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0))).#trainneg
cnf(i_0_10428, negated_conjecture, (v1_finset_1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v5_finset_1(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10429, negated_conjecture, (v3_ordinal1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v3_ordinal1(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10335, negated_conjecture, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_10530, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_10830, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk69_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0))).#trainneg
cnf(i_0_11031, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk4_1(esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_11011, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_11371, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_1900, plain, (r1_tarski(esk70_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k7_domain_1(u1_struct_0(X1),X3,X4))|r2_hidden(esk72_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k7_domain_1(u1_struct_0(X1),X3,X4))|v1_xboole_0(k7_domain_1(u1_struct_0(X1),X3,X4))|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k12_lattice3(X1,X4,X3),X2)|~v1_finset_1(k7_domain_1(u1_struct_0(X1),X3,X4))|~m1_subset_1(k7_domain_1(u1_struct_0(X1),X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_11179, plain, (v1_relat_1(X1)|~r2_hidden(X1,k1_xboole_0))).#trainneg
cnf(i_0_11241, negated_conjecture, (~m1_subset_1(X1,esk3_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),esk3_0))).#trainneg
cnf(i_0_11203, negated_conjecture, (v1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_zfmisc_1(esk3_0))).#trainneg
cnf(i_0_11204, negated_conjecture, (v6_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v6_ordinal1(esk3_0))).#trainneg
cnf(i_0_11205, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_relat_1(esk3_0))).#trainneg
cnf(i_0_11206, negated_conjecture, (v5_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v5_finset_1(esk3_0))).#trainneg
cnf(i_0_11207, negated_conjecture, (v4_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v4_funct_1(esk3_0))).#trainneg
cnf(i_0_11032, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_10906, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,X2),X2)|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,X2),X1)|v1_xboole_0(X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X2),X1)|~v1_finset_1(X2)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_11433, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k1_tarski(X2))|~m1_subset_1(X2,esk3_0))).#trainneg
cnf(i_0_1901, plain, (r2_hidden(esk69_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k7_domain_1(u1_struct_0(X1),X3,X4))|v1_xboole_0(k7_domain_1(u1_struct_0(X1),X3,X4))|m1_subset_1(esk72_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k12_lattice3(X1,X4,X3),X2)|~v1_finset_1(k7_domain_1(u1_struct_0(X1),X3,X4))|~m1_subset_1(k7_domain_1(u1_struct_0(X1),X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_10872, negated_conjecture, (~m1_subset_1(k2_tarski(k1_zfmisc_1(u1_struct_0(esk1_0)),X1),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_10873, negated_conjecture, (~m1_subset_1(k2_tarski(X1,k1_zfmisc_1(u1_struct_0(esk1_0))),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_11190, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v3_ordinal1(k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_11191, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v4_funct_1(k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_11192, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v4_funct_1(k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_11208, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_relat_1(esk3_0)|~v1_funct_1(esk3_0))).#trainneg
cnf(i_0_10871, negated_conjecture, (v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,X1)),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_11332, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0))).#trainneg
cnf(i_0_10798, negated_conjecture, (r2_hidden(esk69_3(esk1_0,X1,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,esk70_3(esk1_0,esk2_0,esk3_0)),X1)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),X1)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11005, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11007, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11010, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk3_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11290, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k6_domain_1(esk3_0,X2))|~m1_subset_1(X2,esk3_0))).#trainneg
cnf(i_0_1902, plain, (r1_tarski(esk70_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k7_domain_1(u1_struct_0(X1),X3,X4))|v1_xboole_0(k7_domain_1(u1_struct_0(X1),X3,X4))|m1_subset_1(esk72_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k12_lattice3(X1,X4,X3),X2)|~v1_finset_1(k7_domain_1(u1_struct_0(X1),X3,X4))|~m1_subset_1(k7_domain_1(u1_struct_0(X1),X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_11035, negated_conjecture, (v1_relat_1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v4_funct_1(esk3_0))).#trainneg
cnf(i_0_11055, negated_conjecture, (v1_relat_1(esk69_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v4_funct_1(esk3_0))).#trainneg
cnf(i_0_11395, negated_conjecture, (v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v5_finset_1(esk3_0))).#trainneg
cnf(i_0_11396, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v3_ordinal1(esk3_0))).#trainneg
cnf(i_0_11457, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0))|~r1_orders_2(esk1_0,esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11009, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11329, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk70_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11333, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10441, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_11491, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_1903, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k12_lattice3(X1,X4,X3))=k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X3,X4),X2))|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_11597, plain, (v1_relat_1(k4_ordinal1)|~r2_hidden(k1_numbers,k1_xboole_0))).#trainneg
cnf(i_0_11475, negated_conjecture, (v1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11476, negated_conjecture, (v6_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v6_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11503, negated_conjecture, (v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_1904, plain, (k12_lattice3(X1,k12_lattice3(X1,X3,X2),k2_yellow_0(X1,X4))=k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X2,X3),X4))|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_11477, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_relat_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11478, negated_conjecture, (v5_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11479, negated_conjecture, (v4_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11462, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_11480, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_relat_1(u1_struct_0(esk1_0))|~v1_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11504, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_11827, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11582, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_11577, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_1906, plain, (k12_lattice3(X1,k12_lattice3(X1,X3,X2),k2_yellow_0(X1,X4))=k2_yellow_0(X1,k2_xboole_0(X4,k7_domain_1(u1_struct_0(X1),X2,X3)))|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,X4)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_11598, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k4_ordinal1))|~r2_hidden(k1_numbers,k1_xboole_0))).#trainneg
cnf(i_0_11865, negated_conjecture, (v1_relat_1(k1_xboole_0)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v4_funct_1(esk3_0))).#trainneg
cnf(i_0_11866, negated_conjecture, (v1_relat_1(k1_xboole_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v4_funct_1(esk3_0))).#trainneg
cnf(i_0_11931, negated_conjecture, (v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0))).#trainneg
cnf(i_0_11397, negated_conjecture, (v1_funct_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v4_funct_1(esk3_0))).#trainneg
cnf(i_0_11398, negated_conjecture, (v1_relat_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v4_funct_1(esk3_0))).#trainneg
cnf(i_0_11884, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11493, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_10333, negated_conjecture, (m1_subset_1(esk4_1(k1_zfmisc_1(esk3_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(esk3_0)),esk2_0)|~v1_finset_1(k1_zfmisc_1(esk3_0))|~m1_subset_1(k1_zfmisc_1(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_11234, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k2_tarski(X2,X3))|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X2,esk3_0))).#trainneg
cnf(i_0_1912, negated_conjecture, (r2_hidden(esk69_3(X1,X2,esk4_1(k1_zfmisc_1(u1_struct_0(X1)))),esk4_1(k1_zfmisc_1(u1_struct_0(X1))))|r1_orders_2(X1,esk72_3(X1,X2,esk4_1(k1_zfmisc_1(u1_struct_0(X1)))),X2)|v1_xboole_0(esk4_1(k1_zfmisc_1(u1_struct_0(X1))))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(X1,k2_yellow_0(X1,esk4_1(k1_zfmisc_1(u1_struct_0(X1)))),X2)|~v1_finset_1(esk4_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(u1_struct_0(X1))),esk2_0)|~v1_finset_1(k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k1_zfmisc_1(u1_struct_0(X1)),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_11273, negated_conjecture, (k6_domain_1(esk3_0,X1)=X2|k1_tarski(X2)!=k1_zfmisc_1(u1_struct_0(esk1_0))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_5901, plain, (v3_ordinal1(X1)|~v3_ordinal1(k2_xboole_0(X2,X3))|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_5902, plain, (v1_funct_1(X1)|~v4_funct_1(k2_xboole_0(X2,X3))|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_11932, negated_conjecture, (v1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_zfmisc_1(esk3_0))).#trainneg
cnf(i_0_11933, negated_conjecture, (v6_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v6_ordinal1(esk3_0))).#trainneg
cnf(i_0_11934, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_relat_1(esk3_0))).#trainneg
cnf(i_0_11935, negated_conjecture, (v5_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v5_finset_1(esk3_0))).#trainneg
cnf(i_0_11581, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(X1,X2),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_11489, negated_conjecture, (v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk3_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11485, negated_conjecture, (r2_hidden(esk69_3(esk1_0,X1,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,esk70_3(esk1_0,esk2_0,esk3_0)),X1)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),X1)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11640, negated_conjecture, (r2_hidden(esk69_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),k2_yellow_0(esk1_0,X1))|v1_xboole_0(X1)|~v1_finset_1(X1)|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_11008, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11883, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11372, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_11580, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(X2,X1),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_1913, negated_conjecture, (r1_tarski(esk70_3(X1,X2,esk4_1(k1_zfmisc_1(u1_struct_0(X1)))),esk4_1(k1_zfmisc_1(u1_struct_0(X1))))|r1_orders_2(X1,esk72_3(X1,X2,esk4_1(k1_zfmisc_1(u1_struct_0(X1)))),X2)|v1_xboole_0(esk4_1(k1_zfmisc_1(u1_struct_0(X1))))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(X1,k2_yellow_0(X1,esk4_1(k1_zfmisc_1(u1_struct_0(X1)))),X2)|~v1_finset_1(esk4_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(u1_struct_0(X1))),esk2_0)|~v1_finset_1(k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k1_zfmisc_1(u1_struct_0(X1)),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5903, plain, (v1_relat_1(X1)|~v4_funct_1(k2_xboole_0(X2,X3))|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_7011, plain, (v3_ordinal1(X1)|~r2_hidden(X1,X2)|~v1_card_1(k2_xboole_0(X2,X3)))).#trainneg
cnf(i_0_11936, negated_conjecture, (v4_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v4_funct_1(esk3_0))).#trainneg
cnf(i_0_11937, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_relat_1(esk3_0)|~v1_funct_1(esk3_0))).#trainneg
cnf(i_0_10430, negated_conjecture, (v1_funct_1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v4_funct_1(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_11036, negated_conjecture, (v1_relat_1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v4_funct_1(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10805, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_11265, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0))).#trainneg
cnf(i_0_11334, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0))).#trainneg
cnf(i_0_11327, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk70_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11579, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_xboole_0(X2,X3),k1_zfmisc_1(esk3_0))|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_1956, plain, (m1_subset_1(esk69_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),k6_domain_1(u1_struct_0(X1),X3))|r2_hidden(esk72_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),k6_domain_1(u1_struct_0(X1),X3))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|~r1_orders_2(X1,X3,X2)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_10326, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v3_ordinal1(esk3_0))).#trainneg
cnf(i_0_10327, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v3_ordinal1(esk3_0))).#trainneg
cnf(i_0_10627, negated_conjecture, (v1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_zfmisc_1(esk3_0))).#trainneg
cnf(i_0_10628, negated_conjecture, (v6_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v6_ordinal1(esk3_0))).#trainneg
cnf(i_0_10629, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_relat_1(esk3_0))).#trainneg
cnf(i_0_11013, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_9586, negated_conjecture, (r2_hidden(esk71_3(esk1_0,esk2_0,X1),esk70_3(esk1_0,esk2_0,X1))|r2_hidden(esk72_3(esk1_0,esk2_0,X1),X1)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12057, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_1958, plain, (esk69_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3))=X4|r2_hidden(esk72_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),k6_domain_1(u1_struct_0(X1),X3))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|k6_domain_1(u1_struct_0(X1),X3)!=k1_tarski(X4)|~r1_orders_2(X1,X3,X2)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_12258, negated_conjecture, (v1_relat_1(k1_xboole_0)|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v4_funct_1(esk3_0))).#trainneg
cnf(i_0_12632, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk71_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk71_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12386, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),k2_yellow_0(esk1_0,esk3_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12622, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(X1)|v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,X1),X1)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~r1_orders_2(esk1_0,esk71_3(esk1_0,esk2_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk71_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12538, negated_conjecture, (v1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12539, negated_conjecture, (v6_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v6_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12540, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_relat_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12541, negated_conjecture, (v5_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12385, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),k2_yellow_0(esk1_0,esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12365, negated_conjecture, (v2_struct_0(esk1_0)|r2_hidden(esk69_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12384, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),k2_yellow_0(esk1_0,esk3_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11814, negated_conjecture, (m1_subset_1(esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk3_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_1977, plain, (r2_hidden(esk5_2(X1,k2_xboole_0(X2,X3)),X3)|r2_hidden(esk5_2(X1,k2_xboole_0(X2,X3)),X2)|esk5_2(X1,k2_xboole_0(X2,X3))=X1|k2_xboole_0(X2,X3)=k1_tarski(X1))).#trainneg
cnf(i_0_12515, plain, (v3_ordinal1(X1)|~v1_card_1(X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_12542, negated_conjecture, (v4_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10630, negated_conjecture, (v5_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v5_finset_1(esk3_0))).#trainneg
cnf(i_0_10631, negated_conjecture, (v4_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v4_funct_1(esk3_0))).#trainneg
cnf(i_0_11187, negated_conjecture, (v1_xxreal_0(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk72_3(esk1_0,esk2_0,esk3_0))|~v3_ordinal1(esk3_0))).#trainneg
cnf(i_0_11188, negated_conjecture, (v1_card_1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk72_3(esk1_0,esk2_0,esk3_0))|~v3_ordinal1(esk3_0))).#trainneg
cnf(i_0_11952, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|r1_orders_2(esk1_0,esk71_3(esk1_0,esk2_0,esk3_0),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12319, negated_conjecture, (m1_subset_1(esk4_1(k2_tarski(X1,X2)),u1_struct_0(esk1_0))|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X2,esk3_0)|~m1_subset_1(X1,esk3_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk2_0)|~m1_subset_1(k2_tarski(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_10111, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11558, negated_conjecture, (m1_subset_1(esk69_3(X1,k2_yellow_0(X1,X2),X2),u1_struct_0(esk1_0))|r2_hidden(esk72_3(X1,k2_yellow_0(X1,X2),X2),X2)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(esk3_0))|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_12358, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk71_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk71_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11317, negated_conjecture, (m1_subset_1(esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk3_0)|m1_subset_1(esk70_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_9623, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,esk70_3(X1,X2,X3)))),X2)|r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3))),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~r2_yellow_0(X1,k1_tarski(esk69_3(X1,X2,X3)))|~r2_yellow_0(X1,esk70_3(X1,X2,X3))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(k2_yellow_0(X1,esk70_3(X1,X2,X3))),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_1979, plain, (r2_hidden(X1,X2)|r2_hidden(X1,X3)|v1_xboole_0(k2_xboole_0(X3,X2))|~m1_subset_1(X1,k2_xboole_0(X3,X2)))).#trainneg
cnf(i_0_11868, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v3_ordinal1(esk3_0))).#trainneg
cnf(i_0_11869, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v3_ordinal1(esk3_0))).#trainneg
cnf(i_0_12049, negated_conjecture, (v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v5_finset_1(esk3_0))).#trainneg
cnf(i_0_12050, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v3_ordinal1(esk3_0))).#trainneg
cnf(i_0_12051, negated_conjecture, (v1_funct_1(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v4_funct_1(esk3_0))).#trainneg
cnf(i_0_11328, negated_conjecture, (m1_subset_1(esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),u1_struct_0(esk1_0))|m1_subset_1(esk70_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11703, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k2_tarski(esk8_2(esk1_0,X1),esk9_2(esk1_0,X1))),k2_tarski(esk8_2(esk1_0,X1),esk9_2(esk1_0,X1)))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k2_tarski(esk8_2(esk1_0,X1),esk9_2(esk1_0,X1))),X1)|v5_waybel_6(X1,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk8_2(esk1_0,X1),esk9_2(esk1_0,X1)),k1_zfmisc_1(esk3_0))|~m1_subset_1(esk9_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10521, negated_conjecture, (m1_subset_1(esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk3_0)|m1_subset_1(esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_9728, plain, (m1_subset_1(esk72_3(X2,X1,k1_tarski(X1)),u1_struct_0(X2))|esk72_3(X2,X1,k1_tarski(X1))=X1|~r1_orders_2(X2,X1,X1)|~r1_orders_2(X2,k2_yellow_0(X2,k1_tarski(X1)),X1)|~m1_subset_1(X1,u1_struct_0(X2))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(X2)))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_12426, negated_conjecture, (esk69_3(esk1_0,k2_yellow_0(esk1_0,X1),X1)=X2|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),k2_yellow_0(esk1_0,X1))|v1_xboole_0(X1)|X1!=k1_tarski(X2)|~v1_finset_1(X1)|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_12281, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|v5_waybel_6(esk2_0,esk1_0)|~r1_tarski(X1,esk4_1(k1_zfmisc_1(esk3_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(esk3_0)),esk2_0)|~v1_finset_1(k1_zfmisc_1(esk3_0))|~m1_subset_1(k1_zfmisc_1(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1980, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X3,X4)|~r2_hidden(esk7_3(X3,X4,k2_xboole_0(X1,X2)),X4)|~r2_hidden(esk7_3(X3,X4,k2_xboole_0(X1,X2)),X1))).#trainneg
cnf(i_0_12052, negated_conjecture, (v1_relat_1(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v4_funct_1(esk3_0))).#trainneg
cnf(i_0_11526, negated_conjecture, (m1_subset_1(esk72_3(X1,X2,k1_tarski(X2)),u1_struct_0(esk1_0))|esk69_3(X1,X2,k1_tarski(X2))=X2|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(esk3_0))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_9558, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9379, negated_conjecture, (r2_hidden(esk69_3(esk1_0,X1,esk3_0),esk3_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,esk3_0),X1)|v1_xboole_0(esk3_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),X1)|~v1_finset_1(esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11826, negated_conjecture, (m1_subset_1(esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),u1_struct_0(esk1_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11486, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,esk70_3(esk1_0,esk2_0,esk3_0)),X1)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),X1)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11238, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k2_tarski(X2,X3)),k2_tarski(X2,X3))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k2_tarski(X2,X3)),X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X2,X3)),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X2,esk3_0))).#trainneg
cnf(i_0_12280, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|v5_waybel_6(esk2_0,esk1_0)|~r2_hidden(X1,esk4_1(k1_zfmisc_1(esk3_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(esk3_0)),esk2_0)|~v1_finset_1(k1_zfmisc_1(esk3_0))|~m1_subset_1(k1_zfmisc_1(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1981, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X3,X4)|~r2_hidden(esk7_3(X3,X4,k2_xboole_0(X1,X2)),X3)|~r2_hidden(esk7_3(X3,X4,k2_xboole_0(X1,X2)),X1))).#trainneg
cnf(i_0_10787, negated_conjecture, (v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10788, negated_conjecture, (v1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10789, negated_conjecture, (v6_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v6_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10790, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_relat_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10791, negated_conjecture, (v5_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10115, negated_conjecture, (r1_tarski(X1,esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_tarski(X1,esk70_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12647, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0))).#trainneg
cnf(i_0_13165, negated_conjecture, (v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk71_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12947, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13248, negated_conjecture, (r1_tarski(esk4_1(k1_zfmisc_1(esk3_0)),u1_struct_0(esk1_0))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(esk3_0)),esk2_0)|~v1_finset_1(k1_zfmisc_1(esk3_0))|~m1_subset_1(k1_zfmisc_1(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_1984, plain, (k2_xboole_0(X1,X2)=k1_tarski(X3)|esk5_2(X3,k2_xboole_0(X1,X2))!=X3|~r2_hidden(esk5_2(X3,k2_xboole_0(X1,X2)),X1))).#trainneg
cnf(i_0_12892, plain, (v3_ordinal1(X1)|~v1_card_1(k2_tarski(X2,X1)))).#trainneg
cnf(i_0_12893, plain, (v3_ordinal1(X1)|~v1_card_1(k2_tarski(X1,X2)))).#trainneg
cnf(i_0_13267, negated_conjecture, (v1_relat_1(k1_xboole_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v4_funct_1(esk3_0))).#trainneg
cnf(i_0_12813, negated_conjecture, (v3_ordinal1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_card_1(esk3_0))).#trainneg
cnf(i_0_12678, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13166, negated_conjecture, (v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,X1),X1)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~v1_finset_1(X1)|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk71_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_13378, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13294, negated_conjecture, (r2_hidden(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),k2_yellow_0(esk1_0,esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12894, plain, (v3_ordinal1(X1)|v1_xboole_0(X2)|~v1_card_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_12329, plain, (v3_ordinal1(X1)|~r2_hidden(X1,X3)|~v1_card_1(k2_xboole_0(X2,X3)))).#trainneg
cnf(i_0_12840, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_card_1(esk3_0))).#trainneg
cnf(i_0_12814, negated_conjecture, (v3_ordinal1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_card_1(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_12842, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_card_1(esk3_0))).#trainneg
cnf(i_0_12964, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),k2_yellow_0(X1,k1_tarski(X2)))|~m1_subset_1(k2_yellow_0(X1,k1_tarski(X2)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v3_orders_2(X1)|~v4_orders_2(X1))).#trainneg
cnf(i_0_13164, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk71_3(esk1_0,esk2_0,X1),esk3_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~v1_finset_1(X1)|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_13141, negated_conjecture, (v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk71_3(esk1_0,esk2_0,esk3_0),esk3_0))).#trainneg
cnf(i_0_13236, plain, (v2_struct_0(X1)|r1_orders_2(X1,esk72_3(X1,X2,k1_tarski(X2)),esk72_3(X1,X2,k1_tarski(X2)))|esk72_3(X1,X2,k1_tarski(X2))=X2|~l1_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X2)|~r1_orders_2(X1,X2,X2)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1))).#trainneg
cnf(i_0_13078, negated_conjecture, (m1_subset_1(esk71_3(esk1_0,esk2_0,X1),esk70_3(esk1_0,esk2_0,X1))|v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~v1_finset_1(X1)|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_12942, negated_conjecture, (esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0))=X1|v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|k1_tarski(X1)!=esk70_3(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13242, negated_conjecture, (r1_tarski(esk4_1(k1_zfmisc_1(esk4_1(k1_zfmisc_1(esk3_0)))),u1_struct_0(esk1_0))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(esk3_0)),esk2_0)|~v1_finset_1(k1_zfmisc_1(esk3_0))|~m1_subset_1(k1_zfmisc_1(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(esk4_1(k1_zfmisc_1(esk3_0)))),esk2_0)|~v1_finset_1(k1_zfmisc_1(esk4_1(k1_zfmisc_1(esk3_0))))|~m1_subset_1(k1_zfmisc_1(esk4_1(k1_zfmisc_1(esk3_0))),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2006, plain, (m1_subset_1(esk70_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),k1_zfmisc_1(k6_domain_1(u1_struct_0(X1),X3)))|r2_hidden(esk72_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),k6_domain_1(u1_struct_0(X1),X3))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|~r1_orders_2(X1,X3,X2)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_10792, negated_conjecture, (v4_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10848, negated_conjecture, (v1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10849, negated_conjecture, (v6_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v6_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11293, negated_conjecture, (r2_hidden(esk69_3(esk1_0,X1,k6_domain_1(esk3_0,X2)),k6_domain_1(esk3_0,X2))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k6_domain_1(esk3_0,X2)),X1)|v1_xboole_0(k6_domain_1(esk3_0,X2))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k6_domain_1(esk3_0,X2)),X1)|~v1_finset_1(k6_domain_1(esk3_0,X2))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,esk3_0))).#trainneg
cnf(i_0_10123, plain, (k2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,k1_tarski(X2)),X2))=k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,k1_tarski(X2))))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,k1_tarski(X2)),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(k1_tarski(k2_yellow_0(X1,k1_tarski(X2))),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_10336, negated_conjecture, (m1_subset_1(k7_domain_1(esk3_0,X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,esk3_0)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_10512, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk12_1(esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~m1_subset_1(esk12_1(esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9557, negated_conjecture, (r2_hidden(esk69_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12557, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~r1_tarski(X1,esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_13681, negated_conjecture, (esk4_1(k2_tarski(X1,k2_yellow_0(esk1_0,k1_tarski(X1))))=X1|esk4_1(k2_tarski(X1,k2_yellow_0(esk1_0,k1_tarski(X1))))=k2_yellow_0(esk1_0,k1_tarski(X1))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X1)),esk2_0)|~m1_subset_1(k2_tarski(X1,k2_yellow_0(esk1_0,k1_tarski(X1))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13597, negated_conjecture, (v3_ordinal1(esk2_0)|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2007, plain, (r1_tarski(X1,k6_domain_1(u1_struct_0(X2),X3))|r2_hidden(esk72_3(X2,X4,k6_domain_1(u1_struct_0(X2),X3)),k6_domain_1(u1_struct_0(X2),X3))|v1_xboole_0(k6_domain_1(u1_struct_0(X2),X3))|~r1_tarski(X1,esk70_3(X2,X4,k6_domain_1(u1_struct_0(X2),X3)))|~r1_orders_2(X2,X3,X4)|~v1_finset_1(k6_domain_1(u1_struct_0(X2),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X2),X3),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_13626, plain, (v3_ordinal1(X1)|~v1_card_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_13596, negated_conjecture, (v3_ordinal1(esk3_0)|~v1_card_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_13618, plain, (v3_ordinal1(X1)|~v1_card_1(k1_tarski(X1)))).#trainneg
cnf(i_0_13590, negated_conjecture, (v3_ordinal1(esk4_1(esk3_0))|v5_waybel_6(esk2_0,esk1_0)|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13397, negated_conjecture, (v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v5_finset_1(esk3_0))).#trainneg
cnf(i_0_13398, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v3_ordinal1(esk3_0))).#trainneg
cnf(i_0_13399, negated_conjecture, (v1_funct_1(esk69_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v4_funct_1(esk3_0))).#trainneg
cnf(i_0_13400, negated_conjecture, (v1_relat_1(esk69_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v4_funct_1(esk3_0))).#trainneg
cnf(i_0_13617, plain, (v3_ordinal1(k2_yellow_0(X1,X2))|v1_xboole_0(u1_struct_0(X1))|~v1_card_1(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_2010, plain, (r2_hidden(esk69_3(X1,X2,k1_tarski(X3)),k1_tarski(X3))|r1_orders_2(X1,esk72_3(X1,X2,k1_tarski(X3)),X2)|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X3)),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_13543, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_card_1(esk3_0))).#trainneg
cnf(i_0_13534, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_card_1(k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_13538, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_card_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_13545, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_card_1(esk3_0))).#trainneg
cnf(i_0_13511, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~v1_finset_1(X1)|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk71_3(esk1_0,esk2_0,X1),esk3_0))).#trainneg
cnf(i_0_13513, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),k2_yellow_0(esk1_0,esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2013, plain, (m1_subset_1(esk69_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k7_domain_1(u1_struct_0(X1),X3,X4))|r2_hidden(esk72_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k7_domain_1(u1_struct_0(X1),X3,X4))|v1_xboole_0(k7_domain_1(u1_struct_0(X1),X3,X4))|~r1_orders_2(X1,k12_lattice3(X1,X3,X4),X2)|~v1_finset_1(k7_domain_1(u1_struct_0(X1),X3,X4))|~m1_subset_1(k7_domain_1(u1_struct_0(X1),X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_13401, negated_conjecture, (esk69_3(esk1_0,esk2_0,esk3_0)=X1|v2_struct_0(esk1_0)|k1_tarski(X1)!=esk3_0)).#trainneg
cnf(i_0_13533, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_card_1(k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_13550, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13537, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_card_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_13436, negated_conjecture, (v1_xxreal_0(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk72_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(esk3_0))).#trainneg
cnf(i_0_13437, negated_conjecture, (v1_card_1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk72_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(esk3_0))).#trainneg
cnf(i_0_13747, plain, (k2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(k2_yellow_0(X1,X2))))=k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2)))|~r2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2)))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_11437, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k1_tarski(X2)),k1_tarski(X2))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k1_tarski(X2)),X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X2)),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,esk3_0))).#trainneg
cnf(i_0_10047, plain, (k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,k2_xboole_0(X3,X2))))=k2_yellow_0(X1,k2_xboole_0(X3,X2))|~m1_subset_1(k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,X3))).#trainneg
cnf(i_0_9575, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,X1)),esk70_3(esk1_0,esk2_0,X1))|r2_hidden(esk69_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,X1)),esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(X1)|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_11244, negated_conjecture, (esk4_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)))=k2_yellow_0(esk1_0,X2)|esk4_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)))=k2_yellow_0(esk1_0,X1)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X1,X2)),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,X2),esk3_0)|~m1_subset_1(k2_yellow_0(esk1_0,X1),esk3_0))).#trainneg
cnf(i_0_9978, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(k1_tarski(k2_yellow_0(X1,X3)),X2))|v1_xboole_0(X3)|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X3)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1))).#trainneg
cnf(i_0_11956, negated_conjecture, (v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk3_0))).#trainneg
cnf(i_0_2015, plain, (esk69_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4))=X5|r2_hidden(esk72_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k7_domain_1(u1_struct_0(X1),X3,X4))|v1_xboole_0(k7_domain_1(u1_struct_0(X1),X3,X4))|k7_domain_1(u1_struct_0(X1),X3,X4)!=k1_tarski(X5)|~r1_orders_2(X1,k12_lattice3(X1,X3,X4),X2)|~v1_finset_1(k7_domain_1(u1_struct_0(X1),X3,X4))|~m1_subset_1(k7_domain_1(u1_struct_0(X1),X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_13544, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_card_1(esk3_0))).#trainneg
cnf(i_0_13636, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(esk3_0))).#trainneg
cnf(i_0_13637, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(esk3_0))).#trainneg
cnf(i_0_10850, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_relat_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10851, negated_conjecture, (v5_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11706, negated_conjecture, (r1_tarski(esk70_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),k2_yellow_0(esk1_0,X1))|v1_xboole_0(X1)|~v1_finset_1(X1)|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_14117, negated_conjecture, (esk72_3(esk1_0,esk2_0,k2_tarski(esk2_0,esk2_0))=esk2_0|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_14204, plain, (r1_orders_2(X1,X2,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_14106, negated_conjecture, (esk72_3(esk1_0,esk2_0,k2_tarski(esk2_0,esk2_0))=esk2_0|esk69_3(esk1_0,esk2_0,k2_tarski(esk2_0,esk2_0))=esk2_0)).#trainneg
cnf(i_0_11193, negated_conjecture, (esk70_3(esk1_0,esk2_0,esk3_0)=X1|v2_struct_0(esk1_0)|k1_tarski(X1)!=k1_zfmisc_1(esk3_0))).#trainneg
cnf(i_0_13992, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v3_ordinal1(esk3_0))).#trainneg
cnf(i_0_13993, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v3_ordinal1(esk3_0))).#trainneg
cnf(i_0_14030, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(esk3_0))).#trainneg
cnf(i_0_14031, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(esk3_0))).#trainneg
cnf(i_0_14131, negated_conjecture, (esk4_1(k2_tarski(X1,X1))=X1|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13409, negated_conjecture, (esk69_3(esk1_0,esk2_0,esk3_0)=X1|v2_struct_0(esk1_0)|k1_tarski(X1)!=u1_struct_0(esk1_0))).#trainneg
cnf(i_0_14036, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(esk3_0))).#trainneg
cnf(i_0_14037, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(esk3_0))).#trainneg
cnf(i_0_10852, negated_conjecture, (v4_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12633, negated_conjecture, (v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v5_finset_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12634, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14198, plain, (k2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(k2_yellow_0(X1,X2))))=k2_yellow_0(X1,X2)|~r2_yellow_0(X1,X2)|~m1_subset_1(k1_tarski(k2_yellow_0(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_14177, negated_conjecture, (r2_hidden(esk72_3(X1,X2,k2_tarski(X2,X2)),k2_tarski(X2,X2))|v1_xboole_0(esk70_3(X1,X2,k2_tarski(X2,X2)))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(X1,X2,k2_tarski(X2,X2))),esk2_0)|~r1_orders_2(X1,esk4_1(esk70_3(X1,X2,k2_tarski(X2,X2))),X2)|~v1_finset_1(esk70_3(X1,X2,k2_tarski(X2,X2)))|~m1_subset_1(esk70_3(X1,X2,k2_tarski(X2,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk4_1(esk70_3(X1,X2,k2_tarski(X2,X2))),u1_struct_0(X1))|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_14307, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14213, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X3)|~r2_yellow_0(X1,k1_tarski(X2))|~r1_orders_2(X1,X2,X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_14156, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(k2_tarski(X2,X2),X3))|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,k2_tarski(X2,X2))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_13953, negated_conjecture, (esk4_1(k2_tarski(X2,k2_yellow_0(esk1_0,k1_tarski(X2))))=X2|v5_waybel_6(esk2_0,esk1_0)|k2_yellow_0(esk1_0,k1_tarski(X2))!=X2|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X2)),esk2_0)|~m1_subset_1(k2_tarski(X2,k2_yellow_0(esk1_0,k1_tarski(X2))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12161, negated_conjecture, (v1_xboole_0(esk4_1(k1_zfmisc_1(esk3_0)))|m1_subset_1(X1,u1_struct_0(esk1_0))|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X1,esk4_1(k1_zfmisc_1(esk3_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(esk3_0)),esk2_0)|~v1_finset_1(k1_zfmisc_1(esk3_0))|~m1_subset_1(k1_zfmisc_1(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_11578, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_xboole_0(X2,X3),k1_zfmisc_1(esk3_0))|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_11465, negated_conjecture, (esk70_3(esk1_0,esk2_0,esk3_0)=X1|v2_struct_0(esk1_0)|k1_tarski(X1)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12635, negated_conjecture, (v1_funct_1(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12636, negated_conjecture, (v1_relat_1(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v4_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13532, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_card_1(k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_13549, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13552, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14099, negated_conjecture, (esk4_1(k2_tarski(X1,X1))=X1|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14180, plain, (k2_yellow_0(X1,k1_tarski(X2))=X2|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_14396, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),k2_yellow_0(esk1_0,esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9382, negated_conjecture, (r1_tarski(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|v1_xboole_0(esk3_0)|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk3_0))).#trainneg
cnf(i_0_10594, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k1_tarski(X2)),k1_tarski(X2))|m1_subset_1(esk72_3(esk1_0,X1,k1_tarski(X2)),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_9484, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0)=X1|v2_struct_0(esk1_0)|k1_tarski(X1)!=esk3_0|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13765, negated_conjecture, (m1_subset_1(esk4_1(k1_zfmisc_1(esk4_1(k1_zfmisc_1(esk3_0)))),k1_zfmisc_1(u1_struct_0(esk1_0)))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(esk4_1(k1_zfmisc_1(esk3_0)))),esk2_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(esk3_0)),esk2_0)|~v1_finset_1(k1_zfmisc_1(esk4_1(k1_zfmisc_1(esk3_0))))|~v1_finset_1(k1_zfmisc_1(esk3_0))|~m1_subset_1(k1_zfmisc_1(esk4_1(k1_zfmisc_1(esk3_0))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_zfmisc_1(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_13912, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_hidden(X1,k7_domain_1(esk3_0,X2,X3))|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X2,esk3_0))).#trainneg
cnf(i_0_9868, negated_conjecture, (X1=X2|k2_xboole_0(u1_struct_0(esk1_0),X3)!=k1_tarski(X2)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_9869, negated_conjecture, (X1=X2|k2_xboole_0(X3,u1_struct_0(esk1_0))!=k1_tarski(X2)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_11435, negated_conjecture, (m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_tarski(X1))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_14060, negated_conjecture, (v1_xxreal_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14061, negated_conjecture, (v1_card_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14062, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14063, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10415, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10032, plain, (k2_yellow_0(X1,k1_tarski(X2))=X2|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X2))|v2_struct_0(X1)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_11871, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9994, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(k1_tarski(k2_yellow_0(X1,X2)),X3))|v1_xboole_0(X2)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1))).#trainneg
cnf(i_0_14805, negated_conjecture, (v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v5_finset_1(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_14806, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v3_ordinal1(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_14807, negated_conjecture, (v1_funct_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v4_funct_1(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_14808, negated_conjecture, (v1_relat_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v4_funct_1(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_14448, negated_conjecture, (r1_orders_2(esk1_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14444, plain, (r1_orders_2(esk42_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk42_0)))).#trainneg
cnf(i_0_14445, plain, (r1_orders_2(esk24_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk24_0)))).#trainneg
cnf(i_0_14809, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_card_1(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10425, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10427, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_14661, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X3)|~r1_orders_2(X1,X2,X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_14644, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,k2_tarski(esk2_0,esk2_0)),k2_tarski(esk2_0,esk2_0))|v1_xboole_0(esk70_3(esk1_0,esk2_0,k2_tarski(esk2_0,esk2_0)))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,k2_tarski(esk2_0,esk2_0))),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,k2_tarski(esk2_0,esk2_0)))|~m1_subset_1(esk70_3(esk1_0,esk2_0,k2_tarski(esk2_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk4_1(esk70_3(esk1_0,esk2_0,k2_tarski(esk2_0,esk2_0))),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk2_0,esk2_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_14658, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13769, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|v5_waybel_6(esk2_0,esk1_0)|~r1_tarski(X1,esk4_1(k1_zfmisc_1(esk4_1(k1_zfmisc_1(esk3_0)))))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(esk4_1(k1_zfmisc_1(esk3_0)))),esk2_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(esk3_0)),esk2_0)|~v1_finset_1(k1_zfmisc_1(esk4_1(k1_zfmisc_1(esk3_0))))|~v1_finset_1(k1_zfmisc_1(esk3_0))|~m1_subset_1(k1_zfmisc_1(esk4_1(k1_zfmisc_1(esk3_0))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_zfmisc_1(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2024, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X3,X4)|~r2_hidden(esk7_3(X3,X4,k2_xboole_0(X1,X2)),X4)|~r2_hidden(esk7_3(X3,X4,k2_xboole_0(X1,X2)),X2))).#trainneg
cnf(i_0_14446, plain, (r1_orders_2(esk23_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk23_0)))).#trainneg
cnf(i_0_14447, plain, (r1_orders_2(esk14_0,X1,X1)|~m1_subset_1(X1,u1_struct_0(esk14_0)))).#trainneg
cnf(i_0_15079, negated_conjecture, (v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v5_finset_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_15080, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v3_ordinal1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_2025, plain, (k2_xboole_0(X1,X2)=k2_xboole_0(X3,X4)|~r2_hidden(esk7_3(X3,X4,k2_xboole_0(X1,X2)),X3)|~r2_hidden(esk7_3(X3,X4,k2_xboole_0(X1,X2)),X2))).#trainneg
cnf(i_0_15081, negated_conjecture, (v1_funct_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v4_funct_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_15082, negated_conjecture, (v1_relat_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v4_funct_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_15083, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_card_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_15092, negated_conjecture, (v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v5_finset_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_15093, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v3_ordinal1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_14994, negated_conjecture, (r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|r2_hidden(esk69_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|v1_xboole_0(X1)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14609, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2))),X3)|~r2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2)))|~r2_yellow_0(X1,X2)|~r1_orders_2(X1,k2_yellow_0(X1,X2),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(k2_yellow_0(X1,X2)),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_13605, plain, (v3_ordinal1(k1_tarski(X1))|v1_xboole_0(X2)|~v1_card_1(k1_zfmisc_1(X2))|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_2028, plain, (k2_xboole_0(X1,X2)=k1_tarski(X3)|esk5_2(X3,k2_xboole_0(X1,X2))!=X3|~r2_hidden(esk5_2(X3,k2_xboole_0(X1,X2)),X2))).#trainneg
cnf(i_0_12520, plain, (v3_ordinal1(X1)|~r2_hidden(X1,X2)|~m1_subset_1(k2_xboole_0(X2,X3),k4_ordinal1))).#trainneg
cnf(i_0_15094, negated_conjecture, (v1_funct_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v4_funct_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_15095, negated_conjecture, (v1_relat_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v4_funct_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_15096, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_card_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10632, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_relat_1(esk3_0)|~v1_funct_1(esk3_0))).#trainneg
cnf(i_0_14664, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(k2_tarski(X2,X2),X3))|~r2_yellow_0(X1,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_11364, negated_conjecture, (r1_tarski(k1_tarski(X1),u1_struct_0(esk1_0))|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_13634, plain, (v3_ordinal1(X1)|~r2_hidden(X1,X3)|~m1_subset_1(k2_xboole_0(X2,X3),k4_ordinal1))).#trainneg
cnf(i_0_10775, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_13535, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_card_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_13536, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_card_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_14064, negated_conjecture, (v1_xxreal_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_14065, negated_conjecture, (v1_card_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_12479, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))))).#trainneg
cnf(i_0_2041, plain, (r1_tarski(esk70_3(X1,X2,k1_tarski(X3)),k1_tarski(X3))|r1_orders_2(X1,esk72_3(X1,X2,k1_tarski(X3)),X2)|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X3)),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_10382, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~r1_tarski(X1,k1_tarski(X2))|~m1_subset_1(X2,esk3_0))).#trainneg
cnf(i_0_11805, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk3_0),k1_numbers)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10601, negated_conjecture, (v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v5_finset_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10602, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v3_ordinal1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_15285, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_15286, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_15395, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2))),X3)|~r2_yellow_0(X1,X2)|~r1_orders_2(X1,k2_yellow_0(X1,X2),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(k1_tarski(k2_yellow_0(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_15517, plain, (k2_yellow_0(X1,k2_xboole_0(k2_tarski(X2,X2),X3))=k12_lattice3(X1,k2_yellow_0(X1,X3),X2)|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~r2_yellow_0(X1,X3)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_15637, plain, (r1_tarski(esk70_3(X1,X2,k1_tarski(X3)),k1_tarski(X3))|r1_orders_2(X1,esk72_3(X1,X2,k1_tarski(X3)),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,X3,X2)|~v5_waybel_6(X2,X1)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_14961, plain, (k12_lattice3(X1,X2,X2)=X2|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_15268, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_15524, plain, (k2_yellow_0(X1,k2_xboole_0(k2_tarski(X2,X2),X3))=k2_yellow_0(X1,k2_tarski(X2,k2_yellow_0(X1,X3)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X3)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_15599, negated_conjecture, (r1_tarski(k1_tarski(X1),u1_struct_0(esk1_0))|v1_xboole_0(esk4_1(k1_zfmisc_1(esk3_0)))|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X1,esk4_1(k1_zfmisc_1(esk3_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(esk3_0)),esk2_0)|~v1_finset_1(k1_zfmisc_1(esk3_0))|~m1_subset_1(k1_zfmisc_1(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_15409, plain, (v3_ordinal1(X1)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k4_ordinal1))).#trainneg
cnf(i_0_2043, plain, (m1_subset_1(esk70_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k1_zfmisc_1(k7_domain_1(u1_struct_0(X1),X3,X4)))|r2_hidden(esk72_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k7_domain_1(u1_struct_0(X1),X3,X4))|v1_xboole_0(k7_domain_1(u1_struct_0(X1),X3,X4))|~r1_orders_2(X1,k12_lattice3(X1,X3,X4),X2)|~v1_finset_1(k7_domain_1(u1_struct_0(X1),X3,X4))|~m1_subset_1(k7_domain_1(u1_struct_0(X1),X3,X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_10603, negated_conjecture, (v1_funct_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v4_funct_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10604, negated_conjecture, (v1_relat_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v4_funct_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10793, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_relat_1(u1_struct_0(esk1_0))|~v1_funct_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13547, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_card_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_13161, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v3_ordinal1(esk3_0))).#trainneg
cnf(i_0_15334, negated_conjecture, (r2_hidden(esk69_3(esk1_0,X1,k1_tarski(X1)),k1_tarski(X1))|r2_hidden(esk72_3(esk1_0,X1,k1_tarski(X1)),k1_tarski(X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_15627, plain, (r1_tarski(esk70_3(X1,X2,k1_tarski(X3)),k1_tarski(X3))|r1_orders_2(X1,esk72_3(X1,X2,k1_tarski(X3)),X2)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_15695, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2044, plain, (r1_tarski(X1,k7_domain_1(u1_struct_0(X2),X3,X4))|r2_hidden(esk72_3(X2,X5,k7_domain_1(u1_struct_0(X2),X3,X4)),k7_domain_1(u1_struct_0(X2),X3,X4))|v1_xboole_0(k7_domain_1(u1_struct_0(X2),X3,X4))|~r1_tarski(X1,esk70_3(X2,X5,k7_domain_1(u1_struct_0(X2),X3,X4)))|~r1_orders_2(X2,k12_lattice3(X2,X3,X4),X5)|~v1_finset_1(k7_domain_1(u1_struct_0(X2),X3,X4))|~m1_subset_1(k7_domain_1(u1_struct_0(X2),X3,X4),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X5,u1_struct_0(X2))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_16228, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_tarski(X1),k4_ordinal1))).#trainneg
cnf(i_0_16231, plain, (v3_ordinal1(X1)|~m1_subset_1(k2_tarski(X2,X1),k4_ordinal1))).#trainneg
cnf(i_0_16232, plain, (v3_ordinal1(X1)|~m1_subset_1(k2_tarski(X1,X2),k4_ordinal1))).#trainneg
cnf(i_0_16147, negated_conjecture, (v3_ordinal1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_16178, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_16148, negated_conjecture, (v3_ordinal1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~m1_subset_1(esk3_0,k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_16149, negated_conjecture, (v3_ordinal1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk3_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_16311, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,k1_tarski(X1)),k1_tarski(X1))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,k1_tarski(X1)),esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_15916, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_15677, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_15624, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))))).#trainneg
cnf(i_0_2047, plain, (r1_tarski(esk70_3(X1,X2,k2_tarski(X4,X3)),k2_tarski(X4,X3))|r2_hidden(esk72_3(X1,X2,k2_tarski(X4,X3)),k2_tarski(X4,X3))|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k12_lattice3(X1,X3,X4),X2)|~m1_subset_1(k2_tarski(X4,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_15607, negated_conjecture, (r1_tarski(k1_tarski(X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_16150, negated_conjecture, (v3_ordinal1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_16179, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_16180, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_13162, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v3_ordinal1(esk3_0))).#trainneg
cnf(i_0_15679, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_15697, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2))),X3)|v1_xboole_0(X2)|~r1_orders_2(X1,k2_yellow_0(X1,X2),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(k1_tarski(k2_yellow_0(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_15911, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))=X1|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|k1_tarski(X1)!=esk70_3(esk1_0,esk2_0,esk3_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16233, plain, (v3_ordinal1(X1)|v1_xboole_0(X2)|~m1_subset_1(X2,k4_ordinal1)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_13640, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(esk3_0))).#trainneg
cnf(i_0_13641, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(esk3_0))).#trainneg
cnf(i_0_14383, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(esk3_0))).#trainneg
cnf(i_0_14384, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(esk3_0))).#trainneg
cnf(i_0_15271, negated_conjecture, (r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(X1,X2,X3)),esk70_3(X1,X2,X3)),esk70_3(X1,X2,X3))|v1_xboole_0(esk70_3(X1,X2,X3))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(X1,X2,X3)),esk70_3(X1,X2,X3)),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(X1,X2,X3)),esk70_3(X1,X2,X3)),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(esk70_3(X1,X2,X3))|~m1_subset_1(k2_yellow_0(esk1_0,esk70_3(X1,X2,X3)),u1_struct_0(esk1_0))|~m1_subset_1(esk70_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_14614, plain, (k2_yellow_0(X1,k2_xboole_0(k6_domain_1(u1_struct_0(X1),X2),k1_tarski(X2)))=X2|v2_struct_0(X1)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_16299, plain, (r1_tarski(esk70_3(X1,esk9_2(X1,X2),k1_tarski(X3)),k1_tarski(X3))|r1_orders_2(X1,esk72_3(X1,esk9_2(X1,X2),k1_tarski(X3)),esk9_2(X1,X2))|v2_struct_0(X1)|v5_waybel_6(X2,X1)|~r1_orders_2(X1,X3,esk9_2(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_15681, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_15120, plain, (r2_hidden(esk69_3(X1,X2,k1_tarski(X3)),k1_tarski(X3))|r1_orders_2(X1,esk72_3(X1,X2,k1_tarski(X3)),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,X3,X2)|~v5_waybel_6(X2,X1)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_16227, negated_conjecture, (v3_ordinal1(esk4_1(X1))|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X1,k4_ordinal1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2073, plain, (r2_hidden(esk69_3(X1,X2,u1_struct_0(X1)),u1_struct_0(X1))|r1_orders_2(X1,esk72_3(X1,X2,u1_struct_0(X1)),X2)|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,u1_struct_0(X1)),X2)|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_14512, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14513, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v3_ordinal1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14691, negated_conjecture, (v1_xxreal_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14692, negated_conjecture, (v1_card_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14693, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14995, negated_conjecture, (r1_tarski(esk70_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|v1_xboole_0(X1)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16376, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16636, negated_conjecture, (v3_ordinal1(esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1))).#trainneg
cnf(i_0_2074, plain, (r1_tarski(esk70_3(X1,X2,u1_struct_0(X1)),u1_struct_0(X1))|r1_orders_2(X1,esk72_3(X1,X2,u1_struct_0(X1)),X2)|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,u1_struct_0(X1)),X2)|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_16635, negated_conjecture, (v3_ordinal1(esk3_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k4_ordinal1))).#trainneg
cnf(i_0_16666, plain, (v3_ordinal1(X1)|~m1_subset_1(k1_zfmisc_1(X1),k4_ordinal1))).#trainneg
cnf(i_0_16629, negated_conjecture, (v3_ordinal1(esk4_1(esk3_0))|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1))).#trainneg
cnf(i_0_16576, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_16567, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(esk3_0),k4_ordinal1))).#trainneg
cnf(i_0_16571, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k4_ordinal1))).#trainneg
cnf(i_0_16587, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0))).#trainneg
cnf(i_0_16866, negated_conjecture, (r1_tarski(X1,X2)|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,X2),X2),X2)|v1_xboole_0(X2)|~r1_tarski(X1,esk70_3(esk1_0,k2_yellow_0(esk1_0,X2),X2))|~v1_finset_1(X2)|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_16850, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16700, negated_conjecture, (esk4_1(k2_tarski(X1,X1))=X1|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X1),k6_domain_1(u1_struct_0(esk1_0),X1)))|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_16507, negated_conjecture, (v2_struct_0(esk1_0)|v1_xboole_0(X1)|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(k2_yellow_0(esk1_0,X1))),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(k2_yellow_0(esk1_0,X1))),esk3_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,X1)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_16354, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,k1_tarski(X1)),k1_zfmisc_1(k1_tarski(X1)))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,k1_tarski(X1)),esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16870, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k1_tarski(X1)),k1_tarski(X1))|r2_hidden(esk72_3(esk1_0,X1,k1_tarski(X1)),k1_tarski(X1))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_16578, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_16566, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(k1_zfmisc_1(esk3_0),k4_ordinal1))).#trainneg
cnf(i_0_16586, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1))).#trainneg
cnf(i_0_16579, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~m1_subset_1(esk3_0,k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_16556, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,k1_tarski(X2)))),X3)|~r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(k1_tarski(k2_yellow_0(X1,k1_tarski(X2))),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_16804, plain, (r1_tarski(esk70_3(X1,esk9_2(X1,X2),k1_tarski(esk9_2(X1,X2))),k1_tarski(esk9_2(X1,X2)))|r1_orders_2(X1,esk72_3(X1,esk9_2(X1,X2),k1_tarski(esk9_2(X1,X2))),esk9_2(X1,X2))|v5_waybel_6(X2,X1)|v2_struct_0(X1)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_16358, negated_conjecture, (r1_tarski(X1,k1_tarski(X2))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,k1_tarski(X2)),esk2_0)|~r1_tarski(X1,esk70_3(esk1_0,esk2_0,k1_tarski(X2)))|~r1_orders_2(esk1_0,X2,esk2_0)|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14919, plain, (r2_hidden(esk69_3(X1,X2,k1_tarski(X2)),k1_tarski(X2))|r2_hidden(esk72_3(X1,X2,k1_tarski(X2)),k1_tarski(X2))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_16137, plain, (k2_yellow_0(X1,k2_xboole_0(k2_tarski(X2,X2),X3))=k2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X3),X2))|~r2_yellow_0(X1,X3)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_15328, negated_conjecture, (esk69_3(esk1_0,k2_yellow_0(esk1_0,X1),X1)=X2|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|v1_xboole_0(X1)|X1!=k1_tarski(X2)|~v1_finset_1(X1)|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_16657, plain, (v3_ordinal1(k2_yellow_0(X1,X2))|v1_xboole_0(u1_struct_0(X1))|~m1_subset_1(u1_struct_0(X1),k4_ordinal1)|~l1_orders_2(X1))).#trainneg
cnf(i_0_2217, plain, (k2_yellow_0(X1,X2)=X3|v1_xboole_0(u1_struct_0(X1))|u1_struct_0(X1)!=k1_tarski(X3)|~l1_orders_2(X1))).#trainneg
cnf(i_0_16346, negated_conjecture, (v1_xxreal_0(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk72_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_16347, negated_conjecture, (v1_card_1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk72_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_16348, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_16349, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_14991, negated_conjecture, (r1_orders_2(esk1_0,X1,k2_yellow_0(esk1_0,k2_tarski(X1,X2)))|r1_orders_2(esk1_0,X2,k2_yellow_0(esk1_0,k2_tarski(X1,X2)))|~v5_waybel_6(k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk1_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(X1,X2)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14555, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(X2,k1_tarski(k2_yellow_0(X1,X2))),X3))|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(k2_yellow_0(X1,X2))))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~r2_yellow_0(X1,X2)|~m1_subset_1(k1_tarski(k2_yellow_0(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1))).#trainneg
cnf(i_0_15105, negated_conjecture, (v2_struct_0(esk1_0)|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(X1)),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(X1)),esk3_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16308, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk4_1(esk3_0),k1_tarski(X1)),k1_tarski(X1))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk4_1(esk3_0),k1_tarski(X1)),esk4_1(esk3_0))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,X1,esk4_1(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11243, negated_conjecture, (esk4_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)))=k2_yellow_0(esk1_0,X2)|esk4_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)))=k2_yellow_0(esk1_0,X1)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,X1)|~r2_yellow_0(esk1_0,X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X2,X1)),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,X2),esk3_0)|~m1_subset_1(k2_yellow_0(esk1_0,X1),esk3_0))).#trainneg
cnf(i_0_5939, negated_conjecture, (esk4_1(X1)=X2|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|k2_xboole_0(k1_numbers,X3)!=k1_tarski(X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_11210, negated_conjecture, (m1_subset_1(X1,esk3_0)|v2_struct_0(esk1_0)|~r2_hidden(X1,esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_16920, negated_conjecture, (v1_xxreal_0(esk4_1(esk3_0))|v5_waybel_6(esk2_0,esk1_0)|~v1_finset_1(esk4_1(esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1))).#trainneg
cnf(i_0_16921, negated_conjecture, (v1_card_1(esk4_1(esk3_0))|v5_waybel_6(esk2_0,esk1_0)|~v1_finset_1(esk4_1(esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1))).#trainneg
cnf(i_0_16922, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_16923, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_17047, negated_conjecture, (v1_finset_1(esk70_3(esk1_0,esk2_0,k1_tarski(X1)))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,k1_tarski(X1)),esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16993, negated_conjecture, (esk4_1(k2_tarski(X1,X1))=X1|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16984, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|m1_subset_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16990, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16819, negated_conjecture, (v1_xxreal_0(esk4_1(X1))|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|~v1_finset_1(esk4_1(X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_6206, negated_conjecture, (esk4_1(X1)=X2|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|k2_xboole_0(X3,k1_numbers)!=k1_tarski(X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_5932, plain, (esk12_1(X1)=X2|v1_xboole_0(X1)|k2_xboole_0(X1,X3)!=k1_tarski(X2))).#trainneg
cnf(i_0_6199, plain, (esk12_1(X1)=X2|v1_xboole_0(X1)|k2_xboole_0(X3,X1)!=k1_tarski(X2))).#trainneg
cnf(i_0_17067, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_17068, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_17149, plain, (v1_xboole_0(X1)|m1_subset_1(esk69_3(X2,X3,k1_tarski(X3)),X1)|r2_hidden(esk72_3(X2,X3,k1_tarski(X3)),k1_tarski(X3))|~m1_subset_1(X3,X1)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_15902, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10907, negated_conjecture, (r2_hidden(esk69_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|v1_xboole_0(X1)|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),u1_struct_0(esk1_0))|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_16820, negated_conjecture, (v1_card_1(esk4_1(X1))|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|~v1_finset_1(esk4_1(X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_2108, plain, (r1_orders_2(X1,k12_lattice3(X1,X2,X3),X4)|r1_orders_2(X1,k2_yellow_0(X1,X5),X4)|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X5,k7_domain_1(u1_struct_0(X1),X2,X3))),X4)|~v5_waybel_6(X4,X1)|~m1_subset_1(k12_lattice3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X5),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,X5)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6669, negated_conjecture, (v1_relat_1(k2_tarski(X1,esk2_0))|~v1_relat_1(u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6828, negated_conjecture, (v5_finset_1(k2_tarski(X1,esk2_0))|~v5_finset_1(u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6954, negated_conjecture, (v4_funct_1(k2_tarski(X1,esk2_0))|~v4_funct_1(u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16565, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_zfmisc_1(esk3_0),k4_ordinal1))).#trainneg
cnf(i_0_16585, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1))).#trainneg
cnf(i_0_16588, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1))).#trainneg
cnf(i_0_17069, negated_conjecture, (v1_xxreal_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k1_zfmisc_1(esk3_0),k4_ordinal1))).#trainneg
cnf(i_0_17070, negated_conjecture, (v1_card_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k1_zfmisc_1(esk3_0),k4_ordinal1))).#trainneg
cnf(i_0_10508, negated_conjecture, (esk69_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)),k2_tarski(esk69_3(esk1_0,esk2_0,esk3_0),esk2_0))=esk2_0|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)),k2_tarski(esk69_3(esk1_0,esk2_0,esk3_0),esk2_0))=esk2_0|esk69_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)),k2_tarski(esk69_3(esk1_0,esk2_0,esk3_0),esk2_0))=esk69_3(esk1_0,esk2_0,esk3_0)|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)),k2_tarski(esk69_3(esk1_0,esk2_0,esk3_0),esk2_0))=esk69_3(esk1_0,esk2_0,esk3_0)|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11689, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k2_tarski(k2_yellow_0(esk1_0,X2),k2_yellow_0(esk1_0,X3))),k2_tarski(k2_yellow_0(esk1_0,X2),k2_yellow_0(esk1_0,X3)))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k2_tarski(k2_yellow_0(esk1_0,X2),k2_yellow_0(esk1_0,X3))),X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X3,X2)),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(k2_yellow_0(esk1_0,X2),k2_yellow_0(esk1_0,X3)),k1_zfmisc_1(esk3_0))|~r2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,X3))).#trainneg
cnf(i_0_14965, plain, (r2_hidden(esk69_3(X1,X2,k1_tarski(X3)),k1_tarski(X3))|r1_orders_2(X1,esk72_3(X1,X2,k1_tarski(X3)),X2)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_12500, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k1_tarski(X2)))|~m1_subset_1(X2,esk3_0))).#trainneg
cnf(i_0_13416, plain, (X1=k1_tarski(X2)|esk5_2(X2,X1)!=X2|~r2_hidden(esk5_2(X2,X1),k1_xboole_0))).#trainneg
cnf(i_0_17071, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1))).#trainneg
cnf(i_0_17072, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1))).#trainneg
cnf(i_0_16568, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k4_ordinal1))).#trainneg
cnf(i_0_16569, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k4_ordinal1))).#trainneg
cnf(i_0_9576, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,X1)),esk70_3(esk1_0,esk2_0,X1))|r2_hidden(esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,X1)),esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(X1)|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_14851, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(k1_tarski(k2_yellow_0(X1,X2)),X3))|v1_xboole_0(X2)|~r2_yellow_0(X1,X3)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(k2_yellow_0(X1,X2)),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_17989, negated_conjecture, (r2_hidden(esk69_3(esk1_0,esk2_0,k1_tarski(X1)),k1_tarski(X1))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,k1_tarski(X1)),esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_17619, negated_conjecture, (v1_finset_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|v2_struct_0(esk1_0)|~v1_finset_1(u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_17620, negated_conjecture, (v1_zfmisc_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|v2_struct_0(esk1_0)|~v1_zfmisc_1(u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_17621, negated_conjecture, (v6_ordinal1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|v2_struct_0(esk1_0)|~v6_ordinal1(u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_17622, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|v2_struct_0(esk1_0)|~v1_relat_1(u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_17977, plain, (r2_hidden(esk69_3(X1,esk9_2(X1,X2),k1_tarski(X3)),k1_tarski(X3))|r1_orders_2(X1,esk72_3(X1,esk9_2(X1,X2),k1_tarski(X3)),esk9_2(X1,X2))|v2_struct_0(X1)|v5_waybel_6(X2,X1)|~r1_orders_2(X1,X3,esk9_2(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_17978, plain, (r2_hidden(esk69_3(X1,esk8_2(X1,X2),k1_tarski(X3)),k1_tarski(X3))|r1_orders_2(X1,esk72_3(X1,esk8_2(X1,X2),k1_tarski(X3)),esk8_2(X1,X2))|v2_struct_0(X1)|v5_waybel_6(X2,X1)|~r1_orders_2(X1,X3,esk8_2(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_17986, negated_conjecture, (r2_hidden(esk69_3(esk1_0,esk4_1(esk3_0),k1_tarski(X1)),k1_tarski(X1))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk4_1(esk3_0),k1_tarski(X1)),esk4_1(esk3_0))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,X1,esk4_1(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_18041, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk3_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2114, plain, (k2_yellow_0(X1,k2_xboole_0(k6_domain_1(u1_struct_0(X1),X2),k7_domain_1(u1_struct_0(X1),X3,X4)))=k12_lattice3(X1,X2,k12_lattice3(X1,X3,X4))|v2_struct_0(X1)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_17623, negated_conjecture, (v5_finset_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|v2_struct_0(esk1_0)|~v5_finset_1(u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_17624, negated_conjecture, (v4_funct_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|v2_struct_0(esk1_0)|~v4_funct_1(u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_17795, plain, (r1_orders_2(X1,k12_lattice3(X1,X2,X3),X4)|r1_orders_2(X1,k2_yellow_0(X1,X5),X4)|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,k2_tarski(X2,X3))|~r2_yellow_0(X1,X5)|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X5,k2_tarski(X2,X3))),X4)|~v5_waybel_6(X4,X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_17993, plain, (r2_hidden(esk69_3(X1,k2_yellow_0(X1,X2),k1_tarski(X3)),k1_tarski(X3))|r1_orders_2(X1,esk72_3(X1,k2_yellow_0(X1,X2),k1_tarski(X3)),k2_yellow_0(X1,X2))|~r1_orders_2(X1,X3,k2_yellow_0(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_17800, plain, (r1_orders_2(X1,k12_lattice3(X1,X2,X3),X4)|r1_orders_2(X1,k2_yellow_0(X1,X5),X4)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,X5)|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X2,X3),X5)),X4)|~v5_waybel_6(X4,X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11688, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k2_tarski(k2_yellow_0(esk1_0,X2),k2_yellow_0(esk1_0,X3))),k2_tarski(k2_yellow_0(esk1_0,X2),k2_yellow_0(esk1_0,X3)))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k2_tarski(k2_yellow_0(esk1_0,X2),k2_yellow_0(esk1_0,X3))),X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X2,X3)),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(k2_yellow_0(esk1_0,X2),k2_yellow_0(esk1_0,X3)),k1_zfmisc_1(esk3_0))|~r2_yellow_0(esk1_0,X3)|~r2_yellow_0(esk1_0,X2))).#trainneg
cnf(i_0_17737, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),u1_struct_0(esk1_0))|v1_xboole_0(X1)|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0))|~v1_finset_1(X1))).#trainneg
cnf(i_0_11237, negated_conjecture, (r2_hidden(esk69_3(esk1_0,X1,k2_tarski(X2,X3)),k2_tarski(X2,X3))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k2_tarski(X2,X3)),X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X2,X3)),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X3,esk3_0)|~m1_subset_1(X2,esk3_0))).#trainneg
cnf(i_0_16863, negated_conjecture, (r1_tarski(esk70_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),u1_struct_0(esk1_0))|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk3_0))|~v1_finset_1(X1)|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2203, negated_conjecture, (esk4_1(X1)=X2|v1_xboole_0(u1_struct_0(esk1_0))|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|k1_tarski(X2)!=u1_struct_0(esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2115, plain, (k2_yellow_0(X1,k2_xboole_0(k1_tarski(X2),k7_domain_1(u1_struct_0(X1),X3,X4)))=k12_lattice3(X1,X2,k12_lattice3(X1,X3,X4))|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))|~r2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_17625, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|v2_struct_0(esk1_0)|~v1_relat_1(u1_struct_0(esk1_0))|~v1_funct_1(u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14694, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14695, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_18172, negated_conjecture, (r2_hidden(esk72_3(X1,X2,X3),X3)|v1_xboole_0(X3)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,k1_tarski(esk69_3(X1,X2,X3))),esk2_0)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~r1_orders_2(X1,esk69_3(esk1_0,esk2_0,k1_tarski(esk69_3(X1,X2,X3))),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,k1_tarski(esk69_3(X1,X2,X3))),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(esk1_0,esk69_3(X1,X2,X3),esk2_0)|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_18780, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X2)),X3)|r1_orders_2(X1,k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X4,X5)),X3)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X4,X5))|~r2_yellow_0(X1,k1_tarski(X2))|~r1_orders_2(X1,k12_lattice3(X1,X2,k12_lattice3(X1,X4,X5)),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_18188, negated_conjecture, (r2_hidden(esk69_3(esk1_0,esk9_2(esk1_0,X1),k1_tarski(esk9_2(esk1_0,X1))),k1_tarski(esk9_2(esk1_0,X1)))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk9_2(esk1_0,X1),k1_tarski(esk9_2(esk1_0,X1))),esk9_2(esk1_0,X1))|v5_waybel_6(X1,esk1_0)|~m1_subset_1(esk9_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_18439, negated_conjecture, (r2_hidden(esk69_3(esk1_0,k2_yellow_0(esk1_0,X1),k1_tarski(k2_yellow_0(esk1_0,X1))),k1_tarski(k2_yellow_0(esk1_0,X1)))|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),k1_tarski(k2_yellow_0(esk1_0,X1))),k2_yellow_0(esk1_0,X1))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_18233, negated_conjecture, (r2_hidden(esk69_3(esk1_0,esk4_1(esk3_0),k1_tarski(esk4_1(esk3_0))),k1_tarski(esk4_1(esk3_0)))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk4_1(esk3_0),k1_tarski(esk4_1(esk3_0))),esk4_1(esk3_0))|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(esk4_1(esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_18190, plain, (r2_hidden(esk69_3(X1,esk9_2(X1,X2),k1_tarski(esk9_2(X1,X2))),k1_tarski(esk9_2(X1,X2)))|r1_orders_2(X1,esk72_3(X1,esk9_2(X1,X2),k1_tarski(esk9_2(X1,X2))),esk9_2(X1,X2))|v5_waybel_6(X2,X1)|v2_struct_0(X1)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_18784, plain, (r1_orders_2(X1,k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3)),X4)|r1_orders_2(X1,X5,X4)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r1_orders_2(X1,k12_lattice3(X1,X5,k12_lattice3(X1,X2,X3)),X4)|~v5_waybel_6(X4,X1)|~m1_subset_1(k1_tarski(X5),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k1_tarski(X5))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_18207, negated_conjecture, (r2_hidden(esk69_3(esk1_0,esk8_2(esk1_0,X1),k1_tarski(esk8_2(esk1_0,X1))),k1_tarski(esk8_2(esk1_0,X1)))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk8_2(esk1_0,X1),k1_tarski(esk8_2(esk1_0,X1))),esk8_2(esk1_0,X1))|v5_waybel_6(X1,esk1_0)|~m1_subset_1(esk8_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_18209, plain, (r2_hidden(esk69_3(X1,esk8_2(X1,X2),k1_tarski(esk8_2(X1,X2))),k1_tarski(esk8_2(X1,X2)))|r1_orders_2(X1,esk72_3(X1,esk8_2(X1,X2),k1_tarski(esk8_2(X1,X2))),esk8_2(X1,X2))|v5_waybel_6(X2,X1)|v2_struct_0(X1)|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_17678, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_18671, negated_conjecture, (r1_tarski(esk70_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),u1_struct_0(esk1_0))|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|v1_xboole_0(X1)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_11482, negated_conjecture, (m1_subset_1(X1,u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|~r2_hidden(X1,esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_2116, plain, (k12_lattice3(X1,k12_lattice3(X1,X2,X3),k12_lattice3(X1,X4,X5))=k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X2,X3),k7_domain_1(u1_struct_0(X1),X4,X5)))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X4,X5))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_14696, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16151, negated_conjecture, (v3_ordinal1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_16352, negated_conjecture, (v1_xxreal_0(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk3_0))|~v1_finset_1(esk72_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_16353, negated_conjecture, (v1_card_1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk3_0))|~v1_finset_1(esk72_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_16481, negated_conjecture, (v1_xxreal_0(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk72_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_11693, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k2_tarski(X2,k2_yellow_0(esk1_0,X3))),k2_tarski(X2,k2_yellow_0(esk1_0,X3)))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k2_tarski(X2,k2_yellow_0(esk1_0,X3))),X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X3,k1_tarski(X2))),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(X2,k2_yellow_0(esk1_0,X3)),k1_zfmisc_1(esk3_0))|~r2_yellow_0(esk1_0,X3)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16300, plain, (r1_tarski(esk70_3(X1,esk8_2(X1,X2),k1_tarski(X3)),k1_tarski(X3))|r1_orders_2(X1,esk72_3(X1,esk8_2(X1,X2),k1_tarski(X3)),esk8_2(X1,X2))|v2_struct_0(X1)|v5_waybel_6(X2,X1)|~r1_orders_2(X1,X3,esk8_2(X1,X2))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_14861, negated_conjecture, (esk4_1(k2_tarski(X1,X1))=X1|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11436, negated_conjecture, (r2_hidden(esk69_3(esk1_0,X1,k1_tarski(X2)),k1_tarski(X2))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k1_tarski(X2)),X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X2)),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,esk3_0))).#trainneg
cnf(i_0_10908, negated_conjecture, (r1_tarski(esk70_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|v1_xboole_0(X1)|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),u1_struct_0(esk1_0))|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_12284, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,esk4_1(k1_zfmisc_1(esk3_0))),esk4_1(k1_zfmisc_1(esk3_0)))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,esk4_1(k1_zfmisc_1(esk3_0))),X1)|v1_xboole_0(esk4_1(k1_zfmisc_1(esk3_0)))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk4_1(k1_zfmisc_1(esk3_0))),X1)|~v1_finset_1(esk4_1(k1_zfmisc_1(esk3_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(esk3_0)),esk2_0)|~v1_finset_1(k1_zfmisc_1(esk3_0))|~m1_subset_1(k1_zfmisc_1(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_15620, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|v2_struct_0(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))))).#trainneg
cnf(i_0_2117, plain, (k12_lattice3(X1,k12_lattice3(X1,X3,X2),k12_lattice3(X1,X4,X5))=k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X2,X3),k7_domain_1(u1_struct_0(X1),X4,X5)))|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X4,X5))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~m1_subset_1(X5,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_11369, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_16482, negated_conjecture, (v1_card_1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk72_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_16483, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_16484, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_16485, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_16486, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_19052, negated_conjecture, (r2_hidden(esk69_3(esk1_0,esk9_2(esk1_0,esk2_0),k1_tarski(esk9_2(esk1_0,esk2_0))),k1_tarski(esk9_2(esk1_0,esk2_0)))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk9_2(esk1_0,esk2_0),k1_tarski(esk9_2(esk1_0,esk2_0))),esk9_2(esk1_0,esk2_0))|v5_waybel_6(esk2_0,esk1_0))).#trainneg
cnf(i_0_19160, negated_conjecture, (r2_hidden(esk69_3(esk1_0,esk8_2(esk1_0,esk2_0),k1_tarski(esk8_2(esk1_0,esk2_0))),k1_tarski(esk8_2(esk1_0,esk2_0)))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk8_2(esk1_0,esk2_0),k1_tarski(esk8_2(esk1_0,esk2_0))),esk8_2(esk1_0,esk2_0))|v5_waybel_6(esk2_0,esk1_0))).#trainneg
cnf(i_0_18830, plain, (r1_orders_2(X1,k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X2)),X3)|r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X4)),X3)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X2))|~r2_yellow_0(X1,k1_tarski(X4))|~r1_orders_2(X1,k12_lattice3(X1,X4,X2),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_19388, negated_conjecture, (r2_hidden(esk69_3(esk1_0,X1,k1_tarski(X2)),k1_tarski(X2))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k1_tarski(X2)),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,esk3_0)|~r1_orders_2(esk1_0,X2,X1)|~v5_waybel_6(X1,esk1_0)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_15121, plain, (esk72_3(X1,X2,k1_tarski(X2))=X2|r2_hidden(esk72_3(X1,X2,k1_tarski(X2)),k1_tarski(X2))|~r1_orders_2(X1,X2,X2)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v5_waybel_6(X2,X1))).#trainneg
cnf(i_0_14282, plain, (k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,k2_xboole_0(X2,X3))))=k2_yellow_0(X1,k2_xboole_0(X2,X3))|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1)))).#trainneg
cnf(i_0_14209, plain, (k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X3,X3),X2))=k12_lattice3(X1,k2_yellow_0(X1,X2),X3)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X3))|~r2_yellow_0(X1,X2)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_12283, negated_conjecture, (r2_hidden(esk69_3(esk1_0,X1,esk4_1(k1_zfmisc_1(esk3_0))),esk4_1(k1_zfmisc_1(esk3_0)))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,esk4_1(k1_zfmisc_1(esk3_0))),X1)|v1_xboole_0(esk4_1(k1_zfmisc_1(esk3_0)))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk4_1(k1_zfmisc_1(esk3_0))),X1)|~v1_finset_1(esk4_1(k1_zfmisc_1(esk3_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(esk3_0)),esk2_0)|~v1_finset_1(k1_zfmisc_1(esk3_0))|~m1_subset_1(k1_zfmisc_1(esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_5943, negated_conjecture, (esk4_1(X1)=X2|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|k2_xboole_0(u1_struct_0(esk1_0),X3)!=k1_tarski(X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_11211, negated_conjecture, (r1_tarski(X1,esk3_0)|v2_struct_0(esk1_0)|~r1_tarski(X1,esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_10381, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~r1_tarski(X1,k6_domain_1(esk3_0,X2))|~m1_subset_1(X2,esk3_0))).#trainneg
cnf(i_0_18379, plain, (r1_orders_2(X1,k12_lattice3(X1,X2,X3),k2_yellow_0(X1,k2_xboole_0(X4,k2_tarski(X2,X3))))|r1_orders_2(X1,k2_yellow_0(X1,X4),k2_yellow_0(X1,k2_xboole_0(X4,k2_tarski(X2,X3))))|v1_xboole_0(k2_xboole_0(X4,k2_tarski(X2,X3)))|~r2_yellow_0(X1,k2_tarski(X2,X3))|~r2_yellow_0(X1,X4)|~v5_waybel_6(k2_yellow_0(X1,k2_xboole_0(X4,k2_tarski(X2,X3))),X1)|~m1_subset_1(k2_yellow_0(X1,k2_xboole_0(X4,k2_tarski(X2,X3))),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(k2_xboole_0(X4,k2_tarski(X2,X3)))|~m1_subset_1(k2_xboole_0(X4,k2_tarski(X2,X3)),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_15430, negated_conjecture, (esk4_1(k2_tarski(k2_yellow_0(esk1_0,k2_tarski(X1,X1)),k2_yellow_0(esk1_0,X2)))=k2_yellow_0(esk1_0,k2_tarski(X1,X1))|esk4_1(k2_tarski(k2_yellow_0(esk1_0,k2_tarski(X1,X1)),k2_yellow_0(esk1_0,X2)))=k2_yellow_0(esk1_0,X2)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,k2_yellow_0(esk1_0,X2)),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(X1,X1)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14559, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(X3,k2_xboole_0(X2,k1_tarski(k2_yellow_0(X1,X2)))))|~r2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(k2_yellow_0(X1,X2))))|~r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~r2_yellow_0(X1,X2)|~m1_subset_1(k1_tarski(k2_yellow_0(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1))).#trainneg
cnf(i_0_10001, plain, (k2_yellow_0(X1,k2_xboole_0(k1_tarski(k2_yellow_0(X1,X2)),k1_tarski(X3)))=k12_lattice3(X1,k2_yellow_0(X1,X2),X3)|v1_xboole_0(X2)|~r2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2)))|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_19666, negated_conjecture, (r2_hidden(esk69_3(esk1_0,esk2_0,k1_tarski(X1)),k1_tarski(X1))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,k1_tarski(X1)),esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_19861, negated_conjecture, (r1_tarski(X1,esk3_0)|v2_struct_0(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))))).#trainneg
cnf(i_0_19585, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_19860, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk3_0),esk3_0)|v2_struct_0(esk1_0))).#trainneg
cnf(i_0_19875, negated_conjecture, (r1_tarski(k6_domain_1(esk3_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_19468, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk12_1(k1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_20103, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2))),X3)|r1_orders_2(X1,X4,X3)|v1_xboole_0(X2)|~r2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2)))|~r1_orders_2(X1,k12_lattice3(X1,k2_yellow_0(X1,X2),X4),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(k1_tarski(X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_16792, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,k1_tarski(X2)))=X2|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X2))|~r2_yellow_0(X1,k1_tarski(X2))|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_11358, negated_conjecture, (r1_tarski(esk70_3(X1,k2_yellow_0(X1,X2),X2),u1_struct_0(esk1_0))|r2_hidden(esk72_3(X1,k2_yellow_0(X1,X2),X2),X2)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(esk3_0))|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_15603, negated_conjecture, (r1_tarski(k1_tarski(X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_tarski(X2))|~m1_subset_1(X2,esk3_0))).#trainneg
cnf(i_0_10370, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk3_0),k1_numbers)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10436, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10443, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10455, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_12054, negated_conjecture, (esk69_3(esk1_0,esk2_0,esk3_0)=X1|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|k1_tarski(X1)!=esk3_0)).#trainneg
cnf(i_0_17129, plain, (m1_subset_1(esk69_3(X1,X2,k1_tarski(X2)),k1_tarski(X2))|r2_hidden(esk72_3(X1,X2,k1_tarski(X2)),k1_tarski(X2))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_9831, negated_conjecture, (m1_subset_1(esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11637, negated_conjecture, (r2_hidden(esk69_3(esk1_0,X1,k2_tarski(esk8_2(esk1_0,X1),esk9_2(esk1_0,X1))),k2_tarski(esk8_2(esk1_0,X1),esk9_2(esk1_0,X1)))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k2_tarski(esk8_2(esk1_0,X1),esk9_2(esk1_0,X1))),X1)|v5_waybel_6(X1,esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk8_2(esk1_0,X1),esk9_2(esk1_0,X1)),k1_zfmisc_1(esk3_0))|~m1_subset_1(esk9_2(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13036, negated_conjecture, (m1_subset_1(esk72_3(X1,k2_yellow_0(X1,X2),X2),X2)|v1_xboole_0(X2)|m1_subset_1(esk69_3(X1,k2_yellow_0(X1,X2),X2),u1_struct_0(esk1_0))|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(esk3_0))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11217, negated_conjecture, (k2_tarski(X1,X2)=X3|k1_tarski(X3)!=k1_zfmisc_1(u1_struct_0(esk1_0))|~m1_subset_1(X2,esk3_0)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_13402, negated_conjecture, (esk69_3(esk1_0,esk2_0,esk3_0)=X1|v2_struct_0(esk1_0)|k2_xboole_0(esk3_0,X2)!=k1_tarski(X1))).#trainneg
cnf(i_0_13403, negated_conjecture, (esk69_3(esk1_0,esk2_0,esk3_0)=X1|v2_struct_0(esk1_0)|k2_xboole_0(X2,esk3_0)!=k1_tarski(X1))).#trainneg
cnf(i_0_12048, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),k1_numbers)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_15615, negated_conjecture, (v1_xxreal_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_15616, negated_conjecture, (v1_card_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_15617, negated_conjecture, (v1_xxreal_0(esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_10409, negated_conjecture, (r2_hidden(esk69_3(esk1_0,X1,k1_tarski(X2)),k1_tarski(X2))|m1_subset_1(esk72_3(esk1_0,X1,k1_tarski(X2)),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_15550, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,k2_tarski(X2,X2)))=k2_yellow_0(X1,k2_tarski(X2,X2))|~r2_yellow_0(X1,k2_tarski(X2,X2))|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11657, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,X1),X1)|v1_xboole_0(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,k2_xboole_0(esk70_3(esk1_0,esk2_0,X1),k1_tarski(esk69_3(esk1_0,esk2_0,X1)))),esk3_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_16888, negated_conjecture, (r1_tarski(esk70_3(esk1_0,k2_yellow_0(esk1_0,X1),k2_xboole_0(X1,k1_tarski(k2_yellow_0(esk1_0,X1)))),k2_xboole_0(X1,k1_tarski(k2_yellow_0(esk1_0,X1))))|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),k2_xboole_0(X1,k1_tarski(k2_yellow_0(esk1_0,X1)))),k2_xboole_0(X1,k1_tarski(k2_yellow_0(esk1_0,X1))))|v1_xboole_0(k2_xboole_0(X1,k1_tarski(k2_yellow_0(esk1_0,X1))))|~v1_finset_1(k2_xboole_0(X1,k1_tarski(k2_yellow_0(esk1_0,X1))))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(k2_xboole_0(X1,k1_tarski(k2_yellow_0(esk1_0,X1))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,X1)),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_13687, plain, (r1_tarski(esk70_3(X1,X2,k2_tarski(X3,k2_yellow_0(X1,k1_tarski(X3)))),k2_tarski(X3,k2_yellow_0(X1,k1_tarski(X3))))|r2_hidden(esk72_3(X1,X2,k2_tarski(X3,k2_yellow_0(X1,k1_tarski(X3)))),k2_tarski(X3,k2_yellow_0(X1,k1_tarski(X3))))|~r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X3)),X2)|~m1_subset_1(k2_tarski(X3,k2_yellow_0(X1,k1_tarski(X3))),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_16908, plain, (r1_tarski(esk70_3(X1,k2_yellow_0(X1,u1_struct_0(X1)),u1_struct_0(X1)),u1_struct_0(X1))|r1_orders_2(X1,esk72_3(X1,k2_yellow_0(X1,u1_struct_0(X1)),u1_struct_0(X1)),k2_yellow_0(X1,u1_struct_0(X1)))|v2_struct_0(X1)|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,u1_struct_0(X1)),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_6210, negated_conjecture, (esk4_1(X1)=X2|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|k2_xboole_0(X3,u1_struct_0(esk1_0))!=k1_tarski(X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_12499, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(k6_domain_1(esk3_0,X2)))|~m1_subset_1(X2,esk3_0))).#trainneg
cnf(i_0_5878, plain, (v1_xboole_0(k2_xboole_0(X1,X2))|~r2_hidden(k2_tarski(k2_xboole_0(X1,X2),X3),X1))).#trainneg
cnf(i_0_20233, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk3_0),k1_numbers)|v2_struct_0(esk1_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20673, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20234, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|v2_struct_0(esk1_0))).#trainneg
cnf(i_0_20161, negated_conjecture, (esk69_3(esk1_0,esk2_0,k1_tarski(X1))=X2|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,k1_tarski(X1)),esk2_0)|k1_tarski(X1)!=k1_tarski(X2)|~r1_orders_2(esk1_0,X1,esk2_0)|~v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_20170, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,esk2_0,k1_tarski(X1)),u1_struct_0(esk1_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,k1_tarski(X1)),esk2_0)|~m1_subset_1(X1,esk3_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~v5_waybel_6(esk2_0,esk1_0))).#trainneg
cnf(i_0_20151, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,esk2_0,k1_tarski(X1)),k1_tarski(X1))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,k1_tarski(X1)),esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_2144, plain, (k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X2,X3),k6_domain_1(u1_struct_0(X1),X4)))=k12_lattice3(X1,k12_lattice3(X1,X2,X3),X4)|v2_struct_0(X1)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_20934, plain, (~r2_hidden(k2_tarski(X1,X2),k1_xboole_0))).#trainneg
cnf(i_0_20504, negated_conjecture, (v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20505, negated_conjecture, (v1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_zfmisc_1(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20560, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk72_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_20654, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0))).#trainneg
cnf(i_0_20608, negated_conjecture, (esk69_3(esk1_0,X1,k2_tarski(esk8_2(esk1_0,X1),esk9_2(esk1_0,X1)))=esk9_2(esk1_0,X1)|esk69_3(esk1_0,X1,k2_tarski(esk8_2(esk1_0,X1),esk9_2(esk1_0,X1)))=esk8_2(esk1_0,X1)|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k2_tarski(esk8_2(esk1_0,X1),esk9_2(esk1_0,X1))),X1)|v5_waybel_6(X1,esk1_0)|~m1_subset_1(k2_tarski(esk8_2(esk1_0,X1),esk9_2(esk1_0,X1)),k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_20570, plain, (m1_subset_1(esk72_3(X1,X2,k1_tarski(X2)),k1_tarski(X2))|m1_subset_1(esk69_3(X1,X2,k1_tarski(X2)),k1_tarski(X2))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_20179, negated_conjecture, (r1_tarski(X1,esk3_0)|v2_struct_0(esk1_0)|~r1_tarski(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))))).#trainneg
cnf(i_0_2145, plain, (k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X2,X3),k1_tarski(X4)))=k12_lattice3(X1,k12_lattice3(X1,X2,X3),X4)|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,k1_tarski(X4))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_20936, plain, (~r2_hidden(k2_tarski(k2_xboole_0(X2,X1),X3),X1))).#trainneg
cnf(i_0_20939, plain, (~r2_hidden(k2_tarski(X3,k2_xboole_0(X1,X2)),X1))).#trainneg
cnf(i_0_20506, negated_conjecture, (v6_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v6_ordinal1(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20507, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_relat_1(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20508, negated_conjecture, (v5_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v5_finset_1(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20509, negated_conjecture, (v4_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v4_funct_1(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_11681, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k2_xboole_0(k1_tarski(X2),k1_tarski(X3))),k2_xboole_0(k1_tarski(X2),k1_tarski(X3)))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k2_xboole_0(k1_tarski(X2),k1_tarski(X3))),X1)|v1_xboole_0(k2_xboole_0(k1_tarski(X2),k1_tarski(X3)))|~r1_orders_2(esk1_0,k12_lattice3(esk1_0,X3,X2),X1)|~v1_finset_1(k2_xboole_0(k1_tarski(X2),k1_tarski(X3)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_xboole_0(k1_tarski(X2),k1_tarski(X3)),k1_zfmisc_1(esk3_0))|~r2_yellow_0(esk1_0,k1_tarski(X2))|~m1_subset_1(X3,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_11312, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),esk69_3(X1,k2_yellow_0(X1,esk3_0),esk3_0))|m1_subset_1(esk72_3(X1,k2_yellow_0(X1,esk3_0),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_12650, negated_conjecture, (esk69_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)),k2_tarski(esk69_3(esk1_0,esk2_0,esk3_0),esk2_0))=esk2_0|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)),k2_tarski(esk69_3(esk1_0,esk2_0,esk3_0),esk2_0))=esk2_0|esk69_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)),k2_tarski(esk69_3(esk1_0,esk2_0,esk3_0),esk2_0))=esk69_3(esk1_0,esk2_0,esk3_0)|esk72_3(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)),k2_tarski(esk69_3(esk1_0,esk2_0,esk3_0),esk2_0))=esk69_3(esk1_0,esk2_0,esk3_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0))).#trainneg
cnf(i_0_17483, negated_conjecture, (v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(X1)),esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_10589, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,k1_tarski(X1)),k1_tarski(X1))|m1_subset_1(esk72_3(esk1_0,esk2_0,k1_tarski(X1)),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16837, plain, (r2_hidden(esk69_3(X1,k2_yellow_0(X1,u1_struct_0(X1)),u1_struct_0(X1)),u1_struct_0(X1))|r1_orders_2(X1,esk72_3(X1,k2_yellow_0(X1,u1_struct_0(X1)),u1_struct_0(X1)),k2_yellow_0(X1,u1_struct_0(X1)))|v2_struct_0(X1)|~v1_finset_1(u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,u1_struct_0(X1)),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_20490, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v3_ordinal1(k1_zfmisc_1(esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20491, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v4_funct_1(k1_zfmisc_1(esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20492, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v4_funct_1(k1_zfmisc_1(esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20493, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_card_1(k1_zfmisc_1(esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20519, negated_conjecture, (v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v5_finset_1(u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_14956, plain, (esk72_3(X1,X2,k1_tarski(X2))=X2|r2_hidden(esk72_3(X1,X2,k1_tarski(X2)),k1_tarski(X2))|~r1_orders_2(X1,X2,X2)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_14410, negated_conjecture, (r1_tarski(X1,X2)|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,X2),X2),k2_yellow_0(esk1_0,X2))|v1_xboole_0(X2)|~r1_tarski(X1,esk70_3(esk1_0,k2_yellow_0(esk1_0,X2),X2))|~v1_finset_1(X2)|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14179, plain, (k2_yellow_0(X1,k2_xboole_0(k2_tarski(X2,X2),k1_tarski(X2)))=k2_yellow_0(X1,k1_tarski(X2))|~r2_yellow_0(X1,k2_tarski(X2,X2))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_20970, negated_conjecture, (esk69_3(esk1_0,esk2_0,k1_tarski(X1))=X1|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,k1_tarski(X1)),esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_20987, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,k1_tarski(X1)),esk69_3(esk1_0,esk2_0,k1_tarski(X1)))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,k1_tarski(X1)),esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_2147, plain, (k12_lattice3(X1,k12_lattice3(X1,X2,X3),k12_lattice3(X1,X5,X4))=k2_yellow_0(X1,k2_xboole_0(k7_domain_1(u1_struct_0(X1),X2,X3),k7_domain_1(u1_struct_0(X1),X4,X5)))|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X4,X5))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1)))).#trainneg
cnf(i_0_20950, negated_conjecture, (v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v5_finset_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20952, negated_conjecture, (v1_funct_1(esk69_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v4_funct_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20953, negated_conjecture, (v1_relat_1(esk69_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v4_funct_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20965, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk3_0),k1_numbers)|v2_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_21513, plain, (m1_subset_1(esk72_3(X1,X2,k1_tarski(X2)),k1_tarski(X2))|esk72_3(X1,X2,k1_tarski(X2))=X2|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_21533, plain, (v1_xboole_0(X1)|m1_subset_1(esk72_3(X2,X3,k1_tarski(X3)),X1)|esk72_3(X2,X3,k1_tarski(X3))=X3|~m1_subset_1(X3,X1)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_21587, negated_conjecture, (r1_tarski(X1,X2)|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,X2),X2),k2_yellow_0(esk1_0,X2))|v1_xboole_0(X2)|~v1_finset_1(X2)|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,k1_zfmisc_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,X2),X2))))).#trainneg
cnf(i_0_21258, negated_conjecture, (r1_tarski(esk70_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),esk3_0)|v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),k2_yellow_0(esk1_0,X1))|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0)))|~v1_finset_1(X1)|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_21494, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,k1_tarski(X1)),k1_zfmisc_1(k1_tarski(X1)))|m1_subset_1(esk72_3(esk1_0,esk2_0,k1_tarski(X1)),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_20944, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|v2_struct_0(esk1_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_21213, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,X1)),esk3_0)|v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk70_3(esk1_0,esk2_0,X1)),esk70_3(esk1_0,esk2_0,X1))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~v1_finset_1(X1)|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_21398, plain, (r1_orders_2(X1,k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3)),X4)|r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X5)),X4)|~r2_yellow_0(X1,k1_tarski(X5))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r1_orders_2(X1,k12_lattice3(X1,k12_lattice3(X1,X2,X3),X5),X4)|~v5_waybel_6(X4,X1)|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_18822, plain, (r1_orders_2(X1,k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3)),X4)|r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X5)),X4)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X3))|~r2_yellow_0(X1,k1_tarski(X5))|~r1_orders_2(X1,k12_lattice3(X1,X5,k12_lattice3(X1,X3,X2)),X4)|~v5_waybel_6(X4,X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X5,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_15326, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),u1_struct_0(esk1_0))|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk3_0))|~v1_finset_1(X1)|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_13382, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk71_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14880, negated_conjecture, (v5_waybel_6(esk2_0,esk1_0)|m1_subset_1(X1,k1_numbers)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_20520, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v3_ordinal1(u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20521, negated_conjecture, (v1_funct_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v4_funct_1(u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_16316, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k1_tarski(X2)),k1_tarski(X2))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k1_tarski(X2)),X1)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_16977, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|m1_subset_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_16261, negated_conjecture, (m1_subset_1(esk69_3(esk1_0,X1,k1_tarski(X1)),k1_tarski(X1))|r2_hidden(esk72_3(esk1_0,X1,k1_tarski(X1)),k1_tarski(X1))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_17109, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk9_2(esk1_0,esk2_0),k1_tarski(esk9_2(esk1_0,esk2_0))),k1_tarski(esk9_2(esk1_0,esk2_0)))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk9_2(esk1_0,esk2_0),k1_tarski(esk9_2(esk1_0,esk2_0))),esk9_2(esk1_0,esk2_0))|v5_waybel_6(esk2_0,esk1_0))).#trainneg
cnf(i_0_21710, negated_conjecture, (v2_struct_0(esk1_0)|esk69_3(esk1_0,esk2_0,k1_tarski(X1))=X1|~m1_subset_1(esk72_3(esk1_0,esk2_0,k1_tarski(X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk72_3(esk1_0,esk2_0,k1_tarski(X1)),esk3_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_22328, negated_conjecture, (v5_waybel_6(esk2_0,esk1_0)|m1_subset_1(X1,k1_numbers)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2159, plain, (r1_tarski(X1,u1_struct_0(X2))|v2_struct_0(X2)|v2_lattice3(X2)|~r1_tarski(X1,esk73_1(X2))|~l1_orders_2(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_21417, plain, (~r2_hidden(k2_tarski(X3,k2_xboole_0(X1,X2)),X2))).#trainneg
cnf(i_0_22079, negated_conjecture, (v1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_zfmisc_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_22080, negated_conjecture, (v6_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v6_ordinal1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_22364, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k1_tarski(esk2_0)),k1_tarski(esk2_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k1_tarski(esk2_0)),X1)|~r1_orders_2(esk1_0,esk2_0,X1)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_22368, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k1_tarski(k2_yellow_0(esk1_0,X2))),k1_tarski(k2_yellow_0(esk1_0,X2)))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k1_tarski(k2_yellow_0(esk1_0,X2))),X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X2),X1)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_22277, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1)=X2|v1_xboole_0(X1)|m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),u1_struct_0(esk1_0))|X1!=k1_tarski(X2)|~v1_finset_1(X1)|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_22365, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k1_tarski(k1_yellow_0(esk1_0,X2))),k1_tarski(k1_yellow_0(esk1_0,X2)))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k1_tarski(k1_yellow_0(esk1_0,X2))),X1)|~r1_orders_2(esk1_0,k1_yellow_0(esk1_0,X2),X1)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_22050, negated_conjecture, (v1_finset_1(esk70_3(esk1_0,esk2_0,k1_tarski(X1)))|m1_subset_1(esk72_3(esk1_0,esk2_0,k1_tarski(X1)),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_22267, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|v1_xboole_0(X1)|m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),u1_struct_0(esk1_0))|~v1_finset_1(X1)|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_14503, negated_conjecture, (v1_finset_1(X1)|v5_waybel_6(esk2_0,esk1_0)|~v5_finset_1(k1_numbers)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X1)),esk2_0)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(k4_ordinal1))|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2178, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~r1_tarski(X1,k1_tarski(X3))|~m1_subset_1(X3,X2))).#trainneg
cnf(i_0_22081, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_relat_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_22082, negated_conjecture, (v5_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v5_finset_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_22083, negated_conjecture, (v4_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v4_funct_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_22068, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_22084, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_relat_1(k1_numbers)|~v1_funct_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_22403, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,X1,k1_tarski(X1)),k1_tarski(X1))|m1_subset_1(esk69_3(esk1_0,X1,k1_tarski(X1)),k1_tarski(X1))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_22119, plain, (r1_orders_2(X1,k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X2)),X3)|r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(X4)),X3)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X2))|~r2_yellow_0(X1,k1_tarski(X4))|~r1_orders_2(X1,k12_lattice3(X1,X2,X4),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_22355, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k1_tarski(esk9_2(esk1_0,X2))),k1_tarski(esk9_2(esk1_0,X2)))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k1_tarski(esk9_2(esk1_0,X2))),X1)|v2_struct_0(esk1_0)|v5_waybel_6(X2,esk1_0)|~r1_orders_2(esk1_0,esk9_2(esk1_0,X2),X1)|~m1_subset_1(X1,esk3_0)|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7223, plain, (r2_hidden(esk69_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),k7_domain_1(u1_struct_0(X1),X3,X4))|r1_orders_2(X1,esk72_3(X1,X2,k7_domain_1(u1_struct_0(X1),X3,X4)),X2)|v1_xboole_0(k7_domain_1(u1_struct_0(X1),X3,X4))|~r1_orders_2(X1,k12_lattice3(X1,X3,X4),X2)|~v1_finset_1(k7_domain_1(u1_struct_0(X1),X3,X4))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_14998, negated_conjecture, (r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|v1_xboole_0(X1)|~r2_hidden(X2,esk70_3(esk1_0,k2_yellow_0(esk1_0,X1),X1))|~r1_orders_2(esk1_0,X2,k2_yellow_0(esk1_0,X1))|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6791, plain, (r1_orders_2(X1,k2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2)),X3)|r1_orders_2(X1,X4,X3)|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X2))|~r1_orders_2(X1,k12_lattice3(X1,X4,X2),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(k1_tarski(X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_9589, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|~m1_subset_1(k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,X1)),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,X1)),esk3_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~v1_finset_1(X1)|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_8368, plain, (r2_hidden(esk69_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|r1_orders_2(X1,esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|v5_waybel_6(X2,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11694, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k2_tarski(X2,k2_yellow_0(esk1_0,X3))),k2_tarski(X2,k2_yellow_0(esk1_0,X3)))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k2_tarski(X2,k2_yellow_0(esk1_0,X3))),X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X2),X3)),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(X2,k2_yellow_0(esk1_0,X3)),k1_zfmisc_1(esk3_0))|~r2_yellow_0(esk1_0,X3)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_6404, plain, (r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X3,k1_tarski(X2))),k2_yellow_0(X1,k2_xboole_0(X3,k1_tarski(X2))))|~m1_subset_1(k2_tarski(X2,k2_yellow_0(X1,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X3)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_2187, plain, (esk61_1(X1)=X2|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|u1_struct_0(X1)!=k1_tarski(X2)|~v2_yellow_0(X1)|~l1_orders_2(X1)|~v5_orders_2(X1))).#trainneg
cnf(i_0_21825, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|v2_struct_0(esk1_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20522, negated_conjecture, (v1_relat_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v4_funct_1(u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_19084, plain, (r1_orders_2(X1,k2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X2)),X3)|r1_orders_2(X1,X4,X3)|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X2,X2))|~r2_yellow_0(X1,k1_tarski(X4))|~r1_orders_2(X1,k12_lattice3(X1,X4,X2),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(k1_tarski(X4),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_14218, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k1_tarski(k2_yellow_0(esk1_0,X2))),k1_tarski(k2_yellow_0(esk1_0,X2)))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k1_tarski(k2_yellow_0(esk1_0,X2))),X1)|v1_xboole_0(X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X2),X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,X2),esk3_0)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_12369, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0)),k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)))|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),u1_struct_0(esk1_0))|~m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_15548, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(X3,k2_tarski(X2,X2)))|~r2_yellow_0(X1,X3)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_22633, negated_conjecture, (v1_xboole_0(X1)|m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),u1_struct_0(esk1_0))|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_2201, plain, (esk9_2(X1,X2)=X3|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|v5_waybel_6(X2,X1)|u1_struct_0(X1)!=k1_tarski(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_22947, negated_conjecture, (v1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_zfmisc_1(k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_22948, negated_conjecture, (v6_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v6_ordinal1(k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_22949, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_relat_1(k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_23182, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),X3)=k2_yellow_0(X1,k2_xboole_0(X2,k2_tarski(X3,X3)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~r2_yellow_0(X1,X2)|~m1_subset_1(k2_tarski(X3,X3),k1_zfmisc_1(u1_struct_0(X1)))|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_22715, negated_conjecture, (r2_hidden(esk72_3(esk1_0,X1,k1_tarski(X1)),k1_tarski(X1))|~r2_hidden(X2,esk70_3(esk1_0,X1,k1_tarski(X1)))|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_22651, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,k1_tarski(X1)),X2)|v1_xboole_0(X2)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,k1_tarski(X1)),esk2_0)|~m1_subset_1(X1,X2)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_22846, plain, (esk69_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk9_2(X1,X2)|esk69_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|r1_orders_2(X1,esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|v5_waybel_6(X2,X1)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_22652, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k1_tarski(X1)),X2)|v1_xboole_0(X2)|r2_hidden(esk72_3(esk1_0,X1,k1_tarski(X1)),k1_tarski(X1))|~m1_subset_1(X1,X2)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_22967, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k1_tarski(X2)),k1_tarski(X2))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k1_tarski(X2)),X1)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X2,esk3_0)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_23064, negated_conjecture, (r2_hidden(esk4_1(k2_xboole_0(X1,k2_tarski(X2,X2))),X1)|r2_hidden(esk4_1(k2_xboole_0(X1,k2_tarski(X2,X2))),k2_tarski(X2,X2))|v1_xboole_0(k2_xboole_0(X1,k2_tarski(X2,X2)))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k12_lattice3(esk1_0,X2,k2_yellow_0(esk1_0,X1)),esk2_0)|~v1_finset_1(k2_xboole_0(X1,k2_tarski(X2,X2)))|~m1_subset_1(k2_xboole_0(X1,k2_tarski(X2,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_22644, negated_conjecture, (v1_finset_1(X1)|v5_waybel_6(esk2_0,esk1_0)|~v5_finset_1(k1_numbers)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2202, plain, (esk8_2(X1,X2)=X3|v1_xboole_0(u1_struct_0(X1))|v2_struct_0(X1)|v5_waybel_6(X2,X1)|u1_struct_0(X1)!=k1_tarski(X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_22950, negated_conjecture, (v5_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v5_finset_1(k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_22951, negated_conjecture, (v4_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v4_funct_1(k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_22952, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|~v1_relat_1(k1_numbers)|~v1_funct_1(k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20523, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_card_1(u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_23196, plain, (r1_orders_2(X1,k2_yellow_0(X1,X2),X3)|r1_orders_2(X1,X4,X3)|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X2,k2_tarski(X4,X4))),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~r2_yellow_0(X1,X2)|~m1_subset_1(k2_tarski(X4,X4),k1_zfmisc_1(u1_struct_0(X1)))|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_22655, plain, (r1_tarski(esk70_3(X1,X2,k1_tarski(X2)),X3)|v1_xboole_0(X3)|r2_hidden(esk72_3(X1,X2,k1_tarski(X2)),k1_tarski(X2))|~m1_subset_1(X2,X3)|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_23705, negated_conjecture, (r2_hidden(esk4_1(k2_tarski(X1,X1)),k2_tarski(X1,X1))|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X1)),esk2_0)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_23631, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk4_1(X1),k1_tarski(X2)),k1_tarski(X2))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk4_1(X1),k1_tarski(X2)),esk4_1(X1))|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,X2,esk4_1(X1))|~m1_subset_1(X2,esk3_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_23632, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,k1_tarski(X1)),k1_tarski(X1))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,k1_tarski(X1)),esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_23578, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,k1_tarski(X1)),k1_zfmisc_1(X2))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,k1_tarski(X1)),esk2_0)|v1_xboole_0(X2)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_23710, negated_conjecture, (v1_finset_1(X1)|v5_waybel_6(esk2_0,esk1_0)|~v5_finset_1(k1_numbers)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2206, plain, (k6_domain_1(X1,X2)=X3|v1_xboole_0(X1)|k1_zfmisc_1(X1)!=k1_tarski(X3)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_20494, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~m1_subset_1(k1_zfmisc_1(esk3_0),k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20510, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_relat_1(esk3_0)|~v1_funct_1(esk3_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_23565, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,k1_tarski(X1)),k1_zfmisc_1(u1_struct_0(esk1_0)))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,k1_tarski(X1)),esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_23428, plain, (k2_yellow_0(X1,k2_xboole_0(X2,k2_tarski(X3,X3)))=k2_yellow_0(X1,k2_tarski(X3,k2_yellow_0(X1,X2)))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X2)|~m1_subset_1(k2_tarski(X3,X3),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_23324, negated_conjecture, (r2_hidden(esk4_1(k2_xboole_0(X1,k2_tarski(X2,X2))),X1)|r2_hidden(esk4_1(k2_xboole_0(X1,k2_tarski(X2,X2))),k2_tarski(X2,X2))|v1_xboole_0(k2_xboole_0(X1,k2_tarski(X2,X2)))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k12_lattice3(esk1_0,k2_yellow_0(esk1_0,X1),X2),esk2_0)|~v1_finset_1(k2_xboole_0(X1,k2_tarski(X2,X2)))|~m1_subset_1(k2_xboole_0(X1,k2_tarski(X2,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_23439, plain, (k2_yellow_0(X1,k2_xboole_0(X2,k2_tarski(X3,X3)))=k2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X2),X3))|~m1_subset_1(k2_yellow_0(X1,X2),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X2)|~m1_subset_1(k2_tarski(X3,X3),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_22705, plain, (r2_hidden(esk69_3(X1,X2,k2_tarski(X3,X4)),k2_tarski(X3,X4))|r1_orders_2(X1,esk72_3(X1,X2,k2_tarski(X3,X4)),X2)|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k12_lattice3(X1,X3,X4),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_2207, plain, (k1_yellow_0(X1,X2)=X3|v1_xboole_0(u1_struct_0(X1))|u1_struct_0(X1)!=k1_tarski(X3)|~l1_orders_2(X1))).#trainneg
cnf(i_0_20524, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20548, negated_conjecture, (v3_ordinal1(esk72_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_20549, negated_conjecture, (v3_ordinal1(esk72_3(esk1_0,esk2_0,esk3_0))|v2_struct_0(esk1_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1))).#trainneg
cnf(i_0_20482, negated_conjecture, (v3_ordinal1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(esk3_0,k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_23900, negated_conjecture, (v1_xxreal_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k1_zfmisc_1(esk3_0),k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_23797, negated_conjecture, (esk4_1(k2_tarski(X1,X1))=X1|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X1)),esk2_0)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_23792, negated_conjecture, (esk4_1(k2_tarski(X1,X1))=X2|v5_waybel_6(esk2_0,esk1_0)|k2_tarski(X1,X1)!=k1_tarski(X2)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X1)),esk2_0)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2208, plain, (k11_lattice3(X1,X2,X3)=X4|v1_xboole_0(u1_struct_0(X1))|u1_struct_0(X1)!=k1_tarski(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_9362, negated_conjecture, (v1_zfmisc_1(k2_tarski(X1,esk3_0))|~v1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9363, negated_conjecture, (v6_ordinal1(k2_tarski(X1,esk3_0))|~v6_ordinal1(k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_23901, negated_conjecture, (v1_card_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k1_zfmisc_1(esk3_0),k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_15618, negated_conjecture, (v1_card_1(esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_card_1(k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_17833, negated_conjecture, (v1_xxreal_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k1_zfmisc_1(esk3_0),k4_ordinal1))).#trainneg
cnf(i_0_17834, negated_conjecture, (v1_card_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k1_zfmisc_1(esk3_0),k4_ordinal1))).#trainneg
cnf(i_0_24601, plain, (r2_hidden(esk69_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))),k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)))|r1_orders_2(X1,esk72_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))),X2)|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X4,X3)),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X4))).#trainneg
cnf(i_0_24600, plain, (r2_hidden(esk69_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))),k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)))|r1_orders_2(X1,esk72_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))),X2)|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X3,X4)),X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X3),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,X3))).#trainneg
cnf(i_0_11659, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,k2_xboole_0(esk70_3(esk1_0,esk2_0,X1),k1_tarski(esk69_3(esk1_0,esk2_0,X1)))),esk3_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_21259, negated_conjecture, (r1_tarski(esk70_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),esk3_0)|v2_struct_0(esk1_0)|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),X1),X1)|v1_xboole_0(X1)|~m1_subset_1(X1,k1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0)))|~v1_finset_1(X1)|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_16793, plain, (k2_yellow_0(X1,k2_tarski(X2,k2_yellow_0(X1,k1_tarski(X2))))=X2|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X2))|~r2_yellow_0(X1,k1_tarski(X2))|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_15718, plain, (r1_orders_2(X1,X2,X3)|r1_orders_2(X1,k2_yellow_0(X1,X4),X3)|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(k2_tarski(X2,X2),X4)),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(k2_yellow_0(X1,X4),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~r2_yellow_0(X1,X4)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_14868, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k1_tarski(X2)),k1_tarski(X2))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k1_tarski(X2)),X1)|~r1_orders_2(esk1_0,X2,X1)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(esk3_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2210, plain, (k7_domain_1(X1,X2,X3)=X4|v1_xboole_0(X1)|k1_zfmisc_1(X1)!=k1_tarski(X4)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_9364, negated_conjecture, (v1_relat_1(k2_tarski(X1,esk3_0))|~v1_relat_1(k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9365, negated_conjecture, (v5_finset_1(k2_tarski(X1,esk3_0))|~v5_finset_1(k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_9366, negated_conjecture, (v4_funct_1(k2_tarski(X1,esk3_0))|~v4_funct_1(k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_17835, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1))).#trainneg
cnf(i_0_17836, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1))).#trainneg
cnf(i_0_17837, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1))).#trainneg
cnf(i_0_17838, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1))).#trainneg
cnf(i_0_20473, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_17926, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k2_tarski(X2,k2_yellow_0(esk1_0,X3))),k2_tarski(X2,k2_yellow_0(esk1_0,X3)))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k2_tarski(X2,k2_yellow_0(esk1_0,X3))),X1)|~r2_yellow_0(esk1_0,k1_tarski(X2))|~r2_yellow_0(esk1_0,X3)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X3,k1_tarski(X2))),X1)|~m1_subset_1(k2_tarski(X2,k2_yellow_0(esk1_0,X3)),k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_18363, plain, (r1_orders_2(X1,k12_lattice3(X1,X2,X3),X4)|r1_orders_2(X1,k2_yellow_0(X1,X5),X4)|~r2_yellow_0(X1,k2_tarski(X2,X3))|~r2_yellow_0(X1,X5)|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(k2_tarski(X2,X3),X5)),X4)|~v5_waybel_6(X4,X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_14917, plain, (k2_yellow_0(X1,k2_xboole_0(k1_tarski(X2),k1_tarski(X3)))=k12_lattice3(X1,X2,X3)|~r2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_5606, negated_conjecture, (r1_tarski(esk70_3(X1,X2,esk4_1(esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))),esk4_1(esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))))|r1_orders_2(X1,esk72_3(X1,X2,esk4_1(esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))),X2)|v1_xboole_0(esk4_1(esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))))|v1_xboole_0(esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(X1,k2_yellow_0(X1,esk4_1(esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))),X2)|~v1_finset_1(esk4_1(esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))),esk2_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))),esk2_0)|~v1_finset_1(esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))|~v1_finset_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_24646, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X1)),esk2_0)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2211, plain, (k2_tarski(X1,X2)=X3|v1_xboole_0(X4)|k1_zfmisc_1(X4)!=k1_tarski(X3)|~m1_subset_1(X2,X4)|~m1_subset_1(X1,X4))).#trainneg
cnf(i_0_15397, plain, (v3_ordinal1(k1_tarski(X1))|v1_xboole_0(X2)|~m1_subset_1(X1,X2)|~m1_subset_1(k1_zfmisc_1(X2),k4_ordinal1))).#trainneg
cnf(i_0_11483, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|v2_struct_0(esk1_0)|~r1_tarski(X1,esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_25016, negated_conjecture, (v1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_zfmisc_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_25017, negated_conjecture, (v6_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v6_ordinal1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_25018, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_relat_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_24766, plain, (r2_hidden(esk69_3(X1,k2_yellow_0(X1,k2_xboole_0(X2,X3)),k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))),k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3)))|r1_orders_2(X1,esk72_3(X1,k2_yellow_0(X1,k2_xboole_0(X2,X3)),k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))),k2_yellow_0(X1,k2_xboole_0(X2,X3)))|v1_xboole_0(k2_xboole_0(X2,X3))|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2)|~m1_subset_1(k2_yellow_0(X1,k2_xboole_0(X2,X3)),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(k2_xboole_0(X2,X3))|~m1_subset_1(k2_xboole_0(X2,X3),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_24714, plain, (r2_hidden(esk69_3(X1,k2_yellow_0(X1,k2_xboole_0(X2,X3)),k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X2))),k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X2)))|r1_orders_2(X1,esk72_3(X1,k2_yellow_0(X1,k2_xboole_0(X2,X3)),k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X2))),k2_yellow_0(X1,k2_xboole_0(X2,X3)))|v1_xboole_0(k2_xboole_0(X2,X3))|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2)|~m1_subset_1(k2_yellow_0(X1,k2_xboole_0(X2,X3)),u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(k2_xboole_0(X2,X3))|~m1_subset_1(k2_xboole_0(X2,X3),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_5605, negated_conjecture, (r2_hidden(esk69_3(X1,X2,esk4_1(esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))),esk4_1(esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))))|r1_orders_2(X1,esk72_3(X1,X2,esk4_1(esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))),X2)|v1_xboole_0(esk4_1(esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))))|v1_xboole_0(esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(X1,k2_yellow_0(X1,esk4_1(esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))),X2)|~v1_finset_1(esk4_1(esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))),esk2_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))),esk2_0)|~v1_finset_1(esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))))|~v1_finset_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))))|~m1_subset_1(esk4_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1)))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(X1))),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2212, plain, (k12_lattice3(X1,X2,X3)=X4|v1_xboole_0(u1_struct_0(X1))|u1_struct_0(X1)!=k1_tarski(X4)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1))).#trainneg
cnf(i_0_9821, negated_conjecture, (esk69_3(esk1_0,esk2_0,esk3_0)=X1|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|k1_tarski(X1)!=esk3_0)).#trainneg
cnf(i_0_10419, negated_conjecture, (esk72_3(esk1_0,esk2_0,esk3_0)=X1|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|k1_tarski(X1)!=esk3_0)).#trainneg
cnf(i_0_25019, negated_conjecture, (v5_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v5_finset_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_25020, negated_conjecture, (v4_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v4_funct_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_25005, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_25021, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_relat_1(k1_numbers)|~v1_funct_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_24636, negated_conjecture, (v1_xxreal_0(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_14378, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(k1_tarski(k2_yellow_0(X1,X3)),X2))|v1_xboole_0(X3)|~r2_yellow_0(X1,X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(k2_yellow_0(X1,X3)),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_25227, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X1)),esk2_0)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2220, plain, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(esk69_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),u1_struct_0(X1))|r2_hidden(esk72_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),k6_domain_1(u1_struct_0(X1),X3))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X3,X2)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_25326, negated_conjecture, (k2_tarski(X1,esk2_0)=X2|k1_tarski(X2)!=k1_zfmisc_1(u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_25325, negated_conjecture, (k2_tarski(X1,esk3_0)=X2|k1_tarski(X2)!=k1_zfmisc_1(k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_24637, negated_conjecture, (v1_card_1(esk69_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_finset_1(esk69_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_20475, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10354, negated_conjecture, (r1_tarski(X1,esk3_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))))).#trainneg
cnf(i_0_2227, plain, (v1_xboole_0(X1)|m1_subset_1(X2,X1)|v1_xboole_0(k6_domain_1(X1,X3))|~m1_subset_1(X3,X1)|~m1_subset_1(X2,k6_domain_1(X1,X3)))).#trainneg
cnf(i_0_25357, plain, (k2_tarski(X1,X2)=X3|k1_zfmisc_1(k1_zfmisc_1(X2))!=k1_tarski(X3)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_11399, negated_conjecture, (esk69_3(esk1_0,esk2_0,esk3_0)=X1|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|k1_tarski(X1)!=esk3_0)).#trainneg
cnf(i_0_22069, negated_conjecture, (esk70_3(esk1_0,esk2_0,esk3_0)=X1|v2_struct_0(esk1_0)|k1_tarski(X1)!=k1_zfmisc_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_11914, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk3_0),k1_numbers)|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_12554, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,esk3_0),k1_numbers)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10053, plain, (k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,k1_tarski(X2))))=k2_yellow_0(X1,k1_tarski(X2))|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_22654, plain, (r1_tarski(esk70_3(X1,X2,k1_tarski(X3)),X4)|v1_xboole_0(X4)|r2_hidden(esk72_3(X1,X2,k1_tarski(X3)),k1_tarski(X3))|~m1_subset_1(X3,X4)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_14537, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(X3,k1_tarski(k2_yellow_0(X1,X3))),X2))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,k2_xboole_0(X3,k1_tarski(k2_yellow_0(X1,X3))))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~r2_yellow_0(X1,X3)|~m1_subset_1(k1_tarski(k2_yellow_0(X1,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1))).#trainneg
cnf(i_0_5538, plain, (r2_hidden(esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|esk69_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk8_2(X1,X2)|esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))=esk9_2(X1,X2)|v5_waybel_6(X2,X1)|~r1_orders_2(X1,k2_yellow_0(X1,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),X2)|~r1_orders_2(X1,X3,X2)|~m1_subset_1(X3,k1_tarski(esk9_2(X1,X2)))|~m1_subset_1(k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_10543, negated_conjecture, (r1_tarski(X1,esk3_0)|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))))).#trainneg
cnf(i_0_2234, plain, (esk5_2(X1,k2_tarski(X2,X3))=X3|esk5_2(X1,k2_tarski(X2,X3))=X2|esk5_2(X1,k2_tarski(X2,X3))=X1|k2_tarski(X2,X3)=k1_tarski(X1))).#trainneg
cnf(i_0_20174, negated_conjecture, (m1_subset_1(X1,k1_zfmisc_1(esk3_0))|v2_struct_0(esk1_0)|~m1_subset_1(X1,k1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))))).#trainneg
cnf(i_0_25348, plain, (k2_tarski(X1,X2)=X3|k1_zfmisc_1(k1_tarski(X2))!=k1_tarski(X3)|~m1_subset_1(X1,k1_tarski(X2)))).#trainneg
cnf(i_0_13895, negated_conjecture, (k7_domain_1(esk3_0,X1,X2)=X3|k1_tarski(X3)!=k1_zfmisc_1(u1_struct_0(esk1_0))|~m1_subset_1(X2,esk3_0)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_15656, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_18029, negated_conjecture, (v1_xxreal_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k4_ordinal1))).#trainneg
cnf(i_0_18030, negated_conjecture, (v1_card_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k4_ordinal1))).#trainneg
cnf(i_0_18031, negated_conjecture, (v1_xxreal_0(esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k4_ordinal1))).#trainneg
cnf(i_0_18032, negated_conjecture, (v1_card_1(esk70_3(esk1_0,esk2_0,esk3_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k4_ordinal1))).#trainneg
cnf(i_0_11294, negated_conjecture, (r1_tarski(esk70_3(esk1_0,X1,k6_domain_1(esk3_0,X2)),k6_domain_1(esk3_0,X2))|r1_orders_2(esk1_0,esk72_3(esk1_0,X1,k6_domain_1(esk3_0,X2)),X1)|v1_xboole_0(k6_domain_1(esk3_0,X2))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k6_domain_1(esk3_0,X2)),X1)|~v1_finset_1(k6_domain_1(esk3_0,X2))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,esk3_0))).#trainneg
cnf(i_0_11335, negated_conjecture, (r1_tarski(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),k2_yellow_0(esk1_0,esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_7503, plain, (esk72_3(X3,X4,k2_tarski(esk8_2(X3,X4),esk9_2(X3,X4)))=esk8_2(X3,X4)|v5_waybel_6(X4,X3)|esk9_2(X3,X4)!=esk8_2(X3,X4)|~r1_orders_2(X3,esk8_2(X3,X4),X4)|~r1_orders_2(X3,esk9_2(X3,X4),X4)|~m1_subset_1(k2_tarski(esk8_2(X3,X4),esk9_2(X3,X4)),k1_zfmisc_1(u1_struct_0(X3)))|~m1_subset_1(esk8_2(X3,X4),u1_struct_0(X3))|~m1_subset_1(esk9_2(X3,X4),u1_struct_0(X3))|~m1_subset_1(X4,u1_struct_0(X3))|~l1_orders_2(X3)|~v2_lattice3(X3)|~v5_orders_2(X3)|~v4_orders_2(X3)|~v3_orders_2(X3))).#trainneg
cnf(i_0_10404, negated_conjecture, (r2_hidden(esk69_3(esk1_0,esk2_0,k1_tarski(X1)),k1_tarski(X1))|m1_subset_1(esk72_3(esk1_0,esk2_0,k1_tarski(X1)),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14165, plain, (k2_yellow_0(X1,k2_xboole_0(k2_tarski(X2,X2),k1_tarski(X3)))=k12_lattice3(X1,X2,X3)|~r2_yellow_0(X1,k2_tarski(X2,X2))|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_26064, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X1,k1_tarski(k2_yellow_0(esk1_0,X1)))),esk2_0)|~m1_subset_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X1)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,X1)),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2288, plain, (m1_subset_1(esk70_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),k1_zfmisc_1(k6_domain_1(u1_struct_0(X1),X3)))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|m1_subset_1(esk72_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),u1_struct_0(X1))|~r1_orders_2(X1,X3,X2)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11363, negated_conjecture, (r1_tarski(k1_tarski(X1),u1_struct_0(esk1_0))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_3274, plain, (esk13_2(X1,X2)=X3|v1_xboole_0(X2)|X2!=k1_tarski(X3)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_11194, negated_conjecture, (esk70_3(esk1_0,esk2_0,esk3_0)=X1|v2_struct_0(esk1_0)|k2_xboole_0(k1_zfmisc_1(esk3_0),X2)!=k1_tarski(X1))).#trainneg
cnf(i_0_26352, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~r2_hidden(esk3_0,esk72_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_26328, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk47_1(esk70_3(esk1_0,esk2_0,esk3_0)),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_26333, negated_conjecture, (v2_struct_0(esk1_0)|m1_subset_1(esk12_1(k1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_26556, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,k1_tarski(X3)))=k12_lattice3(X1,X2,X3)|~r2_yellow_0(X1,k1_tarski(X3))|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k2_tarski(X2,X2))|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_26557, plain, (k12_lattice3(X1,k2_yellow_0(X1,k1_tarski(X3)),X2)=k12_lattice3(X1,X2,X3)|~r2_yellow_0(X1,k1_tarski(X3))|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k2_tarski(X2,X2))|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_26268, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|v1_xboole_0(X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,X1))|~v1_finset_1(X1)|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk70_3(esk1_0,esk2_0,X1),k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_26559, plain, (k2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,k1_tarski(X3)),X2))=k12_lattice3(X1,X2,X3)|~r2_yellow_0(X1,k1_tarski(X3))|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k2_tarski(X2,X2))|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_26251, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(X3,k1_tarski(k2_yellow_0(X1,X3))),X2))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,X3)|~m1_subset_1(k1_tarski(k2_yellow_0(X1,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X3))))).#trainneg
cnf(i_0_12478, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))))).#trainneg
cnf(i_0_8748, negated_conjecture, (v3_ordinal1(X1)|v5_waybel_6(esk2_0,esk1_0)|~v3_ordinal1(k1_tarski(X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X1)),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_2289, plain, (r1_tarski(X1,k6_domain_1(u1_struct_0(X2),X3))|v1_xboole_0(k6_domain_1(u1_struct_0(X2),X3))|m1_subset_1(esk72_3(X2,X4,k6_domain_1(u1_struct_0(X2),X3)),u1_struct_0(X2))|~r1_tarski(X1,esk70_3(X2,X4,k6_domain_1(u1_struct_0(X2),X3)))|~r1_orders_2(X2,X3,X4)|~v1_finset_1(k6_domain_1(u1_struct_0(X2),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X2),X3),k1_zfmisc_1(u1_struct_0(X2)))|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~l1_orders_2(X2)|~v2_lattice3(X2)|~v5_orders_2(X2)|~v4_orders_2(X2)|~v3_orders_2(X2))).#trainneg
cnf(i_0_11195, negated_conjecture, (esk70_3(esk1_0,esk2_0,esk3_0)=X1|v2_struct_0(esk1_0)|k2_xboole_0(X2,k1_zfmisc_1(esk3_0))!=k1_tarski(X1))).#trainneg
cnf(i_0_11255, negated_conjecture, (esk69_3(esk1_0,esk2_0,esk3_0)=X1|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|k1_tarski(X1)!=u1_struct_0(esk1_0))).#trainneg
cnf(i_0_26106, negated_conjecture, (v1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_zfmisc_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_26107, negated_conjecture, (v6_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v6_ordinal1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_26108, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_relat_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_26109, negated_conjecture, (v5_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v5_finset_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_26558, plain, (k2_yellow_0(X1,k2_tarski(X2,k2_yellow_0(X1,k1_tarski(X3))))=k12_lattice3(X1,X2,X3)|~r2_yellow_0(X1,k1_tarski(X3))|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k2_tarski(X2,X2))|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_26248, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),X3)=k2_yellow_0(X1,k2_xboole_0(k2_xboole_0(k2_tarski(X3,X3),k1_tarski(X3)),X2))|~r2_yellow_0(X1,k2_xboole_0(k2_tarski(X3,X3),k1_tarski(X3)))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,k2_tarski(X3,X3))|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_13674, negated_conjecture, (esk4_1(k2_tarski(k2_yellow_0(esk1_0,k1_tarski(X1)),k2_yellow_0(esk1_0,k1_tarski(X1))))=k2_yellow_0(esk1_0,k1_tarski(X1))|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(X1,k2_yellow_0(esk1_0,k1_tarski(X1))))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X1)),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_4801, plain, (r2_hidden(esk69_3(X1,X2,k2_tarski(X2,X2)),k2_tarski(X2,X2))|r2_hidden(esk72_3(X1,X2,k2_tarski(X2,X2)),k2_tarski(X2,X2))|v2_struct_0(X1)|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_26118, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(X1,k1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))))).#trainneg
cnf(i_0_17634, negated_conjecture, (v1_xxreal_0(X1)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X1)),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(X1),k4_ordinal1))).#trainneg
cnf(i_0_11922, negated_conjecture, (esk70_3(esk1_0,esk2_0,esk3_0)=X1|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|k1_tarski(X1)!=k1_zfmisc_1(esk3_0))).#trainneg
cnf(i_0_22937, negated_conjecture, (esk70_3(esk1_0,esk2_0,esk3_0)=X1|v2_struct_0(esk1_0)|k1_tarski(X1)!=k1_zfmisc_1(k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_26110, negated_conjecture, (v4_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v4_funct_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_26091, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v4_funct_1(k1_zfmisc_1(k1_numbers))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_26092, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v4_funct_1(k1_zfmisc_1(k1_numbers))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_26093, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_card_1(k1_zfmisc_1(k1_numbers))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_14160, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(X3,k2_tarski(X2,X2)))|~r2_yellow_0(X1,k2_tarski(X2,X2))|~r2_yellow_0(X1,X3)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v3_orders_2(X1))).#trainneg
cnf(i_0_14966, negated_conjecture, (v2_struct_0(esk1_0)|r1_orders_2(esk1_0,X1,k11_lattice3(esk1_0,X1,X2))|r1_orders_2(esk1_0,X2,k11_lattice3(esk1_0,X1,X2))|~v5_waybel_6(k11_lattice3(esk1_0,X1,X2),esk1_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k11_lattice3(esk1_0,X1,X2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_14971, negated_conjecture, (r1_orders_2(esk1_0,X1,k12_lattice3(esk1_0,X2,X1))|r1_orders_2(esk1_0,X2,k12_lattice3(esk1_0,X2,X1))|~v5_waybel_6(k12_lattice3(esk1_0,X2,X1),esk1_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(k12_lattice3(esk1_0,X2,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_12477, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,k1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))))).#trainneg
cnf(i_0_19590, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|~r1_tarski(X1,X2)|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(X2,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_26732, negated_conjecture, (v1_funct_1(esk47_1(esk70_3(esk1_0,esk2_0,esk3_0)))|v2_struct_0(esk1_0)|~v1_relat_1(esk3_0)|~v1_funct_1(esk3_0))).#trainneg
cnf(i_0_26746, negated_conjecture, (v3_ordinal1(esk12_1(k1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))))|v2_struct_0(esk1_0)|~v1_card_1(k1_zfmisc_1(esk3_0)))).#trainneg
cnf(i_0_26747, negated_conjecture, (v3_ordinal1(esk12_1(k1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))))|v2_struct_0(esk1_0)|~m1_subset_1(k1_zfmisc_1(esk3_0),k4_ordinal1))).#trainneg
cnf(i_0_26666, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X1,k1_tarski(k2_yellow_0(esk1_0,X1)))),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,X1)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13951, negated_conjecture, (r1_tarski(X1,u1_struct_0(esk1_0))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(X1,k1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))))).#trainneg
cnf(i_0_2293, plain, (m1_subset_1(esk69_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),k6_domain_1(u1_struct_0(X1),X3))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|m1_subset_1(esk72_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),u1_struct_0(X1))|~r1_orders_2(X1,X3,X2)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_26094, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_27233, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,k2_xboole_0(X1,X2)))|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X2),k2_yellow_0(esk1_0,k2_xboole_0(X1,X2)))|~v5_waybel_6(k2_yellow_0(esk1_0,k2_xboole_0(X1,X2)),esk1_0)|~m1_subset_1(k2_yellow_0(esk1_0,k2_xboole_0(X1,X2)),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,X1))).#trainneg
cnf(i_0_2915, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,k2_xboole_0(X1,X2)))|r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X2),k2_yellow_0(esk1_0,k2_xboole_0(X1,X2)))|v2_struct_0(esk1_0)|~r2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,X1)|~v5_waybel_6(k2_yellow_0(esk1_0,k2_xboole_0(X1,X2)),esk1_0)|~m1_subset_1(k2_yellow_0(esk1_0,k2_xboole_0(X1,X2)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_9485, negated_conjecture, (esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0)=X1|v2_struct_0(esk1_0)|k2_xboole_0(esk3_0,X2)!=k1_tarski(X1)|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_18167, negated_conjecture, (esk69_3(esk1_0,esk2_0,k1_tarski(X1))=X2|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,k1_tarski(X1)),esk2_0)|k1_tarski(X1)!=k1_tarski(X2)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2295, plain, (esk69_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3))=X4|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|m1_subset_1(esk72_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),u1_struct_0(X1))|k6_domain_1(u1_struct_0(X1),X3)!=k1_tarski(X4)|~r1_orders_2(X1,X3,X2)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_26704, plain, (esk13_2(X1,k1_tarski(X2))=X2|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_26111, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_relat_1(k1_numbers)|~v1_funct_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_26360, negated_conjecture, (v3_ordinal1(esk72_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(esk3_0,k4_ordinal1)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_26351, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_9973, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(k2_yellow_0(X1,X3))))|v1_xboole_0(X3)|~r2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X3)))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~v3_orders_2(X1))).#trainneg
cnf(i_0_2389, plain, (r2_hidden(esk69_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)))|m1_subset_1(esk72_3(X1,X2,k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2))),u1_struct_0(X1))|v5_waybel_6(X2,X1)|~m1_subset_1(k2_tarski(esk8_2(X1,X2),esk9_2(X1,X2)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(esk9_2(X1,X2),u1_struct_0(X1))|~m1_subset_1(esk8_2(X1,X2),u1_struct_0(X1)))).#trainneg
cnf(i_0_2296, plain, (v1_xboole_0(u1_struct_0(X1))|m1_subset_1(esk69_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),u1_struct_0(X1))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|m1_subset_1(esk72_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X3,X2)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_26709, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|v1_xboole_0(esk72_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(esk3_0,esk72_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_11497, negated_conjecture, (v3_ordinal1(esk72_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~v3_ordinal1(esk3_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_3578, plain, (r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|r1_orders_2(X1,esk69_3(X1,X2,X3),X2)|v1_xboole_0(X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|~r2_yellow_0(X1,esk70_3(X1,X2,X3))|~v5_waybel_6(X2,X1)|~m1_subset_1(k1_tarski(esk69_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_23636, negated_conjecture, (r1_tarski(esk70_3(esk1_0,k2_yellow_0(esk1_0,X1),k1_tarski(X2)),k1_tarski(X2))|r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,X1),k1_tarski(X2)),k2_yellow_0(esk1_0,X1))|~r1_orders_2(esk1_0,X2,k2_yellow_0(esk1_0,X1))|~m1_subset_1(X2,esk3_0))).#trainneg
cnf(i_0_12604, plain, (r2_hidden(esk72_3(X1,X2,k1_tarski(X3)),k1_tarski(X3))|m1_subset_1(esk69_3(X1,X2,k1_tarski(X3)),k1_tarski(X3))|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,X3,X2)|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_13045, negated_conjecture, (esk72_3(X1,k2_yellow_0(X1,X2),X2)=X3|v1_xboole_0(X2)|m1_subset_1(esk69_3(X1,k2_yellow_0(X1,X2),X2),u1_struct_0(esk1_0))|X2!=k1_tarski(X3)|~v1_finset_1(X2)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,k1_zfmisc_1(esk3_0))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_14150, plain, (k2_yellow_0(X1,k2_xboole_0(k2_tarski(X3,X3),k1_tarski(X2)))=k12_lattice3(X1,X2,X3)|~r2_yellow_0(X1,k2_tarski(X3,X3))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_27521, plain, (m1_subset_1(X1,X2)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_27519, plain, (m2_subset_1(X2,X1,k1_tarski(X2))|~m1_subset_1(k1_tarski(X2),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_9928, negated_conjecture, (esk69_3(esk1_0,esk2_0,esk3_0)=X1|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|k1_tarski(X1)!=esk3_0)).#trainneg
cnf(i_0_27541, negated_conjecture, (v1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_zfmisc_1(k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_27542, negated_conjecture, (v6_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v6_ordinal1(k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_27543, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_relat_1(k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_27407, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X1)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,k2_yellow_0(esk1_0,k1_tarski(X1))),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k1_tarski(X1))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_27400, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k2_tarski(X1,X1),k1_tarski(X1))),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_27517, negated_conjecture, (esk69_3(esk1_0,esk2_0,k1_tarski(X1))=X1|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,k1_tarski(X1)),esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_27357, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,X2),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k7_domain_1(u1_struct_0(esk1_0),X1,X2),k1_tarski(k12_lattice3(esk1_0,X1,X2)))),esk2_0)|~m1_subset_1(k1_tarski(k12_lattice3(esk1_0,X1,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k12_lattice3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_27405, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X1)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r1_orders_2(esk1_0,k12_lattice3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X1)),X1),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k1_tarski(X1))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_27404, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1))|~r1_orders_2(esk1_0,k12_lattice3(esk1_0,k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1)),X1),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k6_domain_1(u1_struct_0(esk1_0),X1)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_27358, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X2,X1),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k7_domain_1(u1_struct_0(esk1_0),X1,X2),k1_tarski(k12_lattice3(esk1_0,X2,X1)))),esk2_0)|~m1_subset_1(k1_tarski(k12_lattice3(esk1_0,X2,X1)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k12_lattice3(esk1_0,X2,X1),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_22664, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k1_tarski(X3))))).#trainneg
cnf(i_0_10447, negated_conjecture, (esk72_3(esk1_0,esk2_0,esk3_0)=X1|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk3_0))|k1_tarski(X1)!=esk3_0)).#trainneg
cnf(i_0_10459, negated_conjecture, (esk72_3(esk1_0,esk2_0,esk3_0)=X1|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|k1_tarski(X1)!=esk3_0)).#trainneg
cnf(i_0_27544, negated_conjecture, (v5_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v5_finset_1(k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_27545, negated_conjecture, (v4_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v4_funct_1(k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_27674, negated_conjecture, (v1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk72_3(esk1_0,esk2_0,esk3_0))|~v1_zfmisc_1(k1_numbers)|~m1_subset_1(esk3_0,esk72_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_27686, plain, (r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|v1_xboole_0(X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|v1_xboole_0(esk70_3(X1,X2,X3))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(k1_tarski(esk69_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_finset_1(esk70_3(X1,X2,X3))|~m1_subset_1(esk70_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_27389, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X1,X2)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)),k1_tarski(k2_yellow_0(esk1_0,k2_xboole_0(X1,X2))))),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_xboole_0(X1,X2))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_xboole_0(X1,X2)),u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,X1))).#trainneg
cnf(i_0_27390, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X2,X1)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)),k1_tarski(k2_yellow_0(esk1_0,k2_xboole_0(X2,X1))))),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_xboole_0(X2,X1))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_xboole_0(X2,X1)),u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,X1)|~r2_yellow_0(esk1_0,X2))).#trainneg
cnf(i_0_27366, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,X2),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X1),k1_tarski(X2)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k2_xboole_0(k1_tarski(X1),k1_tarski(X2)),k1_tarski(k12_lattice3(esk1_0,X1,X2)))),esk2_0)|~m1_subset_1(k1_tarski(k12_lattice3(esk1_0,X1,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k12_lattice3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,k1_tarski(X2))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_27365, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X2,X1),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X1),k1_tarski(X2)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k2_xboole_0(k1_tarski(X1),k1_tarski(X2)),k1_tarski(k12_lattice3(esk1_0,X2,X1)))),esk2_0)|~m1_subset_1(k1_tarski(k12_lattice3(esk1_0,X2,X1)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k12_lattice3(esk1_0,X2,X1),u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,k1_tarski(X1))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_27953, plain, (k12_lattice3(X1,k2_yellow_0(X1,k1_tarski(X3)),X2)=k12_lattice3(X1,X3,X2)|~r2_yellow_0(X1,k1_tarski(X3))|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k2_tarski(X2,X2))|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_27604, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))=k2_yellow_0(X1,k2_xboole_0(X2,k1_tarski(k2_yellow_0(X1,X3))))|v1_xboole_0(X3)|~r2_yellow_0(X1,X2)|~v1_finset_1(X3)|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(k1_tarski(k2_yellow_0(X1,X3)),k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_2309, negated_conjecture, (v1_xboole_0(X1)|m1_subset_1(esk4_1(k7_domain_1(X1,X2,X3)),X1)|v1_xboole_0(k7_domain_1(X1,X2,X3))|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X3,X1)|~m1_subset_1(X2,X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k7_domain_1(X1,X2,X3)),esk2_0)|~v1_finset_1(k7_domain_1(X1,X2,X3))|~m1_subset_1(k7_domain_1(X1,X2,X3),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_27963, plain, (m1_subset_1(X1,k1_numbers)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_27675, negated_conjecture, (v6_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk72_3(esk1_0,esk2_0,esk3_0))|~v6_ordinal1(k1_numbers)|~m1_subset_1(esk3_0,esk72_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_27676, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk72_3(esk1_0,esk2_0,esk3_0))|~v1_relat_1(k1_numbers)|~m1_subset_1(esk3_0,esk72_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_27677, negated_conjecture, (v5_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk72_3(esk1_0,esk2_0,esk3_0))|~v5_finset_1(k1_numbers)|~m1_subset_1(esk3_0,esk72_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_27678, negated_conjecture, (v4_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk72_3(esk1_0,esk2_0,esk3_0))|~v4_funct_1(k1_numbers)|~m1_subset_1(esk3_0,esk72_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_27529, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~m1_subset_1(k1_zfmisc_1(k1_numbers),k4_ordinal1)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_27978, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,X1),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_28002, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,k2_yellow_0(esk1_0,k1_tarski(X1))),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_yellow_0(esk1_0,k1_tarski(X1))|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_28000, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,k2_yellow_0(esk1_0,k1_tarski(X1)))),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_yellow_0(esk1_0,k1_tarski(X1))|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_28029, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k7_domain_1(u1_struct_0(esk1_0),X1,X2),k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X1,X2))))),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X1,X2))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_27211, negated_conjecture, (v1_xxreal_0(X1)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~v1_finset_1(X1)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(X1),k4_ordinal1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2319, plain, (r2_hidden(esk7_3(k2_xboole_0(X1,X2),X3,X3),X2)|r2_hidden(esk7_3(k2_xboole_0(X1,X2),X3,X3),X1)|k2_xboole_0(k2_xboole_0(X1,X2),X3)=X3|r2_hidden(esk7_3(k2_xboole_0(X1,X2),X3,X3),X3))).#trainneg
cnf(i_0_12528, negated_conjecture, (esk70_3(esk1_0,esk2_0,esk3_0)=X1|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|k1_tarski(X1)!=k1_zfmisc_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_26697, negated_conjecture, (m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_numbers,k1_zfmisc_1(esk3_0))|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_5087, plain, (r1_tarski(k2_tarski(X1,X2),k1_numbers)|~m1_subset_1(X2,k4_ordinal1)|~m1_subset_1(X1,k4_ordinal1))).#trainneg
cnf(i_0_27546, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~v1_relat_1(k1_numbers)|~v1_funct_1(k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_27679, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk72_3(esk1_0,esk2_0,esk3_0))|~v1_relat_1(k1_numbers)|~v1_funct_1(k1_numbers)|~m1_subset_1(esk3_0,esk72_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_27445, negated_conjecture, (v1_xxreal_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_27446, negated_conjecture, (v1_card_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_28034, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k7_domain_1(u1_struct_0(esk1_0),X1,X1),k1_tarski(X1))),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13601, plain, (v3_ordinal1(k1_yellow_0(X1,X2))|v1_xboole_0(u1_struct_0(X1))|~v1_card_1(u1_struct_0(X1))|~l1_orders_2(X1))).#trainneg
cnf(i_0_2320, plain, (esk7_3(k2_tarski(X1,X2),X3,X3)=X2|esk7_3(k2_tarski(X1,X2),X3,X3)=X1|k2_xboole_0(k2_tarski(X1,X2),X3)=X3|r2_hidden(esk7_3(k2_tarski(X1,X2),X3,X3),X3))).#trainneg
cnf(i_0_20241, negated_conjecture, (r1_tarski(k6_domain_1(esk3_0,X1),k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_6251, plain, (r1_tarski(k1_tarski(X1),k1_numbers)|v1_xboole_0(X2)|~m1_subset_1(X2,k1_zfmisc_1(k4_ordinal1))|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_6252, plain, (r1_tarski(k1_tarski(X1),X2)|v1_xboole_0(X3)|~m1_subset_1(X3,k1_zfmisc_1(k1_xboole_0))|~m1_subset_1(X1,X3))).#trainneg
cnf(i_0_10597, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_10807, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_27999, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~r1_orders_2(esk1_0,k12_lattice3(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X1)),X1),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_yellow_0(esk1_0,k1_tarski(X1))|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_28027, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X2,X1)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k7_domain_1(u1_struct_0(esk1_0),X1,X2),k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X2,X1))))),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X2,X1))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_27990, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X1)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(k2_yellow_0(esk1_0,k1_tarski(X1)),X1)),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k1_tarski(X1))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,k1_tarski(X1)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2323, plain, (k2_xboole_0(X1,X2)=X2|r2_hidden(esk7_3(X1,X2,X2),X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_11466, negated_conjecture, (esk70_3(esk1_0,esk2_0,esk3_0)=X1|v2_struct_0(esk1_0)|k2_xboole_0(k1_zfmisc_1(u1_struct_0(esk1_0)),X2)!=k1_tarski(X1))).#trainneg
cnf(i_0_11467, negated_conjecture, (esk70_3(esk1_0,esk2_0,esk3_0)=X1|v2_struct_0(esk1_0)|k2_xboole_0(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))!=k1_tarski(X1))).#trainneg
cnf(i_0_29367, negated_conjecture, (v1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_zfmisc_1(u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_29368, negated_conjecture, (v6_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v6_ordinal1(u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_29369, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_relat_1(u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_28825, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|v1_xboole_0(X1)|~r2_yellow_0(esk1_0,k1_tarski(k2_yellow_0(esk1_0,X1)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X1,k1_tarski(k2_yellow_0(esk1_0,X1)))),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,X1)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_yellow_0(esk1_0,X1)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_28848, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X1)))|~r2_yellow_0(esk1_0,k1_tarski(k2_yellow_0(esk1_0,X1)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X1,k1_tarski(k2_yellow_0(esk1_0,X1)))),esk2_0)|~m1_subset_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X1)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,X1)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_yellow_0(esk1_0,X1))).#trainneg
cnf(i_0_28886, negated_conjecture, (v1_xxreal_0(X1)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~v1_finset_1(X1)|~m1_subset_1(k1_tarski(X1),k4_ordinal1)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2324, plain, (m1_subset_1(esk7_3(X1,X2,X2),X1)|k2_xboole_0(X1,X2)=X2|r2_hidden(esk7_3(X1,X2,X2),X2))).#trainneg
cnf(i_0_29496, plain, (k2_xboole_0(X1,X2)=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_29724, negated_conjecture, (v1_xxreal_0(esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_29725, negated_conjecture, (v1_xxreal_0(k2_yellow_0(esk1_0,esk3_0))|v5_waybel_6(esk2_0,esk1_0)|~v1_finset_1(k2_yellow_0(esk1_0,esk3_0))|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,esk3_0)),k4_ordinal1)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_29370, negated_conjecture, (v5_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v5_finset_1(u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_29371, negated_conjecture, (v4_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v4_funct_1(u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_2326, plain, (esk7_3(X1,X2,X2)=X3|k2_xboole_0(X1,X2)=X2|r2_hidden(esk7_3(X1,X2,X2),X2)|X1!=k1_tarski(X3))).#trainneg
cnf(i_0_30029, plain, (k2_xboole_0(X2,X1)=X2|~v1_xboole_0(X2)|~v1_xboole_0(X1))).#trainneg
cnf(i_0_29351, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v3_ordinal1(k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_29352, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v4_funct_1(k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_29353, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v4_funct_1(k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_29355, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_28831, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X1)),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_28819, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,X2),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k1_tarski(k12_lattice3(esk1_0,X1,X2)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2)),esk2_0)|~m1_subset_1(k1_tarski(k12_lattice3(esk1_0,X1,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2333, plain, (r2_hidden(esk7_3(X1,k2_xboole_0(X2,X3),k2_xboole_0(X2,X3)),X3)|r2_hidden(esk7_3(X1,k2_xboole_0(X2,X3),k2_xboole_0(X2,X3)),X2)|k2_xboole_0(X1,k2_xboole_0(X2,X3))=k2_xboole_0(X2,X3)|r2_hidden(esk7_3(X1,k2_xboole_0(X2,X3),k2_xboole_0(X2,X3)),X1))).#trainneg
cnf(i_0_29722, negated_conjecture, (v1_xxreal_0(esk8_2(esk1_0,esk2_0))|v5_waybel_6(esk2_0,esk1_0)|~v1_finset_1(esk8_2(esk1_0,esk2_0))|~m1_subset_1(k1_tarski(esk8_2(esk1_0,esk2_0)),k4_ordinal1)|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|esk9_2(esk1_0,esk2_0)!=esk8_2(esk1_0,esk2_0)|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_29372, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_relat_1(u1_struct_0(esk1_0))|~v1_funct_1(u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_29340, negated_conjecture, (v1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_zfmisc_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_2334, plain, (esk7_3(X1,k2_tarski(X2,X3),k2_tarski(X2,X3))=X3|esk7_3(X1,k2_tarski(X2,X3),k2_tarski(X2,X3))=X2|k2_xboole_0(X1,k2_tarski(X2,X3))=k2_tarski(X2,X3)|r2_hidden(esk7_3(X1,k2_tarski(X2,X3),k2_tarski(X2,X3)),X1))).#trainneg
cnf(i_0_29341, negated_conjecture, (v6_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v6_ordinal1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_29342, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_relat_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_29343, negated_conjecture, (v5_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v5_finset_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_29344, negated_conjecture, (v4_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v4_funct_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_29700, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k1_tarski(k2_yellow_0(esk1_0,X1)))|~r2_yellow_0(esk1_0,X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X1,k1_tarski(k2_yellow_0(esk1_0,X1)))),esk2_0)|~m1_subset_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X1)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,X1)),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_30373, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X1,X2))))|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2)),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X1,X2))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_30378, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X1)),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2335, plain, (r2_hidden(esk7_3(X1,k1_xboole_0,k1_xboole_0),X2)|X1=k1_xboole_0|r2_hidden(esk7_3(X1,k1_xboole_0,k1_xboole_0),X1))).#trainneg
cnf(i_0_29328, negated_conjecture, (v3_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_zfmisc_1(k1_numbers),k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_29345, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_relat_1(k1_numbers)|~v1_funct_1(k1_numbers)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_30384, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,X2),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2)),esk2_0)|~m1_subset_1(k1_tarski(k12_lattice3(esk1_0,X1,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2337, plain, (k2_xboole_0(X1,X2)=X2|r2_hidden(esk7_3(X1,X2,X2),X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_30337, negated_conjecture, (v1_xxreal_0(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_30338, negated_conjecture, (v1_card_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k1_zfmisc_1(u1_struct_0(esk1_0)),k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_16152, negated_conjecture, (v3_ordinal1(esk72_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k4_ordinal1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_20466, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_28864, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,k2_yellow_0(esk1_0,k1_tarski(X1)))),esk2_0)|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_29178, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X1))|~r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,k12_lattice3(esk1_0,X1,X1)),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_yellow_0(esk1_0,k1_tarski(X1))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2338, plain, (m1_subset_1(esk7_3(X1,X2,X2),X2)|k2_xboole_0(X1,X2)=X2|r2_hidden(esk7_3(X1,X2,X2),X1))).#trainneg
cnf(i_0_2340, plain, (esk7_3(X1,X2,X2)=X3|k2_xboole_0(X1,X2)=X2|r2_hidden(esk7_3(X1,X2,X2),X1)|X2!=k1_tarski(X3))).#trainneg
cnf(i_0_26137, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_26144, negated_conjecture, (m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_30831, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k1_tarski(k2_yellow_0(esk1_0,X1)))|~r2_yellow_0(esk1_0,X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X1,k1_tarski(k2_yellow_0(esk1_0,X1)))),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,X1)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,X1),esk3_0))).#trainneg
cnf(i_0_28483, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X1),k1_tarski(X2)))|~r2_yellow_0(esk1_0,k1_tarski(X2))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k2_xboole_0(k1_tarski(X1),k1_tarski(X2)),k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X1,X2))))),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X1,X2))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_27952, plain, (k12_lattice3(X1,X2,k2_yellow_0(X1,k1_tarski(X3)))=k12_lattice3(X1,X3,X2)|~r2_yellow_0(X1,k1_tarski(X3))|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,k2_tarski(X2,X2))|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_28481, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X2,X1)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X1),k1_tarski(X2)))|~r2_yellow_0(esk1_0,k1_tarski(X2))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k2_xboole_0(k1_tarski(X1),k1_tarski(X2)),k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X2,X1))))),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X2,X1))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_28040, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,X2),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(X1,X2))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k2_tarski(X1,X2),k1_tarski(k12_lattice3(esk1_0,X1,X2)))),esk2_0)|~m1_subset_1(k1_tarski(k12_lattice3(esk1_0,X1,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_31303, negated_conjecture, (v1_zfmisc_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_zfmisc_1(k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_30998, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,X2),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|v1_xboole_0(k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2)),esk2_0)|~m1_subset_1(k1_tarski(k12_lattice3(esk1_0,X1,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v1_finset_1(k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~m1_subset_1(k7_domain_1(u1_struct_0(esk1_0),X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_31116, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X1))|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,k2_yellow_0(esk1_0,k1_tarski(X1))),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_30846, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2)),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X1,X2))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_30994, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,X2),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(X1,X2))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk2_0)|~m1_subset_1(k1_tarski(k12_lattice3(esk1_0,X1,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2355, negated_conjecture, (r2_hidden(esk4_1(k2_xboole_0(esk70_3(esk1_0,esk2_0,X1),k1_tarski(esk69_3(esk1_0,esk2_0,X1)))),esk70_3(esk1_0,esk2_0,X1))|r2_hidden(esk4_1(k2_xboole_0(esk70_3(esk1_0,esk2_0,X1),k1_tarski(esk69_3(esk1_0,esk2_0,X1)))),k1_tarski(esk69_3(esk1_0,esk2_0,X1)))|v1_xboole_0(k2_xboole_0(esk70_3(esk1_0,esk2_0,X1),k1_tarski(esk69_3(esk1_0,esk2_0,X1))))|v5_waybel_6(esk2_0,esk1_0)|r2_hidden(esk72_3(esk1_0,esk2_0,X1),X1)|v1_xboole_0(X1)|~v1_finset_1(k2_xboole_0(esk70_3(esk1_0,esk2_0,X1),k1_tarski(esk69_3(esk1_0,esk2_0,X1))))|~m1_subset_1(k2_xboole_0(esk70_3(esk1_0,esk2_0,X1),k1_tarski(esk69_3(esk1_0,esk2_0,X1))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_31304, negated_conjecture, (v6_ordinal1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v6_ordinal1(k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_31305, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_relat_1(k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_31306, negated_conjecture, (v5_finset_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v5_finset_1(k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_31307, negated_conjecture, (v4_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v4_funct_1(k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_31368, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_xboole_0),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k1_tarski(k2_yellow_0(esk1_0,k1_xboole_0)))|~r2_yellow_0(esk1_0,k1_xboole_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(k2_yellow_0(esk1_0,k1_xboole_0))),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,k1_xboole_0),esk3_0))).#trainneg
cnf(i_0_31110, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X1))|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,k2_yellow_0(esk1_0,k2_tarski(X1,X1))),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2356, negated_conjecture, (r2_hidden(esk4_1(k2_xboole_0(esk70_3(esk1_0,esk2_0,X1),k1_tarski(esk69_3(esk1_0,esk2_0,X1)))),esk70_3(esk1_0,esk2_0,X1))|r2_hidden(esk4_1(k2_xboole_0(esk70_3(esk1_0,esk2_0,X1),k1_tarski(esk69_3(esk1_0,esk2_0,X1)))),k1_tarski(esk69_3(esk1_0,esk2_0,X1)))|v1_xboole_0(k2_xboole_0(esk70_3(esk1_0,esk2_0,X1),k1_tarski(esk69_3(esk1_0,esk2_0,X1))))|v5_waybel_6(esk2_0,esk1_0)|v1_xboole_0(X1)|m1_subset_1(esk72_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0))|~v1_finset_1(k2_xboole_0(esk70_3(esk1_0,esk2_0,X1),k1_tarski(esk69_3(esk1_0,esk2_0,X1))))|~m1_subset_1(k2_xboole_0(esk70_3(esk1_0,esk2_0,X1),k1_tarski(esk69_3(esk1_0,esk2_0,X1))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_31308, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,esk2_0,esk3_0))|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_relat_1(k1_numbers)|~v1_funct_1(k1_numbers)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_31319, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|v1_xboole_0(X1)|~r2_yellow_0(esk1_0,k1_tarski(k2_yellow_0(esk1_0,X1)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(k2_yellow_0(esk1_0,X1))),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,X1),esk3_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_30851, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|v1_xboole_0(k7_domain_1(u1_struct_0(esk1_0),X1,X1))|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X1)),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~v1_finset_1(k7_domain_1(u1_struct_0(esk1_0),X1,X1))|~m1_subset_1(k7_domain_1(u1_struct_0(esk1_0),X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_31322, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X1)),esk2_0)|~m1_subset_1(X1,esk3_0)|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_31429, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X1),k1_tarski(X2)))|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k2_xboole_0(k1_tarski(X1),k1_tarski(X2)),k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X1,X2))))),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X1,X2))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_31403, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X2),k1_tarski(X1)))|~r2_yellow_0(esk1_0,k1_tarski(X2))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k2_xboole_0(k1_tarski(X2),k1_tarski(X1)),k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X1,X2))))),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X1,X2))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2357, negated_conjecture, (r2_hidden(esk4_1(k2_xboole_0(esk70_3(esk1_0,esk2_0,X1),k1_tarski(esk69_3(esk1_0,esk2_0,X1)))),esk70_3(esk1_0,esk2_0,X1))|r2_hidden(esk4_1(k2_xboole_0(esk70_3(esk1_0,esk2_0,X1),k1_tarski(esk69_3(esk1_0,esk2_0,X1)))),k1_tarski(esk69_3(esk1_0,esk2_0,X1)))|v1_xboole_0(k2_xboole_0(esk70_3(esk1_0,esk2_0,X1),k1_tarski(esk69_3(esk1_0,esk2_0,X1))))|v5_waybel_6(esk2_0,esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(X1)|~v1_finset_1(k2_xboole_0(esk70_3(esk1_0,esk2_0,X1),k1_tarski(esk69_3(esk1_0,esk2_0,X1))))|~m1_subset_1(k2_xboole_0(esk70_3(esk1_0,esk2_0,X1),k1_tarski(esk69_3(esk1_0,esk2_0,X1))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_21157, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_30995, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,X2),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,X1)),esk2_0)|~m1_subset_1(k1_tarski(k12_lattice3(esk1_0,X1,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_31129, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X1))|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,k12_lattice3(esk1_0,X1,X1))),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k12_lattice3(esk1_0,X1,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_31461, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(X1,X2))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k2_tarski(X1,X2),k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X1,X2))))),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X1,X2))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_31481, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,X1),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,k12_lattice3(esk1_0,X1,X1)),esk2_0)|~m1_subset_1(k1_tarski(k12_lattice3(esk1_0,X1,X1)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k12_lattice3(esk1_0,X1,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_28175, plain, (r1_orders_2(X1,k2_yellow_0(X1,esk70_3(X1,X2,X3)),X2)|v1_xboole_0(esk70_3(X1,X2,X3))|v1_xboole_0(X3)|m1_subset_1(esk72_3(X1,X2,X3),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(esk70_3(X1,X2,X3))|~v1_finset_1(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(esk70_3(X1,X2,X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1)))).#trainneg
cnf(i_0_2393, plain, (r1_tarski(esk70_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),u1_struct_0(X1))|v1_xboole_0(u1_struct_0(X1))|v1_xboole_0(k6_domain_1(u1_struct_0(X1),X3))|m1_subset_1(esk72_3(X1,X2,k6_domain_1(u1_struct_0(X1),X3)),u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1))|~r1_orders_2(X1,X3,X2)|~v1_finset_1(k6_domain_1(u1_struct_0(X1),X3))|~m1_subset_1(k6_domain_1(u1_struct_0(X1),X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_27684, negated_conjecture, (v1_xxreal_0(esk72_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_finset_1(esk72_3(esk1_0,esk2_0,esk3_0))|~v3_ordinal1(esk3_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_27388, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X2,k2_yellow_0(esk1_0,X1)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_xboole_0(X1,k2_tarski(X2,X2)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k2_xboole_0(X1,k2_tarski(X2,X2)),k1_tarski(k12_lattice3(esk1_0,X2,k2_yellow_0(esk1_0,X1))))),esk2_0)|~m1_subset_1(k1_tarski(k12_lattice3(esk1_0,X2,k2_yellow_0(esk1_0,X1))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k12_lattice3(esk1_0,X2,k2_yellow_0(esk1_0,X1)),u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,X1)|~m1_subset_1(k2_tarski(X2,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_31596, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,X2),esk2_0)|v1_xboole_0(k7_domain_1(u1_struct_0(esk1_0),X1,X2))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2)),esk2_0)|~v1_finset_1(k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~m1_subset_1(k7_domain_1(u1_struct_0(esk1_0),X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k12_lattice3(esk1_0,X1,X2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_31597, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,X2),esk2_0)|v1_xboole_0(k7_domain_1(u1_struct_0(esk1_0),X1,X2))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2)),esk2_0)|~v1_finset_1(k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~m1_subset_1(k7_domain_1(u1_struct_0(esk1_0),X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k12_lattice3(esk1_0,X1,X2),esk3_0))).#trainneg
cnf(i_0_31599, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,k2_yellow_0(esk1_0,k1_tarski(X1))),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_30107, negated_conjecture, (v1_xxreal_0(esk2_0)|~r2_hidden(X1,esk3_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(X1,u1_struct_0(esk1_0))|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_2395, plain, (r1_tarski(X1,X2)|v1_xboole_0(X2)|~m1_subset_1(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(k6_domain_1(X2,X3))))).#trainneg
cnf(i_0_31499, plain, (k2_xboole_0(esk47_1(X1),X2)=X2|~v1_xboole_0(X1))).#trainneg
cnf(i_0_27685, negated_conjecture, (v1_card_1(esk72_3(esk1_0,esk2_0,esk3_0))|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~v1_finset_1(esk72_3(esk1_0,esk2_0,esk3_0))|~v3_ordinal1(esk3_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_21158, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|~m1_subset_1(u1_struct_0(esk1_0),k4_ordinal1))).#trainneg
cnf(i_0_31611, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,X2),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk2_0)|~m1_subset_1(k1_tarski(k12_lattice3(esk1_0,X1,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_31582, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X2,X1),esk2_0)|v1_xboole_0(k7_domain_1(u1_struct_0(esk1_0),X1,X2))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X2)),esk2_0)|~v1_finset_1(k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~m1_subset_1(k1_tarski(k12_lattice3(esk1_0,X2,X1)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k7_domain_1(u1_struct_0(esk1_0),X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2400, plain, (v2_struct_0(X1)|r1_orders_2(X1,esk61_1(X1),esk61_1(X1))|~l1_orders_2(X1)|~v3_orders_2(X1)|~v2_yellow_0(X1)|~v5_orders_2(X1))).#trainneg
cnf(i_0_32077, negated_conjecture, (v1_xxreal_0(esk2_0)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),esk3_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_32099, negated_conjecture, (v1_xxreal_0(esk2_0)|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_32078, negated_conjecture, (v1_xxreal_0(esk2_0)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),k1_numbers)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_32079, negated_conjecture, (v1_xxreal_0(esk2_0)|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_2401, plain, (v2_struct_0(X1)|r1_orders_2(X1,esk12_1(u1_struct_0(X1)),esk12_1(u1_struct_0(X1)))|~l1_orders_2(X1)|~v3_orders_2(X1))).#trainneg
cnf(i_0_32080, negated_conjecture, (v1_xxreal_0(esk2_0)|m1_subset_1(esk69_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_32100, negated_conjecture, (v1_xxreal_0(esk2_0)|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_32101, negated_conjecture, (v1_xxreal_0(esk2_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_32086, negated_conjecture, (v1_xxreal_0(esk2_0)|v2_struct_0(esk1_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1)|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_32082, negated_conjecture, (v1_xxreal_0(esk2_0)|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1)|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_32060, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X1)))|~r2_yellow_0(esk1_0,k1_tarski(k2_yellow_0(esk1_0,X1)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X1,k1_tarski(k2_yellow_0(esk1_0,X1)))),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,X1)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r2_yellow_0(esk1_0,X1))).#trainneg
cnf(i_0_31588, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v1_xboole_0(k7_domain_1(u1_struct_0(esk1_0),X1,X1))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X1)),esk2_0)|~v1_finset_1(k7_domain_1(u1_struct_0(esk1_0),X1,X1))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k7_domain_1(u1_struct_0(esk1_0),X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_32039, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,X2),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk2_0)|~m1_subset_1(k2_tarski(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k12_lattice3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2412, plain, (k2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,k2_xboole_0(X3,X4))))=k2_yellow_0(X1,k2_xboole_0(X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))))|~r2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,X3))).#trainneg
cnf(i_0_32081, negated_conjecture, (v1_xxreal_0(esk2_0)|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_32085, negated_conjecture, (v1_xxreal_0(esk2_0)|m1_subset_1(esk70_3(esk1_0,esk2_0,esk3_0),k1_zfmisc_1(k1_numbers))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1)|~m1_subset_1(u1_struct_0(esk1_0),k1_zfmisc_1(k4_ordinal1)))).#trainneg
cnf(i_0_32083, negated_conjecture, (v1_xxreal_0(esk2_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_32095, negated_conjecture, (v1_xxreal_0(esk2_0)|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_31868, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X1,X1))|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,k2_yellow_0(esk1_0,k2_tarski(X1,X1)))),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(X1,X1)),u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_32599, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,X2),esk2_0)|v1_xboole_0(k7_domain_1(u1_struct_0(esk1_0),X2,X1))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,X1)),esk2_0)|~v1_finset_1(k7_domain_1(u1_struct_0(esk1_0),X2,X1))|~m1_subset_1(k7_domain_1(u1_struct_0(esk1_0),X2,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(k12_lattice3(esk1_0,X1,X2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_32534, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X2,X1)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2))),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_xboole_0(X2,X1))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,X1)|~r2_yellow_0(esk1_0,X2))).#trainneg
cnf(i_0_32198, plain, (k12_lattice3(X1,X3,X2)=k2_yellow_0(X1,k1_tarski(X3))|~r2_yellow_0(X1,k6_domain_1(u1_struct_0(X1),X2))|~m1_subset_1(k1_tarski(X3),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_xboole_0(k6_domain_1(u1_struct_0(X1),X2)))).#trainneg
cnf(i_0_32533, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X1,X2)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2))),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_xboole_0(X1,X2))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,X1))).#trainneg
cnf(i_0_2413, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,k2_xboole_0(X3,X4)))=k2_yellow_0(X1,k2_xboole_0(X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))))|~r2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,X3)|~v3_orders_2(X1))).#trainneg
cnf(i_0_32619, negated_conjecture, (v1_xxreal_0(esk2_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1)|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0))).#trainneg
cnf(i_0_32097, negated_conjecture, (v1_xxreal_0(esk2_0)|r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_32075, negated_conjecture, (v1_xxreal_0(esk2_0)|m1_subset_1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_32708, plain, (r1_orders_2(X1,esk69_3(X1,X2,X3),X2)|r1_orders_2(X1,esk72_3(X1,X2,X3),X2)|v1_xboole_0(X3)|~v5_waybel_6(X2,X1)|~m1_subset_1(k1_tarski(esk69_3(X1,X2,X3)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(esk69_3(X1,X2,X3),u1_struct_0(X1))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r1_orders_2(X1,k2_yellow_0(X1,X3),X2)|~v1_finset_1(X3)|~v1_xboole_0(esk70_3(X1,X2,X3))|~m1_subset_1(X3,k1_zfmisc_1(u1_struct_0(X1))))).#trainneg
cnf(i_0_32694, negated_conjecture, (r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,X1),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(X1)|~m1_subset_1(k1_tarski(esk69_3(esk1_0,esk2_0,X1)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_32877, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X1,X2)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2))),esk2_0)|~m1_subset_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_xboole_0(X1,X2)),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,X1))).#trainneg
cnf(i_0_32878, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X2,X1)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2))),esk2_0)|~m1_subset_1(k2_tarski(k2_yellow_0(esk1_0,X1),k2_yellow_0(esk1_0,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_xboole_0(X2,X1)),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,X1)|~r2_yellow_0(esk1_0,X2))).#trainneg
cnf(i_0_2416, plain, (r2_hidden(esk69_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))),k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)))|m1_subset_1(esk72_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))),u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X3,X4)),X2)|~m1_subset_1(k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,X3))).#trainneg
cnf(i_0_32328, plain, (k2_xboole_0(X2,X1)=X2|~v1_xboole_0(X1))).#trainneg
cnf(i_0_33395, negated_conjecture, (v1_xxreal_0(esk2_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1)|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_33515, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_33487, negated_conjecture, (r1_orders_2(esk1_0,esk69_3(esk1_0,X1,esk3_0),X1)|r2_hidden(esk72_3(esk1_0,X1,esk3_0),esk3_0)|~v5_waybel_6(X1,esk1_0)|~m1_subset_1(k1_tarski(esk69_3(esk1_0,X1,esk3_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,X1,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),X1)|~v1_xboole_0(esk70_3(esk1_0,X1,esk3_0)))).#trainneg
cnf(i_0_33507, negated_conjecture, (r2_hidden(esk72_3(esk1_0,X1,esk3_0),esk3_0)|r1_orders_2(esk1_0,esk69_3(esk1_0,X1,esk3_0),X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),X1)|~v1_xboole_0(esk70_3(esk1_0,X1,esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,X1,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_33516, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,X1),esk2_0)|r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,X1),esk2_0)|v1_xboole_0(X1)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X1),esk2_0)|~v1_finset_1(X1)|~v1_xboole_0(esk70_3(esk1_0,esk2_0,X1))|~m1_subset_1(X1,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk69_3(esk1_0,esk2_0,X1),esk3_0))).#trainneg
cnf(i_0_2417, plain, (r1_tarski(esk70_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))),k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)))|m1_subset_1(esk72_3(X1,X2,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4))),u1_struct_0(X1))|~r1_orders_2(X1,k2_yellow_0(X1,k2_xboole_0(X3,X4)),X2)|~m1_subset_1(k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)),k1_zfmisc_1(u1_struct_0(X1)))|~m1_subset_1(X2,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,X3))).#trainneg
cnf(i_0_33379, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|v1_xxreal_0(esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_33465, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,esk2_0,k1_tarski(X1)),esk2_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X1)),esk2_0)|~v1_xboole_0(esk70_3(esk1_0,esk2_0,k1_tarski(X1)))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,X1,esk2_0)|~v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_33572, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X1),X2)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,k2_yellow_0(esk1_0,X2))),esk2_0)|~m1_subset_1(k2_tarski(X1,k2_yellow_0(esk1_0,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X1),X2)),u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2418, plain, (k12_lattice3(X1,k2_yellow_0(X1,X2),k2_yellow_0(X1,k2_xboole_0(X3,X4)))=k2_yellow_0(X1,k2_xboole_0(k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)),X2))|~r2_yellow_0(X1,X2)|~r2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X3),k2_yellow_0(X1,X4)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,X3)|~v3_orders_2(X1))).#trainneg
cnf(i_0_11495, negated_conjecture, (m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),k1_numbers)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(esk3_0,k1_zfmisc_1(k4_ordinal1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_33636, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X1),X2)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,X2)|~r2_yellow_0(esk1_0,k1_tarski(X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(k2_yellow_0(esk1_0,X2),X1)),esk2_0)|~m1_subset_1(k2_tarski(k2_yellow_0(esk1_0,X2),X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X1),X2)),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_34206, plain, (r1_orders_2(X1,k2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2))),X3)|r1_orders_2(X1,k2_yellow_0(X1,X4),X3)|~r2_yellow_0(X1,k1_tarski(k2_yellow_0(X1,X2)))|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,X2)|~r1_orders_2(X1,k2_yellow_0(X1,X2),X3)|~v5_waybel_6(X3,X1)|~m1_subset_1(X3,u1_struct_0(X1))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~v1_xboole_0(X4))).#trainneg
cnf(i_0_32889, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X2,X1)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk2_0)|~m1_subset_1(k2_tarski(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(X2,X1)),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2419, plain, (k2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,k2_xboole_0(X2,X3)),k2_yellow_0(X1,X4)))=k2_yellow_0(X1,k2_xboole_0(k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3)),X4))|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2))).#trainneg
cnf(i_0_34016, plain, (esk47_1(X1)=X2|~v1_xboole_0(X1)|~v1_xboole_0(X2))).#trainneg
cnf(i_0_32090, negated_conjecture, (v1_xxreal_0(esk2_0)|m1_subset_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),k1_zfmisc_1(u1_struct_0(esk1_0)))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_13546, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~v1_card_1(esk3_0))).#trainneg
cnf(i_0_32550, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X2,k1_tarski(X1))),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,k2_yellow_0(esk1_0,X2))),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_xboole_0(X2,k1_tarski(X1)))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_tarski(X1,k2_yellow_0(esk1_0,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,X2)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_31605, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k7_domain_1(u1_struct_0(esk1_0),X2,X1)),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X1,X2))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_31530, plain, (esk7_3(k2_tarski(X1,X2),X3,X3)=X2|esk7_3(k2_tarski(X1,X2),X3,X3)=X1|k2_xboole_0(k2_tarski(X1,X2),X3)=X3)).#trainneg
cnf(i_0_32841, negated_conjecture, (r1_tarski(esk70_3(esk1_0,esk2_0,k2_tarski(X1,X2)),k2_tarski(X1,X2))|r2_hidden(esk72_3(esk1_0,esk2_0,k2_tarski(X1,X2)),k2_tarski(X1,X2))|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk2_0)|~m1_subset_1(k2_tarski(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k12_lattice3(esk1_0,X1,X2),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2420, plain, (k12_lattice3(X1,k2_yellow_0(X1,k2_xboole_0(X2,X3)),k2_yellow_0(X1,X4))=k2_yellow_0(X1,k2_xboole_0(k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3)),X4))|~r2_yellow_0(X1,X4)|~r2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3)))|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2)|~v3_orders_2(X1))).#trainneg
cnf(i_0_33136, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(X1,k2_tarski(X2,X2))),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(X2,X2))|~r2_yellow_0(esk1_0,X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X2,k2_yellow_0(esk1_0,X1))),esk2_0)|~m1_subset_1(k2_tarski(X2,k2_yellow_0(esk1_0,X1)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_xboole_0(X1,k2_tarski(X2,X2)))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_31924, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(X1,X1))|~r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,X1),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2422, plain, (k12_lattice3(X1,k2_yellow_0(X1,k2_xboole_0(X2,X3)),k2_yellow_0(X1,X4))=k2_yellow_0(X1,k2_xboole_0(X4,k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3))))|~r2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X2),k2_yellow_0(X1,X3)))|~r2_yellow_0(X1,X4)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~r2_yellow_0(X1,X3)|~r2_yellow_0(X1,X2)|~v3_orders_2(X1))).#trainneg
cnf(i_0_13553, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~v1_card_1(u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_32545, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X2,X1)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X2,X1))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_tarski(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_32549, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X1),X2)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,k2_yellow_0(esk1_0,X2))),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X1),X2))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_tarski(X1,k2_yellow_0(esk1_0,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~r2_yellow_0(esk1_0,X2)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_35162, negated_conjecture, (v2_struct_0(esk1_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0))).#trainneg
cnf(i_0_34814, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X1)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(X1,X1)),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_34226, negated_conjecture, (v1_xxreal_0(esk72_3(esk1_0,esk2_0,esk3_0))|v5_waybel_6(esk2_0,esk1_0)|~v1_finset_1(esk72_3(esk1_0,esk2_0,esk3_0))|~m1_subset_1(k1_tarski(esk72_3(esk1_0,esk2_0,esk3_0)),k4_ordinal1)|~m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_34944, negated_conjecture, (v1_finset_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|v1_xxreal_0(esk2_0)|~v1_finset_1(u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_34248, negated_conjecture, (r2_hidden(esk72_3(esk1_0,X1,esk3_0),esk3_0)|r1_orders_2(esk1_0,esk69_3(esk1_0,X1,esk3_0),X1)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),X1)|~v1_xboole_0(esk70_3(esk1_0,X1,esk3_0))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,X1,esk3_0),esk3_0))).#trainneg
cnf(i_0_34252, negated_conjecture, (r2_hidden(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk3_0)|r1_orders_2(esk1_0,esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),k2_yellow_0(esk1_0,esk3_0))|r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk3_0),k2_yellow_0(esk1_0,esk3_0))|~v1_xboole_0(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,esk70_3(esk1_0,esk2_0,esk3_0)),esk2_0)|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_2427, plain, (k2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X2),k12_lattice3(X1,X3,X4)))=k2_yellow_0(X1,k2_xboole_0(X2,k7_domain_1(u1_struct_0(X1),X3,X4)))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X4,u1_struct_0(X1))|~m1_subset_1(X3,u1_struct_0(X1)))).#trainneg
cnf(i_0_34945, negated_conjecture, (v1_zfmisc_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|v1_xxreal_0(esk2_0)|~v1_zfmisc_1(u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_34946, negated_conjecture, (v6_ordinal1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|v1_xxreal_0(esk2_0)|~v6_ordinal1(u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_34947, negated_conjecture, (v1_relat_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|v1_xxreal_0(esk2_0)|~v1_relat_1(u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_34948, negated_conjecture, (v5_finset_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|v1_xxreal_0(esk2_0)|~v5_finset_1(u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_34949, negated_conjecture, (v4_funct_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|v1_xxreal_0(esk2_0)|~v4_funct_1(u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_35195, negated_conjecture, (r1_orders_2(esk1_0,X1,esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,X1),esk2_0)|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_tarski(X1,X1),k1_zfmisc_1(u1_struct_0(esk1_0))))).#trainneg
cnf(i_0_35353, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X1),X2)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,k2_yellow_0(esk1_0,X2))),esk2_0)|~m1_subset_1(k2_tarski(X1,k2_yellow_0(esk1_0,X2)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,X2),u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(X1),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k2_yellow_0(esk1_0,k2_xboole_0(k1_tarski(X1),X2)),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_34994, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r2_yellow_0(esk1_0,k2_tarski(X2,X1))|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X2,X1)),esk2_0)|~m1_subset_1(k1_tarski(k2_yellow_0(esk1_0,k2_tarski(X1,X2))),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_34815, negated_conjecture, (r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(esk9_2(esk1_0,esk2_0),esk8_2(esk1_0,esk2_0))),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(k2_yellow_0(esk1_0,k2_tarski(esk9_2(esk1_0,esk2_0),esk8_2(esk1_0,esk2_0))),u1_struct_0(esk1_0))|~m1_subset_1(k2_tarski(esk8_2(esk1_0,esk2_0),esk9_2(esk1_0,esk2_0)),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(esk9_2(esk1_0,esk2_0),u1_struct_0(esk1_0))|~m1_subset_1(esk8_2(esk1_0,esk2_0),u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_34253, negated_conjecture, (r2_hidden(esk72_3(esk1_0,esk2_0,esk3_0),esk3_0)|r1_orders_2(esk1_0,esk69_3(esk1_0,esk2_0,esk3_0),esk2_0)|m1_subset_1(esk72_3(esk1_0,esk2_0,esk3_0),u1_struct_0(esk1_0))|~v1_xboole_0(esk70_3(esk1_0,esk2_0,esk3_0)))).#trainneg
cnf(i_0_34503, negated_conjecture, (v2_struct_0(esk1_0)|~m1_subset_1(esk72_3(esk1_0,esk2_0,k1_tarski(X1)),u1_struct_0(esk1_0))|~m1_subset_1(esk72_3(esk1_0,esk2_0,k1_tarski(X1)),esk3_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k1_tarski(X1)),esk2_0)|~r1_orders_2(esk1_0,X1,esk2_0)|~v1_xboole_0(esk70_3(esk1_0,esk2_0,k1_tarski(X1)))|~v5_waybel_6(esk2_0,esk1_0)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_32044, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X1,X2),esk2_0)|v1_xboole_0(k7_domain_1(u1_struct_0(esk1_0),X1,X2))|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k12_lattice3(esk1_0,X2,X1),esk2_0)|~v1_finset_1(k7_domain_1(u1_struct_0(esk1_0),X1,X2))|~m1_subset_1(k7_domain_1(u1_struct_0(esk1_0),X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k12_lattice3(esk1_0,X1,X2),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_2428, plain, (k2_yellow_0(X1,k2_tarski(k2_yellow_0(X1,X2),k12_lattice3(X1,X4,X3)))=k2_yellow_0(X1,k2_xboole_0(X2,k7_domain_1(u1_struct_0(X1),X3,X4)))|v1_xboole_0(u1_struct_0(X1))|~r2_yellow_0(X1,k7_domain_1(u1_struct_0(X1),X3,X4))|~r2_yellow_0(X1,X2)|~l1_orders_2(X1)|~v2_lattice3(X1)|~v5_orders_2(X1)|~v4_orders_2(X1)|~v3_orders_2(X1)|~m1_subset_1(X3,u1_struct_0(X1))|~m1_subset_1(X4,u1_struct_0(X1)))).#trainneg
cnf(i_0_34950, negated_conjecture, (v1_funct_1(esk70_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|v1_xxreal_0(esk2_0)|~v1_relat_1(u1_struct_0(esk1_0))|~v1_funct_1(u1_struct_0(esk1_0))|~r1_orders_2(esk1_0,esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),esk2_0)|~v1_finset_1(esk2_0)|~m1_subset_1(k2_yellow_0(esk1_0,esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(k1_tarski(esk2_0),k4_ordinal1))).#trainneg
cnf(i_0_16580, negated_conjecture, (v3_ordinal1(esk69_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0))|m1_subset_1(esk72_3(esk1_0,k2_yellow_0(esk1_0,esk3_0),esk3_0),u1_struct_0(esk1_0))|~m1_subset_1(esk3_0,k4_ordinal1))).#trainneg
cnf(i_0_33443, negated_conjecture, (r1_orders_2(esk1_0,esk72_3(esk1_0,X1,X2),X1)|r1_orders_2(esk1_0,esk69_3(esk1_0,X1,X2),X1)|v1_xboole_0(X2)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,X2),X1)|~v1_finset_1(X2)|~v1_xboole_0(esk70_3(esk1_0,X1,X2))|~v5_waybel_6(X1,esk1_0)|~m1_subset_1(X2,k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(X1,u1_struct_0(esk1_0))|~m1_subset_1(esk69_3(esk1_0,X1,X2),esk3_0))).#trainneg
cnf(i_0_32890, negated_conjecture, (r1_orders_2(esk1_0,k12_lattice3(esk1_0,X2,X1),esk2_0)|v5_waybel_6(esk2_0,esk1_0)|~r1_orders_2(esk1_0,k2_yellow_0(esk1_0,k2_tarski(X1,X2)),esk2_0)|~m1_subset_1(k2_tarski(X1,X2),k1_zfmisc_1(u1_struct_0(esk1_0)))|~m1_subset_1(k12_lattice3(esk1_0,X2,X1),u1_struct_0(esk1_0))|~m1_subset_1(X2,u1_struct_0(esk1_0))|~m1_subset_1(X1,u1_struct_0(esk1_0)))).#trainneg
cnf(i_0_35356, negated_conjecture, (esk70_3(esk1_0,esk2_0,esk3_0)=k1_xboole_0|v2_struct_0(esk1_0))).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 258
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 611
# Removed in clause preprocessing      : 114
# Initial clauses in saturation        : 497
# Processed clauses                    : 4584
# ...of these trivial                  : 18
# ...subsumed                          : 2200
# ...remaining for further processing  : 2366
# Other redundant clauses eliminated   : 43
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 191
# Backward-rewritten                   : 187
# Generated clauses                    : 35102
# ...of the previous two non-trivial   : 14443
# Contextual simplify-reflections      : 521
# Paramodulations                      : 34977
# Factorizations                       : 53
# Equation resolutions                 : 72
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1985
#    Positive orientable unit clauses  : 162
#    Positive unorientable unit clauses: 2
#    Negative unit clauses             : 60
#    Non-unit-clauses                  : 1761
# Current number of unprocessed clauses: 10191
# ...number of literals in the above   : 104463
# Current number of archived formulas  : 0
# Current number of archived clauses   : 2744
# Proof object given clauses           : 218
# Proof search given clauses           : 2366
# Clause-clause subsumption calls (NU) : 3017268
# Rec. Clause-clause subsumption calls : 247122
# Non-unit clause-clause subsumptions  : 2650
# Unit Clause-clause subsumption calls : 17849
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 57
# BW rewrite match successes           : 18
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 1704847

# -------------------------------------------------
# User time                : 22.559 s
# System time              : 0.818 s
# Total time               : 23.377 s
# Maximum resident set size: 214928 pages
22.55user 0.84system 1:02.58elapsed 37%CPU (0avgtext+0avgdata 214928maxresident)k
0inputs+1672outputs (0major+202192minor)pagefaults 0swaps
