2020-05-06 23:21:40.554865: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_8/premsel_enigma_01_2020_T10_loop01_epoch_23
2020-05-06 23:21:40.633087: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-05-06 23:21:40.758713: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-05-06 23:21:41.220694: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-05-06 23:21:42.548664: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 1993796 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_8/premsel_enigma_01_2020_T10_loop01_epoch_23' loaded (query=256; context=1024; weigths=real; len_mult=0.000000).
# Preprocessing time       : 2.548 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(cc8_ordinal1, axiom, ![X1]:(m1_subset_1(X1,k4_ordinal1)=>v7_ordinal1(X1)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', cc8_ordinal1)).
fof(cc17_membered, axiom, ![X1]:(v6_membered(X1)=>![X2]:(m1_subset_1(X2,X1)=>v7_ordinal1(X2))), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', cc17_membered)).
fof(fc6_membered, axiom, v6_membered(k4_ordinal1), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', fc6_membered)).
fof(redefinition_k5_numbers, axiom, k5_numbers=k4_ordinal1, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', redefinition_k5_numbers)).
fof(commutativity_k1_nat_1, axiom, ![X1, X2]:((v7_ordinal1(X1)&m1_subset_1(X2,k5_numbers))=>k1_nat_1(X1,X2)=k1_nat_1(X2,X1)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', commutativity_k1_nat_1)).
fof(redefinition_k1_nat_1, axiom, ![X1, X2]:((v7_ordinal1(X1)&m1_subset_1(X2,k5_numbers))=>k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', redefinition_k1_nat_1)).
fof(spc1_numerals, axiom, (((v2_xxreal_0(np__1)&m2_subset_1(np__1,k1_numbers,k5_numbers))&m1_subset_1(np__1,k5_numbers))&m1_subset_1(np__1,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc1_numerals)).
fof(fc2_nat_1, axiom, ![X1, X2]:((v7_ordinal1(X1)&v7_ordinal1(X2))=>v7_ordinal1(k3_xcmplx_0(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', fc2_nat_1)).
fof(spc2_numerals, axiom, (((v2_xxreal_0(np__2)&m2_subset_1(np__2,k1_numbers,k5_numbers))&m1_subset_1(np__2,k5_numbers))&m1_subset_1(np__2,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc2_numerals)).
fof(spc4_numerals, axiom, (((v2_xxreal_0(np__4)&m2_subset_1(np__4,k1_numbers,k5_numbers))&m1_subset_1(np__4,k5_numbers))&m1_subset_1(np__4,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc4_numerals)).
fof(rqRealMult__k3_xcmplx_0__r4_r4_r16, axiom, k3_xcmplx_0(np__4,np__4)=np__16, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealMult__k3_xcmplx_0__r4_r4_r16)).
fof(spc38_numerals, axiom, (((v2_xxreal_0(np__38)&m2_subset_1(np__38,k1_numbers,k5_numbers))&m1_subset_1(np__38,k5_numbers))&m1_subset_1(np__38,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc38_numerals)).
fof(spc37_numerals, axiom, (((v2_xxreal_0(np__37)&m2_subset_1(np__37,k1_numbers,k5_numbers))&m1_subset_1(np__37,k5_numbers))&m1_subset_1(np__37,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc37_numerals)).
fof(spc36_numerals, axiom, (((v2_xxreal_0(np__36)&m2_subset_1(np__36,k1_numbers,k5_numbers))&m1_subset_1(np__36,k5_numbers))&m1_subset_1(np__36,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc36_numerals)).
fof(spc3_numerals, axiom, (((v2_xxreal_0(np__3)&m2_subset_1(np__3,k1_numbers,k5_numbers))&m1_subset_1(np__3,k5_numbers))&m1_subset_1(np__3,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc3_numerals)).
fof(spc63_numerals, axiom, (((v2_xxreal_0(np__63)&m2_subset_1(np__63,k1_numbers,k5_numbers))&m1_subset_1(np__63,k5_numbers))&m1_subset_1(np__63,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc63_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r1_r2_r3, axiom, k2_xcmplx_0(np__1,np__2)=np__3, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r2_r3)).
fof(rqRealAdd__k2_xcmplx_0__r1_r4_r5, axiom, k2_xcmplx_0(np__1,np__4)=np__5, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r4_r5)).
fof(rqRealAdd__k2_xcmplx_0__r1_r16_r17, axiom, k2_xcmplx_0(np__1,np__16)=np__17, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r16_r17)).
fof(spc16_numerals, axiom, (((v2_xxreal_0(np__16)&m2_subset_1(np__16,k1_numbers,k5_numbers))&m1_subset_1(np__16,k5_numbers))&m1_subset_1(np__16,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc16_numerals)).
fof(spc40_numerals, axiom, (((v2_xxreal_0(np__40)&m2_subset_1(np__40,k1_numbers,k5_numbers))&m1_subset_1(np__40,k5_numbers))&m1_subset_1(np__40,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc40_numerals)).
fof(spc29_numerals, axiom, (((v2_xxreal_0(np__29)&m2_subset_1(np__29,k1_numbers,k5_numbers))&m1_subset_1(np__29,k5_numbers))&m1_subset_1(np__29,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc29_numerals)).
fof(spc30_numerals, axiom, (((v2_xxreal_0(np__30)&m2_subset_1(np__30,k1_numbers,k5_numbers))&m1_subset_1(np__30,k5_numbers))&m1_subset_1(np__30,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc30_numerals)).
fof(spc34_numerals, axiom, (((v2_xxreal_0(np__34)&m2_subset_1(np__34,k1_numbers,k5_numbers))&m1_subset_1(np__34,k5_numbers))&m1_subset_1(np__34,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc34_numerals)).
fof(spc35_numerals, axiom, (((v2_xxreal_0(np__35)&m2_subset_1(np__35,k1_numbers,k5_numbers))&m1_subset_1(np__35,k5_numbers))&m1_subset_1(np__35,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc35_numerals)).
fof(spc61_numerals, axiom, (((v2_xxreal_0(np__61)&m2_subset_1(np__61,k1_numbers,k5_numbers))&m1_subset_1(np__61,k5_numbers))&m1_subset_1(np__61,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc61_numerals)).
fof(spc62_numerals, axiom, (((v2_xxreal_0(np__62)&m2_subset_1(np__62,k1_numbers,k5_numbers))&m1_subset_1(np__62,k5_numbers))&m1_subset_1(np__62,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc62_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r1_r63_r64, axiom, k2_xcmplx_0(np__1,np__63)=np__64, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r63_r64)).
fof(rqRealAdd__k2_xcmplx_0__r38_r1_r39, axiom, k2_xcmplx_0(np__38,np__1)=np__39, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r38_r1_r39)).
fof(rqRealAdd__k2_xcmplx_0__r37_r1_r38, axiom, k2_xcmplx_0(np__37,np__1)=np__38, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r37_r1_r38)).
fof(rqRealAdd__k2_xcmplx_0__r36_r1_r37, axiom, k2_xcmplx_0(np__36,np__1)=np__37, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r36_r1_r37)).
fof(rqRealAdd__k2_xcmplx_0__r1_r3_r4, axiom, k2_xcmplx_0(np__1,np__3)=np__4, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r3_r4)).
fof(spc45_numerals, axiom, (((v2_xxreal_0(np__45)&m2_subset_1(np__45,k1_numbers,k5_numbers))&m1_subset_1(np__45,k5_numbers))&m1_subset_1(np__45,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc45_numerals)).
fof(spc46_numerals, axiom, (((v2_xxreal_0(np__46)&m2_subset_1(np__46,k1_numbers,k5_numbers))&m1_subset_1(np__46,k5_numbers))&m1_subset_1(np__46,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc46_numerals)).
fof(spc44_numerals, axiom, (((v2_xxreal_0(np__44)&m2_subset_1(np__44,k1_numbers,k5_numbers))&m1_subset_1(np__44,k5_numbers))&m1_subset_1(np__44,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc44_numerals)).
fof(spc43_numerals, axiom, (((v2_xxreal_0(np__43)&m2_subset_1(np__43,k1_numbers,k5_numbers))&m1_subset_1(np__43,k5_numbers))&m1_subset_1(np__43,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc43_numerals)).
fof(spc28_numerals, axiom, (((v2_xxreal_0(np__28)&m2_subset_1(np__28,k1_numbers,k5_numbers))&m1_subset_1(np__28,k5_numbers))&m1_subset_1(np__28,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc28_numerals)).
fof(spc27_numerals, axiom, (((v2_xxreal_0(np__27)&m2_subset_1(np__27,k1_numbers,k5_numbers))&m1_subset_1(np__27,k5_numbers))&m1_subset_1(np__27,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc27_numerals)).
fof(spc26_numerals, axiom, (((v2_xxreal_0(np__26)&m2_subset_1(np__26,k1_numbers,k5_numbers))&m1_subset_1(np__26,k5_numbers))&m1_subset_1(np__26,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc26_numerals)).
fof(spc25_numerals, axiom, (((v2_xxreal_0(np__25)&m2_subset_1(np__25,k1_numbers,k5_numbers))&m1_subset_1(np__25,k5_numbers))&m1_subset_1(np__25,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc25_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r1_r40_r41, axiom, k2_xcmplx_0(np__1,np__40)=np__41, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r40_r41)).
fof(rqRealAdd__k2_xcmplx_0__r29_r1_r30, axiom, k2_xcmplx_0(np__29,np__1)=np__30, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r29_r1_r30)).
fof(rqRealAdd__k2_xcmplx_0__r30_r1_r31, axiom, k2_xcmplx_0(np__30,np__1)=np__31, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r30_r1_r31)).
fof(rqRealAdd__k2_xcmplx_0__r34_r1_r35, axiom, k2_xcmplx_0(np__34,np__1)=np__35, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r34_r1_r35)).
fof(rqRealAdd__k2_xcmplx_0__r35_r1_r36, axiom, k2_xcmplx_0(np__35,np__1)=np__36, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r35_r1_r36)).
fof(rqRealAdd__k2_xcmplx_0__r1_r61_r62, axiom, k2_xcmplx_0(np__1,np__61)=np__62, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r61_r62)).
fof(rqRealAdd__k2_xcmplx_0__r1_r62_r63, axiom, k2_xcmplx_0(np__1,np__62)=np__63, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r62_r63)).
fof(spc24_numerals, axiom, (((v2_xxreal_0(np__24)&m2_subset_1(np__24,k1_numbers,k5_numbers))&m1_subset_1(np__24,k5_numbers))&m1_subset_1(np__24,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc24_numerals)).
fof(spc5_numerals, axiom, (((v2_xxreal_0(np__5)&m2_subset_1(np__5,k1_numbers,k5_numbers))&m1_subset_1(np__5,k5_numbers))&m1_subset_1(np__5,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc5_numerals)).
fof(spc21_numerals, axiom, (((v2_xxreal_0(np__21)&m2_subset_1(np__21,k1_numbers,k5_numbers))&m1_subset_1(np__21,k5_numbers))&m1_subset_1(np__21,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc21_numerals)).
fof(spc22_numerals, axiom, (((v2_xxreal_0(np__22)&m2_subset_1(np__22,k1_numbers,k5_numbers))&m1_subset_1(np__22,k5_numbers))&m1_subset_1(np__22,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc22_numerals)).
fof(spc23_numerals, axiom, (((v2_xxreal_0(np__23)&m2_subset_1(np__23,k1_numbers,k5_numbers))&m1_subset_1(np__23,k5_numbers))&m1_subset_1(np__23,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc23_numerals)).
fof(spc47_numerals, axiom, (((v2_xxreal_0(np__47)&m2_subset_1(np__47,k1_numbers,k5_numbers))&m1_subset_1(np__47,k5_numbers))&m1_subset_1(np__47,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc47_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r1_r45_r46, axiom, k2_xcmplx_0(np__1,np__45)=np__46, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r45_r46)).
fof(rqRealAdd__k2_xcmplx_0__r1_r46_r47, axiom, k2_xcmplx_0(np__1,np__46)=np__47, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r46_r47)).
fof(rqRealAdd__k2_xcmplx_0__r1_r44_r45, axiom, k2_xcmplx_0(np__1,np__44)=np__45, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r44_r45)).
fof(rqRealAdd__k2_xcmplx_0__r1_r43_r44, axiom, k2_xcmplx_0(np__1,np__43)=np__44, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r43_r44)).
fof(rqRealAdd__k2_xcmplx_0__r28_r1_r29, axiom, k2_xcmplx_0(np__28,np__1)=np__29, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r28_r1_r29)).
fof(rqRealAdd__k2_xcmplx_0__r27_r1_r28, axiom, k2_xcmplx_0(np__27,np__1)=np__28, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r27_r1_r28)).
fof(rqRealAdd__k2_xcmplx_0__r26_r1_r27, axiom, k2_xcmplx_0(np__26,np__1)=np__27, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r26_r1_r27)).
fof(rqRealAdd__k2_xcmplx_0__r25_r1_r26, axiom, k2_xcmplx_0(np__25,np__1)=np__26, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r25_r1_r26)).
fof(spc7_numerals, axiom, (((v2_xxreal_0(np__7)&m2_subset_1(np__7,k1_numbers,k5_numbers))&m1_subset_1(np__7,k5_numbers))&m1_subset_1(np__7,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc7_numerals)).
fof(spc49_numerals, axiom, (((v2_xxreal_0(np__49)&m2_subset_1(np__49,k1_numbers,k5_numbers))&m1_subset_1(np__49,k5_numbers))&m1_subset_1(np__49,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc49_numerals)).
fof(spc48_numerals, axiom, (((v2_xxreal_0(np__48)&m2_subset_1(np__48,k1_numbers,k5_numbers))&m1_subset_1(np__48,k5_numbers))&m1_subset_1(np__48,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc48_numerals)).
fof(spc20_numerals, axiom, (((v2_xxreal_0(np__20)&m2_subset_1(np__20,k1_numbers,k5_numbers))&m1_subset_1(np__20,k5_numbers))&m1_subset_1(np__20,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc20_numerals)).
fof(spc19_numerals, axiom, (((v2_xxreal_0(np__19)&m2_subset_1(np__19,k1_numbers,k5_numbers))&m1_subset_1(np__19,k5_numbers))&m1_subset_1(np__19,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc19_numerals)).
fof(spc18_numerals, axiom, (((v2_xxreal_0(np__18)&m2_subset_1(np__18,k1_numbers,k5_numbers))&m1_subset_1(np__18,k5_numbers))&m1_subset_1(np__18,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc18_numerals)).
fof(spc6_numerals, axiom, (((v2_xxreal_0(np__6)&m2_subset_1(np__6,k1_numbers,k5_numbers))&m1_subset_1(np__6,k5_numbers))&m1_subset_1(np__6,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc6_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r24_r1_r25, axiom, k2_xcmplx_0(np__24,np__1)=np__25, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r24_r1_r25)).
fof(rqRealAdd__k2_xcmplx_0__r1_r5_r6, axiom, k2_xcmplx_0(np__1,np__5)=np__6, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r5_r6)).
fof(rqRealAdd__k2_xcmplx_0__r21_r1_r22, axiom, k2_xcmplx_0(np__21,np__1)=np__22, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r21_r1_r22)).
fof(rqRealAdd__k2_xcmplx_0__r22_r1_r23, axiom, k2_xcmplx_0(np__22,np__1)=np__23, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r22_r1_r23)).
fof(rqRealAdd__k2_xcmplx_0__r23_r1_r24, axiom, k2_xcmplx_0(np__23,np__1)=np__24, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r23_r1_r24)).
fof(rqRealAdd__k2_xcmplx_0__r1_r47_r48, axiom, k2_xcmplx_0(np__1,np__47)=np__48, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r47_r48)).
fof(spc50_numerals, axiom, (((v2_xxreal_0(np__50)&m2_subset_1(np__50,k1_numbers,k5_numbers))&m1_subset_1(np__50,k5_numbers))&m1_subset_1(np__50,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc50_numerals)).
fof(spc13_numerals, axiom, (((v2_xxreal_0(np__13)&m2_subset_1(np__13,k1_numbers,k5_numbers))&m1_subset_1(np__13,k5_numbers))&m1_subset_1(np__13,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc13_numerals)).
fof(spc14_numerals, axiom, (((v2_xxreal_0(np__14)&m2_subset_1(np__14,k1_numbers,k5_numbers))&m1_subset_1(np__14,k5_numbers))&m1_subset_1(np__14,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc14_numerals)).
fof(spc51_numerals, axiom, (((v2_xxreal_0(np__51)&m2_subset_1(np__51,k1_numbers,k5_numbers))&m1_subset_1(np__51,k5_numbers))&m1_subset_1(np__51,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc51_numerals)).
fof(spc52_numerals, axiom, (((v2_xxreal_0(np__52)&m2_subset_1(np__52,k1_numbers,k5_numbers))&m1_subset_1(np__52,k5_numbers))&m1_subset_1(np__52,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc52_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r1_r7_r8, axiom, k2_xcmplx_0(np__1,np__7)=np__8, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r7_r8)).
fof(rqRealAdd__k2_xcmplx_0__r1_r49_r50, axiom, k2_xcmplx_0(np__1,np__49)=np__50, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r49_r50)).
fof(rqRealAdd__k2_xcmplx_0__r1_r48_r49, axiom, k2_xcmplx_0(np__1,np__48)=np__49, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r48_r49)).
fof(rqRealAdd__k2_xcmplx_0__r20_r1_r21, axiom, k2_xcmplx_0(np__20,np__1)=np__21, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r20_r1_r21)).
fof(rqRealAdd__k2_xcmplx_0__r19_r1_r20, axiom, k2_xcmplx_0(np__19,np__1)=np__20, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r19_r1_r20)).
fof(rqRealAdd__k2_xcmplx_0__r18_r1_r19, axiom, k2_xcmplx_0(np__18,np__1)=np__19, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r18_r1_r19)).
fof(rqRealAdd__k2_xcmplx_0__r1_r6_r7, axiom, k2_xcmplx_0(np__1,np__6)=np__7, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r6_r7)).
fof(spc12_numerals, axiom, (((v2_xxreal_0(np__12)&m2_subset_1(np__12,k1_numbers,k5_numbers))&m1_subset_1(np__12,k5_numbers))&m1_subset_1(np__12,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc12_numerals)).
fof(spc55_numerals, axiom, (((v2_xxreal_0(np__55)&m2_subset_1(np__55,k1_numbers,k5_numbers))&m1_subset_1(np__55,k5_numbers))&m1_subset_1(np__55,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc55_numerals)).
fof(spc54_numerals, axiom, (((v2_xxreal_0(np__54)&m2_subset_1(np__54,k1_numbers,k5_numbers))&m1_subset_1(np__54,k5_numbers))&m1_subset_1(np__54,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc54_numerals)).
fof(spc53_numerals, axiom, (((v2_xxreal_0(np__53)&m2_subset_1(np__53,k1_numbers,k5_numbers))&m1_subset_1(np__53,k5_numbers))&m1_subset_1(np__53,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc53_numerals)).
fof(spc11_numerals, axiom, (((v2_xxreal_0(np__11)&m2_subset_1(np__11,k1_numbers,k5_numbers))&m1_subset_1(np__11,k5_numbers))&m1_subset_1(np__11,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc11_numerals)).
fof(spc10_numerals, axiom, (((v2_xxreal_0(np__10)&m2_subset_1(np__10,k1_numbers,k5_numbers))&m1_subset_1(np__10,k5_numbers))&m1_subset_1(np__10,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc10_numerals)).
fof(spc9_numerals, axiom, (((v2_xxreal_0(np__9)&m2_subset_1(np__9,k1_numbers,k5_numbers))&m1_subset_1(np__9,k5_numbers))&m1_subset_1(np__9,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc9_numerals)).
fof(spc8_numerals, axiom, (((v2_xxreal_0(np__8)&m2_subset_1(np__8,k1_numbers,k5_numbers))&m1_subset_1(np__8,k5_numbers))&m1_subset_1(np__8,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc8_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r1_r50_r51, axiom, k2_xcmplx_0(np__1,np__50)=np__51, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r50_r51)).
fof(rqRealAdd__k2_xcmplx_0__r13_r1_r14, axiom, k2_xcmplx_0(np__13,np__1)=np__14, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r13_r1_r14)).
fof(rqRealAdd__k2_xcmplx_0__r14_r1_r15, axiom, k2_xcmplx_0(np__14,np__1)=np__15, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r14_r1_r15)).
fof(rqRealAdd__k2_xcmplx_0__r1_r51_r52, axiom, k2_xcmplx_0(np__1,np__51)=np__52, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r51_r52)).
fof(rqRealAdd__k2_xcmplx_0__r1_r52_r53, axiom, k2_xcmplx_0(np__1,np__52)=np__53, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r52_r53)).
fof(spc57_numerals, axiom, (((v2_xxreal_0(np__57)&m2_subset_1(np__57,k1_numbers,k5_numbers))&m1_subset_1(np__57,k5_numbers))&m1_subset_1(np__57,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc57_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r12_r1_r13, axiom, k2_xcmplx_0(np__12,np__1)=np__13, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r12_r1_r13)).
fof(rqRealAdd__k2_xcmplx_0__r1_r55_r56, axiom, k2_xcmplx_0(np__1,np__55)=np__56, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r55_r56)).
fof(rqRealAdd__k2_xcmplx_0__r1_r54_r55, axiom, k2_xcmplx_0(np__1,np__54)=np__55, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r54_r55)).
fof(rqRealAdd__k2_xcmplx_0__r1_r53_r54, axiom, k2_xcmplx_0(np__1,np__53)=np__54, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r53_r54)).
fof(rqRealAdd__k2_xcmplx_0__r11_r1_r12, axiom, k2_xcmplx_0(np__11,np__1)=np__12, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r11_r1_r12)).
fof(rqRealAdd__k2_xcmplx_0__r10_r1_r11, axiom, k2_xcmplx_0(np__10,np__1)=np__11, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r10_r1_r11)).
fof(rqRealAdd__k2_xcmplx_0__r9_r1_r10, axiom, k2_xcmplx_0(np__9,np__1)=np__10, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r9_r1_r10)).
fof(rqRealAdd__k2_xcmplx_0__r8_r1_r9, axiom, k2_xcmplx_0(np__8,np__1)=np__9, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r8_r1_r9)).
fof(spc59_numerals, axiom, (((v2_xxreal_0(np__59)&m2_subset_1(np__59,k1_numbers,k5_numbers))&m1_subset_1(np__59,k5_numbers))&m1_subset_1(np__59,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc59_numerals)).
fof(spc58_numerals, axiom, (((v2_xxreal_0(np__58)&m2_subset_1(np__58,k1_numbers,k5_numbers))&m1_subset_1(np__58,k5_numbers))&m1_subset_1(np__58,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc58_numerals)).
fof(spc56_numerals, axiom, (((v2_xxreal_0(np__56)&m2_subset_1(np__56,k1_numbers,k5_numbers))&m1_subset_1(np__56,k5_numbers))&m1_subset_1(np__56,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc56_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r1_r57_r58, axiom, k2_xcmplx_0(np__1,np__57)=np__58, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r57_r58)).
fof(spc60_numerals, axiom, (((v2_xxreal_0(np__60)&m2_subset_1(np__60,k1_numbers,k5_numbers))&m1_subset_1(np__60,k5_numbers))&m1_subset_1(np__60,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc60_numerals)).
fof(spc42_numerals, axiom, (((v2_xxreal_0(np__42)&m2_subset_1(np__42,k1_numbers,k5_numbers))&m1_subset_1(np__42,k5_numbers))&m1_subset_1(np__42,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc42_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r1_r59_r60, axiom, k2_xcmplx_0(np__1,np__59)=np__60, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r59_r60)).
fof(rqRealAdd__k2_xcmplx_0__r58_r1_r59, axiom, k2_xcmplx_0(np__58,np__1)=np__59, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r58_r1_r59)).
fof(rqRealAdd__k2_xcmplx_0__r1_r56_r57, axiom, k2_xcmplx_0(np__1,np__56)=np__57, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r56_r57)).
fof(rqRealAdd__k2_xcmplx_0__r1_r1_r2, axiom, k2_xcmplx_0(np__1,np__1)=np__2, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r1_r2)).
fof(spc41_numerals, axiom, (((v2_xxreal_0(np__41)&m2_subset_1(np__41,k1_numbers,k5_numbers))&m1_subset_1(np__41,k5_numbers))&m1_subset_1(np__41,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc41_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r1_r60_r61, axiom, k2_xcmplx_0(np__1,np__60)=np__61, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r60_r61)).
fof(rqRealAdd__k2_xcmplx_0__r1_r42_r43, axiom, k2_xcmplx_0(np__1,np__42)=np__43, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r42_r43)).
fof(spc33_numerals, axiom, (((v2_xxreal_0(np__33)&m2_subset_1(np__33,k1_numbers,k5_numbers))&m1_subset_1(np__33,k5_numbers))&m1_subset_1(np__33,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc33_numerals)).
fof(spc15_numerals, axiom, (((v2_xxreal_0(np__15)&m2_subset_1(np__15,k1_numbers,k5_numbers))&m1_subset_1(np__15,k5_numbers))&m1_subset_1(np__15,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc15_numerals)).
fof(spc17_numerals, axiom, (((v2_xxreal_0(np__17)&m2_subset_1(np__17,k1_numbers,k5_numbers))&m1_subset_1(np__17,k5_numbers))&m1_subset_1(np__17,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc17_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r1_r41_r42, axiom, k2_xcmplx_0(np__1,np__41)=np__42, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r41_r42)).
fof(spc31_numerals, axiom, (((v2_xxreal_0(np__31)&m2_subset_1(np__31,k1_numbers,k5_numbers))&m1_subset_1(np__31,k5_numbers))&m1_subset_1(np__31,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc31_numerals)).
fof(fc1_nat_1, axiom, ![X1, X2]:((v7_ordinal1(X1)&v7_ordinal1(X2))=>v7_ordinal1(k2_xcmplx_0(X1,X2))), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', fc1_nat_1)).
fof(rqRealAdd__k2_xcmplx_0__r1_r33_r34, axiom, k2_xcmplx_0(np__1,np__33)=np__34, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r33_r34)).
fof(rqRealAdd__k2_xcmplx_0__r1_r15_r16, axiom, k2_xcmplx_0(np__1,np__15)=np__16, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r15_r16)).
fof(rqRealAdd__k2_xcmplx_0__r17_r1_r18, axiom, k2_xcmplx_0(np__17,np__1)=np__18, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r17_r1_r18)).
fof(rqRealAdd__k2_xcmplx_0__r31_r1_r32, axiom, k2_xcmplx_0(np__31,np__1)=np__32, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r31_r1_r32)).
fof(t23_descip_1, conjecture, ![X1]:~(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((r2_hidden(X1,k2_finseq_1(np__64))&X1!=np__1)&X1!=np__2)&X1!=np__3)&X1!=np__4)&X1!=np__5)&X1!=np__6)&X1!=np__7)&X1!=np__8)&X1!=np__9)&X1!=np__10)&X1!=np__11)&X1!=np__12)&X1!=np__13)&X1!=np__14)&X1!=np__15)&X1!=np__16)&X1!=np__17)&X1!=np__18)&X1!=np__19)&X1!=np__20)&X1!=np__21)&X1!=np__22)&X1!=np__23)&X1!=np__24)&X1!=np__25)&X1!=np__26)&X1!=np__27)&X1!=np__28)&X1!=np__29)&X1!=np__30)&X1!=np__31)&X1!=np__32)&X1!=np__33)&X1!=np__34)&X1!=np__35)&X1!=np__36)&X1!=np__37)&X1!=np__38)&X1!=np__39)&X1!=np__40)&X1!=np__41)&X1!=np__42)&X1!=np__43)&X1!=np__44)&X1!=np__45)&X1!=np__46)&X1!=np__47)&X1!=np__48)&X1!=np__49)&X1!=np__50)&X1!=np__51)&X1!=np__52)&X1!=np__53)&X1!=np__54)&X1!=np__55)&X1!=np__56)&X1!=np__57)&X1!=np__58)&X1!=np__59)&X1!=np__60)&X1!=np__61)&X1!=np__62)&X1!=np__63)&X1!=np__64)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', t23_descip_1)).
fof(t4_subset, axiom, ![X1, X2, X3]:((r2_hidden(X1,X2)&m1_subset_1(X2,k1_zfmisc_1(X3)))=>m1_subset_1(X1,X3)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', t4_subset)).
fof(dt_k2_finseq_1, axiom, ![X1]:(v7_ordinal1(X1)=>m1_subset_1(k2_finseq_1(X1),k1_zfmisc_1(k5_numbers))), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', dt_k2_finseq_1)).
fof(rqRealAdd__k2_xcmplx_0__r1_r31_r32, axiom, k2_xcmplx_0(np__1,np__31)=np__32, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r31_r32)).
fof(t1_finseq_1, axiom, ![X1]:(v7_ordinal1(X1)=>![X2]:(v7_ordinal1(X2)=>(r2_hidden(X1,k2_finseq_1(X2))<=>(r1_xxreal_0(np__1,X1)&r1_xxreal_0(X1,X2))))), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', t1_finseq_1)).
fof(spc64_numerals, axiom, (((v2_xxreal_0(np__64)&m2_subset_1(np__64,k1_numbers,k5_numbers))&m1_subset_1(np__64,k5_numbers))&m1_subset_1(np__64,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc64_numerals)).
fof(rqRealAdd__k2_xcmplx_0__r32_r1_r33, axiom, k2_xcmplx_0(np__32,np__1)=np__33, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r32_r1_r33)).
fof(spc32_numerals, axiom, (((v2_xxreal_0(np__32)&m2_subset_1(np__32,k1_numbers,k5_numbers))&m1_subset_1(np__32,k5_numbers))&m1_subset_1(np__32,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc32_numerals)).
fof(l30_descip_1, axiom, ![X1]:(v7_ordinal1(X1)=>![X2]:(v7_ordinal1(X2)=>~((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((r1_xxreal_0(X1,X2)&r1_xxreal_0(X2,k1_nat_1(X1,np__63)))&X2!=X1)&X2!=k1_nat_1(X1,np__1))&X2!=k1_nat_1(X1,np__2))&X2!=k1_nat_1(X1,np__3))&X2!=k1_nat_1(X1,np__4))&X2!=k1_nat_1(X1,np__5))&X2!=k1_nat_1(X1,np__6))&X2!=k1_nat_1(X1,np__7))&X2!=k1_nat_1(X1,np__8))&X2!=k1_nat_1(X1,np__9))&X2!=k1_nat_1(X1,np__10))&X2!=k1_nat_1(X1,np__11))&X2!=k1_nat_1(X1,np__12))&X2!=k1_nat_1(X1,np__13))&X2!=k1_nat_1(X1,np__14))&X2!=k1_nat_1(X1,np__15))&X2!=k1_nat_1(X1,np__16))&X2!=k1_nat_1(X1,np__17))&X2!=k1_nat_1(X1,np__18))&X2!=k1_nat_1(X1,np__19))&X2!=k1_nat_1(X1,np__20))&X2!=k1_nat_1(X1,np__21))&X2!=k1_nat_1(X1,np__22))&X2!=k1_nat_1(X1,np__23))&X2!=k1_nat_1(X1,np__24))&X2!=k1_nat_1(X1,np__25))&X2!=k1_nat_1(X1,np__26))&X2!=k1_nat_1(X1,np__27))&X2!=k1_nat_1(X1,np__28))&X2!=k1_nat_1(X1,np__29))&X2!=k1_nat_1(X1,np__30))&X2!=k1_nat_1(X1,np__31))&X2!=k1_nat_1(X1,np__32))&X2!=k1_nat_1(X1,np__33))&X2!=k1_nat_1(X1,np__34))&X2!=k1_nat_1(X1,np__35))&X2!=k1_nat_1(X1,np__36))&X2!=k1_nat_1(X1,np__37))&X2!=k1_nat_1(X1,np__38))&X2!=k1_nat_1(X1,np__39))&X2!=k1_nat_1(X1,np__40))&X2!=k1_nat_1(X1,np__41))&X2!=k1_nat_1(X1,np__42))&X2!=k1_nat_1(X1,np__43))&X2!=k1_nat_1(X1,np__44))&X2!=k1_nat_1(X1,np__45))&X2!=k1_nat_1(X1,np__46))&X2!=k1_nat_1(X1,np__47))&X2!=k1_nat_1(X1,np__48))&X2!=k1_nat_1(X1,np__49))&X2!=k1_nat_1(X1,np__50))&X2!=k1_nat_1(X1,np__51))&X2!=k1_nat_1(X1,np__52))&X2!=k1_nat_1(X1,np__53))&X2!=k1_nat_1(X1,np__54))&X2!=k1_nat_1(X1,np__55))&X2!=k1_nat_1(X1,np__56))&X2!=k1_nat_1(X1,np__57))&X2!=k1_nat_1(X1,np__58))&X2!=k1_nat_1(X1,np__59))&X2!=k1_nat_1(X1,np__60))&X2!=k1_nat_1(X1,np__61))&X2!=k1_nat_1(X1,np__62))&X2!=k1_nat_1(X1,np__63))))), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', l30_descip_1)).
fof(rqRealAdd__k2_xcmplx_0__r1_r39_r40, axiom, k2_xcmplx_0(np__1,np__39)=np__40, file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', rqRealAdd__k2_xcmplx_0__r1_r39_r40)).
fof(spc39_numerals, axiom, (((v2_xxreal_0(np__39)&m2_subset_1(np__39,k1_numbers,k5_numbers))&m1_subset_1(np__39,k5_numbers))&m1_subset_1(np__39,k1_numbers)), file('/scratch/mptp/hardpred/preds__0.005/t23_descip_1', spc39_numerals)).
fof(c_0_142, plain, ![X7]:(~m1_subset_1(X7,k4_ordinal1)|v7_ordinal1(X7)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc8_ordinal1])])).
fof(c_0_143, plain, ![X10, X11]:(~v6_membered(X10)|(~m1_subset_1(X11,X10)|v7_ordinal1(X11))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[cc17_membered])])])).
cnf(c_0_144, plain, (v6_membered(k4_ordinal1)), inference(split_conjunct,[status(thm)],[fc6_membered])).
cnf(c_0_145, plain, (k5_numbers=k4_ordinal1), inference(split_conjunct,[status(thm)],[redefinition_k5_numbers])).
fof(c_0_146, plain, ![X2, X1]:(epred1_2(X1,X2)<=>~((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((r1_xxreal_0(X1,X2)&r1_xxreal_0(X2,k1_nat_1(X1,np__63)))&X2!=X1)&X2!=k1_nat_1(X1,np__1))&X2!=k1_nat_1(X1,np__2))&X2!=k1_nat_1(X1,np__3))&X2!=k1_nat_1(X1,np__4))&X2!=k1_nat_1(X1,np__5))&X2!=k1_nat_1(X1,np__6))&X2!=k1_nat_1(X1,np__7))&X2!=k1_nat_1(X1,np__8))&X2!=k1_nat_1(X1,np__9))&X2!=k1_nat_1(X1,np__10))&X2!=k1_nat_1(X1,np__11))&X2!=k1_nat_1(X1,np__12))&X2!=k1_nat_1(X1,np__13))&X2!=k1_nat_1(X1,np__14))&X2!=k1_nat_1(X1,np__15))&X2!=k1_nat_1(X1,np__16))&X2!=k1_nat_1(X1,np__17))&X2!=k1_nat_1(X1,np__18))&X2!=k1_nat_1(X1,np__19))&X2!=k1_nat_1(X1,np__20))&X2!=k1_nat_1(X1,np__21))&X2!=k1_nat_1(X1,np__22))&X2!=k1_nat_1(X1,np__23))&X2!=k1_nat_1(X1,np__24))&X2!=k1_nat_1(X1,np__25))&X2!=k1_nat_1(X1,np__26))&X2!=k1_nat_1(X1,np__27))&X2!=k1_nat_1(X1,np__28))&X2!=k1_nat_1(X1,np__29))&X2!=k1_nat_1(X1,np__30))&X2!=k1_nat_1(X1,np__31))&X2!=k1_nat_1(X1,np__32))&X2!=k1_nat_1(X1,np__33))&X2!=k1_nat_1(X1,np__34))&X2!=k1_nat_1(X1,np__35))&X2!=k1_nat_1(X1,np__36))&X2!=k1_nat_1(X1,np__37))&X2!=k1_nat_1(X1,np__38))&X2!=k1_nat_1(X1,np__39))&X2!=k1_nat_1(X1,np__40))&X2!=k1_nat_1(X1,np__41))&X2!=k1_nat_1(X1,np__42))&X2!=k1_nat_1(X1,np__43))&X2!=k1_nat_1(X1,np__44))&X2!=k1_nat_1(X1,np__45))&X2!=k1_nat_1(X1,np__46))&X2!=k1_nat_1(X1,np__47))&X2!=k1_nat_1(X1,np__48))&X2!=k1_nat_1(X1,np__49))&X2!=k1_nat_1(X1,np__50))&X2!=k1_nat_1(X1,np__51))&X2!=k1_nat_1(X1,np__52))&X2!=k1_nat_1(X1,np__53))&X2!=k1_nat_1(X1,np__54))&X2!=k1_nat_1(X1,np__55))&X2!=k1_nat_1(X1,np__56))&X2!=k1_nat_1(X1,np__57))&X2!=k1_nat_1(X1,np__58))&X2!=k1_nat_1(X1,np__59))&X2!=k1_nat_1(X1,np__60))&X2!=k1_nat_1(X1,np__61))&X2!=k1_nat_1(X1,np__62))&X2!=k1_nat_1(X1,np__63)))), introduced(definition)).
fof(c_0_147, plain, ![X33, X34]:(~v7_ordinal1(X33)|~m1_subset_1(X34,k5_numbers)|k1_nat_1(X33,X34)=k1_nat_1(X34,X33)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[commutativity_k1_nat_1])])).
cnf(c_0_148, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k4_ordinal1)), inference(split_conjunct,[status(thm)],[c_0_142])).
fof(c_0_149, plain, ![X24, X25]:(~v7_ordinal1(X24)|~m1_subset_1(X25,k5_numbers)|k1_nat_1(X24,X25)=k2_xcmplx_0(X24,X25)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[redefinition_k1_nat_1])])).
cnf(c_0_150, plain, (v7_ordinal1(X2)|~v6_membered(X1)|~m1_subset_1(X2,X1)), inference(split_conjunct,[status(thm)],[c_0_143])).
cnf(c_0_151, plain, (m1_subset_1(np__1,k5_numbers)), inference(split_conjunct,[status(thm)],[spc1_numerals])).
cnf(c_0_152, plain, (v6_membered(k5_numbers)), inference(rw,[status(thm)],[c_0_144, c_0_145])).
fof(c_0_153, plain, ![X39, X40]:(~v7_ordinal1(X39)|~v7_ordinal1(X40)|v7_ordinal1(k3_xcmplx_0(X39,X40))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc2_nat_1])])).
fof(c_0_154, plain, ![X2, X1]:(epred1_2(X1,X2)=>~((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((r1_xxreal_0(X1,X2)&r1_xxreal_0(X2,k1_nat_1(X1,np__63)))&X2!=X1)&X2!=k1_nat_1(X1,np__1))&X2!=k1_nat_1(X1,np__2))&X2!=k1_nat_1(X1,np__3))&X2!=k1_nat_1(X1,np__4))&X2!=k1_nat_1(X1,np__5))&X2!=k1_nat_1(X1,np__6))&X2!=k1_nat_1(X1,np__7))&X2!=k1_nat_1(X1,np__8))&X2!=k1_nat_1(X1,np__9))&X2!=k1_nat_1(X1,np__10))&X2!=k1_nat_1(X1,np__11))&X2!=k1_nat_1(X1,np__12))&X2!=k1_nat_1(X1,np__13))&X2!=k1_nat_1(X1,np__14))&X2!=k1_nat_1(X1,np__15))&X2!=k1_nat_1(X1,np__16))&X2!=k1_nat_1(X1,np__17))&X2!=k1_nat_1(X1,np__18))&X2!=k1_nat_1(X1,np__19))&X2!=k1_nat_1(X1,np__20))&X2!=k1_nat_1(X1,np__21))&X2!=k1_nat_1(X1,np__22))&X2!=k1_nat_1(X1,np__23))&X2!=k1_nat_1(X1,np__24))&X2!=k1_nat_1(X1,np__25))&X2!=k1_nat_1(X1,np__26))&X2!=k1_nat_1(X1,np__27))&X2!=k1_nat_1(X1,np__28))&X2!=k1_nat_1(X1,np__29))&X2!=k1_nat_1(X1,np__30))&X2!=k1_nat_1(X1,np__31))&X2!=k1_nat_1(X1,np__32))&X2!=k1_nat_1(X1,np__33))&X2!=k1_nat_1(X1,np__34))&X2!=k1_nat_1(X1,np__35))&X2!=k1_nat_1(X1,np__36))&X2!=k1_nat_1(X1,np__37))&X2!=k1_nat_1(X1,np__38))&X2!=k1_nat_1(X1,np__39))&X2!=k1_nat_1(X1,np__40))&X2!=k1_nat_1(X1,np__41))&X2!=k1_nat_1(X1,np__42))&X2!=k1_nat_1(X1,np__43))&X2!=k1_nat_1(X1,np__44))&X2!=k1_nat_1(X1,np__45))&X2!=k1_nat_1(X1,np__46))&X2!=k1_nat_1(X1,np__47))&X2!=k1_nat_1(X1,np__48))&X2!=k1_nat_1(X1,np__49))&X2!=k1_nat_1(X1,np__50))&X2!=k1_nat_1(X1,np__51))&X2!=k1_nat_1(X1,np__52))&X2!=k1_nat_1(X1,np__53))&X2!=k1_nat_1(X1,np__54))&X2!=k1_nat_1(X1,np__55))&X2!=k1_nat_1(X1,np__56))&X2!=k1_nat_1(X1,np__57))&X2!=k1_nat_1(X1,np__58))&X2!=k1_nat_1(X1,np__59))&X2!=k1_nat_1(X1,np__60))&X2!=k1_nat_1(X1,np__61))&X2!=k1_nat_1(X1,np__62))&X2!=k1_nat_1(X1,np__63)))), inference(split_equiv,[status(thm)],[c_0_146])).
cnf(c_0_155, plain, (m1_subset_1(np__2,k5_numbers)), inference(split_conjunct,[status(thm)],[spc2_numerals])).
cnf(c_0_156, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_147])).
cnf(c_0_157, plain, (v7_ordinal1(X1)|~m1_subset_1(X1,k5_numbers)), inference(rw,[status(thm)],[c_0_148, c_0_145])).
cnf(c_0_158, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~v7_ordinal1(X1)|~m1_subset_1(X2,k5_numbers)), inference(split_conjunct,[status(thm)],[c_0_149])).
cnf(c_0_159, plain, (v7_ordinal1(np__1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_151]), c_0_152])])).
cnf(c_0_160, plain, (m1_subset_1(np__4,k5_numbers)), inference(split_conjunct,[status(thm)],[spc4_numerals])).
cnf(c_0_161, plain, (v7_ordinal1(k3_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_153])).
fof(c_0_162, plain, ![X48, X49]:(~epred1_2(X49,X48)|(~r1_xxreal_0(X49,X48)|~r1_xxreal_0(X48,k1_nat_1(X49,np__63))|X48=X49|X48=k1_nat_1(X49,np__1)|X48=k1_nat_1(X49,np__2)|X48=k1_nat_1(X49,np__3)|X48=k1_nat_1(X49,np__4)|X48=k1_nat_1(X49,np__5)|X48=k1_nat_1(X49,np__6)|X48=k1_nat_1(X49,np__7)|X48=k1_nat_1(X49,np__8)|X48=k1_nat_1(X49,np__9)|X48=k1_nat_1(X49,np__10)|X48=k1_nat_1(X49,np__11)|X48=k1_nat_1(X49,np__12)|X48=k1_nat_1(X49,np__13)|X48=k1_nat_1(X49,np__14)|X48=k1_nat_1(X49,np__15)|X48=k1_nat_1(X49,np__16)|X48=k1_nat_1(X49,np__17)|X48=k1_nat_1(X49,np__18)|X48=k1_nat_1(X49,np__19)|X48=k1_nat_1(X49,np__20)|X48=k1_nat_1(X49,np__21)|X48=k1_nat_1(X49,np__22)|X48=k1_nat_1(X49,np__23)|X48=k1_nat_1(X49,np__24)|X48=k1_nat_1(X49,np__25)|X48=k1_nat_1(X49,np__26)|X48=k1_nat_1(X49,np__27)|X48=k1_nat_1(X49,np__28)|X48=k1_nat_1(X49,np__29)|X48=k1_nat_1(X49,np__30)|X48=k1_nat_1(X49,np__31)|X48=k1_nat_1(X49,np__32)|X48=k1_nat_1(X49,np__33)|X48=k1_nat_1(X49,np__34)|X48=k1_nat_1(X49,np__35)|X48=k1_nat_1(X49,np__36)|X48=k1_nat_1(X49,np__37)|X48=k1_nat_1(X49,np__38)|X48=k1_nat_1(X49,np__39)|X48=k1_nat_1(X49,np__40)|X48=k1_nat_1(X49,np__41)|X48=k1_nat_1(X49,np__42)|X48=k1_nat_1(X49,np__43)|X48=k1_nat_1(X49,np__44)|X48=k1_nat_1(X49,np__45)|X48=k1_nat_1(X49,np__46)|X48=k1_nat_1(X49,np__47)|X48=k1_nat_1(X49,np__48)|X48=k1_nat_1(X49,np__49)|X48=k1_nat_1(X49,np__50)|X48=k1_nat_1(X49,np__51)|X48=k1_nat_1(X49,np__52)|X48=k1_nat_1(X49,np__53)|X48=k1_nat_1(X49,np__54)|X48=k1_nat_1(X49,np__55)|X48=k1_nat_1(X49,np__56)|X48=k1_nat_1(X49,np__57)|X48=k1_nat_1(X49,np__58)|X48=k1_nat_1(X49,np__59)|X48=k1_nat_1(X49,np__60)|X48=k1_nat_1(X49,np__61)|X48=k1_nat_1(X49,np__62)|X48=k1_nat_1(X49,np__63))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_154])])).
cnf(c_0_163, plain, (v7_ordinal1(np__2)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_155]), c_0_152])])).
cnf(c_0_164, plain, (k1_nat_1(X1,X2)=k1_nat_1(X2,X1)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_157])).
cnf(c_0_165, plain, (k1_nat_1(np__1,X1)=k2_xcmplx_0(np__1,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_159])).
cnf(c_0_166, plain, (v7_ordinal1(np__4)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_160]), c_0_152])])).
cnf(c_0_167, plain, (k1_nat_1(k3_xcmplx_0(X1,X2),X3)=k1_nat_1(X3,k3_xcmplx_0(X1,X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)|~m1_subset_1(X3,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_161])).
cnf(c_0_168, plain, (k3_xcmplx_0(np__4,np__4)=np__16), inference(split_conjunct,[status(thm)],[rqRealMult__k3_xcmplx_0__r4_r4_r16])).
cnf(c_0_169, plain, (m1_subset_1(np__38,k5_numbers)), inference(split_conjunct,[status(thm)],[spc38_numerals])).
cnf(c_0_170, plain, (m1_subset_1(np__37,k5_numbers)), inference(split_conjunct,[status(thm)],[spc37_numerals])).
cnf(c_0_171, plain, (m1_subset_1(np__36,k5_numbers)), inference(split_conjunct,[status(thm)],[spc36_numerals])).
cnf(c_0_172, plain, (m1_subset_1(np__3,k5_numbers)), inference(split_conjunct,[status(thm)],[spc3_numerals])).
cnf(c_0_173, plain, (X2=X1|X2=k1_nat_1(X1,np__1)|X2=k1_nat_1(X1,np__2)|X2=k1_nat_1(X1,np__3)|X2=k1_nat_1(X1,np__4)|X2=k1_nat_1(X1,np__5)|X2=k1_nat_1(X1,np__6)|X2=k1_nat_1(X1,np__7)|X2=k1_nat_1(X1,np__8)|X2=k1_nat_1(X1,np__9)|X2=k1_nat_1(X1,np__10)|X2=k1_nat_1(X1,np__11)|X2=k1_nat_1(X1,np__12)|X2=k1_nat_1(X1,np__13)|X2=k1_nat_1(X1,np__14)|X2=k1_nat_1(X1,np__15)|X2=k1_nat_1(X1,np__16)|X2=k1_nat_1(X1,np__17)|X2=k1_nat_1(X1,np__18)|X2=k1_nat_1(X1,np__19)|X2=k1_nat_1(X1,np__20)|X2=k1_nat_1(X1,np__21)|X2=k1_nat_1(X1,np__22)|X2=k1_nat_1(X1,np__23)|X2=k1_nat_1(X1,np__24)|X2=k1_nat_1(X1,np__25)|X2=k1_nat_1(X1,np__26)|X2=k1_nat_1(X1,np__27)|X2=k1_nat_1(X1,np__28)|X2=k1_nat_1(X1,np__29)|X2=k1_nat_1(X1,np__30)|X2=k1_nat_1(X1,np__31)|X2=k1_nat_1(X1,np__32)|X2=k1_nat_1(X1,np__33)|X2=k1_nat_1(X1,np__34)|X2=k1_nat_1(X1,np__35)|X2=k1_nat_1(X1,np__36)|X2=k1_nat_1(X1,np__37)|X2=k1_nat_1(X1,np__38)|X2=k1_nat_1(X1,np__39)|X2=k1_nat_1(X1,np__40)|X2=k1_nat_1(X1,np__41)|X2=k1_nat_1(X1,np__42)|X2=k1_nat_1(X1,np__43)|X2=k1_nat_1(X1,np__44)|X2=k1_nat_1(X1,np__45)|X2=k1_nat_1(X1,np__46)|X2=k1_nat_1(X1,np__47)|X2=k1_nat_1(X1,np__48)|X2=k1_nat_1(X1,np__49)|X2=k1_nat_1(X1,np__50)|X2=k1_nat_1(X1,np__51)|X2=k1_nat_1(X1,np__52)|X2=k1_nat_1(X1,np__53)|X2=k1_nat_1(X1,np__54)|X2=k1_nat_1(X1,np__55)|X2=k1_nat_1(X1,np__56)|X2=k1_nat_1(X1,np__57)|X2=k1_nat_1(X1,np__58)|X2=k1_nat_1(X1,np__59)|X2=k1_nat_1(X1,np__60)|X2=k1_nat_1(X1,np__61)|X2=k1_nat_1(X1,np__62)|X2=k1_nat_1(X1,np__63)|~epred1_2(X1,X2)|~r1_xxreal_0(X1,X2)|~r1_xxreal_0(X2,k1_nat_1(X1,np__63))), inference(split_conjunct,[status(thm)],[c_0_162])).
cnf(c_0_174, plain, (k1_nat_1(X1,X2)=k2_xcmplx_0(X1,X2)|~m1_subset_1(X2,k5_numbers)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_157])).
cnf(c_0_175, plain, (m1_subset_1(np__63,k5_numbers)), inference(split_conjunct,[status(thm)],[spc63_numerals])).
cnf(c_0_176, plain, (k1_nat_1(np__2,X1)=k1_nat_1(X1,np__2)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_163])).
cnf(c_0_177, plain, (k1_nat_1(X1,np__1)=k2_xcmplx_0(np__1,X1)|~m1_subset_1(X1,k5_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_165]), c_0_151])])).
cnf(c_0_178, plain, (k2_xcmplx_0(np__1,np__2)=np__3), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r2_r3])).
cnf(c_0_179, plain, (k1_nat_1(np__4,X1)=k1_nat_1(X1,np__4)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_166])).
cnf(c_0_180, plain, (k2_xcmplx_0(np__1,np__4)=np__5), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r4_r5])).
cnf(c_0_181, plain, (k1_nat_1(np__16,X1)=k1_nat_1(X1,np__16)|~m1_subset_1(X1,k5_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_168]), c_0_166])])).
cnf(c_0_182, plain, (k2_xcmplx_0(np__1,np__16)=np__17), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r16_r17])).
cnf(c_0_183, plain, (m1_subset_1(np__16,k5_numbers)), inference(split_conjunct,[status(thm)],[spc16_numerals])).
cnf(c_0_184, plain, (v7_ordinal1(np__38)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_169]), c_0_152])])).
cnf(c_0_185, plain, (v7_ordinal1(np__37)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_170]), c_0_152])])).
cnf(c_0_186, plain, (v7_ordinal1(np__36)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_171]), c_0_152])])).
cnf(c_0_187, plain, (v7_ordinal1(np__3)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_172]), c_0_152])])).
cnf(c_0_188, plain, (m1_subset_1(np__40,k5_numbers)), inference(split_conjunct,[status(thm)],[spc40_numerals])).
cnf(c_0_189, plain, (m1_subset_1(np__29,k5_numbers)), inference(split_conjunct,[status(thm)],[spc29_numerals])).
cnf(c_0_190, plain, (m1_subset_1(np__30,k5_numbers)), inference(split_conjunct,[status(thm)],[spc30_numerals])).
cnf(c_0_191, plain, (m1_subset_1(np__34,k5_numbers)), inference(split_conjunct,[status(thm)],[spc34_numerals])).
cnf(c_0_192, plain, (m1_subset_1(np__35,k5_numbers)), inference(split_conjunct,[status(thm)],[spc35_numerals])).
cnf(c_0_193, plain, (m1_subset_1(np__61,k5_numbers)), inference(split_conjunct,[status(thm)],[spc61_numerals])).
cnf(c_0_194, plain, (m1_subset_1(np__62,k5_numbers)), inference(split_conjunct,[status(thm)],[spc62_numerals])).
cnf(c_0_195, plain, (X1=k1_nat_1(X2,np__1)|X1=k1_nat_1(X2,np__2)|X1=k1_nat_1(X2,np__3)|X1=k1_nat_1(X2,np__4)|X1=k1_nat_1(X2,np__5)|X1=k1_nat_1(X2,np__6)|X1=k1_nat_1(X2,np__7)|X1=k1_nat_1(X2,np__8)|X1=k1_nat_1(X2,np__9)|X1=k1_nat_1(X2,np__10)|X1=k1_nat_1(X2,np__11)|X1=k1_nat_1(X2,np__12)|X1=k1_nat_1(X2,np__13)|X1=k1_nat_1(X2,np__14)|X1=k1_nat_1(X2,np__15)|X1=k1_nat_1(X2,np__16)|X1=k1_nat_1(X2,np__17)|X1=k1_nat_1(X2,np__18)|X1=k1_nat_1(X2,np__19)|X1=k1_nat_1(X2,np__20)|X1=k1_nat_1(X2,np__21)|X1=k1_nat_1(X2,np__22)|X1=k1_nat_1(X2,np__23)|X1=k1_nat_1(X2,np__24)|X1=k1_nat_1(X2,np__25)|X1=k1_nat_1(X2,np__26)|X1=k1_nat_1(X2,np__27)|X1=k1_nat_1(X2,np__28)|X1=k1_nat_1(X2,np__29)|X1=k1_nat_1(X2,np__30)|X1=k1_nat_1(X2,np__31)|X1=k1_nat_1(X2,np__32)|X1=k1_nat_1(X2,np__33)|X1=k1_nat_1(X2,np__34)|X1=k1_nat_1(X2,np__35)|X1=k1_nat_1(X2,np__36)|X1=k1_nat_1(X2,np__37)|X1=k1_nat_1(X2,np__38)|X1=k1_nat_1(X2,np__39)|X1=k1_nat_1(X2,np__40)|X1=k1_nat_1(X2,np__41)|X1=k1_nat_1(X2,np__42)|X1=k1_nat_1(X2,np__43)|X1=k1_nat_1(X2,np__44)|X1=k1_nat_1(X2,np__45)|X1=k1_nat_1(X2,np__46)|X1=k1_nat_1(X2,np__47)|X1=k1_nat_1(X2,np__48)|X1=k1_nat_1(X2,np__49)|X1=k1_nat_1(X2,np__50)|X1=k1_nat_1(X2,np__51)|X1=k1_nat_1(X2,np__52)|X1=k1_nat_1(X2,np__53)|X1=k1_nat_1(X2,np__54)|X1=k1_nat_1(X2,np__55)|X1=k1_nat_1(X2,np__56)|X1=k1_nat_1(X2,np__57)|X1=k1_nat_1(X2,np__58)|X1=k1_nat_1(X2,np__59)|X1=k1_nat_1(X2,np__60)|X1=k1_nat_1(X2,np__61)|X1=k1_nat_1(X2,np__62)|X1=k2_xcmplx_0(X2,np__63)|X2=X1|~epred1_2(X2,X1)|~r1_xxreal_0(X1,k2_xcmplx_0(X2,np__63))|~r1_xxreal_0(X2,X1)|~m1_subset_1(X2,k5_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_174]), c_0_175])])).
cnf(c_0_196, plain, (k2_xcmplx_0(np__1,np__63)=np__64), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r63_r64])).
cnf(c_0_197, plain, (k1_nat_1(np__1,np__2)=np__3), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_177]), c_0_178]), c_0_151]), c_0_155])])).
cnf(c_0_198, plain, (k1_nat_1(np__1,np__4)=np__5), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_177]), c_0_180]), c_0_151]), c_0_160])])).
cnf(c_0_199, plain, (k1_nat_1(np__1,np__16)=np__17), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_177]), c_0_182]), c_0_151]), c_0_183])])).
cnf(c_0_200, plain, (k1_nat_1(np__38,X1)=k2_xcmplx_0(np__38,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_184])).
cnf(c_0_201, plain, (k1_nat_1(np__1,X1)=k1_nat_1(X1,np__1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_159])).
cnf(c_0_202, plain, (k2_xcmplx_0(np__38,np__1)=np__39), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r38_r1_r39])).
cnf(c_0_203, plain, (k1_nat_1(np__37,X1)=k2_xcmplx_0(np__37,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_185])).
cnf(c_0_204, plain, (k2_xcmplx_0(np__37,np__1)=np__38), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r37_r1_r38])).
cnf(c_0_205, plain, (k1_nat_1(np__36,X1)=k2_xcmplx_0(np__36,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_186])).
cnf(c_0_206, plain, (k2_xcmplx_0(np__36,np__1)=np__37), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r36_r1_r37])).
cnf(c_0_207, plain, (k1_nat_1(np__3,X1)=k1_nat_1(X1,np__3)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_187])).
cnf(c_0_208, plain, (k2_xcmplx_0(np__1,np__3)=np__4), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r3_r4])).
cnf(c_0_209, plain, (v7_ordinal1(np__40)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_188]), c_0_152])])).
cnf(c_0_210, plain, (v7_ordinal1(np__29)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_189]), c_0_152])])).
cnf(c_0_211, plain, (v7_ordinal1(np__30)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_190]), c_0_152])])).
cnf(c_0_212, plain, (v7_ordinal1(np__34)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_191]), c_0_152])])).
cnf(c_0_213, plain, (v7_ordinal1(np__35)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_192]), c_0_152])])).
cnf(c_0_214, plain, (v7_ordinal1(np__61)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_193]), c_0_152])])).
cnf(c_0_215, plain, (v7_ordinal1(np__62)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_194]), c_0_152])])).
cnf(c_0_216, plain, (m1_subset_1(np__45,k5_numbers)), inference(split_conjunct,[status(thm)],[spc45_numerals])).
cnf(c_0_217, plain, (m1_subset_1(np__46,k5_numbers)), inference(split_conjunct,[status(thm)],[spc46_numerals])).
cnf(c_0_218, plain, (m1_subset_1(np__44,k5_numbers)), inference(split_conjunct,[status(thm)],[spc44_numerals])).
cnf(c_0_219, plain, (m1_subset_1(np__43,k5_numbers)), inference(split_conjunct,[status(thm)],[spc43_numerals])).
cnf(c_0_220, plain, (m1_subset_1(np__28,k5_numbers)), inference(split_conjunct,[status(thm)],[spc28_numerals])).
cnf(c_0_221, plain, (m1_subset_1(np__27,k5_numbers)), inference(split_conjunct,[status(thm)],[spc27_numerals])).
cnf(c_0_222, plain, (m1_subset_1(np__26,k5_numbers)), inference(split_conjunct,[status(thm)],[spc26_numerals])).
cnf(c_0_223, plain, (m1_subset_1(np__25,k5_numbers)), inference(split_conjunct,[status(thm)],[spc25_numerals])).
cnf(c_0_224, plain, (X1=k1_nat_1(np__1,np__62)|X1=k1_nat_1(np__1,np__61)|X1=k1_nat_1(np__1,np__60)|X1=k1_nat_1(np__1,np__59)|X1=k1_nat_1(np__1,np__58)|X1=k1_nat_1(np__1,np__57)|X1=k1_nat_1(np__1,np__56)|X1=k1_nat_1(np__1,np__55)|X1=k1_nat_1(np__1,np__54)|X1=k1_nat_1(np__1,np__53)|X1=k1_nat_1(np__1,np__52)|X1=k1_nat_1(np__1,np__51)|X1=k1_nat_1(np__1,np__50)|X1=k1_nat_1(np__1,np__49)|X1=k1_nat_1(np__1,np__48)|X1=k1_nat_1(np__1,np__47)|X1=k1_nat_1(np__1,np__46)|X1=k1_nat_1(np__1,np__45)|X1=k1_nat_1(np__1,np__44)|X1=k1_nat_1(np__1,np__43)|X1=k1_nat_1(np__1,np__42)|X1=k1_nat_1(np__1,np__41)|X1=k1_nat_1(np__1,np__40)|X1=k1_nat_1(np__1,np__39)|X1=k1_nat_1(np__1,np__38)|X1=k1_nat_1(np__1,np__37)|X1=k1_nat_1(np__1,np__36)|X1=k1_nat_1(np__1,np__35)|X1=k1_nat_1(np__1,np__34)|X1=k1_nat_1(np__1,np__33)|X1=k1_nat_1(np__1,np__32)|X1=k1_nat_1(np__1,np__31)|X1=k1_nat_1(np__1,np__30)|X1=k1_nat_1(np__1,np__29)|X1=k1_nat_1(np__1,np__28)|X1=k1_nat_1(np__1,np__27)|X1=k1_nat_1(np__1,np__26)|X1=k1_nat_1(np__1,np__25)|X1=k1_nat_1(np__1,np__24)|X1=k1_nat_1(np__1,np__23)|X1=k1_nat_1(np__1,np__22)|X1=k1_nat_1(np__1,np__21)|X1=k1_nat_1(np__1,np__20)|X1=k1_nat_1(np__1,np__19)|X1=k1_nat_1(np__1,np__18)|X1=k1_nat_1(np__1,np__17)|X1=k1_nat_1(np__1,np__15)|X1=k1_nat_1(np__1,np__14)|X1=k1_nat_1(np__1,np__13)|X1=k1_nat_1(np__1,np__12)|X1=k1_nat_1(np__1,np__11)|X1=k1_nat_1(np__1,np__10)|X1=k1_nat_1(np__1,np__9)|X1=k1_nat_1(np__1,np__8)|X1=k1_nat_1(np__1,np__7)|X1=k1_nat_1(np__1,np__6)|X1=k1_nat_1(np__1,np__5)|X1=k1_nat_1(np__1,np__3)|X1=k1_nat_1(np__1,np__1)|np__1=X1|X1=np__64|X1=np__17|X1=np__5|X1=np__3|~epred1_2(np__1,X1)|~r1_xxreal_0(X1,np__64)|~r1_xxreal_0(np__1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_196]), c_0_151])]), c_0_197]), c_0_198]), c_0_199])).
cnf(c_0_225, plain, (k1_nat_1(np__1,np__38)=np__39), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_201]), c_0_202]), c_0_151]), c_0_169])])).
cnf(c_0_226, plain, (k1_nat_1(np__1,np__37)=np__38), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_201]), c_0_204]), c_0_151]), c_0_170])])).
cnf(c_0_227, plain, (k1_nat_1(np__1,np__36)=np__37), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_201]), c_0_206]), c_0_151]), c_0_171])])).
cnf(c_0_228, plain, (k1_nat_1(np__1,np__3)=np__4), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_207]), c_0_208]), c_0_172]), c_0_151])])).
cnf(c_0_229, plain, (k1_nat_1(np__40,X1)=k1_nat_1(X1,np__40)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_209])).
cnf(c_0_230, plain, (k2_xcmplx_0(np__1,np__40)=np__41), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r40_r41])).
cnf(c_0_231, plain, (k1_nat_1(np__29,X1)=k2_xcmplx_0(np__29,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_210])).
cnf(c_0_232, plain, (k2_xcmplx_0(np__29,np__1)=np__30), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r29_r1_r30])).
cnf(c_0_233, plain, (k1_nat_1(np__30,X1)=k2_xcmplx_0(np__30,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_211])).
cnf(c_0_234, plain, (k2_xcmplx_0(np__30,np__1)=np__31), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r30_r1_r31])).
cnf(c_0_235, plain, (k1_nat_1(np__34,X1)=k2_xcmplx_0(np__34,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_212])).
cnf(c_0_236, plain, (k2_xcmplx_0(np__34,np__1)=np__35), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r34_r1_r35])).
cnf(c_0_237, plain, (k1_nat_1(np__35,X1)=k2_xcmplx_0(np__35,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_213])).
cnf(c_0_238, plain, (k2_xcmplx_0(np__35,np__1)=np__36), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r35_r1_r36])).
cnf(c_0_239, plain, (k1_nat_1(np__61,X1)=k1_nat_1(X1,np__61)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_214])).
cnf(c_0_240, plain, (k2_xcmplx_0(np__1,np__61)=np__62), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r61_r62])).
cnf(c_0_241, plain, (k1_nat_1(np__62,X1)=k1_nat_1(X1,np__62)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_215])).
cnf(c_0_242, plain, (k2_xcmplx_0(np__1,np__62)=np__63), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r62_r63])).
cnf(c_0_243, plain, (v7_ordinal1(np__45)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_216]), c_0_152])])).
cnf(c_0_244, plain, (v7_ordinal1(np__46)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_217]), c_0_152])])).
cnf(c_0_245, plain, (v7_ordinal1(np__44)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_218]), c_0_152])])).
cnf(c_0_246, plain, (v7_ordinal1(np__43)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_219]), c_0_152])])).
cnf(c_0_247, plain, (v7_ordinal1(np__28)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_220]), c_0_152])])).
cnf(c_0_248, plain, (v7_ordinal1(np__27)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_221]), c_0_152])])).
cnf(c_0_249, plain, (v7_ordinal1(np__26)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_222]), c_0_152])])).
cnf(c_0_250, plain, (v7_ordinal1(np__25)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_223]), c_0_152])])).
cnf(c_0_251, plain, (m1_subset_1(np__24,k5_numbers)), inference(split_conjunct,[status(thm)],[spc24_numerals])).
cnf(c_0_252, plain, (m1_subset_1(np__5,k5_numbers)), inference(split_conjunct,[status(thm)],[spc5_numerals])).
cnf(c_0_253, plain, (m1_subset_1(np__21,k5_numbers)), inference(split_conjunct,[status(thm)],[spc21_numerals])).
cnf(c_0_254, plain, (m1_subset_1(np__22,k5_numbers)), inference(split_conjunct,[status(thm)],[spc22_numerals])).
cnf(c_0_255, plain, (m1_subset_1(np__23,k5_numbers)), inference(split_conjunct,[status(thm)],[spc23_numerals])).
cnf(c_0_256, plain, (m1_subset_1(np__47,k5_numbers)), inference(split_conjunct,[status(thm)],[spc47_numerals])).
cnf(c_0_257, plain, (X1=k1_nat_1(np__1,np__1)|X1=k1_nat_1(np__1,np__5)|X1=k1_nat_1(np__1,np__6)|X1=k1_nat_1(np__1,np__7)|X1=k1_nat_1(np__1,np__8)|X1=k1_nat_1(np__1,np__9)|X1=k1_nat_1(np__1,np__10)|X1=k1_nat_1(np__1,np__11)|X1=k1_nat_1(np__1,np__12)|X1=k1_nat_1(np__1,np__13)|X1=k1_nat_1(np__1,np__14)|X1=k1_nat_1(np__1,np__15)|X1=k1_nat_1(np__1,np__17)|X1=k1_nat_1(np__1,np__18)|X1=k1_nat_1(np__1,np__19)|X1=k1_nat_1(np__1,np__20)|X1=k1_nat_1(np__1,np__21)|X1=k1_nat_1(np__1,np__22)|X1=k1_nat_1(np__1,np__23)|X1=k1_nat_1(np__1,np__24)|X1=k1_nat_1(np__1,np__25)|X1=k1_nat_1(np__1,np__26)|X1=k1_nat_1(np__1,np__27)|X1=k1_nat_1(np__1,np__28)|X1=k1_nat_1(np__1,np__29)|X1=k1_nat_1(np__1,np__30)|X1=k1_nat_1(np__1,np__31)|X1=k1_nat_1(np__1,np__32)|X1=k1_nat_1(np__1,np__33)|X1=k1_nat_1(np__1,np__34)|X1=k1_nat_1(np__1,np__35)|X1=k1_nat_1(np__1,np__39)|X1=k1_nat_1(np__1,np__40)|X1=k1_nat_1(np__1,np__41)|X1=k1_nat_1(np__1,np__42)|X1=k1_nat_1(np__1,np__43)|X1=k1_nat_1(np__1,np__44)|X1=k1_nat_1(np__1,np__45)|X1=k1_nat_1(np__1,np__46)|X1=k1_nat_1(np__1,np__47)|X1=k1_nat_1(np__1,np__48)|X1=k1_nat_1(np__1,np__49)|X1=k1_nat_1(np__1,np__50)|X1=k1_nat_1(np__1,np__51)|X1=k1_nat_1(np__1,np__52)|X1=k1_nat_1(np__1,np__53)|X1=k1_nat_1(np__1,np__54)|X1=k1_nat_1(np__1,np__55)|X1=k1_nat_1(np__1,np__56)|X1=k1_nat_1(np__1,np__57)|X1=k1_nat_1(np__1,np__58)|X1=k1_nat_1(np__1,np__59)|X1=k1_nat_1(np__1,np__60)|X1=k1_nat_1(np__1,np__61)|X1=k1_nat_1(np__1,np__62)|X1=np__3|X1=np__5|X1=np__17|X1=np__64|np__1=X1|X1=np__4|X1=np__37|X1=np__38|X1=np__39|~epred1_2(np__1,X1)|~r1_xxreal_0(X1,np__64)|~r1_xxreal_0(np__1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_224, c_0_225]), c_0_226]), c_0_227]), c_0_228])).
cnf(c_0_258, plain, (k1_nat_1(np__1,np__40)=np__41), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_229]), c_0_230]), c_0_188]), c_0_151])])).
cnf(c_0_259, plain, (k1_nat_1(np__1,np__29)=np__30), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_201]), c_0_232]), c_0_151]), c_0_189])])).
cnf(c_0_260, plain, (k1_nat_1(np__1,np__30)=np__31), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_201]), c_0_234]), c_0_151]), c_0_190])])).
cnf(c_0_261, plain, (k1_nat_1(np__1,np__34)=np__35), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_201]), c_0_236]), c_0_151]), c_0_191])])).
cnf(c_0_262, plain, (k1_nat_1(np__1,np__35)=np__36), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_201]), c_0_238]), c_0_151]), c_0_192])])).
cnf(c_0_263, plain, (k1_nat_1(np__1,np__61)=np__62), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_239]), c_0_240]), c_0_193]), c_0_151])])).
cnf(c_0_264, plain, (k1_nat_1(np__1,np__62)=np__63), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_241]), c_0_242]), c_0_194]), c_0_151])])).
cnf(c_0_265, plain, (k1_nat_1(np__45,X1)=k1_nat_1(X1,np__45)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_243])).
cnf(c_0_266, plain, (k2_xcmplx_0(np__1,np__45)=np__46), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r45_r46])).
cnf(c_0_267, plain, (k1_nat_1(np__46,X1)=k1_nat_1(X1,np__46)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_244])).
cnf(c_0_268, plain, (k2_xcmplx_0(np__1,np__46)=np__47), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r46_r47])).
cnf(c_0_269, plain, (k1_nat_1(np__44,X1)=k1_nat_1(X1,np__44)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_245])).
cnf(c_0_270, plain, (k2_xcmplx_0(np__1,np__44)=np__45), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r44_r45])).
cnf(c_0_271, plain, (k1_nat_1(np__43,X1)=k1_nat_1(X1,np__43)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_246])).
cnf(c_0_272, plain, (k2_xcmplx_0(np__1,np__43)=np__44), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r43_r44])).
cnf(c_0_273, plain, (k1_nat_1(np__28,X1)=k2_xcmplx_0(np__28,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_247])).
cnf(c_0_274, plain, (k2_xcmplx_0(np__28,np__1)=np__29), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r28_r1_r29])).
cnf(c_0_275, plain, (k1_nat_1(np__27,X1)=k2_xcmplx_0(np__27,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_248])).
cnf(c_0_276, plain, (k2_xcmplx_0(np__27,np__1)=np__28), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r27_r1_r28])).
cnf(c_0_277, plain, (k1_nat_1(np__26,X1)=k2_xcmplx_0(np__26,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_249])).
cnf(c_0_278, plain, (k2_xcmplx_0(np__26,np__1)=np__27), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r26_r1_r27])).
cnf(c_0_279, plain, (k1_nat_1(np__25,X1)=k2_xcmplx_0(np__25,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_250])).
cnf(c_0_280, plain, (k2_xcmplx_0(np__25,np__1)=np__26), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r25_r1_r26])).
cnf(c_0_281, plain, (v7_ordinal1(np__24)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_251]), c_0_152])])).
cnf(c_0_282, plain, (v7_ordinal1(np__5)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_252]), c_0_152])])).
cnf(c_0_283, plain, (v7_ordinal1(np__21)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_253]), c_0_152])])).
cnf(c_0_284, plain, (v7_ordinal1(np__22)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_254]), c_0_152])])).
cnf(c_0_285, plain, (v7_ordinal1(np__23)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_255]), c_0_152])])).
cnf(c_0_286, plain, (v7_ordinal1(np__47)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_256]), c_0_152])])).
cnf(c_0_287, plain, (m1_subset_1(np__7,k5_numbers)), inference(split_conjunct,[status(thm)],[spc7_numerals])).
cnf(c_0_288, plain, (m1_subset_1(np__49,k5_numbers)), inference(split_conjunct,[status(thm)],[spc49_numerals])).
cnf(c_0_289, plain, (m1_subset_1(np__48,k5_numbers)), inference(split_conjunct,[status(thm)],[spc48_numerals])).
cnf(c_0_290, plain, (m1_subset_1(np__20,k5_numbers)), inference(split_conjunct,[status(thm)],[spc20_numerals])).
cnf(c_0_291, plain, (m1_subset_1(np__19,k5_numbers)), inference(split_conjunct,[status(thm)],[spc19_numerals])).
cnf(c_0_292, plain, (m1_subset_1(np__18,k5_numbers)), inference(split_conjunct,[status(thm)],[spc18_numerals])).
cnf(c_0_293, plain, (m1_subset_1(np__6,k5_numbers)), inference(split_conjunct,[status(thm)],[spc6_numerals])).
cnf(c_0_294, plain, (X1=k1_nat_1(np__1,np__60)|X1=k1_nat_1(np__1,np__59)|X1=k1_nat_1(np__1,np__58)|X1=k1_nat_1(np__1,np__57)|X1=k1_nat_1(np__1,np__56)|X1=k1_nat_1(np__1,np__55)|X1=k1_nat_1(np__1,np__54)|X1=k1_nat_1(np__1,np__53)|X1=k1_nat_1(np__1,np__52)|X1=k1_nat_1(np__1,np__51)|X1=k1_nat_1(np__1,np__50)|X1=k1_nat_1(np__1,np__49)|X1=k1_nat_1(np__1,np__48)|X1=k1_nat_1(np__1,np__47)|X1=k1_nat_1(np__1,np__46)|X1=k1_nat_1(np__1,np__45)|X1=k1_nat_1(np__1,np__44)|X1=k1_nat_1(np__1,np__43)|X1=k1_nat_1(np__1,np__42)|X1=k1_nat_1(np__1,np__41)|X1=k1_nat_1(np__1,np__39)|X1=k1_nat_1(np__1,np__33)|X1=k1_nat_1(np__1,np__32)|X1=k1_nat_1(np__1,np__31)|X1=k1_nat_1(np__1,np__28)|X1=k1_nat_1(np__1,np__27)|X1=k1_nat_1(np__1,np__26)|X1=k1_nat_1(np__1,np__25)|X1=k1_nat_1(np__1,np__24)|X1=k1_nat_1(np__1,np__23)|X1=k1_nat_1(np__1,np__22)|X1=k1_nat_1(np__1,np__21)|X1=k1_nat_1(np__1,np__20)|X1=k1_nat_1(np__1,np__19)|X1=k1_nat_1(np__1,np__18)|X1=k1_nat_1(np__1,np__17)|X1=k1_nat_1(np__1,np__15)|X1=k1_nat_1(np__1,np__14)|X1=k1_nat_1(np__1,np__13)|X1=k1_nat_1(np__1,np__12)|X1=k1_nat_1(np__1,np__11)|X1=k1_nat_1(np__1,np__10)|X1=k1_nat_1(np__1,np__9)|X1=k1_nat_1(np__1,np__8)|X1=k1_nat_1(np__1,np__7)|X1=k1_nat_1(np__1,np__6)|X1=k1_nat_1(np__1,np__5)|X1=k1_nat_1(np__1,np__1)|X1=np__39|X1=np__38|X1=np__37|X1=np__4|np__1=X1|X1=np__64|X1=np__17|X1=np__5|X1=np__3|X1=np__63|X1=np__62|X1=np__41|X1=np__36|X1=np__35|X1=np__31|X1=np__30|~epred1_2(np__1,X1)|~r1_xxreal_0(X1,np__64)|~r1_xxreal_0(np__1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_257, c_0_258]), c_0_259]), c_0_260]), c_0_261]), c_0_262]), c_0_263]), c_0_264])).
cnf(c_0_295, plain, (k1_nat_1(np__1,np__45)=np__46), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_265]), c_0_266]), c_0_216]), c_0_151])])).
cnf(c_0_296, plain, (k1_nat_1(np__1,np__46)=np__47), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_267]), c_0_268]), c_0_217]), c_0_151])])).
cnf(c_0_297, plain, (k1_nat_1(np__1,np__44)=np__45), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_269]), c_0_270]), c_0_218]), c_0_151])])).
cnf(c_0_298, plain, (k1_nat_1(np__1,np__43)=np__44), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_271]), c_0_272]), c_0_219]), c_0_151])])).
cnf(c_0_299, plain, (k1_nat_1(np__1,np__28)=np__29), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_273, c_0_201]), c_0_274]), c_0_151]), c_0_220])])).
cnf(c_0_300, plain, (k1_nat_1(np__1,np__27)=np__28), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_275, c_0_201]), c_0_276]), c_0_151]), c_0_221])])).
cnf(c_0_301, plain, (k1_nat_1(np__1,np__26)=np__27), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_277, c_0_201]), c_0_278]), c_0_151]), c_0_222])])).
cnf(c_0_302, plain, (k1_nat_1(np__1,np__25)=np__26), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_279, c_0_201]), c_0_280]), c_0_151]), c_0_223])])).
cnf(c_0_303, plain, (k1_nat_1(np__24,X1)=k2_xcmplx_0(np__24,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_281])).
cnf(c_0_304, plain, (k2_xcmplx_0(np__24,np__1)=np__25), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r24_r1_r25])).
cnf(c_0_305, plain, (k1_nat_1(np__5,X1)=k1_nat_1(X1,np__5)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_282])).
cnf(c_0_306, plain, (k2_xcmplx_0(np__1,np__5)=np__6), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r5_r6])).
cnf(c_0_307, plain, (k1_nat_1(np__21,X1)=k2_xcmplx_0(np__21,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_283])).
cnf(c_0_308, plain, (k2_xcmplx_0(np__21,np__1)=np__22), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r21_r1_r22])).
cnf(c_0_309, plain, (k1_nat_1(np__22,X1)=k2_xcmplx_0(np__22,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_284])).
cnf(c_0_310, plain, (k2_xcmplx_0(np__22,np__1)=np__23), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r22_r1_r23])).
cnf(c_0_311, plain, (k1_nat_1(np__23,X1)=k2_xcmplx_0(np__23,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_285])).
cnf(c_0_312, plain, (k2_xcmplx_0(np__23,np__1)=np__24), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r23_r1_r24])).
cnf(c_0_313, plain, (k1_nat_1(np__47,X1)=k1_nat_1(X1,np__47)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_286])).
cnf(c_0_314, plain, (k2_xcmplx_0(np__1,np__47)=np__48), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r47_r48])).
cnf(c_0_315, plain, (v7_ordinal1(np__7)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_287]), c_0_152])])).
cnf(c_0_316, plain, (v7_ordinal1(np__49)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_288]), c_0_152])])).
cnf(c_0_317, plain, (v7_ordinal1(np__48)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_289]), c_0_152])])).
cnf(c_0_318, plain, (v7_ordinal1(np__20)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_290]), c_0_152])])).
cnf(c_0_319, plain, (v7_ordinal1(np__19)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_291]), c_0_152])])).
cnf(c_0_320, plain, (v7_ordinal1(np__18)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_292]), c_0_152])])).
cnf(c_0_321, plain, (v7_ordinal1(np__6)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_293]), c_0_152])])).
cnf(c_0_322, plain, (m1_subset_1(np__50,k5_numbers)), inference(split_conjunct,[status(thm)],[spc50_numerals])).
cnf(c_0_323, plain, (m1_subset_1(np__13,k5_numbers)), inference(split_conjunct,[status(thm)],[spc13_numerals])).
cnf(c_0_324, plain, (m1_subset_1(np__14,k5_numbers)), inference(split_conjunct,[status(thm)],[spc14_numerals])).
cnf(c_0_325, plain, (m1_subset_1(np__51,k5_numbers)), inference(split_conjunct,[status(thm)],[spc51_numerals])).
cnf(c_0_326, plain, (m1_subset_1(np__52,k5_numbers)), inference(split_conjunct,[status(thm)],[spc52_numerals])).
cnf(c_0_327, plain, (X1=k1_nat_1(np__1,np__1)|X1=k1_nat_1(np__1,np__5)|X1=k1_nat_1(np__1,np__6)|X1=k1_nat_1(np__1,np__7)|X1=k1_nat_1(np__1,np__8)|X1=k1_nat_1(np__1,np__9)|X1=k1_nat_1(np__1,np__10)|X1=k1_nat_1(np__1,np__11)|X1=k1_nat_1(np__1,np__12)|X1=k1_nat_1(np__1,np__13)|X1=k1_nat_1(np__1,np__14)|X1=k1_nat_1(np__1,np__15)|X1=k1_nat_1(np__1,np__17)|X1=k1_nat_1(np__1,np__18)|X1=k1_nat_1(np__1,np__19)|X1=k1_nat_1(np__1,np__20)|X1=k1_nat_1(np__1,np__21)|X1=k1_nat_1(np__1,np__22)|X1=k1_nat_1(np__1,np__23)|X1=k1_nat_1(np__1,np__24)|X1=k1_nat_1(np__1,np__31)|X1=k1_nat_1(np__1,np__32)|X1=k1_nat_1(np__1,np__33)|X1=k1_nat_1(np__1,np__39)|X1=k1_nat_1(np__1,np__41)|X1=k1_nat_1(np__1,np__42)|X1=k1_nat_1(np__1,np__47)|X1=k1_nat_1(np__1,np__48)|X1=k1_nat_1(np__1,np__49)|X1=k1_nat_1(np__1,np__50)|X1=k1_nat_1(np__1,np__51)|X1=k1_nat_1(np__1,np__52)|X1=k1_nat_1(np__1,np__53)|X1=k1_nat_1(np__1,np__54)|X1=k1_nat_1(np__1,np__55)|X1=k1_nat_1(np__1,np__56)|X1=k1_nat_1(np__1,np__57)|X1=k1_nat_1(np__1,np__58)|X1=k1_nat_1(np__1,np__59)|X1=k1_nat_1(np__1,np__60)|X1=np__30|X1=np__31|X1=np__35|X1=np__36|X1=np__41|X1=np__62|X1=np__63|X1=np__3|X1=np__5|X1=np__17|X1=np__64|np__1=X1|X1=np__4|X1=np__37|X1=np__38|X1=np__39|X1=np__26|X1=np__27|X1=np__28|X1=np__29|X1=np__44|X1=np__45|X1=np__46|X1=np__47|~epred1_2(np__1,X1)|~r1_xxreal_0(X1,np__64)|~r1_xxreal_0(np__1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_294, c_0_295]), c_0_296]), c_0_297]), c_0_298]), c_0_299]), c_0_300]), c_0_301]), c_0_302])).
cnf(c_0_328, plain, (k1_nat_1(np__1,np__24)=np__25), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_303, c_0_201]), c_0_304]), c_0_151]), c_0_251])])).
cnf(c_0_329, plain, (k1_nat_1(np__1,np__5)=np__6), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_305]), c_0_306]), c_0_252]), c_0_151])])).
cnf(c_0_330, plain, (k1_nat_1(np__1,np__21)=np__22), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_307, c_0_201]), c_0_308]), c_0_151]), c_0_253])])).
cnf(c_0_331, plain, (k1_nat_1(np__1,np__22)=np__23), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_309, c_0_201]), c_0_310]), c_0_151]), c_0_254])])).
cnf(c_0_332, plain, (k1_nat_1(np__1,np__23)=np__24), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_311, c_0_201]), c_0_312]), c_0_151]), c_0_255])])).
cnf(c_0_333, plain, (k1_nat_1(np__1,np__47)=np__48), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_313]), c_0_314]), c_0_256]), c_0_151])])).
cnf(c_0_334, plain, (k1_nat_1(np__7,X1)=k1_nat_1(X1,np__7)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_315])).
cnf(c_0_335, plain, (k2_xcmplx_0(np__1,np__7)=np__8), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r7_r8])).
cnf(c_0_336, plain, (k1_nat_1(np__49,X1)=k1_nat_1(X1,np__49)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_316])).
cnf(c_0_337, plain, (k2_xcmplx_0(np__1,np__49)=np__50), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r49_r50])).
cnf(c_0_338, plain, (k1_nat_1(np__48,X1)=k1_nat_1(X1,np__48)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_317])).
cnf(c_0_339, plain, (k2_xcmplx_0(np__1,np__48)=np__49), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r48_r49])).
cnf(c_0_340, plain, (k1_nat_1(np__20,X1)=k2_xcmplx_0(np__20,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_318])).
cnf(c_0_341, plain, (k2_xcmplx_0(np__20,np__1)=np__21), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r20_r1_r21])).
cnf(c_0_342, plain, (k1_nat_1(np__19,X1)=k2_xcmplx_0(np__19,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_319])).
cnf(c_0_343, plain, (k2_xcmplx_0(np__19,np__1)=np__20), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r19_r1_r20])).
cnf(c_0_344, plain, (k1_nat_1(np__18,X1)=k2_xcmplx_0(np__18,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_320])).
cnf(c_0_345, plain, (k2_xcmplx_0(np__18,np__1)=np__19), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r18_r1_r19])).
cnf(c_0_346, plain, (k1_nat_1(np__6,X1)=k1_nat_1(X1,np__6)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_321])).
cnf(c_0_347, plain, (k2_xcmplx_0(np__1,np__6)=np__7), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r6_r7])).
cnf(c_0_348, plain, (v7_ordinal1(np__50)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_322]), c_0_152])])).
cnf(c_0_349, plain, (v7_ordinal1(np__13)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_323]), c_0_152])])).
cnf(c_0_350, plain, (v7_ordinal1(np__14)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_324]), c_0_152])])).
cnf(c_0_351, plain, (v7_ordinal1(np__51)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_325]), c_0_152])])).
cnf(c_0_352, plain, (v7_ordinal1(np__52)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_326]), c_0_152])])).
cnf(c_0_353, plain, (m1_subset_1(np__12,k5_numbers)), inference(split_conjunct,[status(thm)],[spc12_numerals])).
cnf(c_0_354, plain, (m1_subset_1(np__55,k5_numbers)), inference(split_conjunct,[status(thm)],[spc55_numerals])).
cnf(c_0_355, plain, (m1_subset_1(np__54,k5_numbers)), inference(split_conjunct,[status(thm)],[spc54_numerals])).
cnf(c_0_356, plain, (m1_subset_1(np__53,k5_numbers)), inference(split_conjunct,[status(thm)],[spc53_numerals])).
cnf(c_0_357, plain, (m1_subset_1(np__11,k5_numbers)), inference(split_conjunct,[status(thm)],[spc11_numerals])).
cnf(c_0_358, plain, (m1_subset_1(np__10,k5_numbers)), inference(split_conjunct,[status(thm)],[spc10_numerals])).
cnf(c_0_359, plain, (m1_subset_1(np__9,k5_numbers)), inference(split_conjunct,[status(thm)],[spc9_numerals])).
cnf(c_0_360, plain, (m1_subset_1(np__8,k5_numbers)), inference(split_conjunct,[status(thm)],[spc8_numerals])).
cnf(c_0_361, plain, (X1=k1_nat_1(np__1,np__60)|X1=k1_nat_1(np__1,np__59)|X1=k1_nat_1(np__1,np__58)|X1=k1_nat_1(np__1,np__57)|X1=k1_nat_1(np__1,np__56)|X1=k1_nat_1(np__1,np__55)|X1=k1_nat_1(np__1,np__54)|X1=k1_nat_1(np__1,np__53)|X1=k1_nat_1(np__1,np__52)|X1=k1_nat_1(np__1,np__51)|X1=k1_nat_1(np__1,np__50)|X1=k1_nat_1(np__1,np__49)|X1=k1_nat_1(np__1,np__48)|X1=k1_nat_1(np__1,np__42)|X1=k1_nat_1(np__1,np__41)|X1=k1_nat_1(np__1,np__39)|X1=k1_nat_1(np__1,np__33)|X1=k1_nat_1(np__1,np__32)|X1=k1_nat_1(np__1,np__31)|X1=k1_nat_1(np__1,np__20)|X1=k1_nat_1(np__1,np__19)|X1=k1_nat_1(np__1,np__18)|X1=k1_nat_1(np__1,np__17)|X1=k1_nat_1(np__1,np__15)|X1=k1_nat_1(np__1,np__14)|X1=k1_nat_1(np__1,np__13)|X1=k1_nat_1(np__1,np__12)|X1=k1_nat_1(np__1,np__11)|X1=k1_nat_1(np__1,np__10)|X1=k1_nat_1(np__1,np__9)|X1=k1_nat_1(np__1,np__8)|X1=k1_nat_1(np__1,np__7)|X1=k1_nat_1(np__1,np__6)|X1=k1_nat_1(np__1,np__1)|X1=np__47|X1=np__46|X1=np__45|X1=np__44|X1=np__29|X1=np__28|X1=np__27|X1=np__26|X1=np__39|X1=np__38|X1=np__37|X1=np__4|np__1=X1|X1=np__64|X1=np__17|X1=np__5|X1=np__3|X1=np__63|X1=np__62|X1=np__41|X1=np__36|X1=np__35|X1=np__31|X1=np__30|X1=np__48|X1=np__25|X1=np__24|X1=np__23|X1=np__22|X1=np__6|~epred1_2(np__1,X1)|~r1_xxreal_0(X1,np__64)|~r1_xxreal_0(np__1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_327, c_0_328]), c_0_329]), c_0_330]), c_0_331]), c_0_332]), c_0_333])).
cnf(c_0_362, plain, (k1_nat_1(np__1,np__7)=np__8), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_334]), c_0_335]), c_0_287]), c_0_151])])).
cnf(c_0_363, plain, (k1_nat_1(np__1,np__49)=np__50), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_336]), c_0_337]), c_0_288]), c_0_151])])).
cnf(c_0_364, plain, (k1_nat_1(np__1,np__48)=np__49), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_338]), c_0_339]), c_0_289]), c_0_151])])).
cnf(c_0_365, plain, (k1_nat_1(np__1,np__20)=np__21), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_340, c_0_201]), c_0_341]), c_0_151]), c_0_290])])).
cnf(c_0_366, plain, (k1_nat_1(np__1,np__19)=np__20), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_342, c_0_201]), c_0_343]), c_0_151]), c_0_291])])).
cnf(c_0_367, plain, (k1_nat_1(np__1,np__18)=np__19), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_344, c_0_201]), c_0_345]), c_0_151]), c_0_292])])).
cnf(c_0_368, plain, (k1_nat_1(np__1,np__6)=np__7), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_346]), c_0_347]), c_0_293]), c_0_151])])).
cnf(c_0_369, plain, (k1_nat_1(np__50,X1)=k1_nat_1(X1,np__50)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_348])).
cnf(c_0_370, plain, (k2_xcmplx_0(np__1,np__50)=np__51), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r50_r51])).
cnf(c_0_371, plain, (k1_nat_1(np__13,X1)=k2_xcmplx_0(np__13,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_349])).
cnf(c_0_372, plain, (k2_xcmplx_0(np__13,np__1)=np__14), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r13_r1_r14])).
cnf(c_0_373, plain, (k1_nat_1(np__14,X1)=k2_xcmplx_0(np__14,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_350])).
cnf(c_0_374, plain, (k2_xcmplx_0(np__14,np__1)=np__15), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r14_r1_r15])).
cnf(c_0_375, plain, (k1_nat_1(np__51,X1)=k1_nat_1(X1,np__51)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_351])).
cnf(c_0_376, plain, (k2_xcmplx_0(np__1,np__51)=np__52), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r51_r52])).
cnf(c_0_377, plain, (k1_nat_1(np__52,X1)=k1_nat_1(X1,np__52)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_352])).
cnf(c_0_378, plain, (k2_xcmplx_0(np__1,np__52)=np__53), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r52_r53])).
cnf(c_0_379, plain, (v7_ordinal1(np__12)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_353]), c_0_152])])).
cnf(c_0_380, plain, (v7_ordinal1(np__55)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_354]), c_0_152])])).
cnf(c_0_381, plain, (v7_ordinal1(np__54)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_355]), c_0_152])])).
cnf(c_0_382, plain, (v7_ordinal1(np__53)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_356]), c_0_152])])).
cnf(c_0_383, plain, (v7_ordinal1(np__11)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_357]), c_0_152])])).
cnf(c_0_384, plain, (v7_ordinal1(np__10)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_358]), c_0_152])])).
cnf(c_0_385, plain, (v7_ordinal1(np__9)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_359]), c_0_152])])).
cnf(c_0_386, plain, (v7_ordinal1(np__8)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_360]), c_0_152])])).
cnf(c_0_387, plain, (m1_subset_1(np__57,k5_numbers)), inference(split_conjunct,[status(thm)],[spc57_numerals])).
cnf(c_0_388, plain, (X1=k1_nat_1(np__1,np__1)|X1=k1_nat_1(np__1,np__8)|X1=k1_nat_1(np__1,np__9)|X1=k1_nat_1(np__1,np__10)|X1=k1_nat_1(np__1,np__11)|X1=k1_nat_1(np__1,np__12)|X1=k1_nat_1(np__1,np__13)|X1=k1_nat_1(np__1,np__14)|X1=k1_nat_1(np__1,np__15)|X1=k1_nat_1(np__1,np__17)|X1=k1_nat_1(np__1,np__31)|X1=k1_nat_1(np__1,np__32)|X1=k1_nat_1(np__1,np__33)|X1=k1_nat_1(np__1,np__39)|X1=k1_nat_1(np__1,np__41)|X1=k1_nat_1(np__1,np__42)|X1=k1_nat_1(np__1,np__50)|X1=k1_nat_1(np__1,np__51)|X1=k1_nat_1(np__1,np__52)|X1=k1_nat_1(np__1,np__53)|X1=k1_nat_1(np__1,np__54)|X1=k1_nat_1(np__1,np__55)|X1=k1_nat_1(np__1,np__56)|X1=k1_nat_1(np__1,np__57)|X1=k1_nat_1(np__1,np__58)|X1=k1_nat_1(np__1,np__59)|X1=k1_nat_1(np__1,np__60)|X1=np__6|X1=np__22|X1=np__23|X1=np__24|X1=np__25|X1=np__48|X1=np__30|X1=np__31|X1=np__35|X1=np__36|X1=np__41|X1=np__62|X1=np__63|X1=np__3|X1=np__5|X1=np__17|X1=np__64|np__1=X1|X1=np__4|X1=np__37|X1=np__38|X1=np__39|X1=np__26|X1=np__27|X1=np__28|X1=np__29|X1=np__44|X1=np__45|X1=np__46|X1=np__47|X1=np__7|X1=np__8|X1=np__19|X1=np__20|X1=np__21|X1=np__49|X1=np__50|~epred1_2(np__1,X1)|~r1_xxreal_0(X1,np__64)|~r1_xxreal_0(np__1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_361, c_0_362]), c_0_363]), c_0_364]), c_0_365]), c_0_366]), c_0_367]), c_0_368])).
cnf(c_0_389, plain, (k1_nat_1(np__1,np__50)=np__51), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_369]), c_0_370]), c_0_322]), c_0_151])])).
cnf(c_0_390, plain, (k1_nat_1(np__1,np__13)=np__14), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_371, c_0_201]), c_0_372]), c_0_151]), c_0_323])])).
cnf(c_0_391, plain, (k1_nat_1(np__1,np__14)=np__15), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_373, c_0_201]), c_0_374]), c_0_151]), c_0_324])])).
cnf(c_0_392, plain, (k1_nat_1(np__1,np__51)=np__52), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_375]), c_0_376]), c_0_325]), c_0_151])])).
cnf(c_0_393, plain, (k1_nat_1(np__1,np__52)=np__53), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_377]), c_0_378]), c_0_326]), c_0_151])])).
cnf(c_0_394, plain, (k1_nat_1(np__12,X1)=k2_xcmplx_0(np__12,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_379])).
cnf(c_0_395, plain, (k2_xcmplx_0(np__12,np__1)=np__13), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r12_r1_r13])).
cnf(c_0_396, plain, (k1_nat_1(np__55,X1)=k1_nat_1(X1,np__55)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_380])).
cnf(c_0_397, plain, (k2_xcmplx_0(np__1,np__55)=np__56), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r55_r56])).
cnf(c_0_398, plain, (k1_nat_1(np__54,X1)=k1_nat_1(X1,np__54)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_381])).
cnf(c_0_399, plain, (k2_xcmplx_0(np__1,np__54)=np__55), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r54_r55])).
cnf(c_0_400, plain, (k1_nat_1(np__53,X1)=k1_nat_1(X1,np__53)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_382])).
cnf(c_0_401, plain, (k2_xcmplx_0(np__1,np__53)=np__54), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r53_r54])).
cnf(c_0_402, plain, (k1_nat_1(np__11,X1)=k2_xcmplx_0(np__11,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_383])).
cnf(c_0_403, plain, (k2_xcmplx_0(np__11,np__1)=np__12), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r11_r1_r12])).
cnf(c_0_404, plain, (k1_nat_1(np__10,X1)=k2_xcmplx_0(np__10,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_384])).
cnf(c_0_405, plain, (k2_xcmplx_0(np__10,np__1)=np__11), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r10_r1_r11])).
cnf(c_0_406, plain, (k1_nat_1(np__9,X1)=k2_xcmplx_0(np__9,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_385])).
cnf(c_0_407, plain, (k2_xcmplx_0(np__9,np__1)=np__10), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r9_r1_r10])).
cnf(c_0_408, plain, (k1_nat_1(np__8,X1)=k2_xcmplx_0(np__8,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_386])).
cnf(c_0_409, plain, (k2_xcmplx_0(np__8,np__1)=np__9), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r8_r1_r9])).
cnf(c_0_410, plain, (v7_ordinal1(np__57)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_387]), c_0_152])])).
cnf(c_0_411, plain, (m1_subset_1(np__59,k5_numbers)), inference(split_conjunct,[status(thm)],[spc59_numerals])).
cnf(c_0_412, plain, (m1_subset_1(np__58,k5_numbers)), inference(split_conjunct,[status(thm)],[spc58_numerals])).
cnf(c_0_413, plain, (m1_subset_1(np__56,k5_numbers)), inference(split_conjunct,[status(thm)],[spc56_numerals])).
cnf(c_0_414, plain, (X1=k1_nat_1(np__1,np__60)|X1=k1_nat_1(np__1,np__59)|X1=k1_nat_1(np__1,np__58)|X1=k1_nat_1(np__1,np__57)|X1=k1_nat_1(np__1,np__56)|X1=k1_nat_1(np__1,np__55)|X1=k1_nat_1(np__1,np__54)|X1=k1_nat_1(np__1,np__53)|X1=k1_nat_1(np__1,np__42)|X1=k1_nat_1(np__1,np__41)|X1=k1_nat_1(np__1,np__39)|X1=k1_nat_1(np__1,np__33)|X1=k1_nat_1(np__1,np__32)|X1=k1_nat_1(np__1,np__31)|X1=k1_nat_1(np__1,np__17)|X1=k1_nat_1(np__1,np__15)|X1=k1_nat_1(np__1,np__12)|X1=k1_nat_1(np__1,np__11)|X1=k1_nat_1(np__1,np__10)|X1=k1_nat_1(np__1,np__9)|X1=k1_nat_1(np__1,np__8)|X1=k1_nat_1(np__1,np__1)|X1=np__50|X1=np__49|X1=np__21|X1=np__20|X1=np__19|X1=np__8|X1=np__7|X1=np__47|X1=np__46|X1=np__45|X1=np__44|X1=np__29|X1=np__28|X1=np__27|X1=np__26|X1=np__39|X1=np__38|X1=np__37|X1=np__4|np__1=X1|X1=np__64|X1=np__17|X1=np__5|X1=np__3|X1=np__63|X1=np__62|X1=np__41|X1=np__36|X1=np__35|X1=np__31|X1=np__30|X1=np__48|X1=np__25|X1=np__24|X1=np__23|X1=np__22|X1=np__6|X1=np__53|X1=np__52|X1=np__51|X1=np__15|X1=np__14|~epred1_2(np__1,X1)|~r1_xxreal_0(X1,np__64)|~r1_xxreal_0(np__1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_388, c_0_389]), c_0_390]), c_0_391]), c_0_392]), c_0_393])).
cnf(c_0_415, plain, (k1_nat_1(np__1,np__12)=np__13), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_394, c_0_201]), c_0_395]), c_0_151]), c_0_353])])).
cnf(c_0_416, plain, (k1_nat_1(np__1,np__55)=np__56), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_396]), c_0_397]), c_0_354]), c_0_151])])).
cnf(c_0_417, plain, (k1_nat_1(np__1,np__54)=np__55), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_398]), c_0_399]), c_0_355]), c_0_151])])).
cnf(c_0_418, plain, (k1_nat_1(np__1,np__53)=np__54), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_400]), c_0_401]), c_0_356]), c_0_151])])).
cnf(c_0_419, plain, (k1_nat_1(np__1,np__11)=np__12), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_402, c_0_201]), c_0_403]), c_0_151]), c_0_357])])).
cnf(c_0_420, plain, (k1_nat_1(np__1,np__10)=np__11), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_404, c_0_201]), c_0_405]), c_0_151]), c_0_358])])).
cnf(c_0_421, plain, (k1_nat_1(np__1,np__9)=np__10), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_406, c_0_201]), c_0_407]), c_0_151]), c_0_359])])).
cnf(c_0_422, plain, (k1_nat_1(np__1,np__8)=np__9), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_408, c_0_201]), c_0_409]), c_0_151]), c_0_360])])).
cnf(c_0_423, plain, (k1_nat_1(np__57,X1)=k1_nat_1(X1,np__57)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_410])).
cnf(c_0_424, plain, (k2_xcmplx_0(np__1,np__57)=np__58), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r57_r58])).
cnf(c_0_425, plain, (v7_ordinal1(np__59)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_411]), c_0_152])])).
cnf(c_0_426, plain, (v7_ordinal1(np__58)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_412]), c_0_152])])).
cnf(c_0_427, plain, (v7_ordinal1(np__56)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_413]), c_0_152])])).
cnf(c_0_428, plain, (m1_subset_1(np__60,k5_numbers)), inference(split_conjunct,[status(thm)],[spc60_numerals])).
cnf(c_0_429, plain, (m1_subset_1(np__42,k5_numbers)), inference(split_conjunct,[status(thm)],[spc42_numerals])).
cnf(c_0_430, plain, (X1=k1_nat_1(np__1,np__1)|X1=k1_nat_1(np__1,np__15)|X1=k1_nat_1(np__1,np__17)|X1=k1_nat_1(np__1,np__31)|X1=k1_nat_1(np__1,np__32)|X1=k1_nat_1(np__1,np__33)|X1=k1_nat_1(np__1,np__39)|X1=k1_nat_1(np__1,np__41)|X1=k1_nat_1(np__1,np__42)|X1=k1_nat_1(np__1,np__56)|X1=k1_nat_1(np__1,np__57)|X1=k1_nat_1(np__1,np__58)|X1=k1_nat_1(np__1,np__59)|X1=k1_nat_1(np__1,np__60)|X1=np__14|X1=np__15|X1=np__51|X1=np__52|X1=np__53|X1=np__6|X1=np__22|X1=np__23|X1=np__24|X1=np__25|X1=np__48|X1=np__30|X1=np__31|X1=np__35|X1=np__36|X1=np__41|X1=np__62|X1=np__63|X1=np__3|X1=np__5|X1=np__17|X1=np__64|np__1=X1|X1=np__4|X1=np__37|X1=np__38|X1=np__39|X1=np__26|X1=np__27|X1=np__28|X1=np__29|X1=np__44|X1=np__45|X1=np__46|X1=np__47|X1=np__7|X1=np__8|X1=np__19|X1=np__20|X1=np__21|X1=np__49|X1=np__50|X1=np__9|X1=np__10|X1=np__11|X1=np__12|X1=np__13|X1=np__54|X1=np__55|X1=np__56|~epred1_2(np__1,X1)|~r1_xxreal_0(X1,np__64)|~r1_xxreal_0(np__1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_414, c_0_415]), c_0_416]), c_0_417]), c_0_418]), c_0_419]), c_0_420]), c_0_421]), c_0_422])).
cnf(c_0_431, plain, (k1_nat_1(np__1,np__57)=np__58), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_423]), c_0_424]), c_0_387]), c_0_151])])).
cnf(c_0_432, plain, (k1_nat_1(np__59,X1)=k1_nat_1(X1,np__59)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_425])).
cnf(c_0_433, plain, (k2_xcmplx_0(np__1,np__59)=np__60), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r59_r60])).
cnf(c_0_434, plain, (k1_nat_1(np__58,X1)=k2_xcmplx_0(np__58,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_426])).
cnf(c_0_435, plain, (k2_xcmplx_0(np__58,np__1)=np__59), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r58_r1_r59])).
cnf(c_0_436, plain, (k1_nat_1(np__56,X1)=k1_nat_1(X1,np__56)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_427])).
cnf(c_0_437, plain, (k2_xcmplx_0(np__1,np__56)=np__57), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r56_r57])).
cnf(c_0_438, plain, (k2_xcmplx_0(np__1,np__1)=np__2), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r1_r2])).
cnf(c_0_439, plain, (v7_ordinal1(np__60)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_428]), c_0_152])])).
cnf(c_0_440, plain, (v7_ordinal1(np__42)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_429]), c_0_152])])).
cnf(c_0_441, plain, (m1_subset_1(np__41,k5_numbers)), inference(split_conjunct,[status(thm)],[spc41_numerals])).
cnf(c_0_442, plain, (X1=k1_nat_1(np__1,np__60)|X1=k1_nat_1(np__1,np__59)|X1=k1_nat_1(np__1,np__58)|X1=k1_nat_1(np__1,np__56)|X1=k1_nat_1(np__1,np__42)|X1=k1_nat_1(np__1,np__41)|X1=k1_nat_1(np__1,np__39)|X1=k1_nat_1(np__1,np__33)|X1=k1_nat_1(np__1,np__32)|X1=k1_nat_1(np__1,np__31)|X1=k1_nat_1(np__1,np__17)|X1=k1_nat_1(np__1,np__15)|X1=k1_nat_1(np__1,np__1)|X1=np__56|X1=np__55|X1=np__54|X1=np__13|X1=np__12|X1=np__11|X1=np__10|X1=np__9|X1=np__50|X1=np__49|X1=np__21|X1=np__20|X1=np__19|X1=np__8|X1=np__7|X1=np__47|X1=np__46|X1=np__45|X1=np__44|X1=np__29|X1=np__28|X1=np__27|X1=np__26|X1=np__39|X1=np__38|X1=np__37|X1=np__4|np__1=X1|X1=np__64|X1=np__17|X1=np__5|X1=np__3|X1=np__63|X1=np__62|X1=np__41|X1=np__36|X1=np__35|X1=np__31|X1=np__30|X1=np__48|X1=np__25|X1=np__24|X1=np__23|X1=np__22|X1=np__6|X1=np__53|X1=np__52|X1=np__51|X1=np__15|X1=np__14|X1=np__58|~epred1_2(np__1,X1)|~r1_xxreal_0(X1,np__64)|~r1_xxreal_0(np__1,X1)), inference(rw,[status(thm)],[c_0_430, c_0_431])).
cnf(c_0_443, plain, (k1_nat_1(np__1,np__59)=np__60), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_432]), c_0_433]), c_0_411]), c_0_151])])).
cnf(c_0_444, plain, (k1_nat_1(np__1,np__58)=np__59), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_434]), c_0_435]), c_0_412]), c_0_151])])).
cnf(c_0_445, plain, (k1_nat_1(np__1,np__56)=np__57), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_436]), c_0_437]), c_0_413]), c_0_151])])).
cnf(c_0_446, plain, (k1_nat_1(np__1,np__1)=np__2), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_201]), c_0_438]), c_0_151])])).
cnf(c_0_447, plain, (k1_nat_1(np__60,X1)=k1_nat_1(X1,np__60)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_439])).
cnf(c_0_448, plain, (k2_xcmplx_0(np__1,np__60)=np__61), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r60_r61])).
cnf(c_0_449, plain, (k1_nat_1(np__42,X1)=k1_nat_1(X1,np__42)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_440])).
cnf(c_0_450, plain, (k2_xcmplx_0(np__1,np__42)=np__43), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r42_r43])).
cnf(c_0_451, plain, (v7_ordinal1(np__41)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_441]), c_0_152])])).
cnf(c_0_452, plain, (m1_subset_1(np__33,k5_numbers)), inference(split_conjunct,[status(thm)],[spc33_numerals])).
cnf(c_0_453, plain, (m1_subset_1(np__15,k5_numbers)), inference(split_conjunct,[status(thm)],[spc15_numerals])).
cnf(c_0_454, plain, (m1_subset_1(np__17,k5_numbers)), inference(split_conjunct,[status(thm)],[spc17_numerals])).
cnf(c_0_455, plain, (X1=k1_nat_1(np__1,np__15)|X1=k1_nat_1(np__1,np__17)|X1=k1_nat_1(np__1,np__31)|X1=k1_nat_1(np__1,np__32)|X1=k1_nat_1(np__1,np__33)|X1=k1_nat_1(np__1,np__39)|X1=k1_nat_1(np__1,np__41)|X1=k1_nat_1(np__1,np__42)|X1=k1_nat_1(np__1,np__60)|X1=np__58|X1=np__14|X1=np__15|X1=np__51|X1=np__52|X1=np__53|X1=np__6|X1=np__22|X1=np__23|X1=np__24|X1=np__25|X1=np__48|X1=np__30|X1=np__31|X1=np__35|X1=np__36|X1=np__41|X1=np__62|X1=np__63|X1=np__3|X1=np__5|X1=np__17|X1=np__64|np__1=X1|X1=np__4|X1=np__37|X1=np__38|X1=np__39|X1=np__26|X1=np__27|X1=np__28|X1=np__29|X1=np__44|X1=np__45|X1=np__46|X1=np__47|X1=np__7|X1=np__8|X1=np__19|X1=np__20|X1=np__21|X1=np__49|X1=np__50|X1=np__9|X1=np__10|X1=np__11|X1=np__12|X1=np__13|X1=np__54|X1=np__55|X1=np__56|X1=np__2|X1=np__57|X1=np__59|X1=np__60|~epred1_2(np__1,X1)|~r1_xxreal_0(X1,np__64)|~r1_xxreal_0(np__1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_442, c_0_443]), c_0_444]), c_0_445]), c_0_446])).
cnf(c_0_456, plain, (k1_nat_1(np__1,np__60)=np__61), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_447]), c_0_448]), c_0_428]), c_0_151])])).
cnf(c_0_457, plain, (k1_nat_1(np__1,np__42)=np__43), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_449]), c_0_450]), c_0_429]), c_0_151])])).
cnf(c_0_458, plain, (k1_nat_1(np__41,X1)=k1_nat_1(X1,np__41)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_451])).
cnf(c_0_459, plain, (k2_xcmplx_0(np__1,np__41)=np__42), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r41_r42])).
cnf(c_0_460, plain, (v7_ordinal1(np__33)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_452]), c_0_152])])).
cnf(c_0_461, plain, (v7_ordinal1(np__15)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_453]), c_0_152])])).
cnf(c_0_462, plain, (v7_ordinal1(np__17)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_454]), c_0_152])])).
cnf(c_0_463, plain, (m1_subset_1(np__31,k5_numbers)), inference(split_conjunct,[status(thm)],[spc31_numerals])).
fof(c_0_464, plain, ![X31, X32]:(~v7_ordinal1(X31)|~v7_ordinal1(X32)|v7_ordinal1(k2_xcmplx_0(X31,X32))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[fc1_nat_1])])).
cnf(c_0_465, plain, (X1=k1_nat_1(np__1,np__41)|X1=k1_nat_1(np__1,np__39)|X1=k1_nat_1(np__1,np__33)|X1=k1_nat_1(np__1,np__32)|X1=k1_nat_1(np__1,np__31)|X1=k1_nat_1(np__1,np__17)|X1=k1_nat_1(np__1,np__15)|X1=np__60|X1=np__59|X1=np__57|X1=np__2|X1=np__56|X1=np__55|X1=np__54|X1=np__13|X1=np__12|X1=np__11|X1=np__10|X1=np__9|X1=np__50|X1=np__49|X1=np__21|X1=np__20|X1=np__19|X1=np__8|X1=np__7|X1=np__47|X1=np__46|X1=np__45|X1=np__44|X1=np__29|X1=np__28|X1=np__27|X1=np__26|X1=np__39|X1=np__38|X1=np__37|X1=np__4|np__1=X1|X1=np__64|X1=np__17|X1=np__5|X1=np__3|X1=np__63|X1=np__62|X1=np__41|X1=np__36|X1=np__35|X1=np__31|X1=np__30|X1=np__48|X1=np__25|X1=np__24|X1=np__23|X1=np__22|X1=np__6|X1=np__53|X1=np__52|X1=np__51|X1=np__15|X1=np__14|X1=np__58|X1=np__61|X1=np__43|~epred1_2(np__1,X1)|~r1_xxreal_0(X1,np__64)|~r1_xxreal_0(np__1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_455, c_0_456]), c_0_457])).
cnf(c_0_466, plain, (k1_nat_1(np__1,np__41)=np__42), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_458]), c_0_459]), c_0_441]), c_0_151])])).
cnf(c_0_467, plain, (k1_nat_1(np__33,X1)=k1_nat_1(X1,np__33)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_460])).
cnf(c_0_468, plain, (k2_xcmplx_0(np__1,np__33)=np__34), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r33_r34])).
cnf(c_0_469, plain, (k1_nat_1(np__15,X1)=k1_nat_1(X1,np__15)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_461])).
cnf(c_0_470, plain, (k2_xcmplx_0(np__1,np__15)=np__16), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r15_r16])).
cnf(c_0_471, plain, (k1_nat_1(np__17,X1)=k2_xcmplx_0(np__17,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_462])).
cnf(c_0_472, plain, (k2_xcmplx_0(np__17,np__1)=np__18), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r17_r1_r18])).
cnf(c_0_473, plain, (v7_ordinal1(np__31)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_463]), c_0_152])])).
cnf(c_0_474, plain, (v7_ordinal1(k2_xcmplx_0(X1,X2))|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_464])).
cnf(c_0_475, plain, (k2_xcmplx_0(np__31,np__1)=np__32), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r31_r1_r32])).
fof(c_0_476, negated_conjecture, ~(![X1]:~(((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((r2_hidden(X1,k2_finseq_1(np__64))&X1!=np__1)&X1!=np__2)&X1!=np__3)&X1!=np__4)&X1!=np__5)&X1!=np__6)&X1!=np__7)&X1!=np__8)&X1!=np__9)&X1!=np__10)&X1!=np__11)&X1!=np__12)&X1!=np__13)&X1!=np__14)&X1!=np__15)&X1!=np__16)&X1!=np__17)&X1!=np__18)&X1!=np__19)&X1!=np__20)&X1!=np__21)&X1!=np__22)&X1!=np__23)&X1!=np__24)&X1!=np__25)&X1!=np__26)&X1!=np__27)&X1!=np__28)&X1!=np__29)&X1!=np__30)&X1!=np__31)&X1!=np__32)&X1!=np__33)&X1!=np__34)&X1!=np__35)&X1!=np__36)&X1!=np__37)&X1!=np__38)&X1!=np__39)&X1!=np__40)&X1!=np__41)&X1!=np__42)&X1!=np__43)&X1!=np__44)&X1!=np__45)&X1!=np__46)&X1!=np__47)&X1!=np__48)&X1!=np__49)&X1!=np__50)&X1!=np__51)&X1!=np__52)&X1!=np__53)&X1!=np__54)&X1!=np__55)&X1!=np__56)&X1!=np__57)&X1!=np__58)&X1!=np__59)&X1!=np__60)&X1!=np__61)&X1!=np__62)&X1!=np__63)&X1!=np__64))), inference(assume_negation,[status(cth)],[t23_descip_1])).
fof(c_0_477, plain, ![X36, X37, X38]:(~r2_hidden(X36,X37)|~m1_subset_1(X37,k1_zfmisc_1(X38))|m1_subset_1(X36,X38)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t4_subset])])).
fof(c_0_478, plain, ![X23]:(~v7_ordinal1(X23)|m1_subset_1(k2_finseq_1(X23),k1_zfmisc_1(k5_numbers))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_k2_finseq_1])])).
cnf(c_0_479, plain, (X1=k1_nat_1(np__1,np__15)|X1=k1_nat_1(np__1,np__17)|X1=k1_nat_1(np__1,np__31)|X1=k1_nat_1(np__1,np__32)|X1=k1_nat_1(np__1,np__33)|X1=k1_nat_1(np__1,np__39)|X1=np__43|X1=np__61|X1=np__58|X1=np__14|X1=np__15|X1=np__51|X1=np__52|X1=np__53|X1=np__6|X1=np__22|X1=np__23|X1=np__24|X1=np__25|X1=np__48|X1=np__30|X1=np__31|X1=np__35|X1=np__36|X1=np__41|X1=np__62|X1=np__63|X1=np__3|X1=np__5|X1=np__17|X1=np__64|np__1=X1|X1=np__4|X1=np__37|X1=np__38|X1=np__39|X1=np__26|X1=np__27|X1=np__28|X1=np__29|X1=np__44|X1=np__45|X1=np__46|X1=np__47|X1=np__7|X1=np__8|X1=np__19|X1=np__20|X1=np__21|X1=np__49|X1=np__50|X1=np__9|X1=np__10|X1=np__11|X1=np__12|X1=np__13|X1=np__54|X1=np__55|X1=np__56|X1=np__2|X1=np__57|X1=np__59|X1=np__60|X1=np__42|~epred1_2(np__1,X1)|~r1_xxreal_0(X1,np__64)|~r1_xxreal_0(np__1,X1)), inference(rw,[status(thm)],[c_0_465, c_0_466])).
cnf(c_0_480, plain, (k1_nat_1(np__1,np__33)=np__34), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_467]), c_0_468]), c_0_452]), c_0_151])])).
cnf(c_0_481, plain, (k1_nat_1(np__1,np__15)=np__16), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_469]), c_0_470]), c_0_453]), c_0_151])])).
cnf(c_0_482, plain, (k1_nat_1(np__1,np__17)=np__18), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_471]), c_0_472]), c_0_454]), c_0_151])])).
cnf(c_0_483, plain, (k1_nat_1(np__31,X1)=k1_nat_1(X1,np__31)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_156, c_0_473])).
cnf(c_0_484, plain, (k2_xcmplx_0(np__1,np__31)=np__32), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r31_r32])).
cnf(c_0_485, plain, (v7_ordinal1(np__32)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_474, c_0_475]), c_0_159]), c_0_473])])).
fof(c_0_486, plain, ![X15, X16]:(((r1_xxreal_0(np__1,X15)|~r2_hidden(X15,k2_finseq_1(X16))|~v7_ordinal1(X16)|~v7_ordinal1(X15))&(r1_xxreal_0(X15,X16)|~r2_hidden(X15,k2_finseq_1(X16))|~v7_ordinal1(X16)|~v7_ordinal1(X15)))&(~r1_xxreal_0(np__1,X15)|~r1_xxreal_0(X15,X16)|r2_hidden(X15,k2_finseq_1(X16))|~v7_ordinal1(X16)|~v7_ordinal1(X15))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[t1_finseq_1])])])])).
fof(c_0_487, negated_conjecture, ((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((((r2_hidden(esk1_0,k2_finseq_1(np__64))&esk1_0!=np__1)&esk1_0!=np__2)&esk1_0!=np__3)&esk1_0!=np__4)&esk1_0!=np__5)&esk1_0!=np__6)&esk1_0!=np__7)&esk1_0!=np__8)&esk1_0!=np__9)&esk1_0!=np__10)&esk1_0!=np__11)&esk1_0!=np__12)&esk1_0!=np__13)&esk1_0!=np__14)&esk1_0!=np__15)&esk1_0!=np__16)&esk1_0!=np__17)&esk1_0!=np__18)&esk1_0!=np__19)&esk1_0!=np__20)&esk1_0!=np__21)&esk1_0!=np__22)&esk1_0!=np__23)&esk1_0!=np__24)&esk1_0!=np__25)&esk1_0!=np__26)&esk1_0!=np__27)&esk1_0!=np__28)&esk1_0!=np__29)&esk1_0!=np__30)&esk1_0!=np__31)&esk1_0!=np__32)&esk1_0!=np__33)&esk1_0!=np__34)&esk1_0!=np__35)&esk1_0!=np__36)&esk1_0!=np__37)&esk1_0!=np__38)&esk1_0!=np__39)&esk1_0!=np__40)&esk1_0!=np__41)&esk1_0!=np__42)&esk1_0!=np__43)&esk1_0!=np__44)&esk1_0!=np__45)&esk1_0!=np__46)&esk1_0!=np__47)&esk1_0!=np__48)&esk1_0!=np__49)&esk1_0!=np__50)&esk1_0!=np__51)&esk1_0!=np__52)&esk1_0!=np__53)&esk1_0!=np__54)&esk1_0!=np__55)&esk1_0!=np__56)&esk1_0!=np__57)&esk1_0!=np__58)&esk1_0!=np__59)&esk1_0!=np__60)&esk1_0!=np__61)&esk1_0!=np__62)&esk1_0!=np__63)&esk1_0!=np__64), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_476])])])).
cnf(c_0_488, plain, (m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))), inference(split_conjunct,[status(thm)],[c_0_477])).
cnf(c_0_489, plain, (m1_subset_1(k2_finseq_1(X1),k1_zfmisc_1(k5_numbers))|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_478])).
cnf(c_0_490, plain, (m1_subset_1(np__64,k5_numbers)), inference(split_conjunct,[status(thm)],[spc64_numerals])).
cnf(c_0_491, plain, (X1=k1_nat_1(np__1,np__39)|X1=k1_nat_1(np__1,np__32)|X1=k1_nat_1(np__1,np__31)|X1=np__42|X1=np__60|X1=np__59|X1=np__57|X1=np__2|X1=np__56|X1=np__55|X1=np__54|X1=np__13|X1=np__12|X1=np__11|X1=np__10|X1=np__9|X1=np__50|X1=np__49|X1=np__21|X1=np__20|X1=np__19|X1=np__8|X1=np__7|X1=np__47|X1=np__46|X1=np__45|X1=np__44|X1=np__29|X1=np__28|X1=np__27|X1=np__26|X1=np__39|X1=np__38|X1=np__37|X1=np__4|np__1=X1|X1=np__64|X1=np__17|X1=np__5|X1=np__3|X1=np__63|X1=np__62|X1=np__41|X1=np__36|X1=np__35|X1=np__31|X1=np__30|X1=np__48|X1=np__25|X1=np__24|X1=np__23|X1=np__22|X1=np__6|X1=np__53|X1=np__52|X1=np__51|X1=np__15|X1=np__14|X1=np__58|X1=np__61|X1=np__43|X1=np__34|X1=np__18|X1=np__16|~epred1_2(np__1,X1)|~r1_xxreal_0(X1,np__64)|~r1_xxreal_0(np__1,X1)), inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[c_0_479, c_0_480]), c_0_481]), c_0_482])).
cnf(c_0_492, plain, (k1_nat_1(np__1,np__31)=np__32), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_483]), c_0_484]), c_0_463]), c_0_151])])).
cnf(c_0_493, plain, (k1_nat_1(np__32,X1)=k2_xcmplx_0(np__32,X1)|~m1_subset_1(X1,k5_numbers)), inference(spm,[status(thm)],[c_0_158, c_0_485])).
cnf(c_0_494, plain, (k2_xcmplx_0(np__32,np__1)=np__33), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r32_r1_r33])).
cnf(c_0_495, plain, (m1_subset_1(np__32,k5_numbers)), inference(split_conjunct,[status(thm)],[spc32_numerals])).
fof(c_0_496, axiom, ![X1]:(v7_ordinal1(X1)=>![X2]:(v7_ordinal1(X2)=>epred1_2(X1,X2))), inference(apply_def,[status(thm)],[l30_descip_1, c_0_146])).
cnf(c_0_497, plain, (r1_xxreal_0(X1,X2)|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_486])).
cnf(c_0_498, negated_conjecture, (r2_hidden(esk1_0,k2_finseq_1(np__64))), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_499, plain, (m1_subset_1(X1,k5_numbers)|~v7_ordinal1(X2)|~r2_hidden(X1,k2_finseq_1(X2))), inference(spm,[status(thm)],[c_0_488, c_0_489])).
cnf(c_0_500, plain, (v7_ordinal1(np__64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_490]), c_0_152])])).
cnf(c_0_501, plain, (r1_xxreal_0(np__1,X1)|~r2_hidden(X1,k2_finseq_1(X2))|~v7_ordinal1(X2)|~v7_ordinal1(X1)), inference(split_conjunct,[status(thm)],[c_0_486])).
cnf(c_0_502, plain, (X1=k1_nat_1(np__1,np__32)|X1=k1_nat_1(np__1,np__39)|X1=np__16|X1=np__18|X1=np__34|X1=np__43|X1=np__61|X1=np__58|X1=np__14|X1=np__15|X1=np__51|X1=np__52|X1=np__53|X1=np__6|X1=np__22|X1=np__23|X1=np__24|X1=np__25|X1=np__48|X1=np__30|X1=np__31|X1=np__35|X1=np__36|X1=np__41|X1=np__62|X1=np__63|X1=np__3|X1=np__5|X1=np__17|X1=np__64|np__1=X1|X1=np__4|X1=np__37|X1=np__38|X1=np__39|X1=np__26|X1=np__27|X1=np__28|X1=np__29|X1=np__44|X1=np__45|X1=np__46|X1=np__47|X1=np__7|X1=np__8|X1=np__19|X1=np__20|X1=np__21|X1=np__49|X1=np__50|X1=np__9|X1=np__10|X1=np__11|X1=np__12|X1=np__13|X1=np__54|X1=np__55|X1=np__56|X1=np__2|X1=np__57|X1=np__59|X1=np__60|X1=np__42|X1=np__32|~epred1_2(np__1,X1)|~r1_xxreal_0(X1,np__64)|~r1_xxreal_0(np__1,X1)), inference(rw,[status(thm)],[c_0_491, c_0_492])).
cnf(c_0_503, plain, (k1_nat_1(np__1,np__32)=np__33), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_493]), c_0_494]), c_0_495]), c_0_151])])).
fof(c_0_504, plain, ![X5, X6]:(~v7_ordinal1(X5)|(~v7_ordinal1(X6)|epred1_2(X5,X6))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[c_0_496])])])).
cnf(c_0_505, negated_conjecture, (r1_xxreal_0(esk1_0,np__64)|~v7_ordinal1(np__64)|~v7_ordinal1(esk1_0)), inference(spm,[status(thm)],[c_0_497, c_0_498])).
cnf(c_0_506, negated_conjecture, (m1_subset_1(esk1_0,k5_numbers)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_499, c_0_498]), c_0_500])])).
cnf(c_0_507, negated_conjecture, (r1_xxreal_0(np__1,esk1_0)|~v7_ordinal1(np__64)|~v7_ordinal1(esk1_0)), inference(spm,[status(thm)],[c_0_501, c_0_498])).
cnf(c_0_508, plain, (X1=k1_nat_1(np__1,np__39)|X1=np__32|X1=np__42|X1=np__60|X1=np__59|X1=np__57|X1=np__2|X1=np__56|X1=np__55|X1=np__54|X1=np__13|X1=np__12|X1=np__11|X1=np__10|X1=np__9|X1=np__50|X1=np__49|X1=np__21|X1=np__20|X1=np__19|X1=np__8|X1=np__7|X1=np__47|X1=np__46|X1=np__45|X1=np__44|X1=np__29|X1=np__28|X1=np__27|X1=np__26|X1=np__39|X1=np__38|X1=np__37|X1=np__4|np__1=X1|X1=np__64|X1=np__17|X1=np__5|X1=np__3|X1=np__63|X1=np__62|X1=np__41|X1=np__36|X1=np__35|X1=np__31|X1=np__30|X1=np__48|X1=np__25|X1=np__24|X1=np__23|X1=np__22|X1=np__6|X1=np__53|X1=np__52|X1=np__51|X1=np__15|X1=np__14|X1=np__58|X1=np__61|X1=np__43|X1=np__34|X1=np__18|X1=np__16|X1=np__33|~epred1_2(np__1,X1)|~r1_xxreal_0(X1,np__64)|~r1_xxreal_0(np__1,X1)), inference(rw,[status(thm)],[c_0_502, c_0_503])).
cnf(c_0_509, plain, (epred1_2(X1,X2)|~v7_ordinal1(X1)|~v7_ordinal1(X2)), inference(split_conjunct,[status(thm)],[c_0_504])).
cnf(c_0_510, negated_conjecture, (r1_xxreal_0(esk1_0,np__64)|~v7_ordinal1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_505, c_0_500])])).
cnf(c_0_511, negated_conjecture, (v7_ordinal1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_506]), c_0_152])])).
cnf(c_0_512, negated_conjecture, (r1_xxreal_0(np__1,esk1_0)|~v7_ordinal1(esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_507, c_0_500])])).
cnf(c_0_513, plain, (X1=k1_nat_1(np__1,np__39)|X1=np__33|X1=np__16|X1=np__18|X1=np__34|X1=np__43|X1=np__61|X1=np__58|X1=np__14|X1=np__15|X1=np__51|X1=np__52|X1=np__53|X1=np__6|X1=np__22|X1=np__23|X1=np__24|X1=np__25|X1=np__48|X1=np__30|X1=np__31|X1=np__35|X1=np__36|X1=np__41|X1=np__62|X1=np__63|X1=np__3|X1=np__5|X1=np__17|X1=np__64|np__1=X1|X1=np__4|X1=np__37|X1=np__38|X1=np__39|X1=np__26|X1=np__27|X1=np__28|X1=np__29|X1=np__44|X1=np__45|X1=np__46|X1=np__47|X1=np__7|X1=np__8|X1=np__19|X1=np__20|X1=np__21|X1=np__49|X1=np__50|X1=np__9|X1=np__10|X1=np__11|X1=np__12|X1=np__13|X1=np__54|X1=np__55|X1=np__56|X1=np__2|X1=np__57|X1=np__59|X1=np__60|X1=np__42|X1=np__32|~r1_xxreal_0(X1,np__64)|~r1_xxreal_0(np__1,X1)|~v7_ordinal1(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_508, c_0_509]), c_0_159])])).
cnf(c_0_514, negated_conjecture, (r1_xxreal_0(esk1_0,np__64)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_510, c_0_511])])).
cnf(c_0_515, negated_conjecture, (r1_xxreal_0(np__1,esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_512, c_0_511])])).
cnf(c_0_516, negated_conjecture, (esk1_0!=np__33), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_517, negated_conjecture, (esk1_0!=np__16), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_518, negated_conjecture, (esk1_0!=np__18), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_519, negated_conjecture, (esk1_0!=np__34), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_520, negated_conjecture, (esk1_0!=np__43), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_521, negated_conjecture, (esk1_0!=np__61), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_522, negated_conjecture, (esk1_0!=np__58), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_523, negated_conjecture, (esk1_0!=np__14), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_524, negated_conjecture, (esk1_0!=np__15), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_525, negated_conjecture, (esk1_0!=np__51), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_526, negated_conjecture, (esk1_0!=np__52), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_527, negated_conjecture, (esk1_0!=np__53), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_528, negated_conjecture, (esk1_0!=np__6), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_529, negated_conjecture, (esk1_0!=np__22), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_530, negated_conjecture, (esk1_0!=np__23), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_531, negated_conjecture, (esk1_0!=np__24), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_532, negated_conjecture, (esk1_0!=np__25), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_533, negated_conjecture, (esk1_0!=np__48), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_534, negated_conjecture, (esk1_0!=np__30), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_535, negated_conjecture, (esk1_0!=np__31), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_536, negated_conjecture, (esk1_0!=np__35), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_537, negated_conjecture, (esk1_0!=np__36), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_538, negated_conjecture, (esk1_0!=np__41), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_539, negated_conjecture, (esk1_0!=np__62), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_540, negated_conjecture, (esk1_0!=np__63), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_541, negated_conjecture, (esk1_0!=np__3), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_542, negated_conjecture, (esk1_0!=np__5), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_543, negated_conjecture, (esk1_0!=np__17), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_544, negated_conjecture, (esk1_0!=np__64), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_545, negated_conjecture, (esk1_0!=np__1), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_546, negated_conjecture, (esk1_0!=np__4), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_547, negated_conjecture, (esk1_0!=np__37), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_548, negated_conjecture, (esk1_0!=np__38), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_549, negated_conjecture, (esk1_0!=np__39), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_550, negated_conjecture, (esk1_0!=np__26), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_551, negated_conjecture, (esk1_0!=np__27), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_552, negated_conjecture, (esk1_0!=np__28), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_553, negated_conjecture, (esk1_0!=np__29), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_554, negated_conjecture, (esk1_0!=np__44), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_555, negated_conjecture, (esk1_0!=np__45), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_556, negated_conjecture, (esk1_0!=np__46), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_557, negated_conjecture, (esk1_0!=np__47), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_558, negated_conjecture, (esk1_0!=np__7), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_559, negated_conjecture, (esk1_0!=np__8), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_560, negated_conjecture, (esk1_0!=np__19), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_561, negated_conjecture, (esk1_0!=np__20), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_562, negated_conjecture, (esk1_0!=np__21), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_563, negated_conjecture, (esk1_0!=np__49), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_564, negated_conjecture, (esk1_0!=np__50), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_565, negated_conjecture, (esk1_0!=np__9), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_566, negated_conjecture, (esk1_0!=np__10), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_567, negated_conjecture, (esk1_0!=np__11), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_568, negated_conjecture, (esk1_0!=np__12), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_569, negated_conjecture, (esk1_0!=np__13), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_570, negated_conjecture, (esk1_0!=np__54), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_571, negated_conjecture, (esk1_0!=np__55), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_572, negated_conjecture, (esk1_0!=np__56), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_573, negated_conjecture, (esk1_0!=np__2), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_574, negated_conjecture, (esk1_0!=np__57), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_575, negated_conjecture, (esk1_0!=np__59), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_576, negated_conjecture, (esk1_0!=np__60), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_577, negated_conjecture, (esk1_0!=np__42), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_578, negated_conjecture, (esk1_0!=np__32), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_579, negated_conjecture, (k1_nat_1(np__1,np__39)=esk1_0), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_513, c_0_514]), c_0_515]), c_0_511])]), c_0_516]), c_0_517]), c_0_518]), c_0_519]), c_0_520]), c_0_521]), c_0_522]), c_0_523]), c_0_524]), c_0_525]), c_0_526]), c_0_527]), c_0_528]), c_0_529]), c_0_530]), c_0_531]), c_0_532]), c_0_533]), c_0_534]), c_0_535]), c_0_536]), c_0_537]), c_0_538]), c_0_539]), c_0_540]), c_0_541]), c_0_542]), c_0_543]), c_0_544]), c_0_545]), c_0_546]), c_0_547]), c_0_548]), c_0_549]), c_0_550]), c_0_551]), c_0_552]), c_0_553]), c_0_554]), c_0_555]), c_0_556]), c_0_557]), c_0_558]), c_0_559]), c_0_560]), c_0_561]), c_0_562]), c_0_563]), c_0_564]), c_0_565]), c_0_566]), c_0_567]), c_0_568]), c_0_569]), c_0_570]), c_0_571]), c_0_572]), c_0_573]), c_0_574]), c_0_575]), c_0_576]), c_0_577]), c_0_578])).
cnf(c_0_580, plain, (k2_xcmplx_0(np__1,np__39)=np__40), inference(split_conjunct,[status(thm)],[rqRealAdd__k2_xcmplx_0__r1_r39_r40])).
cnf(c_0_581, plain, (m1_subset_1(np__39,k5_numbers)), inference(split_conjunct,[status(thm)],[spc39_numerals])).
cnf(c_0_582, negated_conjecture, (esk1_0!=np__40), inference(split_conjunct,[status(thm)],[c_0_487])).
cnf(c_0_583, negated_conjecture, ($false), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_579]), c_0_580]), c_0_581]), c_0_151])]), c_0_582]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 584
# Proof object clause steps            : 426
# Proof object formula steps           : 158
# Proof object conjectures             : 78
# Proof object clause conjectures      : 75
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 208
# Proof object initial formulas used   : 142
# Proof object generating inferences   : 198
# Proof object simplifying inferences  : 522
# Training examples: 0 positive, 0 negative
# Parsed axioms                        : 227
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 488
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 488
# Processed clauses                    : 2329
# ...of these trivial                  : 295
# ...subsumed                          : 460
# ...remaining for further processing  : 1574
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 54
# Generated clauses                    : 17384
# ...of the previous two non-trivial   : 16490
# Contextual simplify-reflections      : 29
# Paramodulations                      : 17384
# Factorizations                       : 0
# Equation resolutions                 : 0
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 1517
#    Positive orientable unit clauses  : 784
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 133
#    Non-unit-clauses                  : 600
# Current number of unprocessed clauses: 14561
# ...number of literals in the above   : 58153
# Current number of archived formulas  : 0
# Current number of archived clauses   : 57
# Clause-clause subsumption calls (NU) : 1889388
# Rec. Clause-clause subsumption calls : 20492
# Non-unit clause-clause subsumptions  : 427
# Unit Clause-clause subsumption calls : 23442
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 44
# BW rewrite match successes           : 41
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 456695

# -------------------------------------------------
# User time                : 26.140 s
# System time              : 1.432 s
# Total time               : 27.572 s
# Maximum resident set size: 316164 pages
26.14user 1.47system 0:30.39elapsed 90%CPU (0avgtext+0avgdata 316164maxresident)k
0inputs+264outputs (0major+413451minor)pagefaults 0swaps
