2020-06-08 18:03:05.117209: I tensorflow/cc/saved_model/reader.cc:31] Reading SavedModel from: models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_49
2020-06-08 18:03:05.225216: I tensorflow/cc/saved_model/reader.cc:54] Reading meta graph with tags { serve }
2020-06-08 18:03:05.445484: I tensorflow/core/platform/cpu_feature_guard.cc:142] Your CPU supports instructions that this TensorFlow binary was not compiled to use: AVX512F
2020-06-08 18:03:06.218216: I tensorflow/cc/saved_model/loader.cc:202] Restoring SavedModel bundle.
2020-06-08 18:03:09.420971: I tensorflow/cc/saved_model/loader.cc:311] SavedModel load for tags { serve }; Status: success. Took 4303747 microseconds.
# ENIGMA: TensorFlow C library version 1.15.0
# ENIGMA: TensorFlow: model 'models_10_bb1/premsel_enigma_01_2020_T30_loop02_epoch_49' loaded (query=128; context=512; weigths=real; len_mult=0.000000).
# Preprocessing time       : 5.858 s

# Proof found!
# SZS status Theorem
# SZS output start CNFRefutation
fof(t12_polynom7, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:((~(v7_struct_0(X2))&(v13_algstr_0(X2)&(v3_rlvect_1(X2)&(v4_rlvect_1(X2)&(v4_vectsp_1(X2)&(v5_vectsp_1(X2)&l6_algstr_0(X2)))))))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))&(v3_polynom7(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,u1_struct_0(X2))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))))=>k5_polynom2(X1,X2,X3,X4)=k6_algstr_0(X2,k3_polynom7(X1,X2,X3),k3_polynom2(X1,k2_polynom7(X1,X2,X3),X2,X4)))))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t24_polynom7', t12_polynom7)).
fof(t16_polynom7, axiom, ![X1, X2]:((~(v2_struct_0(X2))&l2_struct_0(X2))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))&(v4_polynom7(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))))=>(k2_polynom7(X1,X2,X3)=k16_pre_poly(X1)&k3_polynom7(X1,X2,X3)=k3_polynom1(X1,X2,X3,k16_pre_poly(X1))))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t24_polynom7', t16_polynom7)).
fof(dt_l2_algstr_0, axiom, ![X1]:(l2_algstr_0(X1)=>(l2_struct_0(X1)&l1_algstr_0(X1))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t24_polynom7', dt_l2_algstr_0)).
fof(dt_l6_algstr_0, axiom, ![X1]:(l6_algstr_0(X1)=>(l2_algstr_0(X1)&l5_algstr_0(X1))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t24_polynom7', dt_l6_algstr_0)).
fof(dt_l1_algstr_0, axiom, ![X1]:(l1_algstr_0(X1)=>l1_struct_0(X1)), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t24_polynom7', dt_l1_algstr_0)).
fof(t14_polynom2, axiom, ![X1]:(v3_ordinal1(X1)=>![X2]:((~(v7_struct_0(X2))&(v4_vectsp_1(X2)&l6_algstr_0(X2)))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,X1,u1_struct_0(X2))&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))))=>k3_polynom2(X1,k16_pre_poly(X1),X2,X3)=k5_struct_0(X2)))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t24_polynom7', t14_polynom2)).
fof(cc2_struct_0, axiom, ![X1]:(l1_struct_0(X1)=>(~(v7_struct_0(X1))=>~(v2_struct_0(X1)))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t24_polynom7', cc2_struct_0)).
fof(d6_vectsp_1, axiom, ![X1]:((~(v2_struct_0(X1))&l4_algstr_0(X1))=>(v4_vectsp_1(X1)<=>![X2]:(m1_subset_1(X2,u1_struct_0(X1))=>(k6_algstr_0(X1,X2,k5_struct_0(X1))=X2&k6_algstr_0(X1,k5_struct_0(X1),X2)=X2)))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t24_polynom7', d6_vectsp_1)).
fof(dt_l5_algstr_0, axiom, ![X1]:(l5_algstr_0(X1)=>(l4_algstr_0(X1)&l4_struct_0(X1))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t24_polynom7', dt_l5_algstr_0)).
fof(t24_polynom7, conjecture, ![X1]:(v3_ordinal1(X1)=>![X2]:((~(v7_struct_0(X2))&(v13_algstr_0(X2)&(v3_rlvect_1(X2)&(v4_rlvect_1(X2)&(v4_vectsp_1(X2)&(v5_vectsp_1(X2)&l6_algstr_0(X2)))))))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))&(v4_polynom7(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,u1_struct_0(X2))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))))=>k5_polynom2(X1,X2,X3,X4)=k3_polynom7(X1,X2,X3))))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t24_polynom7', t24_polynom7)).
fof(dt_l2_struct_0, axiom, ![X1]:(l2_struct_0(X1)=>l1_struct_0(X1)), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t24_polynom7', dt_l2_struct_0)).
fof(cc3_polynom7, axiom, ![X1, X2]:((~(v2_struct_0(X2))&l2_struct_0(X2))=>![X3]:(m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))=>((v1_funct_1(X3)&(v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))&v4_polynom7(X3,X1,X2)))=>(v1_funct_1(X3)&(v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))&v3_polynom7(X3,X1,X2)))))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t24_polynom7', cc3_polynom7)).
fof(dt_k3_polynom7, axiom, ![X1, X2, X3]:(((~(v2_struct_0(X2))&l2_struct_0(X2))&(v1_funct_1(X3)&(v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))&(v3_polynom7(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))))))=>m1_subset_1(k3_polynom7(X1,X2,X3),u1_struct_0(X2))), file('/local1/mptp/loop2/hardpred_knn256_miz/preds__128/t24_polynom7', dt_k3_polynom7)).
fof(c_0_13, plain, ![X334, X335, X336, X337]:(~v3_ordinal1(X334)|(v7_struct_0(X335)|(~v13_algstr_0(X335)|(~v3_rlvect_1(X335)|(~v4_rlvect_1(X335)|(~v4_vectsp_1(X335)|(~v5_vectsp_1(X335)|~l6_algstr_0(X335))))))|(~v1_funct_1(X336)|(~v1_funct_2(X336,k15_pre_poly(X334),u1_struct_0(X335))|(~v3_polynom7(X336,X334,X335)|~m1_subset_1(X336,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X334),u1_struct_0(X335))))))|(~v1_funct_1(X337)|(~v1_funct_2(X337,X334,u1_struct_0(X335))|~m1_subset_1(X337,k1_zfmisc_1(k2_zfmisc_1(X334,u1_struct_0(X335)))))|k5_polynom2(X334,X335,X336,X337)=k6_algstr_0(X335,k3_polynom7(X334,X335,X336),k3_polynom2(X334,k2_polynom7(X334,X335,X336),X335,X337)))))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t12_polynom7])])])])).
fof(c_0_14, plain, ![X67, X68, X69]:((k2_polynom7(X67,X68,X69)=k16_pre_poly(X67)|(~v1_funct_1(X69)|(~v1_funct_2(X69,k15_pre_poly(X67),u1_struct_0(X68))|(~v4_polynom7(X69,X67,X68)|~m1_subset_1(X69,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X67),u1_struct_0(X68)))))))|(v2_struct_0(X68)|~l2_struct_0(X68)))&(k3_polynom7(X67,X68,X69)=k3_polynom1(X67,X68,X69,k16_pre_poly(X67))|(~v1_funct_1(X69)|(~v1_funct_2(X69,k15_pre_poly(X67),u1_struct_0(X68))|(~v4_polynom7(X69,X67,X68)|~m1_subset_1(X69,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X67),u1_struct_0(X68)))))))|(v2_struct_0(X68)|~l2_struct_0(X68)))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t16_polynom7])])])])])).
fof(c_0_15, plain, ![X34]:((l2_struct_0(X34)|~l2_algstr_0(X34))&(l1_algstr_0(X34)|~l2_algstr_0(X34))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l2_algstr_0])])])).
fof(c_0_16, plain, ![X48]:((l2_algstr_0(X48)|~l6_algstr_0(X48))&(l5_algstr_0(X48)|~l6_algstr_0(X48))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l6_algstr_0])])])).
cnf(c_0_17, plain, (v7_struct_0(X2)|k5_polynom2(X1,X2,X3,X4)=k6_algstr_0(X2,k3_polynom7(X1,X2,X3),k3_polynom2(X1,k2_polynom7(X1,X2,X3),X2,X4))|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v3_polynom7(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v1_funct_1(X4)|~v1_funct_2(X4,X1,u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_13])).
cnf(c_0_18, plain, (k2_polynom7(X1,X2,X3)=k16_pre_poly(X1)|v2_struct_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v4_polynom7(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~l2_struct_0(X2)), inference(split_conjunct,[status(thm)],[c_0_14])).
cnf(c_0_19, plain, (l2_struct_0(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
cnf(c_0_20, plain, (l2_algstr_0(X1)|~l6_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
fof(c_0_21, plain, ![X44]:(~l1_algstr_0(X44)|l1_struct_0(X44)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l1_algstr_0])])).
cnf(c_0_22, plain, (k6_algstr_0(X2,k3_polynom7(X1,X2,X3),k3_polynom2(X1,k2_polynom7(X1,X2,X3),X2,X4))=k5_polynom2(X1,X2,X3,X4)|v7_struct_0(X2)|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v3_polynom7(X3,X1,X2)|~v1_funct_2(X4,X1,u1_struct_0(X2))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_17])).
cnf(c_0_23, plain, (k2_polynom7(X1,X2,X3)=k16_pre_poly(X1)|v2_struct_0(X2)|~v1_funct_1(X3)|~l2_struct_0(X2)|~v4_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_18])).
cnf(c_0_24, plain, (l2_struct_0(X1)|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_19])).
cnf(c_0_25, plain, (l2_algstr_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_20])).
fof(c_0_26, plain, ![X73, X74, X75]:(~v3_ordinal1(X73)|(v7_struct_0(X74)|(~v4_vectsp_1(X74)|~l6_algstr_0(X74))|(~v1_funct_1(X75)|(~v1_funct_2(X75,X73,u1_struct_0(X74))|~m1_subset_1(X75,k1_zfmisc_1(k2_zfmisc_1(X73,u1_struct_0(X74)))))|k3_polynom2(X73,k16_pre_poly(X73),X74,X75)=k5_struct_0(X74)))), inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t14_polynom2])])])])).
fof(c_0_27, plain, ![X85]:(~l1_struct_0(X85)|(v7_struct_0(X85)|~v2_struct_0(X85))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc2_struct_0])])])).
cnf(c_0_28, plain, (l1_struct_0(X1)|~l1_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_21])).
cnf(c_0_29, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_15])).
fof(c_0_30, plain, ![X105, X106]:(((k6_algstr_0(X105,X106,k5_struct_0(X105))=X106|~m1_subset_1(X106,u1_struct_0(X105))|~v4_vectsp_1(X105)|(v2_struct_0(X105)|~l4_algstr_0(X105)))&(k6_algstr_0(X105,k5_struct_0(X105),X106)=X106|~m1_subset_1(X106,u1_struct_0(X105))|~v4_vectsp_1(X105)|(v2_struct_0(X105)|~l4_algstr_0(X105))))&((m1_subset_1(esk12_1(X105),u1_struct_0(X105))|v4_vectsp_1(X105)|(v2_struct_0(X105)|~l4_algstr_0(X105)))&(k6_algstr_0(X105,esk12_1(X105),k5_struct_0(X105))!=esk12_1(X105)|k6_algstr_0(X105,k5_struct_0(X105),esk12_1(X105))!=esk12_1(X105)|v4_vectsp_1(X105)|(v2_struct_0(X105)|~l4_algstr_0(X105))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d6_vectsp_1])])])])])])).
cnf(c_0_31, plain, (k6_algstr_0(X1,k3_polynom7(X2,X1,X3),k3_polynom2(X2,k16_pre_poly(X2),X1,X4))=k5_polynom2(X2,X1,X3,X4)|v7_struct_0(X1)|v2_struct_0(X1)|~v3_polynom7(X3,X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1))))|~v1_funct_2(X3,k15_pre_poly(X2),u1_struct_0(X1))|~v1_funct_2(X4,X2,u1_struct_0(X1))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~v3_ordinal1(X2)|~l2_struct_0(X1)|~v4_polynom7(X3,X2,X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_22, c_0_23])])).
cnf(c_0_32, plain, (l2_struct_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_24, c_0_25])])).
cnf(c_0_33, plain, (v7_struct_0(X2)|k3_polynom2(X1,k16_pre_poly(X1),X2,X3)=k5_struct_0(X2)|~v3_ordinal1(X1)|~v4_vectsp_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))), inference(split_conjunct,[status(thm)],[c_0_26])).
fof(c_0_34, plain, ![X96]:((l4_algstr_0(X96)|~l5_algstr_0(X96))&(l4_struct_0(X96)|~l5_algstr_0(X96))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l5_algstr_0])])])).
cnf(c_0_35, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_27])).
cnf(c_0_36, plain, (l1_struct_0(X1)|~l1_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_28])).
cnf(c_0_37, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_29])).
fof(c_0_38, negated_conjecture, ~(![X1]:(v3_ordinal1(X1)=>![X2]:((~(v7_struct_0(X2))&(v13_algstr_0(X2)&(v3_rlvect_1(X2)&(v4_rlvect_1(X2)&(v4_vectsp_1(X2)&(v5_vectsp_1(X2)&l6_algstr_0(X2)))))))=>![X3]:((v1_funct_1(X3)&(v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))&(v4_polynom7(X3,X1,X2)&m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))))=>![X4]:((v1_funct_1(X4)&(v1_funct_2(X4,X1,u1_struct_0(X2))&m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))))=>k5_polynom2(X1,X2,X3,X4)=k3_polynom7(X1,X2,X3)))))), inference(assume_negation,[status(cth)],[t24_polynom7])).
fof(c_0_39, plain, ![X62]:(~l2_struct_0(X62)|l1_struct_0(X62)), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[dt_l2_struct_0])])).
cnf(c_0_40, plain, (k6_algstr_0(X1,X2,k5_struct_0(X1))=X2|v2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_vectsp_1(X1)|~l4_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_30])).
cnf(c_0_41, plain, (k6_algstr_0(X1,k3_polynom7(X2,X1,X3),k3_polynom2(X2,k16_pre_poly(X2),X1,X4))=k5_polynom2(X2,X1,X3,X4)|v2_struct_0(X1)|v7_struct_0(X1)|~v3_polynom7(X3,X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1))))|~v4_polynom7(X3,X2,X1)|~v1_funct_2(X3,k15_pre_poly(X2),u1_struct_0(X1))|~v1_funct_2(X4,X2,u1_struct_0(X1))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~v3_ordinal1(X2)), inference(csr,[status(thm)],[c_0_31, c_0_32])).
cnf(c_0_42, plain, (k3_polynom2(X1,k16_pre_poly(X1),X2,X3)=k5_struct_0(X2)|v7_struct_0(X2)|~v3_ordinal1(X1)|~v4_vectsp_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))), inference(evalgc,[status(thm)],[c_0_33])).
cnf(c_0_43, plain, (l4_algstr_0(X1)|~l5_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_34])).
cnf(c_0_44, plain, (l5_algstr_0(X1)|~l6_algstr_0(X1)), inference(split_conjunct,[status(thm)],[c_0_16])).
cnf(c_0_45, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v2_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_35])).
cnf(c_0_46, plain, (l1_struct_0(X1)|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_37])])).
fof(c_0_47, plain, ![X78, X79, X80]:((v1_funct_1(X80)|(~v1_funct_1(X80)|(~v1_funct_2(X80,k15_pre_poly(X78),u1_struct_0(X79))|~v4_polynom7(X80,X78,X79)))|~m1_subset_1(X80,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X78),u1_struct_0(X79))))|(v2_struct_0(X79)|~l2_struct_0(X79)))&((v1_funct_2(X80,k15_pre_poly(X78),u1_struct_0(X79))|(~v1_funct_1(X80)|(~v1_funct_2(X80,k15_pre_poly(X78),u1_struct_0(X79))|~v4_polynom7(X80,X78,X79)))|~m1_subset_1(X80,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X78),u1_struct_0(X79))))|(v2_struct_0(X79)|~l2_struct_0(X79)))&(v3_polynom7(X80,X78,X79)|(~v1_funct_1(X80)|(~v1_funct_2(X80,k15_pre_poly(X78),u1_struct_0(X79))|~v4_polynom7(X80,X78,X79)))|~m1_subset_1(X80,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X78),u1_struct_0(X79))))|(v2_struct_0(X79)|~l2_struct_0(X79))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[cc3_polynom7])])])])])).
fof(c_0_48, negated_conjecture, (v3_ordinal1(esk1_0)&((~v7_struct_0(esk2_0)&(v13_algstr_0(esk2_0)&(v3_rlvect_1(esk2_0)&(v4_rlvect_1(esk2_0)&(v4_vectsp_1(esk2_0)&(v5_vectsp_1(esk2_0)&l6_algstr_0(esk2_0)))))))&((v1_funct_1(esk3_0)&(v1_funct_2(esk3_0,k15_pre_poly(esk1_0),u1_struct_0(esk2_0))&(v4_polynom7(esk3_0,esk1_0,esk2_0)&m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(esk1_0),u1_struct_0(esk2_0)))))))&((v1_funct_1(esk4_0)&(v1_funct_2(esk4_0,esk1_0,u1_struct_0(esk2_0))&m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,u1_struct_0(esk2_0))))))&k5_polynom2(esk1_0,esk2_0,esk3_0,esk4_0)!=k3_polynom7(esk1_0,esk2_0,esk3_0))))), inference(skolemize,[status(esa)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_38])])])])).
cnf(c_0_49, plain, (l1_struct_0(X1)|~l2_struct_0(X1)), inference(split_conjunct,[status(thm)],[c_0_39])).
cnf(c_0_50, plain, (k6_algstr_0(X1,X2,k5_struct_0(X1))=X2|v2_struct_0(X1)|~v4_vectsp_1(X1)|~l4_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))), inference(evalgc,[status(thm)],[c_0_40])).
cnf(c_0_51, plain, (k6_algstr_0(X1,k3_polynom7(X2,X1,X3),k5_struct_0(X1))=k5_polynom2(X2,X1,X3,X4)|v2_struct_0(X1)|v7_struct_0(X1)|~v3_polynom7(X3,X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1))))|~v4_polynom7(X3,X2,X1)|~v1_funct_2(X3,k15_pre_poly(X2),u1_struct_0(X1))|~v1_funct_2(X4,X2,u1_struct_0(X1))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~v3_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42])])).
cnf(c_0_52, plain, (l4_algstr_0(X1)|~l5_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_43])).
cnf(c_0_53, plain, (l5_algstr_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[c_0_44])).
cnf(c_0_54, plain, (v7_struct_0(X1)|~v2_struct_0(X1)|~l2_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_46])])).
fof(c_0_55, plain, ![X172, X173, X174]:(v2_struct_0(X173)|~l2_struct_0(X173)|(~v1_funct_1(X174)|(~v1_funct_2(X174,k15_pre_poly(X172),u1_struct_0(X173))|(~v3_polynom7(X174,X172,X173)|~m1_subset_1(X174,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X172),u1_struct_0(X173)))))))|m1_subset_1(k3_polynom7(X172,X173,X174),u1_struct_0(X173))), inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k3_polynom7])])])).
cnf(c_0_56, plain, (v3_polynom7(X1,X2,X3)|v2_struct_0(X3)|~v1_funct_1(X1)|~v1_funct_2(X1,k15_pre_poly(X2),u1_struct_0(X3))|~v4_polynom7(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3))))|~l2_struct_0(X3)), inference(split_conjunct,[status(thm)],[c_0_47])).
cnf(c_0_57, negated_conjecture, (v4_polynom7(esk3_0,esk1_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_58, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(esk1_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_59, negated_conjecture, (v1_funct_2(esk3_0,k15_pre_poly(esk1_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_60, negated_conjecture, (v1_funct_1(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_61, negated_conjecture, (l6_algstr_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_62, negated_conjecture, (~v7_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_63, plain, (l1_struct_0(X1)|~l2_struct_0(X1)), inference(evalgc,[status(thm)],[c_0_49])).
cnf(c_0_64, plain, (k5_polynom2(X2,X1,X3,X4)=k3_polynom7(X2,X1,X3)|v2_struct_0(X1)|v7_struct_0(X1)|~l4_algstr_0(X1)|~m1_subset_1(k3_polynom7(X2,X1,X3),u1_struct_0(X1))|~v4_vectsp_1(X1)|~v3_polynom7(X3,X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1))))|~v4_polynom7(X3,X2,X1)|~v1_funct_2(X3,k15_pre_poly(X2),u1_struct_0(X1))|~v1_funct_2(X4,X2,u1_struct_0(X1))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~v3_ordinal1(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_50, c_0_51])])).
cnf(c_0_65, plain, (l4_algstr_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_53])])).
cnf(c_0_66, plain, (v7_struct_0(X1)|~v2_struct_0(X1)|~l6_algstr_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_25])])).
cnf(c_0_67, plain, (v2_struct_0(X1)|m1_subset_1(k3_polynom7(X3,X1,X2),u1_struct_0(X1))|~l2_struct_0(X1)|~v1_funct_1(X2)|~v1_funct_2(X2,k15_pre_poly(X3),u1_struct_0(X1))|~v3_polynom7(X2,X3,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X3),u1_struct_0(X1))))), inference(split_conjunct,[status(thm)],[c_0_55])).
cnf(c_0_68, plain, (v2_struct_0(X3)|v3_polynom7(X1,X2,X3)|~v1_funct_1(X1)|~l2_struct_0(X3)|~v4_polynom7(X1,X2,X3)|~v1_funct_2(X1,k15_pre_poly(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3))))), inference(evalgc,[status(thm)],[c_0_56])).
cnf(c_0_69, negated_conjecture, (v4_polynom7(esk3_0,esk1_0,esk2_0)), inference(evalgc,[status(thm)],[c_0_57])).
cnf(c_0_70, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(esk1_0),u1_struct_0(esk2_0))))), inference(evalgc,[status(thm)],[c_0_58])).
cnf(c_0_71, negated_conjecture, (v1_funct_2(esk3_0,k15_pre_poly(esk1_0),u1_struct_0(esk2_0))), inference(evalgc,[status(thm)],[c_0_59])).
cnf(c_0_72, negated_conjecture, (v1_funct_1(esk3_0)), inference(evalgc,[status(thm)],[c_0_60])).
cnf(c_0_73, negated_conjecture, (l6_algstr_0(esk2_0)), inference(evalgc,[status(thm)],[c_0_61])).
cnf(c_0_74, negated_conjecture, (~v7_struct_0(esk2_0)), inference(evalgc,[status(thm)],[c_0_62])).
cnf(c_0_75, plain, (v7_struct_0(X1)|~v2_struct_0(X1)|~l2_struct_0(X1)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_45, c_0_63])])).
cnf(c_0_76, plain, (k5_polynom2(X1,X2,X3,X4)=k3_polynom7(X1,X2,X3)|v7_struct_0(X2)|~v3_polynom7(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~m1_subset_1(k3_polynom7(X1,X2,X3),u1_struct_0(X2))|~v4_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_2(X4,X1,u1_struct_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v13_algstr_0(X2)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_64, c_0_65]), c_0_66])).
cnf(c_0_77, plain, (v2_struct_0(X1)|m1_subset_1(k3_polynom7(X3,X1,X2),u1_struct_0(X1))|~v1_funct_1(X2)|~l2_struct_0(X1)|~v3_polynom7(X2,X3,X1)|~v1_funct_2(X2,k15_pre_poly(X3),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X3),u1_struct_0(X1))))), inference(evalgc,[status(thm)],[c_0_67])).
cnf(c_0_78, negated_conjecture, (v3_polynom7(esk3_0,esk1_0,esk2_0)|v2_struct_0(esk2_0)|~l2_struct_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_68, c_0_69]), c_0_70]), c_0_71]), c_0_72])])])).
cnf(c_0_79, negated_conjecture, (l2_struct_0(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_73])])).
cnf(c_0_80, negated_conjecture, (~v2_struct_0(esk2_0)|~l2_struct_0(esk2_0)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_74, c_0_75])])).
cnf(c_0_81, negated_conjecture, (k5_polynom2(esk1_0,esk2_0,esk3_0,esk4_0)!=k3_polynom7(esk1_0,esk2_0,esk3_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_82, plain, (k5_polynom2(X1,X2,X3,X4)=k3_polynom7(X1,X2,X3)|v7_struct_0(X2)|v2_struct_0(X2)|~v3_polynom7(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~v4_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_2(X4,X1,u1_struct_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v13_algstr_0(X2)|~v3_ordinal1(X1)|~l2_struct_0(X2)), inference(evalgc,[status(thm)],[inference(spm,[status(thm)],[c_0_76, c_0_77])])).
cnf(c_0_83, negated_conjecture, (v3_polynom7(esk3_0,esk1_0,esk2_0)|v2_struct_0(esk2_0)), inference(evalgc,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_78, c_0_79])])])).
cnf(c_0_84, negated_conjecture, (~v2_struct_0(esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_80, c_0_79])])).
cnf(c_0_85, negated_conjecture, (v5_vectsp_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_86, negated_conjecture, (v4_vectsp_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_87, negated_conjecture, (v4_rlvect_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_88, negated_conjecture, (v3_rlvect_1(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_89, negated_conjecture, (v13_algstr_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_90, negated_conjecture, (v3_ordinal1(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_91, negated_conjecture, (k3_polynom7(esk1_0,esk2_0,esk3_0)!=k5_polynom2(esk1_0,esk2_0,esk3_0,esk4_0)), inference(evalgc,[status(thm)],[c_0_81])).
cnf(c_0_92, plain, (k5_polynom2(X1,X2,X3,X4)=k3_polynom7(X1,X2,X3)|v7_struct_0(X2)|~v3_polynom7(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~v4_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_2(X4,X1,u1_struct_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v13_algstr_0(X2)|~v3_ordinal1(X1)), inference(csr,[status(thm)],[inference(csr,[status(thm)],[c_0_82, c_0_32]), c_0_66])).
cnf(c_0_93, negated_conjecture, (v3_polynom7(esk3_0,esk1_0,esk2_0)), inference(sr,[status(thm)],[c_0_83, c_0_84])).
cnf(c_0_94, negated_conjecture, (v5_vectsp_1(esk2_0)), inference(evalgc,[status(thm)],[c_0_85])).
cnf(c_0_95, negated_conjecture, (v4_vectsp_1(esk2_0)), inference(evalgc,[status(thm)],[c_0_86])).
cnf(c_0_96, negated_conjecture, (v4_rlvect_1(esk2_0)), inference(evalgc,[status(thm)],[c_0_87])).
cnf(c_0_97, negated_conjecture, (v3_rlvect_1(esk2_0)), inference(evalgc,[status(thm)],[c_0_88])).
cnf(c_0_98, negated_conjecture, (v13_algstr_0(esk2_0)), inference(evalgc,[status(thm)],[c_0_89])).
cnf(c_0_99, negated_conjecture, (v3_ordinal1(esk1_0)), inference(evalgc,[status(thm)],[c_0_90])).
cnf(c_0_100, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_101, negated_conjecture, (v1_funct_2(esk4_0,esk1_0,u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_102, negated_conjecture, (v1_funct_1(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_48])).
cnf(c_0_103, negated_conjecture, (k5_polynom2(esk1_0,esk2_0,esk3_0,X1)!=k5_polynom2(esk1_0,esk2_0,esk3_0,esk4_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,u1_struct_0(esk2_0))))|~v1_funct_2(X1,esk1_0,u1_struct_0(esk2_0))|~v1_funct_1(X1)), inference(evalgc,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_92]), c_0_93]), c_0_70]), c_0_69]), c_0_71]), c_0_72]), c_0_73]), c_0_94]), c_0_95]), c_0_96]), c_0_97]), c_0_98]), c_0_99])]), c_0_74])])).
cnf(c_0_104, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,u1_struct_0(esk2_0))))), inference(evalgc,[status(thm)],[c_0_100])).
cnf(c_0_105, negated_conjecture, (v1_funct_2(esk4_0,esk1_0,u1_struct_0(esk2_0))), inference(evalgc,[status(thm)],[c_0_101])).
cnf(c_0_106, negated_conjecture, (v1_funct_1(esk4_0)), inference(evalgc,[status(thm)],[c_0_102])).
cnf(c_0_107, negated_conjecture, ($false), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(er,[status(thm)],[c_0_103]), c_0_104]), c_0_105]), c_0_106])]), ['proof']).
# SZS output end CNFRefutation
# Proof object total steps             : 108
# Proof object clause steps            : 81
# Proof object formula steps           : 27
# Proof object conjectures             : 43
# Proof object clause conjectures      : 40
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 30
# Proof object initial formulas used   : 13
# Proof object generating inferences   : 15
# Proof object simplifying inferences  : 32
# Training examples: 45 positive, 395 negative
# Training: Positive examples begin
cnf(c_0_99, negated_conjecture, (v3_ordinal1(esk1_0))).# trainpos
cnf(c_0_98, negated_conjecture, (v13_algstr_0(esk2_0))).# trainpos
cnf(c_0_97, negated_conjecture, (v3_rlvect_1(esk2_0))).# trainpos
cnf(c_0_96, negated_conjecture, (v4_rlvect_1(esk2_0))).# trainpos
cnf(c_0_95, negated_conjecture, (v4_vectsp_1(esk2_0))).# trainpos
cnf(c_0_94, negated_conjecture, (v5_vectsp_1(esk2_0))).# trainpos
cnf(c_0_73, negated_conjecture, (l6_algstr_0(esk2_0))).# trainpos
cnf(c_0_72, negated_conjecture, (v1_funct_1(esk3_0))).# trainpos
cnf(c_0_106, negated_conjecture, (v1_funct_1(esk4_0))).# trainpos
cnf(c_0_74, negated_conjecture, (~v7_struct_0(esk2_0))).# trainpos
cnf(c_0_91, negated_conjecture, (k3_polynom7(esk1_0,esk2_0,esk3_0)!=k5_polynom2(esk1_0,esk2_0,esk3_0,esk4_0))).# trainpos
cnf(c_0_69, negated_conjecture, (v4_polynom7(esk3_0,esk1_0,esk2_0))).# trainpos
cnf(c_0_24, plain, (l2_struct_0(X1)|~l2_algstr_0(X1))).# trainpos
cnf(c_0_105, negated_conjecture, (v1_funct_2(esk4_0,esk1_0,u1_struct_0(esk2_0)))).# trainpos
cnf(c_0_71, negated_conjecture, (v1_funct_2(esk3_0,k15_pre_poly(esk1_0),u1_struct_0(esk2_0)))).# trainpos
cnf(c_0_25, plain, (l2_algstr_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_53, plain, (l5_algstr_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_37, plain, (l1_algstr_0(X1)|~l2_algstr_0(X1))).# trainpos
cnf(c_0_63, plain, (l1_struct_0(X1)|~l2_struct_0(X1))).# trainpos
cnf(c_0_36, plain, (l1_struct_0(X1)|~l1_algstr_0(X1))).# trainpos
cnf(c_0_52, plain, (l4_algstr_0(X1)|~l5_algstr_0(X1))).# trainpos
cnf(c_0_104, negated_conjecture, (m1_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,u1_struct_0(esk2_0)))))).# trainpos
cnf(c_0_70, negated_conjecture, (m1_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(esk1_0),u1_struct_0(esk2_0)))))).# trainpos
cnf(c_0_45, plain, (v7_struct_0(X1)|~l1_struct_0(X1)|~v2_struct_0(X1))).# trainpos
cnf(c_0_68, plain, (v2_struct_0(X3)|v3_polynom7(X1,X2,X3)|~v1_funct_1(X1)|~l2_struct_0(X3)|~v4_polynom7(X1,X2,X3)|~v1_funct_2(X1,k15_pre_poly(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3)))))).# trainpos
cnf(c_0_50, plain, (k6_algstr_0(X1,X2,k5_struct_0(X1))=X2|v2_struct_0(X1)|~v4_vectsp_1(X1)|~l4_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).# trainpos
cnf(c_0_77, plain, (v2_struct_0(X1)|m1_subset_1(k3_polynom7(X3,X1,X2),u1_struct_0(X1))|~v1_funct_1(X2)|~l2_struct_0(X1)|~v3_polynom7(X2,X3,X1)|~v1_funct_2(X2,k15_pre_poly(X3),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X3),u1_struct_0(X1)))))).# trainpos
cnf(c_0_23, plain, (k2_polynom7(X1,X2,X3)=k16_pre_poly(X1)|v2_struct_0(X2)|~v1_funct_1(X3)|~l2_struct_0(X2)|~v4_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).# trainpos
cnf(c_0_42, plain, (k3_polynom2(X1,k16_pre_poly(X1),X2,X3)=k5_struct_0(X2)|v7_struct_0(X2)|~v3_ordinal1(X1)|~v4_vectsp_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2)))))).# trainpos
cnf(c_0_22, plain, (k6_algstr_0(X2,k3_polynom7(X1,X2,X3),k3_polynom2(X1,k2_polynom7(X1,X2,X3),X2,X4))=k5_polynom2(X1,X2,X3,X4)|v7_struct_0(X2)|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v3_polynom7(X3,X1,X2)|~v1_funct_2(X4,X1,u1_struct_0(X2))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).# trainpos
cnf(c_0_32, plain, (l2_struct_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_75, plain, (v7_struct_0(X1)|~v2_struct_0(X1)|~l2_struct_0(X1))).# trainpos
cnf(c_0_46, plain, (l1_struct_0(X1)|~l2_algstr_0(X1))).# trainpos
cnf(c_0_65, plain, (l4_algstr_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_31, plain, (k6_algstr_0(X1,k3_polynom7(X2,X1,X3),k3_polynom2(X2,k16_pre_poly(X2),X1,X4))=k5_polynom2(X2,X1,X3,X4)|v7_struct_0(X1)|v2_struct_0(X1)|~v3_polynom7(X3,X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1))))|~v1_funct_2(X3,k15_pre_poly(X2),u1_struct_0(X1))|~v1_funct_2(X4,X2,u1_struct_0(X1))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~v3_ordinal1(X2)|~l2_struct_0(X1)|~v4_polynom7(X3,X2,X1))).# trainpos
cnf(c_0_78, negated_conjecture, (v3_polynom7(esk3_0,esk1_0,esk2_0)|v2_struct_0(esk2_0)|~l2_struct_0(esk2_0))).# trainpos
cnf(c_0_79, negated_conjecture, (l2_struct_0(esk2_0))).# trainpos
cnf(c_0_54, plain, (v7_struct_0(X1)|~v2_struct_0(X1)|~l2_algstr_0(X1))).# trainpos
cnf(c_0_51, plain, (k6_algstr_0(X1,k3_polynom7(X2,X1,X3),k5_struct_0(X1))=k5_polynom2(X2,X1,X3,X4)|v2_struct_0(X1)|v7_struct_0(X1)|~v3_polynom7(X3,X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1))))|~v4_polynom7(X3,X2,X1)|~v1_funct_2(X3,k15_pre_poly(X2),u1_struct_0(X1))|~v1_funct_2(X4,X2,u1_struct_0(X1))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~v3_ordinal1(X2))).# trainpos
cnf(c_0_80, negated_conjecture, (~v2_struct_0(esk2_0)|~l2_struct_0(esk2_0))).# trainpos
cnf(c_0_83, negated_conjecture, (v3_polynom7(esk3_0,esk1_0,esk2_0)|v2_struct_0(esk2_0))).# trainpos
cnf(c_0_66, plain, (v7_struct_0(X1)|~v2_struct_0(X1)|~l6_algstr_0(X1))).# trainpos
cnf(c_0_64, plain, (k5_polynom2(X2,X1,X3,X4)=k3_polynom7(X2,X1,X3)|v2_struct_0(X1)|v7_struct_0(X1)|~l4_algstr_0(X1)|~m1_subset_1(k3_polynom7(X2,X1,X3),u1_struct_0(X1))|~v4_vectsp_1(X1)|~v3_polynom7(X3,X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1))))|~v4_polynom7(X3,X2,X1)|~v1_funct_2(X3,k15_pre_poly(X2),u1_struct_0(X1))|~v1_funct_2(X4,X2,u1_struct_0(X1))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~v3_ordinal1(X2))).# trainpos
cnf(c_0_82, plain, (k5_polynom2(X1,X2,X3,X4)=k3_polynom7(X1,X2,X3)|v7_struct_0(X2)|v2_struct_0(X2)|~v3_polynom7(X3,X1,X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~v4_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_2(X4,X1,u1_struct_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v13_algstr_0(X2)|~v3_ordinal1(X1)|~l2_struct_0(X2))).# trainpos
cnf(c_0_103, negated_conjecture, (k5_polynom2(esk1_0,esk2_0,esk3_0,X1)!=k5_polynom2(esk1_0,esk2_0,esk3_0,esk4_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(esk1_0,u1_struct_0(esk2_0))))|~v1_funct_2(X1,esk1_0,u1_struct_0(esk2_0))|~v1_funct_1(X1))).# trainpos
# Training: Positive examples end
# Training: Negative examples begin
cnf(i_0_173, plain, (v1_xboole_0(k1_xboole_0))).#trainneg
cnf(i_0_77, plain, (k2_finseq_1(np__1)=k1_tarski(np__1))).#trainneg
cnf(i_0_76, plain, (k2_tarski(np__1,np__2)=k2_finseq_1(np__2))).#trainneg
cnf(i_0_20, plain, (X1=k1_xboole_0|X3!=k1_xboole_0|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_167, plain, (X1=k1_xboole_0|~v1_xboole_0(X1))).#trainneg
cnf(i_0_281, plain, (v1_card_1(k1_card_1(X1)))).#trainneg
cnf(i_0_126, plain, (l2_struct_0(X1)|~l4_struct_0(X1))).#trainneg
cnf(i_0_193, plain, (r1_tarski(X1,X1))).#trainneg
cnf(i_0_89, plain, (~v1_xboole_0(k15_pre_poly(X1)))).#trainneg
cnf(i_0_170, plain, (l1_struct_0(X1)|~l3_struct_0(X1))).#trainneg
cnf(i_0_174, plain, (l1_struct_0(X1)|~l3_algstr_0(X1))).#trainneg
cnf(i_0_91, plain, (l4_struct_0(X1)|~l5_algstr_0(X1))).#trainneg
cnf(i_0_152, plain, (l3_struct_0(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_125, plain, (l3_struct_0(X1)|~l4_struct_0(X1))).#trainneg
cnf(i_0_151, plain, (l3_algstr_0(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_287, plain, (v1_finset_1(k1_card_1(X1))|~v1_finset_1(X1))).#trainneg
cnf(i_0_131, plain, (v4_rlvect_1(X1)|k1_algstr_0(X1,esk16_1(X1),k4_struct_0(X1))!=esk16_1(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_150, plain, (k1_finseq_1(X1)=a_1_0_finseq_1(X1)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_84, plain, (r1_tarski(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_252, plain, (u2_struct_0(X1)=k4_struct_0(X1)|~l2_struct_0(X1))).#trainneg
cnf(i_0_35, plain, (v7_ordinal1(esk6_1(X1))|~v1_relat_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_206, plain, (m1_subset_1(k15_pre_poly(X1),k1_zfmisc_1(k15_pre_poly(X1))))).#trainneg
cnf(i_0_239, plain, (m1_subset_1(X1,X2)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_175, plain, (m2_subset_1(k16_pre_poly(X1),k15_pre_poly(X1),k15_pre_poly(X1)))).#trainneg
cnf(i_0_40, plain, (r2_hidden(X1,X3)|X1!=X2|X3!=k1_tarski(X2))).#trainneg
cnf(i_0_41, plain, (X1=X3|X2!=k1_tarski(X3)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_83, plain, (m1_subset_1(X1,k1_zfmisc_1(X2))|~r1_tarski(X1,X2))).#trainneg
cnf(i_0_172, plain, (v4_funct_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k15_pre_poly(X2))))).#trainneg
cnf(i_0_285, plain, (v2_struct_0(X1)|~l1_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1)))).#trainneg
cnf(i_0_244, plain, (v1_funct_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_245, plain, (v1_relat_1(X1)|~v4_funct_1(X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_119, plain, (v1_funct_1(X1)|X2!=k15_pre_poly(X3)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_121, plain, (v1_relat_1(X1)|X2!=k15_pre_poly(X3)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_19, plain, (v1_funct_2(X1,X2,X3)|X3!=k1_xboole_0|X1!=k1_xboole_0|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_117, plain, (v4_valued_0(X1)|X2!=k15_pre_poly(X3)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_34, plain, (k2_finseq_1(esk6_1(X1))=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_31, plain, (k3_finseq_1(X1)=np__1|X1!=k9_finseq_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_166, plain, (k4_finseq_1(X1)=k9_xtuple_0(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_269, plain, (r1_relat_2(X1,X2)|~v1_relat_1(X1)|~r3_orders_1(X1,X2))).#trainneg
cnf(i_0_263, plain, (v1_xboole_0(X2)|~v1_xboole_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_264, plain, (v1_relat_1(X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_268, plain, (r8_relat_2(X1,X2)|~v1_relat_1(X1)|~r3_orders_1(X1,X2))).#trainneg
cnf(i_0_267, plain, (r4_relat_2(X1,X2)|~v1_relat_1(X1)|~r3_orders_1(X1,X2))).#trainneg
cnf(i_0_227, plain, (k1_card_1(X1)=k3_finseq_1(X1)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_33, plain, (v1_finseq_1(X2)|k9_xtuple_0(X2)!=k2_finseq_1(X1)|~v1_relat_1(X2)|~v7_ordinal1(X1))).#trainneg
cnf(i_0_116, plain, (v2_pre_poly(X1)|X2!=k15_pre_poly(X3)|~r2_hidden(X1,X2))).#trainneg
cnf(i_0_231, plain, (v1_xboole_0(X2)|r2_hidden(X1,X2)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_205, plain, (v2_struct_0(X2)|v1_funct_1(k7_polynom1(X1,X2))|~l2_struct_0(X2))).#trainneg
cnf(i_0_266, plain, (r6_relat_2(X1,X2)|~v1_relat_1(X1)|~r3_orders_1(X1,X2))).#trainneg
cnf(i_0_118, plain, (v1_partfun1(X1,X2)|X3!=k15_pre_poly(X2)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_120, plain, (v4_relat_1(X1,X2)|X3!=k15_pre_poly(X2)|~r2_hidden(X1,X3))).#trainneg
cnf(i_0_261, plain, (v5_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X2))))).#trainneg
cnf(i_0_262, plain, (v4_relat_1(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_132, plain, (v4_rlvect_1(X1)|m1_subset_1(esk16_1(X1),u1_struct_0(X1))|~l2_algstr_0(X1))).#trainneg
cnf(i_0_192, plain, (k12_finseq_1(X1,X2)=k9_finseq_1(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_188, plain, (k3_pre_poly(X1,X2)=k9_finseq_1(X2)|v1_xboole_0(X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_148, plain, (m1_subset_1(X1,X3)|~m1_finseq_2(X3,X2)|~m2_finseq_2(X1,X2,X3))).#trainneg
cnf(i_0_250, plain, (k1_relat_1(X1)=X2|~r1_relat_2(X1,X2)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X2))))).#trainneg
cnf(i_0_168, plain, (k20_pre_poly(X1,k16_pre_poly(X1))=k3_pre_poly(k15_pre_poly(X1),k16_pre_poly(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_142, plain, (v1_partfun1(X2,X1)|k1_relset_1(X1,X2)!=X1|~v1_relat_1(X2)|~v4_relat_1(X2,X1))).#trainneg
cnf(i_0_96, plain, (v4_vectsp_1(X1)|v2_struct_0(X1)|m1_subset_1(esk12_1(X1),u1_struct_0(X1))|~l4_algstr_0(X1))).#trainneg
cnf(i_0_147, plain, (m2_finseq_2(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_finseq_2(X2,X3))).#trainneg
cnf(i_0_78, plain, (r2_hidden(np__1,k1_relset_1(k5_numbers,X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(X1)|~r2_hidden(X2,k10_xtuple_0(X1)))).#trainneg
cnf(i_0_282, plain, (v1_funct_1(k3_relat_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_248, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m2_subset_1(X1,X3,X2)|~m1_subset_1(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3)))).#trainneg
cnf(i_0_265, plain, (r3_orders_1(X1,X2)|~v1_relat_1(X1)|~r1_relat_2(X1,X2)|~r8_relat_2(X1,X2)|~r4_relat_2(X1,X2)|~r6_relat_2(X1,X2))).#trainneg
cnf(i_0_30, plain, (k1_funct_1(X1,np__1)=X2|X1!=k9_finseq_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_164, plain, (k1_relset_1(X2,X1)=k9_xtuple_0(X1)|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_141, plain, (k2_relset_1(X2,X1)=k10_xtuple_0(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).#trainneg
cnf(i_0_107, plain, (k4_rlvect_1(X1,k6_finseq_1(u1_struct_0(X1)))=k4_struct_0(X1)|v2_struct_0(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_38, plain, (X2=k1_tarski(X1)|esk7_2(X1,X2)=X1|r2_hidden(esk7_2(X1,X2),X2))).#trainneg
cnf(i_0_222, plain, (v2_struct_0(X2)|v4_polynom7(k7_polynom1(X1,X2),X1,X2)|~l2_struct_0(X2))).#trainneg
cnf(i_0_39, plain, (X2=k1_tarski(X1)|esk7_2(X1,X2)!=X1|~r2_hidden(esk7_2(X1,X2),X2))).#trainneg
cnf(i_0_111, plain, (X2=k15_pre_poly(X1)|v1_funct_1(esk15_2(X1,X2))|r2_hidden(esk15_2(X1,X2),X2))).#trainneg
cnf(i_0_113, plain, (X2=k15_pre_poly(X1)|v1_relat_1(esk15_2(X1,X2))|r2_hidden(esk15_2(X1,X2),X2))).#trainneg
cnf(i_0_109, plain, (X2=k15_pre_poly(X1)|v4_valued_0(esk15_2(X1,X2))|r2_hidden(esk15_2(X1,X2),X2))).#trainneg
cnf(i_0_108, plain, (X2=k15_pre_poly(X1)|v2_pre_poly(esk15_2(X1,X2))|r2_hidden(esk15_2(X1,X2),X2))).#trainneg
cnf(i_0_253, plain, (m1_subset_1(k1_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_260, plain, (m1_subset_1(k2_relset_1(X2,X1),k1_zfmisc_1(X2))|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).#trainneg
cnf(i_0_44, plain, (X1=X2|r2_hidden(esk8_2(X1,X2),X2)|r2_hidden(esk8_2(X1,X2),X1))).#trainneg
cnf(i_0_273, plain, (v7_struct_0(X2)|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v3_vectsp_1(X2)|~v2_struct_0(k11_polynom1(X1,X2)))).#trainneg
cnf(i_0_246, plain, (m1_subset_1(k7_partfun1(X2,X1,X3),X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2))).#trainneg
cnf(i_0_283, plain, (v1_relat_1(k3_relat_1(X1,X2))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_17, plain, (X1=k3_finseq_1(X2)|k9_xtuple_0(X2)!=k2_finseq_1(X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_finseq_1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_29, plain, (X1=k9_finseq_1(X2)|k3_finseq_1(X1)!=np__1|k1_funct_1(X1,np__1)!=X2|~v1_funct_1(X1)|~v1_relat_1(X1)|~v1_finseq_1(X1))).#trainneg
cnf(i_0_90, plain, (X1=X2|~v1_xxreal_0(X2)|~v1_xxreal_0(X1)|~r1_xxreal_0(X2,X1)|~r1_xxreal_0(X1,X2))).#trainneg
cnf(i_0_45, plain, (X1=X2|~r2_hidden(esk8_2(X1,X2),X2)|~r2_hidden(esk8_2(X1,X2),X1))).#trainneg
cnf(i_0_18, plain, (k9_xtuple_0(X2)=k2_finseq_1(X1)|X1!=k3_finseq_1(X2)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v1_finseq_1(X2)|~m2_subset_1(X1,k1_numbers,k5_numbers))).#trainneg
cnf(i_0_249, plain, (v1_xboole_0(X3)|v1_xboole_0(X2)|m1_subset_1(X1,X3)|~m2_subset_1(X1,X2,X3)|~m1_subset_1(X3,k1_zfmisc_1(X2)))).#trainneg
cnf(i_0_189, plain, (v2_struct_0(X3)|v1_polynom1(X1,k15_pre_poly(X2),X3)|~v1_funct_1(X1)|~l2_struct_0(X3)|~v3_polynom7(X1,X2,X3)|~v1_funct_2(X1,k15_pre_poly(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3)))))).#trainneg
cnf(i_0_143, plain, (k1_relset_1(X2,X1)=X2|~v1_relat_1(X1)|~v1_partfun1(X1,X2)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_271, plain, (v7_struct_0(X2)|l6_algstr_0(k11_polynom1(X1,X2))|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v3_vectsp_1(X2))).#trainneg
cnf(i_0_82, plain, (k4_polynom2(X1,X2)=X2|~v3_ordinal1(X1)|~m2_subset_1(X2,k15_pre_poly(X1),k15_pre_poly(X1)))).#trainneg
cnf(i_0_110, plain, (X2=k15_pre_poly(X1)|r2_hidden(esk15_2(X1,X2),X2)|v1_partfun1(esk15_2(X1,X2),X1))).#trainneg
cnf(i_0_112, plain, (X2=k15_pre_poly(X1)|r2_hidden(esk15_2(X1,X2),X2)|v4_relat_1(esk15_2(X1,X2),X1))).#trainneg
cnf(i_0_204, plain, (v2_struct_0(X2)|v1_funct_2(k7_polynom1(X1,X2),k15_pre_poly(X1),u1_struct_0(X2))|~l2_struct_0(X2))).#trainneg
cnf(i_0_251, plain, (k1_relset_1(X1,X2)=X1|~r1_relat_2(X2,X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_257, plain, (r2_hidden(esk21_3(X1,X2,X3),X3)|~v3_ordinal1(X2)|~r2_hidden(X1,a_2_2_bagorder(X2,X3)))).#trainneg
cnf(i_0_258, plain, (k3_bagorder(X2,esk21_3(X1,X2,X3))=X1|~v3_ordinal1(X2)|~r2_hidden(X1,a_2_2_bagorder(X2,X3)))).#trainneg
cnf(i_0_133, plain, (k1_algstr_0(X1,X2,k4_struct_0(X1))=X2|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_203, plain, (v2_struct_0(X2)|m1_subset_1(k7_polynom1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~l2_struct_0(X2))).#trainneg
cnf(i_0_42, plain, (r2_hidden(k1_funct_1(X1,X2),k10_xtuple_0(X1))|~v1_funct_1(X1)|~v1_relat_1(X1)|~r2_hidden(X2,k9_xtuple_0(X1)))).#trainneg
cnf(i_0_21, plain, (X3=k1_xboole_0|v1_funct_2(X2,X1,X3)|k1_relset_1(X1,X2)!=X1|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_184, plain, (v2_struct_0(X3)|X1!=k7_polynom1(X2,X3)|~v1_funct_1(X1)|~l2_struct_0(X3)|~v1_polynom7(X1,X2,X3)|~v1_funct_2(X1,k15_pre_poly(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3)))))).#trainneg
cnf(i_0_256, plain, (r2_hidden(X3,a_2_2_bagorder(X2,X4))|X3!=k3_bagorder(X2,X1)|~v3_ordinal1(X2)|~r2_hidden(X1,X4)|~m2_subset_1(X1,k15_pre_poly(X2),k15_pre_poly(X2)))).#trainneg
cnf(i_0_272, plain, (v7_struct_0(X2)|v36_algstr_0(k11_polynom1(X1,X2))|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v3_vectsp_1(X2))).#trainneg
cnf(i_0_129, plain, (X3=X4|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_2(X4,X1,X2)|~v1_funct_2(X3,X1,X2)|~r2_funct_2(X1,X2,X3,X4)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2))))).#trainneg
cnf(i_0_115, plain, (r2_hidden(X1,X3)|X3!=k15_pre_poly(X2)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v4_valued_0(X1)|~v2_pre_poly(X1)|~v1_partfun1(X1,X2)|~v4_relat_1(X1,X2))).#trainneg
cnf(i_0_69, plain, (r2_hidden(X1,k9_xtuple_0(X2))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,X3))))).#trainneg
cnf(i_0_95, plain, (v4_vectsp_1(X1)|v2_struct_0(X1)|k6_algstr_0(X1,k5_struct_0(X1),esk12_1(X1))!=esk12_1(X1)|k6_algstr_0(X1,esk12_1(X1),k5_struct_0(X1))!=esk12_1(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_68, plain, (r2_hidden(k1_funct_1(X1,X2),k9_xtuple_0(X3))|~v1_funct_1(X3)|~v1_funct_1(X1)|~v1_relat_1(X3)|~v1_relat_1(X1)|~r2_hidden(X2,k9_xtuple_0(k3_relat_1(X1,X3))))).#trainneg
cnf(i_0_128, plain, (r2_funct_2(X3,X4,X1,X2)|X1!=X2|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_funct_2(X2,X3,X4)|~v1_funct_2(X1,X3,X4)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X3,X4)))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X3,X4))))).#trainneg
cnf(i_0_97, plain, (k6_algstr_0(X1,k5_struct_0(X1),X2)=X2|v2_struct_0(X1)|~v4_vectsp_1(X1)|~l4_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_270, plain, (k3_pre_poly(k3_finseq_2(k15_pre_poly(X1)),k4_pre_poly(k15_pre_poly(X1),k16_pre_poly(X1),k16_pre_poly(X1)))=k21_pre_poly(X1,k16_pre_poly(X1))|~v3_ordinal1(X1))).#trainneg
cnf(i_0_22, plain, (X3=k1_xboole_0|k1_relset_1(X2,X1)=X2|~v1_funct_2(X1,X2,X3)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(X2,X3))))).#trainneg
cnf(i_0_134, plain, (k3_polynom7(X1,X2,k4_polynom7(X1,X2,X3))=X3|v2_struct_0(X2)|~l2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_259, plain, (m2_subset_1(esk21_3(X1,X2,X3),k15_pre_poly(X2),k15_pre_poly(X2))|~v3_ordinal1(X2)|~r2_hidden(X1,a_2_2_bagorder(X2,X3)))).#trainneg
cnf(i_0_135, plain, (k2_polynom7(X1,X2,k4_polynom7(X1,X2,X3))=k16_pre_poly(X1)|v2_struct_0(X2)|~l2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_32, plain, (k7_partfun1(X2,X1,X3)=k1_funct_1(X1,X3)|~v1_funct_1(X1)|~v1_relat_1(X1)|~v5_relat_1(X1,X2)|~r2_hidden(X3,k9_xtuple_0(X1)))).#trainneg
cnf(i_0_280, plain, (k15_funct_7(k15_pre_poly(X3),u1_struct_0(X1),k7_polynom1(X3,X1),k16_pre_poly(X3),X2)=k4_polynom7(X3,X1,X2)|v2_struct_0(X1)|~l2_struct_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_146, plain, (k4_rlvect_1(X1,k12_finseq_1(u1_struct_0(X1),X2))=X2|v2_struct_0(X1)|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_23, plain, (k1_funct_1(k3_relat_1(X1,X2),X3)=k1_funct_1(X2,k1_funct_1(X1,X3))|~v1_funct_1(X2)|~v1_funct_1(X1)|~v1_relat_1(X2)|~v1_relat_1(X1)|~r2_hidden(X3,k9_xtuple_0(X1)))).#trainneg
cnf(i_0_93, plain, (k1_polynom6(X1,X2,X3,X4)=X4|X1!=k1_xboole_0|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~m2_subset_1(X4,k15_pre_poly(X2),k15_pre_poly(X2))|~m2_subset_1(X3,k15_pre_poly(X1),k15_pre_poly(X1)))).#trainneg
cnf(i_0_94, plain, (k1_polynom6(X1,X2,X3,X4)=X3|X2!=k1_xboole_0|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~m2_subset_1(X4,k15_pre_poly(X2),k15_pre_poly(X2))|~m2_subset_1(X3,k15_pre_poly(X1),k15_pre_poly(X1)))).#trainneg
cnf(i_0_275, plain, (m2_finseq_1(k20_pre_poly(X1,X2),k15_pre_poly(X1))|~v3_ordinal1(X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v4_valued_0(X2)|~v2_pre_poly(X2)|~v1_partfun1(X2,X1)|~v4_relat_1(X2,X1))).#trainneg
cnf(i_0_199, plain, (v2_struct_0(X3)|v4_polynom7(X1,X2,X3)|~v1_funct_1(X1)|~l2_struct_0(X3)|~v1_funct_2(X1,k15_pre_poly(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3))))|~r2_funct_2(k15_pre_poly(X2),u1_struct_0(X3),X1,k7_polynom1(X2,X3)))).#trainneg
cnf(i_0_195, plain, (v2_struct_0(X3)|v4_polynom7(X1,X2,X3)|k2_polynom1(k15_pre_poly(X2),X3,X1)!=k6_domain_1(k15_pre_poly(X2),k16_pre_poly(X2))|~v1_funct_1(X1)|~l2_struct_0(X3)|~v1_funct_2(X1,k15_pre_poly(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3)))))).#trainneg
cnf(i_0_67, plain, (r2_hidden(X1,k9_xtuple_0(k3_relat_1(X2,X3)))|~v1_funct_1(X3)|~v1_funct_1(X2)|~v1_relat_1(X3)|~v1_relat_1(X2)|~r2_hidden(X1,k9_xtuple_0(X2))|~r2_hidden(k1_funct_1(X2,X1),k9_xtuple_0(X3)))).#trainneg
cnf(i_0_242, plain, (X1=k16_pre_poly(X2)|v2_struct_0(X3)|X1!=k2_polynom7(X2,X3,X4)|k3_polynom1(X2,X3,X4,X1)!=k4_struct_0(X3)|~v1_funct_1(X4)|~v1_funct_1(X1)|~v1_relat_1(X1)|~l2_struct_0(X3)|~v4_valued_0(X1)|~v2_pre_poly(X1)|~v1_partfun1(X1,X2)|~v4_relat_1(X1,X2)|~v3_polynom7(X4,X2,X3)|~v1_funct_2(X4,k15_pre_poly(X2),u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3)))))).#trainneg
cnf(i_0_165, plain, (k7_pre_poly(X2,X1,X3)=k1_xboole_0|~v1_xboole_0(X1)|~v1_relat_2(X3)|~v4_relat_2(X3)|~v8_relat_2(X3)|~v1_partfun1(X3,X2)|~r3_orders_1(X3,X1)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X2))))).#trainneg
cnf(i_0_122, plain, (k3_funct_2(X1,X3,X2,X4)=k1_funct_1(X2,X4)|v1_xboole_0(X1)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_279, plain, (X3=X4|v2_struct_0(X2)|~l2_struct_0(X2)|~m1_subset_1(X4,u1_struct_0(X2))|~m1_subset_1(X3,u1_struct_0(X2))|~r2_funct_2(k15_pre_poly(X1),u1_struct_0(X2),k4_polynom7(X1,X2,X3),k4_polynom7(X1,X2,X4)))).#trainneg
cnf(i_0_114, plain, (X2=k15_pre_poly(X1)|~v1_funct_1(esk15_2(X1,X2))|~v1_relat_1(esk15_2(X1,X2))|~v4_valued_0(esk15_2(X1,X2))|~v2_pre_poly(esk15_2(X1,X2))|~r2_hidden(esk15_2(X1,X2),X2)|~v1_partfun1(esk15_2(X1,X2),X1)|~v4_relat_1(esk15_2(X1,X2),X1))).#trainneg
cnf(i_0_183, plain, (X1=k7_polynom1(X2,X3)|v2_struct_0(X3)|v1_polynom7(X1,X2,X3)|~v1_funct_1(X1)|~l2_struct_0(X3)|~v1_funct_2(X1,k15_pre_poly(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3)))))).#trainneg
cnf(i_0_284, plain, (m2_finseq_1(k7_pre_poly(X2,X1,X3),X2)|~v1_finset_1(X1)|~v1_relat_2(X3)|~v4_relat_2(X3)|~v8_relat_2(X3)|~v1_partfun1(X3,X2)|~m1_subset_1(X1,k1_zfmisc_1(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X2))))).#trainneg
cnf(i_0_56, plain, (k2_relset_1(X1,X2)=X3|X2!=k7_pre_poly(X1,X3,X4)|~v1_finset_1(X3)|~v1_relat_2(X4)|~v4_relat_2(X4)|~v8_relat_2(X4)|~v1_partfun1(X4,X1)|~r3_orders_1(X4,X3)|~m2_finseq_1(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_144, plain, (v1_xboole_0(X1)|m1_subset_1(k3_funct_2(X1,X3,X2,X4),X3)|~v1_funct_1(X2)|~m1_subset_1(X4,X1)|~v1_funct_2(X2,X1,X3)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(X1,X3))))).#trainneg
cnf(i_0_226, plain, (k15_funct_7(k15_pre_poly(X2),u1_struct_0(X3),k7_polynom1(X2,X3),k2_hilbasis(X2,X1),k1_group_1(X3))=k3_hilbasis(X2,X1,X3)|v2_struct_0(X3)|~l5_algstr_0(X3)|~v1_group_1(X3)|~m1_subset_1(X1,X2))).#trainneg
cnf(i_0_154, plain, (v1_xboole_0(X1)|m1_subset_1(k2_polynom1(X1,X2,X3),k1_zfmisc_1(X1))|~v1_funct_1(X3)|~l2_struct_0(X2)|~v1_funct_2(X3,X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2)))))).#trainneg
cnf(i_0_149, plain, (m2_subset_1(k1_polynom6(X1,X2,X3,X4),k15_pre_poly(k10_ordinal2(X1,X2)),k15_pre_poly(k10_ordinal2(X1,X2)))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~m1_subset_1(X4,k15_pre_poly(X2))|~m1_subset_1(X3,k15_pre_poly(X1)))).#trainneg
cnf(i_0_278, plain, (v2_struct_0(X4)|r2_funct_2(k15_pre_poly(X3),u1_struct_0(X4),k4_polynom7(X3,X4,X1),k4_polynom7(X3,X4,X2))|X1!=X2|~l2_struct_0(X4)|~m1_subset_1(X2,u1_struct_0(X4))|~m1_subset_1(X1,u1_struct_0(X4)))).#trainneg
cnf(i_0_153, plain, (k3_polynom1(X3,X1,k7_polynom1(X3,X1),X2)=k4_struct_0(X1)|v2_struct_0(X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|~l2_struct_0(X1)|~v4_valued_0(X2)|~v2_pre_poly(X2)|~v1_partfun1(X2,X3)|~v4_relat_1(X2,X3))).#trainneg
cnf(i_0_58, plain, (v2_struct_0(X2)|v4_polynom7(X3,X1,X2)|esk11_3(X1,X2,X3)!=k16_pre_poly(X1)|~v1_funct_1(X3)|~l2_struct_0(X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_62, plain, (v2_struct_0(X2)|v4_polynom7(X3,X1,X2)|v1_funct_1(esk11_3(X1,X2,X3))|~v1_funct_1(X3)|~l2_struct_0(X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_64, plain, (v2_struct_0(X2)|v4_polynom7(X3,X1,X2)|v1_relat_1(esk11_3(X1,X2,X3))|~v1_funct_1(X3)|~l2_struct_0(X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_60, plain, (v2_struct_0(X2)|v4_polynom7(X3,X1,X2)|v4_valued_0(esk11_3(X1,X2,X3))|~v1_funct_1(X3)|~l2_struct_0(X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_178, plain, (v2_struct_0(X3)|v4_polynom7(X1,X2,X3)|~v1_funct_1(X1)|~l2_struct_0(X3)|~m1_subset_1(X4,u1_struct_0(X3))|~v1_funct_2(X1,k15_pre_poly(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3))))|~r2_funct_2(k15_pre_poly(X2),u1_struct_0(X3),X1,k4_polynom7(X2,X3,X4)))).#trainneg
cnf(i_0_161, plain, (k1_funct_1(X3,X1)=k1_funct_1(X4,X1)|X3!=k1_polynom6(X2,X5,X4,X6)|~v3_ordinal1(X5)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~r2_hidden(X1,X2)|~m2_subset_1(X6,k15_pre_poly(X5),k15_pre_poly(X5))|~m2_subset_1(X4,k15_pre_poly(X2),k15_pre_poly(X2))|~m2_subset_1(X3,k15_pre_poly(k10_ordinal2(X2,X5)),k15_pre_poly(k10_ordinal2(X2,X5))))).#trainneg
cnf(i_0_59, plain, (v2_struct_0(X2)|v4_polynom7(X3,X1,X2)|v2_pre_poly(esk11_3(X1,X2,X3))|~v1_funct_1(X3)|~l2_struct_0(X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_49, plain, (X4=k7_pre_poly(X1,X2,X3)|k2_relset_1(X1,X4)!=X2|~v1_finset_1(X2)|~v1_relat_2(X3)|~v4_relat_2(X3)|~v8_relat_2(X3)|~v1_partfun1(X3,X1)|~r3_orders_1(X3,X2)|~m2_finseq_1(X4,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~r1_xxreal_0(esk10_4(X1,X2,X3,X4),esk9_4(X1,X2,X3,X4)))).#trainneg
cnf(i_0_61, plain, (v2_struct_0(X2)|v4_polynom7(X3,X1,X2)|v1_partfun1(esk11_3(X1,X2,X3),X1)|~v1_funct_1(X3)|~l2_struct_0(X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_219, plain, (v2_struct_0(X2)|v1_funct_1(esk19_3(X1,X2,X3))|~v1_funct_1(X3)|~l2_struct_0(X2)|~v3_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_28, plain, (v1_xboole_0(X3)|X2!=k2_polynom1(X3,X4,X5)|k3_funct_2(X3,u1_struct_0(X4),X5,X1)!=k4_struct_0(X4)|~v1_funct_1(X5)|~l2_struct_0(X4)|~m1_subset_1(X1,X3)|~r2_hidden(X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(X3))|~v1_funct_2(X5,X3,u1_struct_0(X4))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(X4)))))).#trainneg
cnf(i_0_221, plain, (v2_struct_0(X2)|v1_relat_1(esk19_3(X1,X2,X3))|~v1_funct_1(X3)|~l2_struct_0(X2)|~v3_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_63, plain, (v2_struct_0(X2)|v4_polynom7(X3,X1,X2)|v4_relat_1(esk11_3(X1,X2,X3),X1)|~v1_funct_1(X3)|~l2_struct_0(X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_217, plain, (v2_struct_0(X2)|v4_valued_0(esk19_3(X1,X2,X3))|~v1_funct_1(X3)|~l2_struct_0(X2)|~v3_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_216, plain, (v2_struct_0(X2)|v2_pre_poly(esk19_3(X1,X2,X3))|~v1_funct_1(X3)|~l2_struct_0(X2)|~v3_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_218, plain, (v2_struct_0(X2)|v1_partfun1(esk19_3(X1,X2,X3),X1)|~v1_funct_1(X3)|~l2_struct_0(X2)|~v3_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_220, plain, (v2_struct_0(X2)|v4_relat_1(esk19_3(X1,X2,X3),X1)|~v1_funct_1(X3)|~l2_struct_0(X2)|~v3_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_124, plain, (k3_polynom7(X1,X2,k1_polynom7(X1,X2,k4_struct_0(X2),X3))=k4_struct_0(X2)|v2_struct_0(X2)|~v1_funct_1(X3)|~v1_relat_1(X3)|~l2_struct_0(X2)|~v4_valued_0(X3)|~v2_pre_poly(X3)|~v1_partfun1(X3,X1)|~v4_relat_1(X3,X1))).#trainneg
cnf(i_0_123, plain, (k2_polynom7(X1,X2,k1_polynom7(X1,X2,k4_struct_0(X2),X3))=k16_pre_poly(X1)|v2_struct_0(X2)|~v1_funct_1(X3)|~v1_relat_1(X3)|~l2_struct_0(X2)|~v4_valued_0(X3)|~v2_pre_poly(X3)|~v1_partfun1(X3,X1)|~v4_relat_1(X3,X1))).#trainneg
cnf(i_0_182, plain, (v2_struct_0(X2)|m1_subset_1(esk18_3(X1,X2,X3),u1_struct_0(X2))|~v1_funct_1(X3)|~l2_struct_0(X2)|~v4_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_233, plain, (k9_pre_poly(k15_pre_poly(X1),k5_numbers,k4_finseq_2(np__2,k15_pre_poly(X1)),k21_pre_poly(X1,X2),np__1)=k10_finseq_1(k16_pre_poly(X1),X2)|~v3_ordinal1(X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v4_valued_0(X2)|~v2_pre_poly(X2)|~v1_partfun1(X2,X1)|~v4_relat_1(X2,X1))).#trainneg
cnf(i_0_136, plain, (k3_polynom7(X4,X1,k1_polynom7(X4,X1,X2,X3))=X2|v2_struct_0(X1)|~v1_funct_1(X3)|~v1_relat_1(X3)|~l2_struct_0(X1)|~v4_valued_0(X3)|~v2_pre_poly(X3)|~v1_partfun1(X3,X4)|~v4_relat_1(X3,X4)|~m1_subset_1(X2,u1_struct_0(X1)))).#trainneg
cnf(i_0_71, plain, (k3_polynom1(X1,X2,X3,k16_pre_poly(X1))=k3_polynom7(X1,X2,X3)|v2_struct_0(X2)|~v1_funct_1(X3)|~l2_struct_0(X2)|~v4_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_57, plain, (v2_struct_0(X2)|v4_polynom7(X3,X1,X2)|k3_polynom1(X1,X2,X3,esk11_3(X1,X2,X3))!=k4_struct_0(X2)|~v1_funct_1(X3)|~l2_struct_0(X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_139, plain, (v2_struct_0(X2)|r2_funct_2(k15_pre_poly(X1),u1_struct_0(X2),X3,k7_polynom1(X1,X2))|k2_polynom1(k15_pre_poly(X1),X2,X3)!=k1_xboole_0|~v1_funct_1(X3)|~l2_struct_0(X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_140, plain, (k2_polynom1(k15_pre_poly(X1),X2,X3)=k1_xboole_0|v2_struct_0(X2)|~v1_funct_1(X3)|~l2_struct_0(X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~r2_funct_2(k15_pre_poly(X1),u1_struct_0(X2),X3,k7_polynom1(X1,X2)))).#trainneg
cnf(i_0_163, plain, (v2_struct_0(X2)|v1_funct_1(k5_vfunct_1(k15_pre_poly(X1),X2,X3))|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v1_funct_1(X3)|~l2_algstr_0(X2)|~v1_polynom1(X3,k15_pre_poly(X1),X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_160, plain, (k1_funct_1(X4,X1)=k1_funct_1(X5,k5_ordinal3(X1,X2))|X4!=k1_polynom6(X2,X3,X6,X5)|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~m2_subset_1(X6,k15_pre_poly(X2),k15_pre_poly(X2))|~m2_subset_1(X5,k15_pre_poly(X3),k15_pre_poly(X3))|~r2_hidden(X1,k6_subset_1(k10_ordinal2(X2,X3),X2))|~m2_subset_1(X4,k15_pre_poly(k10_ordinal2(X2,X3)),k15_pre_poly(k10_ordinal2(X2,X3))))).#trainneg
cnf(i_0_176, plain, (k2_polynom1(k15_pre_poly(X3),X1,X2)=k1_xboole_0|k1_tarski(k2_polynom7(X3,X1,X2))=k2_polynom1(k15_pre_poly(X3),X1,X2)|v2_struct_0(X1)|~v1_funct_1(X2)|~l2_struct_0(X1)|~v3_polynom7(X2,X3,X1)|~v1_funct_2(X2,k15_pre_poly(X3),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X3),u1_struct_0(X1)))))).#trainneg
cnf(i_0_243, plain, (k2_polynom1(k15_pre_poly(X1),X2,X3)=k1_xboole_0|v2_struct_0(X2)|X4!=k2_polynom7(X1,X2,X3)|k3_polynom1(X1,X2,X3,X4)!=k4_struct_0(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_relat_1(X4)|~l2_struct_0(X2)|~v4_valued_0(X4)|~v2_pre_poly(X4)|~v1_partfun1(X4,X1)|~v4_relat_1(X4,X1)|~v3_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_234, plain, (k2_polynom1(k15_pre_poly(X3),X1,X2)=k1_xboole_0|k2_polynom1(k15_pre_poly(X3),X1,X2)=k6_domain_1(k15_pre_poly(X3),k16_pre_poly(X3))|v2_struct_0(X1)|~v1_funct_1(X2)|~l2_struct_0(X1)|~v4_polynom7(X2,X3,X1)|~v1_funct_2(X2,k15_pre_poly(X3),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X3),u1_struct_0(X1)))))).#trainneg
cnf(i_0_208, plain, (v2_struct_0(X2)|v3_polynom7(X3,X1,X2)|esk20_4(X1,X2,X3,X4)!=X4|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_relat_1(X4)|~l2_struct_0(X2)|~v4_valued_0(X4)|~v2_pre_poly(X4)|~v1_partfun1(X4,X1)|~v4_relat_1(X4,X1)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_138, plain, (v7_struct_0(X2)|v1_finset_1(k2_polynom1(k15_pre_poly(X1),X2,X3))|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X3)|~v1_polynom1(X3,k15_pre_poly(X1),X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_187, plain, (v2_struct_0(X2)|v1_funct_1(k4_polynom1(X1,X2,X3,X4))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l2_algstr_0(X2)|~v1_funct_2(X4,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_212, plain, (v2_struct_0(X2)|v3_polynom7(X3,X1,X2)|v1_funct_1(esk20_4(X1,X2,X3,X4))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_relat_1(X4)|~l2_struct_0(X2)|~v4_valued_0(X4)|~v2_pre_poly(X4)|~v1_partfun1(X4,X1)|~v4_relat_1(X4,X1)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_255, plain, (v7_struct_0(X3)|m1_subset_1(k3_polynom2(X1,X2,X3,X4),u1_struct_0(X3))|~v3_ordinal1(X1)|~v4_vectsp_1(X3)|~l6_algstr_0(X3)|~v1_funct_1(X4)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v4_valued_0(X2)|~v2_pre_poly(X2)|~v1_partfun1(X2,X1)|~v4_relat_1(X2,X1)|~v1_funct_2(X4,X1,u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X3)))))).#trainneg
cnf(i_0_225, plain, (k5_polynom2(X1,X2,k7_polynom1(X1,X2),X3)=k4_struct_0(X2)|v7_struct_0(X2)|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X3)|~v1_funct_2(X3,X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2)))))).#trainneg
cnf(i_0_137, plain, (v7_struct_0(X2)|r3_orders_1(k17_pre_poly(X1),k2_polynom1(k15_pre_poly(X1),X2,X3))|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X3)|~v1_polynom1(X3,k15_pre_poly(X1),X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_26, plain, (X4=k2_polynom1(X1,X2,X3)|v1_xboole_0(X1)|m1_subset_1(esk5_4(X1,X2,X3,X4),X1)|~v1_funct_1(X3)|~l2_struct_0(X2)|~m1_subset_1(X4,k1_zfmisc_1(X1))|~v1_funct_2(X3,X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2)))))).#trainneg
cnf(i_0_169, plain, (k3_polynom1(X3,X1,X2,k2_polynom7(X3,X1,X2))=k3_polynom7(X3,X1,X2)|v2_struct_0(X1)|~v1_funct_1(X2)|~l2_struct_0(X1)|~v3_polynom7(X2,X3,X1)|~v1_funct_2(X2,k15_pre_poly(X3),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X3),u1_struct_0(X1)))))).#trainneg
cnf(i_0_254, plain, (v2_struct_0(X1)|r2_funct_2(k15_pre_poly(X3),u1_struct_0(X1),k4_polynom1(X3,X1,X2,k7_polynom1(X3,X1)),X2)|~v4_rlvect_1(X1)|~v1_funct_1(X2)|~l2_algstr_0(X1)|~v1_funct_2(X2,k15_pre_poly(X3),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X3),u1_struct_0(X1)))))).#trainneg
cnf(i_0_232, plain, (k9_pre_poly(k15_pre_poly(X1),k5_numbers,k4_finseq_2(np__2,k15_pre_poly(X1)),k21_pre_poly(X1,X2),k3_finseq_1(k21_pre_poly(X1,X2)))=k10_finseq_1(X2,k16_pre_poly(X1))|~v3_ordinal1(X1)|~v1_funct_1(X2)|~v1_relat_1(X2)|~v4_valued_0(X2)|~v2_pre_poly(X2)|~v1_partfun1(X2,X1)|~v4_relat_1(X2,X1))).#trainneg
cnf(i_0_181, plain, (v2_struct_0(X2)|r2_funct_2(k15_pre_poly(X1),u1_struct_0(X2),X3,k4_polynom7(X1,X2,esk18_3(X1,X2,X3)))|~v1_funct_1(X3)|~l2_struct_0(X2)|~v4_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_159, plain, (X5=k1_polynom6(X1,X2,X3,X4)|v3_ordinal1(esk17_5(X1,X2,X3,X4,X5))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~m2_subset_1(X4,k15_pre_poly(X2),k15_pre_poly(X2))|~m2_subset_1(X3,k15_pre_poly(X1),k15_pre_poly(X1))|~m2_subset_1(X5,k15_pre_poly(k10_ordinal2(X1,X2)),k15_pre_poly(k10_ordinal2(X1,X2))))).#trainneg
cnf(i_0_27, plain, (k3_funct_2(X1,u1_struct_0(X2),X3,X4)=k4_struct_0(X2)|v1_xboole_0(X1)|r2_hidden(X4,X5)|X5!=k2_polynom1(X1,X2,X3)|~v1_funct_1(X3)|~l2_struct_0(X2)|~m1_subset_1(X4,X1)|~m1_subset_1(X5,k1_zfmisc_1(X1))|~v1_funct_2(X3,X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2)))))).#trainneg
cnf(i_0_202, plain, (k2_polynom1(k15_pre_poly(X2),X3,X1)=k6_domain_1(k15_pre_poly(X2),k16_pre_poly(X2))|v2_struct_0(X3)|r2_funct_2(k15_pre_poly(X2),u1_struct_0(X3),X1,k7_polynom1(X2,X3))|~v1_funct_1(X1)|~l2_struct_0(X3)|~v4_polynom7(X1,X2,X3)|~v1_funct_2(X1,k15_pre_poly(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3)))))).#trainneg
cnf(i_0_53, plain, (X4=k7_pre_poly(X1,X2,X3)|v7_ordinal1(esk9_4(X1,X2,X3,X4))|k2_relset_1(X1,X4)!=X2|~v1_finset_1(X2)|~v1_relat_2(X3)|~v4_relat_2(X3)|~v8_relat_2(X3)|~v1_partfun1(X3,X1)|~r3_orders_1(X3,X2)|~m2_finseq_1(X4,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_52, plain, (X4=k7_pre_poly(X1,X2,X3)|v7_ordinal1(esk10_4(X1,X2,X3,X4))|k2_relset_1(X1,X4)!=X2|~v1_finset_1(X2)|~v1_relat_2(X3)|~v4_relat_2(X3)|~v8_relat_2(X3)|~v1_partfun1(X3,X1)|~r3_orders_1(X3,X2)|~m2_finseq_1(X4,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_162, plain, (v2_struct_0(X2)|v1_polynom1(k5_vfunct_1(k15_pre_poly(X1),X2,X3),k15_pre_poly(X1),X2)|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v1_funct_1(X3)|~l2_algstr_0(X2)|~v1_polynom1(X3,k15_pre_poly(X1),X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_127, plain, (k3_polynom1(X3,X1,X2,X4)=k1_funct_1(X2,X4)|v2_struct_0(X1)|~v1_funct_1(X4)|~v1_funct_1(X2)|~v1_relat_1(X4)|~l1_struct_0(X1)|~v4_valued_0(X4)|~v2_pre_poly(X4)|~v1_partfun1(X4,X3)|~v4_relat_1(X4,X3)|~v1_funct_2(X2,k15_pre_poly(X3),u1_struct_0(X1))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X3),u1_struct_0(X1)))))).#trainneg
cnf(i_0_237, plain, (v2_struct_0(X2)|v1_funct_1(k9_polynom1(X1,X2,X3,X4))|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_2(X4,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_55, plain, (r1_xxreal_0(X4,X3)|X2!=k7_pre_poly(X1,X5,X6)|k7_partfun1(X1,X2,X3)!=k7_partfun1(X1,X2,X4)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_finset_1(X5)|~v1_relat_2(X6)|~v4_relat_2(X6)|~v8_relat_2(X6)|~v1_partfun1(X6,X1)|~r3_orders_1(X6,X5)|~m2_finseq_1(X2,X1)|~m1_subset_1(X5,k1_zfmisc_1(X1))|~r2_hidden(X4,k4_finseq_1(X2))|~r2_hidden(X3,k4_finseq_1(X2))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_230, plain, (v2_struct_0(X2)|v1_funct_1(k6_polynom1(X1,X2,X3,X4))|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~l2_algstr_0(X2)|~v1_funct_2(X4,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_207, plain, (v2_struct_0(X2)|v3_polynom7(X3,X1,X2)|k3_polynom1(X1,X2,X3,esk20_4(X1,X2,X3,X4))!=k4_struct_0(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_relat_1(X4)|~l2_struct_0(X2)|~v4_valued_0(X4)|~v2_pre_poly(X4)|~v1_partfun1(X4,X1)|~v4_relat_1(X4,X1)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_240, plain, (X4=k2_polynom7(X1,X2,X3)|v2_struct_0(X2)|X4!=k16_pre_poly(X1)|k2_polynom1(k15_pre_poly(X1),X2,X3)!=k1_xboole_0|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_relat_1(X4)|~l2_struct_0(X2)|~v4_valued_0(X4)|~v2_pre_poly(X4)|~v1_partfun1(X4,X1)|~v4_relat_1(X4,X1)|~v3_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_214, plain, (v2_struct_0(X2)|v3_polynom7(X3,X1,X2)|v1_relat_1(esk20_4(X1,X2,X3,X4))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_relat_1(X4)|~l2_struct_0(X2)|~v4_valued_0(X4)|~v2_pre_poly(X4)|~v1_partfun1(X4,X1)|~v4_relat_1(X4,X1)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_186, plain, (v2_struct_0(X2)|v1_funct_2(k4_polynom1(X1,X2,X3,X4),k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l2_algstr_0(X2)|~v1_funct_2(X4,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_185, plain, (v2_struct_0(X2)|m1_subset_1(k4_polynom1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l2_algstr_0(X2)|~v1_funct_2(X4,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_210, plain, (v2_struct_0(X2)|v3_polynom7(X3,X1,X2)|v4_valued_0(esk20_4(X1,X2,X3,X4))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_relat_1(X4)|~l2_struct_0(X2)|~v4_valued_0(X4)|~v2_pre_poly(X4)|~v1_partfun1(X4,X1)|~v4_relat_1(X4,X1)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_209, plain, (v2_struct_0(X2)|v3_polynom7(X3,X1,X2)|v2_pre_poly(esk20_4(X1,X2,X3,X4))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_relat_1(X4)|~l2_struct_0(X2)|~v4_valued_0(X4)|~v2_pre_poly(X4)|~v1_partfun1(X4,X1)|~v4_relat_1(X4,X1)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_276, plain, (v7_struct_0(X2)|m1_subset_1(o_3_2_polynom7(X1,X2,X3),k2_polynom1(k15_pre_poly(X1),X2,X3))|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X3)|~v4_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_65, plain, (X4=k16_pre_poly(X2)|k3_polynom1(X2,X3,X1,X4)=k4_struct_0(X3)|v2_struct_0(X3)|~v1_funct_1(X4)|~v1_funct_1(X1)|~v1_relat_1(X4)|~l2_struct_0(X3)|~v4_valued_0(X4)|~v2_pre_poly(X4)|~v1_partfun1(X4,X2)|~v4_relat_1(X4,X2)|~v4_polynom7(X1,X2,X3)|~v1_funct_2(X1,k15_pre_poly(X2),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3)))))).#trainneg
cnf(i_0_25, plain, (X4=k2_polynom1(X1,X2,X3)|k3_funct_2(X1,u1_struct_0(X2),X3,esk5_4(X1,X2,X3,X4))=k4_struct_0(X2)|v1_xboole_0(X1)|~v1_funct_1(X3)|~l2_struct_0(X2)|~m1_subset_1(X4,k1_zfmisc_1(X1))|~v1_funct_2(X3,X1,u1_struct_0(X2))|~r2_hidden(esk5_4(X1,X2,X3,X4),X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2)))))).#trainneg
cnf(i_0_24, plain, (X4=k2_polynom1(X1,X2,X3)|v1_xboole_0(X1)|r2_hidden(esk5_4(X1,X2,X3,X4),X4)|k3_funct_2(X1,u1_struct_0(X2),X3,esk5_4(X1,X2,X3,X4))!=k4_struct_0(X2)|~v1_funct_1(X3)|~l2_struct_0(X2)|~m1_subset_1(X4,k1_zfmisc_1(X1))|~v1_funct_2(X3,X1,u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2)))))).#trainneg
cnf(i_0_51, plain, (X4=k7_pre_poly(X1,X2,X3)|r2_hidden(esk9_4(X1,X2,X3,X4),k4_finseq_1(X4))|k2_relset_1(X1,X4)!=X2|~v1_finset_1(X2)|~v1_relat_2(X3)|~v4_relat_2(X3)|~v8_relat_2(X3)|~v1_partfun1(X3,X1)|~r3_orders_1(X3,X2)|~m2_finseq_1(X4,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_50, plain, (X4=k7_pre_poly(X1,X2,X3)|r2_hidden(esk10_4(X1,X2,X3,X4),k4_finseq_1(X4))|k2_relset_1(X1,X4)!=X2|~v1_finset_1(X2)|~v1_relat_2(X3)|~v4_relat_2(X3)|~v8_relat_2(X3)|~v1_partfun1(X3,X1)|~r3_orders_1(X3,X2)|~m2_finseq_1(X4,X1)|~m1_subset_1(X2,k1_zfmisc_1(X1))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_211, plain, (v2_struct_0(X2)|v3_polynom7(X3,X1,X2)|v1_partfun1(esk20_4(X1,X2,X3,X4),X1)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_relat_1(X4)|~l2_struct_0(X2)|~v4_valued_0(X4)|~v2_pre_poly(X4)|~v1_partfun1(X4,X1)|~v4_relat_1(X4,X1)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_213, plain, (v2_struct_0(X2)|v3_polynom7(X3,X1,X2)|v4_relat_1(esk20_4(X1,X2,X3,X4),X1)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_relat_1(X4)|~l2_struct_0(X2)|~v4_valued_0(X4)|~v2_pre_poly(X4)|~v1_partfun1(X4,X1)|~v4_relat_1(X4,X1)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_158, plain, (X5=k1_polynom6(X1,X2,X3,X4)|r2_hidden(esk17_5(X1,X2,X3,X4,X5),X1)|r2_hidden(esk17_5(X1,X2,X3,X4,X5),k6_subset_1(k10_ordinal2(X1,X2),X1))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~m2_subset_1(X4,k15_pre_poly(X2),k15_pre_poly(X2))|~m2_subset_1(X3,k15_pre_poly(X1),k15_pre_poly(X1))|~m2_subset_1(X5,k15_pre_poly(k10_ordinal2(X1,X2)),k15_pre_poly(k10_ordinal2(X1,X2))))).#trainneg
cnf(i_0_247, plain, (k5_polynom1(X3,X1,X2,X4)=k4_polynom1(X3,X1,X2,X4)|v2_struct_0(X1)|~v4_rlvect_1(X1)|~v1_funct_1(X4)|~v1_funct_1(X2)|~l2_algstr_0(X1)|~v2_rlvect_1(X1)|~v1_funct_2(X4,k15_pre_poly(X3),u1_struct_0(X1))|~v1_funct_2(X2,k15_pre_poly(X3),u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X3),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X3),u1_struct_0(X1)))))).#trainneg
cnf(i_0_238, plain, (k5_polynom1(X3,X1,X2,X4)=k5_polynom1(X3,X1,X4,X2)|v2_struct_0(X1)|~v4_rlvect_1(X1)|~v1_funct_1(X4)|~v1_funct_1(X2)|~l2_algstr_0(X1)|~v2_rlvect_1(X1)|~v1_funct_2(X4,k15_pre_poly(X3),u1_struct_0(X1))|~v1_funct_2(X2,k15_pre_poly(X3),u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X3),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X3),u1_struct_0(X1)))))).#trainneg
cnf(i_0_241, plain, (X4=k2_polynom7(X1,X2,X3)|k3_polynom1(X1,X2,X3,X4)=k4_struct_0(X2)|v2_struct_0(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_relat_1(X4)|~l2_struct_0(X2)|~v4_valued_0(X4)|~v2_pre_poly(X4)|~v1_partfun1(X4,X1)|~v4_relat_1(X4,X1)|~v3_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_215, plain, (X1=esk19_3(X2,X3,X4)|k3_polynom1(X2,X3,X4,X1)=k4_struct_0(X3)|v2_struct_0(X3)|~v1_funct_1(X4)|~v1_funct_1(X1)|~v1_relat_1(X1)|~l2_struct_0(X3)|~v4_valued_0(X1)|~v2_pre_poly(X1)|~v1_partfun1(X1,X2)|~v4_relat_1(X1,X2)|~v3_polynom7(X4,X2,X3)|~v1_funct_2(X4,k15_pre_poly(X2),u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3)))))).#trainneg
cnf(i_0_229, plain, (v2_struct_0(X2)|v1_funct_2(k6_polynom1(X1,X2,X3,X4),k15_pre_poly(X1),u1_struct_0(X2))|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~l2_algstr_0(X2)|~v1_funct_2(X4,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_236, plain, (v2_struct_0(X2)|v1_funct_2(k9_polynom1(X1,X2,X3,X4),k15_pre_poly(X1),u1_struct_0(X2))|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_2(X4,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_155, plain, (X1=k1_polynom6(X2,X3,X4,X5)|k1_funct_1(X1,esk17_5(X2,X3,X4,X5,X1))!=k1_funct_1(X4,esk17_5(X2,X3,X4,X5,X1))|k1_funct_1(X1,esk17_5(X2,X3,X4,X5,X1))!=k1_funct_1(X5,k5_ordinal3(esk17_5(X2,X3,X4,X5,X1),X2))|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~m2_subset_1(X5,k15_pre_poly(X3),k15_pre_poly(X3))|~m2_subset_1(X4,k15_pre_poly(X2),k15_pre_poly(X2))|~m2_subset_1(X1,k15_pre_poly(k10_ordinal2(X2,X3)),k15_pre_poly(k10_ordinal2(X2,X3))))).#trainneg
cnf(i_0_228, plain, (v2_struct_0(X2)|m1_subset_1(k6_polynom1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~l2_algstr_0(X2)|~v1_funct_2(X4,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_48, plain, (X2=k7_pre_poly(X1,X3,X4)|k7_partfun1(X1,X2,esk10_4(X1,X3,X4,X2))=k7_partfun1(X1,X2,esk9_4(X1,X3,X4,X2))|k2_relset_1(X1,X2)!=X3|~v1_finset_1(X3)|~v1_relat_2(X4)|~v4_relat_2(X4)|~v8_relat_2(X4)|~v1_partfun1(X4,X1)|~r3_orders_1(X4,X3)|~m2_finseq_1(X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,X1)))|~r2_hidden(k4_tarski(k7_partfun1(X1,X2,esk9_4(X1,X3,X4,X2)),k7_partfun1(X1,X2,esk10_4(X1,X3,X4,X2))),X4))).#trainneg
cnf(i_0_54, plain, (r1_xxreal_0(X4,X3)|r2_hidden(k4_tarski(k7_partfun1(X1,X2,X3),k7_partfun1(X1,X2,X4)),X5)|X2!=k7_pre_poly(X1,X6,X5)|~v7_ordinal1(X4)|~v7_ordinal1(X3)|~v1_finset_1(X6)|~v1_relat_2(X5)|~v4_relat_2(X5)|~v8_relat_2(X5)|~v1_partfun1(X5,X1)|~r3_orders_1(X5,X6)|~m2_finseq_1(X2,X1)|~m1_subset_1(X6,k1_zfmisc_1(X1))|~r2_hidden(X4,k4_finseq_1(X2))|~r2_hidden(X3,k4_finseq_1(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X1,X1))))).#trainneg
cnf(i_0_130, plain, (k4_polynom1(X3,X1,X2,k5_vfunct_1(k15_pre_poly(X3),X1,X4))=k6_polynom1(X3,X1,X2,X4)|v2_struct_0(X1)|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v1_funct_1(X4)|~v1_funct_1(X2)|~l2_algstr_0(X1)|~v1_funct_2(X4,k15_pre_poly(X3),u1_struct_0(X1))|~v1_funct_2(X2,k15_pre_poly(X3),u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X3),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X3),u1_struct_0(X1)))))).#trainneg
cnf(i_0_235, plain, (v2_struct_0(X2)|m1_subset_1(k9_polynom1(X1,X2,X3,X4),k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v1_funct_2(X4,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_277, plain, (k10_polynom1(X1,X2,X3,X4)=k10_polynom1(X1,X2,X4,X3)|v2_struct_0(X2)|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~v2_rlvect_1(X2)|~v5_group_1(X2)|~v1_funct_2(X4,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_106, plain, (v7_struct_0(X2)|m2_finseq_1(esk13_5(X1,X2,X3,X4,X5),u1_struct_0(X2))|X5!=k5_polynom2(X1,X2,X3,X4)|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X5,u1_struct_0(X2))|~v1_funct_2(X4,X1,u1_struct_0(X2))|~v1_polynom1(X3,k15_pre_poly(X1),X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_104, plain, (k4_rlvect_1(X2,esk13_5(X3,X2,X4,X5,X1))=X1|v7_struct_0(X2)|X1!=k5_polynom2(X3,X2,X4,X5)|~v3_ordinal1(X3)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X1,u1_struct_0(X2))|~v1_funct_2(X5,X3,u1_struct_0(X2))|~v1_polynom1(X4,k15_pre_poly(X3),X2)|~v1_funct_2(X4,k15_pre_poly(X3),u1_struct_0(X2))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X3),u1_struct_0(X2)))))).#trainneg
cnf(i_0_88, plain, (k1_algstr_0(X1,k3_polynom1(X3,X1,X2,X5),k3_polynom1(X3,X1,X4,X5))=k3_polynom1(X3,X1,k4_polynom1(X3,X1,X2,X4),X5)|v2_struct_0(X1)|~v1_funct_1(X5)|~v1_funct_1(X4)|~v1_funct_1(X2)|~v1_relat_1(X5)|~l2_algstr_0(X1)|~v4_valued_0(X5)|~v2_pre_poly(X5)|~v1_partfun1(X5,X3)|~v4_relat_1(X5,X3)|~v1_funct_2(X4,k15_pre_poly(X3),u1_struct_0(X1))|~v1_funct_2(X2,k15_pre_poly(X3),u1_struct_0(X1))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X3),u1_struct_0(X1))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X3),u1_struct_0(X1)))))).#trainneg
cnf(i_0_157, plain, (X1=k1_polynom6(X2,X3,X4,X5)|r2_hidden(esk17_5(X2,X3,X4,X5,X1),X2)|k1_funct_1(X1,esk17_5(X2,X3,X4,X5,X1))!=k1_funct_1(X5,k5_ordinal3(esk17_5(X2,X3,X4,X5,X1),X2))|~v3_ordinal1(X3)|~v3_ordinal1(X2)|~m2_subset_1(X5,k15_pre_poly(X3),k15_pre_poly(X3))|~m2_subset_1(X4,k15_pre_poly(X2),k15_pre_poly(X2))|~m2_subset_1(X1,k15_pre_poly(k10_ordinal2(X2,X3)),k15_pre_poly(k10_ordinal2(X2,X3))))).#trainneg
cnf(i_0_156, plain, (X5=k1_polynom6(X1,X2,X3,X4)|r2_hidden(esk17_5(X1,X2,X3,X4,X5),k6_subset_1(k10_ordinal2(X1,X2),X1))|k1_funct_1(X3,esk17_5(X1,X2,X3,X4,X5))!=k1_funct_1(X5,esk17_5(X1,X2,X3,X4,X5))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~m2_subset_1(X4,k15_pre_poly(X2),k15_pre_poly(X2))|~m2_subset_1(X3,k15_pre_poly(X1),k15_pre_poly(X1))|~m2_subset_1(X5,k15_pre_poly(k10_ordinal2(X1,X2)),k15_pre_poly(k10_ordinal2(X1,X2))))).#trainneg
cnf(i_0_105, plain, (k3_finseq_1(esk13_5(X1,X2,X3,X4,X5))=k3_finseq_1(k7_pre_poly(k15_pre_poly(X1),k2_polynom1(k15_pre_poly(X1),X2,X3),k17_pre_poly(X1)))|v7_struct_0(X2)|X5!=k5_polynom2(X1,X2,X3,X4)|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X5,u1_struct_0(X2))|~v1_funct_2(X4,X1,u1_struct_0(X2))|~v1_polynom1(X3,k15_pre_poly(X1),X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_101, plain, (X5=k5_polynom2(X1,X2,X3,X4)|v7_struct_0(X2)|r1_xxreal_0(np__1,esk14_6(X1,X2,X3,X4,X5,X6))|X5!=k4_rlvect_1(X2,X6)|k3_finseq_1(X6)!=k3_finseq_1(k7_pre_poly(k15_pre_poly(X1),k2_polynom1(k15_pre_poly(X1),X2,X3),k17_pre_poly(X1)))|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X5,u1_struct_0(X2))|~m2_finseq_1(X6,u1_struct_0(X2))|~v1_funct_2(X4,X1,u1_struct_0(X2))|~v1_polynom1(X3,k15_pre_poly(X1),X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_102, plain, (X5=k5_polynom2(X1,X2,X3,X4)|v7_struct_0(X2)|m1_subset_1(esk14_6(X1,X2,X3,X4,X5,X6),k5_numbers)|X5!=k4_rlvect_1(X2,X6)|k3_finseq_1(X6)!=k3_finseq_1(k7_pre_poly(k15_pre_poly(X1),k2_polynom1(k15_pre_poly(X1),X2,X3),k17_pre_poly(X1)))|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X5,u1_struct_0(X2))|~m2_finseq_1(X6,u1_struct_0(X2))|~v1_funct_2(X4,X1,u1_struct_0(X2))|~v1_polynom1(X3,k15_pre_poly(X1),X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_100, plain, (X5=k5_polynom2(X1,X2,X3,X4)|v7_struct_0(X2)|r1_xxreal_0(esk14_6(X1,X2,X3,X4,X5,X6),k3_finseq_1(X6))|X5!=k4_rlvect_1(X2,X6)|k3_finseq_1(X6)!=k3_finseq_1(k7_pre_poly(k15_pre_poly(X1),k2_polynom1(k15_pre_poly(X1),X2,X3),k17_pre_poly(X1)))|~v3_ordinal1(X1)|~v13_algstr_0(X2)|~v3_rlvect_1(X2)|~v4_rlvect_1(X2)|~v4_vectsp_1(X2)|~v5_vectsp_1(X2)|~l6_algstr_0(X2)|~v1_funct_1(X4)|~v1_funct_1(X3)|~m1_subset_1(X5,u1_struct_0(X2))|~m2_finseq_1(X6,u1_struct_0(X2))|~v1_funct_2(X4,X1,u1_struct_0(X2))|~v1_polynom1(X3,k15_pre_poly(X1),X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2)))))).#trainneg
cnf(i_0_103, plain, (k7_partfun1(u1_struct_0(X3),esk13_5(X2,X3,X4,X5,X6),X1)=k6_algstr_0(X3,k7_partfun1(u1_struct_0(X3),k1_partfun1(k5_numbers,k15_pre_poly(X2),k15_pre_poly(X2),u1_struct_0(X3),k7_pre_poly(k15_pre_poly(X2),k2_polynom1(k15_pre_poly(X2),X3,X4),k17_pre_poly(X2)),X4),X1),k3_polynom2(X2,k4_polynom2(X2,k7_partfun1(k15_pre_poly(X2),k7_pre_poly(k15_pre_poly(X2),k2_polynom1(k15_pre_poly(X2),X3,X4),k17_pre_poly(X2)),X1)),X3,X5))|v7_struct_0(X3)|X6!=k5_polynom2(X2,X3,X4,X5)|~v3_ordinal1(X2)|~v13_algstr_0(X3)|~v3_rlvect_1(X3)|~v4_rlvect_1(X3)|~v4_vectsp_1(X3)|~v5_vectsp_1(X3)|~l6_algstr_0(X3)|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X1,k5_numbers)|~r1_xxreal_0(np__1,X1)|~m1_subset_1(X6,u1_struct_0(X3))|~v1_funct_2(X5,X2,u1_struct_0(X3))|~v1_polynom1(X4,k15_pre_poly(X2),X3)|~v1_funct_2(X4,k15_pre_poly(X2),u1_struct_0(X3))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~r1_xxreal_0(X1,k3_finseq_1(esk13_5(X2,X3,X4,X5,X6)))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3)))))).#trainneg
cnf(i_0_99, plain, (X6=k5_polynom2(X3,X1,X4,X5)|v7_struct_0(X1)|X6!=k4_rlvect_1(X1,X2)|k3_finseq_1(X2)!=k3_finseq_1(k7_pre_poly(k15_pre_poly(X3),k2_polynom1(k15_pre_poly(X3),X1,X4),k17_pre_poly(X3)))|k6_algstr_0(X1,k7_partfun1(u1_struct_0(X1),k1_partfun1(k5_numbers,k15_pre_poly(X3),k15_pre_poly(X3),u1_struct_0(X1),k7_pre_poly(k15_pre_poly(X3),k2_polynom1(k15_pre_poly(X3),X1,X4),k17_pre_poly(X3)),X4),esk14_6(X3,X1,X4,X5,X6,X2)),k3_polynom2(X3,k4_polynom2(X3,k7_partfun1(k15_pre_poly(X3),k7_pre_poly(k15_pre_poly(X3),k2_polynom1(k15_pre_poly(X3),X1,X4),k17_pre_poly(X3)),esk14_6(X3,X1,X4,X5,X6,X2))),X1,X5))!=k7_partfun1(u1_struct_0(X1),X2,esk14_6(X3,X1,X4,X5,X6,X2))|~v3_ordinal1(X3)|~v13_algstr_0(X1)|~v3_rlvect_1(X1)|~v4_rlvect_1(X1)|~v4_vectsp_1(X1)|~v5_vectsp_1(X1)|~l6_algstr_0(X1)|~v1_funct_1(X5)|~v1_funct_1(X4)|~m1_subset_1(X6,u1_struct_0(X1))|~m2_finseq_1(X2,u1_struct_0(X1))|~v1_funct_2(X5,X3,u1_struct_0(X1))|~v1_polynom1(X4,k15_pre_poly(X3),X1)|~v1_funct_2(X4,k15_pre_poly(X3),u1_struct_0(X1))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X3,u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X3),u1_struct_0(X1)))))).#trainneg
cnf(i_0_302, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l2_struct_0(X1))).#trainneg
cnf(i_0_521, negated_conjecture, (k1_relset_1(esk1_0,esk4_0)=esk1_0|u1_struct_0(esk2_0)=k1_xboole_0)).#trainneg
cnf(i_0_300, plain, (m1_subset_1(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_315, negated_conjecture, (v1_relat_1(esk4_0))).#trainneg
cnf(i_0_316, negated_conjecture, (v1_relat_1(esk3_0))).#trainneg
cnf(i_0_40, plain, (r2_hidden(X1,X2)|X2!=k1_tarski(X1))).#trainneg
cnf(i_0_326, negated_conjecture, (v4_relat_1(esk4_0,esk1_0))).#trainneg
cnf(i_0_312, negated_conjecture, (v1_xboole_0(esk4_0)|~v1_xboole_0(k2_zfmisc_1(esk1_0,u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_291, plain, (l2_struct_0(X1)|~l5_algstr_0(X1))).#trainneg
cnf(i_0_292, plain, (l1_struct_0(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_293, plain, (l1_struct_0(X1)|~l4_struct_0(X1))).#trainneg
cnf(i_0_522, negated_conjecture, (k1_relset_1(k15_pre_poly(esk1_0),esk3_0)=k15_pre_poly(esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0)).#trainneg
cnf(i_0_831, plain, (X1=k5_polynom2(X2,X3,X4,X5)|r1_xxreal_0(np__1,esk14_6(X2,X3,X4,X5,X1,X6))|v7_struct_0(X3)|k3_finseq_1(X6)!=k3_finseq_1(esk13_5(X2,X3,X4,X7,X8))|X1!=k4_rlvect_1(X3,X6)|~v1_polynom1(X4,k15_pre_poly(X2),X3)|~m2_finseq_1(X6,u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X1,u1_struct_0(X3))|~v1_funct_2(X4,k15_pre_poly(X2),u1_struct_0(X3))|~v1_funct_2(X5,X2,u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v13_algstr_0(X3)|~v3_ordinal1(X2)|X8!=k5_polynom2(X2,X3,X4,X7)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X8,u1_struct_0(X3))|~v1_funct_2(X7,X2,u1_struct_0(X3))|~v1_funct_1(X7))).#trainneg
cnf(i_0_828, plain, (k3_finseq_1(esk13_5(X1,X2,X3,X4,X5))=k3_finseq_1(esk13_5(X1,X2,X3,X6,X7))|v7_struct_0(X2)|X5!=k5_polynom2(X1,X2,X3,X4)|~v1_polynom1(X3,k15_pre_poly(X1),X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~m1_subset_1(X5,u1_struct_0(X2))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_2(X4,X1,u1_struct_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v13_algstr_0(X2)|~v3_ordinal1(X1)|X7!=k5_polynom2(X1,X2,X3,X6)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~m1_subset_1(X7,u1_struct_0(X2))|~v1_funct_2(X6,X1,u1_struct_0(X2))|~v1_funct_1(X6))).#trainneg
cnf(i_0_635, plain, (k1_tarski(k16_pre_poly(X1))=k2_polynom1(k15_pre_poly(X1),X2,X3)|k2_polynom1(k15_pre_poly(X1),X2,X3)=k1_xboole_0|v2_struct_0(X2)|~v3_polynom7(X3,X1,X2)|~l2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~v4_polynom7(X3,X1,X2))).#trainneg
cnf(i_0_830, plain, (X1=k5_polynom2(X2,X3,X4,X5)|r1_xxreal_0(np__1,esk14_6(X2,X3,X4,X5,X1,k7_pre_poly(k15_pre_poly(X2),k2_polynom1(k15_pre_poly(X2),X3,X4),k17_pre_poly(X2))))|v7_struct_0(X3)|X1!=k4_rlvect_1(X3,k7_pre_poly(k15_pre_poly(X2),k2_polynom1(k15_pre_poly(X2),X3,X4),k17_pre_poly(X2)))|~v1_polynom1(X4,k15_pre_poly(X2),X3)|~m2_finseq_1(k7_pre_poly(k15_pre_poly(X2),k2_polynom1(k15_pre_poly(X2),X3,X4),k17_pre_poly(X2)),u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X1,u1_struct_0(X3))|~v1_funct_2(X4,k15_pre_poly(X2),u1_struct_0(X3))|~v1_funct_2(X5,X2,u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v13_algstr_0(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_565, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(k3_polynom7(X2,X1,X3),u1_struct_0(X1))|v2_struct_0(X1)|~v3_polynom7(X3,X2,X1)|~l2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X1))))|~v1_funct_2(X3,k15_pre_poly(X2),u1_struct_0(X1))|~v1_funct_1(X3))).#trainneg
cnf(i_0_539, plain, (k4_rlvect_1(X1,k9_finseq_1(X2))=X2|v2_struct_0(X1)|v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1)|~m1_subset_1(X2,u1_struct_0(X1))|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1))).#trainneg
cnf(i_0_641, negated_conjecture, (k6_domain_1(k15_pre_poly(esk1_0),k16_pre_poly(esk1_0))=k2_polynom1(k15_pre_poly(esk1_0),esk2_0,esk3_0)|k2_polynom1(k15_pre_poly(esk1_0),esk2_0,esk3_0)=k1_xboole_0|v2_struct_0(esk2_0)|~l2_struct_0(esk2_0))).#trainneg
cnf(i_0_324, negated_conjecture, (v5_relat_1(esk4_0,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_295, negated_conjecture, (esk4_0=k1_xboole_0|u1_struct_0(esk2_0)!=k1_xboole_0)).#trainneg
cnf(i_0_128, plain, (r2_funct_2(X1,X2,X3,X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,X2)))|~v1_funct_2(X3,X1,X2)|~v1_funct_1(X3))).#trainneg
cnf(i_0_296, negated_conjecture, (esk3_0=k1_xboole_0|u1_struct_0(esk2_0)!=k1_xboole_0)).#trainneg
cnf(i_0_325, negated_conjecture, (v5_relat_1(esk3_0,u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_313, negated_conjecture, (v1_xboole_0(esk3_0)|~v1_xboole_0(k2_zfmisc_1(k15_pre_poly(esk1_0),u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_327, negated_conjecture, (v4_relat_1(esk3_0,k15_pre_poly(esk1_0)))).#trainneg
cnf(i_0_321, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(esk1_0,u1_struct_0(esk2_0))))|r2_hidden(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_311, plain, (k3_finseq_1(k9_finseq_1(X1))=np__1|~v1_finseq_1(k9_finseq_1(X1))|~v1_relat_1(k9_finseq_1(X1))|~v1_funct_1(k9_finseq_1(X1)))).#trainneg
cnf(i_0_317, plain, (v1_finset_1(k3_finseq_1(X1))|~v1_finset_1(X1)|~v1_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_706, plain, (v3_polynom7(k7_polynom1(X1,X2),X1,X2)|v2_struct_0(X2)|~v2_pre_poly(X3)|~v4_valued_0(X3)|~v4_relat_1(X3,X1)|~v1_partfun1(X3,X1)|~l2_struct_0(X2)|~v1_relat_1(X3)|~m1_subset_1(k7_polynom1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v1_funct_2(k7_polynom1(X1,X2),k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~v1_funct_1(k7_polynom1(X1,X2))|~v2_pre_poly(esk20_4(X1,X2,k7_polynom1(X1,X2),X3))|~v4_valued_0(esk20_4(X1,X2,k7_polynom1(X1,X2),X3))|~v4_relat_1(esk20_4(X1,X2,k7_polynom1(X1,X2),X3),X1)|~v1_partfun1(esk20_4(X1,X2,k7_polynom1(X1,X2),X3),X1)|~v1_relat_1(esk20_4(X1,X2,k7_polynom1(X1,X2),X3))|~v1_funct_1(esk20_4(X1,X2,k7_polynom1(X1,X2),X3)))).#trainneg
cnf(i_0_436, plain, (X1=k9_xtuple_0(X2)|~v4_relat_1(X2,X1)|~v1_relat_1(X2)|~v1_partfun1(X2,X1))).#trainneg
cnf(i_0_639, plain, (k2_polynom1(k15_pre_poly(X1),X2,X3)=k1_xboole_0|v2_struct_0(X2)|k3_polynom1(X1,X2,X3,X4)!=k4_struct_0(X2)|X4!=k16_pre_poly(X1)|~v3_polynom7(X3,X1,X2)|~v2_pre_poly(X4)|~v4_valued_0(X4)|~v4_relat_1(X4,X1)|~v1_partfun1(X4,X1)|~l2_struct_0(X2)|~v1_relat_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v4_polynom7(X3,X1,X2))).#trainneg
cnf(i_0_586, plain, (v1_xboole_0(X1)|r2_hidden(k3_funct_2(X2,X1,X3,X4),X1)|v1_xboole_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X2,X1)))|~m1_subset_1(X4,X2)|~v1_funct_2(X3,X2,X1)|~v1_funct_1(X3))).#trainneg
cnf(i_0_813, plain, (k1_algstr_0(X1,k3_polynom1(X2,X1,X3,X4),k4_struct_0(X1))=k3_polynom1(X2,X1,k4_polynom1(X2,X1,X3,X5),X4)|v2_struct_0(X1)|X4=k16_pre_poly(X2)|~v2_pre_poly(X4)|~v4_valued_0(X4)|~v4_relat_1(X4,X2)|~v1_partfun1(X4,X2)|~l2_algstr_0(X1)|~v1_relat_1(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X1))))|~v1_funct_2(X5,k15_pre_poly(X2),u1_struct_0(X1))|~v1_funct_2(X3,k15_pre_poly(X2),u1_struct_0(X1))|~v1_funct_1(X4)|~v1_funct_1(X5)|~v1_funct_1(X3)|~l2_struct_0(X1)|~v4_polynom7(X5,X2,X1))).#trainneg
cnf(i_0_877, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l2_algstr_0(X1))).#trainneg
cnf(i_0_870, plain, (v4_funct_1(k15_pre_poly(X1)))).#trainneg
cnf(i_0_850, plain, (v1_funct_1(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_882, plain, (l1_struct_0(X1)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_930, negated_conjecture, (k10_xtuple_0(esk4_0)=k2_relset_1(u1_struct_0(esk2_0),esk4_0))).#trainneg
cnf(i_0_278, plain, (r2_funct_2(k15_pre_poly(X1),u1_struct_0(X2),k4_polynom7(X1,X2,X3),k4_polynom7(X1,X2,X3))|v2_struct_0(X2)|~l2_struct_0(X2)|~m1_subset_1(X3,u1_struct_0(X2)))).#trainneg
cnf(i_0_318, plain, (v1_card_1(k3_finseq_1(X1))|~v1_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_847, negated_conjecture, (v1_partfun1(esk4_0,esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0|~v4_relat_1(esk4_0,esk1_0)|~v1_relat_1(esk4_0))).#trainneg
cnf(i_0_949, negated_conjecture, (k9_xtuple_0(esk4_0)=esk1_0|~v1_partfun1(esk4_0,esk1_0))).#trainneg
cnf(i_0_950, negated_conjecture, (k9_xtuple_0(esk3_0)=k15_pre_poly(esk1_0)|~v1_partfun1(esk3_0,k15_pre_poly(esk1_0)))).#trainneg
cnf(i_0_883, plain, (l1_struct_0(X1)|~l5_algstr_0(X1))).#trainneg
cnf(i_0_857, plain, (v1_relat_1(k2_zfmisc_1(X1,X2)))).#trainneg
cnf(i_0_319, plain, (v1_finseq_1(X1)|k9_xtuple_0(X1)!=k9_xtuple_0(X2)|~v7_ordinal1(esk6_1(X2))|~v1_relat_1(X1)|~v1_finseq_1(X2)|~v1_relat_1(X2))).#trainneg
cnf(i_0_320, plain, (v1_finseq_1(X1)|k9_xtuple_0(X1)!=k1_tarski(np__1)|~v7_ordinal1(np__1)|~v1_relat_1(X1))).#trainneg
cnf(i_0_932, plain, (k2_polynom1(k15_pre_poly(X1),X2,k7_polynom1(X1,X2))=k1_xboole_0|v2_struct_0(X2)|~l2_struct_0(X2)|~m1_subset_1(k7_polynom1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v1_funct_2(k7_polynom1(X1,X2),k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(k7_polynom1(X1,X2)))).#trainneg
cnf(i_0_874, plain, (r2_hidden(X1,k1_tarski(X1)))).#trainneg
cnf(i_0_435, plain, (m1_subset_1(k16_pre_poly(X1),k15_pre_poly(X1)))).#trainneg
cnf(i_0_545, plain, (k1_polynom6(X1,X2,X3,k16_pre_poly(X2))=X3|X2!=k1_xboole_0|~m2_subset_1(X3,k15_pre_poly(X1),k15_pre_poly(X1))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_543, plain, (k1_polynom6(X1,X2,X3,k16_pre_poly(X2))=k16_pre_poly(X2)|X1!=k1_xboole_0|~m2_subset_1(X3,k15_pre_poly(X1),k15_pre_poly(X1))|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_668, plain, (k4_polynom1(X1,X2,X3,k7_polynom1(X1,X2))=X3|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~m1_subset_1(k4_polynom1(X1,X2,X3,k7_polynom1(X1,X2)),k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_2(k4_polynom1(X1,X2,X3,k7_polynom1(X1,X2)),k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~v1_funct_1(k4_polynom1(X1,X2,X3,k7_polynom1(X1,X2)))|~l2_algstr_0(X2)|~v4_rlvect_1(X2))).#trainneg
cnf(i_0_475, plain, (v1_relat_1(k7_polynom1(X1,X2))|v2_struct_0(X2)|~l2_struct_0(X2))).#trainneg
cnf(i_0_700, plain, (k3_polynom1(X1,X2,X3,X4)=k1_funct_1(X3,X4)|v2_struct_0(X2)|~v2_pre_poly(X4)|~v4_valued_0(X4)|~v4_relat_1(X4,X1)|~v1_partfun1(X4,X1)|~v1_relat_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l2_struct_0(X2))).#trainneg
cnf(i_0_838, plain, (X1=k5_polynom2(X2,X3,X4,X5)|r1_xxreal_0(esk14_6(X2,X3,X4,X5,X1,k7_pre_poly(k15_pre_poly(X2),k2_polynom1(k15_pre_poly(X2),X3,X4),k17_pre_poly(X2))),k3_finseq_1(k7_pre_poly(k15_pre_poly(X2),k2_polynom1(k15_pre_poly(X2),X3,X4),k17_pre_poly(X2))))|v7_struct_0(X3)|X1!=k4_rlvect_1(X3,k7_pre_poly(k15_pre_poly(X2),k2_polynom1(k15_pre_poly(X2),X3,X4),k17_pre_poly(X2)))|~v1_polynom1(X4,k15_pre_poly(X2),X3)|~m2_finseq_1(k7_pre_poly(k15_pre_poly(X2),k2_polynom1(k15_pre_poly(X2),X3,X4),k17_pre_poly(X2)),u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X1,u1_struct_0(X3))|~v1_funct_2(X4,k15_pre_poly(X2),u1_struct_0(X3))|~v1_funct_2(X5,X2,u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v13_algstr_0(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_687, plain, (X1=k7_polynom1(X2,X3)|k2_polynom1(k15_pre_poly(X2),X3,X1)=k6_domain_1(k15_pre_poly(X2),k16_pre_poly(X2))|v2_struct_0(X3)|~m1_subset_1(k7_polynom1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3))))|~v1_funct_2(k7_polynom1(X2,X3),k15_pre_poly(X2),u1_struct_0(X3))|~v1_funct_2(X1,k15_pre_poly(X2),u1_struct_0(X3))|~v1_funct_1(k7_polynom1(X2,X3))|~v1_funct_1(X1)|~l2_struct_0(X3)|~v4_polynom7(X1,X2,X3))).#trainneg
cnf(i_0_395, plain, (k1_relset_1(k2_zfmisc_1(X1,X2),X3)=k1_xboole_0|X2!=k1_xboole_0|~v1_funct_2(k1_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v4_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_851, plain, (v1_relat_1(X1)|~v4_funct_1(k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_858, plain, (v4_relat_1(k2_zfmisc_1(X1,X2),X1))).#trainneg
cnf(i_0_935, negated_conjecture, (k10_xtuple_0(esk3_0)=k2_relset_1(u1_struct_0(esk2_0),esk3_0))).#trainneg
cnf(i_0_540, plain, (k3_relat_1(X1,X2)=k9_finseq_1(X3)|k1_funct_1(X2,k1_funct_1(X1,np__1))!=X3|k3_finseq_1(k3_relat_1(X1,X2))!=np__1|~v1_finseq_1(k3_relat_1(X1,X2))|~v1_relat_1(k3_relat_1(X1,X2))|~v1_funct_1(k3_relat_1(X1,X2))|~r2_hidden(np__1,k9_xtuple_0(X1))|~v1_relat_1(X2)|~v1_relat_1(X1)|~v1_funct_1(X2)|~v1_funct_1(X1))).#trainneg
cnf(i_0_322, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(esk1_0),u1_struct_0(esk2_0))))|r2_hidden(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(esk1_0),u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_880, negated_conjecture, (r2_hidden(esk4_0,X1)|X1!=k15_pre_poly(esk1_0)|~v2_pre_poly(esk4_0)|~v4_valued_0(esk4_0)|~v1_partfun1(esk4_0,esk1_0))).#trainneg
cnf(i_0_890, negated_conjecture, (v1_partfun1(esk3_0,k15_pre_poly(esk1_0))|u1_struct_0(esk2_0)=k1_xboole_0|~v4_relat_1(esk3_0,k15_pre_poly(esk1_0)))).#trainneg
cnf(i_0_845, negated_conjecture, (k9_xtuple_0(esk4_0)=esk1_0|u1_struct_0(esk2_0)=k1_xboole_0|~v4_relat_1(esk4_0,esk1_0)|~v1_relat_1(esk4_0))).#trainneg
cnf(i_0_862, plain, (v5_relat_1(k2_zfmisc_1(X1,X2),X2))).#trainneg
cnf(i_0_873, plain, (v1_xboole_0(X1)|m2_subset_1(X2,X1,X1)|~m1_subset_1(X2,X1))).#trainneg
cnf(i_0_328, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(esk16_1(X1),u1_struct_0(X1))|v4_rlvect_1(X1)|~l2_algstr_0(X1))).#trainneg
cnf(i_0_329, plain, (v1_funct_1(esk16_1(X1))|v4_rlvect_1(X1)|~v4_funct_1(u1_struct_0(X1))|~l2_algstr_0(X1))).#trainneg
cnf(i_0_878, plain, (k3_polynom1(X1,X2,X3,X4)=k1_funct_1(X3,X4)|v2_struct_0(X2)|~v2_pre_poly(X4)|~v4_valued_0(X4)|~v4_relat_1(X4,X1)|~v1_partfun1(X4,X1)|~v1_relat_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l2_algstr_0(X2))).#trainneg
cnf(i_0_915, negated_conjecture, (k1_tarski(k16_pre_poly(esk1_0))=k2_polynom1(k15_pre_poly(esk1_0),esk2_0,esk3_0)|k2_polynom1(k15_pre_poly(esk1_0),esk2_0,esk3_0)=k1_xboole_0|v2_struct_0(esk2_0)|~l2_struct_0(esk2_0))).#trainneg
cnf(i_0_408, plain, (k2_relset_1(k2_zfmisc_1(X1,X2),X3)=k1_xboole_0|X2!=k1_xboole_0|~v1_funct_2(k2_relset_1(k2_zfmisc_1(X1,X2),X3),X1,X2)|~v5_relat_1(X3,k2_zfmisc_1(X1,X2))|~v1_relat_1(X3))).#trainneg
cnf(i_0_891, plain, (X1=k5_polynom2(X2,X3,X4,X5)|r1_xxreal_0(np__1,esk14_6(X2,X3,X4,X5,X1,esk13_5(X2,X3,X4,X6,X7)))|v7_struct_0(X3)|X7!=k5_polynom2(X2,X3,X4,X6)|X1!=k4_rlvect_1(X3,esk13_5(X2,X3,X4,X6,X7))|~v1_polynom1(X4,k15_pre_poly(X2),X3)|~m2_finseq_1(esk13_5(X2,X3,X4,X6,X7),u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X1,u1_struct_0(X3))|~m1_subset_1(X7,u1_struct_0(X3))|~v1_funct_2(X4,k15_pre_poly(X2),u1_struct_0(X3))|~v1_funct_2(X5,X2,u1_struct_0(X3))|~v1_funct_2(X6,X2,u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X4)|~v1_funct_1(X6)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v13_algstr_0(X3)|~v3_ordinal1(X2))).#trainneg
cnf(i_0_783, plain, (k1_relset_1(k15_pre_poly(X1),k6_polynom1(X1,X2,X3,X4))=k15_pre_poly(X1)|u1_struct_0(X2)=k1_xboole_0|v2_struct_0(X2)|~v1_funct_2(k6_polynom1(X1,X2,X3,X4),k15_pre_poly(X1),u1_struct_0(X2))|~l2_algstr_0(X2)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v1_funct_2(X4,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v13_algstr_0(X2))).#trainneg
cnf(i_0_685, plain, (k3_funct_2(X1,u1_struct_0(X2),X3,X4)=k4_struct_0(X2)|v1_xboole_0(X1)|r2_hidden(X4,k1_relset_1(X1,X5))|k1_relset_1(X1,X5)!=k2_polynom1(X1,X2,X3)|~l2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,u1_struct_0(X2))|~v1_funct_1(X3)|~v4_relat_1(X5,X1)|~v1_relat_1(X5))).#trainneg
cnf(i_0_1023, plain, (v1_funct_1(k16_pre_poly(X1)))).#trainneg
cnf(i_0_1024, plain, (v1_relat_1(k16_pre_poly(X1)))).#trainneg
cnf(i_0_1013, plain, (v1_funct_1(X1)|k1_tarski(X1)!=k15_pre_poly(X2))).#trainneg
cnf(i_0_1017, plain, (m1_subset_1(X1,k1_tarski(X1)))).#trainneg
cnf(i_0_888, negated_conjecture, (k9_xtuple_0(esk3_0)=k15_pre_poly(esk1_0)|u1_struct_0(esk2_0)=k1_xboole_0|~v4_relat_1(esk3_0,k15_pre_poly(esk1_0)))).#trainneg
cnf(i_0_641, negated_conjecture, (k6_domain_1(k15_pre_poly(esk1_0),k16_pre_poly(esk1_0))=k2_polynom1(k15_pre_poly(esk1_0),esk2_0,esk3_0)|k2_polynom1(k15_pre_poly(esk1_0),esk2_0,esk3_0)=k1_xboole_0|v2_struct_0(esk2_0))).#trainneg
cnf(i_0_330, plain, (v1_relat_1(esk16_1(X1))|v4_rlvect_1(X1)|~v4_funct_1(u1_struct_0(X1))|~l2_algstr_0(X1))).#trainneg
cnf(i_0_939, negated_conjecture, (k1_zfmisc_1(k2_zfmisc_1(esk1_0,u1_struct_0(esk2_0)))=k1_xboole_0|r2_hidden(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_849, plain, (v1_xboole_0(k1_zfmisc_1(X1))|r2_hidden(X1,k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_990, plain, (v2_struct_0(X1)|~v1_xboole_0(u1_struct_0(X1))|~l6_algstr_0(X1))).#trainneg
cnf(i_0_331, plain, (v1_xboole_0(u1_struct_0(X1))|r2_hidden(esk12_1(X1),u1_struct_0(X1))|v2_struct_0(X1)|v4_vectsp_1(X1)|~l4_algstr_0(X1))).#trainneg
cnf(i_0_332, plain, (v1_funct_1(esk12_1(X1))|v2_struct_0(X1)|v4_vectsp_1(X1)|~v4_funct_1(u1_struct_0(X1))|~l4_algstr_0(X1))).#trainneg
cnf(i_0_1022, plain, (r2_hidden(k16_pre_poly(X1),k15_pre_poly(X1)))).#trainneg
cnf(i_0_985, plain, (k5_polynom2(X2,X1,X3,X5)=k5_polynom2(X2,X1,X3,X4)|v2_struct_0(X1)|v7_struct_0(X1)|~v3_polynom7(X3,X2,X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X1))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1))))|~v4_polynom7(X3,X2,X1)|~v1_funct_2(X3,k15_pre_poly(X2),u1_struct_0(X1))|~v1_funct_2(X4,X2,u1_struct_0(X1))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l6_algstr_0(X1)|~v5_vectsp_1(X1)|~v4_vectsp_1(X1)|~v4_rlvect_1(X1)|~v3_rlvect_1(X1)|~v13_algstr_0(X1)|~v3_ordinal1(X2)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X1))))|~v1_funct_2(X5,X2,u1_struct_0(X1))|~v1_funct_1(X5))).#trainneg
cnf(i_0_1093, plain, (k1_funct_1(X3,k16_pre_poly(X1))=k3_polynom7(X1,X2,X3)|v2_struct_0(X2)|~l2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v4_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~v2_pre_poly(k16_pre_poly(X1))|~v4_valued_0(k16_pre_poly(X1))|~v4_relat_1(k16_pre_poly(X1),X1)|~v1_partfun1(k16_pre_poly(X1),X1)|~l2_algstr_0(X2)|~v1_relat_1(k16_pre_poly(X1))|~v1_funct_1(k16_pre_poly(X1)))).#trainneg
cnf(i_0_1036, plain, (k1_funct_1(X3,k16_pre_poly(X1))=k3_polynom7(X1,X2,X3)|v2_struct_0(X2)|~l2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v4_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~v2_pre_poly(k16_pre_poly(X1))|~v4_valued_0(k16_pre_poly(X1))|~v4_relat_1(k16_pre_poly(X1),X1)|~v1_partfun1(k16_pre_poly(X1),X1)|~v1_relat_1(k16_pre_poly(X1))|~v1_funct_1(k16_pre_poly(X1)))).#trainneg
cnf(i_0_1031, plain, (k4_polynom1(X1,X2,X3,k7_polynom1(X1,X2))=X3|v2_struct_0(X2)|~l2_algstr_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v1_funct_2(k4_polynom1(X1,X2,X3,k7_polynom1(X1,X2)),k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(k4_polynom1(X1,X2,X3,k7_polynom1(X1,X2)))|~v1_funct_1(X3)|~v4_rlvect_1(X2)|~m1_subset_1(k7_polynom1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v1_funct_2(k7_polynom1(X1,X2),k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(k7_polynom1(X1,X2)))).#trainneg
cnf(i_0_684, plain, (k3_funct_2(X1,u1_struct_0(X2),X3,X4)=k4_struct_0(X2)|v1_xboole_0(X1)|r2_hidden(X4,k2_polynom1(X1,X5,X6))|k2_polynom1(X1,X5,X6)!=k2_polynom1(X1,X2,X3)|~l2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,u1_struct_0(X2))|~v1_funct_1(X3)|~l2_struct_0(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X5))))|~v1_funct_2(X6,X1,u1_struct_0(X5))|~v1_funct_1(X6))).#trainneg
cnf(i_0_661, plain, (v1_xboole_0(X1)|r2_hidden(esk5_4(X1,X2,X3,X4),X1)|X4=k2_polynom1(X1,X2,X3)|~l2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~v1_funct_2(X3,X1,u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_832, plain, (X1=k5_polynom2(X2,X3,X4,X5)|r1_xxreal_0(np__1,esk14_6(X2,X3,X4,X5,X1,esk13_5(X6,X7,X8,X9,X10)))|v7_struct_0(X3)|v7_struct_0(X7)|k3_finseq_1(k7_pre_poly(k15_pre_poly(X6),k2_polynom1(k15_pre_poly(X6),X7,X8),k17_pre_poly(X6)))!=k3_finseq_1(k7_pre_poly(k15_pre_poly(X2),k2_polynom1(k15_pre_poly(X2),X3,X4),k17_pre_poly(X2)))|X1!=k4_rlvect_1(X3,esk13_5(X6,X7,X8,X9,X10))|~v1_polynom1(X4,k15_pre_poly(X2),X3)|~m2_finseq_1(esk13_5(X6,X7,X8,X9,X10),u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X1,u1_struct_0(X3))|~v1_funct_2(X4,k15_pre_poly(X2),u1_struct_0(X3))|~v1_funct_2(X5,X2,u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v13_algstr_0(X3)|~v3_ordinal1(X2)|X10!=k5_polynom2(X6,X7,X8,X9)|~v1_polynom1(X8,k15_pre_poly(X6),X7)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X6),u1_struct_0(X7))))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X6,u1_struct_0(X7))))|~m1_subset_1(X10,u1_struct_0(X7))|~v1_funct_2(X8,k15_pre_poly(X6),u1_struct_0(X7))|~v1_funct_2(X9,X6,u1_struct_0(X7))|~v1_funct_1(X9)|~v1_funct_1(X8)|~l6_algstr_0(X7)|~v5_vectsp_1(X7)|~v4_vectsp_1(X7)|~v4_rlvect_1(X7)|~v3_rlvect_1(X7)|~v13_algstr_0(X7)|~v3_ordinal1(X6))).#trainneg
cnf(i_0_816, plain, (k1_algstr_0(X1,k3_polynom1(X2,X1,X3,k16_pre_poly(X2)),k3_polynom7(X2,X1,X4))=k3_polynom1(X2,X1,k4_polynom1(X2,X1,X3,X4),k16_pre_poly(X2))|v2_struct_0(X1)|~v2_pre_poly(k16_pre_poly(X2))|~v4_valued_0(k16_pre_poly(X2))|~v4_relat_1(k16_pre_poly(X2),X2)|~v1_partfun1(k16_pre_poly(X2),X2)|~l2_algstr_0(X1)|~v1_relat_1(k16_pre_poly(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X1))))|~v1_funct_2(X4,k15_pre_poly(X2),u1_struct_0(X1))|~v1_funct_2(X3,k15_pre_poly(X2),u1_struct_0(X1))|~v1_funct_1(k16_pre_poly(X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l2_struct_0(X1)|~v4_polynom7(X4,X2,X1))).#trainneg
cnf(i_0_523, plain, (k1_relset_1(k15_pre_poly(X1),k7_polynom1(X1,X2))=k15_pre_poly(X1)|u1_struct_0(X2)=k1_xboole_0|v2_struct_0(X2)|~v1_funct_2(k7_polynom1(X1,X2),k15_pre_poly(X1),u1_struct_0(X2))|~l2_struct_0(X2))).#trainneg
cnf(i_0_1012, plain, (X1=k5_polynom2(X2,X3,k7_polynom1(X2,X3),X4)|r1_xxreal_0(np__1,esk14_6(X2,X3,k7_polynom1(X2,X3),X4,X1,k7_pre_poly(k15_pre_poly(X2),k1_xboole_0,k17_pre_poly(X2))))|v7_struct_0(X3)|v2_struct_0(X3)|X1!=k4_rlvect_1(X3,k7_pre_poly(k15_pre_poly(X2),k1_xboole_0,k17_pre_poly(X2)))|~v1_polynom1(k7_polynom1(X2,X3),k15_pre_poly(X2),X3)|~m2_finseq_1(k7_pre_poly(k15_pre_poly(X2),k1_xboole_0,k17_pre_poly(X2)),u1_struct_0(X3))|~m1_subset_1(k7_polynom1(X2,X3),k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X1,u1_struct_0(X3))|~v1_funct_2(k7_polynom1(X2,X3),k15_pre_poly(X2),u1_struct_0(X3))|~v1_funct_2(X4,X2,u1_struct_0(X3))|~v1_funct_1(X4)|~v1_funct_1(k7_polynom1(X2,X3))|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v13_algstr_0(X3)|~v3_ordinal1(X2)|~l2_struct_0(X3))).#trainneg
cnf(i_0_1014, plain, (v1_relat_1(X1)|k1_tarski(X1)!=k15_pre_poly(X2))).#trainneg
cnf(i_0_931, negated_conjecture, (m1_subset_1(k7_partfun1(u1_struct_0(esk2_0),esk4_0,X1),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_1047, negated_conjecture, (k6_domain_1(k15_pre_poly(esk1_0),k16_pre_poly(esk1_0))=k2_polynom1(k15_pre_poly(esk1_0),esk2_0,esk3_0)|k7_polynom1(esk1_0,esk2_0)=esk3_0)).#trainneg
cnf(i_0_1018, plain, (X1=X2|k1_tarski(X1)!=k1_tarski(X2))).#trainneg
cnf(i_0_333, plain, (v1_relat_1(esk12_1(X1))|v2_struct_0(X1)|v4_vectsp_1(X1)|~v4_funct_1(u1_struct_0(X1))|~l4_algstr_0(X1))).#trainneg
cnf(i_0_999, negated_conjecture, (k7_partfun1(X1,esk4_0,X2)=k1_funct_1(esk4_0,X2)|~v5_relat_1(esk4_0,X1)|~r2_hidden(X2,esk1_0)|~v1_partfun1(esk4_0,esk1_0))).#trainneg
cnf(i_0_1101, negated_conjecture, (r2_hidden(k16_pre_poly(esk1_0),k2_polynom1(k15_pre_poly(esk1_0),esk2_0,esk3_0))|k2_polynom1(k15_pre_poly(esk1_0),esk2_0,esk3_0)=k1_xboole_0)).#trainneg
cnf(i_0_1058, negated_conjecture, (k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(esk1_0),u1_struct_0(esk2_0)))=k1_xboole_0|r2_hidden(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(esk1_0),u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_1080, plain, (v1_xboole_0(k1_zfmisc_1(X1))|m2_subset_1(X1,k1_zfmisc_1(X1),k1_zfmisc_1(X1)))).#trainneg
cnf(i_0_335, negated_conjecture, (v1_xboole_0(esk4_0)|m2_subset_1(X1,k2_zfmisc_1(esk1_0,u1_struct_0(esk2_0)),esk4_0)|~m1_subset_1(X1,esk4_0))).#trainneg
cnf(i_0_336, negated_conjecture, (v1_xboole_0(esk3_0)|m2_subset_1(X1,k2_zfmisc_1(k15_pre_poly(esk1_0),u1_struct_0(esk2_0)),esk3_0)|~m1_subset_1(X1,esk3_0))).#trainneg
cnf(i_0_339, plain, (k1_funct_1(k9_finseq_1(X1),np__1)=X1|~v1_finseq_1(k9_finseq_1(X1))|~v1_relat_1(k9_finseq_1(X1))|~v1_funct_1(k9_finseq_1(X1)))).#trainneg
cnf(i_0_420, plain, (esk8_2(X1,X2)=X3|X1=X2|r2_hidden(esk8_2(X1,X2),X1)|X2!=k1_tarski(X3))).#trainneg
cnf(i_0_1030, plain, (k1_polynom6(X1,X2,k16_pre_poly(X1),k16_pre_poly(X2))=k16_pre_poly(X2)|X1!=k1_xboole_0|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_991, plain, (k3_polynom1(X1,X2,X3,X4)=k1_funct_1(X3,X4)|v2_struct_0(X2)|~v2_pre_poly(X4)|~v4_valued_0(X4)|~v4_relat_1(X4,X1)|~v1_partfun1(X4,X1)|~v1_relat_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_1027, plain, (k1_polynom6(X1,X2,k16_pre_poly(X1),k16_pre_poly(X2))=k16_pre_poly(X1)|X2!=k1_xboole_0|~v3_ordinal1(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_820, plain, (k1_algstr_0(X1,k3_polynom7(X2,X1,X3),k3_polynom1(X2,X1,X4,k16_pre_poly(X2)))=k3_polynom1(X2,X1,k4_polynom1(X2,X1,X3,X4),k16_pre_poly(X2))|v2_struct_0(X1)|~v2_pre_poly(k16_pre_poly(X2))|~v4_valued_0(k16_pre_poly(X2))|~v4_relat_1(k16_pre_poly(X2),X2)|~v1_partfun1(k16_pre_poly(X2),X2)|~l2_algstr_0(X1)|~v1_relat_1(k16_pre_poly(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X1))))|~v1_funct_2(X4,k15_pre_poly(X2),u1_struct_0(X1))|~v1_funct_2(X3,k15_pre_poly(X2),u1_struct_0(X1))|~v1_funct_1(k16_pre_poly(X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l2_struct_0(X1)|~v4_polynom7(X3,X2,X1))).#trainneg
cnf(i_0_1173, plain, (k3_funct_2(X1,u1_struct_0(X2),X3,X4)=k4_struct_0(X2)|v1_xboole_0(X1)|r2_hidden(X4,k2_polynom1(X1,X2,X3))|~l2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~m1_subset_1(X4,X1)|~v1_funct_2(X3,X1,u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1243, plain, (k1_funct_1(X3,X4)=k4_struct_0(X2)|X4=k16_pre_poly(X1)|v2_struct_0(X2)|~v2_pre_poly(X4)|~v4_valued_0(X4)|~v4_relat_1(X4,X1)|~v1_partfun1(X4,X1)|~l2_struct_0(X2)|~v1_relat_1(X4)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v4_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l6_algstr_0(X2))).#trainneg
cnf(i_0_1185, plain, (esk5_4(X1,X2,X3,X4)=X5|X4=k2_polynom1(X1,X2,X3)|v1_xboole_0(X1)|X1!=k1_tarski(X5)|~l2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(X1,u1_struct_0(X2))))|~m1_subset_1(X4,k1_zfmisc_1(X1))|~v1_funct_2(X3,X1,u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_1201, plain, (k9_xtuple_0(k7_polynom1(X1,X2))=k15_pre_poly(X1)|u1_struct_0(X2)=k1_xboole_0|v2_struct_0(X2)|~v4_relat_1(k7_polynom1(X1,X2),k15_pre_poly(X1))|~v1_relat_1(k7_polynom1(X1,X2))|~l2_struct_0(X2))).#trainneg
cnf(i_0_1117, plain, (v1_funct_1(X1)|~v4_funct_1(k1_tarski(X1)))).#trainneg
cnf(i_0_936, negated_conjecture, (m1_subset_1(k7_partfun1(u1_struct_0(esk2_0),esk3_0,X1),u1_struct_0(esk2_0)))).#trainneg
cnf(i_0_1120, negated_conjecture, (m1_subset_1(k16_pre_poly(esk1_0),k2_polynom1(k15_pre_poly(esk1_0),esk2_0,esk3_0))|k2_polynom1(k15_pre_poly(esk1_0),esk2_0,esk3_0)=k1_xboole_0)).#trainneg
cnf(i_0_1128, negated_conjecture, (esk4_0=X1|k1_zfmisc_1(k2_zfmisc_1(esk1_0,u1_struct_0(esk2_0)))=k1_xboole_0|k1_tarski(X1)!=k1_zfmisc_1(k2_zfmisc_1(esk1_0,u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_340, plain, (v1_partfun1(X1,X2)|k9_xtuple_0(X1)!=X2|~v4_relat_1(X1,X2)|~v1_relat_1(X1))).#trainneg
cnf(i_0_1118, plain, (v1_relat_1(X1)|~v4_funct_1(k1_tarski(X1)))).#trainneg
cnf(i_0_1066, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(esk1_0,u1_struct_0(esk2_0))))|m2_subset_1(esk4_0,k1_zfmisc_1(k2_zfmisc_1(esk1_0,u1_struct_0(esk2_0))),k1_zfmisc_1(k2_zfmisc_1(esk1_0,u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_1131, plain, (v1_funct_1(X1)|v1_xboole_0(k1_zfmisc_1(X1))|k1_zfmisc_1(X1)!=k15_pre_poly(X2))).#trainneg
cnf(i_0_341, plain, (r2_hidden(np__1,k1_relset_1(k5_numbers,X1))|esk7_2(X2,k10_xtuple_0(X1))=X2|k10_xtuple_0(X1)=k1_tarski(X2)|~v1_finseq_1(X1)|~v1_relat_1(X1)|~v1_funct_1(X1))).#trainneg
cnf(i_0_342, plain, (v1_funct_1(esk7_2(X1,X2))|esk7_2(X1,X2)=X1|X2=k1_tarski(X1)|X2!=k15_pre_poly(X3))).#trainneg
cnf(i_0_1203, plain, (v1_partfun1(k7_polynom1(X1,X2),k15_pre_poly(X1))|u1_struct_0(X2)=k1_xboole_0|v2_struct_0(X2)|~v4_relat_1(k7_polynom1(X1,X2),k15_pre_poly(X1))|~v1_relat_1(k7_polynom1(X1,X2))|~l2_struct_0(X2))).#trainneg
cnf(i_0_478, plain, (v5_relat_1(k7_polynom1(X1,X2),u1_struct_0(X2))|v2_struct_0(X2)|~l2_struct_0(X2))).#trainneg
cnf(i_0_672, plain, (k4_polynom7(X2,X3,esk18_3(X2,X3,X1))=X1|v2_struct_0(X3)|~m1_subset_1(k4_polynom7(X2,X3,esk18_3(X2,X3,X1)),k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3))))|~v1_funct_2(k4_polynom7(X2,X3,esk18_3(X2,X3,X1)),k15_pre_poly(X2),u1_struct_0(X3))|~v1_funct_2(X1,k15_pre_poly(X2),u1_struct_0(X3))|~v1_funct_1(k4_polynom7(X2,X3,esk18_3(X2,X3,X1)))|~v1_funct_1(X1)|~l2_struct_0(X3)|~v4_polynom7(X1,X2,X3))).#trainneg
cnf(i_0_1246, plain, (k1_funct_1(k7_polynom1(X1,X2),X3)=k4_struct_0(X2)|v2_struct_0(X2)|~v2_pre_poly(X3)|~v4_valued_0(X3)|~v4_relat_1(X3,X1)|~v1_partfun1(X3,X1)|~l2_struct_0(X2)|~v1_relat_1(X3)|~v1_funct_1(X3)|~m1_subset_1(k7_polynom1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v1_funct_2(k7_polynom1(X1,X2),k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(k7_polynom1(X1,X2))|~l6_algstr_0(X2))).#trainneg
cnf(i_0_1155, plain, (v2_pre_poly(k16_pre_poly(X1))|k15_pre_poly(X1)!=k15_pre_poly(X2))).#trainneg
cnf(i_0_1154, plain, (v4_valued_0(k16_pre_poly(X1))|k15_pre_poly(X1)!=k15_pre_poly(X2))).#trainneg
cnf(i_0_1059, negated_conjecture, (k7_partfun1(X1,esk4_0,X2)=k1_funct_1(esk4_0,X2)|u1_struct_0(esk2_0)=k1_xboole_0|~v5_relat_1(esk4_0,X1)|~r2_hidden(X2,esk1_0))).#trainneg
cnf(i_0_833, plain, (X1=k5_polynom2(X2,X3,X4,X5)|r1_xxreal_0(np__1,esk14_6(X2,X3,X4,X5,X1,k7_pre_poly(k15_pre_poly(X6),k2_polynom1(k15_pre_poly(X6),X7,X8),k17_pre_poly(X6))))|v7_struct_0(X3)|v7_struct_0(X7)|k3_finseq_1(esk13_5(X6,X7,X8,X9,X10))!=k3_finseq_1(k7_pre_poly(k15_pre_poly(X2),k2_polynom1(k15_pre_poly(X2),X3,X4),k17_pre_poly(X2)))|X1!=k4_rlvect_1(X3,k7_pre_poly(k15_pre_poly(X6),k2_polynom1(k15_pre_poly(X6),X7,X8),k17_pre_poly(X6)))|~v1_polynom1(X4,k15_pre_poly(X2),X3)|~m2_finseq_1(k7_pre_poly(k15_pre_poly(X6),k2_polynom1(k15_pre_poly(X6),X7,X8),k17_pre_poly(X6)),u1_struct_0(X3))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X3))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(X2,u1_struct_0(X3))))|~m1_subset_1(X1,u1_struct_0(X3))|~v1_funct_2(X4,k15_pre_poly(X2),u1_struct_0(X3))|~v1_funct_2(X5,X2,u1_struct_0(X3))|~v1_funct_1(X5)|~v1_funct_1(X4)|~l6_algstr_0(X3)|~v5_vectsp_1(X3)|~v4_vectsp_1(X3)|~v4_rlvect_1(X3)|~v3_rlvect_1(X3)|~v13_algstr_0(X3)|~v3_ordinal1(X2)|X10!=k5_polynom2(X6,X7,X8,X9)|~v1_polynom1(X8,k15_pre_poly(X6),X7)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X6),u1_struct_0(X7))))|~m1_subset_1(X9,k1_zfmisc_1(k2_zfmisc_1(X6,u1_struct_0(X7))))|~m1_subset_1(X10,u1_struct_0(X7))|~v1_funct_2(X8,k15_pre_poly(X6),u1_struct_0(X7))|~v1_funct_2(X9,X6,u1_struct_0(X7))|~v1_funct_1(X9)|~v1_funct_1(X8)|~l6_algstr_0(X7)|~v5_vectsp_1(X7)|~v4_vectsp_1(X7)|~v4_rlvect_1(X7)|~v3_rlvect_1(X7)|~v13_algstr_0(X7)|~v3_ordinal1(X6))).#trainneg
cnf(i_0_1015, plain, (v4_valued_0(X1)|k1_tarski(X1)!=k15_pre_poly(X2))).#trainneg
cnf(i_0_1060, plain, (k2_relset_1(X1,k2_zfmisc_1(X2,X1))=k10_xtuple_0(k2_zfmisc_1(X2,X1)))).#trainneg
cnf(i_0_1067, negated_conjecture, (v1_xboole_0(k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(esk1_0),u1_struct_0(esk2_0))))|m2_subset_1(esk3_0,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(esk1_0),u1_struct_0(esk2_0))),k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(esk1_0),u1_struct_0(esk2_0)))))).#trainneg
cnf(i_0_1232, negated_conjecture, (esk3_0=X1|k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(esk1_0),u1_struct_0(esk2_0)))=k1_xboole_0|k1_tarski(X1)!=k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(esk1_0),u1_struct_0(esk2_0))))).#trainneg
cnf(i_0_343, plain, (v1_relat_1(esk7_2(X1,X2))|esk7_2(X1,X2)=X1|X2=k1_tarski(X1)|X2!=k15_pre_poly(X3))).#trainneg
cnf(i_0_938, negated_conjecture, (r2_hidden(esk3_0,X1)|X1!=k15_pre_poly(k15_pre_poly(esk1_0))|~v2_pre_poly(esk3_0)|~v4_valued_0(esk3_0)|~v1_partfun1(esk3_0,k15_pre_poly(esk1_0)))).#trainneg
cnf(i_0_992, negated_conjecture, (r2_hidden(np__1,k1_relset_1(k5_numbers,esk4_0))|~r2_hidden(X1,k2_relset_1(u1_struct_0(esk2_0),esk4_0))|~v1_finseq_1(esk4_0))).#trainneg
cnf(i_0_1053, negated_conjecture, (r2_hidden(np__1,k1_relset_1(k5_numbers,esk3_0))|~r2_hidden(X1,k2_relset_1(u1_struct_0(esk2_0),esk3_0))|~v1_finseq_1(esk3_0))).#trainneg
cnf(i_0_1132, plain, (v1_relat_1(X1)|v1_xboole_0(k1_zfmisc_1(X1))|k1_zfmisc_1(X1)!=k15_pre_poly(X2))).#trainneg
cnf(i_0_1133, plain, (v4_valued_0(X1)|v1_xboole_0(k1_zfmisc_1(X1))|k1_zfmisc_1(X1)!=k15_pre_poly(X2))).#trainneg
cnf(i_0_344, plain, (v4_valued_0(esk7_2(X1,X2))|esk7_2(X1,X2)=X1|X2=k1_tarski(X1)|X2!=k15_pre_poly(X3))).#trainneg
cnf(i_0_345, plain, (v2_pre_poly(esk7_2(X1,X2))|esk7_2(X1,X2)=X1|X2=k1_tarski(X1)|X2!=k15_pre_poly(X3))).#trainneg
cnf(i_0_616, plain, (k1_funct_1(k16_pre_poly(k10_ordinal2(X1,X2)),X3)=k1_funct_1(X4,X3)|k16_pre_poly(k10_ordinal2(X1,X2))!=k1_polynom6(X1,X2,X4,X5)|~r2_hidden(X3,X1)|~m2_subset_1(X5,k15_pre_poly(X2),k15_pre_poly(X2))|~m2_subset_1(X4,k15_pre_poly(X1),k15_pre_poly(X1))|~v3_ordinal1(X2)|~v3_ordinal1(X1)|~v3_ordinal1(X3))).#trainneg
cnf(i_0_723, plain, (v1_xboole_0(k2_polynom1(k15_pre_poly(X1),X2,X3))|r2_hidden(o_3_2_polynom7(X1,X2,X3),k2_polynom1(k15_pre_poly(X1),X2,X3))|v7_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v4_polynom7(X3,X1,X2)|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(X3)|~l6_algstr_0(X2)|~v5_vectsp_1(X2)|~v4_vectsp_1(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v13_algstr_0(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_626, plain, (k3_polynom7(X1,X2,k7_polynom1(X1,X2))=k4_struct_0(X2)|v2_struct_0(X2)|~v2_pre_poly(k16_pre_poly(X1))|~v4_valued_0(k16_pre_poly(X1))|~v4_relat_1(k16_pre_poly(X1),X1)|~v1_partfun1(k16_pre_poly(X1),X1)|~l2_struct_0(X2)|~v1_relat_1(k16_pre_poly(X1))|~v1_funct_1(k16_pre_poly(X1))|~m1_subset_1(k7_polynom1(X1,X2),k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v4_polynom7(k7_polynom1(X1,X2),X1,X2)|~v1_funct_2(k7_polynom1(X1,X2),k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(k7_polynom1(X1,X2)))).#trainneg
cnf(i_0_923, plain, (k3_polynom7(X1,X2,X3)=X4|v2_struct_0(X2)|u1_struct_0(X2)!=k1_tarski(X4)|~v3_polynom7(X3,X1,X2)|~l2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(X3))).#trainneg
cnf(i_0_803, plain, (k1_relset_1(k15_pre_poly(X1),k9_polynom1(X1,X2,X3,X4))=k15_pre_poly(X1)|u1_struct_0(X2)=k1_xboole_0|v2_struct_0(X2)|~v1_funct_2(k9_polynom1(X1,X2,X3,X4),k15_pre_poly(X1),u1_struct_0(X2))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X1),u1_struct_0(X2))))|~v1_funct_2(X4,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_2(X3,k15_pre_poly(X1),u1_struct_0(X2))|~v1_funct_1(X4)|~v1_funct_1(X3)|~l6_algstr_0(X2)|~v4_rlvect_1(X2)|~v3_rlvect_1(X2)|~v13_algstr_0(X2)|~v3_ordinal1(X1))).#trainneg
cnf(i_0_1251, plain, (k1_algstr_0(X1,k3_polynom1(X2,X1,X3,X4),k1_funct_1(X5,X4))=k3_polynom1(X2,X1,k4_polynom1(X2,X1,X3,X5),X4)|v2_struct_0(X1)|~v2_pre_poly(X4)|~v4_valued_0(X4)|~v4_relat_1(X4,X2)|~v1_partfun1(X4,X2)|~l2_algstr_0(X1)|~v1_relat_1(X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X1))))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(k15_pre_poly(X2),u1_struct_0(X1))))|~v1_funct_2(X5,k15_pre_poly(X2),u1_struct_0(X1))|~v1_funct_2(X3,k15_pre_poly(X2),u1_struct_0(X1))|~v1_funct_1(X4)|~v1_funct_1(X5)|~v1_funct_1(X3)|~l6_algstr_0(X1))).#trainneg
cnf(i_0_1325, plain, (v2_pre_poly(k16_pre_poly(X1)))).#trainneg
cnf(i_0_1326, plain, (v4_valued_0(k16_pre_poly(X1)))).#trainneg
cnf(i_0_1307, negated_conjecture, (v1_partfun1(esk4_0,esk1_0)|k9_xtuple_0(esk4_0)!=esk1_0)).#trainneg
# Training: Negative examples end
# Parsed axioms                        : 129
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 287
# Removed in clause preprocessing      : 15
# Initial clauses in saturation        : 272
# Processed clauses                    : 466
# ...of these trivial                  : 5
# ...subsumed                          : 21
# ...remaining for further processing  : 440
# Other redundant clauses eliminated   : 3
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 1
# Backward-rewritten                   : 10
# Generated clauses                    : 1139
# ...of the previous two non-trivial   : 1070
# Contextual simplify-reflections      : 50
# Paramodulations                      : 1108
# Factorizations                       : 2
# Equation resolutions                 : 29
# Propositional unsat checks           : 0
#    Propositional check models        : 0
#    Propositional check unsatisfiable : 0
#    Propositional clauses             : 0
#    Propositional clauses after purity: 0
#    Propositional unsat core size     : 0
#    Propositional preprocessing time  : 0.000
#    Propositional encoding time       : 0.000
#    Propositional solver time         : 0.000
#    Success case prop preproc time    : 0.000
#    Success case prop encoding time   : 0.000
#    Success case prop solver time     : 0.000
# Current number of processed clauses  : 426
#    Positive orientable unit clauses  : 45
#    Positive unorientable unit clauses: 1
#    Negative unit clauses             : 4
#    Non-unit-clauses                  : 376
# Current number of unprocessed clauses: 866
# ...number of literals in the above   : 9934
# Current number of archived formulas  : 0
# Current number of archived clauses   : 453
# Proof object given clauses           : 45
# Proof search given clauses           : 440
# Clause-clause subsumption calls (NU) : 191400
# Rec. Clause-clause subsumption calls : 3825
# Non-unit clause-clause subsumptions  : 71
# Unit Clause-clause subsumption calls : 4432
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 52
# BW rewrite match successes           : 5
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 84806

# -------------------------------------------------
# User time                : 13.913 s
# System time              : 0.576 s
# Total time               : 14.490 s
# Maximum resident set size: 366192 pages
13.91user 0.64system 0:14.65elapsed 99%CPU (0avgtext+0avgdata 366824maxresident)k
0inputs+264outputs (0major+148089minor)pagefaults 0swaps
